---
id: 6153908a366afb4d57185c8d
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Flexbox ha un asse principale e un asse trasversale. L'asse principale è definito dalla proprietà `flex-direction`, che può assumere quattro valori possibili:

- `row` (predefinito): asse orizzontale con elementi flex da sinistra a destra
- `row-reverse`: asse orizzontale con elementi flex da destra a sinistra
- `column`: asse verticale con elementi flex dall'alto verso il basso
- `column-reverse`: asse verticale con elementi flex dal basso verso l'alto

**Nota**: gli assi e le direzioni saranno diversi a seconda della direzione del testo. I valori mostrati sono per una direzione del testo da sinistra a destra.

Prova i diversi valori per vedere come influenzano il layout.

Quando hai finito, imposta un `flex-direction` esplicito di `row` sull'elemento `.gallery`.

# --hints--

Il tuo selettore `.gallery` dovrebbe avere una proprietà `flex-direction` impostata con valore `row`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection, 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}


.gallery {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
