---
id: 6153897c27f6334716ee5abe
title: Passaggio 12
challengeType: 0
dashedName: step-12
---

# --description--

Flexbox è un layout CSS monodimensionale che può controllare il modo in cui gli elementi sono distribuiti e allineati all'interno di un contenitore.

Per usarlo, assegna a un elemento una proprietà `display` con valore `flex`. Questo trasformerà l'elemento in un <em>contenitore flex</em>. Qualsiasi figlio diretto di un contenitore flex è chiamato <em>elemento flex</em>.

Crea un selettore `.gallery` e rendilo un contenitore flex.



# --hints--

Dovresti avere un selettore `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Il tuo selettore `.gallery` dovrebbe avere una proprietà `display` impostata su `flex` come valore.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
