---
id: 64dd076a1b33c86d84b4232f
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Per visualizzare meglio come sono dimensionati i tuoi elementi, aggiungere un bordo può essere utile.

Dai al tuo elemento `.gallery` un `width` di `50%` e una proprietà `border` impostata su `5px solid red`.

Poi dai ai tuoi elementi `img` un `width` di `100%`, una proprietà `padding` impostata su `5px` e una proprietà `border` impostata su `5px solid blue`.

# --hints--

Dovresti avere un selettore `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Il tuo selettore `.gallery` dovrebbe avere una proprietà `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

Il tuo selettore `.gallery` dovrebbe avere una proprietà `width` impostata su `50%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

Il tuo selettore `.gallery` dovrebbe avere una proprietà `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

Il tuo selettore `.gallery` dovrebbe avere una proprietà `border` impostata su `5px solid red`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

Dovresti avere un selettore `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

Il tuo selettore `img` dovrebbe avere una proprietà `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

Il tuo selettore `img` dovrebbe avere una proprietà `width` impostata su `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

Il tuo selettore `img` dovrebbe avere una proprietà `padding`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

Il tuo selettore `img` dovrebbe avere una proprietà `padding` impostata su `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

Il tuo selettore `img` dovrebbe avere una proprietà `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

Il tuo selettore `img` dovrebbe avere una proprietà `border` impostata su `5px solid blue`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
