---
id: 68a23004e96a105df3a21fc4
title: Construir um Assistente de Contagem de Cartas
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

No jogo de casino Blackjack, um jogador pode determinar se tem uma vantagem sobre a próxima mão da casa, mantendo o número relativo de cartas altas e baixas restantes no baralho.Isso se chama "contar as cartas".

Ter cartas mais altas restantes no baralho favorece o jogador.Quando o contador for positivo, o jogador deve apostar alto.Quando a contagem for zero ou negativa, o jogador deverá apostar baixo.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve usar `let` para declarar uma variável global chamada `count` e defini-la como `0`.
1. Você deve ter uma função chamada `cc`.
1. A função `cc` deve receber um parâmetro `card` que pode ser um número ou uma string.
    - Para valores entre `2` e `10`, o parâmetro `card` será um número.
    - Para todos os outros valores, o parâmetro `card` será uma string.
1. A função `cc` deve modificar a variável global `count` com base em certos critérios.
1. A variável global `count` deve ser incrementada em `1` para os cartões `2`, `3`, `4`, `5` ou `6`
1. A variável global `count` deve permanecer inalterada para as cartas `7`, `8`, `9`.
1. A variável global `count` deve ser decrementada em `1` para as cartas `10`, `"J"`, `"Q"`, `"K"`, `"A"`
1. A função `cc` deve retornar uma string com a contagem atual e a string `Bet` se a contagem for positiva.
1. A função `cc` deve retornar uma string com a contagem atual e a string `Hold` se a contagem for menor ou igual a `0`.
1. Na saída da função, a contagem atual e a decisão do jogador (`Bet` ou `Hold`) devem ser separadas por um espaço. Por exemplo, `-3 Hold`.

# --hints--

Você deve usar `let` para declarar uma variável global chamada `count` e defini-la como `0`. 

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

Você deve ter uma função chamada `cc`. 

```js
assert.isFunction(cc);
```

Sua função deve retornar o valor de `count` e o texto (`Bet` ou `Hold`) com um caractere de espaço entre eles.

```js
count = 0;
let out = cc(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

Após as cartas `2`, `3`, `4`, `5`, então chamar `cc(6)` deve retornar a string `5 Bet`.

```js
count = 0;
cc(2);
cc(3);
cc(4);
cc(5);
const out = cc(6);
assert.strictEqual(out, '5 Bet')
```

Após as cartas `7`, `8`, então chamar `cc(9)` deve retornar a string `0 Hold`.

```js
count = 0;
cc(7);
cc(8);
const out = cc(9);
assert.strictEqual(out, '0 Hold');
```

Após as cartas `10`, `"J"`, `"Q"`, `"K"`, então chamar `cc("A")` deve retornar a string `-5 Hold`.

```js
count = 0;
cc(10);
cc('J');
cc('Q');
cc('K');
const out = cc('A');
assert.strictEqual(out, '-5 Hold');
```

Após as cartas `3`, `7`, `"Q"`, `8`, então chamar `cc("A")` deve retornar a string `-1 Hold`.

```js
count = 0;
cc(3);
cc(7);
cc('Q');
cc(8);
const out = cc('A');
assert.strictEqual(out, '-1 Hold');
```

Após as cartas `2`, `"J"`, `9`, `2`, então chamar `cc(7)` deve retornar a string `1 Bet`.

```js
count = 0;
cc(2);
cc('J');
cc(9);
cc(2);
const out = cc(7);
assert.strictEqual(out, '1 Bet');
```

Após as cartas `2`, `2`, então chamar `cc(10)` deve retornar a string `1 Aposta`.

```js
count = 0;
cc(2);
cc(2);
const out = cc(10);
assert.strictEqual(out, '1 Bet')
```

Após as cartas `3`, `2`, `"A"`, `10`, então chamar `cc("K")` deve retornar a string `-1 Hold`.

```js
count = 0;
cc(3);
cc(2);
cc('A');
cc(10);
const out = cc('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cc(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
