---
id: a6e40f1041b06c996f7b2406
title: 첫 번째 요소 찾기 함수 만들기
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

이 실습에서는 배열을 순차적으로 반복하며 테스트 함수(일명 "진리 테스트")를 통과하는 첫 번째 요소를 반환하는 함수를 만드세요.

함수는 배열을 순차적으로 반복하며 제공된 테스트 함수를 사용해 각 요소를 검사합니다. 마지막에 테스트 함수가 `true`를 반환하는 첫 번째 요소를 반환합니다.

예시:

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `findElement`라는 이름의 함수가 있어야 하며, 배열과 함수를 인수로 받아야 합니다.
2. 함수는 진리 테스트를 통과하는 배열의 첫 번째 항목을 반환해야 합니다. 즉, 전달된 함수 `func`를 요소 `x`에 호출했을 때, `func(x)`가 `true`이면 진리 테스트를 통과한 것입니다.
3. 어떤 요소도 테스트를 통과하지 못하면 함수는 `undefined`를 반환해야 합니다.

# --hints--

`findElement` 함수가 있어야 합니다.

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })`는 `8`를 반환해야 합니다.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })`는 `undefined`를 반환해야 합니다.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })`는 `3`를 반환해야 합니다.

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })`는 `"javascript"`를 반환해야 합니다.

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })`는 `undefined`를 반환해야 합니다.

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })`는 `2`를 반환해야 합니다.

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })`는 `undefined`를 반환해야 합니다.

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
