---
id: 685239f0cad43c7a909cbd98
title: Depurar una página de adopción de mascotas
challengeType: 25
dashedName: lab-debug-pet-adoption-page
---

# --description--

Sally, dueña de una tienda de adopción de mascotas, ha creado su primera página web pero hay algunos problemas.

Tu trabajo es corregir todos los errores para que Sally pueda seguir construyendo su página.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Historias de usuario:**

1. Sally quiere usar una imagen de algunos gatos pero no se muestra correctamente. Necesitarás corregir lo siguiente en el `img` elemento:
   - Reemplaza el atributo `href` con el atributo correcto para la fuente de la imagen.
   - Reemplaza el atributo `att` con el atributo correcto que representa texto breve y descriptivo para las imágenes.
   - Elimina la etiqueta de cierre `</img>` ya que los elementos `img` son elementos vacíos y no tienen etiquetas de cierre.
2. Sally quiere usar algunos enlaces para dirigir a los usuarios a las páginas de perros y gatos. Pero los enlaces no funcionan correctamente. Necesitarás corregir lo siguiente en los elementos `a`:
   - Reemplaza ambos atributos `src` con los atributos correctos usados para especificar URLs.

# --hints--

Tu elemento `img` debería tener un atributo `src` en lugar del atributo `href`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

Tu elemento `img` debería tener un atributo `alt` en lugar del atributo inexistente `att`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

Tu elemento `img` no debería tener una etiqueta de cierre `</img>`.

```js
assert.notMatch(code, /<\/img>/);
```

Tu elemento `a` con el texto `Visitar página de gatos` necesita tener un atributo `href` en lugar de un atributo `src`.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

Tu elemento `a` con el texto `Visitar página de perros` necesita tener un atributo `href` en lugar de un atributo `src`.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Five cats looking around a field."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
