---
id: 66d97c3be4030e8f41a92083
title: Crie o algoritmo de ordenação rápida (quicksort)
challengeType: 14
dashedName: build-the-quicksort-algorithm
---

# --description--

O Quicksort é um algoritmo que usa a abordagem de dividir e conquistar para ordenar arrays.

Ele pode ser implementado da seguinte maneira:

- Escolha um valor pivô entre os elementos do array de entrada (use o primeiro ou o último elemento do array).
- Divida o array de entrada em três subarrays contendo elementos maiores, menores e iguais ao valor do pivô.
- Recursivamente, chame `quicksort` para ordenar os subarrays.

**Objetivo:** atender às histórias de usuário abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve definir uma função chamada `quicksort` para implementar o algoritmo de quicksort.

1. A função `quicksort` deve receber um array de números inteiros como entrada e retornar um array desses números inteiros, ordenado do menor para o maior.

# --hints--

`quicksort` deve ser uma função.

```js
assert.isFunction(quicksort);
```

`quicksort` deve retornar um array ordenado (do menor para o maior elemento).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}

assert(
  isSorted(
    quicksort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`quicksort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` deve retornar um array inalterado, exceto por sua ordenação.

```js
assert.sameMembers(
  quicksort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` não deve usar o método `.sort()` integrado.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  quicksort([0, 1]);
  Array.prototype.sort = temp;
  return !sortUsed;
}

assert(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const quicksort = (array) => {
    if (array.length === 0) {
        return [];
    }
    const pivot = array[0];
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let i of array) {
        if (i < pivot) {
            lesser.push(i);
        } else if (i > pivot) {
            greater.push(i);
        } else {
            equal.push(i);
        }
    }
    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
}

```
