---
id: 6551f4259d66026ee28cb205
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

你也可以從字串的結尾開始存取字串字元。最後一個字元的索引是 `-1`，倒數第二個是 `-2`，依此類推。

現在修改你現有的 `print()` 呼叫，以列印你字串中的最後一個字元。

# --hints--

你仍然應該呼叫 `print()` 函式（程式）。

```js
assert.match(code, /print\s*\(.*\)/)
```

你應該將 `text[-1]` 傳遞給 `print()` 函式（程式），方法是將它放在小括弧內。請注意將函式（程式）呼叫放在行首。

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
