---
id: 6551f22653e6ce6e4f1d1370
title: 步驟11
challengeType: 20
dashedName: step-11
---

# --description--

修改您的代碼`print(shift)`來調用並打印`shift`的數據類型.

# --hints--

你應該保留現有的 `print(type(text))` 調用。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

你應該將 `shift` 傳遞給 `type()` 函數。

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

你應該將 `type(shift)` 傳遞給 `print()` 函數。 注意將函數調用放在行首。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

你的代碼中不應該有 `print(shift)`。

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

你的代碼包含無效語法和/或無效縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
