---
id: 6551f1e8737dbe6e0cb08ea4
title: 步驟 10
challengeType: 20
dashedName: step-10
---

# --description--

現在打印您的新變量。

# --hints--

你不應刪除現有的 `print(type(text))` 調用。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

你應該調用傳入 `shift` 變量的 `print()` 函數。 注意將函數調用放在行首。

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

你的代碼包含無效語法和/或無效縮進

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
