---
id: 6551ecf597038c36637f3759
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

變數是 Python 和任何程式語言中不可或缺的零件。變數是一個名稱，用來取用或指向一個物件。你可以透過在指定運算子 `=` 左邊寫變數名稱，並在指定運算子右邊指定要為該變數指定的值，來宣告一個變數：

```py
variable_name = value
```

創建一個名為 `number` 的變數，並將值 `5` 指定給你的新變數。

# --hints--

你應該宣告一個名為 `number` 的變數。請注意將變數名稱放在行首。

```js
assert.match(code, /^number\s*=/m)
```

你應該將值 `5` 指定到你的 `number` 變數。

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
