---
id: 6551ef504d91936d2d4e54f8
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

*引數* 是在呼叫函式時，傳遞給函式的一個物件或表達式——加在開啟和關閉小括弧之間——：

```py
greet = 'Hello!'
print(greet)
```

上面範例中的程式碼會列印字串 `'Hello!'`，該字串是作為引數傳遞給 `print()` 的變數 `greet` 的值。

將 `text` 變數作為引數傳遞給 `print()` 函式，以將你的 `text` 變數列印到螢幕上。

# --hints--

你應該將 `text` 傳遞給 `print()` 函式（程式），方法是在開啟和關閉小括弧內加入此變數的名稱。請注意將函式（程式）呼叫放在該行的開頭。

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
