---
id: 6551fe3b1df7c9740f13f270
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

Python 中變數命名的關鍵方面有：

- 有些字是保留關鍵字（例如 `for`、`while`、`True`）。它們在 Python 中具有特殊意義，因此你不能將它們用作變數名稱。
- 變數名稱不能以數字開頭，且只能包含英數字元或底線。
- 變數名稱區分大小寫，也就是說 `my_var` 與 `my_Var` 和 `MY_VAR` 是不同的。
- 最後，使用 `snake_case` 來撰寫變數名稱是一個常見的慣例，其中每個空格都被底線字元取代，且字詞皆以小寫字母書寫。

移除對 `print()` 的兩個呼叫，並宣告另一個名為 `alphabet` 的變數。將字串 `'abcdefghijklmnopqrstuvwxyz'` 指定給此變數。

# --hints--

你不應該在你的程式碼中有 `print(type(text))`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

你的程式碼中不應該有 `print(type(shift))`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/)
```

你應該宣告一個名為 `alphabet` 的變數。請注意將變數名稱放在行首。

```js
assert.match(code, /^alphabet\s*=/m)
```

你應該將字串 `'abcdefghijklmnopqrstuvwxyz'` 指定到你的 `alphabet` 變數。請記得使用單引號或雙引號將字串括起來。

```js
assert.match(code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*(#.*)?$/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
