---
id: 6551f3bfed58796ebb1268dc
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

每個字串字元都可以用數字索引來取用。索引計數從零開始。所以字串的第一個字元索引是 `0`。例如，在字串 `'Hello World'` 中，`'H'` 的索引是 `0`，`'e'` 的索引是 `1`，依此類推。 

字串的每個字元都可以使用括號表示法來存取。你需要寫出變數名稱，後面接方括號，並在括號中加入字元的索引值：

```py
text = 'Hello World'
r = text[8]
```

現在，不是列印 `text`，而是只列印索引 `6` 的字元。

# --hints--

你仍然應該呼叫 `print()` 函式（程式）。

```js
assert.match(code, /print\s*\(.*\)/)
```

你應該將 `text[6]` 傳遞給 `print()` 函式（程式），方法是將它放在小括弧之間。請注意將函式（程式）呼叫放在該行的開頭。

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
