---
id: 6551f5756c20146fc94f8675
title: 步驟7
challengeType: 20
dashedName: step-7
---

# --description--

您可以使用內置的 `len()` 函數來統計字符串中的字符數（長度）。

修改你現有的 `print()` 通話，通過 `len(text)` 而不是 `text[-1]`

# --hints--


你應該調用 `len()` 函數。

```js
assert.match(code, /len\s*\(.*\)/)
```

你應該通過將變量 `text` 包含在括號之間來將其傳遞給 `len()` 函數。

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

你應該將 `len(text)` 傳遞給 `print()` 函數，方法是將其包含在括號之間。 注意將函數調用放在行首。

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

你的代碼中不應該有 `print(text[-1])`

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

你的代碼中包含無效的句法和/或無效的縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
