---
id: 6551fb531ce7f773071a0b84
title: 步驟 9
challengeType: 20
dashedName: step-9
---

# --description--

正如您可以看到的那樣，打印 `type(text)`的輸出爲 `<class 'str'>`這意味着你的變量是一個字符串(string)，表示爲 `str`。

現在另起一行並創建一個名爲 `shift`的變量，並賦值爲 `3` 。

# --hints--

你應該聲明一個名爲 `shift` 的變量。 請注意，需把變量的名稱放在該行的開頭位置。

```js
assert.match(code, /^shift\s*=/m)
```

你應該將值 `3` 分配給你的 `shift` 變量。

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

您的代碼中包含無效的句法和/或無效的縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
