---
id: 6551fb531ce7f773071a0b84
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

如你所見，列印 `type(text)` 的輸出是 `<class 'str'>`，這表示你的變數是字串，顯示為 `str`。

現在換行並創建另一個名為 `shift` 的變數，並將值 `3` 指定給此變數。

# --hints--

你應該宣告一個名為 `shift` 的變數。請注意將變數名稱放在行首。

```js
assert.match(code, /^shift\s*=/m)
```

你應該將值 `3` 指定到你的 `shift` 變數。

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
