---
id: 6551f6f9ec6db270eb83f3db
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

你可以在終端機上看到列印的 `11`，因為 `'Hello World'` 包含 11 個字元。

另一個有用的內建函式是 `type()`，它會傳回變數的資料型別。修改你的 `print()` 呼叫以列印 `text` 的資料型別。

# --hints--

你應該呼叫 `type()` 函式（程式）。

```js
assert.match(code, /type\s*\(.*\)/)
```

你應該將 `text` 放在小括弧之間，傳遞給 `type()` 函式（程式）。

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

你應該將 `type(text)` 傳遞給 `print()` 函式（程式），方法是將它放在小括弧內。請注意將函式（程式）呼叫放在行首。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

你不應該在你的程式碼中有 `print(len(text))`。

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
