---
id: 6551f1e8737dbe6e0cb08ea4
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

現在列印你的新變數。

# --hints--

你不應該移除你現有的 `print(type(text))` 呼叫。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

你應該呼叫 `print()` 函式，並傳入 `shift` 變數。請注意將函式呼叫放在該行的開頭。

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
