---
id: 6551f22653e6ce6e4f1d1370
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

修改你的 `print(shift)` 呼叫以列印你的 `shift` 變數的資料型別。

# --hints--

你應該保留你現有的 `print(type(text))` 呼叫。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

你應該將 `shift` 傳遞給 `type()` 函式（程式）。

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

你應該將 `type(shift)` 傳遞給你的 `print()` 函式（程式）。請注意將函式（程式）呼叫放在該行的開頭。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

你的程式碼中不應該有 `print(shift)`。

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
