---
id: 6551f5756c20146fc94f8675
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

你可以使用內建的 `len()` 函式（程式）來存取字串中的字元數量。

修改你現有的 `print()` 呼叫，傳入 `len(text)` 取代 `text[-1]`。

# --hints--


你應該呼叫 `len()` 函式（程式）。

```js
assert.match(code, /len\s*\(.*\)/)
```

你應該將變數 `text` 傳遞給 `len()` 函式（程式），方法是將它放在小括弧之間。

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

你應該將 `len(text)` 傳遞給 `print()` 函式（程式），方法是將它放在小括弧之間。請注意將函式（程式）呼叫放在該行的開頭。

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

你不應該在你的程式碼中有 `print(text[-1])`。

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
