---
id: 65522255d5b9cd80f335c6f2
title: 第 29 步
challengeType: 20
dashedName: step-29
---

# --description--

在你的迴圈主體結尾，宣告一個名為 `new_index` 的變數，並將 `index + shift` 的值指定給這個變數。

# --hints--

你應該在你的 `for` 迴圈內宣告一個名為 `new_index` 的變數。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

你應該在 `for` 迴圈主體的結尾將 `index + shift` 指定到你的新變數。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
