---
id: 6552487e689f6e895f658717
title: 步驟 33
challengeType: 20
dashedName: step-33
---

# --description--

現在你需要在你的迴圈主體結尾處創建一個 `new_char` 變數。將它的值設為 `alphabet[new_index]`。

# --hints--

你應該在你的 `for` 迴圈內創建一個 `new_char` 變數。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_char\s*=/));
```

你應該在你的迴圈主體結尾將你的 `new_char` 變數設為 `alphabet[new_index]`。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
--fcc-editable-region--
```
