---
id: 65521203d7165c7b84b22ad4
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

迴圈允許你系統性地遍歷一系列元素並對每個元素執行動作。

在這種情況下，你會使用 `for` 迴圈。以下是你如何迭代 `text`：

```py
for i in text:
```

`for` 是表示迴圈型別的關鍵字。`i` 是一個變數，依序取得 `text` 中元素的值。述語以冒號 `:` 結尾。

在你宣告 `alphabet` 的那一行下面，寫一個 `for` 迴圈來迭代 `text`。使用 `i` 作為迴圈變數。

這樣做的話，終端機中會出現錯誤。你會在下一步中了解它。

# --hints--

你應該使用 `for` 關鍵字來創建迴圈。請確保將 `for` 關鍵字放在行首，並在關鍵字後留一個空白。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

你應該在 `for` 關鍵字之後撰寫 `i` 變數。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

你應該在 `for i ` 之後寫 `in` 關鍵字。請確保在 `in` 關鍵字的前後留有空格。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

你應該在 `for i in ` 之後寫 `text`。別忘了添加最後的 `:`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

你的 `for` 迴圈應該放在程式碼行 `alphabet = 'abcdefghijklmnopqrstuvwxyz'` 的下方。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
