---
id: 6552008c0d9d9075cbec9772
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

你將使用 `.find()` 方法來尋找你訊息中每個字母在字母表中的位置。方法類似於函式，但它屬於一個物件。

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

上述程式碼中，`.find()` 方法是*在* `sentence`（要搜尋的字串）*上呼叫*，並將 `'r'`（要定位的字元）作為引數傳入。`sentence.find('r')` 呼叫會傳回 `4`，這是 `sentence` 中 `'r'` 第一次出現的位置索引。

在你的程式碼結尾，對 `alphabet` 呼叫 `.find()` 並將 `'z'` 作為引數傳遞給該方法。

# --hints--

你應該呼叫 `.find()` 方法。

```js
assert.match(code, /find\s*\(.*\)/)
```

你應該在 `alphabet` 變數上呼叫 `.find()` 方法。

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


你應該在 `alphabet` 變數上呼叫 `.find()` 方法，並將 `'z'` 傳遞給該方法。請注意將方法呼叫放在該行的開頭。

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
