---
id: 655209a4a27dd37873c4cac3
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

`.find()` 傳回字串中符合字元的索引。如果找不到該字元，則傳回 `-1`。如你所見，`text` 中的第一個字元，大寫的 `'H'`，未被找到，因為 `alphabet` 僅包含小寫字母。

你可以使用 `.lower()` 方法將字串轉換成其小寫等價。為你的程式碼添加另一個 `print()` 呼叫來列印 `text.lower()` 並查看輸出。

# --hints--

你仍然應該在你的程式碼中保留 `print(index)`。請注意將函式呼叫放在該行的開頭。

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

你應該列印 `text.lower()`。請注意將函式呼叫放在該行的開頭。

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
