---
id: 65520f8282faf57a0db4f7fe
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

如你從輸出中所見，`alphabet` 字串中 `'h'` 位於索引 `7`。現在你需要找到索引 `7` 加上 `shift` 值的字母。為此，你可以像進行數學加法一樣使用加法運算子 `+`。

修改你的 `shifted` 變數，使其儲存 `alphabet` 在索引 `index + shift` 的值。

# --hints--

你應該將 `alphabet[index + shift]` 指定到你的 `shifted` 變數。

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
