---
id: 655251308f31958d06cdf267
title: 第 36 步
challengeType: 20
dashedName: step-36
---

# --description--

目前，加密的字元會在每次迴圈中更新。將加密的字串儲存在新的變數中會比較好。在你的 `for` 迴圈之前，宣告一個名為 `encrypted_text` 的變數，並為此變數指定空字串（`''`）。

# --hints--

你應該在你的 `for` 迴圈之前宣告一個名為 `encrypted_text` 的變數。

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

你應該將空字串指定給你的 `encrypted_text` 變數。

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
