---
id: 655261b2e1f2c197093f3993
title: 步驟 41
challengeType: 20
dashedName: step-41
---

# --description--

目前，空格會被加密為 `'c'`。若要維護純文字訊息中的原始間距，你需要一個條件式 `if` 述語。這由 `if` 關鍵字、一個條件和冒號 `:` 組成。

```py
if x != 0:
    print(x)
```

在上述範例中，`if` 述語的條件是 `x != 0`。位於 `if` 述語主體內的程式碼 `print(x)` 僅在條件計算結果為 `True` 時執行（在此範例中，表示 `x` 不等於零）。

在你的 for 迴圈頂端，將 `print(char == ' ')` 替換為一個 `if` 述語。這個 `if` 述語的條件應該在 `char` 是空白字元時評估為 `True`，否則評估為 `False`。在 `if` 主體內，列印字串 `'space!'`。記得縮排這一行。

# --hints--

你不應該在你的程式碼中有 `print(char == ' ')`。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

你應該將 `print(char == ' ')` 替換為當 `char == ' '` 時觸發的 `if` 述語。不要使用小括弧來包住 `if` 條件，並且記得包含最後的冒號。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

你應該在你的新 `if` 述語中列印字串 `'space!'`。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
