---
id: 65521fc818947e800bffe48a
title: 步驟 26
challengeType: 20
dashedName: step-26
---

# --description--

在 `for` 迴圈內，在列印目前字元之前，宣告一個名為 `index` 的變數，並將 `alphabet.find(char)` 傳回的值指定給此變數。

# --hints--

你應該在你的 `for` 迴圈開始時宣告一個名為 `index` 的新變數。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

你應該將 `alphabet.find(char)` 指定到你的新 `index` 變數。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
