---
id: 65525e359ca28d938baa82c5
title: 步驟 39
challengeType: 20
dashedName: step-39
---

# --description--

你可以使用加法指定運算子來取得與 `a = a + b` 相同的效果：

```py
a += b
```

加法指派運算子可讓你將一個值加到變數，然後將結果指定給該變數。

使用 `+=` 運算子將值相加並同時指定到 `encrypted_text`。

# --hints--

你應該在你的迴圈主體內使用加法指定運算子，將 `alphabet[new_index]` 加到 `encrypted_text` 的當前值。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
