---
id: 65520e6f2b9678799977f24d
title: 步驟 18
challengeType: 20
dashedName: step-18
---

# --description--

移除最後一個 `print()` 呼叫。然後，不是傳遞 `text[0]`，而是將 `text[0].lower()` 作為引數傳遞給你的 `.find()` 呼叫，並查看輸出。

# --hints--

你不應該在你的程式碼中有 `print(text.lower())`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

你仍然應該列印你的 `index` 變數。

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

你應該更新你的 `alphabet.find(text[0])` 呼叫，將 `text[0].lower()` 作為引數。請注意將方法呼叫放在該行的開頭。

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
