---
id: 655220953ba90d80514d7ee2
title: 步驟 27
challengeType: 20
dashedName: step-27
---

# --description--

目前，`print()` 函式（程式）接受單一引數 `char`，但它可以接受多個以逗號分隔的引數。

為你的 `print(char)` 添加第二個引數，使其列印該字元及其在字母表中的索引。

# --hints--

你應該將 `index` 作為第二個引數添加到你現有的 `print(char)` 呼叫中。別忘了用逗號分隔引數。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
