---
id: 655208d59b131e7816f18c96
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

`print()` 函式只會在主控台給你輸出，但函式和方法可以有你可以在程式碼中使用的返回值。

現在將 `alphabet.find(text[0])` 指定給名為 `index` 的變數。這樣，`index` 就會儲存 `alphabet.find(text[0])` 傳回的值。

# --hints--

你應該宣告一個名為 `index` 的變數。請注意將變數名稱放在行首。

```js
assert.match(code, /^index\s*=/m)
```

你應該將 `alphabet.find(text[0])` 傳回的值指定給你的 `index` 變數。

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
