---
id: 65524d2a1a253b8bb5197ae2
title: 步驟 28
challengeType: 20
dashedName: step-28
---

# --description--

`find` 再次對大寫字元傳回 `-1`，對空白字元也是如此。你稍後會處理空白字元。

目前，請將 `for` 迴圈改為迭代 `text.lower()`，而不是迭代 `text`。

# --hints--

你應該修改你的 `for` 迴圈，以便對 `text.lower()` 進行迭代，而不是對 `text`。

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
