---
id: 6553f4f66099802c6ae94613
title: 步驟 55
challengeType: 20
dashedName: step-55
---

# --description--

目前，你的程式碼引發了 `TypeError`，因為 `caesar` 函式是用兩個參數（`message` 和 `offset`）定義的，因此它預期會以兩個*引數*來呼叫。

在未提供必要引數的情況下呼叫 `caesar()` 會停止程式碼的執行。

給 `message` 和 `offset` 值，方法是將 `text` 和 `shift` 作為引數傳遞給 `caesar` 函式呼叫。

# --hints--

你應該將 `text` 和 `shift` 作為引數傳遞給你的函式呼叫，方法是將它們包含在小括弧內。別忘了用逗號分隔引數。

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
