---
id: 6553a755879b131a445e664c
title: 步驟 44
challengeType: 20
dashedName: step-44
---

# --description--

嘗試將字串 `'Hello Zaira'` 指定給你的 `text` 變數，並查看終端機中發生了什麼。

你會看到一個 `string index out of range` 例外。別擔心，你很快就會知道如何修正它！

# --hints--

你應該有一個 `text` 變數。

```js
assert.match(code, /^text\s*=/m)
```

你應該將字串 `'Hello Zaira'` 指定到你的 `text` 變數。

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
