---
id: 6553a44b1801991847d8cc69
title: 步驟 47
challengeType: 20
dashedName: step-47
---

# --description--

接著，修改你的 `print()` 呼叫以列印 `'encrypted text:', encrypted_text`，並將它放在 `for` 迴圈外，這樣加密後的字串只會列印一次。

# --hints--

你應該在你的 for 迴圈後列印 `'encrypted text:', encrypted_text`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/m) })
```

你的程式碼中不應該有 `print('char:', char, 'encrypted text:', encrypted_text)`。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.notMatch(commentless_code, /print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
