---
id: 6553f6086add4b2cbb99fd78
title: 步驟 56
challengeType: 20
dashedName: step-56
---

# --description--

在你的程式碼底部，在你現有的 `caesar(text, shift)` 呼叫之後，再次呼叫你的函式。

這次，將 `text` 變數和整數 `13` 作為引數傳遞。

# --hints--

你應該再次呼叫你的函式，這次傳入 `text` 和 `13` 作為引數。

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

你應該保留現有的函式呼叫。

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
