---
id: 6553a572f7a65718f1e42e18
title: 步驟 49
challengeType: 20
dashedName: step-49
---

# --description--

函式（程式）本質上是一個可重複使用的程式碼區塊。你已經接觸過一些內建函式（程式），例如 `print()`、`find()` 和 `len()`。但你也可以像這樣定義訂製函式（程式）：

```py
def function_name():
    <code>
```

函式（程式）宣告以 `def` 關鍵字開始，接著是函式名稱 — 一個有效的變數名稱 — 以及一對小括弧。宣告以冒號結尾。

在你的 `shift` 變數後，宣告一個名為 `caesar` 的函式，並將所有後續行縮排，以給你的新函式一個主體。

# --hints--

你應該使用 `def` 關鍵字來宣告一個新的函式（程式）。

```js
assert.match(code, /^def\s+/m)
```

你應該在 `def` 關鍵字後將函式名稱寫成 `caesar`。記得在 `def` 後面加一個空格。

```js
assert.match(code, /^def\s+caesar/m)
```

你應該在函式名稱後面添加一對小括弧。別忘了最後的冒號。

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

你應該將 `shift = 3` 之後的所有行縮排，使它們成為你的新函式（程式）主體。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
