---
id: 6553f8c570f9982e013a8886
title: 步驟 57
challengeType: 20
dashedName: step-57
---

# --description--

目前，每個字母總是依照指定的偏移量用相同的字母加密。如果每個字母的偏移量都不同呢？那將會更難解密。這個演算法稱為 Vigenère cipher，其中每個字母的偏移量是由另一段文字決定，稱為 *key*。

開始將你的凱薩密碼轉換為維吉尼亞密碼，方法是移除這兩個函式呼叫。

# --hints--

你應該移除你的兩個 `caesar()` 函式呼叫。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
