---
id: 6553995f412dd8122ed38e4a
title: 步驟 43
challengeType: 20
dashedName: step-43
---

# --description--

條件述語也可以有一個 `else` 語句。若 `if` 述語的條件為假，則可以將此語句添加到 `if` 述語的結尾以執行替代的程式碼：

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

如你在輸出中所見，當迴圈迭代到空格時，空格會被添加到加密字串中。然後，`if` 區塊外的程式碼執行，並將 `c` 添加到加密字串。

為了修正它，請在 `encrypted_text += char` 之後添加一個 `else` 語句，並縮排所有後續的程式碼行，除了 `print()` 呼叫。

# --hints--

你應該創建一個 `else` 語句。記得包含最後的冒號。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

你應該縮排 `else` 語句後的程式碼行，但 `print()` 呼叫除外。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
