---
id: 6553f6b4ec51112d44d737c8
title: 第 48 步
challengeType: 20
dashedName: step-48
---

# --description--

在 `print` 呼叫之前，再添加一個並將 `'plain text:', text` 作為引數傳遞給 `print()`。使用相同的縮排。

# --hints--

你應該在你的 for 迴圈之後列印 `'plain text:', text`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*^print/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
