---
id: 6553f4f66099802c6ae94613
title: 步驟 55
challengeType: 20
dashedName: step-55
---

# --description--

目前，你的代碼會引發 `TypeError`，因爲 `caesar` 函數是使用兩個參數（`message` 和 `offset`）定義的，因此預計會使用兩個*參數*來調用。

不帶參數調用 `caesar()` 就停止執行代碼。

給 `message` 和 `offset` 提供內容，通過將 `text` 和 `shift` 作爲參數傳遞給 `caesar`函數調用。

# --hints--

你應該將 `text` 和 `shift` 作爲參數傳遞給函數調用，方法是將它們包含在括號內。 別忘了用逗號分隔參數。

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
