---
id: 6553a7d8d05cbb1ae335a665
title: 步驟 45
challengeType: 20
dashedName: step-45
---

# --description--

當迴圈到達字母 `Z` 時，總和 `index + shift` 超過字串 `alphabet` 的最後索引。因此，`alphabet[new_index]` 嘗試使用無效的索引，這會導致 `IndexError` 被發出。

你可以注意到終端機中的輸出在緊接 `Z` 之前的空格處停止，錯誤發出前的最後一次 `print`。

在這種情況下，取模運算子（`%`）可以用來傳回兩個數字相除後的餘數。 例如：`5 % 2` 相等於 `1`，因為 5 除以 2 的商是 2，餘數是 1。

將 `index + shift` 用小括弧括起，並對該表達式取 `26` 的模數，`26` 是 `alphabet` 的長度。

# --hints--

你應該在你的 `else` 述語中有 `new_index = (index + shift) % 26`。

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
