---
id: 6553ed69ece88d29594748aa
title: 步驟 52
challengeType: 20
dashedName: step-52
---

# --description--

要執行，必須在函式名稱後面加上一對小括弧來呼叫（或喚起）該函式，如下所示：

```py
function_name()
```

在你的程式碼結尾，呼叫你的 `caesar` 函式。注意縮排。

# --hints--

你應該呼叫你的 `caesar` 函式（程式）。請確保在該行的開頭撰寫函式（程式）呼叫。

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
