---
id: 6553980e0527fa115c705646
title: 步驟 42
challengeType: 20
dashedName: step-42
---

# --description--

現在，不要列印 `'space!'`，而是使用加法指定運算子，將空格（目前儲存在 `char` 中）加到 `encrypted_text` 的當前值。

# --hints--

你的程式碼中不應該有 `print('space!')`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

你應該在你的新 `if` 述語中使用 `+=` 運算子，將 `char` 添加到 `encrypted_text` 的當前值。

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
