---
id: 6554930320d70414e7b6acc6
title: 步驟 59
challengeType: 20
dashedName: step-59
---

# --description--

刪除你的 `shift` 變數。然後，宣告一個名為 `custom_key` 的新變數，並將字串 `'python'` 的值指定給此變數。

# --hints--

你應該刪除 `shift` 變數及其值。

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

你應該宣告一個名為 `custom_key` 的變數。

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

你應該將字串 `'python'` 指定到你的 `custom_key` 變數。

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
