---
id: 655491bd5b98b813fa5bedca
title: 步驟 58
challengeType: 20
dashedName: step-58
---

# --description--

現在修改你的函式（程式）宣告：將函式名稱改為 `vigenere`，並將第二個參數改為 `key`。

# --hints--

你應該將你的函式名稱修改為 `vigenere`。

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

你的 `vigenere` 函式（程式）應該接受 `message` 和 `key` 作為參數。

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
