---
id: 6554d25dc5ceaa354307a77e
title: 步驟 70
challengeType: 20
dashedName: step-70
---

# --description--

加密和解密是相反的處理程序，而你的函式只需稍作調整即可同時執行兩者。

為你的函式（程式）定義添加第三個名為 `direction` 的參數。並且，將最後兩行程式碼註解掉，以避免主控台出錯。

# --hints--

你應該將程式碼中的最後兩行變成註解。在每行開頭放置一個 `#`。

```js
({ test: () => assert.match(code, /#\s*encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)\s*#\s*print\s*\(\s*encryption\s*\)/) })
```

你應該將 `direction` 添加為你的函式的第三個參數。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
encryption = vigenere(text, custom_key)
print(encryption)
--fcc-editable-region--
```
