---
id: 6554ad2463b8892748f8efdd
title: 步驟 67
challengeType: 20
dashedName: step-67
---

# --description--

目前，你的函式（程式）列印了一些字串，但這些值無法被程式碼的其他零件用來執行任何動作。

為此目的，你需要使用 `return` 述語：

```python
def foo():
    return 'spam'
```

你需要撰寫 `return`，後面接一個空格和函式應該傳回的值。一旦找到 `return` 述語，該值就會被傳回，函式的執行會停止，接著繼續執行函式呼叫後的下一行程式碼。在上述範例中，`foo` 函式傳回字串 `'spam'`。

從你的函式中移除兩個 `print()` 呼叫並傳回 `encrypted_text`。

# --hints--

你應該從你的函式中移除兩個 `print()` 呼叫。

```js
({ 
  test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/print\s*\(\s*("|')plain\stext:\1\s*,\s*message\s*\)/.test(commentless_code))
    assert.isFalse(/print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/.test(commentless_code))
  }
})
```

你的函式（程式）應該傳回 `encrypted_text`。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*/);
    const regex = /return\s+encrypted_text/
    assert(function_body.match(regex));
    assert.notMatch(block_body, regex);
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
