---
id: 65549561463f0016876e852c
title: 步驟 60
challengeType: 20
dashedName: step-60
---

# --description--

由於你的密鑰比你必須加密的文本短，因此你需要重複它才能生成整個加密文本。 在函數體的開頭，聲明一個 `key_index` 變量並將其設置爲零。

# --hints--

你應該在函數體的開頭聲明一個名爲 `key_index` 的變量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
assert(function_body.match(/^\s*key_index\s*=/));
```

你應該將 `0` 分配給你的 `key_index` 變量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");   
assert(function_body.match(/key_index\s*=\s*0/));
```

你的代碼包含無效語法和/或無效縮進

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
