---
id: 6554a49a4f782f208abcc87e
title: 步驟 63
challengeType: 20
dashedName: step-63
---

# --description--

你需要增加下一次迭代的 `key_index` 計數 爲此，請在剛剛添加的行之後的同一代碼塊中，使用加法賦值運算符將 `key_index` 加 1。

# --hints--

你應該使用 `+=` 運算符將 `1` 添加到 `else` 子句主體內的 `key_index`。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");
assert(block_body.match(/key_char\s*=\s*key\s*\[\s*key_index\s*%\s*len\s*\(\s*key\s*\)\s*\]\s*key_index\s*\+=\s*1/));
```

你的代碼包含無效語法和/或無效縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
            key_char = key[key_index % len(key)]
--fcc-editable-region--
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
