---
id: 6554dfce1683be3c0c9609a6
title: 步驟 72
challengeType: 20
dashedName: step-72
---

# --description--

現在你可以取消註解最後兩行，並修改你的函式呼叫，將 `1` 作為第三個引數傳入。

# --hints--

你應該還原你的程式碼最後兩行，並將 `1` 作為第三個引數傳遞給你的函式呼叫。

```js
({ test: () => assert.match(code, /^encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*,\s*1\s*\)\s*print\s*\(\s*encryption\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
--fcc-editable-region--
#encryption = vigenere(text, custom_key)
#print(encryption)
--fcc-editable-region--
```
