---
id: 65549561463f0016876e852c
title: 步驟 60
challengeType: 20
dashedName: step-60
---

# --description--

由於你的金鑰比你必須加密的文字還短，你需要重複它以產生整個加密文字。  
在你的函式（程式）主體開始時，宣告一個 `key_index` 變數並將它設為零。

# --hints--

你應該在你的函式（程式）主體開始處宣告一個名為 `key_index` 的變數。

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
assert(function_body.match(/^\s*key_index\s*=/));
```

你應該將 `0` 指定給你的 `key_index` 變數。

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");   
assert(function_body.match(/key_index\s*=\s*0/));
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
