---
id: 65549f90cf78131c96ebcf28
title: 步驟 61
challengeType: 20
dashedName: step-61
---

# --description--

撰寫程式碼時，可讀性是關鍵。註解作為有效的筆記，說明你程式碼背後的邏輯。當你在一段時間後返回專案時，它們非常有價值，也有助於你的同事理解程式碼。

在 Python 中，你可以使用 `#` 撰寫註解。`#` 之後的任何內容都不會被執行。

在你的 `if` 述語之前，添加一則註解說明 `Append space to the message`。

# --hints--

你應該添加一個註解說明 `Append space to the message`。別忘了開頭的 `#` 字元。

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
