---
id: 65549f90cf78131c96ebcf28
title: 步驟 61
challengeType: 20
dashedName: step-61
---

# --description--

當編碼時，可讀性是關鍵的。 註釋是有效的備註，解釋了代碼背後的邏輯。 當一段時間後返回項目時，它們會非常有用，並且還能幫助你的同事理解代碼。

在 Python 中，你可以使用 `#` 寫註釋。 `#` 之後出現的任何事情都不會被執行。

在 `if` 語句之前，添加一條註釋：`Append space to the message`。

# --hints--

你應該添加一條註釋，說明`Append space to the message`。 不要忘記開始時的 `#`。

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--

        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
