---
id: 6554d0332949b133a0b35eaa
title: 步驟 68
challengeType: 20
dashedName: step-68
---

# --description--

呼叫你的函式（程式），並傳入 `text` 和 `custom_key` 作為引數。將函式（程式）呼叫的返回值儲存在名為 `encryption` 的變數中。

# --hints--

你應該呼叫 `vigenere`，並傳入 `text` 和 `custom_key` 作為引數。

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

你應該擁有一個 `encryption` 變數。

```js
({ test: () => assert(__userGlobals.has("encryption")) })
```

你應該將你的函式呼叫指定給 `encryption` 變數。

```js
({ test: () => assert.match(code, /encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--

--fcc-editable-region--
```
