---
id: 65552b14e803731fe3c1e4ca
title: 步驟 86
challengeType: 20
dashedName: step-86
---

# --description--

接著，修改你的 `encryption` 和 `decryption` 變數，分別透過呼叫 `encrypt` 和 `decrypt`，取代 `vigenere`。

# --hints--

你的 `encryption` 變數應該具有 `encrypt(text, custom_key)` 的值。

```js
({ test: () => assert.match(code, /^encryption\s*=\s*encrypt\s*\(\s*text\s*,\s*custom_key\s*\)/m) })
```

你的 `decryption` 變數應該具有 `decrypt(encryption, custom_key)` 的值。

```js
({ test: () => assert.match(code, /^decryption\s*=\s*decrypt\s*\(\s*encryption\s*,\s*custom_key\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--
def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
