---
id: 6555e43e783ed31a0532b1b2
title: 步驟 93
challengeType: 20
dashedName: step-93
---

# --description--

換行字元 `\n` 是用來表現新行的特殊序列。你可以在字串中將反斜線 `\` 後接 `n` 寫成一般的字元序列，當程式執行時，它會被取代為輸出中的新行。

在你的第一個 `print` 呼叫開頭放入換行字元，然後查看輸出。

# --hints--

你應該將你的第一個列印呼叫修改為 `print(f'\nEncrypted text: {text}')`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*f("|')\\nEncrypted\stext:\s\{\s*text\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print(f'Encrypted text: {text}')
print(f'Key: {custom_key}')
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
