---
id: 6555201d9b7fc917399f9f0b
title: 步驟 80
challengeType: 20
dashedName: step-80
---

# --description--

當對字串呼叫 `.isalpha()` 方法，若該字串的所有字元皆為字母，則會傳回 `True`。例如，以下程式碼會傳回 `True`：

```py
'freeCodeCamp'.isalpha()
# True
```

刪除條件 `char == ' '` 並透過對 `char` 呼叫 `.isalpha()` 方法來取代它。

# --hints--

你應該在你的程式碼中有 `if char.isalpha():`。

```js
({ test: () => assert.match(code, /if\s+char\.isalpha\s*\(\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--    
        # Append space to the message
        if char == ' ':
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
