---
id: 6555e1bfcd374e18c6be8e58
title: 步驟 91
challengeType: 20
dashedName: step-91
---

# --description--

在 Python 中，有一種簡單格式化字串的方法。*f-strings* 讓你能在字串中插入值。 

插值表示撰寫佔位符，當程式執行時會被指定的值取代。 例如，你可以用 `f'Encrypted text: {text}'` 取得與 `'Encrypted text: ' + text` 相同的結果。你需要在引號前加上 `f` 來創建 f-string，並將你想插值的變數或運算式寫在大括號中。

修改第一個 `print()` 呼叫以使用 f-string。

# --hints--

你應該修改 `print('Encrypted text: ' + text)` 呼叫，使其使用 f-string 列印相同的結果。

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
