---
id: 6555e1bfcd374e18c6be8e58
title: 步驟 91
challengeType: 20
dashedName: step-91
---

# --description--

在 Python 中，可以輕鬆格式化字符串。 *f-strings* 使你能夠在字符串中插入值

插值意味着寫入佔位符，在程序運行時將被指定的值所替代。 例如，你可以使用 `f'Encrypted text: {text}'` 獲得 `'Encrypted text: ' + text` 相同的結果。 你需要在引號前放置一個 `f` 來創建 f 字符串，並編寫要在大括號之間插入的變量或表達式。

修改第一個 `print()` 調用以使用 f 字符串。

# --hints--

你應該修改 `print('Encrypted text: ' + text)` 調用，以便它使用 f 字符串打印相同的結果。

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]

    return final_message

def encrypt(message, key):
    return vigenere(message, key)

def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
