---
id: 655520c20cb1e6177b0641d6
title: 步驟 81
challengeType: 20
dashedName: step-81
---

# --description--

`not` 運算子用於否定一個表達式。當它放在一個會被評估為 `True` 的真值之前時，會傳回 `False`，反之亦然。

在 `if` 條件的開頭添加 `not` 運算子，以檢查該字元是否非字母。

# --hints--

你應該在你的 `if` 述語的條件中使用 `not` 運算子。

```js
({ test: () => assert.match(code, /if\s+not\s+char\.isalpha\s*\(\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char.isalpha():
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
