---
id: 655516e410b8e30fb4fb64e8
title: 步驟 76
challengeType: 20
dashedName: step-76
---

# --description--

也請更新你的註解。你的第二個註解應該將 `encode` 改為 `encode/decode`。你的第三個註解應該將 `encrypted` 改為 `encrypted/decrypted`。

# --hints--

你的第二個評論應該是 `Find the right key character to encode/decode`。

```js
({ test: () => assert.match(code, /#\s*Find\sthe\sright\skey\scharacter\sto\sencode\/decode/) })
```

你的第三個註解應該是 `Define the offset and the encrypted/decrypted letter`。

```js
({ test: () => assert.match(code, /#\s*Define\sthe\soffset\sand\sthe\sencrypted\/decrypted\sletter/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
