---
id: 6555d6e12c95701172b55709
title: 步驟 77
challengeType: 20
dashedName: step-77
---

# --description--

函式（程式）可以使用預設引數來呼叫。預設引數顯示函式（程式）在未傳入該引數時應採用的值。 例如，下面的函式（程式）接受三個引數，但你可以傳入兩個引數來呼叫它。第三個引數將採用預設指定的值：

```py
def foo(a, b, c=value):
    <code>
```

修改 `vigenere` 函式（程式），使其 `direction` 參數具有 `1` 的預設值。

# --hints--

你的 `vigenere` 函式的 `direction` 參數應該有一個預設值為 `1`。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction=1)'
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

--fcc-editable-region--
def vigenere(message, key, direction):
--fcc-editable-region--
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
