---
id: 6555231eac4d9f19bd3d44b5
title: 步驟 83
challengeType: 20
dashedName: step-83
---

# --description--

`pass` 關鍵字可以用作未來程式碼的佔位符。它在你的程式碼中沒有任何效果，但可以避免你因程式碼不完整而產生的錯誤：

```py
def foo():
    pass
```

使用 `1` 呼叫 `vigenere` 來加密，使用 `-1` 來解密是可以的，但可能有點難以理解。創建一個名為 `encrypt` 的新函式（程式），該函式（程式）接受 `message` 和 `key` 參數，並使用 `pass` 來填充函式（程式）主體。


# --hints--

你應該宣告一個名為 `encrypt` 的新函式。記得在函式主體內使用 `pass` 關鍵字。

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

你的 `encrypt` 函式（程式）應該以 `message` 和 `key` 作為參數。順序很重要。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
