---
id: 655619327c7b364166f8dd6f
title: 步驟 46
challengeType: 20
dashedName: step-46
---

# --description--

如果你希望將額外的字元加入到 `alphabet` 字串中，例如數字或特殊字元，你會發現必須手動修改模除操作的右側運算元。

將 `26` 取代為 `len(alphabet)` 以避免此問題。

# --hints--

你應該修改 `new_index` 的值，將 `26` 替換為 `len(alphabet)`。

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
  assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/));
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % 26
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
