---
id: 655a2a7210094920069b117c
title: 步驟 40
challengeType: 20
dashedName: step-40
---

# --description--

比較運算子允許你根據兩個物件的值進行比較。你可以透過將比較運算子放在你想比較的物件之間來使用它們。它們會傳回一個*布林值*——即 `True` 或 `False`——取決於表達式的真實性。

Python 有以下比較運算子：

<table>
  <thead>
    <tr>
      <th>Operator</th>
      <th>Meaning</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>Equal</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>Not equal</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>Greater than</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>Less than</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>Greater than or equal to</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>Less than or equal to</td>
    </tr>
  </tbody>
</table>

在你的迴圈主體開始時，列印比較 `char` 與空格（`' '`）的結果。為此使用等號運算子 `==`。

# --hints--

你應該在你的 `for` 迴圈中使用等號運算子將 `char` 與空白字元比較。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

你應該在你的 `for` 迴圈中列印比較 `char` 與空格的結果。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

你應該在你的迴圈開始時列印比較 `char` 與空白字元的結果。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
