---
id: 655a2a7210094920069b117c
title: 步驟 40
challengeType: 20
dashedName: step-40
---

# --description--

比較操作符允許你根據兩個對象的值進行比較。 你可以使用比較操作符，在你想要比較的對象之間放置它。 它們返回一個 *布爾值* — 即 `True` 或 `False` — 取決於表達式的真實性。

Python 有以下比較操作符：

<table>
  <thead>
    <tr>
      <th>運算符</th>
      <th>含義</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>等於</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>不等於</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>大於</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>小於</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>大於或等於</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>小於或等於</td>
    </tr>
  </tbody>
</table>

在你的循環實體開始時，打印對 `char` 和空格進行比較的結果(`' '`)。 爲此使用相等運算符 `==`。

# --hints--

你應該在 `for` 循環中使用相等運算符將 `char` 與空格進行比較。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

你應該打印 `char` 與 `for` 循環內的空格進行比較的結果。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

你應該在循環開頭打印 `char` 與空格的比較結果。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
