---
id: 65688b5a1655a7a6caede847
title: 步驟 31
challengeType: 20
dashedName: step-31
---

# --description--

當你嘗試像上一步中那樣更改字符串的各個字符時，你會收到 `TypeError`，當你的代碼中使用了不適當類型的對象時，就會出現這種情況。

從錯誤消息中可以看到，字符串不支持項目分配，因爲它們是不可變的。 然而，一個變量可以重新分配另一個字符串：

```py
message = 'Hello World'
message = 'Hello there!'
```

刪除 `text[0]` 行並重新爲 `text` 分配字符串 `'Albatross'`。

# --hints--

你的代碼中不應包含 `text[0] = 'C'`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

你應該重新分配 `text` 字符串 `'Albatross'`。 不要修改 `text = 'Hello World'`

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
