---
id: 65688b5a1655a7a6caede847
title: 步驟 31
challengeType: 20
dashedName: step-31
---

# --description--

當你嘗試像前一步驟那樣更改字串的個別字元時，你會得到一個 `TypeError`，該錯誤發生在你的程式碼中使用了不適當型別的物件。

如你從錯誤訊息中所見，字串不支援項目指派，因為它們是不可變的。然而，變數可以被重新指定另一個字串：

```py
message = 'Hello World'
message = 'Hello there!'
```

刪除 `text[0]` 行並重新指定 `text` 為字串 `'Albatross'`。

# --hints--

你的程式碼中不應該有 `text[0] = 'C'`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

你應該將 `text` 重新指定為字串 `'Albatross'`。不要修改 `text = 'Hello World'`。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
