---
id: 656886ec4e4d9b96ded3c565
title: 第 25 步
challengeType: 20
dashedName: step-25
---

# --description--

迭代變數可以有任何有效的名稱，但給它一個有意義的名稱會比較方便。

將你的 `i` 變數重新命名為 `char`。

# --hints--

你應該將 `i` 取代為 `char`。 

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
