---
id: 65687d2f8c7ee27b0446cef3
title: 第 24 步
challengeType: 20
dashedName: step-24
---

# --description--

每次迭代要執行的程式碼 — 放在 `:` 之後 — 構成迴圈的主體。此程式碼必須縮排。在 Python 中，建議每個縮排層級使用 4 個空格。此縮排層級即為一個程式碼區塊。

```py
for i in text:
    print(i)
```

Python 依賴縮排來顯示程式碼區塊。行尾的冒號是新縮排程式碼區塊即將出現的信號。

因此，當最後一個冒號後找不到縮排區塊時，程式碼執行會停止並丟出 `IndentationError`。此程式碼不會顯示輸出，反而會引發 `IndentationError`：

```py
for i in text:
print(i)
```

為你的 `for` 迴圈添加一個主體，方法是加入對 `print(i)` 的呼叫。記得縮排迴圈主體。

# --hints--

你應該為你的 `for` 迴圈主體添加 `print(i)`。注意縮排。 

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
