---
id: 65688a50e6c998a21d8e41d3
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

字串是不可變的，這表示一旦創建後就無法更改。 例如，你可能會認為以下程式碼會將 `my_string` 的值更改為字串 `'train'`，但這是不正確的：

```py
my_string = 'brain'
my_string[0] = 't'
```

確認使用括號表示法存取 `text` 中的第一個字元並嘗試將其更改為你選擇的字元。你會看到輸出消失並出現錯誤。

# --hints--

你應該仍然在你的程式碼中有 `text = 'Hello World'`。

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

你應該使用 `text[0]` 存取 `text` 中的第一個字母。

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

你應該使用 `=` 運算子將你選擇的字元指定給 `text[0]`。別忘了將字元用單引號或雙引號括起來。

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
