---
id: 6568b7a40c7ba0ccbc4b4425
title: 步驟 32
challengeType: 20
dashedName: step-32
---

# --description--

如你所見，被重新指定的字串中的每個字元都會在迴圈內被列印。

透過刪除 `text` 指派回到原始字串。

# --hints--

你不應該在你的程式碼中有 `text = 'Albatross'`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*=\s*("|')Albatross\1/.test(commentless_code))
```

你應該仍然有一個值為 `'Hello World'` 的 `text` 變數。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /text\s*=\s*("|')Hello World\1/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
