---
id: 659530ca823d6c47f61daa8a
title: 步驟 20
challengeType: 20
dashedName: step-20
---

# --description--

打印你的 `shifted` 變量。

# --hints--

你應該打印你的 `shifted` 變量。 注意將函數調用放在行首。

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

你的代碼包含無效語法和/或無效縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
