---
id: 659530ca823d6c47f61daa8a
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

列印你的 `shifted` 變數。

# --hints--

你應該列印你的 `shifted` 變數。請注意將函式呼叫放在該行的開頭。

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
