---
id: 65951b637ddec63611c482c0
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

你可以使用內建函式（程式）`print()`在終端機上列印你的程式碼輸出。

函式（程式）是可重複使用的程式碼區塊，你可以在需要時呼叫或喚起它們來執行其程式碼。要呼叫函式，只需在其名稱旁邊寫上一對小括弧。你很快就會學到更多關於函式的內容。

現在，換行並為 `print()` 函式添加一個空的呼叫。你應該還看不到任何輸出。

# --hints--

你的程式碼中應該有 `print()`。請注意將函式呼叫放在行首。

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
