---
id: 65952b70924edb449faf1950
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

宣告一個名為 `shifted` 的新變數。使用括號表示法存取 `alphabet` 在索引 `index` 的值，並將其指定給你的新變數。

# --hints--

你應該宣告一個名為 `shifted` 的變數。請注意將變數名稱放在行首。

```js
assert.match(code, /^shifted\s*=/m)
```

你應該將 `alphabet[index]` 指定到你的 `shifted` 變數。

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
