---
id: 659526f3e5424b3eac4b6aa7
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

接著，將 `index` 變數列印到主控台。

# --hints--

你應該列印你的 `index` 變數。請注意將函式呼叫放在該行的開頭。

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
