---
id: 65951b637ddec63611c482c0
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

你可以使用內置函數 `print()` 在終端上打印代碼的輸出。

函數是可重用的代碼塊，你可以在需要時調用或調用它們來運行其代碼。 要調用函數，只需在其名稱旁邊寫一對括號即可。 你將很快了解更多關於函數的信息。

現在轉到一個新行，並在 `print()` 函數中添加一個空調用。 你不應該看到任何輸出。

# --hints--

你的代碼中應該有 `print()` 注意將函數調用放在行首。

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

您的代碼中包含無效的句法和/或無效的縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
