---
id: 659db3e3670d3e712be82593
title: 第 22 步
challengeType: 20
dashedName: step-22
---

# --description--

重複尋找字母在字母表中的位置並為 `text` 中的每個字元決定位移後的字母可能會很繁瑣。幸運的是，你可以使用迴圈來簡化這個過程。

目前，移除 `alphabet` 變數宣告下方的所有程式碼行。

# --hints--

你應該仍然在你的程式碼中有 `text = 'Hello World'`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

你在程式碼中仍然應該有 `shift = 3`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

你仍然應該在你的程式碼中保留 `alphabet = 'abcdefghijklmnopqrstuvwxyz'`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

你應該刪除 `index` 變數及其值。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

你的程式碼中不應該有 `print(index)`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

你應該刪除 `shifted` 變數及其值。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

你的程式碼中不應該有 `print(shifted)`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
