---
id: 65a450e8fb2c9d75c7378d28
title: 步驟 88
challengeType: 20
dashedName: step-88
---

# --description--

由於這次你是從加密字串開始解密，因此你不再需要 `encryption` 變數。

刪除 `encryption` 和 `print(encryption)` 呼叫。並且，將你的程式碼最後兩行註解起來。

# --hints--

你應該刪除 `encryption` 變數及其值。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse(/encryption\s*=\s*encrypt\s*\(\s*text\s*,\s*custom_key\s*\)/.test(commentless_code))
```

你的程式碼中不應該有 `print(encryption)`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse(/print\s*\(\s*encryption\s*\)/.test(commentless_code))
```

你應該將 `decryption = decrypt(encryption, custom_key)` 變成註解。

```js
assert(code.match(/#\s*decryption\s*=/));
```

你應該將 `print(decryption)` 變成註解。

```js
assert(code.match(/#\s*print\s*\(\s*decryption\s*\)/));
```

你的程式碼包含無效的語法和／或無效的縮排。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
encryption = encrypt(text, custom_key)
print(encryption)
decryption = decrypt(encryption, custom_key)
print(decryption)
--fcc-editable-region--
```
