---
id: 6624e0ec0bf47017eed84b1f
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

你將要建置的第一種密碼稱為*Caesar*密碼。具體來說，你會取出訊息中的每個字母，找到它在字母表中的位置，取字母表中位於該位置後 3 個位置的字母，並用新字母取代原本的字母。

要實作這個，你將使用前一步驟中討論的 `.find()` 方法。修改你現有的 `.find()` 呼叫，將 `text[0]` 作為引數傳入，而不是 `'z'`。

# --hints--

你應該修改你現有的 `alphabet.find('z')` 呼叫，將 `text[0]` 傳遞給該方法。

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
