---
id: 6624e0ec0bf47017eed84b1f
title: 步驟14
challengeType: 20
dashedName: step-14
---

# --description--

你要構建的第一種密碼稱爲*Caesar*密碼  具體來說，你將取出消息中的每個字母，找到其在字母表中的位置，取出字母表中 3 個位置之後的字母，然後用新字母替換原始字母。

要實現此目的，你將使用上一步中討論的 `.find()` 方法 修改現有的 `.find()` 調用，將其傳遞爲 `text[0]` 作爲參數，而不是 `'z'`。

# --hints--

你應該修改現有的 `alphabet.find('z')` 調用，將 `text[0]` 傳遞給該方法。

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
