---
id: 66605a64a9822f01ec5ae8ce
title: 步驟 50
challengeType: 20
dashedName: step-50
---

# --description--

在 Python 中，變數的*作用域*決定該變數可以被存取的位置：

- 在函式外定義的變數具有*全域的*範圍，並且它們可以從你的程式碼的任何零件中存取。

- 在函式（程式）內定義的變數是該函式（程式）的*區域的*，且無法在函式（程式）外被存取。

要看到這個效果，請嘗試在你的程式碼結尾列印 `alphabet` 變數。這將引發 `NameError` 例外。

你應該會看到一則錯誤訊息，指出 `alphabet` 未定義。這是因為 `alphabet` 定義在 `caesar` 函式（程式）內部，且無法在其外部存取。

# --hints--

你應該嘗試在 `caesar` 函式外列印 `alphabet` 變數。

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
