---
id: 61695197ac34f0407e339882
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

正如你在之前的项目中看到的，网页应该以 `DOCTYPE html` 声明开头，然后是 `html` 元素。

在文档顶部添加 `DOCTYPE html` 声明，然后添加 `html` 元素。 给 `html` 元素一个值为 `en` 的 `lang` 属性。

# --hints--

你的代码应该包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你应在 `DOCTYPE` 后添加一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你应将文档类型定义为 `html` 。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你应该在 `DOCTYPE` 声明的 `type` 后用 `>` 关闭。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素应该有一个值为 `en` 的 `lang` 属性的开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/gi));
```

你的 `DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
