---
id: 616965351e74d4689eb6de30
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

你可以在一个网页上有多个 `meta` 元素。 每个 `meta` 元素都会添加无法通过其他 HTML 元素来表达的页面信息。

在 `head` 中添加另一个 `meta` 元素。 给它一个 `name` 属性，值为 `"viewport"`，和一个 `content` 属性，值为 `"width=device-width, initial-scale=1.0"`，这样你的页面在各种设备上看起来是一样的。

# --hints--

你应该有两个 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

你新建的 `meta` 元素应为自闭合元素，不应有闭合标签 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```


你的新 `meta` 元素应将 `name` 属性设置为 `"viewport"`，并将 `content` 属性设置为 `"width=device-width, initial-scale=1.0"`。

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

你的两个 `meta` 元素应该放在 `head` 元素中。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
