---
id: 683ecb010f4c6bafd428802c
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви перевірите множину медичних даних, щоб упевнитися, що вони відповідають набору правил.

Оголосіть змінну з іменем `medical_records` і призначте їй порожній список. У наступних кроках ви будете використовувати її для збереження своїх медичних даних.

# --hints--

Ви повинні оголосити змінну з іменем `medical_records`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

Ви повинні призначити порожній список змінній `medical_records`.

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
