---
id: 6842d659a4ae7973dabab3b0
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Завершіть словник, додавши такі три пари ключ-значення:

- Ключ `diagnosis` зі значенням `Hypertension`
- Ключ `medications` зі значенням списку `['Lisinopril']`
- Ключ `last_visit_id` зі значенням `V2301`.

# --hints--

У вашому словнику має бути ключ `diagnosis` зі значенням рядка `Hypertension`.

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

У вашому словнику має бути ключ `medications` зі значенням списку `['Lisinopril']`.

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

У вашому словнику має бути ключ `last_visit_id` зі значенням рядка `V2301`.

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

У вашому словнику має бути ключ `patient_id` зі значенням рядка `P1001`.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

У вашому словнику має бути ключ `age` зі значенням цілого числа `34`.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

У вашому словнику має бути ключ `gender` зі значенням рядка `Female`.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
