---
id: 6842aee5bcaf4ab8883f8191
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Як ви дізналися в попередніх уроках, словник — це структура даних, яка містить пари ключ-значення, де ключі є незмінними типами даних (наприклад, рядками) і унікальні в межах словника:

```py
person = {
    'name': 'John',
    'age': 33
}
```

Список `medical_records` зберігатиме словники, кожен з яких представляє пацієнта. Додайте словник із ключем `patient_id` і значенням рядка `P1001` до списку `medical_records`.

# --hints--

Ваш список `medical_records` повинен містити один словник.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

У вашому словнику має бути ключ `patient_id` зі значенням рядка `P1001`.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
