---
id: 6842d5b537c64266420faf62
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Додайте ключ `age` зі значенням цілого числа `34` і ключ `gender` зі значенням рядка `Female` до свого словника. Не забудьте кому між парами ключ-значення.

# --hints--

Ваш список `medical_records` має містити один словник.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

У вашому словнику має бути ключ `age` зі значенням цілого числа `34`.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

У вашому словнику має бути ключ `gender` зі значенням рядка `Female`.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

У вашому словнику має бути ключ `patient_id` зі значенням рядка `P1001`.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
