---
id: 6842ba0086201089ea268c92
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Використовуючи ту саму структуру, що й у попередніх кроках, список `medical_records` вже заповнений даними інших пацієнтів. Можете ознайомитися з ним.

Далі ви почнете писати функцію для перевірки множини даних. Створіть функцію з назвою `validate` з одним параметром `data`.

Ви хочете переконатися, що ваші дані є або списком, або кортежем. Тому всередині функції `validate` оголосіть змінну з назвою `is_sequence` і призначте їй виклик `isinstance`. Передайте `data` як перший аргумент і кортеж, що містить `list` та `tuple`, як другий аргумент.

# --hints--

Ви повинні мати функцію з назвою `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

Ваша функція `validate` має мати один параметр `data`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

У вашій функції `is_sequence` має бути змінна з назвою `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

Ви повинні призначити `isinstance(data, (list, tuple))` змінній `is_sequence`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--
```
