---
id: 61695197ac34f0407e339882
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Kama ulivyoona katika miradi iliyopita, kurasa za mtandao zinapaswa kuanza na tamko la `DOCTYPE html`, ikifuatiwa na kipengele cha `html`.

Ongeza tamko la `DOCTYPE html` juu kabisa ya hati, na kipengele cha `html` baada yake. Toa kipengele cha `html` sifa ya kipengele ya `lang` yenye thamani ya `en`.

# --hints--

Msimbo wako unapaswa kuwa na rejea la `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Unapaswa kuweka nafasi baada ya rejea la `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Unapaswa kufafanua aina ya hati kuwa `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Unapaswa kufunga tamko la `DOCTYPE` kwa lebo ya kufunga `>` baada ya aina.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufungua yenye sifa ya kipengele `lang` ya `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufunga.

```js
assert(code.match(/<\/html\s*>/gi));
```

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
