---
id: 61696ef7ac756c829f9e4048
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Katika mradi huu utatumia faili ya CSS ya nje kupamba ukurasa. Tayari tumekuandalia faili `styles.css`. Lakini kabla hujatumia, unahitaji kuunganisha na ukurasa.

Panga ndani kipengele cha `link` ndani ya kipengele cha `head`. Kipekee sifa ya kipengele `rel` iwe na thamani `"stylesheet"` na sifa ya kipengele `href` iwe na thamani `"styles.css"`.

# --hints--

Unapaswa kuwa na kipengele kimoja cha `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Kipengele chako cha `link` kinapaswa kuwa kipengele kisicho na lebo ya kufunga `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele `rel` yenye thamani `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele `href` yenye thamani `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele cha `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
