---
id: 616965351e74d4689eb6de30
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Unaweza kuwa na vipengele vingi vya `meta` kwenye ukurasa wa mtandao. Kila kipengele cha `meta` huongeza taarifa kuhusu ukurasa ambayo haiwezi kudhihirishwa na vipengele vingine vya HTML.

Ongeza kipengele kingine cha `meta` ndani ya `head`. Kipe sifa ya `name` iliyowekwa kuwa `"viewport"` na sifa ya `content` iliyowekwa kuwa `"width=device-width, initial-scale=1.0"` ili ukurasa wako uonekane sawa kwenye vifaa vyote.

# --hints--

Unapaswa kuwa na vipengele viwili vya `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

Kipengele chako kipya cha `meta` kinapaswa kuwa kipengele cha void, hakipaswi kuwa na lebo ya kufunga `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```


Kipengele chako kipya cha `meta` kinapaswa kuwa na sifa ya `name` iliyowekwa kuwa `"viewport"`, na sifa ya `content` iliyowekwa kuwa `"width=device-width, initial-scale=1.0"`.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

Vipengele vyako viwili vya `meta` vinapaswa kuwa ndani ya kipengele cha `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
