---
id: 61695d1fbc003856628bf561
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Sifa ya kipengele `charset` inaelezea usimbaji wa herufi unaotumika na hati. `utf-8` (Unicode Transformation Format – 8-bit) ni kiwango cha usimbaji wa herufi kinachotumika kwa mawasiliano ya kielektroniki.

Ndani ya kipengele `head`, panga ndani kipengele `meta` chenye sifa ya kipengele `charset` iliyowekwa kuwa `"utf-8"`.

# --hints--

Unapaswa kuwa na kipengele kimoja cha `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

Kipengele chako cha `meta` kinapaswa kuwa kipengele kisicho na maudhui, hakipaswi kuwa na lebo ya kufunga `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Kipengele chako cha `meta` kinapaswa kuwa na sifa ya kipengele `charset` iliyowekwa kwa thamani `"utf-8"`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Kipengele chako cha `meta` kinapaswa kuwa ndani ya kipengele cha `head`.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
