---
id: 61695ab9f6ffe951c16d03dd
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Panga ndani kipengele cha `head` ndani ya kipengele cha `html`. Mara tu baada ya kipengele cha `head`, ongeza kipengele cha `body`.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `head`.

```js
assert(code.match(/<head\s*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Unapaswa kuwa na lebo ya kufungua ya `body`.

```js
assert(code.match(/<body\s*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Inapaswa kuwepo lebo moja tu ya kufungua ya `head`.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

Inapaswa kuwepo lebo moja tu ya kufunga ya `head`.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

Inapaswa kuwepo lebo moja tu ya kufungua ya `body`.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

Inapaswa kuwepo lebo moja tu ya kufunga ya `body`.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

Kipengele chako cha `head` kinapaswa kuwa tupu.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

Kipengele chako cha `body` kinapaswa kuwa tupu.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

Kipengele chako cha `body` kinapaswa kuwekwa baada ya kipengele chako cha `head`.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

Kipengele chako cha `head` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Kipengele chako cha `body` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
