---
id: 61695c4aad56f95497c19583
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Kumbuka kuwa kipengele cha `title` kinatoa taarifa za ziada kwa injini za utafutaji kuhusu ukurasa. Pia kinaonyesha maudhui ya kipengele cha `title` kwa njia mbili zaidi:

* kwenye upau wa kichwa wakati ukurasa umefunguliwa  
* kwenye kichupo cha kivinjari kwa ukurasa wakati unabeba kipanya juu yake. Hata kama kichupo hicho hakifanyi kazi, mara unapoleta kipanya juu ya kichupo, maandishi ya `title` yanaonyeshwa.

Ndani ya kipengele cha `head`, panga ndani kipengele cha `title` chenye maandishi ya `Colored Markers`.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `title`.

```js
assert(code.match(/<title\s*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Mradi wako unapaswa kuwa na kichwa cha `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Kumbuka, herufi kubwa na ndogo pamoja na tahajia ni muhimu kwa maandishi ya kichwa.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
