---
id: 616d55bd160a95e22453a081
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Sasa umeweka alama moja katikati na rangi, ni wakati wa kuongeza alama zingine.

Katika `container` `div`, ongeza vipengele viwili vya `div` na uwape kila moja darasa la `marker`.

# --hints--

Kipengele chako kipya cha kwanza cha `div` kinapaswa kuwa na lebo ya kufungua.

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

Kipengele chako kipya cha kwanza cha `div` kinapaswa kuwa na lebo ya kufunga.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Kipengele chako kipya cha pili cha `div` kinapaswa kuwa na lebo ya kufungua.

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

Kipengele chako kipya cha pili cha `div` kinapaswa kuwa na lebo ya kufunga.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Vipengele vyako vipya vya `div` vinapaswa kuwa ndani ya `div` yenye darasa la `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

Vipengele vyako vipya vya `div` vyote vinapaswa kuwa na darasa la `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
