---
id: 616d3a67ccf800ad94ec89ae
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Sasa utaongeza vipengele ambavyo hatimaye utavipangilia rangi kama alama.

Kwanza, ndani ya kipengele cha `body`, ongeza kipengele cha `div` na weka sifa ya `class` kuwa `container`. Hakikisha kipengele cha `div` kiko chini ya kipengele cha `h1`.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `div`.

```js
assert(code.match(/<div.*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `div`.

```js
assert(code.match(/<\/div\s*>/i));
```

Kipengele chako cha `div` kinapaswa kuwa ndani ya kipengele chako cha `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Kipengele chako cha `div` kinapaswa kuwa baada ya kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Kipengele chako cha `div` kinapaswa kuwa na sifa ya `class` iliyowekwa kuwa `container`.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
