---
id: 61764c602bee6974e7790f35
title: Hatua 16
challengeType: 0
dashedName: step-16
---

# --description--

Ili kutoa rangi tofauti kwa alama, utahitaji kuongeza darasa la vitu la kipekee kwa kila moja. Madarasa mengi yanaweza kuongezwa kwa kipengele kwa kuyaorodhesha katika sifa ya `class` na kuyatenganisha kwa nafasi. Kwa mfano, yafuatayo yanaongeza madarasa ya `animal` na `dog` kwa kipengele cha `div`.

```html
<div class="animal dog">
```

Ukiongeza madarasa mengi kwa kipengele cha HTML, mitindo ya madarasa ya kwanza unayoorodhesha inaweza kufutwa na madarasa ya baadaye.

Kuanza, ongeza darasa la `one` kwa kipengele cha alama ya kwanza `div`.


# --hints--

Unapaswa kuongeza darasa la `one` kwa kipengele cha alama ya kwanza `div`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

Alama yako ya kwanza `div` inapaswa kuwa na madarasa ya `marker` na `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
