---
id: 61764fdda535587e1fefb3aa
title: Hatua 19
challengeType: 0
dashedName: step-19
---

# --description--

Ongeza darasa la vitu `two` kwa alama ya pili `div`, na ongeza darasa la vitu `three` kwa alama ya tatu `div`.

# --hints--

Unapaswa kuongeza darasa la vitu `two` kwa kipengele cha alama ya pili `div` katika `container` `div`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(containerSecondChild?.classList?.contains('two'));
```

Alama yako ya pili `div` inapaswa kuwa na madarasa ya vitu `marker` na `two`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(containerSecondChild?.classList?.contains('marker') && containerSecondChild.classList?.contains('two'));
```

Unapaswa kuongeza darasa la vitu `three` kwa kipengele cha alama ya tatu `div` katika `container` `div`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(containerThirdChild?.classList?.contains('three'));
```

Alama yako ya tatu `div` inapaswa kuwa na madarasa ya vitu `marker` na `three`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(containerThirdChild?.classList?.contains('marker') && containerThirdChild?.classList?.contains('three'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

```
