---
id: 617bb5624a75e86463b7e638
title: Hatua 28
challengeType: 0
dashedName: step-28
---

# --description--

<dfn>Rangi za sekondari</dfn> ni rangi unazopata unapochanganya rangi za msingi. Huenda umeona baadhi ya rangi za sekondari katika hatua iliyopita ulipobadilisha thamani za nyekundu, kijani, na buluu.

Ili kuunda rangi ya sekondari ya kwanza, ya manjano, sasisha kitendakazi cha `rgb` katika sheria ya `.one` ya CSS ili kuchanganya nyekundu safi na kijani safi.

# --hints--

Sheria yako ya `.one` ya CSS inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(255, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
