---
id: 617b8b38f32bf2080a140675
title: Hatua 22
challengeType: 0
dashedName: step-22
---

# --description--

Kitendakazi ni kipande cha msimbo kinachoweza kupokea ingizo na kutekeleza tendo maalum. Kitendakazi cha CSS `rgb` kinakubali thamani, au <dfn>hoja</dfn>, za nyekundu, kijani, na buluu, na hutengeneza rangi:

```css
rgb(red, green, blue);
```

Kila thamani ya nyekundu, kijani, na buluu ni nambari kutoka `0` hadi `255`. `0` ina maana kwamba hakuna rangi hiyo kabisa (0%), na ni nyeusi. `255` ina maana kwamba kuna asilimia 100 ya rangi hiyo.

Katika sheria ya CSS ya `.one`, badilisha neno msingi la rangi `red` kwa kitendakazi cha `rgb`. Kwa kitendakazi cha `rgb`, weka thamani ya nyekundu kuwa `255`, thamani ya kijani kuwa `0`, na thamani ya buluu kuwa `0`.

# --hints--

Sheria yako ya CSS ya `.one` haipaswi kutumia neno msingi la rangi la `red` kuweka sifa ya `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

Sheria yako ya CSS ya `.one` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
