---
id: 618a132676346ac9f7fd59dd
title: Hatua 39
challengeType: 0
dashedName: step-39
---

# --description--

Tazama jinsi macho yako yanavyovutiwa kwa asili na rangi nyekundu katikati? Unapobuni tovuti, unaweza kutumia athari hii kuvutia umakini kwa vichwa vya sehemu muhimu, vifungo, au viungo.

Kuna mchanganyiko mwingine muhimu wa rangi nje ya rangi za nyongeza, lakini utajifunza hayo baadaye kidogo.

Kwa sasa, tumia kitendakazi cha `rgb` katika sheria ya `.two` ya CSS kuweka `background-color` kuwa nyeusi.

# --hints--

Sheria yako ya `.two` ya CSS inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
