---
id: 61a493ead935c148d2b76312
title: Hatua 50
challengeType: 0
dashedName: step-50
---

# --description--

Kitendakazi cha `linear-gradient` ni rahisi kubadilika sana -- hapa ni sintaksia ya msingi utakayotumia katika mafunzo haya:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` ni mwelekeo wa mstari unaotumika kwa mabadiliko. `color1` na `color2` ni hoja za rangi, ambazo ni rangi zitakazotumika katika mabadiliko yenyewe. Hizi zinaweza kuwa aina yoyote ya rangi, ikiwa ni pamoja na maneno muhimu ya rangi, hex, `rgb`, au `hsl`.

Sasa utaweka mabadiliko ya rangi kutoka nyekundu hadi kijani kwenye mstari wa digrii 90 kwa alama ya kwanza.

Kwanza, katika sheria ya `.red` CSS, weka sifa ya `background` kuwa `linear-gradient()`, na upitishe thamani ya `90deg` kama `gradientDirection`.

# --hints--

Sheria yako ya `.red` CSS inapaswa kuwa na sifa ya `background` yenye thamani `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
