---
id: 68e6698b29e562d21b1d260c
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você vai construir uma lista dos principais navegadores web. O boilerplate HTML foi fornecido para você.

Comece adicionando um título à sua página que diga `List of Major Web Browsers` usando um elemento `h1` dentro do elemento `body`.

# --hints--

Você deve ter um elemento `h1` dentro do elemento `body`.

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

Seu elemento `h1` deve conter o texto `List of Major Web Browsers`.

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
            
    --fcc-editable-region--
    </body>
</html>
```
