---
id: 68ecfa1003d88990bb96a33d
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Listas de descrição são compostas por uma lista de termos e detalhes. O elemento `dt` representa o termo da descrição, enquanto o elemento `dd` representa o detalhe da descrição. 

O primeiro navegador que você adicionará à sua lista de descrição será para o Google Chrome.

Dentro do seu elemento `dl`, adicione um elemento `dt` contendo o texto `Google Chrome`.

Abaixo do seu elemento `dt`, adicione um elemento `dd` com o seguinte texto:

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

Você deve ter um elemento `dt` dentro do seu elemento `dl`.

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

Seu elemento `dt` deve conter o texto `Google Chrome`.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

Você deve ter um elemento `dd` abaixo do seu elemento `dt`.

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

Seu elemento `dd` deve conter o texto `This is a free web browser developed by Google and first released in 2008.`.

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```

