---
id: 68ecfa1003d88990bb96a33e
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Agora é hora de adicionar seu segundo navegador à lista.

Abaixo da sua entrada `Google Chrome`, adicione outro elemento `dt` contendo o texto `Firefox`.

Abaixo do seu elemento `dt`, adicione um elemento `dd` com o seguinte texto:

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

Você deve ter dois elementos `dt` no seu código.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

Seu segundo elemento `dt` deve ter o texto `Firefox`.

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

Você deve ter dois elementos `dd` no seu código.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

Seu segundo elemento `dd` deve conter o texto `This is a free web browser developed by the Mozilla Corporation and first created in 2004.`.

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```

