---
id: 68ecfa0f03d88990bb96a33c
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Como você se lembra de uma lição anterior, listas de descrição são usadas para apresentar termos e definições em um formato organizado e fácil de ler. 

Aqui está um exemplo:

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

Abaixo do elemento `h1`, crie um elemento `dl`. Este conterá a lista de navegadores.

# --hints--

Você deve ter um elemento `dl`.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

Seu elemento `dl` deve estar abaixo do elemento `h1`.

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        --fcc-editable-region--
        <h1>List of Major Web Browsers</h1> 

        --fcc-editable-region--
    </body> 
</html>
```

