---
id: 69038dd51a41ae215fcc030a
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

O terceiro navegador que você adicionará à lista será para o navegador web Safari.

Adicione outro elemento `dt` contendo o texto `Safari`.

Abaixo do seu elemento `dt`, adicione um elemento `dd` com o seguinte texto:

```md
This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.
```

# --hints--

Você deve ter três elementos `dt` no seu código.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 3);
```

Seu terceiro elemento `dt` deve ter o texto `Safari`.

```js
const dt = document.querySelectorAll('dt')[2];
assert.strictEqual(dt?.textContent.trim(), "Safari");
```

Você deve ter três elementos `dd` no seu código.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 3);
```

Seu terceiro elemento `dd` deve conter o texto `This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.`.

```js
const ddElement = document.querySelectorAll('dd')[2];
assert.strictEqual(ddElement?.textContent.trim(), "This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
