---
id: 69038f5efad6f722def4a214
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Agora é hora de adicionar seu quarto navegador à lista.

Adicione outro elemento `dt` contendo o texto `Brave`.

Abaixo do seu elemento `dt`, adicione um elemento `dd` com o seguinte texto:

```md
This is a free web browser first released in 2016 that is based on the Chromium web browser.
```

# --hints--

Você deve ter quatro elementos `dt` no seu código.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 4);
```

Seu quarto elemento `dt` deve ter o texto `Brave`.

```js
const dt = document.querySelectorAll('dt')[3];
assert.strictEqual(dt?.textContent.trim(), "Brave");
```

Você deve ter quatro elementos `dd` no seu código.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 4);
```

Seu quarto elemento `dd` deve conter o texto `This is a free web browser first released in 2016 that is based on the Chromium web browser.`.

```js
const ddElement = document.querySelectorAll('dd')[3];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser first released in 2016 that is based on the Chromium web browser.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            --fcc-editable-region--

            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
