---
id: ac6993d51946422351508a41
title: 實現一個字符串截斷算法
challengeType: 26
dashedName: truncate-a-string
---

# --description--

在本實驗中，你將練習將字符串截斷到一定長度。

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶故事：**

1. 你應該有一個接受兩個參數的函數 `truncateString`，第一個參數是字符串，第二個參數是數字。
1. 如果字符串的長度大於給定的數字，則應截斷字符串以減少長度，使其等於給定的數字，並在截斷字符串的末尾添加 `...`。
1. 如果字符串長度等於或小於給定的數字，則返回的字符串應保持不變。

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)` 應該返回字符串 `A-tisket...`。

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)` 應該返回字符串 `Peter Piper...`。

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)` 應該返回字符串 `A-tisket a-tasket A green and yellow basket`。

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)` 應該返回字符串 `A-tisket a-tasket A green and yellow basket`。

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)` 應該返回字符串 `A...`。

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)` 應該返回字符串 `Ab...`。

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
