---
id: ac6993d51946422351508a41
title: 實作截斷字串演算法
challengeType: 26
dashedName: truncate-a-string
---

# --description--

在本實驗中，你將練習將字串截斷到特定長度。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個函式 `truncateString`，它接受兩個引數，第一個是字串，第二個是數字。
1. 如果字串的長度超過給定的數字，則應該截斷字串以縮短長度，使其與給定的數字相等，並且應在截斷後的字串末尾附加 `...`。
1. 如果字串的長度相等於或小於給定的數字，則應傳回字串本身且不做更改。

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)` 應該傳回字串 `A-tisket...`。

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)` 應該傳回字串 `Peter Piper...`。

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)` 應該傳回字串 `A-tisket a-tasket A green and yellow basket`。

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)` 應該傳回字串 `A-tisket a-tasket A green and yellow basket`。

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)` 應該傳回字串 `A...`。

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)` 應該傳回字串 `Ab...`。

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
