---
id: 615f34ecc1091b4fd5a8a484
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Ndani ya kipengele chako cha `head`, ongeza kipengele cha `link` chenye sifa ya kipengele `rel` iliyowekwa kuwa `stylesheet` na sifa ya kipengele `href` iliyowekwa kuwa `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Hii ita import familia ya fonti ya `Open Sans`, ikiwa na thamani za uzito wa fonti `400`, `700`, na `800`.

Pia ongeza kipengele cha `link` kuunganisha faili lako la `styles.css`.

# --hints--

Msimbo wako unapaswa kuwa na vipengele viwili vya `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Vipengele vyako vya `link` vinapaswa kuwa vipengele visivyo na maelezo ya mwisho (void elements), havipaswi kuwa na lebo ya mwisho `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Vipengele vyako viwili vya `link` vinapaswa kuwa ndani ya kipengele cha `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Vipengele vyako vyote vya `link` vinapaswa kuwa na sifa ya kipengele `rel` iliyowekwa kuwa `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Kimoja cha vipengele vyako vya `link` kinapaswa kuwa na sifa ya kipengele `href` iliyowekwa kuwa `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

Kimoja cha vipengele vyako vya `link` kinapaswa kuwa na sifa ya kipengele `href` iliyowekwa kuwa `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
