---
id: 68e707257ea91c2e12d489fb
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 학부모 상담회 양식을 만들어 라디오 버튼에 사용자 정의 스타일을 추가하는 방법을 연습합니다. HTML 보일러플레이트가 제공되어 있습니다.

`main` 요소를 추가하고 `container`라는 클래스를 지정하는 것부터 시작하세요.

# --hints--

`main` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("main"));
```

`main` 요소에 `container`라는 클래스가 있어야 합니다.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
