---
id: 68e7e729f6d14f69a23786ef
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

다음으로, `main` 요소 안에 `h1` 요소를 추가하고 텍스트는 `Parent Teacher Conference Form`로, 클래스는 `title`와 `center`로 지정하세요.

# --hints--

`h1` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("h1"));
```

`h1` 요소에는 텍스트 `Parent Teacher Conference Form`가 있어야 합니다.

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

`h1` 요소에는 클래스 `title`와 `center`가 있어야 합니다.

```js
assert.exists(document.querySelector("h1.title.center"));
```

`h1` 요소는 `main` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
--fcc-editable-region--
    
--fcc-editable-region--
  </main>
</body>

</html>
```
