---
id: 68ea8c330ed07b60879f3d38
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

이제 학생의 성적 정보를 수집할 폼 요소를 추가할 시간입니다.

텍스트가 `label`이고 `Student Grade: ` 속성이 `for`로 설정된 또 다른 `"grade"` 요소를 추가하는 것부터 시작하세요.

그런 다음 `label` 요소 아래에 `input` 요소를 추가하고 `type` 속성을 `"number"`로 설정하세요. `name`와 `id` 속성은 `"grade"`으로 설정해야 합니다. `placeholder` 속성은 `"E.g., 4"`로 설정하세요. 마지막으로 숫자 `input`는 필수로 지정해야 합니다.

# --hints--

페이지에 `label` 요소가 총 두 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

두 번째 `label` 요소의 텍스트는 `Student Grade: `이어야 합니다. 콜론 뒤에 공백을 잊지 마세요.

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

두 번째 `label` 요소에는 `for` 속성이 있어야 합니다.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

두 번째 `label` 요소의 `for` 속성 값은 `"grade"`이어야 합니다.

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

페이지에 `input` 요소가 두 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

두 번째 `input` 요소는 두 번째 `label` 요소 아래에 있어야 합니다.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

두 번째 `input` 요소에는 `type` 속성이 `"number"`로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

두 번째 `input` 요소에는 `name` 속성이 `"grade"`로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

두 번째 `input` 요소에는 `id` 속성이 `"grade"`로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

두 번째 `input` 요소에는 `placeholder` 속성이 `"E.g., 4"`로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

두 번째 `input` 요소에는 `required` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
