---
id: 68ea8666b8b473455f5b33ef
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

양식의 첫 번째 섹션은 학생의 이름과 학년 같은 정보를 다룹니다.

`form` 요소 안에 `fieldset` 요소를 추가하세요. 그 `fieldset` 요소 안에 `legend` 요소를 추가하고 텍스트 `Student Information`를 넣으세요.

`legend` 요소 아래에 텍스트 `label`와 값이 `Full Name: `인 `for` 속성을 가진 `student-name` 요소를 추가하세요.


# --hints--

`fieldset` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("fieldset"));
```

`fieldset` 요소는 `form` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector("form fieldset"));
```

`legend` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("legend"));
```

`legend` 요소는 `fieldset` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector("fieldset legend"));
```

`legend` 요소는 텍스트 `Student Information`를 가져야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

`label` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("label"));
```

`label` 요소는 텍스트 `Full Name: `를 가져야 합니다. 콜론 뒤에 공백을 잊지 마세요.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

`label` 요소는 `for` 속성을 가져야 합니다.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

`label` 요소의 `for` 속성은 값 `student-name`를 가져야 합니다.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

`label` 요소는 `legend` 요소 아래에 있어야 합니다.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
