---
id: 68ea8e6a87a2816b0bb2eb7d
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

다음 단계는 부모 정보 입력란 섹션입니다.

`fieldset` 요소를 하나 더 추가하는 것부터 시작하세요. 그 `fieldset` 요소 안에 `legend` 요소를 추가하고, 그 안에 `Parent/Guardian Information` 텍스트를 넣으세요.

# --hints--

페이지에 `fieldset` 요소가 총 두 개 있습니다.

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 2);
```

두 번째 `legend` 요소 안에 `fieldset` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

`legend` 요소에는 `Parent/Guardian Information` 텍스트가 있어야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Parent/Guardian Information");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
