---
id: 68ea8fdf2c9b167307e7f671
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

이제 부모님의 정보를 수집할 폼 요소를 추가할 시간입니다.

텍스트가 `label`이고 `Parent/Guardian Name: ` 속성이 `for`로 설정된 `"parent-name"` 요소를 추가하는 것부터 시작하세요.

그런 다음 `label` 요소 아래에 `input` 요소를 추가하고 `type` 속성을 `"text"`로 설정하세요. `name`와 `id` 속성은 `"parent-name"`으로 설정해야 합니다. `placeholder` 속성은 `"E.g., Nancy Doe"`로 설정하세요. 마지막으로 `input`는 필수여야 합니다.

# --hints--

페이지에 `label` 요소가 총 세 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

세 번째 `label` 요소의 텍스트는 `Parent/Guardian Name: `이어야 합니다. 콜론 뒤에 공백을 잊지 마세요.

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

세 번째 `label` 요소에는 `for` 속성이 있어야 합니다.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

세 번째 `label` 요소의 `for` 속성은 `"parent-name"` 값을 가져야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

페이지에 `input` 요소가 세 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

세 번째 `input` 요소는 세 번째 `label` 요소 아래에 있어야 합니다.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

세 번째 `input` 요소에는 `type` 속성이 `"text"`로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

세 번째 `input` 요소에는 `name` 속성이 `"parent-name"`으로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

세 번째 `input` 요소에는 `id` 속성이 `"parent-name"`으로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

세 번째 `input` 요소에는 `placeholder` 속성이 `"E.g., Nancy Doe"`으로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

세 번째 `input` 요소에는 `required` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
