---
id: 68ea9f01c51b66b937793611
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

다음으로, 클래스가 `"contact-method"`이고 `for` 속성이 `"phone"`로 설정된 레이블을 하나 더 추가하세요. 레이블 텍스트는 `Phone: `이어야 합니다.

`label` 요소 아래에 `id`와 `value` 속성이 `"phone"`로 설정된 라디오 버튼을 추가하세요. `name` 속성은 `"contact-method"`로 설정하고 클래스는 `"contact-method-radio-btn"`으로 설정해야 합니다.

# --hints--

마지막 `label` 안에 `fieldset` 요소가 두 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type label"), 2);
```

`label` 요소는 `"contact-method"` 클래스를 가져야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

`label` 요소의 `for` 속성은 `"phone"`로 설정해야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type label[for='phone']"));
```

`label` 요소는 `Phone: ` 텍스트를 가져야 합니다.

```js
const label = document.querySelector("fieldset:last-of-type label[for='phone']");
assert.match(label?.innerText, /Phone:\s?/);
```

마지막 `input` 안에 `fieldset` 요소가 두 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type input"), 2);
```

`input` 요소는 `type` 속성이 `radio`로 설정되어야 합니다.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

`input` 요소는 `id` 속성이 `"phone"`로 설정되어야 합니다.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("id"), "phone");
```

`input` 요소는 `value` 속성이 `"phone"`로 설정되어야 합니다.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("value"), "phone");
```

`input` 요소는 `name` 속성이 `"contact-method"`로 설정되어야 합니다.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

`input` 요소는 `"contact-method-radio-btn"` 클래스를 가져야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type input[id='phone'].contact-method-radio-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
