---
id: 68eaa668d4d9fbe529a0e16b
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

이제 텍스트가 `label`이고 `Any specific concerns or topics you'd like to discuss?` 속성이 `for`로 설정된 `notes` 요소를 추가하세요.

이 `label` 아래에 `textarea` 요소를 추가하세요. `id`를 `notes`로, `name`를 `notes`로, `rows`를 `4`로, `cols`를 `50`로 설정하세요.

# --hints--

마지막 `fieldset`에는 `label` 요소가 포함되어야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

`label` 요소는 텍스트가 `Any specific concerns or topics you'd like to discuss?`이어야 합니다.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

`label` 요소의 `for` 속성은 `"notes"`로 설정해야 합니다.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

마지막 `fieldset`에는 `textarea` 요소가 포함되어야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

`textarea` 요소는 `label` 요소 아래에 있어야 합니다.

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

`textarea` 요소는 `id` 속성이 `"notes"`로 설정되어야 합니다.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

`textarea` 요소는 `name` 속성이 `"notes"`로 설정되어야 합니다.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

`textarea` 요소는 `rows` 속성이 `4`로 설정되어야 합니다.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

`textarea` 요소는 `cols` 속성이 `50`로 설정되어야 합니다.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
