---
id: 68ea7c41f2555e16174e7cfa
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

다음으로, 텍스트 `Please fill out the form below to help schedule your parent-teacher conference.`가 들어간 단락 요소를 추가하세요. 단락 요소에는 `description`과 `center` 클래스도 포함되어야 합니다.

# --hints--

`p` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("p"));
```

`p` 요소에는 텍스트 `Please fill out the form below to help schedule your parent-teacher conference.`가 있어야 합니다.

```js
assert.strictEqual(document.querySelector("p")?.innerText, "Please fill out the form below to help schedule your parent-teacher conference.");
```

`p` 요소에는 `description`과 `center` 클래스가 있어야 합니다.

```js
assert.exists(document.querySelector("p.description.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
