---
id: 68ea8a5db73b155708992687
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

다음 단계는 학생 정보를 위한 관련 `input` 요소를 추가하는 것입니다.

먼저 `input` 속성이 `type`로 설정된 `"text"` 요소를 추가하세요. 그런 다음 `name`와 `id` 속성을 추가하는데, 둘 다 `student-name`로 설정하세요.

다음으로 `placeholder` 속성을 `E.g., Jane Doe`로 설정하세요. 마지막으로 `required` 속성을 추가하세요.

# --hints--

`input` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("fieldset input"));
```

`input` 요소에 `type` 속성이 `"text"`로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

`input` 요소에 `name` 속성이 `"student-name"`로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

`input` 요소에 `id` 속성이 `"student-name"`로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

`input` 요소에 `placeholder` 속성이 `E.g., Jane Doe`로 설정되어야 합니다.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

`input` 요소에 `required` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
