---
id: 68ea95c833f0518ae6e366cd
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

다음 몇 단계에서는 사용자가 선호하는 연락 방법을 수집하는 폼 요소를 추가할 것입니다.

`fieldset` 요소를 하나 더 추가하고 그 안에 `legend` 요소를 중첩하세요. `legend` 요소에는 `Preferred Contact Method` 텍스트가 있어야 합니다.

# --hints--

페이지에 `fieldset` 요소가 총 세 개 있습니다.

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 3);
```

세 번째 `legend` 요소 안에 `fieldset` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

`legend` 요소에는 `Preferred Contact Method` 텍스트가 있어야 합니다.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Preferred Contact Method");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
