---
id: 68ea975c7cbbb894a386e402
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

다음 단계는 이메일 연락 방법을 위한 `label`와 `input` 요소를 추가하는 것입니다.

`"contact-method"` 클래스를 가진 레이블과 `for` 속성이 `"email"`로 설정된 레이블을 추가하는 것부터 시작하세요. 레이블 텍스트는 `Email: `여야 합니다.

`label` 요소 아래에 `id`와 `value` 속성이 `"email"`로 설정된 라디오 버튼을 추가하세요. `name` 속성은 `"contact-method"`로 설정하고 클래스는 `"contact-method-radio-btn"`로 설정해야 합니다.

마지막으로, 이 라디오 버튼이 기본적으로 선택되어 있는지 확인하세요.

# --hints--


마지막 필드셋에는 `label` 요소가 포함되어야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

이 새로운 `label` 요소(이메일용)는 `"contact-method"` 클래스를 가져야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

새로운 `label` 요소의 `for` 속성은 `"email"`로 설정되어야 합니다.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

`label` 요소에는 `Email: ` 텍스트가 있어야 합니다. 콜론 뒤에 공백이 있는지 확인하세요.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

마지막 `input` 안에 `fieldset` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

`input` 요소에는 `type` 속성이 `radio`로 설정되어야 합니다.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

`input` 요소에는 `id` 속성이 `"email"`로 설정되어야 합니다.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

`input` 요소에는 `value` 속성이 `"email"`로 설정되어야 합니다.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

`input` 요소에는 `name` 속성이 `"contact-method"`로 설정되어야 합니다.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

`input` 요소에는 `"contact-method-radio-btn"` 클래스가 있어야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

`radio` 버튼은 기본적으로 `checked` 상태여야 합니다.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
