---
id: 68ea81da9f687732368a5322
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

이제 부모와 학생 정보를 나타낼 `form`와 `input` 요소를 추가할 차례입니다.

`form` 요소 아래에 `p` 요소를 추가하는 것부터 시작하세요.

# --hints--

`form` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector(".container form"));
```

form 요소는 `p` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector(".container > p+form"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
