---
id: 68eafe1bae175f3b2a87a90d
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

마지막으로, 폼에 제출 버튼을 추가하세요.

마지막 `fieldset` 요소 아래에 `button` 요소를 추가하고, 클래스는 `submit-btn`로, `type` 속성은 `"submit"`로 설정하세요. 버튼의 텍스트 내용은 `Submit Form`여야 합니다.

이로써 HTML 구조가 완성됩니다!

# --hints--

마지막 `button` 요소 아래에 `fieldset` 요소를 추가하세요.

```js
assert.exists(document.querySelector("fieldset:last-of-type + button"));
```

`button` 요소는 `"submit-btn"` 클래스를 가져야 합니다.

```js
assert.exists(document.querySelector("fieldset:last-of-type + button.submit-btn"));
```

`button` 요소는 `type` 속성을 `"submit"`로 설정해야 합니다.

```js
const button = document.querySelector("fieldset:last-of-type + button");
assert.strictEqual(button?.getAttribute("type"), "submit");
```

`button` 요소는 `Submit Form` 텍스트를 가져야 합니다.

```js
const button = document.querySelector("fieldset:last-of-type + button");
assert.strictEqual(button?.innerText, "Submit Form");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
