---
id: 68eb01876d31f24efecb0475
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

HTML 구조가 완성되었으니 이제 CSS 스타일링으로 넘어가세요.

`styles.css` 파일을 여세요. 먼저 `body`에 어두운 배경색을 설정하고 텍스트에는 밝은 색을 지정하세요.

`body` 요소를 선택한 후 `background-color`를 `midnightblue`로, `color`를 `whitesmoke`로 설정하세요.

# --hints--

`body` 요소를 대상으로 지정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

`background-color` 속성을 `midnightblue`로 설정하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(style?.backgroundColor, 'midnightblue');
```

`color` 속성을 `whitesmoke`로 설정하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(style?.color, 'whitesmoke');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
