---
id: 68eb058081504963447bfa60
title: 17단계
challengeType: 0
dashedName: step-17
---

# --description--

다음으로, 메인 컨테이너를 스타일링하세요. `container` 클래스를 가진 요소를 선택하고 그 `background-color` 속성을 16진수 코드 `#ffffff1a`로 설정하세요.

16진수 코드는 HTML과 CSS에서 색상을 나타내기 위해 숫자와 문자로 이루어진 6자리 조합입니다. `#`로 시작하며 세 쌍의 숫자가 뒤따릅니다:

- 처음 두 자리는 빨간색을 나타냅니다
- 다음 두 자리는 초록색을 나타냅니다
- 마지막 두 자리는 파란색을 나타냅니다

예를 들어, `#ffffff`는 빨간색, 초록색, 파란색의 최대값을 가지므로 순수한 흰색입니다.

또한 불투명도를 조절하는 두 자리 숫자를 끝에 추가할 수 있는데, 이를 알파 채널이라고 합니다. `#ffffff1a`에서 `1a`는 흰색을 반투명하게 만듭니다. 알파 값이 낮을수록 색상이 더 투명하게 보입니다. 16진수 코드에 대해서는 나중에 더 배우게 될 것입니다!

# --hints--

`container` 클래스를 가진 요소를 선택하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

`background-color` 속성을 `#ffffff1a`로 설정하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('.container')
assert.strictEqual(style?.backgroundColor, 'rgba(255, 255, 255, 0.1)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```


```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

--fcc-editable-region--

--fcc-editable-region--
```
