---
id: 68eb206a76e5d0b9c0b40eb7
title: 23단계
challengeType: 0
dashedName: step-23
---

# --description--

`fieldset` 요소에 스타일을 지정하여 구별되는 테두리를 만드세요.

`fieldset` 요소를 선택하세요. `border`를 `1px solid gray`로, `border-radius`를 `5px`로, `margin`를 `20px 0`로, `padding`를 `20px`로 설정하세요.

# --hints--

`fieldset` 요소를 대상으로 지정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('fieldset'));
```

`border` 속성을 `1px solid gray`로 설정하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('fieldset');
assert.oneOf(style?.border, ['1px solid gray', '1px solid grey']);
```

`border-radius` 속성을 `5px`로 설정하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('fieldset');
assert.strictEqual(style?.borderRadius, '5px');
```

`margin` 속성을 `20px 0`로 설정하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('fieldset');
assert.strictEqual(style?.margin, '20px 0px');
```

`padding` 속성을 `20px`로 설정하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('fieldset');
assert.strictEqual(style?.padding, '20px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

.container {
    background-color: #ffffff1a;
    width: 80%;
    max-width: 600px;
    border-radius: 10px;
    margin: 20px auto;
    padding: 10px 20px;
    box-shadow: 0 5px 15px black;  
}

.center {
  text-align: center;
}

.description {
  font-size: 1.2rem;
}

--fcc-editable-region--

--fcc-editable-region--
```
