---
id: 68eb30a907dc8813531084df
title: 31단계
challengeType: 0
dashedName: step-31
---

# --description--

기본적으로 라디오 버튼 같은 폼 요소는 레이블과 완벽하게 정렬되지 않을 수 있습니다. 특히 레이블 텍스트가 더 크거나 다른 폰트 설정을 가진 경우 그렇습니다.

라디오 버튼을 레이블과 수직으로 정렬하려면 `vertical-align` 속성을 사용할 수 있습니다. 이 속성은 인라인 또는 인라인 블록 요소가 주변 텍스트와 수직으로 어떻게 정렬되는지 제어합니다.

이제 라디오 버튼에 `vertical-align`을 `bottom`로 설정하세요. `middle`이나 `top` 같은 다른 값도 실험해 보면서 정렬에 어떤 영향을 주는지 확인할 수 있지만, `bottom`가 보통 레이블 옆 라디오 버튼에 잘 맞습니다.

# --hints--

계속해서 `.contact-method-radio-btn` 클래스를 대상으로 하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn'));
```

`vertical-align` 속성을 `bottom`로 설정하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn');
assert.strictEqual(style?.verticalAlign, 'bottom');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

.container {
    background-color: #ffffff1a;
    width: 80%;
    max-width: 600px;
    border-radius: 10px;
    margin: 20px auto;
    padding: 10px 20px;
    box-shadow: 0 5px 15px black;  
}

.center {
  text-align: center;
}

.description {
  font-size: 1.2rem;
}

fieldset {
  border: 1px solid gray;
  border-radius: 5px;
  margin: 20px 0;
  padding: 20px;
}

fieldset legend {
  font-size: 1.3rem;
  font-weight: 600;
}

label {
  font-size: 1.2rem;
}

label:not(.contact-method) {
  display: block;
  margin: 10px 0;
}

input:not(.contact-method-radio-btn),
textarea {
  background-color: #ffffff1a;
  width: 95%;
  border: 1px solid gray;
  border-radius: 5px;
  padding: 10px;
}

input,
input::placeholder,
textarea {
  color: whitesmoke;
}


.contact-method-radio-btn {
  appearance: none;
  width: 20px;
  height: 20px;
  border-radius: 50%;
  border: 2px solid gray;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
