---
id: 66c07238b01053abaf812065
title: 팩토리얼 계산기 만들기
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

팩토리얼은 정수와 그 아래 모든 정수의 곱입니다. 예를 들어, `5`의 팩토리얼은 `5 * 4 * 3 * 2 * 1 = 120`입니다. 이 실습에서는 사용자로부터 숫자를 입력받아 그 숫자의 팩토리얼을 계산하는 팩토리얼 계산기를 만듭니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `num`라는 변수를 선언하고 원하는 숫자를 할당하세요. 할당된 숫자는 `1`부터 `20`까지 포함된 범위여야 합니다.
1. `factorialCalculator`라는 이름의 함수를 만들고, 숫자를 인수로 받아 그 숫자의 팩토리얼을 반환하세요.
1. 함수 내부에서 `result` 변수를 선언하고 `1` 값을 할당하세요. 반복문을 사용해 `1`부터 입력 숫자(포함)까지 모든 숫자를 순회하며, 각 숫자마다 `result` 변수를 현재 숫자와 곱하고 결과를 `result` 변수에 할당하세요. 여기서는 `for` 반복문, `while` 반복문 또는 `do...while` 반복문 중 하나를 선택해 사용할 수 있습니다.
1. `factorialCalculator` 함수를 `num`를 인수로 호출하고 결과를 `factorial` 변수에 할당하세요.
1. 최종 출력은 `Factorial of [num] is [factorial]` 형식으로 저장하고 `resultMsg` 변수에 할당하세요.
1. `resultMsg` 값을 콘솔에 출력하세요.

# --hints--

`num` 변수가 있어야 합니다.

```js
assert.isDefined(num);
```

`num` 변수에 값을 할당해야 합니다.

```js
assert.isNotNull(num);
```

`num` 값은 `1`과 `20` 사이여야 합니다.

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

`num` 값은 숫자여야 합니다.

```js
assert.isNumber(num);
```

`factorialCalculator`라는 함수가 있어야 합니다.

```js
assert.isFunction(factorialCalculator);
```

`factorialCalculator` 함수는 숫자를 인수로 받아야 합니다.

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

`5`의 팩토리얼은 `120`를 반환해야 합니다.

```js
assert.strictEqual(factorialCalculator(5), 120);
```

`7`의 팩토리얼은 `5040`를 반환해야 합니다.

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

`factorialCalculator` 함수를 `num` 변수를 인수로 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s*$/m);
```

`factorial` 변수를 정의하고 `factorialCalculator` 함수의 결과를 할당해야 합니다.

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

`factorialCalculator`은 올바른 결과를 내야 합니다.

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

`resultMsg`은 `Factorial of [num] is [factorial]` 형식의 문자열을 포함해야 합니다.

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

`resultMsg` 값을 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```
