---
id: 641da5462576784453146ec2
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

现在你有了所需的数据，可以使用它来填充 UI。 但是获取的数据包含 26 位作者的数组，如果将他们全部同时添加到页面，可能会导致性能不佳。

你应该一次添加 8 位作者，然后有一个按钮来添加另外 8 位作者，直到没有更多数据可显示。

使用 `let` 创建 2 个名为 `startingIndex` 和 `endingIndex` 的变量，并分别向它们分配数字值 `0` 和 `8`。 另外，使用 `let` 创建一个 `authorDataArr` 变量并将其设置为一个空数组。

# --hints--

你应该使用 `let` 来声明一个名为 `startingIndex` 的变量。

```js
assert.match(code, /let\s+startingIndex/)
```

你应该将 `startingIndex` 变量设置为 `0`。

```js
assert.match(code, /let\s+startingIndex\s*=\s*0\s*;?/)
```

你应该使用 `let` 来声明一个名为 `endingIndex` 的变量。

```js
assert.match(code, /let\s+endingIndex/)
```

你应该将 `endingIndex` 变量设置为 `8`。

```js
assert.match(code, /let\s+endingIndex\s*=\s*8\s*;?/)
```

你应该使用 `let` 来声明一个名为 `authorDataArr` 的变量。

```js
assert.match(code, /let\s+authorDataArr/)
```

你应该将 `authorDataArr` 变量设置为空数组 (`[]`)。

```js
assert.match(code, /let\s+authorDataArr\s*=\s*\[\s*\]\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
