---
id: 63ee3f71381756f9716727ef
title: CSS-Grundlagen Übung A
challengeType: 25
dashedName: css-foundations-exercise-a
---

# --description--

**Ziel:**  
In dieser Übung üben Sie, CSS zu einer HTML-Datei hinzuzufügen, und zwar mit allen drei Methoden: externes CSS, internes CSS und inline CSS. Sie sollten für diese Übung nur Typselektoren verwenden, wenn Sie Styles über die externe und interne Methode hinzufügen. Verwenden Sie außerdem Farbschlüsselwörter (z. B. „blue“) anstelle von `RGB` oder `HEX` Werten.

## User Stories

- Sie sollten ein `div`-Element mit etwas Text sehen.  
  - Es sollte einen `red`-Hintergrund, `white`-Text, eine Schriftgröße von `32px`, zentrierten Text und `bold` haben.  
  - Das CSS für das `div`-Element sollte extern hinzugefügt werden und einen Typselektor verwenden.  
- Sie sollten ein `p`-Element mit etwas Text sehen.  
  - Es sollte einen `green`-Hintergrund, `white`-Text und eine Schriftgröße von `18px` haben.  
  - Das CSS für das `p`-Element sollte intern hinzugefügt werden und einen Typselektor verwenden.  
- Sie sollten ein `button`-Element mit etwas Text sehen.  
  - Das `button`-Element sollte einen `orange`-Hintergrund und eine Schriftgröße von `18px` haben.  
  - Das CSS für das `button`-Element sollte mit inline Styles hinzugefügt werden.

# --hints--

Sie sollten ein `div`-Element mit etwas Text haben.

```js
const divElement = document.querySelector('div');

assert.isNotNull(divElement);
assert.isAtLeast(divElement?.innerText.length, 1);
```

Sie sollten ein externes Stylesheet haben, das die Styles für das `div`-Element enthält.

```js
const styleSheet = new __helpers.CSSHelp(document).getStyleSheet();
const isExternal = styleSheet?.ownerNode.classList.contains('fcc-injected-styles');
const divStyle = new __helpers.CSSHelp(document).getStyle('div');

assert.isTrue(isExternal);
assert.isNotNull(divStyle);
```

Ihr `div`-Element sollte sein CSS nicht mit internen oder inline Styles erhalten.

```js
const styleElement = document.querySelector('style:not([class])');

assert.isNotTrue(styleElement?.innerText.includes('div'));
assert.isNotTrue(document.querySelector('div')?.hasAttribute('style'));
```

Ihr `div`-Element sollte eine `background-color` von `red` und eine `color` von `white` haben.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const divBGColor = divStyle?.getPropertyValue('background-color');
const divColor = divStyle?.getPropertyValue('color');

assert.equal(divBGColor, 'red');
assert.equal(divColor, 'white');
```

Ihr `div`-Element sollte `font-weight` auf `bold`, `font-size` auf `32px` und `text-align` auf `center` gesetzt haben.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const textAlign = divStyle?.getPropertyValue('text-align');
const fontSize = divStyle?.getPropertyValue('font-size');
const fontWeight = divStyle?.getPropertyValue('font-weight');

assert.equal(textAlign, 'center');
assert.equal(fontSize, '32px');
assert.equal(fontWeight,'bold');
```

Sie sollten ein `p`-Element haben, das etwas Text enthält.

```js
const pElement = document.querySelector('p');

assert.isNotNull(pElement);
assert.isAtLeast(pElement?.innerText.length, 1)
```

Ihr `p`-Element sollte seine Styles intern mit einem `style`-Element erhalten.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let isStyled = false;

if (rules && rules.selectorText === 'p') {
  isStyled = true;
}

assert.isTrue(isStyled);
```

Ihr `p`-Element sollte eine `font-size` von `18px` haben und `color` auf `white` gesetzt sein.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let fontSize, color;

if (rules && rules.selectorText === 'p') {
  fontSize = rules.style.fontSize;
  color = rules.style.color;
}

assert.equal(fontSize, "18px");
assert.equal(color, 'white');
```

Sie sollten ein `button`-Element mit etwas Text haben.

```js
const btnElement = document.querySelector('button');

assert.isNotNull(btnElement);
assert.isAtLeast(btnElement?.innerText.length, 1);
```

Ihr `button`-Element sollte einen inline Style haben.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('style'));
```
 
Ihr `button`-Element sollte sein `background-color` auf `orange` gesetzt haben.

```js
assert.equal(document.querySelector('button')?.style.backgroundColor, 'orange')
```

Ihr `button`-Element sollte sein `font-size` auf `18px` gesetzt haben.

```js
assert.equal(document.querySelector('button')?.style.fontSize, '18px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>

  </body>
</html>
```

```css
/* styles.css */
```

# --solutions--

```html
<!DOCTYPE html>
<html>
<head>
  <title>My Styling Example</title>
  <link rel="stylesheet" type="text/css" href="styles.css">
  <style>
    p {
      background-color: green;
      color: white;
      font-size: 18px;
    }
  </style>
</head>
<body>
  <div>Hello World!</div>
  <p>This is a paragraph.</p>
  <button style="background-color: orange; font-size: 18px;">Click Me</button>
</body>
</html>
```

```css
div {
  background-color: red;
  color: white;
  font-size: 32px;
  text-align: center;
  font-weight: bold;
}
```
