---
id: 63ee3ff1381756f9716727f2
title: CSS-Grundlagen Übung D
challengeType: 25
dashedName: css-foundations-exercise-d
---

# --description--

Bei dieser Übung haben wir Ihnen eine fertige HTML-Datei bereitgestellt, sodass Sie nur die CSS-Datei bearbeiten müssen. Für diese Übung ist es wichtiger zu verstehen, wie die Verkettung verschiedener Selektoren funktioniert, als tatsächlich die Attribute hinzuzufügen.

1. Sie sollten eine `width` von `300px` auf der `avatar`- und `proportioned`-Klasse sehen.
1. Sie sollten ihr eine Höhe geben, damit sie ihre ursprünglichen quadratischen Proportionen behält (verwenden Sie keinen festen Pixelwert für die Höhe!).
1. Sie sollten den Elementen mit sowohl der `avatar`- als auch der `distorted`-Klasse eine `width` von `200px` geben.
1. Sie sollten ihr eine `height` geben, die doppelt so groß ist wie ihre Breite.

# --hints--

Sie sollten eine `width` von `300px` auf der `avatar`- und `proportioned`-Klasse haben.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.proportioned`) || new __helpers.CSSHelp(document).getStyle(`.proportioned.avatar`);
assert(style?.width === '300px');
```

Sie sollten eine Höhe von `auto` auf der `avatar`- und `proportioned`-Klasse haben.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.proportioned`) || new __helpers.CSSHelp(document).getStyle(`.proportioned.avatar`)
assert(style?.height === 'auto');
```

Sie sollten einen verketteten Selektor auf der `avatar`- und `proportioned`-Klasse verwenden.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.proportioned`) || new __helpers.CSSHelp(document).getStyle(`.proportioned.avatar`);
assert(style);
```

Sie sollten eine `width` von `200px` auf der `avatar`- und `distorted`-Klasse haben.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.distorted`) || new __helpers.CSSHelp(document).getStyle(`.distorted.avatar`);
assert(style?.width === '200px');
```

Sie sollten einen verketteten Selektor auf der `avatar`- und `distorted`-Klasse verwenden.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.distorted`) || new __helpers.CSSHelp(document).getStyle(`.distorted.avatar`);
assert(style);
```

Sie sollten eine `height` haben, die doppelt so groß wie die Breite auf der `avatar`- und `distorted`-Klasse ist.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.distorted`) || new __helpers.CSSHelp(document).getStyle(`.distorted.avatar`);
assert(style?.height === '400px');
```

# --seed--

## --seed-contents--

```css

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Chaining Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <!-- Use the classes BELOW this line -->
    <div>
      <img class="avatar proportioned" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-01.jpg" alt="Two kittens on a blanket.">
      <img class="avatar distorted" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-02.jpg" alt="A cute kitten in the woods.">
    </div>
    <!-- Use the classes ABOVE this line -->
    <div>
      <img class="original proportioned" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-01.jpg" alt="Two kittens on a blanket.">
      <img class="original distorted" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-02.jpg" alt="A cute kitten in the woods.">
    </div>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Chaining Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <!-- Use the classes BELOW this line -->
    <div>
      <img class="avatar proportioned" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-01.jpg" alt="Two kittens on a blanket.">
      <img class="avatar distorted" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-02.jpg" alt="A cute kitten in the woods.">
    </div>
    <!-- Use the classes ABOVE this line -->
    <div>
      <img class="avatar proportioned" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-01.jpg" alt="Two kittens on a blanket.">
      <img class="avatar distorted" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-02.jpg" alt="A cute kitten in the woods.">
  </body>
</html>
```

```css
.avatar.proportioned {
  height: auto;
  width: 300px;
}

.avatar.distorted {
  height: 400px;
  width: 200px;
}
```

