---
id: 63ee3fe9381756f9716727f1
title: CSS-Grundlagen Übung C
challengeType: 25
dashedName: css-foundations-exercise-c
---

# --description--

Bauen wir ein wenig auf der vorherigen Übung auf, in der Sie mehreren Klassen zu einem einzelnen Element hinzugefügt haben, um zwei verschiedene Regeln darauf anzuwenden.

1. Sie sollten einen schwarzen Hintergrund und weißen Text beim ersten `button`-Element sehen.
1. Sie sollten einen gelben Hintergrund beim zweiten `button`-Element sehen.
1. Sie sollten eine Schriftgröße von `28px` mit einer Selektorliste festlegen.
1. Sie sollten eine Liste von Schriftarten haben, die `Helvetica` und `Times New Roman` enthält, mit `sans-serif` als Fallback in der Selektorliste.
1. Sie sollten einen eindeutigen Klassennamen bei jedem Element haben.
1. Sie sollten eine Selektorliste für Stile haben, die beide Elemente gemeinsam haben.

# --hints--

Sie sollten einen `black`-Hintergrund beim ersten Element haben.

```js
const classes = document.querySelectorAll('button')?.[0].classList;

const style = new __helpers.CSSHelp(document).getStyle(`.${classes[0]}`);

assert.equal(style?.backgroundColor, 'black');

```

Ihr erstes Element sollte eine Textfarbe von `white` haben.

```js

const classes = document.querySelectorAll('button')?.[0].classList;

const style = new __helpers.CSSHelp(document).getStyle(`.${classes[0]}`);

assert.equal(style?.color, 'white');

```

Sie sollten die `font-size` für beide Elemente mit einer Selektorliste auf `28px` setzen.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}
assert.equal(eitherOr()?.fontSize, '28px');
```

Sie sollten einen eindeutigen Klassennamen bei jedem Element haben.

```js
const elementOneClasses = document.querySelectorAll('button')?.[0].classList;
const elementTwoClasses = document.querySelectorAll('button')?.[1].classList;

for(let i = 0; i < elementOneClasses.length; i++){
    assert(![...elementTwoClasses].includes(elementOneClasses[i]));
}

```

Sie sollten eine Selektorliste für Stile haben, die beide Elemente gemeinsam haben.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}
assert.exists(eitherOr());
```

Sie sollten eine Liste von Schriftarten haben, die `Helvetica` und `Times New Roman` enthält, mit `sans-serif` als Fallback in der Selektorliste.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}

assert.equal(eitherOr()?.fontFamily, 'Helvetica, "Times New Roman", sans-serif');
```

Jedes Element sollte nur eine Klasse haben.

```js
const elementOneClasses = document.querySelectorAll('button')?.[0].classList;
const elementTwoClasses = document.querySelectorAll('button')?.[1].classList;

assert(elementOneClasses.length === 1 && elementTwoClasses.length === 1);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Grouping Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <button>Click Me!</button>
    <button>No, Click Me!</button>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
  <html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Grouping Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <button class="inverted">Click Me!</button>
    <button class="fancy">No, Click Me!</button>
  </body>
</html>
```

```css
.inverted,
.fancy {
  font-family: Helvetica, "Times New Roman", sans-serif;
  font-size: 28px;
}

.inverted {
  background-color: black;
  color: white;
}

.fancy {
  background-color: yellow;
}
```
