---
id: 68af274e60e009597dd861e4
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

在前面的步骤中，你创建了一个基本的电影对象。现在，你将添加一些验证，以确保电影不能用无效的属性实例化。

在 `__init__` 方法开始处，创建一个 `if` 语句，用于查看 `title` 是否为空或仅包含空白字符（你可以使用 `strip` 方法来构建条件）。

如果是这样，则引起带有消息 `Title cannot be empty` 的 `ValueError`。

# --hints--

当 `title` 为空或仅包含空白时，你应该引起一个带有消息 `Title cannot be empty` 的 `ValueError`。

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

当 `title` 是有效的字串时，你不应引起 `ValueError`。

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
