---
id: 68af278a60e009597dd861ea
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

`MediaCatalogue` 类需要一个方法来为目录添加媒体项。创建一个 `add` 方法，该方法接受一个项 `media_item` 并将其追加到 `items` 列表中。

# --hints--

你的 `MediaCatalogue` 应该有一个 `add` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").has_function("add")`)
})
```

你的 `add` 方法应有两个参数：`self` 和 `media_item`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("add").has_args("self, media_item")`)
})
```

你的 `add` 方法应将传入的项追加到 `items` 属性。

```js
({
    test: () => runPython(`
    _catalogue = MediaCatalogue()
    _catalogue.add("first")
    assert _catalogue.items == ["first"]
    _catalogue.add("second")
    assert _catalogue.items == ["first", "second"]
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
