---
id: 68af277560e009597dd861e8
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

现在删除你的 `print(movie1)` 行。然后，将你的 `movie1` 赋值放入一个 `try` 块中，并创建一个 `except` 子句来捕获 `ValueError as e`，并使用 f 字串打印 `Validation Error:`，后跟一个空格和错误信息。

随意更改传递给实例化 `movie1` 的参数以查看不同的错误信息。

# --hints--

你的代码中不应有 `print(movie1)`。

```js
({
  test: () => runPython(`assert not _Node(_code).has_call("print(movie1)")`)
})
```

你应该有一个包含 `movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)` 的 `try` 块。

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().has_stmt("movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)")`)
})
```

你应该有一个 `except` 子句来捕获 `ValueError as e`。

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].has_except("ValueError", "e")`)
})
```

你应该在 `except` 块中使用 f-string 打印 `Validation Error:` 后跟一个空格和错误信息。

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_except("ValueError", "e").has_call("print(f'Validation Error: {e}')")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
--fcc-editable-region--
```
