---
id: 68af277e60e009597dd861e9
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

现在创建一个名为 `MediaCatalogue` 的新建类，并为它添加一个 `__init__` 方法，该方法将 `items` 属性设置为空列表。这将表现目录中的项列表。 

# --hints--

你应该有一个名为 `MediaCatalogue` 的类。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("MediaCatalogue")`)
})
```

你的 `MediaCatalogue` 类应有一个 `__init__` 方法。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_class("MediaCatalogue").has_function("__init__")
    `)
})
```

你的 `__init__` 方法应将 `items` 属性设置为空列表。

```js
({
    test: () => runPython(`
    _catalogue = MediaCatalogue()
    assert _catalogue.items == []
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
