---
id: 68af274560e009597dd861e3
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

现在，创建一个 `Movie` 实例，传入你选择的参数，并将其赋值给名为 `movie1` 的变量。然后，将 `movie1` 打印到控制台。

# --hints--

你应该有一个名为 `movie1` 的变量。

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

你应该将一个 `Movie` 实例赋值给 `movie1` 变量。

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

你应该将 `movie1` 打印到控制台。

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
```
