---
id: 68af276360e009597dd861e6
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

创建另一个 `if` 语句来查看 `director` 是否为空或仅包含空白，并引起带有消息 `Director cannot be empty` 的 `ValueError`。

# --hints--

当 `director` 为空或仅包含空白时，你应该引起一个带有消息 `Director cannot be empty` 的 `ValueError`。

```js
({
  test: () => runPython(`
  _msg = 'Director cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1999, ' ', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 1999, '  ', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

当 `director` 是有效的字串时，你不应引起 `ValueError`。

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid director'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
