---
id: 68af272f60e009597dd861e1
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次研讨会中，你将创建一个媒体目录。首先定义一个名为 `Movie` 的类。

为你的类添加一个带有以下参数的 `__init__` 方法：`self`、`title`、`year`、`director` 和 `duration`。在 `__init__` 方法内部，将 `title`、`year`、`director` 和 `duration` 参数赋值给同名的实例属性。

# --hints--

你应该创建一个名为 `Movie` 的类。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

你的 `Movie` 类应有一个 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

你的 `__init__` 方法应包含以下参数：`self`、`title`、`year`、`director` 和 `duration`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

你应该在你的 `__init__` 方法内将 `title` 赋值给 `self.title`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

你应该在你的 `__init__` 方法内将 `year` 赋值给 `self.year`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

你应该在你的 `__init__` 方法内将 `director` 赋值给 `self.director`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

你应该在你的 `__init__` 方法内将 `duration` 赋值给 `self.duration`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
