---
id: 68b6d9692a6526838265ad7f
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

`TVSeries` 对象将相对于其父类 `Movie` 实例化时添加额外的属性。因此，创建一个带有以下参数的 `__init__` 方法：`self`、`title`、`year`、`director`、`duration`、`seasons` 和 `total_episodes`。

在 `__init__` 方法内，将 `seasons` 和 `total_episodes` 赋值给同名属性。

# --hints--

你的 `TVSeries` 类应该有一个 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").has_function("__init__")`)
})
```

你的 `__init__` 方法应包含以下参数：`self`、`title`、`year`、`director`、`duration`、`seasons` 和 `total_episodes`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_args("self, title, year, director, duration, seasons, total_episodes")`)
})
```

你应该在你的 `__init__` 方法中将 `seasons` 赋值给 `self.seasons`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("self.seasons = seasons")`)
})
```

你应该在你的 `__init__` 方法中将 `total_episodes` 赋值给 `self.total_episodes`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("self.total_episodes = total_episodes")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class TVSeries(Movie):
    pass
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
