---
id: 68b6d5244e6ce639e52fd1e7
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

如果你想表现不同类型的媒体项，比如说一个电视串行，你可以创建另一个类并让它继承自 `Movie`，以访问它的所有方法和属性。

创建一个名为 `TVSeries` 的空类，并使其继承自 `Movie`。以下是如何使一个 `Child` 类继承自一个 `Parent` 类。

```py
class Parent:
    ...

class Child(Parent):
    ...
```

# --hints--

你应该有一个名为 `TVSeries` 的类。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("TVSeries")`)
})
```

你的 `TVSeries` 类应继承自 `Movie`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").inherits_from("Movie")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
