---
id: 68b702892910ab7b495de081
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

你需要一种合适的方式来直观表现目录。所以在你的 `MediaCatalogue` 类中创建一个 `__str__` 方法。

如果 `items` 列表为空，确保 `__str__` 方法返回 `Media Catalogue (empty)`。

# --hints--

你的 `MediaCatalogue` 类中应该有一个 `__str__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").has_function("__str__")`)
})
```

当 `self.items` 为空时，你的 `__str__` 方法应返回 `Media Catalogue (empty)`。

```js
({
  test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    print(MediaCatalogue())
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Media Catalogue (empty)" in output
  `)
})
```

你的 `__str__` 方法在 `self.items` 不为空时不应返回 `Media Catalogue (empty)`。

```js
({
  test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
    _catalogue = MediaCatalogue()
    _catalogue.add(Movie('The Matrix', 1999, 'The Wachowskis', 136))
    try:
      print(_catalogue)
    except TypeError:
      pass
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Media Catalogue (empty)" not in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
