---
id: 68b804f3be853a355f055fef
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

`super` 函数用于在不显式命名的情况下参考父类。例如，`super().__init__(arg1, arg2, ...)` 调用父类中定义的 `__init__` 方法。

目前，你实际上已经重写了从 `Movie` 继承的 `__init__` 方法。这意味着 `TVSeries` 对象将只有一个 `seasons` 属性和一个 `total_episodes` 属性。这不是你想要的。

为了保持所有在 `Movie` 类的 `__init__` 中定义的代码，使用 `super` 来调用父类的 `__init__`。请注意，你需要传入该方法所需的参数。在本例中为 `title`、`year`、`director` 和 `duration`。

# --hints--

你应该在你的 `__init__` 方法内调用 `super().__init__(title, year, director, duration)`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("super().__init__(title, year, director, duration)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        
        self.seasons = seasons
        self.total_episodes = total_episodes
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
