---
id: 68b95c952d50ed52abe02ed3
title: 步骤 31
challengeType: 20
dashedName: step-31
---

# --description--

为你的 `MediaError` 类添加一个带有以下参数的 `__init__` 方法：`self`、`message` 和 `obj`。

在其中，调用 `super().__init__(message)` 以确保用于存储和显示错误信息的机制（已存在于 `Exception` 类中）能够正常工作。

之后，将 `obj` 赋值给 `self.obj`。

# --hints--

你的 `MediaError` 类应有一个带有 `self`、`message` 和 `obj` 参数的 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaError").find_function("__init__").has_args("self, message, obj")`)
})
```

你应该在你的 `__init__` 方法内调用 `super().__init__(message)`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaError").find_function("__init__").has_stmt("super().__init__(message)")`)
})
```

你应该在你的 `__init__` 方法中将 `obj` 赋值给 `self.obj`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaError").find_function("__init__").has_stmt("self.obj = obj")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MediaError(Exception):
    """Custom exception for media-related errors."""


--fcc-editable-region--
class Movie:
    """Parent class representing a movie."""
    
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    """Child class representing an entire TV series."""

    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)

        if seasons < 1:
            raise ValueError('Seasons must be 1 or greater')
        if total_episodes < 1:
            raise ValueError('Total episodes must be 1 or greater')
        
        self.seasons = seasons
        self.total_episodes = total_episodes

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.seasons} seasons, {self.total_episodes} episodes, {self.duration} min avg, {self.director}'

class MediaCatalogue:
    """A catalogue that can store different types of media items."""

    def __init__(self):
        self.items = []

    def add(self, media_item):
        if not isinstance(media_item, Movie):
            raise TypeError('Only Movie or TVSeries instances can be added')
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)

    series1 = TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182)
    catalogue.add(series1)
    series2 = TVSeries('Breaking Bad', 2008, 'Vince Gilligan', 47, 5, 62)
    catalogue.add(series2)

    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
