---
id: 68b970cf5f000fafa537dbbd
title: 步骤35
challengeType: 20
dashedName: step-35
---

# --description--

由于你将导致异常的对象存储在异常对象中，因此你可以对其进行一定的利用。

在你现有的 `print` 调用之后，添加 `print(f'Unable to add {e.obj}: {type(e.obj)}')`。

这种消息在你尝试调试你的代码时非常有用。

# --hints--

你应该在现有的 `print(f'Media Error: {e}')` 调用之后添加 `print(f'Unable to add {e.obj}: {type(e.obj)}')`。

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_except("MediaError", "e").is_ordered("print(f'Media Error: {e}')", "print(f'Unable to add {e.obj}: {type(e.obj)}')")`)
})
```

# --seed--

## --seed-contents--

```py
class MediaError(Exception):
    """Custom exception for media-related errors."""

    def __init__(self, message, obj):
        super().__init__(message)
        self.obj = obj

class Movie:
    """Parent class representing a movie."""
    
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    """Child class representing an entire TV series."""

    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)

        if seasons < 1:
            raise ValueError('Seasons must be 1 or greater')
        if total_episodes < 1:
            raise ValueError('Total episodes must be 1 or greater')
        
        self.seasons = seasons
        self.total_episodes = total_episodes

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.seasons} seasons, {self.total_episodes} episodes, {self.duration} min avg, {self.director}'

class MediaCatalogue:
    """A catalogue that can store different types of media items."""

    def __init__(self):
        self.items = []

    def add(self, media_item):
        if not isinstance(media_item, Movie):
            raise MediaError('Only Movie or TVSeries instances can be added', media_item)
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)

    series1 = TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182)
    catalogue.add(series1)
    series2 = TVSeries('Breaking Bad', 2008, 'Vince Gilligan', 47, 5, 62)

    catalogue.add('series2')

    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
--fcc-editable-region--
except MediaError as e:
    print(f'Media Error: {e}')

--fcc-editable-region--
```
