---
id: 6718d2d59337c822ecb697ff
title: 銀行口座管理プログラムを作成する
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `BankAccount` という名前のクラスを定義してください。コンストラクタは、デフォルトの `balance` を `0` に設定し、取引記録をオブジェクトとして保存する空の配列 `transactions` を初期化する必要があります。

2. `transactions` 配列に保存される各取引は、`type` と `amount` という2つのプロパティを持つオブジェクトである必要があります。`type` プロパティは `deposit` または `withdraw` のいずれかであり、`amount` プロパティは入金または出金された金額です。

3. `deposit` という名前のメソッドを定義してください。このメソッドは入金額をパラメータとして受け取ります。入金額が `0` より大きい場合、以下を行う必要があります：

   - 取引配列に、タイプが `deposit` で入金額を含む新しいオブジェクトを追加する。
   - 残高を更新する。   
   - `"Successfully deposited $[amount]. New balance: $[balance]"` を返す。

4. 金額が `0` 以下の場合、`deposit` メソッドは `"Deposit amount must be greater than zero."` を返す必要があります。

5. `withdraw` という名前のメソッドを定義してください。このメソッドは金額をパラメータとして受け取ります。出金に応じて現在の残高を更新します。出金額が `0` より大きく、かつ現在の残高以下の場合、以下を行う必要があります：

   - 取引配列に、タイプが `withdraw` で出金額を含む新しいオブジェクトを追加する。
   - 残高を更新する。
   - `"Successfully withdrew $[amount]. New balance: $[balance]"` を返す。

6. 出金額が `0` 以下、または現在の残高を超える場合、`withdraw` メソッドは `"Insufficient balance or invalid amount."` を返す必要があります。

7. `checkBalance` という名前のメソッドを定義してください。このメソッドは `"Current balance: $[balance]"` の形式で現在の残高を返します。

8. `listAllDeposits` という名前のメソッドを定義してください。このメソッドは取引配列を反復して、すべての入金を `"Deposits: amount,amount,..."` の形式で返します。

9. `listAllWithdrawals` という名前のメソッドを定義してください。このメソッドは取引配列を反復して、すべての出金を `"Withdrawals: amount,amount,..."` の形式で返します。

10. `BankAccount` の新しいインスタンスを `myAccount` という名前で作成してください。

11. `myAccount` 銀行口座には少なくとも5件の取引がある必要があります。

12. `myAccount` 銀行口座には少なくとも2件の入金がある必要があります。

13. `myAccount` 銀行口座には少なくとも2件の出金がある必要があります。

14. `myAccount` 銀行口座の残高は `$100` より大きい必要があります。

# --hints--

`BankAccount` という名前のクラスを定義してください。

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

`BankAccount` オブジェクトは最初に `balance` が `0` で、取引記録を保存する空の配列 `transactions` を持つ必要があります。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

入金額をパラメータとして受け取る `deposit` メソッドを用意してください。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

出金額をパラメータとして受け取る `withdraw` メソッドを用意してください。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

現在の残高を確認する `checkBalance` メソッドを用意してください。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

すべての入金を一覧表示する `listAllDeposits` メソッドを用意してください。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

すべての出金を一覧表示する `listAllWithdrawals` メソッドを用意してください。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` は `"Successfully deposited $100. New balance: $100"` を返す必要があります。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` は `"Deposit amount must be greater than zero."` を返す必要があります。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` は `"Deposit amount must be greater than zero."` を返す必要があります。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

口座残高が `100` の場合、`BankAccount.withdraw(150)` は `"Insufficient balance or invalid amount."` を返す必要があります。

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` は `"Insufficient balance or invalid amount."` を返す必要があります。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` は `"Insufficient balance or invalid amount."` を返す必要があります。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

口座残高が `200` の場合、`BankAccount.withdraw(150)` は `"Successfully withdrew $150. New balance: $50"` を返す必要があります。

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

口座残高が `200` の場合、`BankAccount.checkBalance()` は `"Current balance: $200"` を返す必要があります。

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

`10`、`35`、`90` を入金すると、`listAllDeposits` メソッドは `"Deposits: 10,35,90"` を返す必要があります。

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

`20`、`50`、`100` を出金すると、`listAllWithdrawals` メソッドは `"Withdrawals: 20,50,100"` を返す必要があります。

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

`BankAccount` のインスタンスを `myAccount` という名前で用意してください。

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

`myAccount` 銀行口座には少なくとも5件の取引がある必要があります。

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

`myAccount` 銀行口座には少なくとも2件の入金がある必要があります。

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

`myAccount` 銀行口座には少なくとも2件の出金がある必要があります。

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```

`myAccount` 銀行口座の残高は `$100` より大きい必要があります。

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
