---
id: 587d825a367417b2b2512c89
title: Implementar el algoritmo Quicksort
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

En este laboratorio implementarás el algoritmo de ordenamiento: quicksort.

Quicksort es un enfoque eficiente, recursivo y de divide y vencerás para ordenar un arreglo. En este método, se elige un valor pivote en el arreglo original. Luego, el arreglo se divide en dos subarreglos con valores menores y mayores que el valor pivote. Después, se combina el resultado de llamar recursivamente al algoritmo quicksort en ambos subarreglos. Esto continúa hasta que se alcanza el caso base de un arreglo vacío o con un solo elemento, que se devuelve. El desenrollado de las llamadas recursivas devuelve el arreglo ordenado.

Si bien la elección del valor del pivote es importante, cualquier pivote lo será para nuestros propósitos. Por simplicidad, el primero o el último elemento podría ser utilizado.

Quicksort is a very efficient sorting method, providing *O(nlog(n))* performance on average. También es relativamente fácil de implementar. Estos atributos lo convierten en un método de ordenación popular y útil.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función `quicksort` que tome un arreglo de enteros como entrada.
1. La función `quicksort` debe devolver un arreglo que contenga los mismos enteros que la entrada, pero en orden de menor a mayor.
1. La función `quicksort` debe llamarse recursivamente a sí misma para ordenar el arreglo.

# --hints--

Debes tener una función `quicksort`.

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` debe devolver un arreglo que no cambie excepto por el orden.

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` debe devolver un arreglo ordenado (de menor a mayor).

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` no debe usar el método incorporado `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
