---
id: 6551fe3b1df7c9740f13f270
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

Pythonにおける変数名の重要なポイントは次の通りです：

- 一部の単語は予約キーワード（例：`for`、`while`、`True`）で、Pythonで特別な意味を持つため変数名として使えません。
- 変数名は数字で始めることはできず、英数字またはアンダースコアのみを含めることができます。
- 変数名は大文字と小文字を区別します。つまり、`my_var`は`my_Var`や`MY_VAR`とは異なります。
- 最後に、変数名は`snake_case`を使って書くのが一般的な慣例で、各スペースはアンダースコアに置き換えられ、単語は小文字で書かれます。

`print()`の呼び出しを両方とも削除し、`alphabet`という名前の別の変数を宣言してください。この変数に文字列`'abcdefghijklmnopqrstuvwxyz'`を代入してください。

# --hints--

コードに`print(type(text))`が含まれていないことを確認してください。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

コードに`print(type(shift))`が含まれていないことを確認してください。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/)
```

`alphabet`という名前の変数を宣言してください。変数名は行の先頭に置くよう注意してください。

```js
assert.match(code, /^alphabet\s*=/m)
```

`'abcdefghijklmnopqrstuvwxyz'`変数に文字列`alphabet`を代入してください。文字列はシングルクォートか二重引用符で囲むことを忘れないでください。

```js
assert.match(code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*(#.*)?$/m)
```

コードに無効な構文や無効なインデントが含まれています。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
