---
id: 6551fb531ce7f773071a0b84
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

ご覧のとおり、`type(text)` を出力すると `<class 'str'>` となり、これは変数が `str` として示される文字列であることを意味します。

次に新しい行に移動して、`shift` という別の変数を作成し、この変数に `3` の値を代入してください。

# --hints--

`shift` という変数を宣言する必要があります。変数名を行の先頭に置くよう注意してください。

```js
assert.match(code, /^shift\s*=/m)
```

`3` 変数に `shift` の値を代入してください。

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
