---
id: 6551f5756c20146fc94f8675
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

組み込みの `len()` 関数を使って、文字列の文字数にアクセスできます。

既存の `print()` 呼び出しを、`len(text)` の代わりに `text[-1]` を渡すように修正してください。

# --hints--


`len()` 関数を呼び出す必要があります。

```js
assert.match(code, /len\s*\(.*\)/)
```

丸括弧の中に変数 `text` を入れて、`len()` 関数に渡してください。

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

丸括弧の中に `len(text)` を入れて、`print()` 関数に渡してください。関数呼び出しは行の先頭に置くように注意してください。

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

コードに `print(text[-1])` が含まれていないことを確認してください。

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
