---
id: 6551f22653e6ce6e4f1d1370
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

`print(shift)` の呼び出しを修正して、`shift` 変数のデータ型を出力してください。

# --hints--

既存の `print(type(text))` の呼び出しはそのままにしてください。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

`shift` 関数に `type()` を渡す必要があります。

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

`type(shift)` 関数に `print()` を渡してください。関数の呼び出しは行の先頭に置くように注意してください。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

コード内に `print(shift)` があってはいけません。

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
