---
id: 6551f1e8737dbe6e0cb08ea4
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

そして、新しい変数を出力してください。

# --hints--

既存の`print(type(text))`の呼び出しは削除しないでください。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

`print()`変数を渡して`shift`関数を呼び出す必要があります。関数の呼び出しは行の先頭に置くように注意してください。

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
