---
id: 6551f4259d66026ee28cb205
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

文字列の末尾から文字にアクセスすることもできます。最後の文字のインデックスは`-1`で、その一つ前は`-2`、というように続きます。

では、既存の`print()`の呼び出しを修正して、文字列の最後の文字を出力してください。

# --hints--

`print()`関数は引き続き呼び出してください。

```js
assert.match(code, /print\s*\(.*\)/)
```

丸括弧の間に`text[-1]`を入れて`print()`関数に渡してください。関数の呼び出しは行の先頭に置くように注意してください。

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

コードに無効な構文やインデントの誤りがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
