---
id: 6551ef504d91936d2d4e54f8
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

*引数*とは、関数が呼び出されるときに、開き丸括弧と閉じ丸括弧の間に渡されるオブジェクトや式のことです。

```py
greet = 'Hello!'
print(greet)
```

上の例のコードは、`'Hello!'`に引数として渡された変数`greet`の値である文字列`print()`を出力します。

`text`関数に`text`変数を引数として渡して、`print()`変数を画面に出力してください。

# --hints--

`text`を`print()`関数に渡すには、この変数名を開き丸括弧と閉じ丸括弧の間に追加してください。関数呼び出しは行の先頭に置くように注意しましょう。

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
