---
id: 6551ecf597038c36637f3759
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

変数はPythonやあらゆるプログラミング言語において不可欠な要素です。変数とは、オブジェクトを参照または指し示す名前のことです。変数は、代入演算子 `=` の左側に変数名を書き、右側にその変数に代入する値を指定することで宣言できます。

```py
variable_name = value
```

`number` という名前の変数を作成し、新しい変数に値 `5` を代入してください。

# --hints--

`number` という名前の変数を宣言してください。変数名は行の先頭に置くよう注意してください。

```js
assert.match(code, /^number\s*=/m)
```

`5` の値を `number` 変数に代入してください。

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

コードに無効な構文やインデントの誤りがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
