---
id: 6551f6f9ec6db270eb83f3db
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

`11` がターミナルに出力されているのは、`'Hello World'` に11文字が含まれているからです。

もう一つ便利な組み込み関数は `type()` で、これは変数のデータ型を返します。`print()` の呼び出しを修正して、`text` のデータ型を出力してください。

# --hints--

`type()` 関数を呼び出す必要があります。

```js
assert.match(code, /type\s*\(.*\)/)
```

丸括弧の間に `text` を入れて、`type()` 関数に渡してください。

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

丸括弧の間に `type(text)` を入れて、`print()` 関数に渡してください。関数呼び出しは行の先頭に置くように注意してください。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

コードに `print(len(text))` が含まれていてはいけません。

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
