---
id: 6551f3bfed58796ebb1268dc
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

文字列の各文字は数値のインデックスで参照できます。インデックスのカウントはゼロから始まります。つまり、文字列の最初の文字はインデックス `0` を持ちます。例えば、文字列 `'Hello World'` では、`'H'` はインデックス `0` にあり、`'e'` はインデックス `1` にあり、以下同様です。

文字列の各文字にはブラケット記法を使ってアクセスできます。変数名の後に大括弧を付け、その括弧の間に文字のインデックスを入れてください。

```py
text = 'Hello World'
r = text[8]
```

今度は `text` を出力する代わりに、インデックス `6` の文字だけを出力してください。

# --hints--

`print()` 関数は引き続き呼び出してください。

```js
assert.match(code, /print\s*\(.*\)/)
```

丸括弧の間に `text[6]` を入れて `print()` 関数に渡してください。関数呼び出しは行の先頭に置くように注意してください。

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
