---
id: 65520e6f2b9678799977f24d
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

最後の `print()` の呼び出しを削除してください。それから、`text[0]` の代わりに `text[0].lower()` を引数として `.find()` の呼び出しに渡し、出力を確認してください。

# --hints--

コードに `print(text.lower())` が含まれていないことを確認してください。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

`index` 変数は引き続き出力してください。

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

`alphabet.find(text[0])` の呼び出しを `text[0].lower()` を引数に使うように更新してください。メソッド呼び出しは行の先頭に置くことに注意してください。

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

コードに無効な構文やインデントの誤りがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
