---
id: 65522255d5b9cd80f335c6f2
title: ステップ 29
challengeType: 20
dashedName: step-29
---

# --description--

ループ本体の最後で、`new_index` という変数を宣言し、この変数に `index + shift` の値を代入してください。

# --hints--

`new_index` ループの中で `for` という変数を宣言する必要があります。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

`index + shift` ループ本体の最後で、新しい変数に `for` を代入してください。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
