---
id: 6552008c0d9d9075cbec9772
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

`.find()` メソッドを使って、メッセージ内の各文字のアルファベット上の位置を見つけます。メソッドは関数に似ていますが、オブジェクトに属しています。

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

上の例では、`.find()` メソッドが `sentence`（検索対象の文字列）に対して *呼び出され*、`'r'`（位置を特定する文字）が引数として渡されています。`sentence.find('r')` の呼び出しは `4` を返します。これは `'r'` の中で `sentence` が最初に現れるインデックスです。

コードの最後で、`.find()` に対して `alphabet` を呼び出し、`'z'` をメソッドの引数として渡してください。

# --hints--

`.find()` メソッドを呼び出す必要があります。

```js
assert.match(code, /find\s*\(.*\)/)
```

`.find()` 変数に対して `alphabet` メソッドを呼び出す必要があります。

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


`.find()` 変数に対して `alphabet` メソッドを呼び出し、`'z'` をメソッドに渡す必要があります。メソッド呼び出しは行の先頭に置くことに注意してください。

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


コードに無効な構文やインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
