---
id: 655209a4a27dd37873c4cac3
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

`.find()` は文字列内で一致する文字のインデックスを返します。文字が見つからない場合は `-1` を返します。ご覧の通り、`text` の最初の文字である大文字の `'H'` は見つかりません。なぜなら `alphabet` は小文字のみを含んでいるからです。

文字列を小文字に変換するには `.lower()` メソッドを使います。もう一度 `print()` を呼び出して `text.lower()` を出力し、結果を確認してください。

# --hints--

コードにはまだ `print(index)` があるはずです。関数呼び出しが行の先頭にあることに注意してください。

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

`text.lower()` を出力してください。関数呼び出しが行の先頭にあることに注意してください。

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

コードに無効な構文やインデントの誤りがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
