---
id: 65521fc818947e800bffe48a
title: ステップ 26
challengeType: 20
dashedName: step-26
---

# --description--

`for` ループの中で、現在の文字を出力する前に、`index` という変数を宣言し、`alphabet.find(char)` が返す値をこの変数に代入してください。

# --hints--

`index` ループの最初に、`for` という新しい変数を宣言する必要があります。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

新しい `alphabet.find(char)` 変数に `index` を代入してください。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
