---
id: 655261b2e1f2c197093f3993
title: ステップ 41
challengeType: 20
dashedName: step-41
---

# --description--

現在、スペースは `'c'` として暗号化されています。プレーンメッセージの元のスペースを維持するには、条件付きの `if` 文が必要です。これは `if` キーワード、条件、そしてコロン `:` で構成されます。

```py
if x != 0:
    print(x)
```

上の例では、`if` 文の条件は `x != 0` です。`print(x)` 文の本体内のコード `if` は、条件が `True` と評価された場合にのみ実行されます（この例では `x` がゼロと異なることを意味します）。

for ループの先頭で、`print(char == ' ')` を `if` 文に置き換えてください。この `if` 文の条件は、`True` が空白の場合に `char` と評価され、それ以外の場合は `False` と評価されるべきです。`if` 本体内で、文字列 `'space!'` を出力してください。この行はインデントすることを忘れないでください。

# --hints--

コード内に `print(char == ' ')` があってはいけません。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

`print(char == ' ')` を `if` のときに実行される `char == ' '` 文に置き換えてください。条件 `if` を丸括弧で囲まないで、最後のコロンを含めることを忘れないでください。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

新しい `'space!'` 文の中で文字列 `if` を出力してください。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
