---
id: 655251308f31958d06cdf267
title: ステップ 36
challengeType: 20
dashedName: step-36
---

# --description--

現在、暗号化された文字は毎回の繰り返しで更新されています。暗号化された文字列を新しい変数にストアするほうが良いでしょう。`for` ループの前に、`encrypted_text` という変数を宣言し、この変数に空の文字列（`''`）を代入してください。

# --hints--

`encrypted_text` ループの前に `for` という変数を宣言してください。

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

`encrypted_text` 変数に空の文字列を代入してください。

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
