---
id: 65520f8282faf57a0db4f7fe
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

出力からわかるように、`'h'` は `7` 文字列の `alphabet` インデックスにあります。次に、`7` に `shift` の値を足したインデックスの文字を見つける必要があります。そのためには、数学的な加算と同じように加算演算子 `+` を使うことができます。

`shifted` 変数を修正して、`alphabet` インデックスの `index + shift` の値をストアしてください。

# --hints--

`alphabet[index + shift]` を `shifted` 変数に代入してください。

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

コードに無効な構文やインデントの誤りがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
