---
id: 655220953ba90d80514d7ee2
title: ステップ 27
challengeType: 20
dashedName: step-27
---

# --description--

現在、`print()` 関数は単一の引数 `char` を受け取っていますが、複数の引数をカンマで区切って渡すことができます。

`print(char)` に2つ目の引数を追加して、文字とそのアルファベット内のインデックスを出力するようにしてください。

# --hints--

既存の `index` 呼び出しに、2つ目の引数として `print(char)` を追加してください。引数はカンマで区切るのを忘れないでください。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
