---
id: 65524d2a1a253b8bb5197ae2
title: ステップ 28
challengeType: 20
dashedName: step-28
---

# --description--

`find` は大文字の文字に対して再び `-1` を返していますし、スペース文字に対しても同様です。スペースについては後で対応します。

今は、`text` を繰り返す代わりに、`for` のループを `text.lower()` を繰り返すように変更してください。

# --hints--

`for` のループを `text.lower()` ではなく `text` を繰り返すように変更してください。

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
