---
id: 65524b790ba8558a2f1c9fe5
title: ステップ 35
challengeType: 20
dashedName: step-35
---

# --description--

出力を少しきれいにしてください。`print(char, index)` を削除し、最後の `print()` の呼び出しを `print('char:', char, 'new char:', new_char)` に変更してください。

# --hints--

コードに `print(char, index)` が含まれていないことを確認してください。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse(/print\s*\(\s*char\s*,\s*index\s*\)/.test(commentless_code))
```

`print(new_char)` を `print('char:', char, 'new char:', new_char)` に変更してください。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse(/print\s*\(\s*new_char\s*\)/.test(commentless_code));
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')new\schar:\2\s*,\s*new_char\s*\)\s*$/));
```

コードに無効な構文や無効なインデントが含まれています。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
    new_char = alphabet[new_index]
    print(new_char)
--fcc-editable-region--
```
