---
id: 655208d59b131e7816f18c96
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

`print()` 関数はコンソールに出力するだけですが、関数やメソッドはコード内で使える戻り値を持つことができます。

今、`alphabet.find(text[0])` を `index` という名前の変数に代入してください。こうすることで、`index` は `alphabet.find(text[0])` が返す値をストアします。

# --hints--

`index` という名前の変数を宣言してください。変数名は行の先頭に置くよう注意してください。

```js
assert.match(code, /^index\s*=/m)
```

`alphabet.find(text[0])` が返す値を `index` 変数に代入してください。

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

コードに無効な構文やインデントの誤りがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
