---
id: 65525e359ca28d938baa82c5
title: ステップ 39
challengeType: 20
dashedName: step-39
---

# --description--

加算代入演算子を使うことで、`a = a + b` と同じ効果を得られます。

```py
a += b
```

加算代入演算子は、変数に値を加えて、その結果をその変数に代入できます。

`+=` 演算子を使って、値を加えながら同時に `encrypted_text` に代入してください。

# --hints--

ループの中で `alphabet[new_index]` の現在の値に `encrypted_text` を加えるには、加算代入演算子を使うことをお勧めします。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
