---
id: 6553efd6ada3f42aa2d75448
title: ステップ 53
challengeType: 20
dashedName: step-53
---

# --description--

もう一度出力が表示されるはずです。この方法は動作しますが、コードの再利用性を大幅に向上させるわけではありません。関数を繰り返し呼び出しても同じ結果になります。しかし、関数は*パラメータ*を使って宣言することで、多様性やカスタマイズが可能になります。

```py
def function_name(param_1, param_2):
    <code>
```

パラメータは関数内で使える変数です。関数は異なる数のパラメータで宣言できます。上の例では、`param_1` と `param_2` がパラメータです。

関数の宣言を修正して、`message` と `offset` という2つのパラメータを受け取るようにしてください。

その後、ターミナルにエラーが表示されます。次のステップでその解決方法を説明します。

# --hints--

`caesar` 関数は `message` と `offset` をパラメータとして受け取る必要があります。パラメータはカンマで区切ることを忘れないでください。

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
