---
id: 6553ed69ece88d29594748aa
title: ステップ 52
challengeType: 20
dashedName: step-52
---

# --description--

関数を実行するには、その名前の後に丸括弧を付けて呼び出す（またはインボークする）必要があります。例えば次のようにします。

```py
function_name()
```

コードの最後に、`caesar` 関数を呼び出してください。インデントに注意してください。

# --hints--

`caesar` 関数を呼び出す必要があります。関数呼び出しは行の先頭に書いてください。

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
