---
id: 6553995f412dd8122ed38e4a
title: ステップ 43
challengeType: 20
dashedName: step-43
---

# --description--

条件文には `else` 節を付けることもできます。この節は `if` 文の末尾に追加して、`if` 文の条件が false の場合に代わりのコードを実行できます。

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

出力を見ると、ループの繰り返しがスペースに達したときに、暗号化された文字列にスペースが追加されます。その後、`if` ブロックの外のコードが実行され、暗号化された文字列に `c` が追加されます。

これを修正するには、`else` の後に `encrypted_text += char` 節を追加し、`print()` の呼び出しを除くすべての後続のコード行をインデントしてください。

# --hints--

`else` 節を作成してください。最後のコロンを忘れないでください。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

`else` 節の後のコード行を、`print()` の呼び出しを除いてインデントしてください。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
