---
id: 6553a44b1801991847d8cc69
title: ステップ 47
challengeType: 20
dashedName: step-47
---

# --description--

次に、`print()` の呼び出しを修正して `'encrypted text:', encrypted_text` を出力し、`for` のループの外に置いて、暗号化された文字列が1回だけ出力されるようにしてください。

# --hints--

forループの後に `'encrypted text:', encrypted_text` を出力してください。

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/m) })
```

コード内に `print('char:', char, 'encrypted text:', encrypted_text)` を含めないでください。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.notMatch(commentless_code, /print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
