---
id: 6553f4f66099802c6ae94613
title: ステップ 55
challengeType: 20
dashedName: step-55
---

# --description--

現在、コードは`TypeError`を発生させています。なぜなら`caesar`関数は2つのパラメータ（`message`と`offset`）で定義されているため、2つの引数で呼び出されることを期待しているからです。

必要な引数なしで`caesar()`を呼び出すと、コードの実行がストップします。

`message`と`offset`に値を与えるために、`text`と`shift`を`caesar`関数の呼び出しに引数として渡してください。

# --hints--

関数呼び出しの丸括弧内に`text`と`shift`を含めて、引数として渡してください。引数はカンマで区切るのを忘れないでください。

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
