---
id: 6553a7d8d05cbb1ae335a665
title: ステップ 45
challengeType: 20
dashedName: step-45
---

# --description--

ループが文字 `Z` に達したとき、合計 `index + shift` は文字列 `alphabet` の最後のインデックスを超えています。したがって、`alphabet[new_index]` は無効なインデックスを使おうとしており、それが `IndexError` のスローを引き起こします。

ターミナルの出力がエラーがスローされる直前の最後の `print` である `Z` の直前の空白で停止していることに気づけます。

この場合、モジュラス演算子（`%`）を使って2つの数の割り算の余りを返すことができます。例えば：`5 % 2` は `1` と等しいです。なぜなら5を2で割ると商が2で余りが1だからです。

`index + shift` を丸括弧で囲み、式を `26` でモジュラスしてください。`26` は `alphabet` の長さです。

# --hints--

`new_index = (index + shift) % 26` を `else` 文に含める必要があります。

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
