---
id: 6553f6086add4b2cbb99fd78
title: ステップ 56
challengeType: 20
dashedName: step-56
---

# --description--

コードの一番下で、既存の`caesar(text, shift)`の呼び出しの後に、もう一度関数を呼び出してください。

今回は、`text`変数と整数の`13`を引数として渡してください。

# --hints--

もう一度関数を呼び出す必要があります。今回は`text`と`13`を引数として渡してください。

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

既存の関数呼び出しはそのままにしてください。

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
