---
id: 6553980e0527fa115c705646
title: ステップ 42
challengeType: 20
dashedName: step-42
---

# --description--

今度は、`'space!'` を出力する代わりに、加算代入演算子を使って、現在 `char` に格納されている値に `encrypted_text`（スペース）を加えてください。

# --hints--

コードに `print('space!')` を含めてはいけません。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

新しい `+=` 文の中で、`char` 演算子を使って `encrypted_text` の現在の値に `if` を加えてください。

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

コードに無効な構文やインデントの誤りがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
