---
id: 6553a572f7a65718f1e42e18
title: ステップ 49
challengeType: 20
dashedName: step-49
---

# --description--

関数とは、本質的に再利用可能なコードのブロックです。すでに `print()`、`find()`、`len()` のような組み込み関数に触れています。しかし、このようにカスタム関数を定義することもできます。

```py
def function_name():
    <code>
```

関数の宣言は `def` キーワードで始まり、その後に関数名（有効な変数名）と丸括弧のペアが続きます。宣言はコロンで終わります。

`shift` 変数のすぐ後に、`caesar` という名前の関数を宣言し、新しい関数の本体となるすべての行をインデントしてください。

# --hints--

新しい関数を宣言するには `def` キーワードを使う必要があります。

```js
assert.match(code, /^def\s+/m)
```

`caesar` キーワードの後に関数名として `def` を書いてください。`def` の後にスペースを入れるのを忘れないでください。

```js
assert.match(code, /^def\s+caesar/m)
```

関数名の後に丸括弧のペアを追加してください。最後のコロンも忘れないでください。

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

`shift = 3` の後のすべての行をインデントして、新しい関数の本体にしてください。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
