---
id: 6553f8c570f9982e013a8886
title: ステップ 57
challengeType: 20
dashedName: step-57
---

# --description--

現在、すべての文字は指定されたオフセットに応じて常に同じ文字で暗号化されています。もし文字ごとにオフセットが異なっていたらどうでしょうか？それはずっと解読が難しくなります。このアルゴリズムはヴィジュネル暗号と呼ばれ、各文字のオフセットは「キー」と呼ばれる別のテキストによって決まります。

シーザー暗号をヴィジュネル暗号に変換し始めるために、2つの関数呼び出しを削除してください。

# --hints--

2つの`caesar()`関数呼び出しを削除してください。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
