---
id: 6553a755879b131a445e664c
title: ステップ 44
challengeType: 20
dashedName: step-44
---

# --description--

文字列 `'Hello Zaira'` を `text` 変数に代入して、ターミナルで何が起きるか試してみてください。

`string index out of range` 例外が表示されます。心配しないでください、すぐに修正方法がわかります！

# --hints--

`text` 変数があるはずです。

```js
assert.match(code, /^text\s*=/m)
```

文字列 `'Hello Zaira'` を `text` 変数に代入してください。

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
