---
id: 6554e0adc7bb193cbfdb36d5
title: ステップ 73
challengeType: 20
dashedName: step-73
---

# --description--

関数が文字列を元のプレーンテキストに復号できるか確認してください。

`decryption` という別の変数を宣言し、`vigenere(encryption, custom_key, -1)` を代入してください。

# --hints--

`vigenere`、`encryption`、`custom_key` を引数として渡して `-1` を呼び出す必要があります。

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/) })
```

`decryption` という変数を宣言してください。

```js
({ test: () => assert(__userGlobals.has("decryption")) })
```

`decryption` 変数には `vigenere(encryption, custom_key, -1)` の値を代入してください。

```js
({ test: () => assert.match(code, /^decryption\s*=\s*vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--    
encryption = vigenere(text, custom_key, 1)
print(encryption)

--fcc-editable-region--
```
