---
id: 655491bd5b98b813fa5bedca
title: ステップ 58
challengeType: 20
dashedName: step-58
---

# --description--

関数宣言を修正してください：関数名を`vigenere`に、2番目のパラメータを`key`に変更してください。

# --hints--

関数名を`vigenere`に変更してください。

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

`vigenere`関数は`message`と`key`をパラメータとして受け取る必要があります。

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
