---
id: 6554ad2463b8892748f8efdd
title: ステップ 67
challengeType: 20
dashedName: step-67
---

# --description--

現在、関数は文字列を出力していますが、これらの値はコードの他の部分で何らかの処理に使うことができません。

そのためには、`return` 文を使う必要があります。

```python
def foo():
    return 'spam'
```

関数が返すべき値の後にスペースを入れて `return` と書いてください。`return` 文が見つかると、その値が返されて関数の実行はストップし、関数呼び出しの後のコードの次の行に進みます。上の例では、`foo` 関数は文字列 `'spam'` を返しています。

関数から2つの `print()` 呼び出しを削除して、`encrypted_text` を返してください。

# --hints--

関数から2つの `print()` 呼び出しを削除してください。

```js
({ 
  test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/print\s*\(\s*("|')plain\stext:\1\s*,\s*message\s*\)/.test(commentless_code))
    assert.isFalse(/print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/.test(commentless_code))
  }
})
```

関数は `encrypted_text` を返すべきです。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*/);
    const regex = /return\s+encrypted_text/
    assert(function_body.match(regex));
    assert.notMatch(block_body, regex);
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
