---
id: 6554930320d70414e7b6acc6
title: ステップ 59
challengeType: 20
dashedName: step-59
---

# --description--

`shift` 変数を削除してください。その後、`custom_key` という新しい変数を宣言し、文字列 `'python'` の値をこの変数に代入してください。

# --hints--

`shift` 変数とその値を削除する必要があります。

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

`custom_key` という変数を宣言してください。

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

文字列 `'python'` を `custom_key` 変数に代入してください。

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
