---
id: 6554d25dc5ceaa354307a77e
title: ステップ 70
challengeType: 20
dashedName: step-70
---

# --description--

暗号化と復号化は逆の処理であり、関数に少し手を加えれば両方を実行できます。

関数定義に `direction` という3つ目のパラメータを追加してください。また、コンソールのエラーを避けるために、コードの最後の2行をコメントアウトしてください。

# --hints--

コードの最後の2行をコメントアウトにしてください。各行の先頭に `#` を付けてください。

```js
({ test: () => assert.match(code, /#\s*encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)\s*#\s*print\s*\(\s*encryption\s*\)/) })
```

関数の3つ目のパラメータとして `direction` を追加してください。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
encryption = vigenere(text, custom_key)
print(encryption)
--fcc-editable-region--
```
