---
id: 6554d0332949b133a0b35eaa
title: ステップ 68
challengeType: 20
dashedName: step-68
---

# --description--

`text` と `custom_key` を引数として渡して関数を呼び出してください。関数呼び出しの戻り値を `encryption` という変数にストアしてください。

# --hints--

`vigenere` を呼び出す際に `text` と `custom_key` を引数として渡す必要があります。

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

`encryption` という変数が必要です。

```js
({ test: () => assert(__userGlobals.has("encryption")) })
```

関数呼び出しの結果を `encryption` という変数に代入してください。

```js
({ test: () => assert.match(code, /encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--

--fcc-editable-region--
```
