---
id: 6554a88d5af937226f4a9121
title: ステップ 65
challengeType: 20
dashedName: step-65
---

# --description--

`.index()` メソッドは `.find()` メソッドと同じですが、部分文字列が見つからない場合は `ValueError` 例外をスローします。

`ValueError` は、適切な型だが不適切な値の引数が関数に渡されたときに発生する組み込みの例外です。

`key_index` をインクリメントした後、`offset` という名前の変数を宣言してください。`key_char` の中で `alphabet` が持つインデックスを見つけて `offset` に代入してください。インデックスを見つけるには `.index()` を使います。

# --hints--

`offset` という名前の変数を宣言してください。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=/));
```

`offset` 変数には `alphabet.index(key_char)` の値を格納してください。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=\s*alphabet\.index\s*\(\s*key_char\s*\)/));
```

コードに無効な構文やインデントの誤りがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
            
--fcc-editable-region--            
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
