---
id: 65549f90cf78131c96ebcf28
title: ステップ 61
challengeType: 20
dashedName: step-61
---

# --description--

コードを書くときは、読みやすさが重要です。コメントは効果的なメモとして機能し、コードの背後にあるロジックを説明します。時間が経ってからプロジェクトに戻る際に役立ち、同僚がコードを理解するのにも役立ちます。

Pythonでは、`#`を使ってコメントを書くことができます。`#`の後に続くものは実行されません。

`if`の文の前に、「`Append space to the message`」というコメントを追加してください。

# --hints--

「`Append space to the message`」というコメントを追加してください。先頭に`#`の文字を忘れずに付けてください。

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
