---
id: 6555e1bfcd374e18c6be8e58
title: ステップ 91
challengeType: 20
dashedName: step-91
---

# --description--

Python には文字列を簡単にフォーマットする方法があります。*f-strings* を使うと、文字列の中に値を補間できます。

補間とは、プログラム実行時に指定した値で置き換えられるプレースホルダーを書くことです。例えば、`'Encrypted text: ' + text` と同じ結果を `f'Encrypted text: {text}'` で得ることができます。f-string を作るには、引用符の前に `f` を付けて、補間したい変数や式を波括弧の間に書きます。

最初の `print()` の呼び出しを f-string を使うように変更してください。

# --hints--

`print('Encrypted text: ' + text)` の呼び出しを変更して、f-string を使い同じ結果を出力するようにしてください。

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
