---
id: 6555201d9b7fc917399f9f0b
title: ステップ 80
challengeType: 20
dashedName: step-80
---

# --description--

`.isalpha()` メソッドは、呼び出された文字列のすべての文字が文字である場合に `True` を返します。例えば、以下のコードは `True` を返します。

```py
'freeCodeCamp'.isalpha()
# True
```

条件 `char == ' '` を削除し、`.isalpha()` に対して `char` メソッドを呼び出すように置き換えてください。

# --hints--

コード内に `if char.isalpha():` があるはずです。

```js
({ test: () => assert.match(code, /if\s+char\.isalpha\s*\(\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--    
        # Append space to the message
        if char == ' ':
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
