---
id: 6555231eac4d9f19bd3d44b5
title: ステップ 83
challengeType: 20
dashedName: step-83
---

# --description--

`pass` キーワードは将来のコードのプレースホルダーとして使えます。コードには何の影響もありませんが、不完全なコードで発生するエラーを防ぐことができます。

```py
def foo():
    pass
```

`vigenere` を `1` で暗号化し、`-1` で復号する呼び出しは問題ありませんが、少しわかりにくいかもしれません。`encrypt` という新しい関数を作成し、`message` と `key` のパラメータを受け取り、関数の本体には `pass` を使ってください。


# --hints--

`encrypt` という新しい関数を宣言してください。関数の本体には `pass` キーワードを使うことを忘れないでください。

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

`encrypt` 関数は `message` と `key` をパラメータとして受け取る必要があります。順番が重要です。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
