---
id: 6555d6e12c95701172b55709
title: ステップ 77
challengeType: 20
dashedName: step-77
---

# --description--

関数はデフォルト引数を使って呼び出すことができます。デフォルト引数とは、引数が渡されなかった場合に関数が取るべき値を示します。例えば、以下の関数は3つの引数を受け取りますが、2つの引数だけを渡して呼び出すこともできます。3つ目の引数はデフォルトで指定された値を取ります。

```py
def foo(a, b, c=value):
    <code>
```

`vigenere` 関数の `direction` パラメータに `1` のデフォルト値を設定してください。

# --hints--

`direction` 関数の `vigenere` パラメータは `1` をデフォルト値にする必要があります。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction=1)'
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

--fcc-editable-region--
def vigenere(message, key, direction):
--fcc-editable-region--
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
