---
id: 655619327c7b364166f8dd6f
title: ステップ 46
challengeType: 20
dashedName: step-46
---

# --description--

`alphabet` 文字列に数字や特殊文字などの追加の文字を組み込みたい場合は、モジュラス演算の右辺を手動で変更する必要があります。

この問題を避けるには、`26` を `len(alphabet)` に置き換えてください。

# --hints--

`new_index` の値を変更し、`26` を `len(alphabet)` に置き換える必要があります。

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
  assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/));
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % 26
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
