---
id: 655a2a7210094920069b117c
title: ステップ 40
challengeType: 20
dashedName: step-40
---

# --description--

比較演算子は、2つのオブジェクトの値を基に比較することを可能にします。比較したいオブジェクトの間に比較演算子を置くことで使えます。
これらは式の真偽に応じて、*ブール値*、つまり `True` または `False` を返します。

Pythonには以下の比較演算子があります。

<table>
  <thead>
    <tr>
      <th>演算子</th>
      <th>意味</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>等しい</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>等しくない</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>より大きい</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>より小さい</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>以上</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>以下</td>
    </tr>
  </tbody>
</table>

ループ本体の最初で、`char` と空白（`' '`）を比較した結果を出力してください。その際、等値演算子 `==` を使います。

# --hints--

`char` と空白を等値演算子で比較し、`for` ループ内で使う必要があります。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

`char` と空白を比較した結果を `for` ループ内で出力してください。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

ループの最初で `char` と空白を比較した結果を出力してください。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
