---
id: 65688b5a1655a7a6caede847
title: ステップ 31
challengeType: 20
dashedName: step-31
---

# --description--

前のステップで行ったように文字列の個々の文字を変更しようとすると、`TypeError`が発生します。これはコード内で不適切な型のオブジェクトが使われた場合に起こります。

エラーメッセージからわかるように、文字列は不変であるため、アイテムの代入はサポートされていません。しかし、変数には別の文字列を再代入できます。

```py
message = 'Hello World'
message = 'Hello there!'
```

`text[0]`の行を削除し、`text`に文字列`'Albatross'`を再代入してください。

# --hints--

コード内に`text[0] = 'C'`があってはいけません。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

`text`に文字列`'Albatross'`を再代入してください。`text = 'Hello World'`は変更しないでください。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
