---
id: 65687d2f8c7ee27b0446cef3
title: ステップ 24
challengeType: 20
dashedName: step-24
---

# --description--

各イテレーションで実行されるコードは、`:`の後に配置され、ループの本体を構成します。このコードはインデントする必要があります。Pythonでは、インデントレベルごとに4つのスペースを使うことをお勧めします。このインデントされたレベルがコードブロックです。

```py
for i in text:
    print(i)
```

Pythonはインデントによってコードブロックを示します。行の末尾にコロンがあると、新しいインデントされたコードブロックが続くことを示します。

したがって、最後のコロンの後にインデントされたブロックが見つからない場合、コードの実行はストップし、`IndentationError`がスローされます。このコードは出力を表示せず、代わりに`IndentationError`を発生させます：

```py
for i in text:
print(i)
```

`for`ループに`print(i)`の呼び出しを追加して本体を作成してください。ループ本体はインデントすることを忘れないでください。

# --hints--

`print(i)`を`for`ループ本体に追加してください。インデントに注意してください。

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

コードに無効な構文および／または無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
