---
id: 656886ec4e4d9b96ded3c565
title: ステップ 25
challengeType: 20
dashedName: step-25
---

# --description--

繰り返し変数には有効な名前なら何でも使えますが、意味のある名前を付けると便利です。

`i` 変数の名前を `char` に変更してください。

# --hints--

`i` を `char` に置き換えてください。

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
