---
id: 65688a50e6c998a21d8e41d3
title: ステップ 30
challengeType: 20
dashedName: step-30
---

# --description--

文字列は不変であり、一度作成されると変更できません。たとえば、次のコードが `my_string` の値を文字列 `'train'` に変更すると考えるかもしれませんが、これは正しくありません。

```py
my_string = 'brain'
my_string[0] = 't'
```

ブラケット記法を使って `text` の最初の文字にアクセスし、お好きな文字に変更しようとしてみてください。出力が消え、エラーが表示されるのが確認できます。

# --hints--

コード内にまだ `text = 'Hello World'` があるはずです。

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

`text` の最初の文字には `text[0]` でアクセスしてください。

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

`=` 演算子を使って、お好きな文字を `text[0]` に代入してください。文字はシングルクォートか二重引用符で囲むのを忘れないでください。

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
