---
id: 6568b7a40c7ba0ccbc4b4425
title: ステップ 32
challengeType: 20
dashedName: step-32
---

# --description--

ご覧のとおり、再代入された文字列の各文字がループ内で出力されています。

`text` の再代入を削除して、元の文字列に戻してください。

# --hints--

コード内に `text = 'Albatross'` があってはいけません。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*=\s*("|')Albatross\1/.test(commentless_code))
```

`text` という変数は、値 `'Hello World'` のままである必要があります。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /text\s*=\s*("|')Hello World\1/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
