---
id: 6568c86dc9193000d11ca5e0
title: ステップ 38
challengeType: 20
dashedName: step-38
---

# --description--

`alphabet[new_index]` に `encrypted_text` を代入する代わりに、この変数に `encrypted_text` と `alphabet[new_index]` の現在の値の合計を代入してください。

# --hints--

`encrypted_text + alphabet[new_index]` 変数に `encrypted_text` を代入する必要があります。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*=\s*encrypted_text\s*\+\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
