---
id: 659526f3e5424b3eac4b6aa7
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

次に、`index` 変数をコンソールに出力してください。

# --hints--

`index` 変数を出力してください。関数の呼び出しは行の先頭に置くように注意してください。

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
