---
id: 659530ca823d6c47f61daa8a
title: ステップ 20
challengeType: 20
dashedName: step-20
---

# --description--

`shifted` 変数を出力してください。

# --hints--

`shifted` 変数を出力してください。関数の呼び出しは行の先頭に置くよう注意してください。

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
