---
id: 65951b637ddec63611c482c0
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

組み込み関数の `print()` を使って、コードの出力をターミナルに表示できます。

関数は再利用可能なコードのブロックで、必要なときに呼び出してコードを実行できます。関数を呼び出すには、名前の隣に丸括弧を1組書くだけです。関数についてはすぐにもっと詳しく学びます。

まずは新しい行に移動して、`print()` 関数を空の呼び出しで追加してください。まだ何も出力されないはずです。

# --hints--

コードに `print()` があるはずです。関数呼び出しは行の先頭に置くように注意してください。

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
