---
id: 65952b70924edb449faf1950
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

`shifted` という名前の新しい変数を宣言してください。ブラケット記法を使って、`alphabet` のインデックス `index` にある値にアクセスし、その値を新しい変数に代入してください。

# --hints--

`shifted` という名前の変数を宣言する必要があります。変数名は行の先頭に置くように注意してください。

```js
assert.match(code, /^shifted\s*=/m)
```

`alphabet[index]` を `shifted` 変数に代入してください。

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

コードに無効な構文や無効なインデントがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
