---
id: 659db3e3670d3e712be82593
title: ステップ 22
challengeType: 20
dashedName: step-22
---

# --description--

`text` の各文字について、アルファベット内の文字を見つけてシフト後の文字を決定する作業を繰り返すのは面倒です。幸い、ループを使えばこれを簡単にできます。

まずは、`alphabet` 変数の宣言より下のすべてのコード行を削除してください。

# --hints--

コード内にまだ `text = 'Hello World'` があるはずです。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

コード内にまだ `shift = 3` があるはずです。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

コード内にまだ `alphabet = 'abcdefghijklmnopqrstuvwxyz'` があるはずです。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

`index` 変数とその値は削除してください。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

コード内に `print(index)` はないはずです。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

`shifted` 変数とその値は削除してください。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

コード内に `print(shifted)` はないはずです。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

コードに無効な構文やインデントの誤りがあります。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
