---
id: 6624e0ec0bf47017eed84b1f
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

最初に作る暗号は「シーザー」暗号と呼ばれます。具体的には、メッセージの各文字を取り、その文字のアルファベット上の位置を見つけ、アルファベットで3つ後ろにある文字を取得し、元の文字を新しい文字に置き換えます。

これを実装するには、前のステップで説明した`.find()`メソッドを使います。既存の`.find()`の呼び出しを修正し、`text[0]`の代わりに`'z'`を引数として渡してください。

# --hints--

既存の`alphabet.find('z')`の呼び出しを修正し、`text[0]`をメソッドに渡してください。

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
