---
id: 666064e915eba7aa1de03f6b
title: ステップ 51
challengeType: 20
dashedName: step-51
---

# --description--

今、`alphabet` 関数の外で `caesar` 変数を出力しようとしている行を削除して、エラーを修正してください。

# --hints--

`print(alphabet)` の行を削除してください。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--
print(alphabet)
--fcc-editable-region--
```
