---
id: 66605a64a9822f01ec5ae8ce
title: ステップ 50
challengeType: 20
dashedName: step-50
---

# --description--

Pythonでは、変数のスコープによって、その変数がどこでアクセスできるかが決まります。

- 関数の外で定義された変数はグローバルスコープを持ち、コードのどの部分からでもアクセスできます。

- 関数の中で定義された変数はその関数にローカルであり、関数の外からはアクセスできません。

これを実際に確認するために、コードの最後で`alphabet`変数を出力してみてください。そうすると`NameError`例外が発生します。

`alphabet`が定義されていないというエラーメッセージが表示されるはずです。これは`alphabet`が`caesar`関数の中で定義されており、関数の外からアクセスできないためです。

# --hints--

caesar関数の外で`alphabet`変数を出力してみてください。

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
