---
id: 694c84330cbe604821b9deef
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Il centro di controllo ha bisogno di un metodo per aggiungere nuovi membri all'elenco. Crea una funzione vuota chiamata `addCrewMember` che accetti due parametri, `crew` e `astronaut`.

# --hints--

Dovresti creare una funzione chiamata `addCrewMember`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addCrewMember);
```

La tua funzione `addCrewMember` dovrebbe avere i parametri `crew` e `astronaut`.

```js
const explorer = await __helpers.Explorer(code);
const { addCrewMember } = explorer.allFunctions;
assert.lengthOf(addCrewMember?.parameters, 2);
assert.equal(addCrewMember?.parameters?.[0]?.toString(), "crew");
assert.equal(addCrewMember?.parameters?.[1]?.toString(), "astronaut");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
