---
id: 694c40d86fe6a78d20b6559c
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop sfrutterai JavaScript per creare un elenco di astronauti per una missione spaziale.

Per cominciare, crea un array vuoto per il tuo elenco chiamato `squad`.

# --hints--

Dovresti creare una variabile `squad`.

```js
assert.isDefined(squad);
```

La tua variabile `squad` dovrebbe essere un array.

```js
assert.isArray(squad);
```

Il tuo array `squad` dovrebbe essere vuoto.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
