---
id: 694ce8609fb805ba21cebc02
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Il centro di controllo ha bisogno di un metodo per scambiare le posizioni di due membri dell'equipaggio in un elenco. Crea una funzione vuota chiamata `swapCrewMembers` con tre parametri:

- `crew`: un array di oggetti astronauta

- `fromIndex`: l'indice del primo astronauta da scambiare

- `toIndex`: l'indice del secondo astronauta da scambiare

# --hints--

Dovresti creare una funzione chiamata `swapCrewMembers`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.swapCrewMembers);
```

La tua funzione `swapCrewMembers` dovrebbe avere i parametri `crew`, `fromIndex` e `toIndex`.

```js
const explorer = await __helpers.Explorer(code);
const { swapCrewMembers } = explorer.allFunctions;
assert.lengthOf(swapCrewMembers?.parameters, 3);
assert.equal(swapCrewMembers?.parameters?.[0]?.toString(), "crew");
assert.equal(swapCrewMembers?.parameters?.[1]?.toString(), "fromIndex");
assert.equal(swapCrewMembers?.parameters?.[2]?.toString(), "toIndex");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

--fcc-editable-region--

--fcc-editable-region--
```
