---
id: 694c84330cbe604821b9deef
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Il centro di controllo ha bisogno di un metodo per aggiungere nuovi membri all'elenco. Crea una funzione vuota chiamata `addCrewMember` che accetti due parametri, `crew` e `astronaut`.

# --hints--

Dovresti creare una funzione chiamata `addCrewMember`.

```js
assert.isFunction(addCrewMember);
```

La tua funzione `addCrewMember` dovrebbe avere i parametri `crew` e `astronaut`.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
