---
id: 694c9c186c089f7bdff0f3c6
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

In generale, è buona pratica convalidare i tuoi input all'interno di una funzione. Usa un loop `for` per iterare su un array `crew` e controlla se qualche elemento ha già lo stesso `id` dell'input `astronaut`. Se trovi un duplicato, registra il messaggio di errore `console.log("Duplicate ID: " + astronaut.id)` e poi chiama `return` per uscire anticipatamente dalla funzione.

Ecco un esempio di questa tecnica comunemente usata:

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

La tua funzione dovrebbe usare un loop `for` per scansionare i membri `crew` esistenti (qualsiasi modulo `for (...)` valido va bene).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

All'interno del loop `for`, dovresti controllare la presenza di un duplicato `id`. Se ne trovi uno, registra `"Duplicate ID: " + astronaut.id"` nella console e poi chiama `return` per uscire dalla funzione.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Quando riceve input validi, la tua funzione `addCrewMember` non dovrebbe invocare un messaggio di errore `console.log()` né un'uscita anticipata `return`.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
