---
id: 694c84106e9023f289c03b84
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Il tuo equipaggio inizia con un solo astronauta comandante. Crea un oggetto chiamato `firstAstronaut` con le seguenti proprietà:

| Chiave | Valore |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

Dovresti creare una variabile `firstAstronaut`.

```js
assert.isDefined(firstAstronaut);
```

Il tuo `firstAstronaut` dovrebbe essere un oggetto.

```js
assert.isObject(firstAstronaut);
```

Il tuo oggetto `firstAstronaut` dovrebbe includere i valori corretti delle proprietà.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
