---
id: 694d8ddd106340237efb0c80
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Il metodo `splice()` può modificare gli array aggiungendo o rimuovendo elementi in qualsiasi posizione, incluso il mezzo. Poiché `splice()` restituisce un array contenente gli elementi rimossi, puoi usarlo per scambiare due elementi in un array senza modificare l'originale. Questo può essere fatto in una riga usando la seguente tecnica:

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

Ecco un esempio:

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

La tecnica, applicata nella terza riga, funziona così:

- `splice(3, 1, copyArray[1])` rimuove l'elemento all'indice `3` (`55`)

- Inserisce l'elemento da `copyArray[1]` (`97`) all'indice `3`

- `splice()` restituisce `[55]`, un array contenente l'elemento rimosso

- `[0]` estrae `55` da quell'array

- Quel valore (`55`) viene assegnato di nuovo a `copyArray[1]`, completando lo scambio

Usa questa tecnica con l'array `updatedCrew` per scambiare gli astronauti agli indici `fromIndex` e `toIndex`.

# --hints--

Dovresti eseguire lo scambio chiamando il metodo `splice()`.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

Dovresti chiamare `splice()` con l'array `updatedCrew` per eseguire lo scambio. Dovresti usare la tecnica in una riga.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

La tua funzione `swapCrewMembers` non dovrebbe modificare l'array di input originale `crew`.

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
