---
id: 6718d2d59337c822ecb697ff
title: 建置一個銀行帳戶管理程式
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。 

**使用者故事：**

1. 你應該定義一個名為 `BankAccount` 的類別並包含一個建構式。建構式應該將預設的 `balance` 設為 `0`，並初始化一個名為 `transactions` 的空陣列，用來儲存作為物件的交易記錄。

2. 每筆儲存在 `transactions` 陣列中的交易都應該是一個具有兩個屬性的物件：`type` 和 `amount`。`type` 屬性應該是 `deposit` 或 `withdraw`，而 `amount` 屬性應該是存入或提取的金額。

3. 你應該定義一個名為 `deposit` 的方法，該方法接受存款金額作為參數。當存款金額大於 `0` 時，應該：

   - 將一個新的物件推入 `transactions` 陣列，該物件的型別為 `deposit` 且包含存入的金額。
   - 更新餘額。
   - 傳回 `"Successfully deposited $[amount]. New balance: $[balance]"`。

4. 如果金額小於或相等於 `0`，`deposit` 方法應該傳回 `"Deposit amount must be greater than zero."`。

5. 你應該定義一個名為 `withdraw` 的方法，該方法接受一個金額作為參數。當要提取的金額大於 `0` 且小於或相等於當前餘額時，該方法應根據提款更新當前餘額。

   - 將一個新的物件推入 `transactions` 陣列，型別為 `withdraw` 且金額為提取的數額。
   - 更新餘額。
   - 傳回 `"Successfully withdrew $[amount]. New balance: $[balance]"`。

6. 如果提領金額小於或相等於 `0` 或大於目前餘額，`withdraw` 方法應該傳回 `"Insufficient balance or invalid amount."`。


7. 你應該定義一個名為 `checkBalance` 的方法，該方法以 `"Current balance: $[balance]"` 格式傳回目前餘額。

8. 你應該定義一個名為 `listAllDeposits` 的方法，該方法迭代 `transactions` 陣列並以 `"Deposits: amount,amount,..."` 格式傳回所有存款。

9. 你應該定義一個名為 `listAllWithdrawals` 的方法，該方法會迭代 `transactions` 陣列並以 `"Withdrawals: amount,amount,..."` 格式傳回所有提款。

10. 你應該創建一個名為 `myAccount` 的 `BankAccount` 實體。

11. 你的 `myAccount` 銀行帳戶應該至少有五筆交易。

12. 你的 `myAccount` 銀行帳戶應該至少有兩筆存款。

13. 你的 `myAccount` 銀行帳戶應該至少有兩筆提款。

14. 你的 `myAccount` 銀行帳戶餘額應該大於 `$100`。

# --hints--

你應該定義一個名為 `BankAccount` 的類別。

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

`BankAccount` 物件初始時應該有一個 `balance` 為 `0`，以及一個空的陣列 `transactions` 用來儲存交易記錄。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

你應該有一個 `deposit` 方法，該方法以存款金額作為參數。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

你應該有一個 `withdraw` 方法，該方法接受提款金額作為參數。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

你應該有一個 `checkBalance` 方法，用來檢查目前的餘額。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

你應該有一個 `listAllDeposits` 方法，該方法會列出所有存款。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

你應該有一個 `listAllWithdrawals` 方法，列出所有提款。

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` 應該傳回 `"Successfully deposited $100. New balance: $100"`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` 應該傳回 `"Deposit amount must be greater than zero."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` 應該傳回 `"Deposit amount must be greater than zero."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

當帳戶餘額為 `100` 時，`BankAccount.withdraw(150)` 應該傳回 `"Insufficient balance or invalid amount."`。

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` 應該傳回 `"Insufficient balance or invalid amount."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` 應該傳回 `"Insufficient balance or invalid amount."`。

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

當帳戶餘額為 `200` 時，`BankAccount.withdraw(150)` 應該傳回 `"Successfully withdrew $150. New balance: $50"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

當帳戶餘額為 `200` 時，`BankAccount.checkBalance()` 應該傳回 `"Current balance: $200"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

當你存入 `10`、`35`、`90` 時，`listAllDeposits` 方法應該傳回 `"Deposits: 10,35,90"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

當你提領 `20`、`50`、`100` 時，`listAllWithdrawals` 方法應該傳回 `"Withdrawals: 20,50,100"`。

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

你應該有一個名為 `myAccount` 的 `BankAccount` 實體。

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

你的 `myAccount` 銀行帳戶應該至少有五筆交易。

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

你的 `myAccount` 銀行帳戶應該至少有兩筆存款。

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

你的 `myAccount` 銀行帳戶應該至少有兩筆提款。

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```


你的 `myAccount` 銀行帳戶餘額應該大於 `$100`。

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
