---
id: 655cd899f8de09431eabb40c
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

マージソートは、分割統治の原理を使ってデータの集合をソートするアルゴリズムです。つまり、集合をより小さな部分に「分割」し、それぞれの部分を独立してソートして「制覇」し、最後にソート済みの部分をマージします。

`merge_sort` という名前の関数を宣言し、パラメータ `array` を受け取るところから始めましょう。この関数は数値のリストをソートする役割を担います。関数の本体では `pass` キーワードを使ってください。

# --hints--

`merge_sort` という名前の関数を、パラメータ `array` を1つ受け取る形で宣言してください。関数の本体では `pass` キーワードを使うのを忘れないでください。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
