---
id: 6564aee9c077774ea49c3faf
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

マージソートアルゴリズムは次の3つの動作を行います。

- ソートされていないアイテムのシーケンスをサブパーツに分割する
- サブパーツ内のアイテムをソートする
- ソートされたサブパーツをマージする

これらは再帰的に行われ、サブパーツが完全なソート済みシーケンスにマージされるまで続きます。まずはシーケンスを分割することから始めましょう。

最初に、`pass` キーワードを変数 `middle_point` に置き換えてください。整数除算演算子 (`//`) を使って `array` リストの長さを2で割り、その結果を新しい `middle_point` 変数に代入してください。コードのインデントを忘れないでください。

# --hints--

関数の本体に `pass` を含めてはいけません。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

`middle_point` 変数を宣言し、`len(array) // 2` を代入してください。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
