---
id: 656627b47bd2d2a4afbc945d
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

前のステップで、中央の位置を取得しました。それを使って `array` を2つに分け、それぞれを新しい変数に代入できます。

スライス構文を使って `array` の左半分を抽出し、`left_part` という名前の変数に代入してください。

# --hints--

`left_part` 関数内に `merge_sort` という名前の変数があるはずです。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

配列の左半分を抽出するにはスライス構文を使い、`left_part` 変数に代入してください。

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
