---
id: 656639771fed09b5c0e8fe71
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`array` リストを2つの別々のリストに分割したので、各リストをさらに分割して、すべての要素がそれぞれ単独のリストになるまで続けます。1つの数字だけのリストは常にソート済みです。

そのために、関数内で再帰的に `merge_sort` を呼び出し、呼び出しの引数として `left_part` を渡してください。

# --hints--

関数本体の最後で `merge_sort` を引数にして `left_part` を呼び出す必要があります。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
