---
id: 656680b0fc79f2c38a34d90e
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

いよいよリスト（`left_part` と `right_part`）をソートして元のリストにマージする時です。

両方のリストの要素を比較し、小さい方の要素をメインリストにマージすることでこれを行います。`left_part` と `right_part` のすべてのインデックスに対してこの比較を行います。

3つの変数 `left_array_index`、`right_array_index`、`sorted_index` を作成し、それぞれに `0` を設定してください。これらの変数はソート処理中に各インデックスを追跡するのに役立ちます。

# --hints--

`left_array_index` 関数の中に `merge_sort` という名前の変数があるはずです

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

`left_array_index` の値は `0` であるべきです。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

`right_array_index` 関数の中に `merge_sort` という名前の変数があるはずです

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

`right_array_index` の値は `0` であるべきです。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

`sorted_index` 関数の中に `merge_sort` という名前の変数があるはずです

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

`sorted_index` の値は `0` であるべきです。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)
--fcc-editable-region--
    
--fcc-editable-region--
```
