---
id: 65670d1ef177e7e2b76d9528
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

前のステップで作成した `if` と `else` の文は、ソートされた配列に要素を代入します。

ソートされた配列に代入された各要素はリストのインデックスを占めます。したがって、代入のたびにソートされた配列の次のインデックスに移動する必要があります。

`if`/`else` ブロックの下、しかしまだ `while` ループ内で、`sorted_index` を `1` だけ増やしてください。

# --hints--

`sorted_index` 節の後に、`1` の現在の値を `else` だけ増やす必要があります。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_bodies()[0].find_body()[1]
    cond.is_equivalent("sorted_index += 1") or cond.is_equivalent("sorted_index = sorted_index + 1") or cond.is_equivalent("sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
--fcc-editable-region--
```
