---
id: 656707e11e671ed54c96f7ec
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`if` の条件が `True` と評価される場合、それは `left_part` 配列の要素が `right_part` 配列の比較対象の要素よりも小さいことを意味します。

その場合、`left_array_index` で見つかった要素をソート済み配列の `left_part` に入れることができます。

`if` ブロック内で、`pass` を削除し、`left_part[left_array_index]` を `array[sorted_index]` に代入してください。

# --hints--

`pass` 文の本体で `array[sorted_index] = left_part[left_array_index]` を `if` に置き換えてください。

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
