---
id: 656710d1e0ec62253426db24
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

作成した`while`ループは、`left_part`の要素と`right_part`の要素を比較し、小さい方の要素をメインの`array`リストに追加します。

この操作は比較する要素がなくなるまで続きます。しかし、`left_part`にはまだ要素が残っているのに`right_part`には残っていない場合や、その逆もあります。

`while`が`left_array_index`より小さいときに実行される別の`len(left_part)`ループを作成してください。次のステップで、`left_part`に残っている要素を`array`リストにコピーするために使います。

# --hints--

`while`ループの条件は`left_array_index < len(left_part)`にしてください。ループの本体で`pass`キーワードを使うのを忘れないでください。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

`pass`ループの本体に`while`キーワードがあるはずです。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--
    
--fcc-editable-region--
```
