---
id: 656709e50ed928da35260276
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

前のステップでは、`left_part` の要素が小さいため `array` リストに代入しました。しかし、これは常にそうとは限りません。比較の中には、右側の要素の方が小さい場合もあります。

`else` の要素が `left_part` の要素より小さくない場合に実行される `right_part` 節を作成してください。

`else` ブロック内で、`right_part[right_array_index]` を `array[sorted_index]` に代入してください。

# --hints--

`array[sorted_index] = right_part[right_array_index]` 節の本文に `else` があるはずです。

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_bodies()[1].find_body()[0].is_equivalent("array[sorted_index] = right_part[right_array_index]")
`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
