---
id: 656702f8b4cbd8cbf0a433c6
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

次に、`left_part` と `right_part` の両方に要素が残っている限り続くループを作成する必要があります。

そのために、2つの条件を持つ `while` ループを作成してください。1つは `left_array_index` が `left_part` の長さより小さいかどうかをチェックし、もう1つは `right_array_index` が `right_part` の長さより小さいかどうかをチェックします。

# --hints--

`while` ループは、`left_array_index` が `left_part` の長さより小さいかどうか、そして `right_array_index` が `right_part` の長さより小さいかどうかをチェックする必要があります。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

ループの本体で `pass` キーワードを使ってください。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
