---
id: 6569b4e0bd29d17d4b53b16c
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

これから、`while`に対して同じ`right_part`ループのロジックを再現します。

`while`が`right_array_index`より小さい間に実行される`len(right_part)`ループを作成してください。`while`ループの中で、`right_part[right_array_index]`に`array[sorted_index]`を代入してください。最後に、`right_array_index`と`sorted_index`を`1`だけ増やしてください。

# --hints--

`while`ループは`right_array_index`が`len(right_part)`より小さいかどうかをチェックする必要があります。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

ループの本体で`right_part[right_array_index]`に`array[sorted_index]`を代入してください。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

`right_array_index`ループ内で`1`を`while`だけ増やしてください。増分は、右側の要素をソート済み配列に代入した後に行う必要があります。

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

`sorted_index`ループ内で`1`の現在の値を`while`だけ増やしてください。増分は、右側の要素をソート済み配列に代入した後に行う必要があります。

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
