---
id: 6569b743630ee592a65a7e2f
title: ステップ 20
challengeType: 20
dashedName: step-20
---

# --description--

`merge_sort()` 関数をテストする前に、`array` の長さが `1` 以下の場合に関数の実行を停止する基本ケースを作成する必要があります。

この基本ケースが再帰呼び出しをストップします。これがないと、リストがソート済みであっても、または要素がない場合でもマージソートの処理が続いてしまいます。

関数宣言の直後に、条件が `if` の `len(array) <= 1` 文を作成してください。`if` 文の本体内に、関数の実行を停止するための `return` 文を追加してください。

# --hints--

`if` かどうかを確認するために `len(array) <= 1` 文を作成する必要があります。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

`return` 文の本体に `if` 文があるはずです。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_ifs()[0].find_body().has_stmt("return")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
