---
id: 6577254891048c8f2c19e961
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

関数本体の一番下で、`merge_sort()` 関数をもう一度呼び出してください。今回は、`right_part` を引数として関数呼び出しに渡してください。

# --hints--

関数本体の一番下で、`merge_sort()` を引数 `right_part` で呼び出す必要があります。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
