---
id: 6577562501feabdf0984a184
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

`left_part` の要素をソート済み配列に代入した後、`left_array_index` を `1` だけ増やしてください。

# --hints--

要素をソート済み配列に挿入した後、現在の `1` の値に `left_array_index` を加える必要があります。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body()[1]

    cond.is_equivalent("left_array_index += 1") or cond.is_equivalent("left_array_index = left_array_index + 1") or cond.is_equivalent("left_array_index = 1 + left_array_index")
    `)) 
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]

--fcc-editable-region--
```
