---
id: 657823a9f4f372518614c8b7
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

`else` ブロック内で、`right_array_index` を `1` だけ増やしてください。

# --hints--

ソート済み配列に要素を挿入した後、現在の `1` の値に `right_array_index` を加える必要があります。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_bodies()[1].find_body()[1]

    cond.is_equivalent("right_array_index += 1") or cond.is_equivalent("right_array_index = right_array_index + 1") or cond.is_equivalent("right_array_index = 1 + right_array_index")
    `)) 
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        
--fcc-editable-region--
        else:
            array[sorted_index] = right_part[right_array_index]

--fcc-editable-region--
```
