---
id: 657f3a661730891aa64f3568
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

slice 構文を使って `array` の右半分を抽出し、`right_part` という名前の変数に代入してください。

# --hints--

`right_part` という名前の変数があるはずです。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

slice 構文を使って配列の右半分を抽出し、`right_part` という変数に代入してください。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
