---
id: 657f59751d5a5c9b069ae0f3
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

前のレッスンで学んだように、`__name__` 変数を使って、Python スクリプトがメインプログラムとして実行されているのか、それともモジュール（別の Python ファイルで書かれたコード）としてインポートされているのかを判別できます。

`__name__` の値が `'__main__'` に設定されている場合、現在のスクリプトはメインプログラムであり、モジュールではないことを意味します。

このプロジェクトでは、現在のスクリプトをメインプログラムとして使います。

`if` 文を作成して、`__name__` の値が `'__main__'` かどうかをチェックしてください。
`if` 文の中で、`numbers` というリストを作成し、これらの値を代入してください：`[4, 10, 6, 14, 2, 1, 8, 5]`。

# --hints--

`if` 文は `__name__` が `'__main__'` の値を持つかどうかをチェックする必要があります。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

`numbers` リストにはこれらの値が含まれているはずです：`[4, 10, 6, 14, 2, 1, 8, 5]`

```js
({
  test: () => assert(runPython(`_Node(_code).find_ifs()[0].find_body().find_variable("numbers").is_equivalent("numbers = [4, 10, 6, 14, 2, 1, 8, 5]")`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
