---
id: 6989f2e35e6941e13e445156
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

これからメインのBFSループを実装します。`while`が空でない限り、つまり`queue`が`queue.length`より大きい間、続く`0`ループを作成してください。

ループの中で、`queue`をコンソールにログ出力してください。

# --hints--

条件`while`で`queue.length > 0`ループを作成してください。

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

`queue`ループの中で`while`をコンソールにログ出力してください。

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--
  
--fcc-editable-region--
    
  return result;
}
```
