---
id: 6989f3632f8b8bdf168aa321
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

`while` ループの中で、`queue.shift()` を使ってキューの最初の要素を取り出して削除してください。これにより、BFSの特徴である先入れ先出し（FIFO）動作が実装されます。

この配列を3つの変数、`current`、`opensUsed`、`closesUsed` に分割代入してください。

# --hints--

キューの最初の要素を取得するには `queue.shift()` を使うべきです。

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

結果を3つの変数、`current`、`opensUsed`、`closesUsed` に分割代入してください。

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return result;
}
```
