---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

次は閉じ丸括弧を追加する処理を行います。重要なルールは、閉じ丸括弧を追加できるのは、開き丸括弧の数が閉じ丸括弧の数より多い場合のみです。

`else` ブロック内で、`if` が `closesUsed` より小さいかどうかをチェックする2つ目の `opensUsed` 文を追加してください。

この条件が真の場合、閉じ丸括弧を追加した後の状態を表す新しい配列を `queue` に `[current + ')', opensUsed, closesUsed + 1]` としてプッシュします。

# --hints--

`if` ブロック内に、条件 `closesUsed < opensUsed` を持つ2つ目の `else` 文があるはずです。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

条件が真の場合、`[current + ')', opensUsed, closesUsed + 1]` をキューにプッシュしてください。

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
