---
id: 6989eb65324d97775c4965fe
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、幅優先探索（BFS）を使って丸括弧のすべての有効な組み合わせを生成する関数を実装します。たとえば、2組の丸括弧の有効な組み合わせは `(())` と `()()` です。

まず、`genParentheses` という単一のパラメータを持つ `pairs` という名前の関数を作成してください。今は関数から空の配列を返してください。

# --hints--

`genParentheses` という名前の関数を定義してください。

```js
assert.isFunction(genParentheses);
```

`genParentheses` 関数は `pairs` という単一のパラメータを持つべきです。

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

関数は空の配列を返すべきです。

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
