---
id: 6989f207306332c06b0add39
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

幅優先探索の方法では、キューを使って丸括弧の組み合わせを作る過程で異なる状態を追跡します。各状態は3つの要素を含む配列で表されます：

- 現在作成中の文字列
- これまでに使った開き丸括弧の数
- これまでに使った閉じ丸括弧の数

`queue` という名前の変数を作成し、空の文字列と丸括弧が0個使われている状態を表す配列 `['', 0, 0]` を含む配列で初期化してください。

# --hints--

`queue` 関数内で `genParentheses` という名前の変数を宣言してください。

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.queue);
```

`queue` を `['', 0, 0]` を含む配列で初期化してください。

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { queue } = genParentheses?.variables ?? {};
assert.isFalse(queue?.value.isEmpty());
const queueValue = eval(queue?.value.toString());
assert.deepEqual(queueValue, [['', 0, 0]]);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  
  return result;
}
```
