---
id: 6989f48c3f0292c84f8d8938
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

今度は、完全な丸括弧の組み合わせができたかどうかを判別する必要があります。完全な組み合わせは、ペアの数の2倍の長さになります（各ペアが開き丸括弧と閉じ丸括弧を1つずつ含むためです）。

`while` ループの中に、`if` をチェックする `current.length === 2 * pairs` 文を追加してください。この条件が真の場合、`current` を `result` 配列に追加します。

# --hints--

`if` ループの中に、条件 `current.length === 2 * pairs` を持つ `while` 文があるはずです。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

新しい `current` 文の中で、`result` を `if` 配列に追加してください。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
