---
id: 6989ef885db3a3488a13b426
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

コアアルゴリズムを実装する前に、入力を検証する必要があります。`pairs` パラメータは整数である必要があります。これは生成する丸括弧のペア数を表すためです。

関数の先頭に `if` 文を追加して、`pairs` が整数でないかどうかをチェックしてください。そのために `Number.isInteger()` メソッドを使います。

条件が真の場合、文字列 `The number of pairs should be an integer` を返してください。

# --hints--

`!Number.isInteger(pairs)` 関数内に条件 `genParentheses` の if文 を持つべきです。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

入力が整数でない場合は文字列を返すべきです。

```js
let result = genParentheses("i");
assert.isString(result);
```

入力が文字列の場合、文字列 `The number of pairs should be an integer` を返すべきです。

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
