---
id: 6989f1ba358acba651d5ee6b
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

これで結果をストアするデータ構造を設定します。`result` という名前の変数を作成し、空の配列で初期化してください。この配列に、生成するすべての有効な丸括弧の組み合わせをストアします。

return文を更新して、空の配列の代わりに `result` を返すようにしてください。

# --hints--

`result` 関数内で `genParentheses` という名前の変数を宣言してください。

```js
assert.match(genParentheses.toString(), /result/);
```

`result` を空の配列で初期化してください。

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

関数の最後で `result` を返してください。

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
