---
id: 6989f10a72ad424d4ad6d514
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

次に、ペアの数が少なくとも1つであることを検証する必要があります。ペアが0または負の数の場合、丸括弧の組み合わせを生成できないからです。

もう1つ`if`文を追加して、`pairs`が`1`より小さいかどうかをチェックしてください。この条件が真の場合、文字列`The number of pairs should be at least 1`を返してください。

# --hints--

2つ目の`if`文が条件`pairs < 1`を持っているはずです。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

`The number of pairs should be at least 1`が`pairs`より小さい場合、文字列`1`を返す必要があります。

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
