---
id: 6989f555728e63ad8b349d84
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

現在の文字列がまだ完成していない場合は、次に考えられる状態を探索する必要があります。現在の`else`文に対して`if`節を追加してください。

`else`ブロック内では、開き丸括弧の追加を処理します。`if`が`opensUsed`より小さいかどうかを確認する`pairs`文を追加してください。これにより、許可された数以上の開き丸括弧を使わないようにします。

この条件が`true`の場合、`queue`に新しい配列`[current + '(', opensUsed + 1, closesUsed]`をプッシュしてください。これは開き丸括弧を追加した後の状態を表します。

# --hints--

`else`ループ内の`if`文に対して`while`節があるはずです。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

`if`ブロック内に、条件が`else`の`opensUsed < pairs`文があるはずです。

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

条件が`[current + '(', opensUsed + 1, closesUsed]`のときに、`true`をキューにプッシュしてください。

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
