---
id: 68e6a99a0fd7386fb66e156f
title: خطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتدرب على تنسيق مربعات الاختيار من خلال بناء لوحة إعدادات اللعبة.

تم توفير كل قالب HTML الجاهز (`DOCTYPE`، `html`، `head` و `body`) لك.

داخل `body`، أنشئ عنصر `div` بداخله عنصر `h1` مسكن. اجعل نص عنصر `h1` هو `Game Settings`

# --hints--

يجب أن يكون لديك عنصر `div`.

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

يجب أن يكون لديك عنصر `h1`.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

يجب أن يكون عنصر `h1` داخل عنصر `div` الخاص بك.

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

يجب أن يحتوي عنصر `h1` على النص `Game Settings`.

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
