---
id: 68e9a73075dc26207320811a
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

الآن بعد أن أصبح بإمكانك رؤية مربع الاختيار مرة أخرى، ستنهي بعض خيارات التنسيق له. امنحه حافة مستديرة بإضافة `border-radius` بقيمة `4px` إلى محدد `input[type="checkbox"]`.

ثم امنحه `background-color` بقيمة `white` لجعل مركز مربع الاختيار يبرز عن خلفية الحاوية.

بما أنك ستُعد انتقالًا مخصصًا عند نقر المستخدم على مربعات الاختيار، اضبط `transition` بقيمة `all` و `0.3s` ليحدث الانتقال بسلاسة خلال 0.3 ثانية بدلاً من أن يكون فوريًا.

# --hints--

يجب أن تضيف خاصية `border-radius` إلى محدد `input[type="checkbox"]` الخاص بك.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.borderRadius);
```

يجب أن تحتوي على خاصية `border-radius` بقيمة `4px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.borderRadius, "4px");
```

يجب أن تضيف خاصية `background-color` إلى محدد `input[type="checkbox"]` الخاص بك.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.backgroundColor);
```

يجب أن تحتوي على خاصية `background-color` بقيمة `white`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.backgroundColor, ["white", "rgb(255,255,255)", "#ffffff"]);
```

يجب أن تضيف خاصية `transition` إلى محدد `input[type="checkbox"]` الخاص بك.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.transition);
```

يجب أن تحتوي على `transition` بقيمة `all 0.3s`.

```js
/**
 * Because CSS is simply the most lovely thing to ever exist, the behaviour of the `all` shorthand does wonky things and will not show up
 * in the computed stylesheet. Instead, we need to query the actual style element and extract the selector, checking the properties as written
 * by the camper.
 * I really hate CSS.
 * <3 Naomi
 */
const { props } = document.querySelector('.fcc-injected-styles')?.innerText?.match(/input\[type=('|")checkbox\1\]\s*\{(?<props>.*)\}/s)?.groups;
assert.match(props, /transition\s*\:\s*all\s+0.3\s*/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
