---
id: 68e9a3194160f8f5620ddf21
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

بعد عنصر `h1` مباشرة، أنشئ أربعة عناصر `label`.

داخل كل عنصر من عناصر `label`، أنشئ عنصر `input` مع تعيين خاصية `type` إلى `"checkbox"`.

بعد كل عنصر `input`، أضف إلى عناصر `label` النصوص التالية بالترتيب: `Sound Effects`، `Background Music`، `Hard Mode`، و`Haptic Feedback`.

# --hints--

يجب أن يكون لديك أربعة عناصر `label`.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

يجب أن يحتوي أول عنصر `label` على عنصر `input`.

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

يجب أن تحتوي أول عنصر `input` على خاصية `type` بقيمة `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

يجب أن يحتوي أول عنصر `label` على النص `Sound Effects`.

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

يجب أن يأتي نص `Sound Effects` بعد عنصر `input`.

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

يجب أن يحتوي ثاني عنصر `label` على عنصر `input`.

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

يجب أن تحتوي ثاني عنصر `input` على خاصية `type` بقيمة `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

يجب أن يحتوي ثاني عنصر `label` على النص `Background Music`.

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

يجب أن يأتي نص `Background Music` بعد عنصر `input`.

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

يجب أن يحتوي ثالث عنصر `label` على عنصر `input`.

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

يجب أن تحتوي ثالث عنصر `input` على خاصية `type` بقيمة `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

يجب أن يحتوي ثالث عنصر `label` على النص `Hard Mode`.

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

يجب أن يأتي نص `Hard Mode` بعد عنصر `input`.

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

يجب أن يحتوي رابع عنصر `label` على عنصر `input`.

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

يجب أن تحتوي رابع عنصر `input` على خاصية `type` بقيمة `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

يجب أن يحتوي رابع عنصر `label` على النص `Haptic Feedback`.

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

يجب أن يأتي نص `Haptic Feedback` بعد عنصر `input`.

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
