---
id: 68e9a6bf1a74721a4e43e217
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

الآن ستزيل مربع الاختيار الافتراضي الذي تطبقه المتصفحات. داخل محدِّد `input[type='checkbox']`، أضف خاصية `appearance` بقيمة `none`.

تعيين خاصية `appearance` إلى `none` سيمسح المظهر الذي يطبقه المتصفح على مربعات الاختيار، مما يسمح لك بعرض النمط الذي تريده.

بعد ذلك، وبما أن مربع الاختيار لن يكون مرئيًا بعد الآن، عرّف حدودًا (`border`) بسُمك `2px`، ونمط `solid`، ورمز اللون السداسي `#f1be32`.

# --hints--

يجب أن تضيف خاصية `appearance` إلى محدِّد `input[type="checkbox"]` الخاص بك.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

يجب أن تحتوي على خاصية `appearance` بقيمة `none`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

يجب أن تضيف خاصية `border` إلى محدِّد `input[type="checkbox"]` الخاص بك.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

يجب أن تحتوي على خاصية `border` بقيمة `2px solid #f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
