---
id: 68e9a834dfef752dc35f6f90
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

الآن بعد أن تغير مربعات الاختيار لونها عند نقر المستخدم عليها، ستطبق بعض التنسيقات لتوضيح أن مربع الاختيار تم تحديده.

يمكن ربط الفئات المزيفة في CSS. يمكنك إضافة فئة مزيفة أو عنصر مزيف بعد الآخر لاستهداف العناصر التي تحقق شروطًا متعددة.

أنشئ محددًا جديدًا لـ `input[type="checkbox"]`. أضف فئة مزيفة `:checked` تليها عنصر مزيف `::after`. هذا يسمح لك بإضافة مؤشر بصري عند تحديد مربع الاختيار.

بعد ذلك، أضف الخاصية `content` بقيمة `"✓"`.

# --hints--

يجب أن يكون لديك محدد نوع مدمج مع محدد فئة مزيفة ومحدد عنصر مزيف على الشكل `input[type="checkbox"]:checked::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

يجب أن تضيف خاصية `content` إلى محدد النوع المدمج مع محدد الفئة المزيفة والمحدد العنصر المزيف `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

يجب أن تحتوي خاصية `content` على القيمة `"✓"`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```

