---
id: 68e9a6201e2efa142d62bfe1
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

التالي، ستقوم بتكبير مربعات الاختيار لتحسين مستوى الظهور.

ابدأ بإعداد محدِّد لـ `input`، مع استهداف محدد لـ `[type="checkbox"]`.

داخل المحدِّد، عيّن `width` و `height` إلى `20px`. هذا يجعلها أكبر مما كانت عليه سابقًا.

ولتتوافق مع إعداد `cursor` الذي عُيّن على المسميات، أضف `cursor` وعيّن له القيمة `pointer`. بعد ذلك، عند تحويم المؤشر فوق مربعات الاختيار سيظهر مؤشر يد.

# --hints--

يجب أن يكون لديك محدِّد في CSS الخاص بك هو `input[type="checkbox"]`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

يجب أن تضيف خاصية `width` إلى المحدِّد `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

يجب أن تحتوي على خاصية `width` بقيمة `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

يجب أن تضيف خاصية `height` إلى المحدِّد `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

يجب أن تحتوي على خاصية `height` بقيمة `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

يجب أن تضيف خاصية `cursor` إلى المحدِّد `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

يجب أن تحتوي على خاصية `cursor` بقيمة `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
