---
id: 68ee86e3ee07b01626a392cc
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

داخل محدِّد `.settings-card` عرّف خاصية `margin` بقيمة `auto`.

تعيين خاصية `margin` إلى `auto` يضبط هوامش العنصر تلقائيًا لتوزيع المساحة المتبقية بالتساوي داخل الحاوية، ويُستخدم عادةً لمركزة العناصر ذات المستوى الكتلي أفقيًا.

وأخيرًا، عرّف خاصية `text-align` بقيمة `left`. هذا سيُحاذي المحتوى المضمَّن، مثل النص، إلى الجانب الأيسر من العنصر الحاوي.

# --hints--

يجب إضافة خاصية `margin` إلى محدِّد الصنف `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

يجب أن تحتوي على خاصية `margin` بقيمة `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

يجب إضافة خاصية `text-align` إلى محدِّد الصنف `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

يجب أن تحتوي على خاصية `text-align` بقيمة `left`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
