---
id: 68f15ef0446464325ce1aa45
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

تحتاج إلى بعض التباعد بين مربعات الاختيار والمسمّيات. ابدأ باستخدام المحدِّد الخاص بـ `label` في CSS واضبط خاصية `display` على القيمة `block`.

بعد ذلك، اضبط خاصية `margin` على القيمة `8px auto` لإضافة تباعد عمودي بين العناصر. هذا سيُنشئ تباعدًا ثابتًا فوق وتحت كل عنصر، مما يساعد على فصل المحتوى وتحسين قابلية القراءة.

أخيرًا، اضبط خاصية `cursor` على القيمة `pointer` لعناصر `label`. هذا سيغير مؤشر الفأرة إلى أيقونة يد عند مرور المستخدم فوق `label`، مما يشير إلى أن العنصر قابل للنقر ويحسن تجربة المستخدم بشكل عام.

# --hints--

يجب أن يكون لديك محدِّد نوع في CSS الخاص بك لـ `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

يجب أن تضيف خاصية `display` إلى محدِّد `label` الخاص بك.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

يجب أن تحتوي خاصية `display` على القيمة `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

يجب أن تضيف خاصية `margin` إلى محدِّد `label` الخاص بك.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

يجب أن تحتوي خاصية `margin` على القيمة `8px auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

يجب أن تضيف خاصية `cursor` إلى محدِّد `label` الخاص بك.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

يجب أن تحتوي خاصية `cursor` على القيمة `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
