---
id: 68f1560aff1214103ab6b59b
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

الآن ستبدأ بتجميل الصفحة ببعض تنسيقات CSS. ابدأ بإنشاء محدِّد `body`.

عيّن للـ `body` خاصية `height` بقيمة `100vh` وخاصية `background-color` بقيمة `#f0f0f0`.

تجعل خاصية `height` بقيمة `100vh` الـ `body` يحتل كامل ارتفاع إطار العرض في المتصفح، بينما يوفر لون الخلفية الرمادي الفاتح قاعدة ناعمة للصفحة.

أخيرًا، عيّن خاصية `text-align` بقيمة `center`. هذا سيُوسّط كل المحتوى المضمَّن داخل الصفحة ما لم يتجاوز عنصر فرعي ذلك بمحاذاته الخاصة.

# --hints--

يجب أن تستخدم محدِّد النوع `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

يجب أن تضيف خاصية `height` داخل محدِّد النوع `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

يجب أن تحتوي على خاصية `height` بقيمة `100vh`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

يجب أن تضيف خاصية `background-color` داخل محدِّد النوع `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

يجب أن تحتوي على خاصية `background-color` بقيمة `#f0f0f0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

يجب أن تضيف خاصية `text-align` داخل محدِّد النوع `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

يجب أن تحتوي على خاصية `text-align` بقيمة `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
