---
id: 68f29753fd5524530c6e65c8
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

بعد ذلك، عرّف خاصية `vertical-align` بقيمة `middle`.

تتحكم خاصية `vertical-align` في كيفية محاذاة العناصر `inline` أو `inline-block` عموديًا مع النص المحيط أو العناصر المضمَّنة الأخرى. تُستخدم غالبًا لضبط الموضع العمودي لعناصر مثل الصور أو الأيقونات أو النص داخل السطر.

ثم أنهِ مربع الاختيار بخاصية `margin` بقيمة `15px`.

# --hints--

يجب إضافة خاصية `vertical-align` إلى محدِّد `input[type="checkbox"]` الخاص بك.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign);
```

يجب أن تحتوي على خاصية `vertical-align` بقيمة `middle`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign, "middle");
```

يجب إضافة خاصية `margin-right` إلى محدِّد `input[type="checkbox"]` الخاص بك.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight);
```

يجب أن تحتوي على خاصية `margin-right` بقيمة `15px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight, "15px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
