---
id: 68f29dbf9125476b4808d6fb
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

حان الوقت الآن لتنسيق حاوية `settings-card`. أنشئ محدد صنف لـ `settings-card`.

هذا هو المكان الذي ستضع فيه كل تنسيقك للحاوية. عرّف `max-width` بقيمة `250px` لتحديد الحجم الكلي للحاوية.

بعد ذلك، عرّف `padding` بقيمة `20px` ليكون هناك مساحة بين المحتوى وحدود الحاوية.

بعد ذلك، أنشئ حافة مستديرة بتعيين `border-radius` إلى `10px`.

ثم عرّف `box-shadow` بالقيم `0 2px 6px rgba(0,0,0,0.2)`. هذا سيخلق مظهرًا خفيفًا "مرتفعًا" يضيف عمقًا للحاوية.

# --hints--

يجب أن يكون لديك محدد صنف في ملف CSS الخاص بك لـ `settings-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

يجب أن تضيف خاصية `max-width` إلى محدد الصنف `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

يجب أن تحتوي على خاصية `max-width` بقيمة `250px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

يجب أن تضيف خاصية `padding` إلى محدد الصنف `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

يجب أن تحتوي على خاصية `padding` بقيمة `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

يجب أن تضيف خاصية `border-radius` إلى محدد الصنف `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

يجب أن تحتوي على خاصية `border-radius` بقيمة `10px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

يجب أن تضيف خاصية `box-shadow` إلى محدد الصنف `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

يجب أن تحتوي على خاصية `box-shadow` بالقيمة `0 2px 6px rgba(0,0,0,0.2)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
