---
id: 61329b210dac0b08047fd6ab
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

繼續使用 `meta` 元素，`viewport` 的定義告訴瀏覽器如何呈現頁面。包含一個可以提升行動裝置上的視覺可及性，並改善 *SEO*（搜尋引擎優化）。

為你的 `viewport` 添加一個 `content` 屬性，詳細說明頁面的 `width` 和 `initial-scale`。

# --hints--

你應該在 `head` 中創建另一個 `meta` 元素。

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

你應該給 `meta` 一個 `name` 屬性為 `viewport`。

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

你應該給 `meta` 一個 `content` 屬性，值為 `width=device-width, initial-scale=1`。

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
