---
id: 614090d5a22b6f0a5a6b464c
title: 步驟 13
challengeType: 0
dashedName: step-13
---

<!-- TODO: I purposefully added the `nav` styles without Camper input -->

# --description--

子代組合器選擇器 `>` 用於選擇器之間，以標的僅符合第二個選擇器且為第一個選擇器的直接子元素的元素。

當你有深度巢狀的元素並且想要控制你的樣式範圍時，這會很有幫助。

使用 `>` 選擇器來標的 `nav` 元素內的無序列表元素，並使用 *Flexbox* 均勻分布子元素。

# --hints--

你應該使用 `nav > ul` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('nav > ul'));
```

你應該為 `nav > ul` 元素設定 `display` 為 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('nav > ul')?.display, 'flex');
```

你應該為 `nav > ul` 元素設定 `justify-content` 為 `space-evenly`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('nav > ul')?.justifyContent, 'space-evenly');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

--fcc-editable-region--

--fcc-editable-region--
```
