---
id: 6140883f74224508174794c0
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

讓 `header` 佔滿其父容器的整個寬度，將其 `height` 設為 `50px`，並將 `background-color` 設為 `#1b1b32`。然後，將 `display` 設為使用 *Flexbox*。

# --hints--

你應該使用 `header` 元素選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('header'));
```

你應該給 `header` 一個 `100%` 的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.width, '100%');
```

你應該給 `header` 一個 `50px` 的 `height`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.height, '50px');
```

你應該為 `header` 設定 `background-color` 為 `#1b1b32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.backgroundColor, 'rgb(27, 27, 50)');
```

你應該給 `header` 一個 `display` 為 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

```
