---
id: 6140827cff96e906bd38fc2b
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

如同 <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCamp 風格指南</a> 所述，標誌應保持 `35 / 4` 的長寬比，並在文字周圍保留內距。

首先，將 `background-color` 更改為 `#0a0a23`，這樣你就可以看到標誌。然後，使用 `aspect-ratio` 屬性將所需的長寬比設為 `35 / 4`。最後，為四周添加 `0.4rem` 的 `padding`。

# --hints--

你應該為 `#logo` 設定 `background-color` 為 `#0a0a23`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

你應該使用 `aspect-ratio` 屬性。

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

你不應該使用 `height` 屬性。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

你不應該更改 `width` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

你應該給 `img` 一個 `aspect-ratio` 為 `35 / 4`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

你應該給 `img` 一個 `0.4rem` 的 `padding`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--
#logo {
  width: max(10rem, 18vw);

}
--fcc-editable-region--

```
