---
id: 6145f14f019a4b0adb94b051
title: 步驟 37
challengeType: 0
dashedName: step-37
---

# --description--

在第一個 `div.answer` 元素內，巢狀一個必填的 `select` 元素，該元素包含三個 `option` 元素。

為第一個 `option` 元素指定 `value` 為 `""`，以及文字為 `Select an option`。為第二個 `option` 元素指定 `value` 為 `yes`，以及文字為 `Yes`。為第三個 `option` 元素指定 `value` 為 `no`，以及文字為 `No`。

# --hints--

你應該在第一個 `div.answer` 元素中巢狀一個 `select` 元素。

```js
assert.exists(document.querySelector('div.answer > select'));
```

你應該在 `select` 元素中巢狀三個 `option` 元素。

```js
assert.equal(document.querySelectorAll('div.answer > select > option')?.length, 3);
```

你應該給第一個 `option` 元素一個 `value` 為 `""`。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(1)')?.value, '');
```

你應該為第一個 `option` 元素提供 `Select an option` 的文字內容。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(1)')?.textContent.trim(), 'Select an option');
```

你應該為第二個 `option` 元素指定 `value` 為 `yes`。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(2)')?.value, 'yes');
```

你應該給第二個 `option` 元素一個 `Yes` 的文字內容。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(2)')?.textContent.trim(), 'Yes');
```

你應該給第三個 `option` 元素一個 `value` 為 `no`。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(3)')?.value, 'no');
```

你應該給第三個 `option` 元素一個 `No` 的文字內容。

```js
assert.equal(document.querySelector('div.answer > select > option:nth-child(3)')?.textContent.trim(), 'No');
```

你應該為 `select` 添加一個 `required` 屬性。

```js
assert.isTrue(document.querySelector('div.answer > select')?.required);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label for="student-name">Name:</label>
            <input type="text" name="student-name" id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">Email:</label>
            <input type="email" name="student-email" id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">Date of Birth:</label>
            <input type="date" name="birth-date" id="birth-date" />
          </div>
        </section>
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
          <div class="question-block">
            <h3><span class="sr-only">Question</span>1</h3>
            <fieldset class="question" name="html-question-one">
              <legend>
                The legend element represents a caption for the content of its
                parent fieldset element
              </legend>
              <ul class="answers-list">
                <li>
                  <label for="q1-a1">
                    <input type="radio" id="q1-a1" name="q1" value="true" />
                    True
                  </label>
                </li>
                <li>
                  <label for="q1-a2">
                    <input type="radio" id="q1-a2" name="q1" value="false" />
                    False
                  </label>
                </li>
              </ul>
            </fieldset>
          </div>
          <div class="question-block">
            <h3><span class="sr-only">Question</span>2</h3>
            <fieldset class="question" name="html-question-two">
              <legend>
                A label element nesting an input element is required to have a
                for attribute with the same value as the input's id
              </legend>
              <ul class="answers-list">
                <li>
                  <label for="q2-a1">
                    <input type="radio" id="q2-a1" name="q2" value="true" />
                    True
                  </label>
                </li>
                <li>
                  <label for="q2-a2">
                    <input type="radio" id="q2-a2" name="q2" value="false" />
                    False
                  </label>
                </li>
              </ul>
            </fieldset>
          </div>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
          <div class="formrow">
            <div class="question-block">
              <label>Can the CSS margin property accept negative values?</label>
            </div>
--fcc-editable-region--
            <div class="answer">
              
            </div>
--fcc-editable-region--
            <div class="question-block">
              <label>Do you have any questions:</label>
            </div>
            <div class="answer">
            </div>
          </div>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

h3::before {
  content: "Question #";
}

.sr-only {
  position: absolute;
  width: 1px;
  height: 1px;
  padding: 0;
  margin: -1px;
  overflow: hidden;
  clip: rect(0, 0, 0, 0);
  white-space: nowrap;
  border: 0;
}

```
