---
id: 614ccc21ea91ef1736b9b578
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

為這個工作坊，你已經獲得一個基本的樣板程式碼。

開始這段無障礙旅程，請為你的 `html` 元素提供 `lang` 屬性。這將協助螢幕閱讀器識別頁面的語言。

此外，記得將你的樣式表聯結到頁面。

# --hints--

你應該為 `html` 元素指定一個 `lang` 屬性。*提示：你可以使用值 `en` 代表英文。*

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

你應該在你的 `head` 元素中有一個 `link` 元素。

```js
assert.exists(document.querySelector("head>link"));
```

你的 `link` 元素應該有 `rel="stylesheet"` 和 `href="styles.css"`。

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
