---
id: 6823f9df49cc206af5471a30
title: 調試 Camperbot 的個人資料頁面
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
---

# --description--

Camperbot 正在嘗試建立一個個人資料頁面。 他們請了一位朋友來檢查他們的代碼，並且表示其中存在一些錯誤。

你的任務是修復Camperbot的所有錯誤，以便他們能夠繼續構建他們的個人資料頁面。 完成以下用戶需求中的各項內容，然後點擊“運行測試”查看是否修復了所有錯誤。

**用戶故事：**

1. Camperbot 嘗試使用 `heading2` 元素，但該元素不存在。 修復這些標籤，使其使用正確的二級標題元素。
2. Camperbot 嘗試添加兩個帶有 `pp` 的段落，但這些元素也不存在。 修復它們，以便使用正確的段落標籤。
3. Camperbot 在對 `Background and Interests` 子標題使用 `h3` 元素，但有一個語法錯誤。 找到問題並解決它。

# --hints--

你在頁面上應該只有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

你不應該從頁面刪除 `<h1>Hello from Camperbot!</h1>` 元素。

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

你在頁面上應該只有一個 `h2` 元素。

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

你應該有一個 `h2` 元素，其文本爲 `About` 。 這是一個示例：`<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

你應該有一個段落元素，文本爲 `My name is Camperbot and I love learning new things.` 。這是一個示例：`<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

你應該有一個段落元素，文本爲 `I enjoy solving puzzles.` 。這是一個示例：`<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

你在頁面上應該恰好有兩個段落元素。

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

你在頁面上應該只有一個 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

你應該有一個 `h3` 元素，其文本爲 `Background and Interests` 。 這是一個示例：`<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

你在頁面上應該總共有五個元素。

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
