---
id: 6823f9df49cc206af5471a30
title: 偵錯 Camperbot 的評測頁面
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Camperbot 正在嘗試建置一個評測頁面。他們請朋友檢查他們的程式碼，朋友說裡面有一些錯誤。

你的工作是修正 Camperbot 的所有錯誤，以便他們能繼續建置他們的評測頁面。完成以下使用者故事中的項目，然後點擊「Run the Tests」以查看你是否修正了所有錯誤。

**使用者故事：**

1. Camperbot 正嘗試使用 `heading2` 元素，但該元素不存在。修正那些標記，使其使用正確的第二層級標題元素。
2. Camperbot 正在嘗試使用 `pp` 添加兩個段落，但那些也不存在。請修正它們，使其使用正確的段落標籤。
3. Camperbot 使用 `h3` 元素作為 `Background and Interests` 子標題，但它有語法錯誤。找出問題並加以決議。

# --hints--

你應該在頁面上有且只有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

你不應該從頁面中移除 `<h1>Hello from Camperbot!</h1>` 元素。

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

你應該在頁面上有且只有一個 `h2` 元素。

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

你應該有一個 `h2` 元素，裡面的文字是 `About`。這裡有一個範例：`<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

你應該有一個段落元素，裡面包含文字 `My name is Camperbot and I love learning new things.`。這裡有一個範例：`<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

你應該有一個段落元素，裡面包含文字 `I enjoy solving puzzles.`。這裡有一個範例：`<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

你應該在頁面上有正好兩個段落元素。

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

你應該在頁面上有且只有一個 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

你應該有一個 `h3` 元素，其文字為 `Background and Interests`。以下是一個範例：`<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

你應該在頁面上總共擁有 5 個元素。

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
