---
id: a6b0bb188d873cb2c8729495
title: implementare un convertitore di entità HTML
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

Questo laboratorio riguarda la conversione di caratteri speciali in una stringa con le corrispondenti entità HTML.

**Obiettivo**: Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione `convertHTML` che accetta una stringa come argomento.
1. La funzione `convertHTML` dovrebbe restituire una nuova stringa convertendo i caratteri speciali nella stringa argomento nelle corrispondenti entità HTML.

   - `&` dovrebbe essere convertito in `&amp;`.
   - `<` dovrebbe essere convertito in `&lt;`.
   - `>` dovrebbe essere convertito in `&gt;`.
   - `"` dovrebbe essere convertito in `&quot;`.
   - `'` dovrebbe essere convertito in `&apos;`.

# --hints--

Dovresti avere una funzione `convertHTML`.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` dovrebbe restituire la stringa `Dolce &amp; Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` dovrebbe restituire la stringa `Hamburgers &lt; Pizza &lt; Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` dovrebbe restituire la stringa `Sixty &gt; twelve`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` dovrebbe restituire la stringa `Stuff in &quot;quotation marks&quot;`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` dovrebbe restituire la stringa `Schindler&apos;s List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` dovrebbe restituire la stringa `&lt;&gt;`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` dovrebbe restituire la stringa `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
