---
id: 6a89c69c99298438d6fbe60b
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Fügen Sie nun den rekursiven Fall hinzu. Erstellen Sie nach der `if`-Anweisung eine Variable namens `count_list`. Weisen Sie ihr das Ergebnis des Aufrufs von `countup` mit `number - 1` zu.

Jeder Aufruf verwendet eine kleinere Zahl, sodass die Rekursion sich dem Basisfall nähert.

# --hints--

Sie sollten `countup(number - 1)` einer Variable namens `count_list` zuweisen.

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
