---
id: 6a89c69c6af682856558d78b
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Der rekursive Aufruf gibt eine Liste zurück, die die kleineren Zahlen enthält. Verwenden Sie `append()`, um `number` zu `count_list` hinzuzufügen, und geben Sie dann `count_list` zurück.

# --hints--

Sie sollten `number` an `count_list` anhängen.

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

Sie sollten `count_list` aus Ihrer `countup`-Funktion zurückgeben.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

Ihre `countup`-Funktion sollte `[1, 2, 3]` zurückgeben, wenn sie mit `3` aufgerufen wird.

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
