---
id: 6a89c69c111694edc6d1c628
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Testen Sie Ihre `countup`-Funktion, indem Sie ihr `5` innerhalb eines `print()`-Aufrufs übergeben. Die Ausgabe sollte `[1, 2, 3, 4, 5]` sein.

Damit ist der Workshop zur countup-Funktion abgeschlossen.

# --hints--

Sie sollten `countup(5)` innerhalb der `print()`-Funktion aufrufen.

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
