---
id: 6a89c69c8dbfda74151fa0b6
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop erstellen Sie eine `countup`-Funktion, die Rekursion verwendet, um eine Liste von Zahlen von `1` bis zu einer angegebenen Zahl zu erzeugen.

Beginnen Sie damit, eine Funktion namens `countup` mit einem Parameter namens `number` zu definieren. Fügen Sie das Schlüsselwort `pass` in den Funktionskörper ein, damit die Funktion gültig ist, während sie noch leer ist.

# --hints--

Sie sollten eine Funktion namens `countup` mit einem Parameter namens `number` haben.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

Ihre `countup`-Funktion sollte das Schlüsselwort `pass` enthalten.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
