---
id: 6a89c69c91bc7365fbd4b733
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Fügen Sie Ihrer `countup`-Funktion den Basisfall hinzu. Ersetzen Sie `pass` durch eine `if`-Anweisung, die überprüft, ob `number` kleiner als `1` ist. Geben Sie innerhalb der `if`-Anweisung eine leere Liste zurück.

# --hints--

Ihre `countup`-Funktion sollte eine leere Liste zurückgeben, wenn `number` kleiner als `1` ist.

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

Ihre `countup`-Funktion sollte keine leere Liste zurückgeben, wenn `number` größer oder gleich `1` ist.

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
