---
id: 61968df2acd5550bf1616c34
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`body` 요소의 `margin`와 `padding`를 제거하여 페이지 크기를 표준화하세요.

# --hints--

`body`에 `margin` 값을 가진 `0`를 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

`body`에 `padding` 값을 가진 `0`를 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.padding, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  
}
--fcc-editable-region--
```
