---
id: 6196d1ac33c68d27dcda5796
title: 23단계
challengeType: 0
dashedName: step-23
---

# --description--

`.back-mountain` 요소를 `45deg` 시계 방향으로 회전하세요. 그런 다음 `left` 값을 가진 `110px` 속성과 `top` 값을 가진 `225px` 속성을 부여하세요.

# --hints--

요소를 회전하려면 `transform` 속성을 사용하세요.

```js
const backMountain = new __helpers.CSSHelp(document).getStyle('.back-mountain');

if (backMountain?.transform) {
  assert.notEmpty(backMountain?.transform);
} else {
  assert.notEmpty(backMountain?.rotate);
}
```

`.back-mountain`에 `transform` 값을 가진 `--fcc-expected--`을 부여해야 하지만 `--fcc-actual--`이 발견되었습니다.

```js
const backMountain = new __helpers.CSSHelp(document).getStyle('.back-mountain');

if (backMountain?.transform) {
  assert.equal(backMountain?.transform, 'rotate(45deg)');
} else {
  assert.equal(backMountain?.rotate, '45deg');
}
```

`.back-mountain`에 `left` 속성을 부여해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.back-mountain')?.left);
```

`.back-mountain`에 `left` 값을 가진 `--fcc-expected--` 속성을 부여해야 하지만 `--fcc-actual--`이 발견되었습니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.back-mountain')?.left, '110px');
```

`.back-mountain`에 `top` 속성을 부여해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.back-mountain')?.top);
```

`.back-mountain`에 `top` 값을 가진 `--fcc-expected--` 속성을 부여해야 하지만 `--fcc-actual--`이 발견되었습니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.back-mountain')?.top, '225px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

--fcc-editable-region--
.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;

}
--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
