---
id: 619665c9abd72906f3ad30f9
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

행복한 Flappy Penguin을 만들면서 CSS 변형과 애니메이션을 더 깊이 탐구할 것입니다.

기본 HTML 보일러플레이트부터 시작하세요. `DOCTYPE` 선언, 언어가 영어로 설정된 `html` 요소, 적절한 `meta` 태그, `head`, `body`, `title` 요소를 포함하세요. 또한 스타일시트를 페이지에 연결하세요.

# --hints--

코드에 `<!DOCTYPE html>` 선언이 있어야 합니다.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

코드에 `html` 요소가 있어야 합니다.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

`html` 요소는 `lang` 속성이 `en`으로 설정된 여는 태그를 가져야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 요소는 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/html\s*>/gi));
```

코드에 `head` 요소가 `html` 요소 안에 있어야 합니다.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

코드에 `body` 요소가 `html` 요소 안에 있어야 합니다.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

`head` 요소는 `body` 요소보다 앞에 와야 합니다.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

`meta` 요소가 두 개 있어야 합니다.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

한 `meta` 요소는 `name`이 `viewport`으로, `content`가 `width=device-width, initial-scale=1.0`으로 설정되어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

다른 `meta` 요소는 `charset` 속성이 `UTF-8`으로 설정되어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

코드에 `title` 요소가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title`에는 텍스트가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

코드에 `link` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('link'));
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

`link` 요소는 `rel` 속성이 `stylesheet` 값으로 있어야 합니다.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 요소는 `href` 속성이 `styles.css` 값으로 있어야 합니다.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
