---
id: 6196990f966e8f10a40094f6
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

`.penguin` 요소를 선택하고, 그 `width`과 `height`를 `300px`로 설정하세요.

# --hints--

`.penguin` 선택자를 사용하세요.

```js
assert.match(code, /\.penguin\s*\{/);
```

`.penguin`에 `width`를 `300px`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin')?.width, '300px');
```

`.penguin`에 `height`를 `300px`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin')?.height, '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
