---
id: 6196928658b6010f28c39484
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`.ground` 요소에 시계 방향으로 90도 기울어진 선형 그라디언트를 `background`로 지정하세요. 그라디언트는 `rgb(88, 175, 236)`에서 시작해 `rgb(182, 255, 255)`에서 끝나야 합니다.

# --hints--

`.ground`에 `background`인 `linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255))`를 지정해야 합니다.

```js
assert.include(['linear-gradient(90deg,rgb(88,175,236),rgb(182,255,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(88,175,236),rgb(182,255,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.ground')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.ground {
  width: 100vw;
  height: 400px;

}
--fcc-editable-region--
```
