---
id: 61968e9243a4090cc805531c
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

가로 및 세로 스크롤바를 모두 하나의 속성만 사용해 제거하세요.

# --hints--

`body`에 `overflow` 값을 가진 `--fcc-expected--`을 지정해야 합니다. 하지만 `--fcc-actual--`이(가) 발견되었습니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;

}
--fcc-editable-region--
```
