---
id: 61967e74a8e3690ab6292daa
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`body` 요소를 선택하여 `background`을 시계 방향으로 45도 기울어진 선형 그라디언트로 설정하세요. 시작점은 `rgb(118, 201, 255)`이고 끝점은 `rgb(247, 255, 222)`입니다.

# --hints--

`body` 요소 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`background` 선택자에서 `body` 속성을 사용하세요.

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

`background`을 `linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))`로 설정하세요.

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
