---
id: 61969e7451455614217e901b
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

산이 더 산처럼 보이게 하려면, 두 개의 인수를 받는 `skew` 변환 함수를 사용할 수 있습니다. 첫 번째 인수는 x축을 전단할 각도이고, 두 번째 인수는 y축을 전단할 각도입니다.

`transform` 속성을 사용해 산을 x축에서 `0deg`, y축에서 `44deg`만큼 기울이세요.

# --hints--

`.left-mountain`에 `transform` 속성을 지정해야 합니다.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.transform);
```

`skew`에 `transform` 함수를 사용해야 합니다.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.transform, 'skew');
```

`.left-mountain`에 `transform`인 `skew(0deg, 44deg)`를 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.getPropVal('transform', true), 'skew(0deg,44deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;

}
--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
