---
id: 619699c10a0f6e11591d73c4
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

`margin` 속성을 사용해 `.penguin` 요소를 수평으로 가운데 정렬하고, `margin-top`를 `75px`로 설정하세요.

# --hints--

`.penguin`에 `margin`의 `--fcc-expected--`를 지정해야 하지만 `--fcc-actual--`이(가) 발견되었습니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin')?.marginRight, 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin')?.marginLeft, 'auto');
```

`.penguin`에 `margin-top`의 `--fcc-expected--`를 지정해야 하지만 `--fcc-actual--`이(가) 발견되었습니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin')?.marginTop, '75px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.penguin {
  width: 300px;
  height: 300px;

}
--fcc-editable-region--

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
