---
id: 61a8fe15a6a31306e60d1e89
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

페이지를 정규화하려면 `width`를 `100%`로, `height`를 `100vh`로 설정하세요.

# --hints--

`body`에 `width`의 `--fcc-expected--`를 지정해야 하지만 `--fcc-actual--`가 발견되었습니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

`body`에 `height`의 `--fcc-expected--`를 지정해야 하지만 `--fcc-actual--`가 발견되었습니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;

}
--fcc-editable-region--
```
