---
id: 699b5370a43e299b4673bcaa
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

سيكون من المفيد رؤية خصائص كائن `profile` في وحدة التحكم أثناء بنائه.

أضف `console.log` يسجل `profile` في وحدة التحكم.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تسجل `profile` في وحدة التحكم.

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
