---
id: 699b71d98c820ac7ab9eea2a
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

إذا فتحت وحدة التحكم مرة أخرى، سترى رسالة خطأ جديدة:

```md
Type 'number' is not assignable to type 'string'.
```

يحدث هذا لأنك تحاول إضافة خاصية `guest` بقيمة رقمية.

لإزالة خطأ TypeScript، غيّر قيمة خاصية `guest` إلى `"read-only"`.

مع هذا التغيير الأخير، تكون ورشة عمل ملف المستخدم قد اكتملت!

# --hints--

يجب أن تكون قيمة خاصية `guest` في كائن `userRoles` هي `"read-only"`.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
