---
id: 699b5b7cb68b6d35afca61b7
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

بينما يمكنك الاستمرار في إضافة خصائص بهذا الشكل إلى كائن `profile`، سيكون من الأفضل تقييد أنواع الخصائص المسموح بها. من الأفضل تحديد شكل كائن `profile` مسبقًا.

في الدروس السابقة، تعلمت كيفية العمل مع أنواع الكائنات بهذا الشكل:

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

في هذا المثال، `composer` هو كائن بنوع كائن مضمَّن. يحتوي هذا الكائن على ثلاث خصائص مطلوبة هي `name` و`instrument` و`isActive`. يجب أن تكون الخاصيتان الأوليان من نوع `string` ويجب أن تكون `isActive` من نوع `boolean`.

حدّث كائن `profile` الحالي ليشمل نوع كائن مضمَّن. يجب أن تكون خاصية `username` من نوع `string`، و`age` من نوع `number`، و`isLoggedIn` من نوع `boolean`.

# --hints--

يجب أن يحتوي كائن `profile` على نوع كائن مضمَّن مع خاصية `username` من نوع `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

يجب أن يحتوي كائن `profile` على نوع كائن مضمَّن مع خاصية `age` من نوع `number`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

يجب أن يحتوي كائن `profile` على نوع كائن مضمَّن مع خاصية `isLoggedIn` من نوع `boolean`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
--fcc-editable-region--
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}

console.log(profile);
```
