---
id: 699b65caf17a8e8d340af27d
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

حان الوقت الآن لإضافة الخاصية الأخيرة إلى كائن `profile`.

في الدروس السابقة، تعلمت كيفية التعامل مع أنواع المصفوفات. تعرفت على أنواع مختلفة من المصفوفات مثل `string[]` أو `number[]`:

```ts
const scores: number[] = [95, 87, 100];
```

في هذا المثال، مصفوفة `scores` تقبل الأرقام فقط. إذا حاولت إضافة نوع آخر مثل `string` أو `null`، فسيُصدر TypeScript خطأ.

حدّث كائن `profile` لإضافة خاصية `programmingLanguages`، مع تحديد نوعها كـ `string[]`. داخل كائن التعريف، عيّن هذه الخاصية إلى `["JavaScript", "Python", "C++"]`.

# --hints--

يجب أن تحتوي على خاصية `programmingLanguages` داخل كائن `profile`.

```js
assert.property(profile, "programmingLanguages");
```

يجب أن تكون قيمة خاصية `programmingLanguages` هي `["JavaScript", "Python", "C++"]`.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

يجب أن تكون خاصية `programmingLanguages` محددة النوع كـ `string[]`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
--fcc-editable-region--
  
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
  
--fcc-editable-region--
};

console.log(profile);
```
