---
id: 699b70b10ae344c802d106b0
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

حان الوقت الآن لإضافة خاصيتين إضافيتين إلى كائن `userRoles`.

أضف خاصية `moderator` بقيمة `"medium-access"`.

ثم أسفل ذلك، أضف خاصية `guest` بقيمة `3`.

# --hints--

يجب أن يحتوي كائن `userRoles` على خاصية `moderator`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

يجب أن تكون قيمة خاصية `moderator` هي `"medium-access"`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

يجب أن يحتوي كائن `userRoles` على خاصية `guest`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

يجب أن تكون قيمة خاصية `guest` هي الرقم `3`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
--fcc-editable-region--
  viewer: "read-only"
  
--fcc-editable-region--
};

console.log(userRoles);
```
