---
id: 699b6dd1a2961bd38491ef2e
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

كان كائن `profile` يحتوي على مفاتيح وأنواع صارمة. إذا حاولت إضافة خاصية غير معرفة، أو تعيين نوع غير متوافق، كان TypeScript يرمي خطأ.

سيكون من الجيد لو لم يكن كائن `userRoles` صارمًا جدًا وكان لديه مفاتيح أكثر مرونة.

كما تتذكر من الدروس السابقة، يمكنك تحقيق هذه النتيجة باستخدام نوع `Record`:

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

في مثال `studentGrades` هذا، لا يحتاج TypeScript إلى معرفة كل المواد مسبقًا. هذا يسمح لك بإضافة مواد أخرى مثل `Art` أو `Music` لاحقًا. يهم TypeScript فقط أن تكون القيم لكل مفتاح أرقامًا.

حدّث كائن `userRoles` الخاص بك بإضافة نوع `Record` مضمّن. عيّنه إلى `<string, string>` بحيث يمكن أن يكون كل مفتاح أي سلسلة نصية، ويجب أن تكون كل قيمة سلسلة نصية.

# --hints--

يجب أن يحتوي كائن `userRoles` الخاص بك على نوع `Record` مضمّن مضبوط على `Record<string, string>`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
