---
id: 699b64096301fd75e3efeb1b
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

إذا فتحت وحدة التحكم مرة أخرى، سترى رسالة خطأ مختلفة:

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

كما تتذكر من الدروس السابقة، يمكنك تمييز الخصائص كاختيارية باستخدام الرمز `?` بجانب اسم الخاصية هكذا:

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

في هذا المثال، تم تمييز `numberOfAlbums` كاختيارية لذلك لن يرمي TypeScript خطأ في هذه الحالة.

حدّث خاصية `bio` لتكون مميزة كاختيارية.

# --hints--

يجب أن تكون خاصية `bio` مميزة كاختيارية.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
