---
id: 699b61ca8c9cb9bc1968364e
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

إذا فتحت وحدة التحكم، سترى رسالة الخطأ التالية:

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

في JavaScript العادي، كان مسموحًا لك إضافة أي خصائص وقيم تريدها إلى كائن `profile`. لكن في TypeScript، يتحقق المترجم من أن الكائن يحتوي فقط على الخصائص التي عرّفتها (`username`، `age`، و`isLoggedIn` في هذه الحالة). بما أن `mood` لم تُدرج في النوع، يعتبرها TypeScript خاصية غير صالحة ويرمي خطأ.

احذف `mood: null` من شفرتك.

# --hints--

يجب ألا يكون لديك `mood: null` في شفرتك بعد الآن.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
