---
id: 699b62c99ec8c63470313365
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

حان الوقت الآن لإضافة خاصية أخرى إلى كائن `profile` الخاص بك.

ابدأ بتحديث نوع الكائن المضمَّن ليشمل خاصية `bio` مع تعيين نوعها إلى `string`.

# --hints--

يجب عليك تحديث نوع الكائن المضمَّن ليشمل خاصية `bio` مع تعيينها إلى `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
