---
id: 699b4e284291adcfcb90df47
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستتدرب على العمل مع الأنواع الأولية بالإضافة إلى أنواع الكائنات والمصفوفات من خلال بناء ملف تعريف مستخدم آمن النوع.

ابدأ بإنشاء متغير يسمى `profile` وقم بتعيين كائن له. يجب أن يحتوي كائنك على هذه الخصائص والقيم الثلاث:

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

يجب أن يكون لديك متغير يسمى `profile`.

```js
assert.exists(profile);
```

يجب أن يكون متغير `profile` كائنًا.

```js
assert.isObject(profile);
```

يجب أن يحتوي كائن `profile` على خاصية `username`.

```js
assert.property(profile, "username");
```

يجب أن تكون قيمة خاصية `username` هي `"codeLearner"`.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

يجب أن يحتوي كائن `profile` على خاصية `age`.

```js
assert.property(profile, "age");
```

يجب أن تكون قيمة خاصية `age` هي `25`.

```js
assert.propertyVal(profile, "age", 25);
```

يجب أن يحتوي كائن `profile` على خاصية `isLoggedIn`.

```js
assert.property(profile, "isLoggedIn");
```

يجب أن تكون قيمة خاصية `isLoggedIn` هي `false`.

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
