---
id: 699b68f94a0d154c2a7a5ce0
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

الآن بعد أن اكتمل كائن `profile`، حان الوقت لبناء كائن جديد يمثل أدوار المستخدمين.

أنشئ متغيرًا جديدًا اسمه `userRoles` وعيّن له كائنًا. يجب أن يحتوي كائنك على الخصائص والقيم التالية:

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

يجب أن يكون لديك متغير اسمه `userRoles`.

```js
assert.isDefined(userRoles);
```

يجب أن يكون متغير `userRoles` كائنًا.

```js
assert.isObject(userRoles);
```

يجب أن يحتوي كائن `userRoles` على خاصية `admin`.

```js
assert.property(userRoles, "admin");
```

يجب أن تكون قيمة خاصية `admin` هي `"full-access"`.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

يجب أن يحتوي كائن `userRoles` على خاصية `editor`.

```js
assert.property(userRoles, "editor");
```

يجب أن تكون قيمة خاصية `editor` هي `"limited-access"`.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

يجب أن يحتوي كائن `userRoles` على خاصية `viewer`.

```js
assert.property(userRoles, "viewer");
```

يجب أن تكون قيمة خاصية `viewer` هي `"read-only"`.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
