---
id: 680b3ef395479b0e449ecb6e
title: "implementare l'algoritmo di ordinamento selection sort"
challengeType: 27
dashedName: implement-selection-sort-algorithm
---

# --description--

Selection sort è un altro algoritmo di ordinamento molto diffuso, insegnato nella maggior parte dei corsi di informatica.

Questo algoritmo funziona trovando ripetutamente l'elemento più piccolo nella parte non ordinata della lista e scambiandolo con il primo elemento non ordinato. Inizia selezionando il valore minimo nell'intera lista e scambiandolo con il primo elemento. Poi passa alla seconda posizione, trova il valore più piccolo tra gli elementi rimanenti non ordinati e lo scambia con il secondo elemento. Questo processo continua, avanzando di un elemento alla volta, fino a quando l'intera lista è ordinata.

Selection sort ha una complessità temporale quadratica nei casi migliore, medio e peggiore. La complessità spaziale sarà costante `O(1)` perché l'ordinamento avviene in loco e viene usata una quantità costante di memoria indipendentemente dalla dimensione della lista.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti definire una funzione chiamata `selection_sort`.
1. La tua funzione `selection_sort` dovrebbe avere un parametro che rappresenta la lista di elementi.
1. La tua funzione `selection_sort` dovrebbe prendere una lista e ordinare gli elementi in loco dal più piccolo al più grande.
1. La tua funzione `selection_sort` dovrebbe modificare la lista di input in loco e restituirla una volta ordinata.
1. La tua funzione `selection_sort` dovrebbe seguire l'algoritmo selection sort, scambiando l'elemento più piccolo della parte non ordinata della lista con il primo elemento non ordinato.
1. La tua funzione `selection_sort` non dovrebbe eseguire scambi inutili quando l'elemento più piccolo è già nella posizione corretta.
1. La tua funzione `selection_sort` non dovrebbe usare né il metodo integrato `sort()` né la funzione `sorted()`.

# --hints--

Dovresti avere una funzione chiamata `selection_sort`.

```js
({test: () => {
    runPython(`
assert _Node(_code).has_function('selection_sort')
    `)
}})
```

La tua funzione `selection_sort` dovrebbe avere un parametro.

```js
(
    {test: () => {
        runPython(`
            import inspect
            sig = inspect.signature(selection_sort)
            assert len(sig.parameters) == 1
        `)
    }}
)
```

Non dovresti importare nessun modulo né usare metodi di ordinamento integrati nel tuo codice.

```js
(
    {
        test: () => runPython(`
            assert len(_Node(_code).find_imports()) == 0
            assert not _Node(_code).block_has_call("sort")
            assert not _Node(_code).block_has_call("sorted")
        `)
    }
)
```

La tua `selection_sort` dovrebbe restituire la stessa lista passata in input.

```js
(
    {
        test: () => runPython(`
            input_list = [32, 0, 88, 1, 203]
            assert selection_sort(input_list) is input_list
        `)
    }
)
```

La tua `selection_sort` dovrebbe modificare la lista di input in loco. Non dovresti usare metodi che aggiungono o rimuovono elementi dalla lista.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)

                def append(self, *args, **kwargs):
                    self._record.append(f'Appending to list: {args}, {kwargs}')
                    super().append(*args, **kwargs)

                def insert(self, *args, **kwargs):
                    self._record.append(f'Inserting to list: {args}, {kwargs}')
                    super().insert(*args, **kwargs)

                def pop(self, *args, **kwargs):
                    self._record.append(f'Popping from list: {args}, {kwargs}')
                    super().pop(*args, **kwargs)

                def remove(self, *args, **kwargs):
                    self._record.append(f'Removing from list: {args}, {kwargs}')
                    super().remove(*args, **kwargs)


            list_to_sort = CustomList([32, 0, 88, 1, 203])

            selection_sort(list_to_sort)
            assert all(
                'Setting item' in record
                for record in list_to_sort._record
            )
        `)
    }
)
```

La tua funzione `selection_sort` dovrebbe seguire l'algoritmo selection sort, scambiando il valore minimo nella parte non ordinata della lista con il primo elemento non ordinato. Evita scambi inutili quando il valore minimo è già nella posizione corretta.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)


            list_to_sort = CustomList([33, 1, 89, 2, 67, 245])

            swap_pairs = [
                ((0, 1, 33), (1, 33, 1)),
                ((1, 2, 33), (3, 33, 2)),
                ((2, 33, 89), (3, 89, 33)),
                ((3, 67, 89), (4, 89, 67)),
            ]

            selection_sort(list_to_sort)
            record = list_to_sort._record

            actual_pairs = [
                (record[offset * 2], record[offset * 2 + 1])
                for offset, _ in enumerate(record[::2])
            ]
            assert len(swap_pairs) == len(actual_pairs)

            for expected_pair, actual_pair in zip(swap_pairs, actual_pairs):
                assert (
                    set(f'Setting item {key} to {value}. Previous value: {old}' for key, value, old in expected_pair)
                    == set(actual_pair)
                )
        `)
    }
)
```

`selection_sort([33, 1, 89, 2, 67, 245])` dovrebbe restituire `[1, 2, 33, 67, 89, 245]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([33, 1, 89, 2, 67, 245]) == [1, 2, 33, 67, 89, 245]
            
        `)
    }
)
```

`selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3])` dovrebbe restituire `[3, 5, 12, 15, 16, 23, 72, 99, 567]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3]) == [3, 5, 12, 15, 16, 23, 72, 99, 567]
            
        `)
    }
)
```

`selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` dovrebbe restituire `[1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]) == [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
            
        `)
    }
)
```

La tua funzione `selection_sort` dovrebbe ordinare correttamente qualsiasi lista di numeri.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([42, 17, 93, 8, 61, 29]) == [8, 17, 29, 42, 61, 93]
            assert selection_sort([11, 4, 78, 23, 55, 198, 65, 90, 2]) == [2, 4, 11, 23, 55, 65, 78, 90, 198]
            assert selection_sort([9, 27, 3, 7, 101, 66, 34, 52, 87, 42, 12, 29]) == [3, 7, 9, 12, 27, 29, 34, 42, 52, 66, 87, 101]
            assert selection_sort([5, 14, 33, 77, 2, 18, 92, 1, 100, 45, 73, 64, 28, 56]) == [1, 2, 5, 14, 18, 28, 33, 45, 56, 64, 73, 77, 92, 100]
            
        `)
    }
)     
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def selection_sort(nums):
    for i, _ in enumerate(nums):
        min_index = i

        for j in range(i + 1, len(nums)):
            if nums[j] < nums[min_index]:
                min_index = j

        if min_index != i:
            nums[i], nums[min_index] = nums[min_index], nums[i]

    return nums

```
