---
id: 657efce98e958b75df86b305
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Dentro do corpo da declaração `if`, você vai adicionar o caractere convertido à lista que criou anteriormente.

Para isso, o método `.append()` será usado.Este método adiciona um objeto dado ao final da lista em que ele é invocado.

Use o método `.append()` em `snake_cased_char_list` para adicionar o `converted_character` à lista.

# --hints--

Você deve usar o método `.append()` para adicionar o caractere convertido à `snake_cased_char_list`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\s*\(\s*converted_character\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
