---
id: 657ee28cefc4945568287673
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Com a lista vazia pronta, agora você pode começar a iterar pela string de entrada e convertê-la em snake case.

Dentro da função, abaixo da lista que você acabou de criar, adicione um loop `for` para iterar através da `pascal_or_camel_cased_string`.Certifique-se de nomear a variável de destino como `char`.Por enquanto, adicione uma instrução `pass` como um espaço reservado no corpo do loop.

# --hints--

Você deve escrever um novo loop `for` com a variável alvo chamada `char`.Não esqueça do dois-pontos no final e de usar `pass` para preencher o corpo do loop.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
