---
id: 657efa642593c5746acc5c81
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Dentro do corpo da declaração `if`, você precisa converter qualquer caractere maiúsculo para minúsculo e adicionar um sublinhado antes desse caractere minúsculo.

Use o método de string `.lower()` para converter caracteres maiúsculos em caracteres minúsculos. Então, adicione um sublinhado antes do caractere. Atribua os resultados a uma variável chamada `converted_character`.

# --hints--

Você deve atribuir o caractere modificado a uma variável chamada `converted_character`.

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

Você não deve ter `pass` na sua declaração `if`.

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
