---
id: 657e928716b77b2277980276
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste projeto, você vai construir um programa que recebe uma string formatada em `camelCase` ou `PascalCase` como entrada e converte-a para uma string formatada em `snake_case` usando duas abordagens. Primeiro, você usará um loop for e depois list comprehension para alcançar os mesmos resultados. Você verá como list comprehension pode tornar seu código mais conciso.

Comece definindo uma nova função chamada `convert_to_snake_case()` que aceita uma string chamada `pascal_or_camel_cased_string` como entrada. Por enquanto, adicione uma instrução `pass` dentro da função.

# --hints--

Você deve definir uma nova função chamada `convert_to_snake_case()`. Não esqueça do dois-pontos no final e use `pass` para preencher o corpo da função.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

Sua função deve receber um parâmetro chamado `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
