---
id: 657effaa2a5e0277d71f9cbe
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Neste ponto, a variável `snake_cased_char_list` contém a lista de caracteres convertidos. Para combinar esses caracteres em uma única string, você pode utilizar o método `.join()`.

O método `join` funciona concatenando cada elemento de uma lista em uma string, separado por uma string designada, conhecida como separador.

```py
result_string = ''.join(characters)
```

O exemplo acima junta os elementos da lista `characters` em uma única string onde cada elemento é concatenado usando uma string vazia como separador.

Agora, logo após o loop `for`, use o método `.join()` para unir os elementos em `snake_cased_char_list` usando uma string vazia como separador. Atribua o resultado a uma nova variável chamada `snake_cased_string`.


# --hints--

Você deve juntar os caracteres na lista `snake_cased_char_list` em uma string usando o método `.join()` e atribuir a string a uma variável chamada `snake_cased_string` após o seu loop `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
