---
id: 657ef2a86d4e545cec9a85fb
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Tanto em camel case quanto em pascal case, caracteres maiúsculos indicam o início de novas palavras. Para converter a string de entrada para snake case, você precisará verificar se os caracteres na string de entrada estão em maiúsculas. 

Você pode usar o método de string `.isupper()` para verificar se um caractere está em maiúscula. Este método retorna `True` se o caractere for maiúsculo e `False` se não for.

Dentro do loop `for`, adicione uma instrução `if` para verificar se o caractere atual está em maiúscula. Mova a declaração `pass` para dentro da nova declaração `if`.

# --hints--

Você deve escrever uma nova instrução if com char.isupper() como condição. Lembre-se de adicionar os dois pontos no final e usar a palavra-chave pass para preencher o corpo da instrução if.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

Você deve substituir a instrução `pass` no corpo do loop pela instrução `if`.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
