---
id: 657f465f8e718b19c5105ae5
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

List comprehensions aceitam declarações condicionais, para avaliar a expressão fornecida somente se certas condições forem atendidas:

```py
spam = [i * 2 for i in iterable if i > 0]
```

Como você pode ver na saída, a lista de caracteres gerada a partir de `pascal_or_camel_cased_string` foi unida. Como a expressão dentro da list comprehension é avaliada para cada caractere, o resultado é uma string em minúsculas com todos os caracteres separados por um underscore.

Siga o exemplo acima para adicionar uma cláusula `if` à sua list comprehension para que a expressão seja executada somente se o caractere for maiúsculo.

# --hints--

Você deve adicionar uma cláusula `if` com a condição `char.isupper()` à sua list comprehension.

```js
({ 
    test: () => assert(runPython(`
    ifs = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_ifs()
    len(ifs) == 1 and ifs[0].is_equivalent("char.isupper()")
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
