---
id: 657f456223b8c1187b461987
title: Passo 19
challengeType: 20
dashedName: step-19
---

# --description--

Em Python, uma <dfn>list comprehension</dfn> é uma construção que permite gerar uma nova lista aplicando uma expressão a cada item em um iterável existente e opcionalmente filtrando itens com uma condição. Além de serem mais breves, as list comprehensions frequentemente executam mais rápido.

Uma compreensão de lista básica consiste em uma expressão seguida por uma cláusula `for`:

```py
spam = [i * 2 for i in iterable]
```

O acima usa a variável `i` para iterar sobre `iterable`. Cada elemento da lista resultante é obtido avaliando a expressão `i * 2` na iteração atual.

Nesta etapa, você precisa preencher a lista vazia `snake_cased_char_list` usando a sintaxe de list comprehension.

Transforme sua lista vazia em uma list comprehension que converte cada caractere em `pascal_or_camel_cased_string` em um caractere minúsculo e adiciona um underscore antes dele (o código que você comentou antes pode ajudar a escrever a expressão). Use `char` para iterar sobre `pascal_or_camel_cased_string`. 

# --hints--

Você deve transformar `snake_cased_char_list` em uma list comprehension que itere sobre `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

Sua compreensão de lista deve usar `char` para iterar sobre `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

Sua list comprehension deve avaliar `'_' + char.lower()` para cada `char` em `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
