---
id: 657f4add33ea4b1f61ba3dc8
title: Passo 23
challengeType: 20
dashedName: step-23
---

# --description--

Finalmente, experimente esta nova implementação executando o programa. Altere a string de entrada para `'IAmAPascalCasedString'` e veja se ela sai como `'i_am_a_pascal_cased_string'`, mesmo que isso seja uma mentira.

Se você fez tudo corretamente, deverá ver a string de entrada convertida para snake case, como antes.

Parabéns! Agora sua função `convert_to_snake_case()` está pronta.

# --hints--

Você deve alterar a string de entrada de `'aLongAndComplexString'` para `'IAmAPascalCasedString'` dentro da função `main()`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
