---
id: 657f0044be09db790b1eb1c5
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Em pascal case, strings começam com uma letra maiúscula.Após converter todos os caracteres para minúsculas e adicionar um underscore a eles, há a chance de ter um underscore extra no início da sua string.

A maneira mais fácil de corrigir isso é usando o método de string `.strip()`, que remove de uma string quaisquer caracteres iniciais ou finais entre um conjunto de caracteres passados como argumento.Por exemplo:

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

O método `strip()` é aplicado a `original_string`.Isso remove qualquer underscore inicial e final.O resultado do exemplo acima seria a string `'example_string'`.

Declare uma nova variável chamada `clean_snake_cased_string` e atribua a ela o resultado do método `.strip()` aplicado a `snake_cased_string`, passando `'_'` como argumento para o método.

# --hints--

Você deve atribuir a string resultante de `snake_cased_string.strip('_')` a uma variável chamada `clean_snake_cased_string`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
