---
id: 657f47b12c51e41b3149e584
title: Passo 21
challengeType: 20
dashedName: step-21
---

# --description--

Ainda assim, o resultado final não é exatamente o que você quer alcançar.Você precisa executar uma expressão diferente para os caracteres filtrados pela cláusula `if`.Você usará uma cláusula `else` para isso:

```py
spam = [i * 2 if i > 0 else -1 for i in iterable]
```

Observe que, diferentemente da cláusula `if`, a construção `if`/`else` deve ser colocada entre a expressão e a palavra-chave `for`.

Modifique sua list comprehension para que, quando um caractere não for maiúsculo, ele permaneça inalterado.

# --hints--

Você deve modificar sua list comprehension para avaliar a expressão `'_' + char.lower()` se `char.isupper()` e char caso contrário.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_expr().is_equivalent("'_' + char.lower() if char.isupper() else char")
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string if char.isupper()]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
