---
id: 657f025ec86c3d7c4177b6be
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Com a função completa, você agora pode usá-la dentro de outra função. 

Crie uma nova função chamada `main()` com `pass` como o corpo da função.

# --hints--

Você deve definir uma nova função chamada `main()`. Não esqueça do dois-pontos no final e use `pass` para preencher o corpo da função.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(main)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--


--fcc-editable-region--
```
