---
id: 657f04ed0035f47ed04d0f1f
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Para exibir a saída da função `convert_to_snake_case()`, você precisa chamar a função `main()`. 

No mesmo nível das duas funções existentes, adicione uma chamada para a função `main()`. Você deve ver a string em camel case ou pascal case dada convertida para snake case após a execução.

# --hints--

Você deve adicionar uma chamada para a função `main()`.

```js
({
  test: () => {
    const pythonCode = `_Node(_code.replace('\\\\r', '')).has_call("main()")`;
    const result = runPython(pythonCode);
    assert(result);
  }
});
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

--fcc-editable-region--
```
