---
id: 657f425dbab54e11993c80f0
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Até agora, neste projeto você usou um loop `for` para iterar sobre sua string de entrada e convertê-la na saída desejada. Agora você começará a transição de um loop `for` para uma list comprehension.

Comece comentando todas as linhas de código dentro da função `convert_to_snake_case()`. Não os exclua pois serão úteis quando você implementar a lógica usando uma list comprehension.

Lembre-se de adicionar a palavra-chave `pass` ao corpo da função para evitar que o código falhe durante os testes.

# --hints--

Você deve comentar todas as 10 linhas de código dentro da função `convert_to_snake_case()` e nada mais. Adicione `pass` para preencher temporariamente o corpo da função e evitar um erro.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
