---
id: 663b16e62fee463b4caf46e9
title: Passo 18
challengeType: 20
dashedName: step-18
---

# --description--

Após juntar os elementos da lista `snake_cased_char_list`, você precisará remover quaisquer underscores iniciais ou finais da string resultante.Para isso, use o método `strip` com o caractere underscore `_` como argumento.

Chamadas de método podem ser encadeadas, o que significa que o resultado de uma chamada de método pode ser usado como o objeto para outra chamada de método.

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

No exemplo acima, o método `.upper()` está encadeado a `' '.join(words_list)`, portanto `.upper()` é chamado no resultado da chamada `.join()`.

Modifique a declaração `return` encadeando a `''.join(snake_cased_char_list)` uma chamada ao método `.strip()` para remover quaisquer underscores iniciais ou finais.

# --hints--

Você deve modificar a instrução de retorno para incluir o método `strip` com o caractere underscore `_` como argumento.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
