---
id: 5d822fd413a79914d39e98c9
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the CSS Variables Skyline project! Start by adding the `!DOCTYPE html` declaration at the top of the document so the browser knows what type of document it's reading.

# --hints--

يجب أن يحتوي كودك على `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب عليك تضمين مسافة بعد `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب عليك تحديد نوع المستند ليكون `html`.

```js
assert(code.match(/html/gi));
```

يجب عليك إغلاق تعريف `DOCTYPE` باستخدام `>` بعد تحديد النوع.

```js
assert(code.match(/html\s*>/gi));
```

يجب أن يكون تعريف `DOCTYPE` في بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
