---
id: 5d822fd413a79914d39e98cc
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

Within the `head`, nest a `meta` element with a `charset` of `UTF-8`, a `title` element with a title of `City Skyline`, and a `link` element that links your `styles.css` file.

# --hints--

يجب عليك إنشاء عنصر `meta` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('meta'));
```

يجب عليك إعطاء عنصر `meta` خاصية `charset` بقيمة `UTF-8`.

```js
const linkElement = document.querySelector('meta');
const charset = linkElement?.getAttribute("charset").toLowerCase();
assert.strictEqual(charset, 'utf-8');
```

يجب أن يحتوي كودك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون عنصر `title` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

يجب أن يكون لمشروعك العنوان `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

تذكر أن الحالة الحرف سواء كبيرا أو صغيرا, والإملاء مهمان للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

يجب أن يحتوي كودك على عنصر `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

لا يجب عليك تغيير عنصر `head` الحالي. تحقق أنك لم تقم بحذف وسم الإغلاق.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

You should have one void `link` element.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

يجب أن يكون عنصرك `link` داخل عنصرك `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

يجب أن يحتوي عنصرك `link` على سمة `rel` بقيمة `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

يجب أن يحتوي عنصرك `link` على خاصية `href` مع القيمة `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
