---
id: 587d78af367417b2b2512b04
title: Erstellen Sie eine Produkt-Zielseite
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301144
dashedName: build-a-product-landing-page
---

# --description--

**Ziel:** Erstellen Sie eine App, die funktional ähnlich ist wie <a href="https://product-landing-page.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://product-landing-page.freecodecamp.rocks</a>. **Kopieren Sie dieses Demo-Projekt nicht**.

**User Stories:**

1. Ihre Produkt-Zielseite sollte ein `header`-Element mit einem entsprechenden `id="header"` haben
1. Sie sollten ein Bild innerhalb des `header`-Elements mit einem entsprechenden `id="header-img"` sehen (Ein Logo wäre hier ein gutes Bild)
1. Innerhalb des `#header`-Elements sollten Sie ein `nav`-Element mit einem entsprechenden `id="nav-bar"` sehen
1. Sie sollten mindestens drei anklickbare Elemente innerhalb des `nav`-Elements sehen, von denen jedes die Klasse `nav-link` hat
1. Wenn Sie auf eine `.nav-link`-Schaltfläche im `nav`-Element klicken, werden Sie zum entsprechenden Abschnitt der Zielseite geführt
1. Sie können ein eingebettetes Produktvideo mit `id="video"` ansehen
1. Ihre Zielseite hat ein `form`-Element mit einem entsprechenden `id="form"`
1. Innerhalb des Formulars gibt es ein `input`-Feld mit `id="email"`, in das Sie eine E-Mail-Adresse eingeben können
1. Das `#email`-Eingabefeld sollte einen Platzhaltertext haben, der den Nutzenden zeigt, wofür das Feld gedacht ist
1. Das `#email`-Eingabefeld verwendet HTML5-Validierung, um zu bestätigen, dass der eingegebene Text eine E-Mail-Adresse ist
1. Innerhalb des Formulars gibt es eine Absende-`input` mit einem entsprechenden `id="submit"`
1. Wenn Sie auf das `#submit`-Element klicken, wird die E-Mail an eine statische Seite gesendet (verwenden Sie diese Beispiel-URL: `https://www.freecodecamp.com/email-submit`)
1. Die Navigationsleiste sollte immer oben im Viewport sein
1. Ihre Produkt-Zielseite sollte mindestens eine Media Query enthalten
1. Ihre Produkt-Zielseite sollte mindestens einmal CSS Flexbox verwenden

Erfüllen Sie die User Stories und bestehen Sie alle untenstehenden Tests, um dieses Projekt abzuschließen. Verleihen Sie ihm Ihren eigenen Stil. Viel Spaß beim Programmieren!

**Hinweis:** Stellen Sie sicher, dass Sie `<link rel="stylesheet" href="styles.css">` in Ihrem HTML hinzufügen, um Ihr Stylesheet zu verknüpfen und Ihr CSS anzuwenden

# --hints--

Sie sollten ein `header`-Element mit einem `id` von `header` haben.

```js
const el = document.getElementById('header')
assert(!!el && el.tagName === 'HEADER')
```

Sie sollten ein `img`-Element mit einem `id` von `header-img` haben.

```js
const el = document.getElementById('header-img')
assert(!!el && el.tagName === 'IMG')
```

Ihr `#header-img` sollte ein Nachfahre des `#header` sein.

```js
const els = document.querySelectorAll('#header #header-img')
assert(els.length > 0)
```

Ihr `#header-img` sollte ein `src`-Attribut haben.

```js
const el = document.getElementById('header-img')
assert(!!el && !!el.src)
```

Der `#header-img`-Wert Ihres `src` sollte eine gültige URL sein (beginnt mit `http`).

```js
const el = document.getElementById('header-img')
assert(!!el && /^http/.test(el.src))
```

Sie sollten ein `nav`-Element mit einem `id` von `nav-bar` haben.

```js
const el = document.getElementById('nav-bar')
assert(!!el && el.tagName === 'NAV')
```

Ihr `#nav-bar` sollte ein Nachfahre des `#header` sein.

```js
const els = document.querySelectorAll('#header #nav-bar')
assert(els.length > 0)
```

Sie sollten mindestens 3 `.nav-link`-Elemente innerhalb des `#nav-bar` haben.

```js
const els = document.querySelectorAll('#nav-bar .nav-link')
assert(els.length >= 3)
```

Jedes `.nav-link`-Element sollte ein `href`-Attribut haben.

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  if (!el.href) assert(false)
})
assert(els.length > 0)
```

Jedes `.nav-link`-Element sollte mit einem entsprechenden Element auf der Zielseite verknüpft sein (hat ein `href` mit dem Wert der id eines anderen Elements, z. B. `#footer`).

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  const linkDestination = el.getAttribute('href').slice(1)
  if (!document.getElementById(linkDestination)) assert(false)
})
assert(els.length > 0)
```

Sie sollten ein `video`- oder `iframe`-Element mit einem `id` von `video` haben.

```js
const el = document.getElementById('video')
assert(!!el && (el.tagName === 'VIDEO' || el.tagName === 'IFRAME'))
```

Ihr `#video` sollte ein `src`-Attribut haben.

```js
let el = document.getElementById('video')
const sourceNode = el.children;
let sourceElement = null;
if (sourceNode.length) {
  sourceElement = [...video.children].filter(el => el.localName === 'source')[0];
}
if (sourceElement) {
  el = sourceElement;
}
assert(el.hasAttribute('src'));
```

Sie sollten ein `form`-Element mit einem `id` von `form` haben.

```js
const el = document.getElementById('form')
assert(!!el && el.tagName === 'FORM')
```

Sie sollten ein `input`-Element mit einem `id` von `email` haben.

```js
const el = document.getElementById('email')
assert(!!el && el.tagName === 'INPUT')
```

Ihr `#email` sollte ein Nachfahre des `#form` sein.

```js
const els = document.querySelectorAll('#form #email')
assert(els.length > 0)
```

Ihr `#email` sollte das `placeholder`-Attribut mit Platzhaltertext haben.

```js
const el = document.getElementById('email')
assert(!!el && !!el.placeholder && el.placeholder.length > 0)
```

Ihr `#email` sollte HTML5-Validierung verwenden, indem Sie sein `type` auf `email` setzen.

```js
const el = document.getElementById('email')
assert(!!el && el.type === 'email')
```

Sie sollten ein `input`-Element mit einem `id` von `submit` haben.

```js
const el = document.getElementById('submit')
assert(!!el && el.tagName === 'INPUT')
```

Ihr `#submit` sollte ein Nachfahre des `#form` sein.

```js
const els = document.querySelectorAll('#form #submit')
assert(els.length > 0)
```

Ihr `#submit` sollte ein `type` von `submit` haben.

```js
const el = document.getElementById('submit')
assert(!!el && el.type === 'submit')
```

Ihr `#form` sollte ein `action`-Attribut von `https://www.freecodecamp.com/email-submit` haben.

```js
const el = document.getElementById('form')
assert(!!el && el.action === 'https://www.freecodecamp.com/email-submit')
```

Ihr `#email` sollte ein `name`-Attribut von `email` haben.

```js
const el = document.getElementById('email')
assert(!!el && el.name === 'email')
```

Ihr `#nav-bar` sollte immer oben im Viewport sein.

```js
  const timeout = (milliseconds) => new Promise((resolve) => setTimeout(resolve, milliseconds));

  const header = document.getElementById('header');
  const headerChildren = header.children;
  const navbarCandidates = [header, ...headerChildren];
  
  // Return smallest top position of all navbar candidates
  const getNavbarPosition = (candidates = []) => {
    return candidates.reduce(
      (min, candidate) =>
        Math.min(min, Math.abs(candidate?.getBoundingClientRect().top)),
      Infinity
    );
  };
  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the viewport '
  );

  window.scroll(0, 500);
  await timeout(1);

  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the ' +
      'viewport even after scrolling '
  );
    
  window.scroll(0, 0);
```

Ihre Produkt-Zielseite sollte mindestens eine Media Query verwenden.

```js
const htmlSourceAttr = Array.from(document.querySelectorAll('source')).map(el => el.getAttribute('media'))
const cssCheck = new __helpers.CSSHelp(document).getCSSRules('media')
assert(cssCheck.length > 0 || htmlSourceAttr.length > 0);
```

Ihre Produkt-Zielseite sollte mindestens einmal CSS Flexbox verwenden.

```js
const hasFlex = (rule) => ["flex", "inline-flex"].includes(rule.style?.display)
const stylesheet = new __helpers.CSSHelp(document).getStyleSheet()
const cssRules = new __helpers.CSSHelp(document).styleSheetToCssRulesArray(stylesheet)
const mediaRules = new __helpers.CSSHelp(document).getCSSRules('media')
const usesFlex = cssRules.find(rule => hasFlex(rule))
const usesFlexMedia = mediaRules.find(mediaRule => {
  return [...mediaRule.cssRules].find(rule => hasFlex(rule))
})
assert(usesFlex || usesFlexMedia)
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Product Landing Page</title>
  </head>
  <body>
    <header id="header">
      <nav id="nav-bar">
        <img
          id="header-img"
          src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
          max-height="50px"
        />
        <a href="#Features" class="nav-link">Features</a> |
        <a href="#Video" class="nav-link">See our facility!</a> |
        <a href="#Pricing" class="nav-link">Purchase</a>
        <hr />
      </nav>
    </header>
    <main>
      <h1>
        Pokemon Daycare Service
      </h1>
      <section id="Features">
        <h2>What we offer</h2>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">Guaranteed friendly and loving staff!</div>
        </div>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">
            Comfortable environment for Pokemon to explore and play!
          </div>
        </div>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">
            Multiple membership plans to fit your lifestyle!
          </div>
        </div>
      </section>
      <section id="Video">
        <h2>Check us out!</h2>
        A sneak peek into our facility:
        <br />
        <iframe
          id="video"
          width="520"
          height="347"
          src="https://www.youtube.com/embed/Nw-ksH2r6AQ"
          frameborder="0"
          allowfullscreen
          alt="A video tour of our facility"
        >
        </iframe>
      </section>
      <section id="Pricing">
        <h2>Membership Plans</h2>
        <div class="flex-mem">
          <div class="flex-mem-box">
            <font size="+2">Basic Membership</font><br />
            <ul>
              <li>One Pokemon</li>
              <li>Food and berries provided</li>
            </ul>
            <em>$9.99/month</em>
          </div>
          <div class="flex-mem-box">
            <font size="+2">Silver Membership</font><br />
            <ul>
              <li>Up to Three Pokemon</li>
              <li>Food and berries provided</li>
              <li>Grooming and accessories included</li>
            </ul>
            <em>$19.99/month</em>
          </div>
          <div class="flex-mem-box">
            <font size="+2">Gold Membership</font><br />
            <ul>
              <li>Up to six Pokemon!</li>
              <li>Food and berries provided</li>
              <li>Grooming and accessories included</li>
              <li>Personal training for each Pokemon</li>
              <li>Breeding and egg hatching</li>
            </ul>
            <em>$29.99/month</em>
          </div>
        </div>
      </section>
      <form id="form" action="https://www.freecodecamp.com/email-submit">
        <p>Sign up for our newsletter!</p>
        <label for="email"><p>Email:</p><input name="email" id="email" type="email" placeholder="johnsmith@email.com" required></label>
        <input type="submit" id="submit">
      </form>
      <footer>
        <a href="../">Return to Project List</a> |
        <a href="https://www.nhcarrigan.com">Return to HomePage</a>
      </footer>
    </main>
  </body>
</html>
```

```css
body {
  background-color: #3a3240;
  color: white;
}
main {
  max-width: 750px;
  margin: 50px auto;
}
input {
  background-color: #92869c;
}
a:not(.nav-link) {
  color: white;
}
#header-img {
  max-height: 25px;
}
#nav-bar {
  position: fixed;
  width: 100%;
  text-align: center;
  top: 0%;
  background-color: #92869c;
}
h1 {
  text-align: center;
}
body {
  text-align: center;
}
footer {
  text-align: center;
}
#bullet {
  max-height: 25px;
}
.flex-here {
  display: flex;
  justify-content: center;
}
.flex-left {
  height: 25px;
}
.flex-mem {
  display: flex;
  justify-content: center;
}
.flex-mem-box {
  background-color: #92869c;
  border-color: black;
  border-width: 5px;
  border-style: solid;
  margin: 10px;
  padding: 10px;
  color: black;
}
@media (max-width: 350px) {
  #video {
    width: 300;
    height: 200;
  }
}
```
