---
id: 68e76b1ce8216249d04df670
title: 행성 클래스 만들기
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `Planet`라는 이름의 클래스를 만들어야 합니다.
1. `Planet` 클래스는 다음을 수행하는 `__init__` 메서드를 가져야 합니다:
   - 네 개의 매개변수 `self`, `name`, `planet_type`, `star`를 가집니다.
   - 전달된 인수 중 하나라도 문자열 타입이 아니면 메시지 `TypeError`와 함께 `name, planet type, and star must be strings` 예외를 발생시킵니다.
   - 전달된 인수 중 하나라도 빈 문자열이면 메시지 `ValueError`와 함께 `name, planet_type, and star must be non-empty strings` 예외를 발생시킵니다.
   - 전달된 값을 인스턴스 속성 `name`, `planet_type`, `star`에 할당합니다.
1. `Planet` 클래스는 `orbit` 메서드를 가져야 하며, `{name} is orbiting around {star}...` 형식의 문자열을 반환해야 합니다.
1. `Planet` 클래스는 `__str__` 메서드를 가져야 하며, `Planet: {name} | Type: {planet_type} | Star: {star}` 형식의 문자열을 반환해야 합니다.
1. `Planet` 클래스의 인스턴스 세 개를 각각 `planet_1`, `planet_2`, `planet_3`라는 이름으로 만들어야 합니다.
1. 각 행성 객체를 출력하여 `__str__` 메서드의 출력을 확인해야 합니다.
1. 각 행성 객체에서 `orbit` 메서드를 호출하고 결과를 출력해야 합니다.

# --hints--

`Planet`라는 이름의 클래스를 만들어야 합니다.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

`Planet` 클래스는 `__init__` 메서드를 가져야 합니다.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

`__init__` 메서드는 네 개의 매개변수 `self`, `name`, `planet_type`, `star`를 이 순서대로 가져야 합니다.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

`__init__` 메서드는 첫 번째 인수가 문자열 타입이 아니면 메시지 `TypeError`와 함께 `name, planet type, and star must be strings` 예외를 발생시켜야 합니다.

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`__init__` 메서드는 두 번째 인수가 문자열 타입이 아니면 메시지 `TypeError`와 함께 `name, planet type, and star must be strings` 예외를 발생시켜야 합니다.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`__init__` 메서드는 세 번째 인수가 문자열 타입이 아니면 메시지 `TypeError`와 함께 `name, planet type, and star must be strings` 예외를 발생시켜야 합니다.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`__init__` 메서드는 첫 번째 인수가 빈 문자열이면 메시지 `ValueError`와 함께 `name, planet_type, and star must be non-empty strings` 예외를 발생시켜야 합니다.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`__init__` 메서드는 두 번째 인수가 빈 문자열이면 메시지 `ValueError`와 함께 `name, planet_type, and star must be non-empty strings` 예외를 발생시켜야 합니다.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`__init__` 메서드는 세 번째 인수가 빈 문자열이면 메시지 `ValueError`와 함께 `name, planet_type, and star must be non-empty strings` 예외를 발생시켜야 합니다.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

`name`를 `self.name`에 `__init__` 메서드에서 할당해야 합니다.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

`planet_type`를 `self.planet_type`에 `__init__` 메서드에서 할당해야 합니다.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

`star`를 `self.star`에 `__init__` 메서드에서 할당해야 합니다.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

`Planet` 클래스는 `orbit` 메서드를 가져야 합니다.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

`orbit` 메서드는 `{name} is orbiting around {star}...` 형식의 문자열을 반환해야 합니다.

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

`Planet` 클래스는 `__str__` 메서드를 가져야 합니다.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

`__str__` 메서드는 `Planet: {name} | Type: {planet_type} | Star: {star}` 형식의 문자열을 반환해야 합니다.

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

`Planet` 클래스의 인스턴스 세 개를 각각 `planet_1`, `planet_2`, `planet_3`라는 이름으로 만들어야 합니다.

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

각 행성 객체를 출력하여 `__str__` 메서드의 출력을 확인해야 합니다.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

각 행성 객체에서 `orbit` 메서드를 호출하고 결과를 출력해야 합니다.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
