---
id: 68a47dcdbce12193caf77012
title: реалізувати пошук індексу для вставки в відсортований масив
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

У цій лабораторній роботі ви створите функцію, яка повертає найнижчий індекс, за яким значення слід вставити у масив після його сортування за зростанням.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати `getIndexToIns` функцію, яка приймає два аргументи: масив і число.  
2. Ви повинні використовувати метод `sort` для сортування масиву за зростанням.
3. Ваша функція `getIndexToIns` повинна повертати найнижчий індекс, за яким число слід вставити, використовуючи метод `findIndex`.
4. Ваша функція `getIndexToIns` завжди повинна повертати число.

**Підказка:**  
Метод `findIndex` — це вбудований метод масиву в JavaScript. Він приймає функцію зворотного виклику і повертає індекс першого елемента, який задовольняє умову. І `findIndex`, і `sort` — це функції вищого порядку.

**Приклади:**

- `getIndexToIns([1, 2, 3, 4], 1.5)` має повернути `1`, тому що `1.5` більше за `1` (індекс `0`) і менше за `2` (індекс `1`).
- `getIndexToIns([20, 3, 5], 19)` має повернути `2`, тому що після сортування до `[3, 5, 20]` `19` менше за `20` (індекс `2`) і більше за `5` (індекс `1`).


# --hints--

Ви повинні мати `getIndexToIns` функцію.

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` завжди має повертати число.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

Ваша функція має використовувати метод `sort`.

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


Ваша функція має використовувати метод `findIndex`.

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` має повертати `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` має повертати `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` має повертати `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` має повертати `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` має повертати `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` має повертати `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` має повертати `3`

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` має повертати `0`

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
```
