---
id: 662693f82c91a66be46c881b
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

先生が生徒のテストの採点を終え、クラスの平均点を計算する手助けが必要です。

テストの点数の配列を受け取り、平均点を返す`getAverage`関数を完成させてください。

平均点は、すべての点数を合計し、点数の総数で割ることで計算されます。

```md
average = sum of all scores / total number of scores
```

コードをテストできるように、いくつかの関数呼び出しが用意されています。

**ヒント**

- ループを使って`scores`配列を繰り返し、すべての点数を合計できます。
- `length`プロパティを使って点数の総数を取得できます。

# --hints--

あなたの`getAverage`関数は数値を返す必要があります。

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])`は`71.7`を返すべきです。

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])`は`85.4`を返すべきです。

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])`は`92.4`を返すべきです。

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

あなたの`getAverage`関数はテストの点数の平均点を返す必要があります。

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
