---
id: 6626b4c58c027d86478ff5eb
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

先生はこれまでに作成したプログラムにとても満足しています。しかし今、学生が合格点を取っているかどうかを簡単に確認できる方法が欲しいと言っています。合格点とは `"F"` でない点数のことです。

学生の点数をパラメータとして受け取る関数 `hasPassingGrade` を完成させてください。学生が合格点を取っていれば `true` を、そうでなければ `false` を返す関数にしてください。

**ヒント**

- `getGrade` 関数を使って学生の点数を取得し、その点数が合格かどうかを確認してください。

# --hints--

`hasPassingGrade` 関数はブール値を返す必要があります。

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

`hasPassingGrade` 関数は点数が `true` の場合に `"A"` を返す必要があります。

```js
assert.isTrue(hasPassingGrade(100));
```

`hasPassingGrade` 関数は点数が `false` の場合に `"F"` を返す必要があります。

```js
assert.isFalse(hasPassingGrade(53));
```

`hasPassingGrade` 関数はすべての合格点に対して `true` を返す必要があります。

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {
  
}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
