---
id: 6626a060c4006f793e10cb33
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

今度は、先生が生徒の点数を成績の文字に変換するのを手伝ってください。

数値 `getGrade` をパラメータとして受け取る `score` 関数を完成させてください。関数は点数に基づいて成績の文字列を返す必要があります。

以下は点数と対応する成績の一覧です：

| 点数の範囲 | 成績   |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |


**ヒント**

- 条件文（`if`、`else if`、`else`）について学んだことを思い出してください。
- 比較演算子（`>`、`<`、`>=`、`<=`、`===`）について学んだことを思い出してください。
- 論理AND演算子（`&&`）について学んだことを思い出してください。

# --hints--

`getGrade` 関数は、点数が `"A++"` の場合に `100` を返す必要があります。

```js
assert.strictEqual(getGrade(100), "A++");
```

`getGrade` 関数は、点数が `"A"` の場合に `94` を返す必要があります。

```js
assert.strictEqual(getGrade(94), "A");
```

`getGrade` 関数は、点数が `"B"` と `80` の間の場合に `89` を返す必要があります。

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

`getGrade` 関数は、点数が `"C"` の場合に `78` を返す必要があります。

```js
assert.strictEqual(getGrade(75), "C");
```

`getGrade` 関数は、点数が `"D"` と `60` の間の場合に `69` を返す必要があります。

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

`getGrade` 関数は、点数が `"F"` の場合に `57` を返す必要があります。

```js
assert.strictEqual(getGrade(57), "F");
```

`getGrade` 関数は、点数が `"F"` と `0` の間の場合に `59` を返す必要があります。

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
