---
id: 6626b8dcf5057f896f948440
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

先生は必要な情報をすべて手に入れたので、結果を学生にメッセージで伝えたいと考えています。

`studentMsg` 関数を `totalScores` と `studentScore` をパラメータとして完成させてください。この関数は学生へのメッセージを表す文字列を返す必要があります。

学生がコースに合格した場合、文字列は次の形式に従うべきです：

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

学生がコースに不合格だった場合、文字列は次の形式に従うべきです：

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

`average-goes-here` をクラスの総合スコアの平均に置き換えてください。  
`grade-goes-here` を学生の成績に置き換えてください。

**ヒント**

- クラスの平均を取得するには `getAverage` 関数を使ってください。  
- 学生の成績を取得するには `getGrade` 関数を使ってください。  
- メッセージを作成するには文字列連結（`+`）を使ってください。  
- メッセージ内の句読点やスペースに注意してください。  

# --hints--

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` は次のメッセージを返すべきです：`"Class average: 71.7. Your grade: F. You failed the course."`。

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` は次のメッセージを返すべきです：`"Class average: 50.8. Your grade: A++. You passed the course."`。

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

あなたの `studentMsg` 関数は学生のスコアとクラスの平均に基づいて正しいメッセージを返すべきです。

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
