---
id: 6564a9fe51964c229d5b7f4c
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Du benötigst einen Parameter, der die zu sortierende Datensammlung angibt. Create a parameter called `array` in the `merge_sort` function.

# --hints--

Deine `merge_sort`-Funktion solte einen einzigen Parameter: `array` verwenden.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
