---
id: 66a833e88d08593618f22285
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro de tu elemento `head`, anida un elemento `meta` con el atributo `charset` establecido al valor `"UTF-8"`.

Debajo de ese elemento `meta`, agrega un elemento `title`. El texto del elemento `title` debe ser `Hotel Feedback Form`.

# --hints--

Debes tener un elemento `meta`.

```js
assert.isNotNull(document.querySelector("meta"));
```

El elemento `meta` es un elemento vacío, no debe tener una etiqueta final `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Tu etiqueta `meta` debe tener un atributo `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Tu atributo `charset` debe tener un valor de `"UTF-8"`.

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

Tu elemento `meta` debe estar anidado dentro de tu elemento `head`.

```js
assert.match(code, /<head>.*\s*<meta.*charset.*>.*\s*<\/head>/si);
```

Debes tener una etiqueta `title` de apertura.

```js
assert.match(code, /<title>/i);
```

Debes tener una etiqueta `title` de cierre.

```js
assert.match(code, /<\/title>/i);
```

Tu elemento `title` debe estar anidado dentro del elemento `head`.

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

Tu elemento `title` debe tener el texto `Hotel Feedback Form`. Talvez debas revisar tu ortografía.

```js
const titleText = document.querySelector('title')?.innerText
assert.strictEqual(titleText, "Hotel Feedback Form");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    --fcc-editable-region--
    <head>

    </head>
    --fcc-editable-region--
</html>
```
