---
id: 66a83bdcf425e7446900b7c4
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

En los vídeos de la lección anterior, aprendiste a trabajar con el elemento `form` así:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

El atributo `action` se utiliza para especificar dónde se deben enviar los datos del formulario cuando este se envíe.

El atributo `method` se utiliza para especificar el método HTTP a usar al enviar los datos del formulario. Los métodos más comunes son `GET` y `POST`.

**NOTA**: Aprenderás cómo funcionan los métodos HTTP en módulos posteriores.

Dentro de tu elemento `main`, agrega un elemento `form` con un atributo `action` configurado en `"https://hotel-feedback.freecodecamp.org"` y un atributo `method` configurado en `"POST"`.

# --hints--

Debes tener una etiqueta de apertura de `form`.

```js
assert.match(code, /<form.*>/i);
```

Debes tener una etiqueta de cierre de `form`.

```js
assert.match(code, /<\/form>/i);
```

Tu elemento `form` debe tener un atributo `action` configurado en `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

Tu elemento `form` debe tener un atributo `method` configurado en `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--

    --fcc-editable-region--
    </main>
  </body>
</html>
```
