---
id: 66a8347f97a33e36ffc81b9f
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Para comenzar a crear contenido real, añade un elemento `body` debajo del elemento `head`.

# --hints--

Debe haber un tag de apertura para `body`.

```js
assert.match(code, /<body>/i);
```

Debe haber un tag de cierre para `body`.

```js
assert.match(code, /<\/body>/i);
```

No debes modificar el elemento `head`. Asegúrese de que no eliminaste tu etiqueta de cierre.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Tu elemento `body` debe estar después del elemento `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>

  --fcc-editable-region--
</html>
```
