---
id: 66a83fec026a7a4631e084d2
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Cuando se trabaja con elementos `fieldset`, es común usar una leyenda para describir el grupo de entradas. Puedes usar el elemento `legend` para esto.

Aquí hay un ejemplo de uso de un elemento `legend`:

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

Dentro de tu elemento `fieldset`, añade un elemento `legend` con el texto `Personal Information`.

# --hints--

Deberías tener una etiqueta de apertura `legend`.

```js
assert.match(code, /<legend>/);
```

Deberías tener una etiqueta de cierre `legend`.

```js
assert.match(code, /<\/legend>/);
```

Tu elemento `legend` debería estar dentro del elemento `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

Tu elemento `legend` debería contener el texto `Información Personal`.

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText, 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>

        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
