---
id: 66a8290a27c2c625e2355042
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, practicarás trabajar con formularios HTML construyendo un Formulario de Retroalimentación del Hotel.

Comienza agregando el `<!DOCTYPE html>` seguido de un elemento `html` con un atributo `lang` de `en`.

Dentro de tu elemento `html`, agrega un elemento `head`.


# --hints--

Deberías tener el `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Deberías tener una etiqueta `html` de apertura con el idioma configurado en inglés.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Debes tener una etiqueta de cierre `html`.

```js
assert.match(code, /<\/html>/i);
```

Tu `DOCTYPE` debería estar antes del elemento `html`.

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

Debe haber un tag de apertura para `head`.

```js
assert.match(code, /<head>/i);
```

Debe haber un tag de cierre para `head`.

```js
assert.match(code, /<\/head>/i);
```

Tu etiqueta de apertura `head` debería venir antes de la etiqueta de cierre `head`.

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

Tu elemento `head` debería estar dentro del elemento `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[.\n\s]*<head>[.\n\s]*<\/head>[.\n\s]*<\/html>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
