---
id: 66a83e5e491625454b6f62c3
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Los formularios consisten en `inputs` donde los usuarios pueden ingresar sus datos. Puedes agrupar entradas relacionadas usando el elemento `fieldset`.

Aquí hay un ejemplo de uso de un elemento `fieldset`:

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

Dentro de tu elemento `form`, añade un elemento `fieldset`.

# --hints--

Deberías tener una etiqueta de apertura de `fieldset`.

```js
assert.match(code, /<fieldset>/i);
```

Deberías tener una etiqueta de cierre de `fieldset`.

```js
assert.match(code, /<\/fieldset>/i);
```

Tu elemento `fieldset` debería estar dentro del elemento `form`.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">

      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
