---
id: 66a84111965a0c46df6bbd0a
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

In the previous lectures, you learned how to associate a `label` element with an `input` like this:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

El atributo `for` en el elemento `label` debe coincidir con el atributo `id` en el elemento `input`. Esto se conoce como una asociación explícita.

Debajo de su elemento `legend`, agregue un elemento `label` con el texto de `Name (required):`. Establezca su atributo `for` al valor de `"full-name"`.

Luego, debajo de su elemento `label`, agregue un elemento `input` sin atributos. En los próximos pasos, agregarás los atributos necesarios.

# --hints--

Debe tener una etiqueta de apertura `label`.

```js
assert.match(code, /<label.*>/);
```

Debe tener una etiqueta de cierre `label`.

```js
assert.match(code, /<\/label>/);
```

El elemento `label` debe estar ubicado debajo del `legend` dentro del elemento `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

El elemento `label` debe tener un atributo `for` con el valor de `"full-name"`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

El elemento `label` debe tener el texto de `Name (required):`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText, 'Name (required):');
```

Debe tener un elemento `input` debajo de su elemento `label`.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>

        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
