---
id: 66a95d0eff8168747805f1f3
title: Paso 22
challengeType: 0
dashedName: step-22
---

# --description--

Cuando deseas que un usuario seleccione múltiples opciones de una lista, puedes usar casillas de verificación.

Aquí hay un ejemplo de cómo trabajar con casillas de verificación al tratar con opciones de alimentos:

```html
<fieldset>
  <legend>Food Options</legend>
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="burger">Burger</label>
  <input type="checkbox" id="burger" name="food" value="burger">
</fieldset>
```

El atributo `value` se utiliza para especificar el valor que se enviará al servidor cuando el formulario se envíe.

Debajo de tu elemento `legend`, agrega un elemento `label` con el texto de `anuncios de redes sociales`. El atributo `for` debe establecerse en `"ads"`.

Debajo de tu elemento `label`, agrega un `input` de tipo checkbox con los atributos `id`, `name` y `value` establecidos en `"ads"`.

# --hints--

Debes tener un elemento `label` debajo de tu `legend`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + label"));
```

Tu elemento `label` debe tener el texto de `anuncios de redes sociales`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + label")?.textContent, "Social Media Ads");
```

Tu elemento `label` debe tener un atributo `for` establecido en `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + label")?.getAttribute("for"), "ads");
```

Debes tener un `input` de tipo checkbox debajo de tu `label`.

```js
const inputElement = document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']");
assert.strictEqual(inputElement?.parentElement.tagName, "FIELDSET");
```

Tu casilla de verificación debe tener un atributo `id` establecido en `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.id, "ads");
```

Tu casilla de verificación debe tener un atributo `name` establecido en `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.name, "ads");
```

Tu casilla de verificación debe tener un atributo `value` establecido en `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.value, "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        --fcc-editable-region--
        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
