---
id: 66a937e74920ba68ebe5e86d
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

In the previous lectures, you learned how to work with the `placeholder` and `required` attributes like this:

```html
<input type="text" placeholder="Ex. John Doe" required>
```

Para tu elemento de `input` existente, añade un atributo `placeholder` con el valor de `"Ej. John Doe"`.

Además, añade el atributo `required` al elemento de `input`.

# --hints--

No debes eliminar el elemento de `input` existente.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Tu `input` debería tener un atributo `placeholder` con el valor de `"Ej. John Doe"`.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'Ex. John Doe');
```

Tu `input` debería tener un atributo `required`.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
