---
id: 66a9521bc70162712caf118d
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Si quieres que los usuarios seleccionen una opción de una lista de opciones, puedes usar un conjunto de botones de opción.

Aquí tienes un ejemplo de dos botones de opción:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

En este ejemplo, los botones de opción están agrupados utilizando el mismo valor del atributo `name`. Esto significa que solo se puede seleccionar un botón de opción a la vez.

Debajo de tu elemento `legend`, agrega un elemento `label` con el texto `Yes` y un atributo `for` configurado a `"yes-option"`.

Debajo de tu elemento `label`, agrega un botón de opción con el `id` configurado a `"yes-option"`, y el atributo `name` configurado a `"hotel-stay"`.

# --hints--

Deberías tener un elemento `label` debajo de tu elemento `legend`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + label'));
```

Tu elemento `label` debería tener un atributo `for` configurado a `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + label')?.htmlFor, 'yes-option');
```

Tu elemento `label` debería tener el texto `Yes`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + label')?.innerText, 'Yes');
```

Deberías tener un botón de opción debajo de tu elemento `label`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) label + input[type="radio"]'));
```

Tu botón de opción debería tener un atributo `id` configurado a `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input[type="radio"]')?.id, 'yes-option');
```

Tu botón de opción debería tener un atributo `name` configurado a `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input[type="radio"]')?.name, 'hotel-stay');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        --fcc-editable-region--
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
