---
id: 642db8c409d9991d0b3b2f0d
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Il tuo progetto inizia con un contenitore HTML di base e il CSS corrispondente. Il tuo primo compito sarà quello di generare programmaticamente le celle per il foglio di calcolo.

L'oggetto globale `window` rappresenta la finestra (o scheda) del browser. Ha una proprietà `onload` che consente di definire il comportamento quando la finestra ha caricato l'intera pagina, inclusi fogli di stile e gli script.

Inizia impostando la proprietà `onload` di `window` su una funzione freccia senza parametri. In the function, declare a `container` variable and assign it the value of getting the element by the `id` of `"container"`.

# --hints--

Dovresti accedere alla proprietà `onload` dell'oggetto `window`.

```js
assert.match(code, /window\.onload/);
```

Dovresti impostare la proprietà `onload` su una funzione.

```js
assert.isFunction(window.onload);
```

Dovresti usare la sintassi freccia.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

La funzione `onload` non dovrebbe prendere alcun parametro.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

Dovresti dichiarare una variabile `container` nella funzione `onload`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

La variabile `container` deve essere dichiarata con `const`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

Dovresti usare `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

You should get the element with the `id` of `"container"`.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

Dovresti assegnare l'elemento `#container` a `container`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
