---
id: 66d75dd0aa65a71600dc669b
title: Ein Programm zur Bestandsverwaltung erstellen
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

In diesem Labor erstellen Sie ein Programm zur Bestandsverwaltung, mit dem Sie Produkte zum Bestand hinzufügen, aktualisieren, finden und entfernen können. Sie verwenden ein Array von Objekten, um Ihren Bestand darzustellen, wobei jedes Objekt die Schlüssel name und quantity besitzt.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein leeres Array namens `inventory` deklarieren, das Produktobjekte speichert, die einen Schlüssel `name` mit dem Wert eines Strings in Kleinbuchstaben und einen Schlüssel `quantity` mit dem Wert einer Ganzzahl besitzen.
1. Sie sollten eine Funktion namens `findProductIndex` erstellen, die den Produktnamen als Argument nimmt und den Index des entsprechenden Produktobjekts im `inventory` Array zurückgibt. Die Funktion sollte immer die Kleinbuchstabenform des Produktnamens für die Suche verwenden. Wenn das Produkt nicht gefunden wird, sollte die Funktion `-1` zurückgeben.
1. Sie sollten eine Funktion namens `addProduct` erstellen, die ein Produktobjekt als Argument nimmt.
1. Wenn das Produkt bereits im Bestand vorhanden ist, sollte die `addProduct` Funktion den Wert von `quantity` aktualisieren, indem sie die übergebene Menge zur aktuellen `quantity` addiert, und den Produktnamen gefolgt von einem Leerzeichen und `quantity updated` in der Konsole ausgeben.
1. Wenn das Produkt nicht im Bestand vorhanden ist, sollte die `addProduct` Funktion das Produkt zum `inventory` Array hinzufügen und den Produktnamen gefolgt von einem Leerzeichen und `added to inventory` in der Konsole ausgeben.
1. Sie sollten eine Funktion namens `removeProduct` erstellen, die den Produktnamen und die Menge als Argumente nimmt.
1. Die `removeProduct` Funktion sollte die übergebene Menge vom entsprechenden Produktobjekt im Bestand subtrahieren und den String `Remaining <product-name> pieces: <product-quantity>` in der Konsole ausgeben, wobei `<product-name>` durch den Produktnamen und `<product-quantity>` durch die Produktmenge ersetzt wird.
1. Wenn die Menge nach der Subtraktion null ist, sollte `removeProduct` das Produktobjekt aus dem Bestand entfernen. Wenn die Menge im Bestand nicht ausreicht, um die Subtraktion durchzuführen, sollte die `removeProduct` Funktion den String `Not enough <product-name> available, remaining pieces: <product-quantity>` in der Konsole ausgeben.
1. Wenn das zu entfernende Produkt nicht im Bestand vorhanden ist, sollte die `removeProduct` Funktion `<product-name> not found` in der Konsole ausgeben.

**Hinweis:** Um Konflikte zu vermeiden, behalten Sie beim Ausführen der Tests nur die in den User Stories genannten Ausgaben bei.

# --hints--

Sie sollten ein leeres Array namens `inventory` deklarieren.

```js
const explorer = await __helpers.Explorer(code);
const inventory = explorer.variables.inventory;
assert.exists(inventory);
assert.isTrue(inventory.value.matches('[]'));
```

Sie sollten eine Funktion namens `findProductIndex` haben.

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")` sollte den Index des Objekts zurückgeben, dessen `name` gleich `"flour"` im `inventory` Array ist.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")` sollte den Index des Objekts zurückgeben, dessen `name` gleich `"flour"` im `inventory` Array ist.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

`findProductIndex("Flour")` sollte `-1` zurückgeben, wenn kein Objekt mit `name` gleich `"flour"` im `inventory` Array gefunden wird.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

Sie sollten eine Funktion namens `addProduct` haben.

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})` sollte `5` zum `quantity` Wert des Objekts hinzufügen, dessen `name` gleich `"flour"` im `inventory` Array ist.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})` sollte `flour quantity updated` ausgeben, wenn ein Objekt mit `name` gleich `"flour"` im `inventory` Array gefunden wird.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})` sollte `{name: "flour", quantity: 5}` zum `inventory` Array hinzufügen, wenn kein Objekt mit `name` gleich `"flour"` im Bestand gefunden wird.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})` sollte `flour added to inventory` ausgeben, wenn kein Objekt mit `name` gleich `"flour"` im Bestand gefunden wird.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

Sie sollten eine Funktion namens `removeProduct` erstellen.

```js
assert.isFunction(removeProduct);
```

`removeProduct("FLOUR", 5)` sollte `flour not found` ausgeben, wenn kein Objekt mit `name` gleich `"flour"` im `inventory` Array gefunden wird.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)` sollte `5` vom `quantity` Wert des Objekts subtrahieren, dessen `name` gleich `"flour"` im `inventory` Array ist.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)` sollte `Remaining flour pieces: 15` in der Konsole ausgeben, wenn `inventory` gleich `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]` ist.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

Wenn der `quantity` nach der Subtraktion null ist, sollte `removeProduct` das Produktobjekt aus dem Bestand entfernen.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)` sollte `Not enough flour available, remaining pieces: 5` ausgeben, wenn `inventory` gleich `[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]` ist.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
