---
id: 5a24c314108439a4d4036157
title: Escrever um contador com Redux
challengeType: 6
forumTopicId: 301453
dashedName: write-a-counter-with-redux
---

# --description--

Agora vocĂȘ aprendeu todos os princĂ­pios fundamentais do Redux! VocĂȘ viu como criar criadores de aĂ§Ă”es e aĂ§Ă”es, criar uma store do Redux, despachar suas aĂ§Ă”es para a store e projetar atualizaĂ§Ă”es de estado com reducers puros. VocĂȘ atĂ© viu como gerenciar o estado complexo com a composiĂ§ĂŁo do reducer e a lidar com aĂ§Ă”es assĂ­ncronas. Estes exemplos sĂŁo simplistas, mas estes conceitos sĂŁo os princĂ­pios fundamentais do Redux. Se vocĂȘ os entendeu bem, vocĂȘ estĂĄ pronto para comeĂ§ar a construir seu prĂłprio aplicativo Redux. Os prĂłximos desafios cobrem alguns detalhes sobre imutabilidade do `state`, mas primeiro, aqui estĂĄ uma revisĂŁo de tudo o que vocĂȘ aprendeu atĂ© agora.

# --instructions--

Nesta liĂ§ĂŁo, vocĂȘ implementarĂĄ um contador simples com o Redux do zero. O bĂĄsico Ă© fornecido no editor de cĂłdigo, mas vocĂȘ terĂĄ que preencher os detalhes! Use os nomes que sĂŁo fornecidos e defina os criadores de aĂ§ĂŁo `incAction` e `decAction`, os tipos de aĂ§ĂŁo `counterReducer()`, `INCREMENT` e `DECREMENT` e, finalmente, a `store` do Redux. Depois de ter terminado, vocĂȘ deve ser capaz de despachar as aĂ§Ă”es `INCREMENT` ou `DECREMENT` para incrementar ou decrementar o estado mantido na `store`. Boa sorte ao construir seu primeiro aplicativo Redux!

# --hints--

O criador de aĂ§ĂŁo `incAction` deve retornar um objeto de aĂ§ĂŁo com `type` igual ao valor de `INCREMENT`

```js
assert(incAction().type === INCREMENT);
```

O criador de aĂ§ĂŁo `decAction` deve retornar um objeto de aĂ§ĂŁo com `type` igual ao valor de `DECREMENT`

```js
assert(decAction().type === DECREMENT);
```

Executar `store.getState()` deve retornar um nĂșmero

```js
assert(typeof store.getState() === 'number');
```

A store do Redux deve inicializar com o `state` igual a 0.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    return _store.getState() === 0;
  })()
);
```

Despachar `incAction` na store do Redux deve incrementar o `state` em 1.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(incAction());
    const incState = _store.getState();
    return initialState + 1 === incState;
  })()
);
```

Despachar `decAction` na store do Redux deve decrementar o `state` em 1.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(decAction());
    const decState = _store.getState();
    return initialState - 1 === decState;
  })()
);
```

`counterReducer` deve ser uma funĂ§ĂŁo

```js
assert(typeof counterReducer === 'function');
```

# --seed--

## --seed-contents--

```js
const INCREMENT = null; // Define a constant for increment action types
const DECREMENT = null; // Define a constant for decrement action types

const counterReducer = null; // Define the counter reducer which will increment or decrement the state based on the action it receives

const incAction = null; // Define an action creator for incrementing

const decAction = null; // Define an action creator for decrementing

const store = null; // Define the Redux store here, passing in your reducers
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const incAction = () => {
  return {
    type: INCREMENT
  }
};

const decAction = () => {
  return {
    type: DECREMENT
  }
};

const store = Redux.createStore(counterReducer);
```
