---
id: 68af276c60e009597dd861e7
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Infine, crea un'ultima istruzione `if` per verificare se `duration` non è maggiore di `0` e genera un'eccezione `ValueError` con il messaggio `Duration must be positive`.

# --hints--

Dovresti generare un'eccezione `ValueError` con il messaggio `Duration must be positive` quando `duration` non è un numero positivo.

```js
({
  test: () => runPython(`
  _msg = 'Duration must be positive'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1999, 'The Wachowskis', 0)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 1999, 'The Wachowskis', -1)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Non dovresti generare un'eccezione `ValueError` quando `duration` è un numero valido.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid duration'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
