---
id: 68af272f60e009597dd861e1
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop creerai un catalogo multimediale. Inizia definendo una classe chiamata `Movie`.

Dai alla tua classe un metodo `__init__` con i seguenti parametri: `self`, `title`, `year`, `director` e `duration`. All'interno del metodo `__init__`, assegna i parametri `title`, `year`, `director` e `duration` ad attributi dell'istanza con lo stesso nome.

# --hints--

Dovresti creare una classe chiamata `Movie`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

La tua classe `Movie` dovrebbe avere un metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

Il tuo metodo `__init__` dovrebbe avere i seguenti parametri: `self`, `title`, `year`, `director` e `duration`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

Dovresti assegnare `title` a `self.title` all'interno del tuo metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

Dovresti assegnare `year` a `self.year` all'interno del tuo metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

Dovresti assegnare `director` a `self.director` all'interno del tuo metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

Dovresti assegnare `duration` a `self.duration` all'interno del tuo metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
