---
id: 68af274560e009597dd861e3
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Ora, crea un'istanza di `Movie` passando argomenti a tua scelta e assegnala a una variabile chiamata `movie1`. Poi, stampa `movie1` sulla console.

# --hints--

Dovresti avere una variabile chiamata `movie1`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

Dovresti assegnare un'istanza di `Movie` alla variabile `movie1`.

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

Dovresti stampare `movie1` sulla console.

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
```
