---
id: 68af275860e009597dd861e5
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Il primo film è stato realizzato nel 1895. Crea un'altra istruzione `if` che verifichi se `year` è minore di `1895` e sollevi un `ValueError` con il messaggio `Year must be 1895 or later`.

# --hints--

Dovresti sollevare un `ValueError` con il messaggio `Year must be 1895 or later` quando `year` è minore di `1895`.

```js
({
  test: () => runPython(`
  _msg = 'Year must be 1895 or later'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1894, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 0, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Non dovresti sollevare un `ValueError` quando `year` è un numero valido.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1895, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid year'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
