---
id: 68af273960e009597dd861e2
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Vuoi una visualizzazione leggibile per eseguire il debug e stampare. Crea un metodo `__str__` che restituisca una stringa con il formato `title (year) - duration min, director` (dove `title`, `year`, `duration` e `director` devono essere sostituiti con i valori degli attributi corrispondenti) in modo che stampare un film dia un bel riepilogo.

# --hints--

La tua classe `Movie` dovrebbe avere un metodo `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

Il tuo metodo `__str__` dovrebbe restituire una stringa nel formato `title (year) - duration min, director`, dove `title`, `year`, `duration` e `director` devono essere sostituiti con i valori degli attributi corrispondenti.

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    
--fcc-editable-region--
```
