---
id: 68af274e60e009597dd861e4
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Nei passaggi precedenti, hai creato un oggetto movie di base. Ora, aggiungerai della validazione per assicurarti che i film non possano essere istanziati con attributi non validi.

All'inizio del metodo `__init__`, crea un'istruzione `if` che verifichi se `title` è vuoto o contiene solo spazi bianchi (puoi usare il metodo `strip` per costruire la condizione).

Se è così, genera un'eccezione `ValueError` con il messaggio `Title cannot be empty`.

# --hints--

Dovresti generare un'eccezione `ValueError` con il messaggio `Title cannot be empty` quando `title` è vuoto o contiene solo spazi bianchi.

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Non dovresti generare un'eccezione `ValueError` quando `title` è una stringa valida.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
