---
id: 68b6cc518177b0a11c3fe3b1
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Prima del blocco `try`, crea un'istanza di `MediaCatalogue` e assegnala a una variabile chiamata `catalogue`.

Poi, usa il metodo `add` per aggiungere `movie1` al catalogo. Infine stampa `catalogue` sulla console.

# --hints--

Dovresti assegnare un'istanza di `MediaCatalogue` a una variabile chiamata `catalogue` prima del tuo blocco `try`.

```js
({
  test: () => runPython(`assert _Node(_code).has_stmt("catalogue = MediaCatalogue()")`)
})
```

Dovresti avere `catalogue.add(movie1)` nel tuo blocco `try`.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().has_stmt("catalogue.add(movie1)")`)
})
```

Dovresti stampare `catalogue` alla fine del tuo blocco `try`.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().is_ordered("catalogue.add(movie1)", "print(catalogue)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result
--fcc-editable-region--


try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)

--fcc-editable-region--
except ValueError as e:
    print(f'Validation Error: {e}')
```
