---
id: 68b702892910ab7b495de081
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Ti serve un modo adeguato per rappresentare visivamente il catalogo. Quindi crea un metodo `__str__` nella tua classe `MediaCatalogue`.

Se la lista `items` è vuota, assicurati che il metodo `__str__` restituisca `Media Catalogue (empty)`.

# --hints--

Dovresti avere un metodo `__str__` nella tua classe `MediaCatalogue`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").has_function("__str__")`)
})
```

Il tuo metodo `__str__` dovrebbe restituire `Media Catalogue (empty)` quando `self.items` è vuota.

```js
({
  test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    print(MediaCatalogue())
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Media Catalogue (empty)" in output
  `)
})
```

Il tuo metodo `__str__` non dovrebbe restituire `Media Catalogue (empty)` quando `self.items` non è vuota.

```js
({
  test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
    _catalogue = MediaCatalogue()
    _catalogue.add(Movie('The Matrix', 1999, 'The Wachowskis', 136))
    try:
      print(_catalogue)
    except TypeError:
      pass
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Media Catalogue (empty)" not in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
