---
id: 68b804f3be853a355f055fef
title: Passaggio 18
challengeType: 20
dashedName: step-18
---

# --description--

La funzione `super` serve per fare riferimento alla classe genitore senza nominarla esplicitamente. Per esempio, `super().__init__(arg1, arg2, ...)` chiama il metodo `__init__` definito nella classe genitore.

Al momento, hai effettivamente sovrascritto il metodo `__init__` ereditato da `Movie`. Questo significa che un oggetto `TVSeries` avrà solo un attributo `seasons` e un attributo `total_episodes`. Questo non è quello che vuoi.

Per mantenere tutto il codice definito nel metodo `Movie` della classe `__init__`, usa `super` per chiamare il `__init__` del genitore. Nota che dovrai passare gli argomenti richiesti dal metodo. In questo caso, `title`, `year`, `director` e `duration`.

# --hints--

Dovresti chiamare `super().__init__(title, year, director, duration)` all'interno del tuo metodo `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("super().__init__(title, year, director, duration)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        
        self.seasons = seasons
        self.total_episodes = total_episodes
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
