---
id: 587d7daf367417b2b2512b80
title: プロトタイプを変更するときはコンストラクタプロパティを設定するのを忘れないでください
challengeType: 1
forumTopicId: 301323
dashedName: remember-to-set-the-constructor-property-when-changing-the-prototype
---

# --description--

プロトタイプを手動で新しいオブジェクトに設定するときに重要な副作用があります。それは`constructor`プロパティが消えてしまうことです！このプロパティはどのコンストラクタ関数がインスタンスを作成したかを確認するために使えますが、プロパティが上書きされてしまったため、誤った結果を返します。

```js
duck.constructor === Bird;
duck.constructor === Object;
duck instanceof Bird;
```

順に、これらの式は`false`、`true`、`true`と評価されます。

これを修正するには、プロトタイプを手動で新しいオブジェクトに設定するときは、必ず`constructor`プロパティを定義してください。

```js
Bird.prototype = {
  constructor: Bird,
  numLegs: 2,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name); 
  }
};
```

# --instructions--

`constructor` `Dog`の`prototype`プロパティを定義してください。

# --hints--

`Dog.prototype`は`constructor`プロパティを設定する必要があります。

```js
assert(Dog.prototype.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
Dog.prototype = {

  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
  constructor: Dog,
  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```
