---
id: 587d7db2367417b2b2512b8b
title: 即時実行関数式（IIFE）を理解する
challengeType: 1
forumTopicId: 301328
dashedName: understand-the-immediately-invoked-function-expression-iife
---

# --description--

JavaScriptでよく使われるパターンの一つに、関数を宣言すると同時に実行する方法があります。

```js
(function () {
  console.log("Chirp, chirp!");
})();
```

これは無名関数式で、すぐに実行され、`Chirp, chirp!`を即座に出力します。

関数には名前がなく、変数にも格納されていないことに注意してください。関数式の末尾にある丸括弧()が、関数を即時に実行または呼び出す原因となっています。このパターンは<dfn>即時実行関数式</dfn>または<dfn>IIFE</dfn>として知られています。

# --instructions--

関数`makeNest`を書き換えて、その呼び出しを削除し、代わりに無名の即時実行関数式（IIFE）にしてください。

# --hints--

関数は無名である必要があります。

```js
assert(/\((function|\(\))(=>|\(\)){?/.test(__helpers.removeJSComments(code).replace(/\s/g, '')));
```

関数式の末尾に丸括弧を付けて、即座に呼び出すようにしてください。

```js
assert(/\(.*(\)\(|\}\(\))\)/.test(__helpers.removeJSComments(code).replace(/[\s;]/g, '')));
```

# --seed--

## --seed-contents--

```js
function makeNest() {
  console.log("A cozy nest is ready");
}

makeNest();
```

# --solutions--

```js
(function () {
  console.log("A cozy nest is ready");
})();
```

---

```js
(function () {
  console.log("A cozy nest is ready");
}());
```

---

```js
(() => {
  console.log("A cozy nest is ready");
})();
```

---

```js
(() =>
  console.log("A cozy nest is ready")
)();
```
