---
id: 587d7db1367417b2b2512b85
title: 子のプロトタイプを親のインスタンスに設定する
challengeType: 1
forumTopicId: 301325
dashedName: set-the-childs-prototype-to-an-instance-of-the-parent
---

# --description--

前のチャレンジでは、スーパタイプ（または親）`Animal`から振る舞いを継承する最初のステップ、すなわち`Animal`の新しいインスタンスを作成する方法を学びました。

このチャレンジでは次のステップを扱います：サブタイプ（または子）`prototype`の`Bird`を`Animal`のインスタンスに設定してください。

```js
Bird.prototype = Object.create(Animal.prototype);
```

`prototype`はオブジェクトを作るための「レシピ」のようなものだということを覚えておいてください。ある意味で、`Bird`のレシピには今や`Animal`の重要な「材料」がすべて含まれています。

```js
let duck = new Bird("Donald");
duck.eat();
```

`duck`は`Animal`のすべてのプロパティを継承し、`eat`メソッドも含まれています。

# --instructions--

`Dog`のインスタンスが`Animal`から継承するようにコードを修正してください。

# --hints--

`Dog.prototype`は`Animal`のインスタンスである必要があります。

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }

// Only change code below this line


let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);

let beagle = new Dog();
beagle.eat();
```
