---
id: 587d7db1367417b2b2512b87
title: 継承後にメソッドを追加する
challengeType: 1
forumTopicId: 301315
dashedName: add-methods-after-inheritance
---

# --description--

スーパタイプのコンストラクタ関数から`prototype`オブジェクトを継承するコンストラクタ関数は、継承されたメソッドに加えて独自のメソッドを持つことができます。

例えば、`Bird`は`prototype`から`Animal`を継承するコンストラクタです：

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
Bird.prototype.constructor = Bird;
```

`Animal`から継承されたものに加えて、`Bird`オブジェクトに固有の動作を追加したいです。ここで、`Bird`は`fly()`関数を持ちます。関数はどのコンストラクタ関数でも同じ方法で`Bird's` `prototype`に追加されます：

```js
Bird.prototype.fly = function() {
  console.log("I'm flying!");
};
```

これで`Bird`のインスタンスは`eat()`と`fly()`の両方のメソッドを持つことになります：

```js
let duck = new Bird();
duck.eat();
duck.fly();
```

`duck.eat()`はコンソールに文字列`nom nom nom`を出力し、`duck.fly()`は文字列`I'm flying!`を出力します。

# --instructions--

`Dog`オブジェクトが`Animal`から継承し、`Dog`の`prototype`コンストラクタが`Dog`に設定されるように必要なコードをすべて追加してください。それから、`bark()`オブジェクトに`Dog`メソッドを追加して、`beagle`が`eat()`と`bark()`の両方を持てるようにしてください。`bark()`メソッドはコンソールに`Woof!`を出力する必要があります。

# --hints--

`Animal`は`bark()`メソッドに反応してはいけません。

```js
assert(typeof Animal.prototype.bark == 'undefined');
```

`Dog`は`eat()`から`Animal`メソッドを継承する必要があります。

```js
assert(typeof Dog.prototype.eat == 'function');
```

`Dog`のプロトタイプは`bark()`メソッドを持つべきです。

```js
assert('bark' in Dog.prototype);
```

`beagle`は`instanceof` `Animal`であるべきです。

```js
assert(beagle instanceof Animal);
```

`beagle`のコンストラクタは`Dog`に設定されるべきです。

```js
assert(beagle.constructor === Dog);
```

`beagle.eat()`は文字列`nom nom nom`をログに出力するべきです。

```js
const spy = __helpers.spyOn(console, 'log');
beagle.eat();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'nom nom nom');
```

`beagle.bark()`は文字列`Woof!`をログに出力するべきです。

```js
const spy = __helpers.spyOn(console, 'log');
beagle.bark();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'Woof!');
```

# --seed--

## --seed-contents--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }

// Only change code below this line




// Only change code above this line

let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Dog.prototype.bark = function () {
  console.log('Woof!');
};
let beagle = new Dog();

beagle.eat();
beagle.bark();
```
