---
id: 587d7dad367417b2b2512b75
title: オブジェクトにメソッドを作成する
challengeType: 1
forumTopicId: 301318
dashedName: create-a-method-on-an-object
---

# --description--

オブジェクトは、<dfn>メソッド</dfn>と呼ばれる特別な種類のプロパティを持つことができます。

メソッドは関数であるプロパティです。これによりオブジェクトに異なる動作が追加されます。以下はメソッドを持つ`duck`の例です。

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + duck.name + ".";}
};
duck.sayName();
```

この例では、`sayName`メソッドが追加されています。これは`duck`の名前を返す関数です。return文で`name`を使って`duck.name`プロパティにアクセスしていることに注目してください。次のチャレンジではこれとは別の方法を扱います。

# --instructions--

`dog`オブジェクトを使って、`sayLegs`という名前のメソッドを作成してください。そのメソッドは`This dog has 4 legs.`という文を返すようにしてください。

# --hints--

`dog.sayLegs()`は関数である必要があります。

```js
assert(typeof dog.sayLegs === 'function');
```

`dog.sayLegs()`は指定された文字列を返す必要があります。句読点やスペースも正確にしてください。

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,

};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
