---
id: 587d7dac367417b2b2512b74
title: オブジェクトのプロパティにドット記法でアクセスする
challengeType: 1
forumTopicId: 301333
dashedName: use-dot-notation-to-access-the-properties-of-an-object
---

# --description--

前のチャレンジでは、さまざまなプロパティを持つオブジェクトを作成しました。今回は、そのプロパティの値にアクセスする方法を見ていきます。例を示します。

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
console.log(duck.name);
```

ドット記法は、オブジェクト名 `duck` の後にプロパティ名 `name` を続けて使い、`Aflac` の値にアクセスします。

# --instructions--

`dog` オブジェクトの両方のプロパティをコンソールに出力してください。

# --hints--

コードでは `console.log` を使って、`name` オブジェクトの `dog` プロパティの値を出力してください。

```js
assert(/console.log\(.*dog\.name.*\)/g.test(__helpers.removeJSComments(code)));
```

コードでは `console.log` を使って、`numLegs` オブジェクトの `dog` プロパティの値を出力してください。

```js
assert(/console.log\(.*dog\.numLegs.*\)/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
// Only change code below this line
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
console.log(dog.name);
console.log(dog.numLegs);
```
