---
id: 587d7dad367417b2b2512b77
title: コンストラクタ関数を定義する
challengeType: 1
forumTopicId: 16804
dashedName: define-a-constructor-function
---

# --description--

<dfn>コンストラクタ</dfn>は新しいオブジェクトを作成する関数です。新しいオブジェクトに属するプロパティや動作を定義します。新しいオブジェクトを作るための設計図のようなものだと考えてください。

以下はコンストラクタの例です。

```js
function Bird() {
  this.name = "Albert";
  this.color = "blue";
  this.numLegs = 2;
}
```

このコンストラクタは`Bird`オブジェクトを定義し、プロパティ`name`、`color`、`numLegs`をそれぞれAlbert、blue、2に設定しています。コンストラクタにはいくつかの慣例があります。

<ul><li>コンストラクタは他の<code>関数</code>と区別するために名前の最初の文字を大文字にします。</li><li>コンストラクタは作成するオブジェクトのプロパティを設定するために<code>this</code>キーワードを使います。コンストラクタ内の<code>this</code>は作成される新しいオブジェクトを指します。</li><li>コンストラクタは他の関数のように値を返すのではなく、プロパティや動作を定義します。</li></ul>

# --instructions--

プロパティ`Dog`、`name`、`color`をそれぞれ文字列、文字列、数値に設定するコンストラクタ`numLegs`を作成してください。

# --hints--

`Dog`は文字列に設定された`name`プロパティを持つべきです。

```js
assert(typeof new Dog().name === 'string');
```

`Dog`は文字列に設定された`color`プロパティを持つべきです。

```js
assert(typeof new Dog().color === 'string');
```

`Dog`は数値に設定された`numLegs`プロパティを持つべきです。

```js
assert(typeof new Dog().numLegs === 'number');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function Dog (name, color, numLegs) {
  this.name = 'name';
  this.color = 'color';
  this.numLegs = 4;
}
```
