---
id: 587d7db0367417b2b2512b83
title: 継承を使ってコードの繰り返しを避ける
challengeType: 1
forumTopicId: 301334
dashedName: use-inheritance-so-you-dont-repeat-yourself
---

# --description--

プログラミングには「<dfn>繰り返しを避ける（Don't Repeat Yourself、DRY）</dfn>」という原則があります。コードが繰り返される問題は、変更があるたびに複数箇所のコードを修正しなければならず、プログラマーの作業が増え、ミスが起こりやすくなることです。

以下の例では、`describe` メソッドが `Bird` と `Dog` に共通して使われていることに注目してください。

```js
Bird.prototype = {
  constructor: Bird,
  describe: function() {
    console.log("My name is " + this.name);
  }
};

Dog.prototype = {
  constructor: Dog,
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

`describe` メソッドは2箇所で繰り返されています。DRYの原則に従うために、`supertype` という `Animal`（親）を作成してコードを修正できます。

```js
function Animal() { };

Animal.prototype = {
  constructor: Animal, 
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

`Animal` に `describe` メソッドが含まれているので、`Bird` と `Dog` からは削除できます。

```js
Bird.prototype = {
  constructor: Bird
};

Dog.prototype = {
  constructor: Dog
};
```

# --instructions--

`eat` メソッドが `Cat` と `Bear` の両方で繰り返されています。DRYの精神に従い、`eat` メソッドを `Animal` `supertype` に移動してコードを編集してください。

# --hints--

`Animal.prototype` は `eat` プロパティを持つべきです。

```js
assert(Animal.prototype.hasOwnProperty('eat'));
```

`Bear.prototype` は `eat` プロパティを持ってはいけません。

```js
assert(!Bear.prototype.hasOwnProperty('eat'));
```

`Cat.prototype` は `eat` プロパティを持ってはいけません。

```js
assert(!Cat.prototype.hasOwnProperty('eat'));
```

# --seed--

## --seed-contents--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,

};
```

# --solutions--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
```
