---
id: 587d7db0367417b2b2512b81
title: オブジェクトのプロトタイプがどこから来るのかを理解する
challengeType: 1
forumTopicId: 301330
dashedName: understand-where-an-objects-prototype-comes-from
---

# --description--

人が親から遺伝子を受け継ぐのと同じように、オブジェクトはそれを作成したコンストラクタ関数から直接`prototype`を受け継ぎます。例えば、ここでは`Bird`コンストラクタが`duck`オブジェクトを作成しています。

```js
function Bird(name) {
  this.name = name;
}

let duck = new Bird("Donald");
```

`duck`は`prototype`コンストラクタ関数から`Bird`を受け継ぎます。この関係は`isPrototypeOf`メソッドで示せます。

```js
Bird.prototype.isPrototypeOf(duck);
```

これは`true`を返します。

# --instructions--

`isPrototypeOf`を使って`prototype`の`beagle`を確認してください。

# --hints--

`Dog.prototype`が`prototype`の`beagle`であることを示す必要があります。

```js
assert(/Dog\.prototype\.isPrototypeOf\(beagle\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
```
