---
id: 587d7dac367417b2b2512b73
title: 基本的なJavaScriptオブジェクトを作成する
challengeType: 1
forumTopicId: 301317
dashedName: create-a-basic-javascript-object
---

# --description--

人々が毎日目にするもの、例えば車、店、鳥について考えてみましょう。これらはすべて<dfn>オブジェクト</dfn>です。人々が観察し、やり取りできる具体的なものです。

これらのオブジェクトにはどんな特徴がありますか？車には車輪があります。店は商品を販売します。鳥には翼があります。

これらの特徴、つまり<dfn>プロパティ</dfn>はオブジェクトを構成するものを定義します。似たオブジェクトは同じプロパティを共有しますが、そのプロパティの値は異なる場合があります。例えば、すべての車に車輪がありますが、車輪の数はすべての車で同じとは限りません。

JavaScriptのオブジェクトは現実世界のオブジェクトをモデル化するために使われ、現実の対応物と同じようにプロパティや動作を持たせることができます。以下はこれらの概念を使って`duck`オブジェクトを作成する例です。

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
```

この`duck`オブジェクトは2つのプロパティと値のペアを持っています。`name`が`Aflac`で、`numLegs`が2です。

# --instructions--

`dog`オブジェクトを作成し、`name`と`numLegs`のプロパティを持たせ、それぞれ文字列と数値に設定してください。

# --hints--

`dog`はオブジェクトである必要があります。

```js
assert(typeof dog === 'object');
```

`dog`は文字列に設定された`name`プロパティを持つ必要があります。

```js
assert(typeof dog.name === 'string');
```

`dog`は数値に設定された`numLegs`プロパティを持つ必要があります。

```js
assert(typeof dog.numLegs === 'number');
```

# --seed--

## --seed-contents--

```js
let dog = {

};
```

# --solutions--

```js
let dog = {
  name: '',
  numLegs: 4
};
```
