---
id: 587d7dae367417b2b2512b7b
title: "Own Properties を理解する"
challengeType: 1
forumTopicId: 301326
dashedName: understand-own-properties
---

# --description--

次の例では、`Bird` コンストラクタが 2 つのプロパティ、`name` と `numLegs` を定義しています。

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let duck = new Bird("Donald");
let canary = new Bird("Tweety");
```

`name` と `numLegs` は <dfn>own properties</dfn>（自身のプロパティ）と呼ばれます。これはインスタンスオブジェクトに直接定義されているためです。つまり、`duck` と `canary` はそれぞれこれらのプロパティの別々のコピーを持っています。実際、`Bird` のすべてのインスタンスはこれらのプロパティの独自のコピーを持ちます。次のコードは `duck` のすべての自身のプロパティを配列 `ownProps` に追加しています。

```js
let ownProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  }
}

console.log(ownProps);
```

コンソールには値 `["name", "numLegs"]` が表示されます。

# --instructions--

`canary` の自身のプロパティを配列 `ownProps` に追加してください。

# --hints--

`ownProps` には値 `numLegs` と `name` が含まれている必要があります。

```js
assert(ownProps.indexOf('name') !== -1 && ownProps.indexOf('numLegs') !== -1);
```

組み込みメソッド `Object.keys()` を使わずにこのチャレンジを解いてください。

```js
assert(!/Object(\.keys|\[(['"`])keys\2\])/.test(__helpers.removeJSComments(code)));
```

`ownProps` 配列をハードコーディングせずにこのチャレンジを解いてください。

```js
assert(
  !/\[\s*(?:'|")(?:name|numLegs)|(?:push|concat)\(\s*(?:'|")(?:name|numLegs)/.test(
    __helpers.removeJSComments(code)
  )
);
```

# --seed--

## --seed-contents--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
let ownProps = [];
// Only change code below this line
```

# --solutions--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
function getOwnProps (obj) {
  const props = [];

  for (let prop in obj) {
    if (obj.hasOwnProperty(prop)) {
      props.push(prop);
    }
  }

  return props;
}

const ownProps = getOwnProps(canary);
```
