---
id: 587d7db2367417b2b2512b8c
title: IIFEを使ってモジュールを作成する
challengeType: 1
forumTopicId: 301332
dashedName: use-an-iife-to-create-a-module
---

# --description--

即時実行関数式（IIFE）は、関連する機能を1つのオブジェクトや<dfn>モジュール</dfn>にまとめるためによく使われます。例えば、以前のチャレンジでは2つのミックスインを定義しました。

```js
function glideMixin(obj) {
  obj.glide = function() {
    console.log("Gliding on the water");
  };
}
function flyMixin(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  };
}
```

これらのミックスインを次のようにモジュールにまとめることができます。

```js
let motionModule = (function () {
  return {
    glideMixin: function(obj) {
      obj.glide = function() {
        console.log("Gliding on the water");
      };
    },
    flyMixin: function(obj) {
      obj.fly = function() {
        console.log("Flying, wooosh!");
      };
    }
  }
})();
```

即時実行関数式（IIFE）がオブジェクトを返していることに注目してください `motionModule`。この返されたオブジェクトには、すべてのミックスインの動作がオブジェクトのプロパティとして含まれています。モジュールパターンの利点は、すべての動作を1つのオブジェクトにまとめてパッケージ化でき、それをコードの他の部分で使えることです。以下はその使用例です。

```js
motionModule.glideMixin(duck);
duck.glide();
```

# --instructions--

`funModule`という名前のモジュールを作成し、2つのミックスイン `isCuteMixin` と `singMixin` をラップしてください。`funModule` はオブジェクトを返す必要があります。

# --hints--

`funModule` は定義され、オブジェクトを返す必要があります。

```js
assert(typeof funModule === 'object');
```

`funModule.isCuteMixin` は関数にアクセスする必要があります。

```js
assert(typeof funModule.isCuteMixin === 'function');
```

`funModule.singMixin` は関数にアクセスする必要があります。

```js
assert(typeof funModule.singMixin === 'function');
```

# --seed--

## --seed-contents--

```js
let isCuteMixin = function(obj) {
  obj.isCute = function() {
    return true;
  };
};
let singMixin = function(obj) {
  obj.sing = function() {
    console.log("Singing to an awesome tune");
  };
};
```

# --solutions--

```js
const funModule = (function () {
  return {
    isCuteMixin: obj => {
      obj.isCute = () => true;
    },
    singMixin: obj => {
      obj.sing = () => console.log("Singing to an awesome tune");
    }
  };
})();
```
