---
id: 587d7dad367417b2b2512b78
title: "コンストラクタを使ってオブジェクトを作成する"
challengeType: 1
forumTopicId: 18233
dashedName: use-a-constructor-to-create-objects
---

# --description--

前のチャレンジで使った`Bird`コンストラクタはこちらです：

```js
function Bird() {
  this.name = "Albert";
  this.color  = "blue";
  this.numLegs = 2;
}

let blueBird = new Bird();
```

**注意：** コンストラクタ内の`this`は常に作成されるオブジェクトを指します。

`new`演算子がコンストラクタを呼び出すときに使われていることに注目してください。これはJavaScriptに`Bird`の新しいインスタンスを作成し、`blueBird`と呼ぶよう指示しています。`new`演算子がなければ、コンストラクタ内の`this`は新しく作成されたオブジェクトを指さず、予期しない結果になります。これで`blueBird`は`Bird`コンストラクタ内で定義されたすべてのプロパティを持つことになります：

```js
blueBird.name;
blueBird.color;
blueBird.numLegs;
```

他のオブジェクトと同様に、そのプロパティはアクセスおよび変更できます：

```js
blueBird.name = 'Elvira';
blueBird.name;
```

# --instructions--

前のレッスンの`Dog`コンストラクタを使って`Dog`の新しいインスタンスを作成し、変数`hound`に代入してください。

# --hints--

`hound`は`Dog`コンストラクタを使って作成する必要があります。

```js
assert(hound instanceof Dog);
```

コードでは`new`演算子を使って`Dog`のインスタンスを作成してください。

```js
assert(__helpers.removeJSComments(code).match(/new/g));
```

# --seed--

## --seed-contents--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
// Only change code below this line
```

# --solutions--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
const hound = new Dog();
```
