---
id: 587d7db1367417b2b2512b86
title: 継承されたコンストラクタのプロパティをリセットする
challengeType: 1
forumTopicId: 301324
dashedName: reset-an-inherited-constructor-property
---

# --description--

オブジェクトが別のオブジェクトから`prototype`を継承するとき、スーパタイプのコンストラクタのプロパティも継承されます。

例を示します：

```js
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
let duck = new Bird();
duck.constructor
```

しかし`duck`と`Bird`のすべてのインスタンスは、`Bird`ではなく`Animal`によって構築されたことを示すべきです。そのためには、`Bird`のコンストラクタのプロパティを手動で`Bird`オブジェクトに設定できます：

```js
Bird.prototype.constructor = Bird;
duck.constructor
```

# --instructions--

`duck.constructor`と`beagle.constructor`がそれぞれのコンストラクタを返すようにコードを修正してください。

# --hints--

`Bird.prototype`は`Animal`のインスタンスであるべきです。

```js
assert(Animal.prototype.isPrototypeOf(Bird.prototype));
```

`duck.constructor`は`Bird`を返すべきです。

```js
assert(duck.constructor === Bird);
```

`Dog.prototype`は`Animal`のインスタンスであるべきです。

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

`beagle.constructor`は`Dog`を返すべきです。

```js
assert(beagle.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Animal() { }
function Bird() { }
function Dog() { }

Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);

// Only change code below this line



let duck = new Bird();
let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
function Bird() { }
function Dog() { }
Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Bird.prototype.constructor = Bird;
let duck = new Bird();
let beagle = new Dog();
```
