---
id: 587d7daf367417b2b2512b7f
title: プロトタイプを新しいオブジェクトに変更する
challengeType: 1
forumTopicId: 301316
dashedName: change-the-prototype-to-a-new-object
---

# --description--

これまでは `prototype` にプロパティを個別に追加してきました。

```js
Bird.prototype.numLegs = 2;
```

プロパティが数個を超えると、これは面倒になります。

```js
Bird.prototype.eat = function() {
  console.log("nom nom nom");
}

Bird.prototype.describe = function() {
  console.log("My name is " + this.name);
}
```

より効率的な方法は、`prototype` をすでにプロパティを含む新しいオブジェクトに設定することです。こうすれば、プロパティが一度にまとめて追加されます。

```js
Bird.prototype = {
  numLegs: 2, 
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --instructions--

`numLegs` を新しいオブジェクトに設定して、`eat()` の `describe()` にプロパティ `prototype` とメソッド `Dog`、`prototype` の2つを追加してください。

# --hints--

`Dog.prototype` は新しいオブジェクトに設定してください。

```js
assert(/Dog\.prototype\s*?=\s*?{/.test(__helpers.removeJSComments(code)));
```

`Dog.prototype` はプロパティ `numLegs` を持つべきです。

```js
assert(Dog.prototype.numLegs !== undefined);
```

`Dog.prototype` はメソッド `eat()` を持つべきです。

```js
assert(typeof Dog.prototype.eat === 'function');
```

`Dog.prototype` はメソッド `describe()` を持つべきです。

```js
assert(typeof Dog.prototype.describe === 'function');
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype = {
  // Only change code below this line

};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
numLegs: 4,
  eat () {
    console.log('nom nom nom');
  },
  describe () {
    console.log('My name is ' + this.name);
  }
};
```
