---
id: 587d7daf367417b2b2512b7d
title: すべてのプロパティを繰り返す
challengeType: 1
forumTopicId: 301320
dashedName: iterate-over-all-properties
---

# --description--

これまでに2種類のプロパティを見てきました：<dfn>own properties</dfn> と `prototype` プロパティです。Own properties はオブジェクトインスタンス自体に直接定義されています。そして prototype プロパティは `prototype` に定義されています。

```js
function Bird(name) {
  this.name = name;  //own property
}

Bird.prototype.numLegs = 2; // prototype property

let duck = new Bird("Donald");
```

ここでは `duck` の own properties を配列 `ownProps` に、`prototype` のプロパティを配列 `prototypeProps` に追加する方法を示します：

```js
let ownProps = [];
let prototypeProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  } else {
    prototypeProps.push(property);
  }
}

console.log(ownProps);
console.log(prototypeProps);
```

`console.log(ownProps)` はコンソールに `["name"]` を表示し、`console.log(prototypeProps)` は `["numLegs"]` を表示します。

# --instructions--

`beagle` のすべての own properties を配列 `ownProps` に追加してください。`prototype` のすべての `Dog` プロパティを配列 `prototypeProps` に追加してください。

# --hints--

`ownProps` 配列には `name` のみを含めるべきです。

```js
assert.deepEqual(ownProps, ['name']);
```

`prototypeProps` 配列には `numLegs` のみを含めるべきです。

```js
assert.deepEqual(prototypeProps, ['numLegs']);
```

組み込みメソッド `Object.keys()` を使わずにこのチャレンジを解いてください。

```js
assert(!/\Object.keys/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];
for (let prop in beagle) {
  if (beagle.hasOwnProperty(prop)) {
    ownProps.push(prop);
  } else {
    prototypeProps.push(prop);
  }
}
```
