---
id: 587d7dae367417b2b2512b7c
title: 重複したコードを減らすためにプロトタイプのプロパティを使う
challengeType: 1
forumTopicId: 301336
dashedName: use-prototype-properties-to-reduce-duplicate-code
---

# --description--

`numLegs` はおそらくすべての `Bird` のインスタンスで同じ値になるため、実質的に各 `numLegs` インスタンスの中に重複した変数 `Bird` を持っていることになります。

インスタンスが2つだけなら問題ないかもしれませんが、もし何百万ものインスタンスがあったらどうでしょうか。大量の重複した変数が存在することになります。

より良い方法は `prototype` の `Bird` を使うことです。`prototype` のプロパティは `Bird` のすべてのインスタンスで共有されます。`numLegs` に `Bird prototype` を追加する方法は次の通りです。

```js
Bird.prototype.numLegs = 2;
```

これで `Bird` のすべてのインスタンスに `numLegs` プロパティが備わりました。

```js
console.log(duck.numLegs);
console.log(canary.numLegs);
```

すべてのインスタンスが自動的に `prototype` 上のプロパティを持つため、`prototype` はオブジェクトを作成するための「レシピ」と考えてください。`prototype` と `duck` の `canary` は `Bird` コンストラクタの一部として `Bird.prototype` に含まれていることに注意してください。

# --instructions--

`numLegs` の `prototype` の `Dog` プロパティを追加してください。

# --hints--

`beagle` は `numLegs` プロパティを持つべきです。

```js
assert(beagle.numLegs !== undefined);
```

`beagle.numLegs` は数値であるべきです。

```js
assert(typeof beagle.numLegs === 'number');
```

`numLegs` は自身のプロパティではなく `prototype` のプロパティであるべきです。

```js
assert(beagle.hasOwnProperty('numLegs') === false);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}



// Only change code above this line
let beagle = new Dog("Snoopy");
```

# --solutions--

```js
function Dog (name) {
  this.name = name;
}
Dog.prototype.numLegs = 4;
let beagle = new Dog("Snoopy");
```
