---
id: 65ef1be2b958cabf712296cd
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

現在你將透過尋找目前區間的中點並將中點的平方與標的值比較，反覆縮小區間。

為此，在 `else` 區塊內，創建一個執行最多 `max_iterations` 次的 `for` 迴圈。

在你的迴圈中，使用 `range` 函式（程式），它會產生一個你可以迭代的數字序列。語法是 `range(start, stop, step)`，其中 `start` 是起始整數（包含），`stop` 是最後一個整數（不包含），而 `step` 是序列中一個數字與前一個數字之間的差異。

此外，使用 `_` 作為迴圈變數。`_` 作為佔位符，當你需要使用變數但實際上不需要其值時非常有用。

# --hints--

你應該創建一個 `for` 迴圈來迭代 `range(max_iterations)`。使用 `_` 作為迴圈變數。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
