---
id: 65ef198fde24dfb7ff675b42
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

`raise` 述語允許你強制引發特定的異常。它由 `raise` 關鍵字後接異常型別組成，並且啟用你提供訂製的錯誤訊息：

```py
raise ValueError("Invalid value")
```

當上述程式碼執行時，會引發 `ValueError`，並向使用者顯示訊息 `"Invalid value"`。

如果 `square_target` 小於 `0`，則無法計算實數平方根。因此，引發帶有訊息 `'Square root of negative number is not defined in real numbers'` 的 `ValueError`。別忘了移除 `pass` 關鍵字。

# --hints--

你應該移除 `pass` 關鍵字。

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

你應該在 `if` 區塊內引發一個帶有訊息 `'Square root of negative number is not defined in real numbers'` 的 `ValueError`。

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
