---
id: 65ef181dad3d96b14a21a03e
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個項目中，你將使用二分法找到給定數字的近似平方根。

二分法是一種尋找實值函數根的技術。 它的工作原理是縮小平方根所在的區間，直到它收斂到指定公差範圍內的值。

首先創建一個名爲 `square_root_bisection` 的函數。 暫時將參數留空。

# --hints--

你應該定義不帶參數的函數 `square_root_bisection`。

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
