---
id: 65ef1be2b958cabf712296cd
title: 步驟11
challengeType: 20
dashedName: step-11
---

# --description--

現在，你將通過找到當前間隔的中點並將中點的平方與目標值進行比較來反覆縮小間隔。

爲此，在 `else` 塊內，創建一個 `for` 循環，最多運行 `max_iterations` 次。

對於循環，使用 `range` 函數，它會生成一個你可以迭代的數字序列。 語法爲 `range(start, stop, step)`，其中 `start` 是起始整數（含），`stop` 是最後一個整數（不含），`step` 是序列中某個數字與前一個數字之間的差值。

另外，使用 `_` 作爲循環變量。 `_` 充當佔位符，當你需要使用變量但實際上不需要其值時很有用。

# --hints--

你應該創建一個 `for` 循環來迭代 `range(max_iterations)`。 使用 `_` 作爲循環變量。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None

--fcc-editable-region--

--fcc-editable-region--
```
