---
id: 65ef19ec318e2fb9066381e3
title: 步驟5
challengeType: 20
dashedName: step-5
---

# --description--

你將爲 `square_target` 是 `0` 或 `1` 時創建單獨的案例。

首先創建一個 `if` 語句來檢查 `square_target` 是否等於 `1`。

# --hints--

你應該創建一個 `if` 語句來檢查 `square_target == 1`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[0].is_equivalent("square_target == 1")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')

--fcc-editable-region--
```
