---
id: 65ef190c6b51e9b5a5f7ed29
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

為 `square_root_bisection` 函式（程式）提供以下參數：

- `square_target`：你想要尋找平方根的數字。
- `tolerance`（選用）：近似根值的平方與實際標的值之間可接受的差異（預設值為 `1e-7`）。容差 `1e-7` 表示解法將在真實值的 `0.0000001` 以內達到準確度，是一個在準確度與效率間取得良好平衡的預設值。
- `max_iterations`（選用）：要執行的最大迭代次數（預設值為 `100`）。如果該方法在此限制內未收斂，你將假設未找到解決方案。
  
# --hints--

你的函式（程式）應該有這些參數：`square_target`、`tolerance = 1e-7` 和 `max_iterations = 100`。順序很重要。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```

