---
id: 65ef19425d1b27b6c930bae6
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

如果你想要尋找平方根的數字是負數，程式碼應該引發錯誤，因為負數的平方根在實數中未定義。

移除 `pass` 述語並創建一個 `if` 述語以檢查 `square_target` 是否小於 `0`。

# --hints--

你應該移除 `pass` 關鍵字。

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

你應該創建一個 `if` 述語來檢查 `square_target < 0`。

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
