---
id: 65ef1a720d6e65bbb0d22463
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--
接下來，你將處理 `square_target` 是除 `1` 或 `0` 以外的正數的情況。

創建一個 `else` 語句來控制代碼這些情況。 

# --hints--

創建一個 `else` 語句，並且不要忘記為你的 `pass` 關鍵字添加。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[2] == _Node()`))
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')
--fcc-editable-region--


--fcc-editable-region--
```
