---
id: 65ef181dad3d96b14a21a03e
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個專案中，你將使用二分法來找到給定數字的近似平方根。

二分法是一種用於尋找實值函式根的技術。它透過縮小平方根所在的區間，直到收斂到指定容差範圍內的值。

首先建立一個名為 `square_root_bisection` 的函式（程式）。目前先將參數留空。

# --hints--

你應該定義沒有參數的函式（程式）`square_root_bisection`。

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
