---
id: 65ef1d5e3d2927c5e0f4997b
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

在 Python 中，`is` 關鍵字用來檢查物件身分。它用於判斷兩個變數是否指向記憶體中相同的物件。與 `is` 不同，等號運算子（`==`）用來判斷兩個物件的值是否相同，無論它們是否為記憶體中相同的物件。

在 for 迴圈外，創建一個 `if` 述語來檢查 `root is None`。如果是，則列印訊息 `'Failed to converge within {max_iterations} iterations.'`。記得使用 f-string 格式化該訊息。

# --hints--

你應該創建一個 `if` 述語來檢查 `root is None`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

你應該將引數 `f'Failed to converge within {max_iterations} iterations.'` 傳遞給 `print` 函式（程式）。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
