---
id: 65ef1c34a541afc12dbb849d
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

`abs()` 函式（程式）傳回數字的絕對值，該值永遠為正數，無論數字的符號為何。你會使用它來檢查估計的平方根是否足夠接近實際值。

現在，創建一個 `if` 述語來檢查 `square_mid` 和 `square_target` 之間差值的絕對值是否在指定的 `tolerance` 範圍內。

# --hints--

你應該在 `for` 迴圈的主體內有一個條件為 `abs(square_mid - square_target) < tolerance:` 的 `if` 述語。

```js
({
    test: () =>
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[0].is_equivalent("abs(square_mid - square_target) < tolerance")`))
    } 
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2
            
--fcc-editable-region--
```
