---
id: 65ef1a50049cf9bada13266f
title: 步驟7
challengeType: 20
dashedName: step-7
---

# --description--

創建一個 `elif` 語句來檢查 `square_target` 是否等於 `0`。 如果是，則將值 `0` 分配給 `root` 變量。 另外，打印消息 `'The square root of {square_target} is 0'`。 記得使用 f 字符串來格式化消息。

# --hints--

你應該有一個 `elif` 語句來檢查條件 `square_target == 0`。


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

你應該將值 `0` 分配給 `root` 變量，並將參數 `f'The square root of {square_target} is 0'` 傳遞給 `print` 調用。


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')

--fcc-editable-region--
```
