---
id: 61537bb9b1a29430ac15ad38
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Within your `head` element, add a `title` element with the text set to `Photo Gallery`, and a `link` element to add your `styles.css` file to the page.

# --hints--

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tu proyecto debe tener el título `Photo Gallery`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

Recuerda que las mayúsculas, minúsculas y la ortografía son importantes para el título.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
