---
id: 65ef1a50049cf9bada13266f
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Crie uma declaração `elif` para verificar se `square_target` é igual a `0`. Se for, atribua o valor `0` à variável `root`. Além disso, imprima a mensagem `'The square root of {square_target} is 0'`.  Remember to format the message using an f-string.

# --hints--

Você deve ter uma declaração `elif` para verificar a condição `square_target == 0`.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

Você deve atribuir o valor `0` à variável `root` e passar o argumento `f'The square root of {square_target} is 0'` para uma chamada de `print`.


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
