---
id: 65ef19425d1b27b6c930bae6
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Se o número para o qual você quer encontrar a raiz quadrada for negativo, o código deve gerar um erro, pois a raiz quadrada de um número negativo não está definida nos números reais.

Remova a declaração `pass` e crie uma declaração `if` para verificar se `square_target` é menor que `0`.

# --hints--

Você deve remover a palavra-chave `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

Você deve criar uma declaração `if` para verificar se `square_target < 0`.

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
