---
id: 65ef1be2b958cabf712296cd
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Agora você repetidamente reduzirá o intervalo encontrando o ponto médio do intervalo atual e comparando o quadrado do ponto médio com o valor alvo.

Para isso, dentro do bloco `else`, crie um loop `for` que execute até `max_iterations` vezes.

Para seu loop, use a função `range`, que gera uma sequência de números sobre a qual você pode iterar. A sintaxe é `range(start, stop, step)`, onde `start` é o inteiro inicial (inclusivo), `stop` é o último inteiro (não inclusivo) e `step` é a diferença entre um número e o anterior na sequência.

Além disso, use `_` como uma variável de loop. O `_` atua como um espaço reservado e é útil quando você precisa usar uma variável mas não precisa realmente do seu valor.

# --hints--

Você deve criar um loop `for` para iterar sobre `range(max_iterations)`. Use `_` como a variável do loop.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
