---
id: 65ef181dad3d96b14a21a03e
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste projeto, você encontrará a raiz quadrada aproximada de um número dado usando o método da bissecção.

O método da bissecção é uma técnica para encontrar as raízes de uma função de valor real.Funciona restringindo um intervalo onde a raiz quadrada se encontra até que converja para um valor dentro de uma tolerância especificada.

Comece criando uma função chamada `square_root_bisection`.Deixe os parâmetros vazios por enquanto.

# --hints--

Você deve definir a função `square_root_bisection` sem parâmetros.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
