---
id: 65ef190c6b51e9b5a5f7ed29
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Forneça à função `square_root_bisection` os seguintes parâmetros:

- `square_target`: O número para o qual você deseja encontrar a raiz quadrada.
- `tolerance` (opcional): A diferença aceitável entre o quadrado do valor aproximado da raiz e o valor real alvo (o padrão é `1e-7`). A tolerância `1e-7` implica que a solução será precisa dentro de `0.0000001` do valor verdadeiro e é uma boa escolha padrão que equilibra precisão e desempenho.
- `max_iterations` (opcional): O número máximo de iterações a serem realizadas (padrão é `100`). Se o método não convergir dentro desse limite, você assumirá que a solução não foi encontrada.
  
# --hints--

Sua função deve ter esses parâmetros: `square_target`, `tolerance = 1e-7` e `max_iterations = 100`.A ordem importa.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```

