---
id: 65ef1cda150a59c3b8306944
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

Se a diferença não estiver dentro da `tolerance` especificada, crie uma declaração `elif` para verificar se `square_mid` é menor que `square_target`.

Atribua o valor de `mid` para `low` pois a raiz quadrada agora estará entre `low` e `mid`.

# --hints--

Você deve criar uma declaração `elif` que verifique se `square_mid < square_target`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[1].is_equivalent("square_mid < square_target")`))     
                  
    }

})
```

Você deve atribuir o valor de `mid` para `low`.

```js
({
    test: () => 
    {
          assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid")`))
    }

})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

--fcc-editable-region--
            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

--fcc-editable-region--
```
