---
id: 65ef1d5e3d2927c5e0f4997b
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Em Python, a palavra-chave `is` verifica a identidade do objeto. É usado para determinar se duas variáveis apontam para o mesmo objeto na memória. Em contraste com `is`, o operador de igualdade (`==`) determina se os valores de dois objetos são os mesmos, independentemente de serem o mesmo objeto na memória.

Fora do loop for, crie uma declaração `if` para verificar se `root is None`. Se for, imprima a mensagem `'Failed to converge within {max_iterations} iterations.'`.  Remember to format the message using an f-string.

# --hints--

Você deve criar uma declaração `if` que verifica se `root is None`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

Você deve passar o argumento `f'Failed to converge within {max_iterations} iterations.'` para a função `print`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
