---
id: 65ef1e10268f8ec9121446d1
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Fora da definição da função, crie uma variável `N` e atribua o valor `16` a ela.

# --hints--

Você deve ter a variável `N = 16` fora da definição da função.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("N").is_equivalent("N = 16")`))
    }
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

        if root is None:
            print(f"Failed to converge within {max_iterations} iterations.")
        
        else:   
            print(f'The square root of {square_target} is approximately {root}')
    
    return root

--fcc-editable-region--

--fcc-editable-region--

```
