---
id: 65ef1a720d6e65bbb0d22463
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--
Em seguida, você vai trabalhar nos casos em que o `square_target` é um número positivo diferente de `1` ou `0`.

Crie uma cláusula `else` para tratar esses casos. 

# --hints--

Crie uma cláusula `else` e não se esqueça de adicionar a palavra-chave `pass`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[2] == _Node()`))
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')
--fcc-editable-region--


--fcc-editable-region--
```
