---
id: 587d7dad367417b2b2512b78
title: Utiliza un constructor para crear objetos
challengeType: 1
forumTopicId: 18233
dashedName: use-a-constructor-to-create-objects
---

# --description--

Aquí tenemos el constructor `Bird` del desafío anterior:

```js
function Bird() {
  this.name = "Albert";
  this.color  = "blue";
  this.numLegs = 2;
}

let blueBird = new Bird();
```

**NOTA:** `this` dentro del constructor siempre se refiere al objeto que se está creando.

Observa que el operador `new` se utiliza cuando llamamos a un constructor. Esto le dice a JavaScript que cree una nueva instancia de `Bird` llamada `blueBird`. Sin el operador `new`, `this` dentro del constructor no haría referencia al nuevo objeto por lo que daría resultados inesperados. Ahora `blueBird` tiene todas las propiedades que han sido definidas dentro del constructor `Bird`:

```js
blueBird.name;
blueBird.color;
blueBird.numLegs;
```

Como con cualquier otro objecto, se pueden acceder a sus propiedades y modificarlas:

```js
blueBird.name = 'Elvira';
blueBird.name;
```

# --instructions--

Usa el constructor `Dog` de la anterior lección para crear una nueva instancia de `Dog` y asignarlo a la variable `hound`.

# --hints--

`hound` se debe crear usando el constructor `Dog`.

```js
assert(hound instanceof Dog);
```

Tu código debe usar el operador `new` para crear una instancia de `Dog`.

```js
assert(__helpers.removeJSComments(code).match(/new/g));
```

# --seed--

## --seed-contents--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
// Only change code below this line
```

# --solutions--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
const hound = new Dog();
```
