---
id: 6853e58d52ec64349a5760ca
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste workshop, você irá implementar o algoritmo do caminho mais curto. Você escreverá uma função Python que calcula o caminho mais curto entre os nós em um grafo e também retorna o caminho percorrido.

Por exemplo, dado um grafo onde cidades estão conectadas por estradas com diferentes distâncias, o algoritmo encontrará a rota mais curta de uma cidade para outra. Se você quiser viajar da Cidade A para a Cidade D, o algoritmo pode descobrir que ir de A ⇨ B ⇨ C ⇨ D (total: 15km) é mais curto do que ir diretamente de A ⇨ D (20km).

Para começar, defina uma variável chamada `INF` e atribua a ela o valor `float('inf')`, que representa infinito positivo. Mais tarde, você usará isso para indicar uma distância infinita entre dois nós.

# --hints--

Você deve definir a variável `INF` no topo do seu arquivo.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

Você deve atribuir `float('inf')` à variável `INF`.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
