---
id: 686253f1b051998ad4904e3e
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Você agora criará a função principal que aceita três parâmetros: a matriz de adjacência, o nó inicial e um nó alvo opcional.

Crie uma função chamada `shortest_path` que recebe três parâmetros: `matrix`, `start_node` e `target_node`. Atribua `None` como valor padrão para `target_node`.

O parâmetro `target_node` é definido como `None` por padrão, indicando que se nenhum nó alvo for especificado, a função deve calcular os caminhos mais curtos do nó inicial para todos os outros nós no grafo.

Adicione uma instrução `pass` dentro do corpo da função por enquanto.

# --hints--

Você deve criar uma função chamada `shortest_path`.

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

A função deve receber três parâmetros: `matrix`, `start_node` e `target_node` com `target_node` tendo um valor padrão de `None`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
