---
id: 686251fb49ebc489ba60ccb5
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Você precisará criar uma lista 2D para representar a matriz de adjacência do grafo. Essa matriz será usada para representar os pesos das arestas entre os nós no grafo.

Crie uma variável chamada `adj_matrix` e atribua a ela uma lista 2D que representa o grafo com os seguintes pesos:

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Você deve criar uma variável chamada `adj_matrix`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` deve ser uma lista 2D contendo os valores fornecidos.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
