---
id: 68625ad34aa3f58d7f3b9b07
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Você precisa armazenar o número de nós no grafo. Para isso você precisará de uma variável que corresponda ao comprimento da matriz de adjacência.

Dentro da função `shortest_path`, crie uma variável `n` e defina-a como o comprimento da `matrix`.

# --hints--

Você deve inicializar `n` com o comprimento da `matrix`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("n").is_equivalent("n = len(matrix)")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    pass

--fcc-editable-region--
```
