---
id: 686b73007fc582db0132c361
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Você precisa acompanhar a menor distância conhecida do nó inicial para cada outro nó no grafo.

Para fazer isso, crie uma variável chamada `distances` e inicialize-a como uma lista contendo um único elemento: `INF`.  


# --hints--

Você deve criar uma variável chamada `distances`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("distances")`
        )
      );
    }
});
```

A lista `distances` deve ser inicializada como `[INF]`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF]")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)

--fcc-editable-region--
```
