---
id: 686b78bc13d518e1267448e4
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Além de rastrear distâncias, você também precisa acompanhar os caminhos reais percorridos para alcançar cada nó.

Você criará uma lista onde cada entrada armazena o caminho percorrido para alcançar esse nó. Inicialmente, o caminho de cada nó conterá apenas ele mesmo.  

List comprehensions fornecem uma forma concisa de criar listas. Por exemplo:

```py
[x * 2 for x in range(3)]
```

Crie uma variável chamada `paths` e inicialize-a usando uma list comprehension que cria uma lista contendo `[node_no]` para cada `node_no` em `range(n)`.

# --hints--

Você deve ter uma variável chamada `paths`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

Você deve inicializar `paths` usando uma list comprehension.

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

Sua compreensão de lista deve usar `node_no` para iterar sobre `range(n)`.

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

Sua compreensão de lista deve avaliar `[node_no]` para cada `node_no` em `range(n)`.  

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0

--fcc-editable-region--
```
