---
id: 68760671d950767b055bf2f9
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Agora você precisa verificar cada nó para encontrar aquele com a menor distância conhecida que ainda não foi visitado.

Para fazer isso, adicione um loop `for` dentro do loop principal.

O loop deve iterar por cada `node_no` em `range(n)`, onde `n` é o número de nós no grafo. Adicione `pass` como o corpo do loop por enquanto.

# --hints--

Você deve ter um segundo loop `for` dentro da função `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
len(outer_loop.find_bodies()[0].find_for_loops()) >= 1
      `)
      );
    }
});
```

Seu loop `for` deve iterar sobre `range(n)` com `node_no` como a variável de iteração.

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
inner_loop = outer_loop.find_bodies()[0].find_for_loops()[0]
assert inner_loop.find_for_vars().is_equivalent("node_no")
assert inner_loop.find_for_iter().is_equivalent("range(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1

--fcc-editable-region--
```
