---
id: 68760464eb9c2e79cc913fac
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Nesta etapa, você adicionará um loop que será executado uma vez para cada nó no grafo. Este loop permitirá que o algoritmo atualize distâncias e caminhos ao longo de várias passagens.

Crie um loop `for` que execute `n` vezes. Use `_` como a variável do loop já que você não precisa usar o valor da iteração.  

Dentro do loop, você precisa se preparar para selecionar o próximo nó a ser processado criando duas variáveis: 

- um para armazenar a menor distância encontrada até agora na iteração atual
- e outro para armazenar o índice do nó que possui essa menor distância.

Crie as variáveis `min_distance` e `current` e defina-as como `INF` e `-1`, respectivamente. 

# --hints--

Você deve ter um loop `for` dentro da função `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

Seu loop `for` deve iterar sobre `range(n)`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

Seu loop for deve usar `_` como a variável de iteração.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

Você deve criar uma variável chamada `min_distance` e definir ela como `INF` dentro do laço for.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

Você deve criar uma variável chamada `current` e definir como `-1` dentro do loop for.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
