---
id: 6876089f7a500b7c782e5b83
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Você precisa decidir se o nó atual é uma escolha melhor do que aquele que você já encontrou (se houver). Para fazer isso, você adicionará uma instrução condicional dentro do loop.

A condição deve fazer duas coisas:

- Verifique se o nó ainda não foi visitado.
- Compare a distância conhecida com a atual `min_distance`.

Dentro do `for` interno, adicione uma instrução `if` que verifica se `node_no` não foi visitado e se `distances[node_no]` é menor que `min_distance`.

Adicione `pass` como um espaço reservado dentro do bloco condicional.

# --hints--

Você deve adicionar uma declaração `if` dentro do seu `for` interno.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

Sua declaração `if` deve verificar se o nó não foi visitado (`visited[node_no]` é falso) e `distances[node_no]` é menor que `min_distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            pass

--fcc-editable-region--
```
