---
id: 68760c4c2648e7832c5eef16
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Após o loop que encontra o nó não visitado mais próximo, você precisa verificar se um nó válido foi realmente encontrado.

Se nenhum nó desse tipo existir, isso significa que os nós restantes são inalcançáveis a partir do nó inicial e o algoritmo deve parar antecipadamente.

No mesmo nível do loop `for` aninhado, adicione uma instrução `if` que verifica se `current == -1` e interrompe o loop se for verdadeiro.

# --hints--

Você deve ter uma declaração `if` que verifica se `current` ainda é `-1` após o `for` interno.

```js
({
    test: () => {
      runPython(`
      cond = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
      assert cond.is_equivalent("current == -1") or cond.is_equivalent("-1 == current")
      `);
    }
});
```

Você deve usar uma instrução `break` para sair do loop.

```js
({
    test: () => {
      runPython(`
      assert _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("break")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

--fcc-editable-region--

--fcc-editable-region--
```
