---
id: 68777816e95d4936f129819f
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Antes de tentar atualizar a distância para um vizinho, você precisa verificar se o vizinho é alcançável e não visitado. Então você calculará qual seria a distância total para alcançar esse vizinho através do nó atual.

Dentro do loop `for`, adicione uma declaração `if` que verifica:

- A `distance` não é igual a `INF` (significando que há uma aresta entre os nós)
- O vizinho `node_no` ainda não foi visitado

Dentro do bloco `if`, crie uma variável chamada `new_distance` e atribua a soma de `distances[current]` (a menor distância para o nó atual) e `distance` (a distância do nó atual para o vizinho).

# --hints--

Você deve ter uma declaração `if` que verifica se `distance` não é igual a `INF` e se o nó vizinho não foi visitado.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

Você deve criar uma variável chamada `new_distance` dentro do bloco `if`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

Você deve atribuir a `new_distance` a soma de `distances[current]` e `distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
