---
id: 687b3ef0c2fe185b2abc4654
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Uma vez que o algoritmo tenha terminado de executar, você precisa decidir para qual(is) nó(s) exibir os resultados.

Se um `target_node` específico foi fornecido, você mostrará apenas a distância e o caminho para esse nó. Caso contrário, você mostrará os resultados para todos os nós.

Após o término do loop externo `for _ in range(n):`, crie uma variável chamada `targets`. Use uma expressão condicional para atribuir a ela `[target_node]` se `target_node` não for `None`, caso contrário atribua `range(n)`.

# --hints--

Você deve criar uma variável chamada `targets`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("targets")`
        )
      );
    }
});
```

Sua variável `targets` deve conter uma expressão condicional.

```js
({
    test: () => {
      runPython(`
        import ast
        val =_Node(_code).find_function("shortest_path").find_variable("targets").tree.value
        assert isinstance(val, ast.IfExp)
      `);
    }
});
```

A expressão condicional deve verificar se `target_node is not None`.

```js
({
    test: () => {
        runPython(`
        compare = _Node(_Node(_code).find_function("shortest_path").find_variable("targets").tree.value.test)
        assert compare.is_equivalent("target_node is not None")
        `);
    }
});
```

Você deve atribuir `[target_node]` quando um alvo for fornecido e `range(n)` caso contrário.

```js
({
  test: () => {
      runPython(`
      import ast
      val = _Node(_code).find_function("shortest_path").find_variable("targets").tree.value
      assert _Node(val.body).is_equivalent("[target_node]")
      assert _Node(val.orelse).is_equivalent("range(n)")
    `);
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no  

        if current == -1:  
            break  
        visited[current] = True  

        for node_no in range(n):  
            distance = matrix[current][node_no]  
            if distance != INF and not visited[node_no]:  
                new_distance = distances[current] + distance  
                if new_distance < distances[node_no]:  
                    distances[node_no] = new_distance  
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--

--fcc-editable-region--
```
