---
id: 68e4c0ff9edcc92771d7549a
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

No momento, a lista `distances` tem apenas um elemento. No entanto, você precisa de um valor de distância para cada nó no grafo.

Em Python, você pode multiplicar uma lista por um inteiro para repeti-la. Por exemplo, `[0] * 3` cria `[0, 0, 0]`.

Atualize a lista `distances` para conter `n` cópias de `INF` multiplicando `[INF]` por `n`.

# --hints--

Você deve atualizar `distances` para conter `n` cópias de `INF`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
