---
id: 69046179fba35c03dec84fb6
title: Passo 24
challengeType: 20
dashedName: step-24
---

# --description--

Agora você usará o método `join()` para combinar as representações em string dos números dos nós em um único caminho legível.

O método `join()` recebe um iterável (como sua expressão geradora) e combina todos os elementos em uma única string, colocando o separador entre cada elemento. Por exemplo:

```py
numbers = ['1', '2', '3']
route = ' -> '.join(numbers) # route will be '1 -> 2 -> 3'
```

Crie uma variável chamada `path` e atribua a ela o resultado da junção de `string_path` usando `' -> '` como separador.

# --hints--

Você deve criar uma variável chamada `path`.

```js
({
test: () => {
  assert(
    runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
for_loop.find_bodies()[0].has_variable("path")
      `)
);
    }
});
```

Você usa `join` para unir os itens em `string_path` com `' -> '` como separador e atribui o resultado a `path`.

```js
({
  test: () => {
    runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
assert for_loop.find_bodies()[0].find_variable("path").is_equivalent('path = " -> ".join(string_path)')
    `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no

        if current == -1:
            break
        visited[current] = True

        for node_no in range(n):
            distance = matrix[current][node_no]
            if distance != INF and not visited[node_no]:
                new_distance = distances[current] + distance
                if new_distance < distances[node_no]:
                    distances[node_no] = new_distance
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
    targets = [target_node] if target_node is not None else range(n)
    for node_no in targets:
        if node_no == start_node or distances[node_no] == INF:
            continue
        string_path = (str(n) for n in paths[node_no])
        
--fcc-editable-region--
```
