---
id: 698b42492a7f5f737895892e
title: Побудова конвертера з матриці суміжності у список суміжності
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

У цій лабораторній роботі ви створите функцію, яка конвертує представлення графа у вигляді списку суміжності у матрицю суміжності. Список суміжності — це об’єкт, де кожен ключ представляє вузол, а відповідне значення — масив вузлів, з якими цей вузол з’єднаний. Матриця суміжності — це двовимірний масив, де елемент у позиції `[i][j]` дорівнює `1`, якщо існує ребро від вузла `i` до вузла `j`, і `0` в іншому випадку.

Наприклад, для списку суміжності:

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

Відповідна матриця суміжності буде:

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні визначити функцію з ім’ям `adjacencyListToMatrix`, яка конвертує список суміжності у матрицю суміжності.
2. Функція повинна приймати об’єкт, що представляє список суміжності неваженого (орієнтованого або неорієнтованого) графа, як свій параметр.
3. Функція повинна:
   - Конвертувати список суміжності у матрицю суміжності.
   - Вивести кожен рядок матриці суміжності.
   - Повернути матрицю суміжності.

Наприклад, `adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})` має вивести:

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

і повернути `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`.

# --hints--

Визначте функцію з ім’ям `adjacencyListToMatrix`.

```js
assert.isFunction(adjacencyListToMatrix);
```

Функція `adjacencyListToMatrix` повинна мати один параметр.

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

Функція повинна правильно визначати кількість вузлів зі списку суміжності.

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

Функція повинна правильно встановлювати значення матриці в `1` для існуючих ребер.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Функція повинна виводити кожен рядок матриці.

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

Функція повинна повертати матрицю суміжності.

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

Якщо функції передати список суміжності `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}`, вона повинна повернути `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`.

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Якщо функції передати список суміжності `{0: [1], 1: [0]}`, вона повинна повернути `[[0, 1], [1, 0]]`.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Якщо функції передати список суміжності `{0: [], 1: [], 2: []}`, вона повинна повернути `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
