---
id: 66325a250690a3612c1db0f6
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

CamperBot ha creado una nueva variable llamada `bgHexCodeSpanElement` para almacenar la referencia al elemento `span` con el `id` de `bg-hex-code`. Sin embargo, cuando intentan registrar esa variable en la consola, obtienen `null`.

`null` es un valor especial en JavaScript que representa la ausencia de un valor. Esto puede suceder cuando intentas acceder a una propiedad de un objeto que no existe.

En este caso, CamperBot no está pasando el selector correcto al método `document.querySelector`.

Corrige la línea `document.querySelector("bg-hex-code")` para que seleccione correctamente el elemento con el `id` de `bg-hex-code`.

# --hints--

Debe corregir `"bg-hex-code"` para que haga referencia correctamente al `id` de `bg-hex-code`. Recuerda que los selectores de `id` en CSS comienzan con un símbolo de almohadilla (`#`).

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

--fcc-editable-region--
const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("bg-hex-code");

console.log(bgHexCodeSpanElement);
--fcc-editable-region--
```
