---
id: 669e2f60e83c011754f711f9
title: Erstellen Sie eine Reisebüroseite
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `DOCTYPE`-Deklaration haben.
1. Sie sollten ein `html`-Element mit `lang` auf `en` gesetzt haben.
1. Sie sollten ein `head`-Element haben, das ein `meta`-Void-Element mit `charset` auf `utf-8` und einen `title` mit dem Text `Travel Agency Page` enthält.
1. Sie sollten ein `meta`-Tag in Ihrem `head`-Element haben, das eine kurze Beschreibung Ihrer Website für SEO enthält.
1. Sie sollten ein `h1`-Element haben, um Ihre Reiseziele zu präsentieren. Es sollte das einzige `h1`-Element auf Ihrer Seite sein.
1. Sie sollten einen Absatz unter dem `h1`-Element haben, der die Reisemöglichkeiten vorstellt.
1. Sie sollten ein `h2`-Element mit dem Text `Packages` haben.
1. Sie sollten ein `p`-Element haben, das die verschiedenen Pakete kurz vorstellt.
1. Sie sollten ein ungeordnetes Listen-Element mit zwei Listeneinträgen haben. Die beiden Listeneinträge sollten jeweils den Text `Group Travels` und `Private Tours` enthalten. Der Text jedes Listeneintrags sollte von einem Ankerelement umschlossen sein.
1. Sie sollten ein `h2`-Element mit dem Text `Top Itineraries` haben.
1. Sie sollten mindestens drei `figure`-Elemente haben, von denen jedes ein Ankerelement und ein `figcaption`-Element enthält, das das Bild der Abbildung beschreibt.
1. Die drei Ankerelemente sollten jeweils ein `img`-Element mit einem passenden `alt`-Attribut und einem `src`-Attribut, das auf ein gültiges Bild zeigt, enthalten. Sie können `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` und `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` verwenden, wenn Sie möchten.
1. Alle Ihre fünf Ankerelemente sollten ein `href`-Attribut mit dem Wert `https://www.freecodecamp.org/learn` und ein `target`-Attribut mit dem Wert `_blank` haben.

# --hints--

Ihre Reisebüroseite sollte eine `!DOCTYPE html` Deklaration haben.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Sie sollten ein `html` Element mit `lang` gesetzt auf `en` haben.

```js
assert.match(code, /<\s*html(?=[\s>])[^>]*\slang\s*=\s*('|")en\1[^>]*>/i);
```

Ihr `html` Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/html\s*>/i);
```

Sie sollten ein `head` Element innerhalb des `html` Elements haben.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>/i
);
```

Ihr `head`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/head\s*>/i);
```

Sie sollten ein `title`-Element innerhalb Ihres `head`-Elements haben.

```js
assert.isNotNull(document.querySelector('head title'));
```

Ihr `title`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/title\s*>/i);
```

Sie sollten zwei `meta` Elemente innerhalb Ihres `head` Elements haben.

```js
assert.isAtLeast(document.querySelectorAll('head meta').length, 2);
```

Ein `meta` Element sollte ein `name` Attribut mit dem Wert `description` und ein nicht-leeres `content` Attribut haben.

```js
assert(
  Array.from(document.querySelectorAll('head meta')).some(
    (meta) =>
      meta.getAttribute('name')?.toLowerCase() === 'description' &&
      meta.getAttribute('content')?.trim()
  )
);
```

Ein `meta`-Element sollte sein `charset`-Attribut auf `utf-8` gesetzt haben.

```js
const charsetMeta = document.querySelector('head meta[charset]');
assert.equal(charsetMeta?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Ihr `title`-Element sollte den Namen Ihrer Reiseagentur enthalten.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Sie sollten ein `body` Element innerhalb Ihres `html` Elements haben.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>(?:(?!<\/html\s*>)[\s\S])*<\s*body\s*>/i
);
```

Ihr `body`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/body\s*>/i);
```

Sie sollten ein `h1`-Element haben, um Ihre Reiseziele zu präsentieren.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Sie sollten nur ein `h1` Element haben.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Sie sollten ein `p`-Element unter dem `h1`-Element haben.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

Ihr erstes `p`-Element sollte die Reisemöglichkeiten vorstellen.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Ihr erstes `h2`-Element sollte den Text `Packages` enthalten.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

Sie sollten ein `p`-Element unter Ihrem ersten `h2`-Element haben.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Ihr zweites `p`-Element sollte die verschiedenen Pakete kurz vorstellen.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

Sie sollten ein ungeordnetes Listen-Element unter Ihrem zweiten `p`-Element haben.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

Sie sollten zwei Elemente in Ihrer ungeordneten Liste haben.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Beide Listenelemente sollten ein Ankerelement enthalten.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

Das Ankerelement Ihres ersten Listenelements sollte den Text `Group Travels` umschließen.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

Das Ankerelement Ihres zweiten Listenelements sollte den Text `Private Tours` umschließen.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

Sie sollten ein `h2`-Element nach Ihrer ungeordneten Liste haben.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

Ihr zweites `h2`-Element sollte den Text `Top Itineraries` enthalten.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

Sie sollten mindestens drei `figure`-Elemente haben.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Jedes `figure`-Element sollte als erstes verschachteltes Element ein Ankerelement enthalten.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Jedes `figure`-Element sollte als zweites verschachteltes Element ein `figcaption`-Element enthalten.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Jedes `figcaption` sollte einen Text enthalten.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Jedes der in Ihren `figure`-Elementen verschachtelten `a`-Elemente sollte ein Bild enthalten.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Jedes `img`-Element sollte ein gültiges `src`-Attribut haben.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Jedes `img`-Element sollte ein `alt`-Attribut mit einem passenden Wert haben.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Jedes `a`-Element sollte ein `href`-Attribut mit dem Wert `https://www.freecodecamp.org/learn` haben. Vergessen Sie nicht die Links in den Listenelementen.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Jedes `a`-Element sollte ein `target`-Attribut mit dem Wert `_blank` haben. Vergessen Sie nicht die Links in den Listenelementen.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
