---
id: 63b606f09a14cc1781aea1fb
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

폼에서 사용자가 일일 칼로리 예산을 나타내는 숫자를 입력할 수 있게 하세요.

`label` 요소를 만들고, `for` 속성을 `budget`로 설정한 다음 텍스트 `Budget`를 넣으세요. 그다음 `input` 요소를 만들고 `id`를 `budget`로 설정하세요.

# --hints--

`label` 안에 `form` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('form label'));
```

`label` 요소에 `for` 속성을 `budget`로 설정해야 합니다.

```js
assert.equal(document.querySelector('form label')?.getAttribute('for'), 'budget');
```

`label` 요소에 텍스트 `Budget`가 있어야 합니다.

```js
assert.equal(document.querySelector('form label')?.innerText, 'Budget');
```

`input` 안에 `form` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('form input'));
```

`input` 요소는 `label` 요소 다음에 와야 합니다.

```js
assert.equal(document.querySelector('form input')?.previousElementSibling?.tagName, "LABEL");
```

`input` 요소에 `id`가 `budget`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('form input')?.id, "budget");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
--fcc-editable-region--
        <form id="calorie-counter">

        </form>
--fcc-editable-region--
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
