---
id: 63b607af6fcdb119aae9b16a
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

폼에서 사용자가 다양한 식사 유형과 해당 칼로리 수치를 추가할 수 있어야 합니다.

`fieldset` 요소를 만들고 `id` 속성을 `breakfast`로 설정하세요.

그 요소 안에 텍스트가 `legend`인 `Breakfast`를 만들고, `div` 속성이 `class`로 설정된 빈 `input-container`를 만드세요.

# --hints--

`fieldset` 안에 `form` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('form fieldset'));
```

`fieldset` 요소 다음에 `input` 요소가 와야 합니다.

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

`fieldset` 요소에 `id`가 `breakfast`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

`legend` 안에 `fieldset` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('form fieldset legend'));
```

`legend` 요소에 텍스트 `Breakfast`가 있어야 합니다.

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

`div` 안에 `fieldset` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('form fieldset div'));
```

`div` 요소 다음에 `legend` 요소가 와야 합니다.

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

`div` 요소에 `class`가 `input-container`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
