---
id: 63b60aaaa65f8922bfce6b7e
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

`span` 요소 안에 `label`용 `entry-dropdown` 요소를 만들고 텍스트 `Add food or exercise:`을 넣으세요. 그다음 `select`가 `id`로 설정된 `entry-dropdown` 요소와 `name`가 `options`로 설정된 요소를 만드세요. 그 아래에 `button`이 `id`로 설정되고 텍스트 `add-entry`가 들어간 `Add Entry` 요소를 추가하세요.

자동 폼 제출을 막기 위해 `button` 요소에 `type` 속성을 `button`로 설정하세요.

# --hints--

`label` 요소에 `span` 요소를 추가해야 합니다.

```js
assert.exists(document.querySelector('.controls > span > label'));
```

새로운 `label` 요소에 `for` 속성을 `entry-dropdown`로 설정해야 합니다.

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

새로운 `label` 요소에 텍스트 `Add food or exercise:`를 넣어야 합니다.

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

`select` 요소에 `span` 요소를 추가해야 합니다.

```js
assert.exists(document.querySelector('.controls > span > select'));
```

`select` 요소는 `label` 요소 다음에 와야 합니다.

```js
assert.equal(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName, 'LABEL');
```

새로운 `select` 요소에 `id` 속성을 `entry-dropdown`로 설정해야 합니다.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

새로운 `select` 요소에 `name` 속성을 `options`로 설정해야 합니다.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

`select` 요소는 비어 있어야 하며 닫는 `</select>` 태그가 있어야 합니다.

```js
assert.equal(document.querySelector('.controls > span > select')?.innerHTML?.trim(), '');
```

`button` 요소에 `span` 요소를 추가해야 합니다.

```js
assert.exists(document.querySelector('.controls > span > button'));
```

`button` 요소는 `select` 요소 다음에 와야 합니다.

```js
assert.equal(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName, 'SELECT');
```

새로운 `button` 요소에 `id` 속성을 `add-entry`로 설정해야 합니다.

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

새로운 `button` 요소에 텍스트 `Add Entry`를 넣어야 합니다.

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
