---
id: 587d7b86367417b2b2512b3b
title: اكتشاف أخطاء الانحراف بمقدار واحد عند استخدام الفهرسة
challengeType: 1
forumTopicId: 301189
dashedName: catch-off-by-one-errors-when-using-indexing
---

# --description--

<dfn>أخطاء الانحراف بمقدار واحد</dfn> (تُسمى أحيانًا OBOE) تظهر عندما تحاول استهداف فهرس محدد في سلسلة نصية أو مصفوفة (لقطع جزء أو الوصول إلى مقطع)، أو عند التكرار على فهارسها. تبدأ فهارس JavaScript من الصفر، وليس من الواحد، مما يعني أن الفهرس الأخير دائمًا أقل بواحد من طول العنصر. إذا حاولت الوصول إلى فهرس يساوي الطول، قد يرمي البرنامج خطأ مرجع "الفهرس خارج النطاق" أو يطبع `undefined`.

عندما تستخدم طرق السلاسل النصية أو المصفوفات التي تأخذ نطاقات الفهارس كمعلمات، من المفيد قراءة الوثائق وفهم ما إذا كانت شاملة (العنصر عند الفهرس المعطى جزء مما يُعاد) أم لا. إليك بعض الأمثلة على أخطاء الانحراف بمقدار واحد:

```js
let alphabet = "abcdefghijklmnopqrstuvwxyz";
let len = alphabet.length;
for (let i = 0; i <= len; i++) {
  console.log(alphabet[i]);
}
for (let j = 1; j < len; j++) {
  console.log(alphabet[j]);
}
for (let k = 0; k < len; k++) {
  console.log(alphabet[k]);
}
```

المثال الأول هنا يكرر التكرار مرة أكثر من اللازم، والثاني يكرر مرة أقل من اللازم (يفتقد الفهرس الأول، 0). المثال الثالث صحيح.

# --instructions--

صحح خطأي الفهرسة في الدالة التالية حتى تُطبع جميع الأرقام من 1 إلى 5 في وحدة التحكم.

# --hints--

يجب أن يحدد شفرتك الشرط الابتدائي للحلقة بحيث تبدأ من الفهرس الأول.

```js
assert(__helpers.removeJSComments(code).match(/i\s*?=\s*?0\s*?;/g).length == 1);
```

يجب أن تصلح شفرتك الشرط الابتدائي للحلقة بحيث يبدأ الفهرس من 0.

```js
assert(!__helpers.removeJSComments(code).match(/i\s?=\s*?1\s*?;/g));
```

يجب أن يحدد شفرتك الشرط النهائي للحلقة بحيث تتوقف عند الفهرس الأخير.

```js
assert(__helpers.removeJSComments(code).match(/i\s*<\s*len\s*;|i\s*<=\s*len\s*-\s*1\s*;/g).length == 1);
```

يجب أن تصلح شفرتك الشرط النهائي للحلقة بحيث تتوقف قبل الطول بواحد.

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?<=\s*?len;/g));
```

# --seed--

## --seed-contents--

```js
function countToFive() {
  let firstFive = "12345";
  let len = firstFive.length;
  // Only change code below this line
  for (let i = 1; i <= len; i++) {
  // Only change code above this line
    console.log(firstFive[i]);
  }
}

countToFive();
```

# --solutions--

```js
function countToFive() {
 let firstFive = "12345";
 let len = firstFive.length;
 // Only change code below this line
 for (let i = 0; i < len; i++) {
 // Only change code above this line
   console.log(firstFive[i]);
 }
}

countToFive();
```
