---
id: 587d7b84367417b2b2512b37
title: اكتشاف الاستخدام المختلط لعلامات الاقتباس المفردة والمزدوجة
challengeType: 1
forumTopicId: 301188
dashedName: catch-mixed-usage-of-single-and-double-quotes
---

# --description--

تسمح JavaScript باستخدام كل من علامات الاقتباس المفردة (`'`) والمزدوجة (`"`) لإعلان سلسلة نصية. عادةً ما يعتمد اختيار أيهما على التفضيل الشخصي، مع بعض الاستثناءات.

وجود خيارين مفيد عندما تحتوي السلسلة النصية على اختصارات أو قطعة أخرى من النص داخل علامات اقتباس. فقط احرص على ألا تغلق السلسلة النصية مبكرًا، لأن ذلك يسبب خطأ نحوي.

إليك بعض الأمثلة على خلط علامات الاقتباس:

```js
const grouchoContraction = "I've had a perfectly wonderful evening, but this wasn't it.";
const quoteInString = "Groucho Marx once said 'Quote me as saying I was mis-quoted.'";
const uhOhGroucho = 'I've had a perfectly wonderful evening, but this wasn't it.';
```

الأولان صحيحان، لكن الثالث غير صحيح.

بالطبع، لا بأس باستخدام نمط واحد فقط من علامات الاقتباس. يمكنك الهروب من علامات الاقتباس داخل السلسلة النصية باستخدام حرف الهروب الشرطة المائلة العكسية (`\`):

```js
const allSameQuotes = 'I\'ve had a perfectly wonderful evening, but this wasn\'t it.';
```

# --instructions--

صحح السلسلة النصية بحيث تستخدم إما علامات اقتباس مختلفة لقيمة `href`، أو اهرب العلامات الموجودة. احتفظ بعلامات الاقتباس المزدوجة حول السلسلة النصية بأكملها.

# --hints--

يجب أن يصلح شفرتك علامات الاقتباس حول قيمة `href` `#Home` إما بتغييرها أو الهروب منها.

```js
assert(__helpers.removeJSComments(code).match(/<a href=\s*?('|\\")#Home\1\s*?>/g));
```

يجب أن تحافظ شفرتك على علامات الاقتباس المزدوجة حول السلسلة النصية بأكملها.

```js
assert(__helpers.removeJSComments(code).match(/"<p>.*?<\/p>";/g));
```

# --seed--

## --seed-contents--

```js
let innerHtml = "<p>Click here to <a href="#Home">return home</a></p>";
console.log(innerHtml);
```

# --solutions--

```js
let innerHtml = "<p>Click here to <a href=\"#Home\">return home</a></p>";
console.log(innerHtml);
```
