---
id: 587d7b85367417b2b2512b39
title: اكتشاف فقدان الأقواس العادية المفتوحة والمغلقة بعد استدعاء دالة
challengeType: 1
forumTopicId: 301185
dashedName: catch-missing-open-and-closing-parenthesis-after-a-function-call
---

# --description--

عندما لا تأخذ دالة أو طريقة أي معلمات، قد تنسى تضمين الأقواس العادية (الفارغة) المفتوحة والمغلقة عند استدعائها. غالبًا ما يُحفظ ناتج استدعاء الدالة في متغير لاستخدامه لاحقًا في شفرتك. يمكن اكتشاف هذا الخطأ بتسجيل قيم المتغيرات (أو أنواعها) في وحدة التحكم ورؤية أن أحدها مضبوط على مرجع دالة، بدلاً من القيمة المتوقعة التي تُرجعها الدالة.

المتغيرات في المثال التالي مختلفة:

```js
function myFunction() {
  return "You rock!";
}
let varOne = myFunction;
let varTwo = myFunction();
```

هنا `varOne` هي الدالة `myFunction`، و`varTwo` هي السلسلة النصية `You rock!`.

# --instructions--

صحح الشفرة بحيث يُضبط المتغير `result` على القيمة التي تُرجعها دالة `getNine` عند استدعائها.

# --hints--

يجب أن تصحح شفرتك المتغير `result` ليُضبط على الرقم الذي تُرجعه الدالة `getNine`.

```js
assert(result == 9);
```

يجب أن تستدعي شفرتك الدالة `getNine`.

```js
assert(__helpers.removeJSComments(code).match(/getNine\(\)/g).length == 2);
```

# --seed--

## --seed-contents--

```js
function getNine() {
  let x = 6;
  let y = 3;
  return x + y;
}

let result = getNine;
console.log(result);
```

# --solutions--

```js
function getNine() {
 let x = 6;
 let y = 3;
 return x + y;
}

let result = getNine();
console.log(result);
```
