---
id: 587d7b86367417b2b2512b3d
title: منع الحلقات اللانهائية بشرط نهائي صحيح
challengeType: 1
forumTopicId: 301192
dashedName: prevent-infinite-loops-with-a-valid-terminal-condition
---

# --description--

الموضوع الأخير هو الحلقة اللانهائية المخيفة. الحلقات أدوات رائعة عندما تحتاج أن يُنفذ برنامجك كتلة شفرة عددًا معينًا من المرات أو حتى يتحقق شرط ما، لكنها تحتاج إلى شرط نهائي يُنهي التكرار. الحلقات اللانهائية غالبًا ما تُجمد المتصفح أو تُعطله، وتتسبب في فوضى عامة في تنفيذ البرنامج، وهذا ما لا يريده أحد.

كان هناك مثال على حلقة لانهائية في مقدمة هذا القسم - لم يكن لديها شرط نهائي للخروج من حلقة `while` داخل `loopy()`. لا تستدعِ هذه الدالة!

```js
function loopy() {
  while(true) {
    console.log("Hello, world!");
  }
}
```

مهمة المبرمج هي التأكد من الوصول في النهاية إلى الشرط النهائي، الذي يُخبر البرنامج متى يخرج من شفرة الحلقة. أحد الأخطاء هو زيادة أو نقصان متغير العداد في الاتجاه الخاطئ بالنسبة للشرط النهائي. خطأ آخر هو إعادة تعيين متغير العداد أو الفهرس داخل شفرة الحلقة عن طريق الخطأ، بدلًا من زيادته أو إنقاصه.

# --instructions--

تحتوي الدالة `myFunc()` على حلقة لانهائية لأن الشرط النهائي `i != 4` لن يُقيّم أبدًا إلى `false` (ولذلك لا يكسر التكرار) - حيث تزداد قيمة `i` بمقدار 2 في كل مرة، وتتجاوز 4 مباشرة لأن `i` فردي في البداية. أصلح عامل المقارنة في الشرط النهائي بحيث تعمل الحلقة فقط عندما تكون قيمة `i` أقل من أو تساوي 4.

# --hints--

يجب أن يغير شفرتك عامل المقارنة في الشرط النهائي (الجزء الأوسط) من حلقة `for`.

```js
assert(__helpers.removeJSComments(code).match(/i\s*?<=\s*?4;/g).length == 1);
```

يجب أن تصلح شفرتك عامل المقارنة في الشرط النهائي للحلقة.

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?!=\s*?4;/g));
```

# --seed--

## --seed-contents--

```js
function myFunc() {
  for (let i = 1; i != 4; i += 2) {
    console.log("Still going!");
  }
}
```

# --solutions--

```js
function myFunc() {
 for (let i = 1; i <= 4; i += 2) {
   console.log("Still going!");
 }
}
```
