---
id: 587d7b83367417b2b2512b37
title: فهم الفروقات بين وحدة تحكم freeCodeCamp ووحدة تحكم المتصفح
challengeType: 1
forumTopicId: 301193
dashedName: understanding-the-differences-between-the-freecodecamp-and-browser-console
---

# --description--

ربما لاحظت أن بعض تحديات freeCodeCamp تتضمن وحدة تحكم خاصة بها. هذه الوحدة تتصرف بشكل مختلف قليلاً عن وحدة تحكم المتصفح.

هناك العديد من الطرق لاستخدامها مع `console` لإخراج الرسائل. مثل `log` و `warn` و `clear` على سبيل المثال. وحدة تحكم freeCodeCamp تعرض فقط رسائل `log`، بينما وحدة تحكم المتصفح تعرض جميع الرسائل. عند تعديل شفرتك، تُنفذ تلقائيًا وتُظهر السجلات. ثم تُمسح وحدة تحكم freeCodeCamp في كل مرة تُنفذ فيها الشفرة.

# --instructions--

أولًا، افتح وحدة تحكم المتصفح لتتمكن من رؤية السجلات. للقيام بذلك، انقر بزر الفأرة الأيمن على شريط التنقل الخاص بـ freeCodeCamp في الأعلى ثم انقر على `inspect` في معظم المتصفحات. بعد ذلك، ابحث عن تبويب `console` في النافذة التي تفتح.

بعد ذلك، استخدم `console.log` لتسجيل المتغير `output`. شاهد الوحدتين لترى السجل. أخيرًا، استخدم `console.clear` بعد تسجيلك لمسح وحدة تحكم المتصفح. لاحظ الفرق بين الوحدتين.

# --hints--

يجب أن تستخدم `console.log()` لطباعة المتغير `output`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/console\.log\(output\)/));
```

يجب أن تستخدم `console.clear()` لمسح وحدة تحكم المتصفح.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/console.clear\(\)/)
);
```

يجب مسح وحدة التحكم بعد تسجيل السجل.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/console\.log\(output\)[\s\S]*console.clear\(\)/)
);
```

# --seed--

## --seed-contents--

```js
let output = "Get this to show once in the freeCodeCamp console and not at all in the browser console";

```

# --solutions--

```js
let output = "Get this to show once in the freeCodeCamp console and not at all in the browser console";

console.log(output);
console.clear();
```
