---
id: 587d7b84367417b2b2512b36
title: "التقاط الأقواس العادية، الأقواس، الأقواس المتعرجة والاقتباسات غير المغلقة"
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--

خطأ نحوي آخر يجب الانتباه إليه هو أن جميع الأقواس العادية، والأقواس، والأقواس المتعرجة، والاقتباسات المفتوحة يجب أن يكون لها زوج إغلاق. نسيان جزء منها يحدث عادةً عند تعديل كود موجود وإدخال عناصر باستخدام أحد أنواع الأزواج. انتبه أيضًا عند تسكين كتل الكود داخل أخرى، مثل إضافة دالة استدعاء راجع كمعلمة لطريقة.

طريقة لتجنب هذا الخطأ هي أنه بمجرد كتابة الحرف المفتوح، أضف فورًا الزوج المغلق، ثم حرك المؤشر بينهما واستمر في كتابة الكود. لحسن الحظ، معظم محررات الكود الحديثة تولد النصف الثاني من الزوج تلقائيًا.

# --instructions--

صحح خطأي الزوج المفقودين في الكود.

# --hints--

يجب أن يصلح كودك الجزء المفقود من المصفوفة.

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

يجب أن يصلح كودك الجزء المفقود من طريقة `.reduce()`. يجب أن يظهر الإخراج في وحدة التحكم أن `Sum of array values is: 6`.

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```
