---
id: 587d7b85367417b2b2512b3a
title: التقاط المعلمات الممررة بترتيب خاطئ عند استدعاء دالة
challengeType: 1
forumTopicId: 301184
dashedName: catch-arguments-passed-in-the-wrong-order-when-calling-a-function
---

# --description--

مواصلة النقاش حول استدعاء الدوال، الخطأ التالي الذي يجب الحذر منه هو عندما تُمرر معلمات الدالة بترتيب غير صحيح. إذا كانت المعلمات من أنواع مختلفة، مثل دالة تتوقع مصفوفة وعددًا صحيحًا، فمن المحتمل أن يؤدي ذلك إلى رمي خطأ اثناء التشغيل. إذا كانت المعلمات من نفس النوع (جميعها أعداد صحيحة، على سبيل المثال)، فلن يكون منطق الشفرة منطقيًا. تأكد من تمرير جميع المعلمات المطلوبة بالترتيب الصحيح لتجنب هذه المشاكل.

# --instructions--

الدالة `raiseToPower` ترفع أساسًا إلى أس. للأسف، لم تُستدعَ بشكل صحيح - أصلح الشفرة بحيث تكون قيمة `power` هي 8 المتوقعة.

# --hints--

يجب أن تصلح شفرتك المتغير `power` ليكون مساويًا لـ 2 مرفوعًا للأس 3، وليس 3 مرفوعًا للأس 2.

```js
assert(power == 8);
```

يجب أن تستخدم شفرتك الترتيب الصحيح للمعلمات في استدعاء الدالة `raiseToPower`.

```js
assert(__helpers.removeJSComments(code).match(/raiseToPower\(\s*?base\s*?,\s*?exp\s*?\);/g));
```

# --seed--

## --seed-contents--

```js
function raiseToPower(b, e) {
  return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(exp, base);
console.log(power);
```

# --solutions--

```js
function raiseToPower(b, e) {
 return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(base, exp);
console.log(power);
```
