---
id: 587d7b84367417b2b2512b34
title: استخدام typeof لفحص نوع المتغير
challengeType: 1
forumTopicId: 18374
dashedName: use-typeof-to-check-the-type-of-a-variable
---

# --description--

يمكنك استخدام `typeof` لفحص هيكل البيانات، أو نوع، المتغير. هذا مفيد في تصحيح الأخطاء عند العمل مع أنواع بيانات متعددة. إذا ظننت أنك تجمع رقمين، لكن أحدهما في الواقع سلسلة نصية، قد تكون النتائج غير متوقعة. أخطاء النوع قد تختبئ في العمليات الحسابية أو استدعاءات الدوال. كن حذرًا خاصة عند وصولك والعمل مع بيانات خارجية على شكل كائن JavaScript Object Notation (JSON).

إليك بعض الأمثلة باستخدام `typeof`:

```js
console.log(typeof "");
console.log(typeof 0);
console.log(typeof []);
console.log(typeof {});
```

بالترتيب، ستعرض وحدة التحكم السلاسل النصية `string`، و`number`، و`object`، و`object`.

يتعرف JavaScript على سبعة أنواع بيانات أولية (غير قابلة للتغيير): `Boolean`، و`Null`، و`Undefined`، و`Number`، و`String`، و`Symbol` (جديد مع ES6)، و`BigInt` (جديد مع ES2020)، ونوع واحد للعناصر القابلة للتغيير: `Object`. لاحظ أن المصفوفات في JavaScript هي تقنيًا نوع من الكائنات.

# --instructions--

أضف بيانين `console.log()` لفحص `typeof` لكل من المتغيرين `seven` و`three` في الشفرة.

# --hints--

يجب أن يستخدم شفرتك `typeof` في بيانين `console.log()` لفحص نوع المتغيرات.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\s*\(typeof[\( ].*\)?\)/g).length == 2);
```

يجب أن يستخدم شفرتك `typeof` لفحص نوع المتغير `seven`.

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]seven\)?/g));
```

يجب أن يستخدم شفرتك `typeof` لفحص نوع المتغير `three`.

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]three\)?/g));
```

# --seed--

## --seed-contents--

```js
let seven = 7;
let three = "3";
console.log(seven + three);
// Only change code below this line
```

# --solutions--

```js
let seven = 7;let three = "3";console.log(typeof seven);
console.log(typeof three);
```
