---
id: 587d7b83367417b2b2512b33
title: استخدام وحدة التحكم في JavaScript لفحص قيمة متغير
challengeType: 1
forumTopicId: 18372
dashedName: use-the-javascript-console-to-check-the-value-of-a-variable
---

# --description--

يحتوي كل من Chrome وFirefox على وحدات تحكم JavaScript ممتازة، تُعرف أيضًا باسم أدوات المطوِّر، لتصحيح أخطاء شفرتك.

يمكنك العثور على أدوات المطوِّر في قائمة Chrome أو وحدة تحكم الويب في قائمة Firefox. إذا كنت تستخدم متصفحًا مختلفًا أو هاتفًا محمولًا، ننصح بشدة بالتحويل إلى Firefox أو Chrome على سطح المكتب.

طريقة `console.log()`، التي "تطبع" ناتج ما بداخل أقواسها العادية إلى وحدة التحكم، ستكون على الأرجح أداة تصحيح الأخطاء الأكثر فائدة. وضعها في نقاط استراتيجية في شفرتك يمكن أن يُظهر لك القيم الوسيطة للمتغيرات. من الممارسات الجيدة أن تكون لديك فكرة عن الناتج المتوقع قبل النظر إلى ما هو عليه. وجود نقاط فحص لرؤية حالة حساباتك طوال شفرتك سيساعد في تضييق مكان المشكلة.

إليك مثالًا لطباعة السلسلة النصية `Hello world!` إلى وحدة التحكم:

```js
console.log('Hello world!');
```

# --instructions--

استخدم طريقة `console.log()` لطباعة قيمة المتغير `a` حيث هو مذكور في الشفرة.

# --hints--

يجب أن تستخدم شفرتك `console.log()` لفحص قيمة المتغير `a`.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\(a\)/g));
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 1;
a++;
// Only change code below this line


let sumAB = a + b;
console.log(sumAB);
```

# --solutions--

```js
var a = 5; console.log(a);
```
