---
id: 587d7b85367417b2b2512b38
title: اكتشاف استخدام معامل التعيين بدلًا من معامل المساواة
challengeType: 1
forumTopicId: 301191
dashedName: catch-use-of-assignment-operator-instead-of-equality-operator
---

# --description--

تعتمد البرامج المتفرعة، أي التي تنفذ أفعالًا مختلفة إذا تحققت شروط معينة، على عبارات `if` و `else if` و `else` في JavaScript. أحيانًا يكون الشرط على شكل اختبار ما إذا كانت النتيجة تساوي قيمة معينة.

يُقال هذا المنطق (بالإنجليزية على الأقل) كـ "if x equals y, then ..."، والذي يمكن ترجمته حرفيًا إلى شفرة باستخدام `=`، أو معامل التعيين. هذا يؤدي إلى تدفق تحكم غير متوقع في برنامجك.

كما تم تغطيته في التحديات السابقة، معامل التعيين (`=`) في JavaScript يُعيّن قيمة إلى اسم متغير. أما المعاملان `==` و `===` فيختبران المساواة (المعامل الثلاثي `===` يختبر المساواة الصارمة، أي أن القيمة والنوع متطابقان).

الشفرة أدناه تُعيّن `x` لتكون 2، والتي تُقيّم كـ `true`. تقريبًا كل قيمة بمفردها في JavaScript تُقيّم كـ `true`، ما عدا القيم المعروفة باسم "القيم الزائفة": `false`، `0`، `""` (سلسلة نصية فارغة)، `NaN`، `undefined`، و `null`.

```js
let x = 1;
let y = 2;
if (x = y) {

} else {

}
```

في هذا المثال، كتلة الشفرة داخل عبارة `if` ستُنفذ لأي قيمة من `y`، ما لم تكن `y` زائفة. كتلة `else`، التي نتوقع تنفيذها هنا، لن تُنفذ فعليًا.

# --instructions--

صحح الشرط ليُشغّل البرنامج الفرع الصحيح، ويُعيّن القيمة المناسبة إلى `result`.

# --hints--

يجب أن تُصلح شفرتك الشرط ليختبر المساواة بدلًا من استخدام التعيين.

```js
assert(result == 'Not equal!');
```

يجب أن يستخدم الشرط إما `==` أو `===` لاختبار المساواة.

```js
assert(__helpers.removeJSComments(code).match(/x\s*?===?\s*?y/g));
```

# --seed--

## --seed-contents--

```js
let x = 7;
let y = 9;
let result = "to come";

if(x = y) {
  result = "Equal!";
} else {
  result = "Not equal!";
}

console.log(result);
```

# --solutions--

```js
let x = 7;
let y = 9;
let result = "to come";

if(x === y) {
 result = "Equal!";
} else {
 result = "Not equal!";
}

console.log(result);
```
