---
id: 587d7b86367417b2b2512b3c
title: استخدم الحذر عند إعادة تهيئة المتغيرات داخل حلقة
challengeType: 1
forumTopicId: 301194
dashedName: use-caution-when-reinitializing-variables-inside-a-loop
---

# --description--

أحيانًا تحتاج إلى حفظ معلومات، زيادة العدادات، أو إعادة تعيين المتغيرات داخل حلقة. المشكلة المحتملة هي عندما يجب إعادة تهيئة المتغيرات ولا تُعاد، أو العكس. هذا خطر بشكل خاص إذا أعدت تعيين المتغير المستخدم في شرط النهاية عن طريق الخطأ، مما يسبب حلقة لا نهائية.

طباعة قيم المتغيرات مع كل دورة من الحلقة باستخدام `console.log()` يمكن أن تكشف عن سلوك خاطئ متعلق بإعادة التعيين أو الفشل في إعادة تعيين متغير.

# --instructions--

الدالة التالية من المفترض أن تنشئ مصفوفة ثنائية الأبعاد تحتوي على `m` صفوف و`n` أعمدة من الأصفار. للأسف، لا تعطي النتيجة المتوقعة لأن المتغير `row` لا يُعاد تهيئته (يُعاد تعيينه إلى مصفوفة فارغة) في الحلقة الخارجية. أصلح الشفرة لتُرجع مصفوفة صحيحة بحجم 3x2 من الأصفار، والتي تبدو كـ `[[0, 0], [0, 0], [0, 0]]`.

# --hints--

يجب أن تُعيّن شفرتك المتغير `matrix` إلى مصفوفة تحتوي على 3 صفوف و2 عمود من الأصفار في كل منها.

```js
assert(JSON.stringify(matrix) == '[[0,0],[0,0],[0,0]]');
```

يجب أن يحتوي المتغير `matrix` على 3 صفوف.

```js
assert(matrix.length == 3);
```

يجب أن يحتوي المتغير `matrix` على 2 عمود في كل صف.

```js
assert(
  matrix[0].length == 2 && matrix[1].length === 2 && matrix[2].length === 2
);
```

يجب أن تُرجع `zeroArray(4,3)` مصفوفة تحتوي على 4 صفوف و3 أعمدة من الأصفار في كل منها.

```js
assert(JSON.stringify(zeroArray(4,3)) == '[[0,0,0],[0,0,0],[0,0,0],[0,0,0]]');
```

# --seed--

## --seed-contents--

```js
function zeroArray(m, n) {
  // Creates a 2-D array with m rows and n columns of zeroes
  let newArray = [];
  let row = [];
  for (let i = 0; i < m; i++) {
    // Adds the m-th row into newArray

    for (let j = 0; j < n; j++) {
      // Pushes n zeroes into the current row to create the columns
      row.push(0);
    }
    // Pushes the current row, which now has n zeroes in it, to the array
    newArray.push(row);
  }
  return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```

# --solutions--

```js
function zeroArray(m, n) {
 // Creates a 2-D array with m rows and n columns of zeroes
 let newArray = [];
 for (let i = 0; i < m; i++) {
   let row = [];
   // Adds the m-th row into newArray

   for (let j = 0; j < n; j++) {
     // Pushes n zeroes into the current row to create the columns
     row.push(0);
   }
   // Pushes the current row, which now has n zeroes in it, to the array
   newArray.push(row);
 }
 return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```
