---
id: 587d7b84367417b2b2512b35
title: اكتشاف أسماء المتغيرات والدوال المكتوبة بشكل خاطئ
challengeType: 1
forumTopicId: 301186
dashedName: catch-misspelled-variable-and-function-names
---

# --description--

تُعد الطريقتان `console.log()` و `typeof` الطريقتين الرئيسيتين لفحص القيم الوسيطة وأنواع ناتج البرنامج. حان الوقت الآن للتعرف على الأشكال الشائعة التي تتخذها الأخطاء البرمجية. إحدى المشكلات على مستوى الصيغة النحوية التي يمكن لمن يكتبون بسرعة التعاطف معها هي خطأ التهجئة البسيط.

الأحرف المنقلبة أو المفقودة أو المكتوبة بحروف كبيرة أو صغيرة بشكل خاطئ في اسم متغير أو دالة تجعل المتصفح يبحث عن كائن غير موجود - ويُظهر شكوى على شكل خطأ مرجع. أسماء المتغيرات والدوال في JavaScript حساسة لحالة الأحرف.

# --instructions--

صحح خطأي التهجئة في الشفرة حتى يعمل حساب `netWorkingCapital` بشكل صحيح.

# --hints--

تحقق من تهجئة المتغيرين المستخدمين في حساب netWorkingCapital، يجب أن يظهر ناتج وحدة التحكم "رأس المال العامل الصافي هو: 2".

```js
assert(netWorkingCapital === 2);
```

يجب ألا توجد أي حالات لمتغيرات مكتوبة بشكل خاطئ في الشفرة.

```js
assert(!__helpers.removeJSComments(code).match(/recievables/g));
```

يجب إعلان واستخدام المتغير `receivables` بشكل صحيح في الشفرة.

```js
assert(__helpers.removeJSComments(code).match(/receivables/g).length == 2);
```

يجب ألا توجد أي حالات لمتغيرات مكتوبة بشكل خاطئ في الشفرة.

```js
assert(!__helpers.removeJSComments(code).match(/payable;/g));
```

يجب إعلان واستخدام المتغير `payables` بشكل صحيح في الشفرة.

```js
assert(__helpers.removeJSComments(code).match(/payables/g).length == 2);
```

# --seed--

## --seed-contents--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = recievables - payable;
console.log(`Net working capital is: ${netWorkingCapital}`);
```

# --solutions--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = receivables - payables;
console.log(`Net working capital is: ${netWorkingCapital}`);
```
