---
id: 587d7b84367417b2b2512b36
title: "التقاط الأقواس العادية، الأقواس، الأقواس المتعرجة والاقتباسات غير المغلقة"
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--

خطأ نحوي آخر يجب الانتباه إليه هو أن جميع الأقواس العادية، الأقواس، الأقواس المتعرجة، والاقتباسات يجب أن تكون لها زوج إغلاق. نسيان جزء منها يحدث عادةً عند تعديل شفرة موجودة وإدخال عناصر من نوع أحد الأزواج. أيضًا، انتبه عند التسكين داخل كتل شفرة أخرى، مثل إضافة دالة استدعاء راجع كمعلمة لطريقة.

طريقة لتجنب هذا الخطأ هي أنه بمجرد كتابة حرف الفتح، أدرج فورًا الزوج المغلق، ثم حرّك المؤشر بينهما واستمر في الترميز. لحسن الحظ، معظم محررات الشفرة الحديثة تولد النصف الثاني من الزوج تلقائيًا.

# --instructions--

صحح خطأي الزوجين في الشفرة.

# --hints--

يجب أن تصلح شفرتك الجزء المفقود من المصفوفة.

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

يجب أن تصلح شفرتك الجزء المفقود من طريقة `.reduce()`. يجب أن يظهر في وحدة التحكم أن `Sum of array values is: 6`.

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```
