---
id: 69a56b5069ca99f7317e6e19
title: Erstellen Sie einen Frachtmanifest-Validator
challengeType: 26
dashedName: lab-cargo-manifest-validator
---

# --description--

In diesem Labor verwenden Sie JavaScript, um Frachtmanifeste zu normalisieren und zu validieren. Ein Frachtmanifest ist ein Dokument, das typischerweise die transportierten Güter auflistet (zum Beispiel per Schiff oder Zug) und Details zu diesen Gütern enthält.

Jedes Frachtmanifest wird als Objekt mit den folgenden Eigenschaften dargestellt:

- `containerId`: eine positive ganze Zahl, die den zugehörigen Frachtcontainer identifiziert.
- `destination`: ein nicht leerer String (nach Entfernen von Leerzeichen), der das Ziel der Fracht angibt.
- `weight`: eine positive Zahl, die das Gewicht der Fracht darstellt.
- `unit`: ein String, der die Einheit für die Gewichtseigenschaft der Fracht beschreibt (entweder `"kg"` für Kilogramm oder `"lb"` für Pfund).
- `hazmat`: ein Boolean-Wert, der angibt, ob eine Handhabung von Gefahrgut erforderlich ist.

Beispiel für ein Frachtmanifest-Objekt:

```js
{
  containerId: 1,
  destination: "Monterey, California, USA",
  weight: 831,
  unit: "lb",
  hazmat: false
}
```

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `normalizeUnits` mit einem `manifest`-Parameter implementieren.
    - Die Funktion darf das ursprüngliche Manifestobjekt nicht verändern und muss immer ein neues Objekt zurückgeben, bei dem `weight` auf Kilogramm normalisiert ist und `unit` auf `"kg"` gesetzt wird.
    - Wenn das Gewicht des Manifestobjekts in Pfund (`unit: "lb"`) angegeben ist, sollte die Funktion das `weight` mit der ungefähren Umrechnung 1 lb = 0,45 kg in Kilogramm umrechnen und die `unit` entsprechend aktualisieren.
    - Wenn das Gewicht bereits in Kilogramm (`unit: "kg"`) angegeben ist, sollten `weight` und `unit` unverändert bleiben.
2. Sie sollten eine Funktion namens `validateManifest` mit einem `manifest`-Parameter implementieren.
    - Die Funktion darf das ursprüngliche Manifestobjekt nicht verändern und muss immer ein neues Objekt zurückgeben.
    - Wenn das Eingabe-Manifest gültig ist (keine fehlenden oder ungültigen Eigenschaften), sollte die Funktion ein leeres Objekt zurückgeben.
    - Wenn das Eingabe-Manifest nicht gültig ist, sollte die Funktion ein Objekt zurückgeben, das Einträge für jede fehlende oder ungültige Eigenschaft enthält. Fehlende Eigenschaften sollten den Wert `"Missing"` haben und ungültige Eigenschaften den Wert `"Invalid"`.

    Beispiel für einen Rückgabewert, bei dem das Eingabeobjekt die Eigenschaft `destination` vermisst und eine ungültige `weight` Eigenschaft hat:

    ```js
    {
      destination: "Missing",
      weight: "Invalid"
    }
    ```

3. Sie sollten eine Funktion namens `processManifest` mit einem `manifest`-Parameter implementieren. Die Funktion sollte Folgendes ausgeben:
    - Wenn das Manifestobjekt gültig ist, `Validation success: ${containerId}` und dann das Gewicht des Manifests in Kilogramm, so: `Total weight: ${weight} kg`. Verwenden Sie `normalizeUnits()` für diese Umrechnung.
    - Wenn das Manifestobjekt nicht gültig ist, `Validation error: ${containerId}` und dann das Objekt, das durch den Aufruf von `validateManifest()` mit dem Manifestobjekt zurückgegeben wird.
    **Hinweis:** Jeder dieser beiden Fälle sollte zwei `console.log()`-Aufrufe enthalten.

# --before-each--

```js
const _validManifests = [
  {
    containerId: 4,
    destination: "Monterey, California, USA",
    weight: 831,
    unit: "lb",
    hazmat: false
  },
  {
    containerId: 5,
    destination: "Montreal, Quebec, Canada",
    weight: 151,
    unit: "kg",
    hazmat: false
  }
];

const _invalidManifests = [
  {
    containerId: -6,
    destination: 123,
    weight: 0,
    unit: "pounds",
    hazmat: true
  },
  {
    containerId: 0,
    weight: -21,
    unit: "KG"
  },
  {
    containerId: "eight",
    destination: "Guadalajara, Jalisco, Mexico",
    weight: 9001,
    hazmat: "no"
  },
  {
    destination: " ",
    weight: NaN,
    unit: "kg",
    hazmat: 1
  },
  {
    containerId: 10,
    destination: " ",
  },
];

const _testObj = {
  containerId: 1.51,
  destination: "Lavender Town",
  weight: NaN,
};

console.log = () => {};
```

# --hints--

Sie sollten eine Funktion namens `normalizeUnits` mit einem `manifest`-Parameter haben.

```js
assert.isFunction(normalizeUnits);
const funcString = normalizeUnits.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

Der Aufruf von `normalizeUnits()` mit `{ containerId: 68, destination: "Salinas", weight: 101, unit: "lb", hazmat: true }` sollte das neue Objekt `{ containerId: 68, destination: "Salinas", weight: 45.45, unit: "kg", hazmat: true }` zurückgeben, ohne die ursprüngliche Eingabe zu verändern.

```js
const testObj = {
  containerId: 68,
  destination: "Salinas", 
  weight: 101,
  unit: "lb",
  hazmat: true
};

const expected = {
  containerId: 68,
  destination: "Salinas",
  weight: 45.45,
  unit: "kg",
  hazmat: true
};

const copy = normalizeUnits(testObj);

assert.isObject(copy);
assert.deepEqual(copy, expected, "normalizeUnits() did not return the expected normalized object");
assert.notStrictEqual(copy, testObj, "normalizeUnits() should return a new object.");
assert.strictEqual(testObj.containerId, 68, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, "Salinas", "Original object should not be mutated.");
assert.strictEqual(testObj.weight, 101, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "lb", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, true, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 5, "Original object should not be mutated.");
```

Ihre `normalizeUnits`-Funktion sollte eine Kopie des Eingabe-Manifestobjekts zurückgeben, bei der `weight` auf Kilogramm normalisiert und `unit` auf `"kg"` gesetzt ist. Verwenden Sie die ungefähre Umrechnung `1 lb = 0.45 kg` für die Gewichtsumrechnung.

```js
_validManifests.forEach((obj, i) => {
  const result = normalizeUnits(obj);
  assert.isObject(result, "normalizeUnits should return an object");
  assert.notStrictEqual(result, obj, "normalizeUnits must return a new object, not modify the original");
  const expectedWeight = obj.unit === "lb" ? obj.weight * 0.45 : obj.weight;
  const expected = {
    ...obj,
    weight: expectedWeight,
    unit: "kg"
  };

  assert.deepEqual(
    result,
    expected,
    `normalizeUnits did not return the expected normalized object for manifest at index ${i}`
  );
});
```

Ihre `normalizeUnits`-Funktion sollte eine neue Kopie des Eingabe-`manifest`-Objekts zurückgeben, ohne das Original zu verändern.

```js
const original = {
  containerId: 151,
  destination: "Saffron City",
  weight: 151,
  unit: "lb",
  hazmat: false
};

const expected = {
  containerId: 151,
  destination: "Saffron City",
  weight: 67.95,
  unit: "kg",
  hazmat: false
};

const copy = normalizeUnits(original);
assert.isObject(copy);
assert.deepEqual(copy, expected);
assert.notStrictEqual(copy, original, "normalizeUnits() should return a new object.");
assert.strictEqual(original.weight, 151, "Original object should not be mutated.");
assert.strictEqual(original.unit, "lb", "Original object should not be mutated.");
assert.equal(Object.keys(original).length, 5, "Original object should not be mutated.");
```

Sie sollten eine Funktion namens `validateManifest` mit einem `manifest`-Parameter haben.

```js
assert.isFunction(validateManifest);
const funcString = validateManifest.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

Der Aufruf von `validateManifest()` mit `{ containerId: 1, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }` sollte das neue Objekt `{}` zurückgeben.

```js
const testObj = {
  containerId: 1,
  destination: "Santa Cruz", 
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {};

const copy = validateManifest(testObj);

assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for a valid manifest");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, 1, "Original object should not be mutated");
assert.strictEqual(testObj.destination, "Santa Cruz", "Original object should not be mutated");
assert.strictEqual(testObj.weight, 304, "Original object should not be mutated");
assert.strictEqual(testObj.unit, "kg", "Original object should not be mutated");
assert.strictEqual(testObj.hazmat, false, "Original object should not be mutated");
```

Wenn das Eingabe-Manifestobjekt gültig ist, sollte Ihre `validateManifest`-Funktion ein leeres Objekt `{}` zurückgeben.

```js
const solution = {}
for (const obj of _validManifests) {
  assert.deepEqual(
    validateManifest(obj),
    solution,
    "validateManifest() should return an empty object for valid manifests."
  );
}
```

Der Aufruf von `validateManifest()` mit `{}` sollte das neue Objekt `{ containerId: "Missing", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` zurückgeben, ohne die ursprüngliche Eingabe zu verändern.

```js
const testObj = {

};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.equal(Object.keys(testObj).length, 0, "validateManifest() should return a new object, not mutate the original")
```

Der Aufruf von `validateManifest()` mit `{ containerId: null, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }` sollte das neue Objekt `{ containerId: "Invalid" }` zurückgeben, ohne die ursprüngliche Eingabe zu verändern.

```js
const testObj = {
  containerId: null,
  destination: "Santa Cruz",
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {
  containerId: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for null containerId");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
```

Der Aufruf von `validateManifest()` mit `{ containerId: 0, destination: 405, weight: -84, unit: "pounds", hazmat: "no" }` sollte das neue Objekt `{ containerId: "Invalid", destination: "Invalid", weight: "Invalid", unit: "Invalid", hazmat: "Invalid" }` zurückgeben, ohne die ursprüngliche Eingabe zu verändern.

```js
const testObj = {
  containerId: 0,
  destination: 405,
  weight: -84,
  unit: "pounds",
  hazmat: "no"
};

const expected = {
  containerId: "Invalid",
  destination: "Invalid",
  weight: "Invalid",
  unit: "Invalid",
  hazmat: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, 0, "Original object should not be mutated");
assert.strictEqual(testObj.destination, 405, "Original object should not be mutated");
assert.strictEqual(testObj.weight, -84, "Original object should not be mutated");
assert.strictEqual(testObj.unit, "pounds", "Original object should not be mutated");
assert.strictEqual(testObj.hazmat, "no", "Original object should not be mutated");
```

Der Aufruf von `validateManifest()` mit `{ containerId: -2 }` sollte das neue Objekt `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` zurückgeben, ohne die ursprüngliche Eingabe zu verändern.

```js
const testObj = {
  containerId: -2,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, -2, "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

Der Aufruf von `validateManifest()` mit `{ containerId: 3.50 }` sollte das neue Objekt `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` zurückgeben, ohne die ursprüngliche Eingabe zu verändern. Sie können `Number.isInteger()` verwenden, um Ganzzahlen zu validieren.

```js
const testObj = {
  containerId: 3.50,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, 3.50, "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

Der Aufruf von `validateManifest()` mit `{ destination: "  " }` sollte das neue Objekt `{ containerId: "Missing", destination: "Invalid", weight: "Missing", unit: "Missing", hazmat: "Missing" }` zurückgeben, ohne die ursprüngliche Eingabe zu verändern. Sie können `String.trim()` verwenden, um Leerzeichen aus einem String zu entfernen.

```js
const testObj = {
  destination: "  ",
};

const expected = {
  containerId: "Missing",
  destination: "Invalid",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.destination, "  ", "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

Der Aufruf von `validateManifest()` mit `{ weight: NaN }` sollte das neue Objekt `{ containerId: "Missing", destination: "Missing", weight: "Invalid", unit: "Missing", hazmat: "Missing" }` zurückgeben, ohne die ursprüngliche Eingabe zu verändern. Sie können `Number.isNaN()` verwenden, um NaN-Werte zu validieren.

```js
const testObj = {
  weight: NaN,
};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Invalid",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.ok(Number.isNaN(testObj.weight), "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

Wenn das Eingabe-Manifestobjekt nicht gültig ist, sollte Ihre `validateManifest`-Funktion ein Objekt zurückgeben, das fehlende und/oder ungültige Eigenschaften beschreibt.

```js
const solutions = [
  { containerId: 'Invalid', destination: 'Invalid', weight: 'Invalid', unit: 'Invalid' },
  { containerId: 'Invalid', destination: 'Missing', weight: 'Invalid', unit: 'Invalid', hazmat: 'Missing' },
  { containerId: 'Invalid', unit: 'Missing', hazmat: 'Invalid' },
  { containerId: 'Missing', destination: 'Invalid', weight: 'Invalid', hazmat: 'Invalid' },
  { destination: 'Invalid', weight: 'Missing', unit: 'Missing', hazmat: 'Missing' }
];

let i = 0;
for (const obj of _invalidManifests) {
  const result = validateManifest(obj);
  assert.deepEqual(
    result,
    solutions[i],
    "validateManifest() should return an object with errors for invalid manifests."
  );
  i += 1;
}
```

Ihre `validateManifest`-Funktion sollte ein neues Objekt zurückgeben, ohne das Original zu verändern.

```js
const invalidOriginal = {
  containerId: -151,
  destination: "",
  weight: -151,
  unit: "kilograms",
  hazmat: "idk"
};

const validOriginal = {
  containerId: 151,
  destination: "Pallet Town",
  weight: 151,
  unit: "kg",
  hazmat: false
};

const res0 = validateManifest(invalidOriginal);
assert.isObject(res0);
assert.notStrictEqual(res0, invalidOriginal, "validateManifest() should return a new object.");
assert.strictEqual(invalidOriginal.containerId, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.destination, "", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.weight, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.unit, "kilograms", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.hazmat, "idk", "Original object should not be mutated.");
assert.equal(Object.keys(invalidOriginal).length, 5, "Original object should not be mutated.");

const res1 = validateManifest(validOriginal);
assert.isObject(res1);
assert.notStrictEqual(res1, validOriginal, "validateManifest() should return a new object.");
assert.strictEqual(validOriginal.containerId, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.destination, "Pallet Town", "Original object should not be mutated.");
assert.strictEqual(validOriginal.weight, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.unit, "kg", "Original object should not be mutated.");
assert.strictEqual(validOriginal.hazmat, false, "Original object should not be mutated.");
assert.equal(Object.keys(validOriginal).length, 5, "Original object should not be mutated.");
```

Sie sollten eine Funktion namens `processManifest` mit einem `manifest`-Parameter haben.

```js
assert.isFunction(processManifest);
const funcString = processManifest.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

Der Aufruf von `processManifest()` mit `{ containerId: 55, destination: "Carmel", weight: 400, unit: "lb", hazmat: false }` sollte zuerst `"Validation success: 55"` ausgeben und dann `"Total weight: 180 kg"`.

```js
const testObj = {
  containerId: 55,
  destination: "Carmel",
  weight: 400,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);

  const expectedCalls = [
    ["Validation success: 55"],
    ["Total weight: 180 kg"]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation success message");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the normalized weight in kg");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Wenn das Eingabe-Manifestobjekt gültig ist, sollte Ihre `processManifest`-Funktion zuerst die Erfolgsmeldung `Validation success: ${containerId}` ausgeben.

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ]
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Wenn das Eingabe-Manifestobjekt gültig ist, sollte Ihre `processManifest`-Funktion es mit `normalizeUnits()` auf Kilogramm normalisieren und dann Folgendes ausgeben: `Total weight: ${weight} kg`.

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ]
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Wenn das Eingabe-Manifestobjekt gültig ist, sollte Ihre `processManifest`-Funktion eine Erfolgsmeldung mit dem `containerId` des Objekts ausgeben und dann das `weight` des Objekts in Kilogramm ausgeben. Sie sollten `normalizeUnits()` für die Umrechnung verwenden und zwei `console.log()`-Aufrufe haben.

```js
const originalNormalize = normalizeUnits;
let called = false;
try {
  normalizeUnits = function (...args) {
    called = true;
    return originalNormalize(...args);
  };

  processManifest(_validManifests[0]);

  assert.isTrue(
    called,
    "processManifest should call normalizeUnits() when the manifest is valid"
  );

} catch (err) {
  assert.fail(err);
} finally {
  normalizeUnits = originalNormalize;
}

const spy = __helpers.spyOn(console, "log");

try {
  for (const obj of _validManifests) {
    processManifest(obj);
  }
  const expectedCalls = [
    ["Validation success: 4"],
    ["Total weight: 373.95 kg"],
    ["Validation success: 5"],
    ["Total weight: 151 kg"],
  ]
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Der Aufruf von `processManifest()` mit `{ containerId: -88, destination: "Soledad", weight: NaN }` sollte zuerst `Validation error: -88` ausgeben und dann das Objekt `{ containerId: "Invalid", weight: "Invalid", unit: "Missing", hazmat: "Missing" }`.

```js
const testObj = {
  containerId: -88,
  destination: "Soledad",
  weight: NaN
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: -88"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Der Aufruf von `processManifest()` mit `{ destination: "Watsonville", hazmat: true }` sollte zuerst `Validation error: undefined` ausgeben und dann das Objekt `{ containerId: "Missing", weight: "Missing", unit: "Missing" }`.

```js
const testObj = {
  destination: "Watsonville",
  hazmat: true
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: undefined"],
    [{
        containerId: "Missing",
        weight: "Missing",
        unit: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Wenn das Eingabe-Manifestobjekt nicht gültig ist, sollte Ihre `processManifest`-Funktion zuerst die Fehlermeldung `Validation error: ${containerId}` ausgeben.

```js
const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);   
} finally {
  spy.restore();
}
```

Wenn das Eingabe-Manifestobjekt nicht gültig ist, sollte Ihre `processManifest`-Funktion auch das Objekt ausgeben, das durch den Aufruf von `validateManifest()` mit dem ursprünglichen Manifestobjekt zurückgegeben wird. Rufen Sie `console.log()` direkt mit dem zurückgegebenen Objekt auf.

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest should call validateManifest() when the manifest is invalid"
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);   
} finally {
  spy.restore();
}
```

Wenn das Eingabe-Manifestobjekt nicht gültig ist, sollte Ihre `processManifest`-Funktion eine Fehlermeldung mit dem `containerId` des Objekts ausgeben und dann das Objekt ausgeben, das durch den Aufruf von `validateManifest()` mit dem Eingabeobjekt zurückgegeben wird. Dies sollte zwei `console.log()`-Aufrufe beinhalten.

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest should call validateManifest() when the manifest is invalid"
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);   
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function normalizeUnits(manifest) {
  const normalized = {
    ...manifest
  }
  if (normalized.unit === "lb") {
    normalized.weight = normalized.weight * 0.45;
    normalized.unit = "kg";
  }
  return normalized;
}

function validateManifest(manifest) {
  const result = {}

  if (manifest.containerId === undefined) {
    result.containerId = "Missing";
  } else if (
    typeof manifest.containerId !== "number" ||
    !Number.isInteger(manifest.containerId) ||
    manifest.containerId <= 0
  ) {
    result.containerId = "Invalid";
  }
  if (manifest.destination === undefined) {
    result.destination = "Missing";
  } else if (typeof manifest.destination !== "string" || manifest.destination.trim() === "") {
    result.destination = "Invalid";
  }
  if (manifest.weight === undefined) {
    result.weight = "Missing";
  } else if (typeof manifest.weight !== "number" || manifest.weight <= 0 || Number.isNaN(manifest.weight)) {
    result.weight = "Invalid";
  }
  if (manifest.unit === undefined) {
    result.unit = "Missing";
  } else if (manifest.unit !== "kg" && manifest.unit !== "lb") {
    result.unit = "Invalid";
  }
  if (manifest.hazmat === undefined) {
    result.hazmat = "Missing";
  } else if (typeof manifest.hazmat !== "boolean") {
    result.hazmat = "Invalid";
  }

  return result;
}

function processManifest(manifest) {
  const validation = validateManifest(manifest);

  if (Object.keys(validation).length > 0) {
    console.log(`Validation error: ${manifest.containerId}`);
    console.log(validation);
  } else {
    const normalized = normalizeUnits(manifest);
    console.log(`Validation success: ${manifest.containerId}`);
    console.log(`Total weight: ${normalized.weight} kg`);
  }
}
```
