---
id: 6606b224a69a293f98f8db8f
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Метод `.index()` викликає виняток `ValueError`, коли значення не знайдено. Щоб запобігти зупинці виконання програми, ви маєте вкласти цей рядок коду всередину блоку `try`. Інструкція `try` використовується для інкапсуляції коду, який може викликати виняток. Умова `except`, натомість, пропонує альтернативний код для виконання, якщо виникає виняток:

```py
try:
    <code>
except:
    <code>
```

Помістіть присвоєння `col` всередину блоку `try`. Потім створіть умову `except` і заповніть її тіло `pass`.

# --hints--

Ви маєте помістити присвоєння `col` всередину блоку `try` і створити умову `except`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
