---
id: 6606a219f9efbf38ad496f67
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Ви також можете використовувати точкову нотацію, щоб отримати доступ до атрибута екземпляра.

Поза межами класу `Board`, після ініціалізації об’єкта `gameboard`, використайте `gameboard.board`, щоб отримати доступ до атрибута `board` вашого об’єкта `gameboard` і виведіть результат на екран.

# --hints--

Ви повинні вивести `gameboard.board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.board)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
