---
id: 66069167b3307b2f4067b22b
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Новий екземпляр класу створюється за допомогою нотації функції, яка полягає у додаванні пари круглих дужок до імені класу.

Поза межами визначення класу створіть екземпляр класу `Board` і призначте його змінній з ім’ям `gameboard`.

# --hints--

Ви повинні оголосити змінну `gameboard` поза межами класу `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

Значення вашої змінної `gameboard` має бути `Board()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
