---
id: 66069b992c1c5e3451f3deb0
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Інстанціювання створює порожній об’єкт. Метод `__init__` — це спеціальний метод, який дозволяє інстанціювати об’єкт у кастомному стані. Коли клас реалізує метод `__init__`, `__init__` автоматично викликається під час інстанціювання.

У вашому класі `Board` видаліть метод `spam` і замініть його на метод `__init__`, який містить параметр `self`.

# --hints--

У вашому класі `spam` не має бути методу `Board`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

Ви маєте визначити метод `__init__` у вашому класі `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

Ваш метод `__init__` має мати параметр `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
