---
id: 6606927d010be4300a4e5330
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Інстанціювання створює порожній об’єкт. Але класи можуть мати методи, які є локальними функціями для кожного екземпляра. Усередині класу методи оголошують так:

```py
class ClassName:
    def method_name():
        pass
```

У класі `Board` замініть `pass` на порожній метод `spam`.

# --hints--

Ви маєте визначити метод з ім’ям `spam` всередині класу `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
