---
id: 6606d8c323d6205890fbbd54
title: Крок 46
challengeType: 20
dashedName: step-46
---

# --description--

У класі `Board` створіть ще один метод `is_valid` і надайте йому три параметри: `self`, `empty` та `num`. Де `empty` — це кортеж, що представляє індекси рядка та стовпця порожньої клітинки, а `num` — число, яке потрібно перевірити.

Цей метод перевірятиме, чи є задане число допустимим вибором для порожньої клітинки на дошці судоку, перевіряючи його сумісність із рядком, стовпцем і квадратом 3x3, що відповідає вказаній порожній клітинці.

# --hints--

Ви повинні створити новий метод із назвою `is_valid` у класі `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("is_valid")`)) })
```

Ваш метод `is_valid` має мати три параметри: `self`, `empty` та `num`, у цьому порядку.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("is_valid").has_args("self, empty, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start = (col // 3) * 3
        for row_no in range(row_start, row_start + 3):
            for col_no in range(col_start, col_start + 3):
                if self.board[row_no][col_no] == num:
                    return False
        return True
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
