---
id: 6606933d6813a8308c962dd1
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Щоб бути методом екземпляра, метод потребує спеціального параметра, який за конвенцією називається `self`. Цей параметр є референсом на екземпляр класу і завжди має бути першим параметром.

Додайте параметр `self` до вашого методу `spam`.

# --hints--

Ваш метод `spam` має мати параметр `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
