---
id: 66068fb0bfddba2b7977eb60
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви дізнаєтеся про класи та об’єкти, створюючи розв’язувач головоломки судоку.

У Python клас — це шаблон для створення об’єктів. Об’єкти, створені з класу, є екземплярами цього класу. Ви можете створити клас за допомогою такого синтаксису:

```py
class ClassName:
    pass
```

Де `class` — ключове слово, необхідне для визначення класу, а `ClassName` — ім’я класу, яке за конвенцією пишеться у *PascalCase*.

Почніть зі створення класу `Board`.

# --hints--

Ви повинні створити клас із назвою `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
