---
id: 6606a4641ec48b3a9fe8c2fc
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Вбудована функція `enumerate` приймає ітерований об’єкт як аргумент і повертає об’єкт enumerate, по якому ви можете ітеруватися.
Вона надає лічильник (який за замовчуванням починається з нуля) і значення з ітерованого об’єкта.

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

Цикл із прикладу вище виведе кортежі `0, a`, `1, b` та `2, c`.

У методі `find_empty_cell` замініть `pass` на цикл `for`, який використовує функцію `enumerate()` для ітерування над кожним рядком у дошці судоку. Використовуйте `row` як індекс поточного рядка і `contents` для елементів поточного рядка.


# --hints--

Ви повинні замінити `pass` на цикл `for`, який ітерує над `enumerate(self.board)`. Використовуйте `row` та `contents` як змінні циклу.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
