---
id: 66069f86f58f85371d47123e
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Атрибут — це змінна, пов’язана з об’єктом, яка використовується для збереження даних як звичайні змінні.

Всередині методу `__init__` призначте параметр `board` (який передається під час створення екземпляра класу `Board`) атрибуту екземпляра `board` за допомогою `self.board`.

`self.board` посилається на атрибут `board` екземпляра класу. Це змінна, що належить об’єкту, створеному з класу `Board`.

# --hints--

Ви повинні видалити `pass` і призначити параметр `board` атрибуту `self.board` всередині методу `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
