---
id: 61696ef7ac756c829f9e4048
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

في هذا المشروع ستعمل مع ملف CSS خارجي لتنسيق الصفحة. لقد تم إنشاء ملف `styles.css` لك مسبقًا. لكن قبل أن تستخدمه، يجب ربطه بالصفحة.

قم بتسكين عنصر `link` داخل عنصر `head`. أعطه خاصية `rel` بقيمة `"stylesheet"` وخاصية `href` بقيمة `"styles.css"`.

# --hints--

يجب أن يكون لديك عنصر `link` واحد.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

يجب أن يكون عنصر `link` عنصرًا فارغًا، فلا يجب أن يحتوي على وسم إغلاق `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بقيمة `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

يجب أن يكون عنصر `link` داخل عنصر `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
