---
id: 616968c2c94c8071b349c146
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

داخل `body`، قم بتسكين عنصر `h1` يحتوي على النص `CSS Color Markers`.

# --hints--

يجب أن تحتوي شفرتك على عنصر `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

يجب أن يكون لديك وسم افتتاحي `h1`.

```js
assert.match(code, /<h1\s*>/i);
```

يجب أن يكون لديك وسم إغلاق `h1`.

```js
assert.match(code, /<\/h1\s*>/i);
```

يجب أن يكون عنصر `h1` داخل عنصر `body`.

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

يجب أن يحتوي عنصر `h1` على النص `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

تذكر، أن حالة الأحرف والإملاء مهمان لنص `h1`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
