---
id: 616d595270d902f0e7086e18
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

رغم وجود ثلاثة عناصر `div` منفصلة للعلامة، إلا أنها تبدو كمستطيل كبير واحد. أضف بعض المسافة بينها لتسهيل رؤية كل عنصر.

عندما تحتوي خاصية `margin` المختصرة على قيمتين، فإنها تضبط `margin-top` و`margin-bottom` على القيمة الأولى، و`margin-left` و`margin-right` على القيمة الثانية.

في قاعدة CSS الخاصة بـ `.marker`، اضبط خاصية `margin` على `10px auto`.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.marker` على خاصية `margin` مضبوطة على `10px auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

```
