---
id: 616d55bd160a95e22453a081
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

الآن بعد أن أصبح لديك مؤشر واحد في المنتصف مع لون، حان الوقت لإضافة المؤشرات الأخرى.

في عنصر `.container`، أضف عنصرين `div` آخرين وأعط كلًا منهما صنفًا `marker`.

# --hints--

يجب أن يحتوي أول عنصر `div` جديد على وسم البداية.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][2]);
```

يجب أن يحتوي أول عنصر `div` جديد على وسم الإغلاق.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][2]);
```

يجب أن يحتوي ثاني عنصر `div` جديد على وسم البداية.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][3]);
```

يجب أن يحتوي ثاني عنصر `div` جديد على وسم الإغلاق.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][3]);
```

يجب أن تكون عناصر `div` الجديدة داخل `div` الذي يحمل الصنف `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
containerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(containerChildren, 3);
```

يجب أن يحمل كلا عنصرَي `div` الجديدين صنفًا `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert.isNotEmpty(containerChildren);
containerChildren.forEach(child => assert.isTrue(child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
