---
id: 616d525007b8c5d8b3308b1c
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

سيبدو مؤشر موقعك أفضل إذا كان في وسط الصفحة. طريقة سهلة للقيام بذلك هي استخدام الخاصية المختصرة `margin`.

يمكنك تعيين مساحة الهامش للعناصر بشكل منفصل باستخدام خصائص مثل `margin-top` و`margin-left`، أما الخاصية المختصرة `margin` فتسهل تعيين عدة مساحات هامش في نفس الوقت.

لتمركز المؤشر في الصفحة، عيّن خاصية `margin` إلى `auto`. هذا يضبط `margin-top` و`margin-right` و`margin-bottom` و`margin-left` كلها إلى `auto`.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.marker` على خاصية `margin` مضبوطة على `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
