---
id: 616d3a67ccf800ad94ec89ae
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

الآن أضف بعض العناصر التي ستُنسق لاحقًا لتصبح علامات لونية.

أولًا، داخل عنصر `body`، أضف عنصر `div` مع خاصية `class` مضبوطة على `container`. تأكد من وضع عنصر `div` أسفل عنصر `h1`.

# --hints--

يجب أن يكون لديك وسم افتتاحي `div`.

```js
assert.match(code, /<div.*>/i);
```

يجب أن يكون لديك وسم إغلاق `div`.

```js
assert.match(code, /<\/div\s*>/i);
```

يجب أن يكون عنصر `div` داخل عنصر `body` الخاص بك.

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

يجب أن يكون عنصر `div` بعد عنصر `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

يجب أن يحتوي عنصر `div` على خاصية `class` مضبوطة على `container`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}
```
