---
id: 61764fdda535587e1fefb3aa
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

أضف الصنف `two` إلى عنصر `div` الثاني للعلامة، والصنف `three` إلى عنصر `div` الثالث للعلامة.

# --hints--

يجب إضافة الصنف `two` إلى عنصر `div` الثاني للعلامة داخل عنصر `.container`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('two'));
```

يجب أن يحتوي عنصر `div` الثاني للعلامة على الأصناف `marker` و `two`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('marker'));
assert.isTrue(containerSecondChild.classList?.contains('two'));
```

يجب إضافة الصنف `three` إلى عنصر `div` الثالث للعلامة داخل عنصر `.container`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('three'));
```

يجب أن يحتوي عنصر `div` الثالث للعلامة على الأصناف `marker` و `three`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('marker'));
assert.isTrue(containerThirdChild?.classList?.contains('three'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

```
