---
id: 617b92b9de349513466f6156
title: الخطوة 19
challengeType: 0
dashedName: step-19
---

# --description--

بينما تبدو العلامات الحمراء والزرقاء متشابهة، فإن العلامة الخضراء أخف بكثير مما كانت عليه سابقًا. ذلك لأن كلمة اللون المفتاحية `green` في الواقع درجة أغمق، وتقع تقريبًا في منتصف المسافة بين الأسود وأقصى قيمة للون الأخضر.

في قاعدة CSS الخاصة بـ `.two`، عيّن قيمة اللون الأخضر في دالة `rgb` إلى `127` لتقليل شدته.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.two` على خاصية `background-color` مضبوطة على `rgb(0, 127, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
