---
id: 617b8e0d93a8d10d9a90e720
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

لاحظ أن `background-color` للعلامة الخاصة بك لا تزال حمراء. هذا لأنك ضبطت قيمة الأحمر في دالة `rgb` إلى الحد الأقصى `255`، أي 100% أحمر، وضبطت كل من القيم الخضراء والزرقاء إلى `0`.

الآن استخدم دالة `rgb` لضبط الألوان الأخرى.

في قاعدة CSS `.two`، استخدم دالة `rgb` لضبط `background-color` إلى القيمة القصوى للأخضر، و`0` للقيم الأخرى. وفي قاعدة CSS `.three`، استخدم دالة `rgb` لضبط `background-color` إلى القيمة القصوى للأزرق، و`0` للقيم الأخرى.

# --hints--

يجب ألا تستخدم قاعدة CSS `.two` كلمة لون مفتاحية `green` لضبط خاصية `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

يجب أن تحتوي قاعدة CSS `.two` على خاصية `background-color` مضبوطة إلى `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

يجب ألا تستخدم قاعدة CSS `.three` كلمة لون مفتاحية `blue` لضبط خاصية `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

يجب أن تحتوي قاعدة CSS `.three` على خاصية `background-color` مضبوطة إلى `rgb(0, 0, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
