---
id: 617bd6ec666b1da2587e4e37
title: الخطوة 29
challengeType: 0
dashedName: step-29
---

# --description--

هناك ثلاثة ألوان ثالثية أخرى: الأخضر الشارتروز (الأخضر + الأصفر)، الأزرق السماوي (الأزرق + السماوي)، والوردي (الأحمر + الماجنتا).

لإنشاء الأخضر الشارتروز، حدّث دالة `rgb` في قاعدة CSS `.one` بحيث يكون الأحمر عند `127`، واضبط الأخضر على القيمة القصوى.

بالنسبة للأزرق السماوي، حدّث دالة `rgb` في قاعدة CSS `.two` بحيث يكون الأخضر عند `127` والأزرق عند القيمة القصوى.

أما الوردي، الذي يُسمى أحيانًا الوردي الفاقع، فحدّث دالة `rgb` في قاعدة CSS `.three` بحيث يكون الأزرق عند `127` والأحمر عند القيمة القصوى.

# --hints--

يجب أن تحتوي قاعدة CSS `.one` على خاصية `background-color` مضبوطة على `rgb(127, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(127, 255, 0)');
```

يجب أن تحتوي قاعدة CSS `.two` على خاصية `background-color` مضبوطة على `rgb(0, 127, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 255)');
```

يجب أن تحتوي قاعدة CSS `.three` على خاصية `background-color` مضبوطة على `rgb(255, 0, 127)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
