---
id: 617b8b38f32bf2080a140675
title: الخطوة 17
challengeType: 0
dashedName: step-17
---

# --description--

كما تعلمت في الدروس السابقة، الدالة هي قطعة شفرة يمكنها أخذ مُدخَل وتنفيذ إجراء محدد. دالة CSS `rgb` تقبل قيمًا، أو <dfn>معلمات</dfn>، للأحمر والأخضر والأزرق، وتنتج لونًا:

```css
rgb(red, green, blue);
```

كل قيمة من الأحمر والأخضر والأزرق هي رقم من `0` إلى `255`. تعني `0` أن نسبة ذلك اللون 0%، وهو اللون الأسود. تعني `255` أن نسبة ذلك اللون 100%.

في قاعدة CSS `.one`، استبدل كلمة اللون المفتاحية `red` بدالة `rgb`. في دالة `rgb`، عيّن قيمة الأحمر إلى `255`، وقيمة الأخضر إلى `0`، وقيمة الأزرق إلى `0`.

# --hints--

يجب ألا تستخدم قاعدة CSS `.one` كلمة اللون المفتاحية `red` لتعيين خاصية `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'red');
```

يجب أن تحتوي قاعدة CSS `.one` على خاصية `background-color` مضبوطة على `rgb(255, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: red;
--fcc-editable-region--
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
