---
id: 617b954d9f4f6217a749380e
title: الخطوة 20
challengeType: 0
dashedName: step-20
---

# --description--

أضف الآن مساحة رأسية أكثر قليلاً بين العلامات وحافة عنصر `container` الذي توجد فيه.

في قاعدة CSS الخاصة بـ `.container`، استخدم الخاصية المختصرة `padding` لإضافة حشو بمقدار `10px` في الأعلى والأسفل، واضبط الحشو الأيسر والأيمن على `0`. هذا يعمل بطريقة مشابهة للخاصية المختصرة `margin` التي استخدمتها سابقًا.

# --hints--

لا يجب إزالة خاصية `background-color` وقيمتها من قاعدة CSS الخاصة بـ `.container`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

يجب أن تحتوي قاعدة CSS الخاصة بـ `.container` على خاصية `padding` مضبوطة على `10px 0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.padding, '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
--fcc-editable-region--

--fcc-editable-region--
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
