---
id: 618a132676346ac9f7fd59dd
title: الخطوة 34
challengeType: 0
dashedName: step-34
---

# --description--

هل لاحظت كيف تنجذب عيناك طبيعيًا إلى اللون الأحمر في الوسط؟ عند تصميم موقع، يمكنك استخدام هذا التأثير لجذب الانتباه إلى العناوين المهمة أو الأزرار أو الروابط.

هناك عدة تركيبات ألوان مهمة أخرى بخلاف الألوان التكاملية، لكنك ستتعلمها لاحقًا قليلاً.

في الوقت الحالي، استخدم الدالة `rgb` في قاعدة CSS `.two` لتعيين خاصية `background-color` إلى اللون الأسود.

# --hints--

يجب أن تحتوي قاعدة CSS `.two` على خاصية `background-color` مضبوطة على `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 0);
}

```
