---
id: 619b761916dac02643940022
title: الخطوة 41
challengeType: 0
dashedName: step-41
---

# --description--

طريقة شائعة جدًا لتطبيق اللون على عنصر باستخدام CSS هي باستخدام القيم <dfn>السداسية عشرية</dfn> أو القيم hex. رغم أن القيم hex تبدو معقدة، فهي في الواقع مجرد شكل آخر من قيم RGB.

تبدأ قيم اللون السداسية عشرية برمز `#` وتتكون من ستة أحرف من 0-9 و A-F. الزوج الأول من الأحرف يمثل الأحمر، والزوج الثاني يمثل الأخضر، والزوج الثالث يمثل الأزرق. على سبيل المثال، `#4B5320`.

في محدد الصنف `.green`، عرّف الخاصية `background-color` إلى رمز لون سداسي عشري بقيم `00` للأحمر، و`FF` للأخضر، و`00` للأزرق.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.green` على خاصية `background-color` مضبوطة على `#00FF00`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

.blue {
  background-color: rgb(0, 0, 0);
}

```
