---
id: 619b7c3c83de403126b69c1e
title: الخطوة 42
challengeType: 0
dashedName: step-42
---

# --description--

قد تكون بالفعل على دراية بالقيم العشرية، أو القيم ذات الأساس 10، التي تتراوح من 0 إلى 9. القيم الست عشرية، أو ذات الأساس 16، تتراوح من 0 إلى 9، ثم من A إلى F:

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

في ألوان hex، تمثل `00` نسبة 0% من ذلك اللون، وتمثل `FF` نسبة 100%. لذا فإن `#00FF00` تعني 0% أحمر، 100% أخضر، و0% أزرق، وهي نفسها `rgb(0, 255, 0)`.

اخفض شدة اللون الأخضر بتعيين قيمة اللون الأخضر في اللون الست عشري إلى `7F`.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.green` على خاصية `background-color` مضبوطة على `#007F00`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
--fcc-editable-region--
  background-color: #00FF00;
--fcc-editable-region--
}

.blue {
  background-color: rgb(0, 0, 0);
}

```
