---
id: 61a49d15bdbb5e57cc6fd280
title: الخطوة 49
challengeType: 0
dashedName: step-49
---

# --description--

تسمح لك توقفات اللون بضبط دقيق لمواقع الألوان على طول خط التدرج. تُستخدم في دالة `linear-gradient` لتحديد مكان انتهاء لون معين وبداية الانتقال إلى اللون التالي. يمكن استخدام نسبة مئوية أو قيمة طول لتحديد موقع نقطة توقف اللون.

على سبيل المثال، في هذا التدرج من الأحمر إلى الأسود، يحدث الانتقال من الأحمر إلى الأسود عند نقطة 90% على طول خط التدرج، لذا يشغل اللون الأحمر معظم المساحة المتاحة:

```css
linear-gradient(90deg, red 90%, black);
```

في دالة `linear-gradient`، أضف نقطة توقف لون `75%` بعد أول معلمة لون أحمر. لا تضف نقاط توقف للألوان الأخرى.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.red` على خاصية `background` مضبوطة على `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
