---
id: 61a5d594b887335228ee6533
title: الخطوة 62
challengeType: 0
dashedName: step-62
---

# --description--

استخدم الدالة `linear-gradient`، ومرر الدالة `hsl` مع القيم `186` للون، و`76%` للتشبع، و`16%` لـفاتح اللون كأول معلمة للون.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.blue` على خاصية `background` بالقيمة `linear-gradient(hsl(186, 76%, 16%))`.

```js
assert.match(
  code,
  /\.blue\s*\{\s*background:\s*linear-gradient\(hsl\(\s*186(,\s*|\s+)76%(?:\1)16%\s*\)\)/
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
--fcc-editable-region--
  background: hsl(240, 100%, 50%);
--fcc-editable-region--
}

```
