---
id: 61addaf7e83988b59a7d18ff
title: الخطوة 77
challengeType: 0
dashedName: step-77
---

# --description--

تسمح لك الخاصية المختصرة `border-left` بتعيين عرض وأسلوب ولون الحد الأيسر في نفس الوقت.

إليك الصيغة النحوية:

```css
border-left: width style color;
```

في قاعدة CSS الخاصة بـ `.sleeve`، استبدل الخصائص `border-left-width` و`border-left-style` و`border-left-color` بالخاصية المختصرة `border-left`. يجب أن تكون القيم الخاصة بعرض وأسلوب ولون الحد الأيسر متطابقة.

# --hints--

يجب ألا تحتوي قاعدة CSS الخاصة بـ `.sleeve` على خاصية وقيمة `border-left-width`.

```js
assert.notMatch(code, /border-left-width/g);
```

يجب ألا تحتوي قاعدة CSS الخاصة بـ `.sleeve` على خاصية وقيمة `border-left-style`.

```js
assert.notMatch(code, /border-left-style/g);
```

يجب ألا تحتوي قاعدة CSS الخاصة بـ `.sleeve` على خاصية وقيمة `border-left-color`.

```js
assert.notMatch(code, /border-left-color/);
```

يجب أن تحتوي قاعدة CSS الخاصة بـ `.sleeve` على الخاصية المختصرة `border-left` بقيمة `10px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.borderLeft, '10px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
--fcc-editable-region--
  border-left-width: 10px;
  border-left-style: solid;
  border-left-color: black;
--fcc-editable-region--
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
