---
id: 61b0a44a6b865738ba49b9fb
title: الخطوة 81
challengeType: 0
dashedName: step-81
---

# --description--

آخر شيء ستفعله هو إضافة ظل خفيف لكل علامة لجعلها تبدو أكثر واقعية.

تسمح لك خاصية `box-shadow` بتطبيق ظل واحد أو أكثر حول عنصر. هذه هي الصيغة النحوية الأساسية:

```css
box-shadow: offsetX offsetY color;
```

إليك كيف تعمل قيم `offsetX` و `offsetY`:

* كل من `offsetX` و `offsetY` يقبلان قيم رقمية بوحدة `px` ووحدات CSS أخرى
* القيمة الموجبة لـ `offsetX` تحرك الظل إلى اليمين والقيمة السالبة تحركه إلى اليسار
* القيمة الموجبة لـ `offsetY` تحرك الظل إلى الأسفل والقيمة السالبة تحركه إلى الأعلى
* إذا أردت قيمة صفر (`0`) لأي من `offsetX` أو `offsetY` أو كلاهما، لا تحتاج إلى إضافة وحدة. كل المتصفحات تفهم أن الصفر يعني عدم التغيير.

يحدد ارتفاع وعرض الظل ارتفاع وعرض العنصر الذي يُطبق عليه. يمكنك أيضًا استخدام قيمة اختيارية `spreadRadius` لتوسيع مدى انتشار الظل. سنتحدث عن ذلك لاحقًا.

ابدأ بإضافة ظل بسيط للعلامة الحمراء.

في قاعدة CSS الخاصة بـ `.red`، أضف خاصية `box-shadow` بالقيم `5px` لـ `offsetX`، و`5px` لـ `offsetY`، و`red` لـ `color`.

# --hints--

يجب أن تحتوي قاعدة CSS الخاصة بـ `.red` على خاصية مختصرة `box-shadow` بالقيمة `5px 5px red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow, 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
--fcc-editable-region--
  
--fcc-editable-region--
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
