---
id: 587d8247367417b2b2512c36
title: Instalar e solicitar o Helmet
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

Trabalhar nesses desafios vai fazer com que você escreva seu código usando um dos seguintes métodos:

- Clone <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">este repositório do GitHub</a> e complete esses desafios localmente.
- Use um construtor de sites de sua escolha para completar o projeto. Certifique-se de incorporar todos os arquivos do nosso repositório no GitHub.

O Helmet ajuda você a proteger seus aplicativos do Express, definindo vários cabeçalhos HTTP.

# --instructions--

Todo o código para estas aulas vai para o arquivo `myApp.js` entre as linhas de código que fornecemos para você começar. Não altere nem exclua o código que adicionamos para você.

A versão `3.21.3` do Helmet já foi instalada, então solicite-a como `helmet` no `myApp.js`.

# --hints--

A versão do `helmet` `3.21.3` deve estar no `package.json`

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```

