---
id: 587d8247367417b2b2512c37
title: Esconder informações potencialmente perigosas usando o helmet.hidePoweredBy()
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

Como lembrete, este projeto está sendo construído com base no seguinte projeto inicial clonado do <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Hackers podem explorar vulnerabilidades conhecidas no Express/Node se perceberem que seu site é alimentado pelo Express. `X-Powered-By: Express` é enviado em toda requisição proveniente do Express por padrão. Use o middleware `helmet.hidePoweredBy()` para remover o cabeçalho X-Powered-By.

# --hints--

O middleware helmet.hidePoweredBy() deve ser montado corretamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```

