---
id: 68c18b1d0fce081f85312d97
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Em uma lição anterior, você aprendeu que um atributo prefixado com um único underscore é destinado ao uso interno por convenção.

Modifique os atributos `name` e `level` para `_name` e `_level`, já que eles não devem ser modificados de fora da sua classe.

Observe que isso não impede que o atributo seja acessado ou modificado fora da classe. Além disso, em Python sempre há uma maneira de acessar atributos privados (prefixados com um duplo sublinhado) também.

# --hints--

Seu método `__init__` deve definir `self._name` como `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

Seu método `__init__` deve definir `self._level` como `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
