---
id: 68c15a2a1f0e70ca7154aba5
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste workshop, você vai construir um sistema de acompanhamento de salário para funcionários.

Comece criando uma classe chamada `Employee`. Dentro dela crie um método `__init__` com os parâmetros `self`, `name` e `level`. Dentro do método `__init__`, atribua `name` e `level` aos atributos da instância com o mesmo nome.

# --hints--

Você deve ter uma classe chamada `Employee`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

Sua classe `Employee` deve ter um método `__init__` com três parâmetros `self`, `name` e `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

Seu método `__init__` deve definir `self.name` como `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Seu método `__init__` deve definir `self.level` como `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
