---
id: 68c16924498728d259e4781c
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Agora crie uma instância da classe `Employee` passando as strings `Charlie Brown` e `trainee`. Atribua a instância a uma variável chamada `charlie_brown`.

# --hints--

Você deve ter uma variável chamada `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

Você deve atribuir `Employee('Charlie Brown', 'trainee')` a `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level


--fcc-editable-region--
```
