---
id: 68c281cecda94ba2c6838e17
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Agora que você definiu um getter para `name`, você pode acessar o atributo `_name` através da propriedade `name`. Então imprima `charlie_brown.name` no console. 

# --hints--

Você deve imprimir `charlie_brown.name` no console.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
