---
id: 68c27fa1697ceb7c578b2c9f
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

O decorador `@property` é usado em Python para transformar um método em uma propriedade. Ele é tipicamente usado para definir métodos getter, que são métodos usados para recuperar o valor de um atributo:

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

Crie um método chamado `name` com um parâmetro `self` e decore-o com `@property`. Dentro do método, retorne `self._name`.
    
# --hints--

Sua classe `Employee` deve ter um método `name` com um parâmetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

Seu método `name` deve ser decorado com `@property`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

Seu método `name` deve retornar `self._name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
