---
id: 68c28d659d345270d4a61084
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

O método `__repr__` é um método especial que deve retornar uma representação em string do objeto que pode ser usada para instanciá-lo.

Por exemplo, o método `__repr__` de `Employee('Charlie', 'developer')` deve retornar a string `Employee('Charlie', 'developer')`, que é a mesma string usada para criar o objeto.

Dê à sua classe `Employee` um método `__repr__` com um parâmetro `self` e faça com que ele retorne uma string que possa ser usada para instanciar o objeto.

# --hints--

Sua classe `Employee` deve ter um método `__repr__` com um parâmetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

Seu método `__repr__` deve retornar a string usada para instanciar o objeto.

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
