---
id: 68c28c6517e7de5e85f8ebe5
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Agora que você tem getters para os atributos `_name` e `_level`, atualize a string retornada por `__str__` para usar `self.name` e `self.level`. Isso chamará os getters em vez de acessar diretamente os atributos.

# --hints--

Seu método `__str__` deve retornar `f'{self.name}: {self.level}'`

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_return("f'{self.name}: {self.level}'")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)
print(charlie_brown.level)
```
