---
id: 68c27eb292bfec6b9402d58f
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Adicione um método `__str__` à classe `Employee`. Faça-o retornar uma f-string com o formato `name: level`, substituindo `name` e `level` pelos atributos correspondentes.

Depois disso, imprima `charlie_brown` no console.

# --hints--

Sua classe `Employee` deve ter um método `__str__` com um parâmetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

Seu método `__str__` deve retornar uma f-string com o formato `name: level`, substituindo `name` e `level` pelos valores de `self._name` e `self._level`, respectivamente.

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

Você deve imprimir `charlie_brown` no console.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--
```
