---
id: 68c2d31a86d5672ca49d4521
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Agora é hora de adicionar alguma validação ao método `__init__`. No início do método, crie uma declaração `if` que verifique se `name` ou `level` não são instâncias de `str`.

Dentro da declaração `if`, levante um `TypeError` com a mensagem `'name' e 'level' attribute must be of type 'str'.`.

# --hints--

Você deve ter uma declaração `if` no seu método `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0]`)) })
```

Sua declaração `if` deve verificar se `name` e `level` não são instâncias de `str`.

```js
({ test: () => runPython(`
  cond = _Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_conditions()[0]
  conditions = [
    "not (isinstance(name, str) and isinstance(level, str))",
    "not isinstance(name, str) or not isinstance(level, str)",
    "not (isinstance(level, str) and isinstance(name, str))",
    "not isinstance(level, str) or not isinstance(name, str)",
  ]
  assert any(cond.is_equivalent(c) for c in conditions)
`) })
```

Sua declaração `if` deve gerar um `TypeError` com a mensagem `'name' e 'level' attribute must be of type 'str'.`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError("\\'name\\' and \\'level\\' attribute must be of type \\'str\\'.")')`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
--fcc-editable-region--
    def __init__(self, name, level):
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
