---
id: 68c28b7cc3ab8f4daf144b5e
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Seguindo o que você fez nas etapas anteriores, crie um getter para o atributo `_level`.

# --hints--

Sua classe `Employee` deve ter o método `level` com um parâmetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_args("self")`)) })
```

Seu método `level` deve ser decorado com `@property`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_decorators("property")`)) })
```

Seu método `level` deve retornar `self._level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_return("self._level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

```
