---
id: 68c7f7012a700243eff1cbc0
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

O nível usado para instanciar um funcionário deve ser escolhido entre níveis específicos. Você usará `_base_salaries` para validar o nível e definir o salário correto para o funcionário.

Após sua declaração `if` existente, crie outro `if` que verifica se `level` não está em `Employee._base_salaries`.

Dentro da nova declaração `if`, levante um `ValueError` com a mensagem `Invalid value '{level}' for 'level' attribute.`, onde `{level}` deve ser substituído pelo valor do argumento `level`.

# --hints--

Quando `level` não estiver em `Employee._base_salaries`, você deve lançar um `ValueError` com a mensagem `Invalid value '{level}' for 'level' attribute.`, onde `{level}` deve ser substituído pelo valor do argumento `level`.

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

Você não deve lançar nenhuma exceção quando `level` estiver em `Employee._base_salaries`.

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
