---
id: 68c8002beb157fe47d2e7699
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Um setter é um método usado para definir o valor de um atributo, permitindo verificações de validação e restrições. Você pode criar um setter usando o decorador `@propertyName.setter`, onde `propertyName` deve corresponder ao nome da propriedade a ser definida:

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, value):
        self._name = value

p = Person('Alice')
p.name = 'Abigail' # Calls the setter
print(p.name) # Abigail
```

Após seu método getter, crie um método `name` com os parâmetros `self` e `new_name`. Decore o método com `@name.setter`. Dentro do método, defina `self._name` como `new_name`.

# --hints--

Sua classe `Employee` deve ter um método `name` com os parâmetros `self` e `new_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_args("self, new_name")`)) })
```

Seu método `name` deve ser decorado com `@name.setter`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_decorators("name.setter")`)) })
```

Seu método `name` deve definir `self._name` como `new_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_stmt("self._name = new_name")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary


charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
