---
id: 68c9c3aa714bc326e026b826
title: Passo 26
challengeType: 20
dashedName: step-26
---

# --description--

O novo nível não pode ser definido sem verificar se é um nível válido. No início do seu setter, crie uma declaração `if` que levante um `ValueError` quando `new_level` não for uma chave de `Employee._base_salaries`.

Para a mensagem de erro, use `Invalid value '{new_level}' for 'level' attribute.`, onde `{new_level}` deve ser substituído pelo argumento passado para o setter.

# --hints--

Você deve ter uma declaração `if` dentro do seu setter `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[0]`)) })
```

Quando `new_level` não for uma chave de `Employee._base_salaries`, você deve lançar um `ValueError` com a mensagem `Invalid value '{new_level}' for 'level' attribute.`, onde `{new_level}` deve ser substituído pelo argumento passado para o setter.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "dreamer"
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

Você não deve lançar nenhuma exceção quando `new_level` for uma chave de `Employee._base_salaries`

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  for i in Employee._base_salaries:
    try:
      emp.level = i
    except Exception:
      assert False, "Expected not to raise any exception with valid new_level"    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level
--fcc-editable-region--
    @level.setter
    def level(self, new_level):
        
        self._level = new_level
    
--fcc-editable-region--
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
