---
id: 68dbc6b7e8f794a2ae69bc74
title: Algorithmus für das N-Damen-Problem implementieren
challengeType: 27
dashedName: implement-the-n-queens-algorithm
---

# --description--

Das N-Damen-Problem verlangt, dass Sie N Damen auf einem N×N-Schachbrett so platzieren, dass sich keine zwei Damen gegenseitig angreifen (keine zwei teilen sich eine Reihe, Spalte oder Diagonale).

Zum Beispiel, wenn es ein 4x4-Brett gibt, ist eine gültige Anordnung:

```md
[1, 3, 0, 2]
```

Das bedeutet, dass in Reihe 0 die Dame in Spalte 1 steht; in Reihe 1 die Dame in Spalte 3; in Reihe 2 die Dame in Spalte 0; und in Reihe 3 die Dame in Spalte 2.

Visuell sieht diese Anordnung so aus:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Dabei steht `Q` für eine Dame und `.` für ein leeres Feld.

In diesem Labor implementieren Sie den Lösungsalgorithmus für das N-Damen-Problem mit dem Tiefensuche-Ansatz.


**Ziel**: Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `dfs_n_queens` haben.
2. Die Funktion sollte genau ein Argument akzeptieren: eine ganze Zahl `n`.
3. Wenn `n` kleiner als `1` ist, sollte die Funktion eine leere Liste (`[]`) zurückgeben.
4. Die Funktion sollte eine Liste von Lösungen zurückgeben; jede Lösung ist selbst eine Liste der Länge `n`, wobei das Element am Index `i` den Spaltenindex (0-basiert) der Dame in der Reihe `i` darstellt.

# --hints--

Sie sollten eine Funktion mit dem Namen `dfs_n_queens` haben, die ein Argument annimmt.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs_n_queens)
  sig = inspect.signature(dfs_n_queens)
  assert len(sig.parameters) == 1
`) })
```

Wenn `n` kleiner als `1` ist, sollte die Funktion eine leere Liste zurückgeben.

```js
({ test: () => runPython(`
  assert dfs_n_queens(0) == []
  assert dfs_n_queens(-1) == []
  assert dfs_n_queens(-5) == []
`) })
```

Die Funktion sollte eine Liste von Lösungen zurückgeben, wobei jede Lösung eine Liste der Länge `n` ist.

```js
({ test: () => runPython(`
  result = dfs_n_queens(4)
  assert isinstance(result, list)
  for solution in result:
    assert isinstance(solution, list)
    assert len(solution) == 4
`) })
```

`dfs_n_queens(1)` sollte `[[0]]` zurückgeben.

```js
({ test: () => runPython(`
  assert dfs_n_queens(1) == [[0]]
`) })
```

`dfs_n_queens(2)` sollte `[]` zurückgeben.

```js
({ test: () => runPython(`
  assert dfs_n_queens(2) == []
`) })
```

`dfs_n_queens(3)` sollte `[]` zurückgeben.

```js
({ test: () => runPython(`
  assert dfs_n_queens(3) == []
`) })
```

`dfs_n_queens(4)` sollte `[[1, 3, 0, 2], [2, 0, 3, 1]]` zurückgeben.

```js
({ test: () => runPython(`
  assert dfs_n_queens(4) == [[1, 3, 0, 2], [2, 0, 3, 1]]
`) })
```

`dfs_n_queens(5)` sollte `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]` zurückgeben.

```js
({ test: () => runPython(`
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`len(dfs_n_queens(5))` sollte `10` sein.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(5)) == 10
`) })
```

`len(dfs_n_queens(8))` sollte `92` sein.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(8)) == 92
`) })
```

`dfs_n_queens` sollte das korrekte Ergebnis zurückgeben.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(7)) == 40
  assert dfs_n_queens(6) == [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs_n_queens(n):
    if n < 1:
        return []

    results = []

    def is_safe(queens, row, col):
        for r, c in enumerate(queens):
            if c == col or abs(row - r) == abs(col - c):
                return False
        return True

    def dfs(queens):
        row = len(queens)
        if row == n:
            results.append(queens[:])
            return
        for col in range(n):
            if is_safe(queens, row, col):
                queens.append(col)
                dfs(queens)
                queens.pop()

    dfs([])
    return results
```
