---
id: 686ccc2c8b967e17ab18d593
title: 이분법 구현하기
challengeType: 27
dashedName: implement-the-bisection-method
---

# --description--

이분법, 또는 이진 탐색법이라고도 하는 이 방법은 이진 탐색을 사용해 실수 함수의 근을 찾습니다. 제곱근이 존재하는 구간을 점점 좁혀가면서 지정한 허용 오차 내의 값으로 수렴할 때까지 작동합니다.

예를 들어, 허용 오차가 `0.01`라면, 이분법은 상한과 하한의 차이가 `0.01` 이하가 될 때까지 구간을 계속 반으로 나눕니다.

이번 실습에서는 이분법을 사용해 숫자의 제곱근을 찾는 함수를 구현하세요.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `square_root_bisection`라는 이름의 함수를 세 개의 매개변수로 정의해야 합니다:
    - 제곱근을 구하려는 숫자
    - 결과에 허용할 오차 범위인 허용 오차. 기본값을 설정해야 합니다.
    - 수행할 최대 반복 횟수. 기본값을 설정해야 합니다.

1. `square_root_bisection` 함수는 다음을 수행해야 합니다:
    - 함수에 전달된 숫자가 음수일 경우 `ValueError` 메시지와 함께 `Square root of negative number is not defined in real numbers`를 발생시켜야 합니다.
    - `0`와 `1`인 숫자에 대해서는 `The square root of [number] is [number]` 메시지를 출력하고 그 숫자 자체를 제곱근으로 반환해야 합니다.
    - 그 외의 양수에 대해서는 `The square root of [square_target] is approximately [root]` 메시지와 함께 근사 제곱근을 출력하고 계산된 근 값을 반환해야 합니다.
    - 허용 오차 조건을 만족하는 값이 없으면 `Failed to converge within [maximum] iterations` 실패 메시지를 출력하고 `None`을 반환해야 합니다.

**참고:** 이번 실습에서는 어떤 모듈도 가져올 수 없습니다.

# --hints--

어떤 모듈도 가져오면 안 됩니다.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_imports()) == 0`)) })
```

`square_root_bisection`라는 이름의 함수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("square_root_bisection")`)) })
```

`square_root_bisection` 함수는 세 개의 매개변수를 가져야 합니다.

```js
({ test: () => runPython(`
    import inspect 
    sig = inspect.signature(square_root_bisection)
    assert len(sig.parameters) == 3
`) })
```

허용 오차와 최대 반복 횟수에 기본값을 설정해야 합니다.

```js
({ test: () => runPython(`
try:
  import inspect 
  sig = inspect.signature(square_root_bisection)
  assert len(sig.parameters) == 3
  square_root_bisection(4)
except TypeError:
  assert False
`) })
```

`square_root_bisection` 함수는 함수에 전달된 숫자가 음수일 때 `ValueError` 메시지와 함께 `Square root of negative number is not defined in real numbers`를 발생시켜야 합니다.

```js
({ test: () => runPython(`
try:
  square_root_bisection(-6)
except ValueError as e:
  assert str(e) == "Square root of negative number is not defined in real numbers"
else:
  assert False
`) })
```

`square_root_bisection(0)`은 `0`을 반환해야 합니다.

```js
({ test: () => runPython(`assert square_root_bisection(0) == 0`) })
```

`square_root_bisection(0)`는 `The square root of 0 is 0`을 출력해야 합니다.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(0)
assert "The square root of 0 is 0" in _out
`) })
```

`square_root_bisection(0.001, 1e-7, 50)`은 `0.03162267660168379`와 `0.031622876601683794` 사이의 숫자를 반환해야 합니다.

```js
({ test: () => runPython(`assert 0.03162267660168379 <= square_root_bisection(0.001, 1e-7, 50) <= 0.031622876601683794`) })
```

`square_root_bisection(0.001, 1e-7, 50)`는 `The square root of 0.001 is approximately X`를 출력해야 하며, 여기서 `X`은 `0.03162267660168379`과 `0.031622876601683794` 사이의 숫자입니다.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.001, 1e-7, 50)

assert 0.03162267660168379 <= _root <= 0.031622876601683794
assert f"The square root of 0.001 is approximately {_root}" in _out
`) })
```

`square_root_bisection(0.25, 1e-7, 50)`는 `0.4999999`과 `0.5000001` 사이의 숫자를 반환해야 합니다.

```js
({ test: () => runPython(`assert 0.4999999 <= square_root_bisection(0.25, 1e-7, 50) <= 0.5000001`) })
```

`square_root_bisection(0.25, 1e-7, 50)`는 `The square root of 0.25 is approximately X`를 출력해야 하며, 여기서 `X`는 `0.4999999`와 `0.5000001` 사이의 숫자입니다.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.25, 1e-7, 50)

assert 0.4999999 <= _root <= 0.5000001
assert f"The square root of 0.25 is approximately {_root}" in _out
`) })
```

`square_root_bisection(1)`은 `1`을 반환해야 합니다.

```js
({ test: () => runPython(`assert square_root_bisection(1) == 1`) })
```

`square_root_bisection(1)`는 `The square root of 1 is 1`를 출력해야 합니다.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(1)
assert "The square root of 1 is 1" in _out
`) })
```

`square_root_bisection(81, 1e-3, 50)`는 `8.999`와 `9.001` 사이의 숫자를 반환해야 합니다.

```js
({ test: () => runPython(`assert 8.999 <= square_root_bisection(81, 1e-3, 50) <= 9.001`) })
```

`square_root_bisection(81, 1e-3, 50)`는 `The square root of 81 is approximately X`를 출력해야 하며, 여기서 `X`은 `8.999`과 `9.001` 사이의 숫자입니다.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(81, 1e-3, 50)

assert 8.999 <= _root <= 9.001
assert f"The square root of 81 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-3, 100)`는 `14.999`과 `15.001` 사이의 숫자를 반환해야 합니다.

```js
({ test: () => runPython(`assert 14.999 <= square_root_bisection(225, 1e-3, 100) <= 15.001`) })
```

`square_root_bisection(225, 1e-3, 100)`는 `The square root of 225 is approximately X`를 출력해야 하며, 여기서 `X`는 `14.999`와 `15.001` 사이의 숫자입니다.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-3, 100)

assert 14.999 <= _root <= 15.001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-5, 100)`는 `14.99999`와 `15.00001` 사이의 숫자를 반환해야 합니다.

```js
({ test: () => runPython(`assert 14.99999 <= square_root_bisection(225, 1e-5, 100) <= 15.00001`) })
```

`square_root_bisection(225, 1e-5, 100)`는 `The square root of 225 is approximately X`를 출력해야 하며, 여기서 `X`는 `14.99999`과 `15.00001` 사이의 숫자입니다.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-5, 100)

assert 14.99999 <= _root <= 15.00001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 100)`는 `14.9999999`와 `15.0000001` 사이의 숫자를 반환해야 합니다.

```js
({ test: () => runPython(`assert 14.9999999 <= square_root_bisection(225, 1e-7, 100) <= 15.0000001`) })
```

`square_root_bisection(225, 1e-7, 100)`는 `The square root of 225 is approximately X`를 출력해야 하며, 여기서 `X`는 `14.9999999`과 `15.0000001` 사이의 숫자입니다.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-7, 100)

assert 14.9999999 <= _root <= 15.0000001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 10)`는 `None`을 반환해야 합니다.

```js
({ test: () => runPython(`assert square_root_bisection(225, 1e-7, 10) is None`) })
```

`square_root_bisection(225, 1e-7, 10)`는 `Failed to converge within 10 iterations`를 출력해야 합니다.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(225, 1e-7, 10)
assert "Failed to converge within 10 iterations" in _out
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = square_target if square_target < 1 else 1
        high = 1 if square_target < 1 else square_target
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if high - low <= tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

        if root is None:
            print(f"Failed to converge within {max_iterations} iterations")

        else:   
            print(f'The square root of {square_target} is approximately {root}')

    return root
```
