---
id: 6582267ef39f5047411c1743
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

다음으로, `print_expenses`라는 이름의 함수를 정의하세요. 이 함수는 하나의 매개변수 `expenses`를 받습니다. 이 함수는 나중에 목록에 있는 각 지출 항목을 표시하는 데 사용됩니다.

새 함수의 본문을 `pass` 문장으로 채우세요.

# --hints--

`print_expenses`라는 이름의 새 함수를 정의해야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

`print_expenses` 함수는 하나의 매개변수 `expenses`를 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
