---
id: 65821c1c40ccac44adf38e6b
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`expenses`라는 빈 리스트를 만드세요. 각 지출을 저장하는 데 사용할 것입니다.

# --hints--

`expenses`라는 새 변수를 선언하세요.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

`expenses` 변수에 빈 리스트를 할당하세요.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
