---
id: 65823cfc74aa564ffc460489
title: 26단계
challengeType: 20
dashedName: step-26
---

# --description--

이제 `map()`를 호출하면서 `lambda` 함수를 첫 번째 인수로, `expenses` 리스트를 두 번째 인수로 전달하세요.

# --hints--

`map()` 함수 안에서 `total_expenses` 함수를 호출해야 합니다.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("total_expenses").find_calls("map")) == 1`)) })
```

`lambda` 호출에 첫 번째 인수로 `map()` 함수를 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[0].is_equivalent("lambda expense: expense['amount']")`)) })
```

`expenses` 호출에 두 번째 인수로 `map()` 리스트를 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
