---
id: 658225d560369446ece5518b
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

`expense` 리스트에 추가하는 사전에 또 다른 키-값 쌍을 추가하세요. 문자열 `'category'`을 키로, `category` 매개변수를 값으로 사용하세요.

# --hints--

`.append()` 메서드에 인수로 전달하는 사전에 두 번째 키-값 쌍을 추가해야 합니다.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

`{'amount': amount, 'category': category}` 메서드에 인수로 `.append()`를 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
