---
id: 6582575b8089f85b8b92d7c8
title: 42단계
challengeType: 20
dashedName: step-42
---

# --description--

비용 세부 정보를 얻으면 `add_expense` 함수를 호출하여 새 비용을 `expenses` 목록에 추가해야 합니다.

`amount`를 사용해 `category`와 `input()`을 얻은 후, 세 개의 인수 `add_expense`, `expenses`, `amount`을 전달하여 `category` 함수를 호출하세요.

- `expenses`는 이 프로젝트의 메인 함수에서 이전에 생성한 빈 목록입니다.
- `amount`는 비용 금액입니다.
- `category`는 비용 카테고리입니다.

# --hints--

`add_expense()` 문장 내에서 `if` 함수를 호출해야 합니다.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

`add_expense()` 호출에 세 개의 인수 `expenses`, `amount`, `category`를 전달해야 합니다. 순서가 중요합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

`add_expense()` 호출은 `amount`와 `category` 할당 이후에 와야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            
--fcc-editable-region--
```
