---
id: 65821fcc010c3245718f2a06
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

`expenses` 함수의 `add_expense` 매개변수는 비용 목록입니다. 목록 끝에 항목을 추가할 수 있어야 합니다. 이를 위해 `.append()` 리스트 메서드를 사용하세요.

`.append()` 리스트에 `expenses` 메서드를 호출하세요. 지금은 `.append()`에 인수를 전달하지 마세요.

# --hints--

함수에 `expenses.append()`가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

함수에 `pass`가 없어야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
