---
id: 65822308ded1ba4632f66aa3
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

딕셔너리는 Python의 또 다른 내장 데이터 유형입니다. 딕셔너리는 *키*-*값* 쌍 형태의 데이터 모음입니다. 딕셔너리는 중괄호(`{}`)로 정의하며, 쉼표로 구분된 키-값 쌍을 포함합니다. 각 키 다음에는 콜론(`:`)이 오고 그 뒤에 값이 옵니다.

```py
{'amount': 50.0, 'category': 'Food'}
```

위 예제에서 `'amount'`와 `'category'`는 키이고, `50.0`와 `'Food'`는 각각 대응하는 값입니다.

키가 `'amount'`이고 값이 `amount` 매개변수인 딕셔너리를 만들고, 새 딕셔너리를 `.append()` 호출에 전달하세요.

# --hints--

딕셔너리를 `.append()` 메서드에 인수로 전달해야 합니다.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

`{'amount': amount}`을 `.append()` 메서드에 인수로 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
