---
id: 65824dfdb6815d563b2d3256
title: 37단계
challengeType: 20
dashedName: step-37
---

# --description--

`input()` 함수는 사용자 입력을 받아 문자열 형태로 반환합니다.

`while` 반복문 안에서 `input()` 함수를 호출하고, 문자열 `'Enter your choice: '` 을 인수로 전달한 뒤 결과를 `choice`라는 변수에 할당하세요.

# --hints--

`choice` 반복문 안에 `while`라는 변수를 선언해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].has_variable("choice")
`)) })
```

문자열 `input()`을 인수로 하여 `'Enter your choice: '` 함수를 호출한 결과를 `choice` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_variable("choice").is_equivalent("choice = input('Enter your choice: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
--fcc-editable-region--
```
