---
id: 65822bd82d708c4895080c35
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

Python에서 중요한 점은 문자열을 정의할 때 사용한 따옴표와 같은 종류의 따옴표는 문자열 내부에서 사용할 수 없다는 것입니다. 예를 들어, 문자열 `'I'm a string!'`은 유효하지 않습니다. 문자열 내부에서 작은따옴표를 사용하려면 다음 중 하나를 해야 합니다:

- 백슬래시를 앞에 붙여서 따옴표를 이스케이프하세요: `'I\'m a string!'`
- 또는 문자열을 정의할 때 큰따옴표를 사용하세요: `"I'm a string!"` (권장).

딕셔너리의 값은 키를 통해 접근할 수 있습니다. 대괄호 표기법을 사용하고 대괄호 사이에 키를 포함해야 합니다:

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

현재 f-string에서 `expense` 딕셔너리를 보간하고 있습니다. f-string 표현식을 수정하여 `'amount'` 딕셔너리에서 `'category'` 키와 `expense` 키의 값을 접근하세요.

# --hints--

`f'Amount: {expense["amount"]}, Category: {expense["category"]}'`를 `print()` 호출에 전달해야 합니다. 작은따옴표로 감싼 f-string 내부에서는 큰따옴표를, 큰따옴표로 감싼 f-string 내부에서는 작은따옴표를 사용하세요.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
