---
id: 65824872894f59546e3084e2
title: 34단계
challengeType: 20
dashedName: step-34
---

# --description--

이전 단계에서 만든 `while` 루프는 무한 루프로, 사용자가 종료를 결정할 때까지 프로그램이 메뉴 옵션을 계속해서 보여주도록 합니다.

`print()` 호출 뒤에 `print()` 호출을 하나 더 추가하여 문자열 `'1. Add an expense'`을 출력하세요.

# --hints--

기존 `'1. Add an expense'` 호출 뒤에 `while` 루프 안에서 `print()`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`
call_lst = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")
test_lst = ["print('\\\\nExpense Tracker')", "print('1. Add an expense')"]
all(call_lst[n].is_equivalent(i) for n, i in enumerate(test_lst))
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
--fcc-editable-region--
    while True:
        print('\nExpense Tracker')
        
--fcc-editable-region--
```
