---
id: 658252f6b1526d57b103d48a
title: 39단계
challengeType: 20
dashedName: step-39
---

# --description--

`if` 문장 안에서 사용자가 지출 금액을 입력하도록 요청하고, 이를 변수에 저장하세요.

빈 `'Enter amount: '` 호출에 문자열 `input()`을 전달하여 지출을 저장할 수 있게 하세요.

# --hints--

`'Enter amount: '` 변수에 할당된 `input()` 호출에 `amount`을 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input('Enter amount: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = input()
--fcc-editable-region--
```
