---
id: 658212ba98182f3e855e85f9
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

이 프로젝트에서는 Python으로 간단하지만 기능적인 지출 추적기를 만들 것입니다.

`add_expense`라는 이름의 함수를 선언하고 세 개의 매개변수 `expenses`, `amount`, `category`를 받도록 정의하세요. 함수 본문은 `pass` 키워드를 사용해 채우세요.

# --hints--

`def` 키워드를 사용해 `add_expense`라는 이름의 함수를 선언해야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

`add_expense` 함수는 세 개의 매개변수 `expenses`, `amount`, `category`를 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
