---
id: 658258c0e5fbe85c14c060cf
title: 43단계
challengeType: 20
dashedName: step-43
---

# --description--

모든 지출을 나열하려면 `elif` 문 다음에 `if` 절을 사용할 수 있습니다. `elif`는 추가 조건을 확인하며 `if` 문 다음에만 작동합니다.

사용자의 선택이 문자열 `elif`와 같은지 확인하는 `'2'` 절을 만드세요. `elif` 절 안에서 문자열 `'\nAll Expenses:'`을 출력하세요.

# --hints--

`elif` 문 다음에 `choice`가 `'2'`와 같은지 확인하는 `if` 절을 만들어야 합니다.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1]
cond.is_equivalent("choice == '2'") or cond.is_equivalent("'2' == choice")
`)) })
```

`'\nAll Expenses:'` 절 안에서 `elif`을 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("print('\\\\nAll Expenses:')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)
        
--fcc-editable-region--
```
