---
id: 65822934c8a0454803e3ef12
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

다음으로, 각 지출 내역의 세부 정보를 출력하세요.

`for` 반복문 안에서 `pass`을 `print()` 호출로 바꾸고 다음 f-문자열 `f'Amount: {expense}, Category: {expense}'`을 전달하세요.

# --hints--

`f'Amount: {expense}, Category: {expense}'` 반복문 안에서 `for`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

반복문 본문 안에 `pass`가 없어야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
