---
id: 65822773f4c60a479ec89f66
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

`print_expenses` 함수 안에 `for` 리스트의 각 항목을 반복하는 `expenses` 반복문을 만드세요. `expense`을 반복 변수로 사용하고 `pass`를 반복문 본문 안으로 옮기세요.

# --hints--

`for` 함수에서 `expenses`를 반복하는 `print_expenses` 반복문을 만들어야 합니다. 반복문 본문 안에서 `pass`를 사용하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

반복 변수로 `expense`를 사용해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

`pass`는 반복문 본문 안에만 있어야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
