---
id: 6582352cbdbcde4d0c4f7b0b
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

람다 함수를 호출하려면 변수 이름 뒤에 소괄호 한 쌍을 붙여 일반 함수 구문을 사용하면 됩니다.

`test` 람다 함수를 호출하고 `3`을 인수로 전달하세요. 그런 다음 결과를 출력하세요.

# --hints--

코드에 `print(test(3))`가 있어야 합니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
