---
id: 6582687859366a618424d84b
title: 50단계
challengeType: 20
dashedName: step-50
---

# --description--

프로그램을 종료할 방법을 제공하려면, 또 다른 `elif` 절을 사용해 `choice`이 문자열 `'5'`와 같은지 확인하세요.

새로운 `elif` 절 안에서, 프로그램이 실행을 종료하고 있음을 보여주기 위해 문자열 `'Exiting the program.'`를 출력하세요.

# --hints--

`elif`인지 확인하는 `choice == '5'` 절을 만들어야 합니다.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[4]
cond.is_equivalent("choice == '5'") or cond.is_equivalent("'5' == choice")
`)) })
```

새로운 `'Exiting the program.'` 절에서 `elif`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[4].has_call("print('Exiting the program.')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)

        elif choice == '2':
            print('\nAll Expenses:')
            print_expenses(expenses)

        elif choice == '3':
            print('\nTotal Expenses: ', total_expenses(expenses))
--fcc-editable-region--
        elif choice == '4':
            category = input('Enter category to filter: ')
            print(f'\nExpenses for {category}:')
            expenses_from_category = filter_expenses_by_category(expenses, category)
            print_expenses(expenses_from_category)
        
--fcc-editable-region--
```
