---
id: 65823634a308cb4d6bcc5fa6
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

람다 함수는 리스트와 같은 객체 모음의 각 요소에 대해 지정된 함수를 실행하는 `map()` 함수와 함께 유용하게 결합할 수 있습니다.

```py
map(lambda x: x * 2, [1, 2, 3])
```

실행할 함수는 첫 번째 인수로 전달되고, 반복 가능한 객체는 두 번째 인수로 전달됩니다.

위 예제의 결과는 `[2, 4, 6]`이 될 것이며, `map()`에 전달된 리스트의 각 항목이 람다 함수의 작용으로 두 배가 되었습니다.

`print()` 호출을 수정하여 `map()`를 `test`를 첫 번째 인수로, `[2, 3, 5, 8]`를 두 번째 인수로 호출한 결과를 출력하세요. 아직 읽기 쉬운 출력 결과는 보이지 않을 것입니다.

# --hints--

코드에 `print(map(test, [2, 3, 5, 8]))`가 있어야 합니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
