---
id: 6582324b5d37fe4b622bc3c0
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

Lambda 함수는 Python에서 간단하고 일회성 작업에 적합한 짧고 익명인 함수입니다. `lambda` 키워드로 정의하며, 다음 구문을 사용합니다:

```py
lambda x: expr
```

위 예시에서 `x`은 표현식 `expr`에 사용되는 매개변수를 나타내며, 일반 함수의 매개변수처럼 작동합니다. `expr`은 lambda 함수가 호출될 때 평가되어 반환되는 표현식입니다.

`test`라는 변수를 만들고, `x` 매개변수를 받아 `x * 2`을 반환하는 lambda 함수를 할당하세요.

# --hints--

`test`라는 변수를 선언해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

`lambda x: x * 2`을 `test` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
