---
id: 658238f7604f154ea9a23e1e
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

`sum()` 함수는 인수로 전달된 반복 가능한 항목들의 합을 반환합니다. `sum()`을 `map()` 및 `lambda` 함수와 함께 사용하여 총 지출 금액을 구할 것입니다.

지금은 간단한 테스트를 하고 현재 `print()` 호출에서 `list()` 호출을 `sum()` 함수 호출로 바꾸되, 인수로 현재 `map()` 호출을 전달하세요.

# --hints--

`sum(map(test, [2, 3, 5, 8]))`를 출력하세요.

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
