---
id: 658254db6e12485a48130f57
title: 40단계
challengeType: 20
dashedName: step-40
---

# --description--

계산을 수행하기 전에 비용 금액을 변환해야 합니다. `float()` 함수는 문자열이나 정수 숫자를 인수로 받아 부동 소수점 숫자를 반환합니다.

`input('Enter amount: ')`을 `float()` 함수에 전달하세요.

# --hints--

`input('Enter amount: ')` 할당에서 `float()`을 `amount` 함수에 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = float(input('Enter amount: '))")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
        
        if choice == '1':
            amount = input('Enter amount: ')
--fcc-editable-region--
```
