---
id: 65823dde36392f505a39f7c7
title: 27단계
challengeType: 20
dashedName: step-27
---

# --description--

마지막으로, `map()` 호출을 `sum()` 함수에 전달하여 총 지출을 구하고 결과를 반환하세요.

# --hints--

현재 `sum()` 호출을 인수로 전달하여 `map()` 함수를 호출하고, `total_expenses` 함수에서 결과를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").has_return("sum(map(lambda expense: expense['amount'], expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    map(lambda expense: expense['amount'], expenses)
--fcc-editable-region--

expenses = []
```
