---
id: 65823ff0d4b991510fade1a8
title: 29단계
challengeType: 20
dashedName: step-29
---

# --description--

`filter_expenses_by_category` 함수 내에서 `pass`을 `lambda` 함수로 교체하세요. `expense`을 매개변수로 사용하고 `'category'` 딕셔너리의 `expense` 키 값과 `category` 사이의 비교를 동등 연산자로 평가하세요.

# --hints--

`lambda`이라는 매개변수를 사용하고 `expense` 함수 내에서 `expense['category'] == category` 표현식을 평가하는 `filter_expenses_by_category` 함수를 만들어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

`pass` 함수 내에 `filter_expenses_by_category`이 없어야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
