---
id: 65823040bdb32949d4c1a96f
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

총 지출 금액을 계산하는 함수를 만들어야 합니다.

하나의 매개변수 `total_expenses`를 받는 `expenses`라는 이름의 함수를 정의하세요. 함수 본문은 우선 `pass` 문장으로 채우세요.

# --hints--

`total_expenses`라는 이름의 함수를 정의해야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

`total_expenses` 함수는 하나의 매개변수 `expenses`를 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
