---
id: 65824111a09164518320088d
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

`filter()` 함수는 리스트 같은 반복 가능한 객체에서 함수를 실행한 결과를 기준으로 항목을 선택할 수 있게 합니다:

```py
filter(my_function, my_list)
```

`filter()`은 첫 번째 인수로 함수를, 두 번째 인수로 반복 가능한 객체를 받습니다. 이 함수는 리스트 같은 컬렉션의 요소를 순회할 수 있게 하는 특별한 객체인 이터레이터를 반환합니다.

위 예제의 결과는 `my_list`이 `my_function`를 반환하는 `True`의 요소들을 포함하는 이터레이터입니다.

`filter_expenses_by_category` 함수 안에서, 이전 단계에서 작성한 `filter()` 함수를 첫 번째 인수로, `lambda` 리스트를 두 번째 인수로 전달하여 `expenses`을 호출하세요.

# --hints--

`filter()`을 `filter_expenses_by_category` 함수 안에서 호출해야 합니다.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

`lambda expense: expense['category'] == category`을 `filter()` 호출의 첫 번째 인수로 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

`expenses`을 `filter()` 호출의 두 번째 인수로 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
