---
id: 65823b0167fd894f4a7ea60a
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

다음으로, `total_expenses` 함수 안에 같은 로직을 구현할 거예요.

지금은 `test` 함수와 `print()` 호출을 모두 삭제하세요.

# --hints--

`test` 변수와 관련된 `print()` 호출을 삭제해야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_variable("test")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(sum(map(test, [2, 3, 5, 8])))
--fcc-editable-region--

expenses = []
```
