---
id: 65824561f3478e5371a33ae5
title: 32단계
challengeType: 20
dashedName: step-32
---

# --description--

다음 단계는 대화형 지출 추적기 프로그램의 진입점이 될 메인 함수를 정의하는 것입니다.

매개변수 없이 `main`라는 이름의 함수를 정의하세요. 이 프로젝트 시작 부분에서 만든 `expenses` 목록으로 함수 본문을 채우세요. 이 목록은 지출 기록을 저장하는 데 사용할 것입니다.

# --hints--

매개변수 없이 `main()`라는 이름의 함수를 정의해야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

`expenses` 목록을 `main()` 함수 본문 안으로 옮겨야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
