---
id: 658246c28575d653d1f89b59
title: 33단계
challengeType: 20
dashedName: step-33
---

# --description--

`while` 루프는 지정한 조건이 `True` 동안 코드 일부를 실행하는 또 다른 종류의 루프입니다. 조건이 `False` 되면 루프가 종료됩니다:

```py
while condition:
    <code>
```

`expenses` 목록 아래에 `while` 루프를 만드세요. 조건으로 `True` 를 사용하고, 루프 본문 안에 문자열 `'\nExpense Tracker'` 를 출력하여 프로그램 제목을 표시하세요.

# --hints--

`while` 함수 내에서 조건으로 `True` 을 사용하여 `main()` 루프를 만들어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

`'\nExpense Tracker'` 루프 안에서 `while` 을 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
