---
id: 65825cda2668995d5168e37c
title: 46단계
challengeType: 20
dashedName: step-46
---

# --description--

`elif`인지 확인하는 또 다른 `choice == '4'` 절을 만드세요. 새 `elif` 안에 변수 `category`을 선언하고 비용 카테고리를 필터링하기 위해 `input('Enter category to filter: ')`를 할당하세요.

# --hints--

`elif`인지 확인하는 `choice == '4'` 절을 만드세요.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[3]
cond.is_equivalent("choice == '4'") or cond.is_equivalent("'4' == choice")
`)) })
```

새 `category` 절 안에 `elif`이라는 변수를 선언하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[3].has_variable("category")`)) })
```

`input('Enter category to filter: ')`를 `category` 변수에 할당하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[3].find_variable("category").is_equivalent("category = input('Enter category to filter: ')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
       
        choice = input('Enter your choice: ')

        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)

        elif choice == '2':
            print('\nAll Expenses:')
            print_expenses(expenses)
--fcc-editable-region--           
        elif choice == '3':
            print('\nTotal Expenses: ', total_expenses(expenses))
        
--fcc-editable-region--
```
