---
id: 65823e87c9741750a22085a7
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

다음으로, `filter_expenses_by_category`라는 이름의 함수를 정의하세요. 이 함수는 두 개의 매개변수 `expenses`과 `category`를 받습니다. 함수 본문은 `pass`를 사용하세요.

# --hints--

`filter_expenses_by_category`라는 이름의 함수를 정의해야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

`filter_expenses_by_category` 함수는 `expenses`와 `category`를 매개변수로 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
