---
id: 66314f2add0a8b57da5d6d91
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

콘솔에 `<map object at 0xd273a8>`와 비슷한 문자열이 출력될 것입니다. 이는 `map()`가 반환한 map 객체의 문자열 표현입니다.

읽기 쉬운 출력을 얻으려면 map 객체를 리스트로 변환해야 합니다. `map()` 호출을 `list()` 함수의 인수로 전달하여 그렇게 하세요.

# --hints--

`map()` 호출을 `list()` 함수의 인수로 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
