---
id: 666951d2e4b27e3a3c72ca5f
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

단일 요소에 접근해 그 값을 가져올 수도 있습니다.

리스트는 문자열처럼 0부터 시작하는 인덱스를 가집니다. 즉, 첫 번째 요소는 인덱스 `0`에 있고, 두 번째 요소는 인덱스 `1`에 있습니다.

요소에 접근하려면 대괄호 표기법을 사용하세요. 예를 들어, `example_list[1]`는 `1`의 두 번째 요소인 인덱스 `example_list`에 있는 요소에 접근하는 것입니다.

`my_list`의 첫 번째 요소를 출력하세요.

# --hints--

리스트의 첫 번째 요소를 `print(my_list[0])`로 출력해야 합니다.

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
