---
id: 66694eaaeb7d253049fa7883
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

목록에 항목을 추가하려면 대괄호 사이에 값을 쉼표로 구분하여 작성하세요.

목록 안에 숫자 `1`와 `2`을 추가하세요.

# --hints--

`my_list`의 첫 번째 요소는 `1`이어야 합니다.

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

`my_list`의 두 번째 요소는 `2`이어야 합니다.

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list`에는 두 개의 요소가 있어야 합니다.

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

`[1, 2]`을 `my_list`에 직접 할당해야 합니다.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
