---
id: 66694fc4bba24f33ca01fa5b
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

Python에는 몇 가지 리스트 메서드가 있습니다. 예를 들어 리스트 항목을 추가하거나 제거하는 메서드가 있습니다.

`append()` 메서드를 사용해 리스트 끝에 항목을 추가할 수 있습니다. 예를 들어:

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list`는 이제 `[4, 5, 6, 7]`입니다.

`append()` 메서드를 사용해 `3`를 `my_list`에 추가해 보세요. 그런 다음 리스트를 출력하세요.

# --hints--

`my_list`의 마지막 요소는 `3`이어야 합니다.

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list`의 길이는 3이어야 합니다.

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

리스트 선언은 변경하지 마세요.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

리스트를 출력해야 합니다.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

`append()`를 사용한 후 리스트를 출력해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
