---
id: 6669539c1379793f9cb8917c
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

Python 리스트는 변경 가능하여 리스트 항목의 값을 바꿀 수 있습니다. 대괄호 표기법을 사용해 요소의 값을 변경할 수 있습니다.

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

이렇게 하면 `example_list`의 값이 `[4, 'oh', 6, 7]`이 됩니다.

`my_list`의 첫 번째 요소를 `0`으로 바꾸고, 리스트를 출력하여 값을 확인하세요.

# --hints--

리스트 선언은 변경하지 마세요.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

`my_list`의 첫 번째 요소는 0이어야 합니다.

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

`my_list[0] = 0`를 사용해 첫 번째 요소를 변경해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

리스트를 출력해야 합니다.

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

첫 번째 요소를 변경한 후 리스트를 출력해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
