---
id: 66694bfa435cef241f6f9a68
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

프로젝트를 시작하기 전에 리스트에 대해 배워보세요.

숫자와 문자열과 달리, <dfn>리스트</dfn>는 다른 값들을 담는 컨테이너 역할을 하는 데이터 타입입니다.

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

리스트는 모든 값 주위에 대괄호가 있고, 값들 사이에 쉼표가 있는 것이 특징입니다. 예를 들어 `["A", "happy", "list"]`와 같습니다. 만약 리스트에 값이 하나도 없다면, 빈 리스트입니다: `[]`.

리스트는 서로 다른 데이터 타입을 담을 수 있습니다: `[1, "Up", ["Down", "Twice"]]`. 여기에는 모든 가능한 데이터 타입이 포함됩니다. 또 다른 리스트를 담을 수도 있습니다!

`my_list`라는 변수를 만들고 빈 리스트를 할당하세요.

# --hints--

`my_list`라는 변수가 있어야 합니다.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

변수 `my_list`는 빈 리스트 값을 가져야 합니다.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
