---
id: 666ab9013cd5563f2b39947a
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

리스트를 살펴봤으니, 프로젝트를 시작할 수 있도록 검토를 위해 작성한 모든 코드를 삭제하세요.

# --hints--

모든 것을 삭제해야 합니다.

```js
({
    test: () => runPython(`
    assert not _Node(_code)
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)

my_list.pop()
print(my_list)

--fcc-editable-region--
```
