---
id: 666aae6a5d30a71f1fd7749f
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

`insert` 메서드는 리스트의 어느 위치에든 요소를 추가할 수 있습니다. 첫 번째 인수는 요소를 추가해야 하는 위치이고, 두 번째 인수는 추가할 요소입니다. 예를 들어, `example_list`의 세 번째 위치에 새 요소를 추가하는 방법은 다음과 같습니다:

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

`insert()`를 사용하여 `1`에 `my_list`를 올바른 위치에 추가해 숫자가 오름차순이 되도록 한 다음, 리스트를 출력하세요.

# --hints--

`1`를 사용하여 `my_list`에 `1`를 인덱스 `insert()`에 추가해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

`my_list`에는 4개의 요소가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list`은 `[0, 1, 2, 3]`과 같아야 합니다.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

새 요소를 추가한 후 리스트를 출력해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
