---
id: 666ab82c874add3c95145213
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

`pop()` 메서드는 리스트에서 요소를 제거하는 데 사용할 수 있습니다. 기본적으로 리스트의 마지막 요소를 제거합니다. 메서드에 인덱스를 인수로 전달하면 해당 인덱스에 있는 요소를 제거합니다.

```py
fruits_list = ["cherry", "lemon", "tomato", "apple", "orange"]

fruits_list.pop(2)

print(fruits_list) # ["cherry", "lemon", "apple", "orange"]
```

이 경우 `fruits_list.pop(2)`은 리스트의 `2` 인덱스에 있는 요소를 제거합니다.

`pop()`을 사용해 `my_list`에서 마지막 요소를 제거한 다음 `my_list`를 출력하세요.

# --hints--

`pop()`을 사용하세요.

```js
({
    test: () => runPython(`
    assert len(_Node(_code).find_calls('pop')) > 0
    `)
})
```

`my_list`은 `[0, 1, 2]`이어야 합니다.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2]
    `)
})
```

`my_list`을 사용한 후 `pop()`를 출력하세요.

```js
({
    test: () => runPython(`
    code = _Node(_code)
    assert(
        any(
            code.is_ordered(pop, 'print(my_list)') for pop in ['my_list.pop()', 'my_list.pop(3)']
        )
    )
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)


--fcc-editable-region--
```
