---
id: 5dc174fcf86c76b9248c6eb2
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Елементи HTML мають початковий та кінцевий теги з вмістом між ними.

Ось основний синтаксис:

```html
<openingTag>content</closingTag>
```

Перший елемент, про який ви дізнаєтеся, — елемент `h1`. Елемент `h1` є елементом заголовка, який використовується для головного заголовка вебсторінки.

```html
<h1>This is a main heading</h1>
```

Змініть текст елемента `h1` нижче з `Hello World` на `CatPhotoApp` та перегляньте зміну в попередньому перегляді браузера.

Коли закінчите, натисніть на кнопку «Перевірити код», щоб переконатись, що все правильно.

# --hints--

Код повинен містити текст `CatPhotoApp`. Можливо, потрібно перевірити написання.

```js
assert(code.match(/catphotoapp/i));
```

Елемент `h1` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Елемент `h1` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/h1\>/));
```

Ви маєте більше ніж один елемент `h1`. Видаліть зайвий елемент `h1`.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Текстом елемента `h1` має бути `CatPhotoApp`. Ви пропустили текст, маєте друкарську помилку, або його немає між початковим та кінцевим тегами елемента `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Здається, ви використовуєте розширення браузера, яке змінює сторінку. Переконайтеся, що вимкнули усі розширення браузера.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

