---
id: 5dc23f9bf86c76b9248c6eba
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Ви можете додати зображення до свого сайту за допомогою елемента `img`. Елементи `img` мають початковий тег, але не мають кінцевого. Елемент без кінцевого тега відомий як <dfn>порожній елемент</dfn>.

Додайте елемент `img` під елементом `p`. Наразі жодне зображення не буде відтворюватись в браузері.

# --hints--

Елемент `img` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('img'));
```

Елемент `img` не повинен мати кінцевий тег `</img>`.

```js
assert(!code.match(/<\/img\>/));
```

Ви повинні мати тільки один елемент `img`. Видаліть зайві.

```js
assert(document.querySelectorAll('img').length === 1);
```

Елемент `img` повинен бути під елементом `p`. Вони знаходяться в неправильному порядку.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```

