---
id: 5ddbd81294d8ddc1510a8e56
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Текст посилання потрібно писати між початковим та кінцевим тегами елемента посилання (`a`).

Ось приклад посилання з текстом `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Додайте текст `link to cat pictures` до елемента посилання. Це буде текстом посилання.

# --hints--

Елемент посилання (`a`) повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Елемент посилання (`a`) повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/a\>/));
```

Текстом елемента посилання (`a`) має бути `link to cat pictures`. Переконайтеся, що розмістили текст посилання між початковим та кінцевим тегами елемента посиланння (`a`).

```js
assert(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'link to cat pictures'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

