---
id: 5dfb6250eacea3f48c6300b2
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

Після невпорядкованого списку додайте нове зображення з таким значенням атрибута `src`:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

Та з таким значенням атрибута `alt`:

`A slice of lasagna on a plate.`

# --hints--

Елемент `img` повинен бути відразу після кінцевого тега `</ul>`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.lastElementChild?.nodeName, 'IMG');
```

Нове зображення не має атрибута `alt`. Перевірте, чи є пробіл після назви початкового тега та/або є пробіли перед усіма назвами атрибутів.

```js
assert.isTrue(document.querySelectorAll('section')?.[1]?.lastElementChild?.hasAttribute('alt'));
```

Нове зображення повинне мати `alt` зі значенням `A slice of lasagna on a plate.` Переконайтеся, що значення атрибута `alt` написане в лапках.

```js
assert(
  document.querySelectorAll('section')?.[1]
    ?.lastElementChild?.getAttribute('alt')
    .replace(/\s+/g, ' ')
    .match(/^A slice of lasagna on a plate\.?$/i)
);
```

Нове зображення не має атрибута `src`. Перевірте, чи є пробіл після назви початкового тега та/або є пробіли перед усіма назвами атрибутів.

```js
assert.isTrue(document.querySelectorAll('section')?.[1]?.lastElementChild?.hasAttribute('src'));
```

Нове зображення повинне мати `src` зі значенням `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. Переконайтеся, що значення атрибута `src` написане в лапках.

```js
assert.strictEqual(
  document.querySelectorAll('section')?.[1]?.lastElementChild?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Хоча ви встановили атрибут `src` нового зображення як правильну URL-адресу, краще завжди писати значення атрибута в лапках.

```js
assert.isTrue(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

