---
id: 5ef9b03c81a63668521804dc
title: Крок 46
challengeType: 0
dashedName: step-46
---

# --description--

Ви можете використовувати перемикач для запитань, де хочете тільки одну відповідь з декількох варіантів.

Ось приклад радіокнопки з текстом `cat`:

```html
<input type="radio"> cat
```

Пам’ятайте, що елемент `input` є порожнім елементом.

Перед полем тексту додайте радіокнопку з текстом:

`Indoor`

# --hints--

Створіть елемент `input` для радіокнопки. Перевірте синтаксис.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

Елемент `input` повинен мати початковий тег, але не кінцевий тег.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

Додайте лише один елемент `input` для радіокнопки. Видаліть зайві.

```js
assert(document.querySelectorAll('form > input'), 2);
```

Новий елемент `input` має бути над наявним `input` з атрибутом `type` зі значенням `text`. Вони знаходяться в неправильному порядку.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Новий елемент `input` не має атрибута `type`. Перевірте, чи є пробіл після назви початкового тега.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Новий елемент `input` повинен мати лише один атрибут `type`. Видаліть зайві.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Новий елемент `input` повинен мати атрибут `type` зі значенням `radio`. Ви або не написали значення, або маєте друкарську помилку. Пам'ятайте, що значення атрибутів повинні бути в лапках.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

Хоча ви встановили атрибут `type` нового елемента `input` як `radio`, краще завжди писати значення атрибута в лапках.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Текст `Indoor` кнопки `radio` повинен бути розташованим після неї, а не перед нею.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

Текст `Indoor` повинен бути розташованим праворуч від кнопки `radio`. Ви або не написали текст, або маєте друкарську помилку.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

