---
id: 5ef9b03c81a63668521804d6
title: Крок 37
challengeType: 0
dashedName: step-37
---

# --description--

Зараз ви додасте вебформу для збору інформації від користувачів.

Елемент `form` використовують, щоб отримати інформацію від користувачів (наприклад, ім’я, електронну пошту тощо).

Після заголовка `Cat Form` додайте елемент `form`.

# --hints--

Елемент `form` повинен мати початковий та кінцевий теги. Можливо, вам не вистачає одного/обох необхідних тегів або вони розміщені в неправильному порядку.

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

Теги елемента `form` розташовані в неправильному порядку.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

Елемент `form`, вкладений в останній елемент `section`, повинен бути під елементом `h2`. Елементи `h2` та `form` розташовані в неправильному порядку.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Елемент `form` не повинен мати вмісту. Видаліть будь-які елементи HTML або текст між тегами елемента `form`.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

