---
id: 5f07c98cdb9413cbd4b16750
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Настав час додати нову секцію. Додайте другий елемент `section` під наявним елементом `section`.

# --hints--

Елемент `section` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 2);
```

Додайте лише один початковий тег `section`. Будь ласка, видаліть всі зайві.

```js
assert(document.querySelectorAll('section').length === 2);
```

Елемент `section` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/section>/g).length >= 2);
```

Додайте лише один кінцевий тег `section`. Будь ласка, видаліть всі зайві.

```js
assert(code.match(/<\/section>/g).length === 2);
```

Другий елемент `section` не повинен бути вкладеним в перший елемент `section`.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

Обидва елементи `section` повинні бути між початковим та кінцевим тегами елемента `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

