---
id: 5f07be6ef7412fbad0c5626b
title: Крок 18
challengeType: 0
dashedName: step-18
---

# --description--

Перш ніж додавати будь-який новий вміст, вам слід використати елемент `section`, щоб відокремити вміст фотографій котів від майбутнього вмісту.

Елемент `section` використовують, щоб визначити розділи в документі (наприклад, частини, верхні та нижні колонтитули або будь-який інший розділ документа). Це семантичний елемент, який допомагає з SEO та доступністю.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Візьміть елементи `h2`, `p` та посилання (`a`) і вкладіть їх в елемент `section`.

# --hints--

Елемент `section` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('section'));
```

Елемент `section` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/section\s*>/i));
```

Весь елемент `section` повинен бути між початковим та кінцевим тегами елемента `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

Наявні елементи `h2`, `p` та посилання (`a`) мають бути між початковим і кінцевим тегами елемента `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4);
```

Елемент `h1` не повинен бути вкладеним в елементі `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

