---
id: 671141feba228a35cefba82d
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Перетворіть наявний текст `cute cats` на елемент посилання, що направляє на:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

У першому елементі `p` має бути новий елемент посилання.

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

Текстом посилання має бути `cute cats`.

```js
assert.strictEqual(document.querySelector('a')?.innerText.trim(), "cute cats");
```

Атрибутом `href` нового елемента посилання має бути `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`.

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

Код повинен містити фразу `cute cats` лише один раз.

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

Текстом елемента `p` досі має бути `Everyone loves cute cats online!`.

```js
assert.strictEqual(document.querySelector('p')?.innerText.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
