---
id: 671141d8e32fe934c26fa1be
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Додайте слова `See more` перед елементом посилання та `in our gallery` після елемента посилання.

# --hints--

Використайте наявний елемент посилання, не створюйте новий.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Не змінюйте текст в межах елемента посилання, ним має бути `cat photos`.

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

Код повинен містити `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery`.

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
