---
id: 7cf9b03d81a65668421804c3
title: Крок 41
challengeType: 0
dashedName: step-41
---

# --description--

Для того, щоб дані форми були доступні з місця розташування, зазначеного в атрибуті `action`, ви повинні надати текстовому полю атрибут `name` та призначити йому значення для представлення даних, що надсилаються.

Ось приклад елемента `input` з атрибутом `name`:

```html
<input type="text" name="name">
```

Додайте атрибут `name` зі значенням `catphotourl` до текстового поля.

# --hints--

Ви або видалили елемент `input`, або він має недійсний синтаксис. Всі значення атрибутів повинні бути в лапках.

```js
assert(document.querySelectorAll('input').length);
```

`form` має містити лише елемент `input`. Видаліть будь-які додаткові елементи HTML або текст в елементі `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Елемент `input` не має атрибута `name`. Перевірте, чи є пробіл після назви початкового тега та/або є пробіли перед усіма назвами атрибутів.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

Елемент `input` повинен мати атрибут `name` зі значенням `catphotourl`.  Ви або не написали значення, або маєте друкарську помилку. Пам'ятайте, що значення атрибутів повинні бути в лапках.

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

Хоча ви встановили атрибут `name` елемента `input` як `catphotourl`, краще завжди писати значення атрибута в лапках.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

