---
id: 60a3e3396c7b40068ad69974
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Kila uchoraji unahitaji fremu.

Panga ndani kipengele cha `.canvas` katika kipengele kingine cha `div`. Mpe `div` darasa la `frame`.
# --hints--

Unapaswa kuongeza kipengele kipya cha `div`.

```js
assert(document.querySelectorAll('div').length === 2)
```

Kipengele chako cha `.canvas` kinapaswa kupangwa ndani ya kipengele kipya cha `div`.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

`div` mpya inapaswa kuwa na `class` yenye thamani ya `frame`.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

`div` mpya inapaswa kuwa ndani ya kipengele chako cha `body`.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
