---
id: 60a3e3396c7b40068ad6996e
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Ukingo ni eneo lililo nje ya kisanduku, na linaweza kutumika kudhibiti nafasi kati ya visanduku vingine au vipengele.

Hapa kipengele cha chini kina ukingo mkubwa wa juu, kinachokisukuma chini zaidi kwenye ukurasa.

Sasa kwa kuwa umeelewa mfano wa kisanduku cha CSS, tuanze na uchoraji wa Rothko.

Ondoa kipengele cha `<img>`.

# --hints--

Haupaswi kuwa na kipengele cha `img` katika msimbo wako.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
