---
id: 60a3e3396c7b40068ad6996d
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Yaliyomo yamezungukwa na nafasi inayoitwa nafasi ya ndani, kama vile vifuniko vya bubble wrap vinavyotenganisha kipengele kutoka kwenye sanduku lililokizunguka.

Fikiria mpaka kama sanduku la karatasi ambalo kipengele chako kilitumwa ndani yake.

Badilisha sifa ya kipengele `src` kuwa `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

Kipengele cha `img` kinapaswa kuwa na sifa ya kipengele `src` ya `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
