---
id: 60a3e3396c7b40068ad6996b
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Ndani ya kipengele cha `head`, ongeza lebo ya `meta` ambayo inaweka `charset` kuwa `UTF-8`, na kipengele cha `title` chenye thamani ya `Rothko Painting`.

Ndani ya kipengele cha `body`, ongeza kipengele cha `img` chenye `src` ya `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.
# --hints--

Msimbo wako unapaswa kuwa na lebo ya `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Lebo ya `meta` inapaswa kuweka `charset` kuwa `UTF-8`.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Msimbo wako unapaswa kuwa na kipengele cha `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` inapaswa kuwa `Rothko Painting`.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

Msimbo wako unapaswa kuwa na kipengele cha `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

Kipengele cha `img` kinapaswa kuwa na `src` ya `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
