---
id: 60a3e3396c7b40068ad6996f
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Ongeza kipengele cha `div` ndani ya `body`.

Weka sifa ya kipengele ya `class` sawa na `canvas`.

Hii itatumika kama turubai ya kuchora kwa uchoraji wako.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `div`.

```js
assert(document.querySelectorAll('div').length === 1)
```

Kipengele cha `div` kinapaswa kuwa na sifa ya kipengele ya `class` yenye thamani ya `canvas`.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
