---
id: 66ea7e7399244c6d4b6052a1
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

工作坊的下一部分將回顧 `Math.random()` 方法的運作方式。

首先為你的 `console.log()` 添加另一個，將字串 `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` 記錄到主控台。
 
# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該有一個 `console.log()`，用來記錄字串 `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."`。

```js
assert.equal(
  getLogs()[1],
  "The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."
);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```
