---
id: 66ea7fc048a6e26e33c598fe
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

現在，是時候使用 `Math.random()` 方法來產生隨機數。

創建一個名為 `randomNum` 的變數，並將呼叫 `Math.random()` 方法的結果指定給它。

然後，將 `randomNum` 變數記錄到主控台。

嘗試在程式碼中添加一個空格來重新執行機器人，並查看產生的不同隨機數。

# --hints--

你應該有一個名為 `randomNum` 的變數。

```js
assert.isNotNull(randomNum);
```

你應該將呼叫 `Math.random()` 方法的結果指定給 `randomNum` 變數。

```js
assert.match(code, /randomNum\s*=\s*Math\.random\(\);?/);
```

你不應該為 `randomNum` 變數硬編碼隨機數。請確保使用 `Math.random()` 方法來產生隨機數。

```js
assert.notMatch(code, /randomNum\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum\s*=\s*\d+;?/);
```

你應該將 `randomNum` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
