---
id: 66ea8239b9eeb76f12299030
title: 步驟4
challengeType: 1
dashedName: step-4
---

# --description--

工作坊的下一部分內容是回顧如何在兩個數值之間生成隨機數。

首先添加另一個 `console.log()`，將信息 `"Now, generate a random number between two values."` 記錄到控制檯。

然後，創建一個名爲 `min` 的變量，爲其賦值 `1`，創建一個名爲 `max` 的變量，爲其賦值 `100` 。

下一步，你將在這兩個值之間生成一個隨機數。

# --hints--

你應該有一個 `console.log()`，記錄信息 `"Now, generate a random number between two values."`。

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

你應該有一個名爲 `min` 的變量。

```js
assert.isNotNull(min);
```

你的 `min` 變量應該是一個數字。

```js
assert.isNumber(min);
```

你應將數字 `1` 賦值給 `min` 變量。

```js
assert.equal(min, 1);
```

你應該有一個名爲 `max` 的變量。

```js
assert.isNotNull(max);
```

你的 `max` 變量應該是一個數字。

```js
assert.isNumber(max);
```

你應將數字 `100` 賦值給 `max` 變量。

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
