---
id: 66ea7e7399244c6d4b6052a1
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

工作坊的下一部分將回顧 `Math.random()` 方法的運作方式。 

首先為你的 `console.log()` 添加另一個，將字串 `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` 記錄到主控台。
 
# --hints--

你應該有一個 `console.log()`，用來記錄字串 `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."`。

```js
assert.match(code, /console\.log\((`|'|")\s*The\s+Math\.random\(\)\s+method\s+returns\s+a\s+pseudo\s+random\s+number\s+greater\s+than\s+or\s+equal\s+to\s+0\s+and\s+less\s+than\s+1\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```
