---
id: 66ea8239b9eeb76f12299030
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

工作坊的下一部分是檢視如何在兩個值之間產生隨機數。 

首先為你的程式碼添加另一個 `console.log()`，將訊息 `"Now, generate a random number between two values."` 記錄到主控台。

然後，創建一個名為 `min` 的變數並指定其值為 `1`，再創建一個名為 `max` 的變數並指定其值為 `100`。 

在下一步中，你將在這兩個值之間產生一個隨機數。

# --hints--

你應該有一個 `console.log()`，用來記錄訊息 `"Now, generate a random number between two values."`。

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

你應該有一個名為 `min` 的變數。

```js
assert.isNotNull(min);
```

你的 `min` 變數應該是一個數字。

```js
assert.isNumber(min);
```

你應該將數字 `1` 指定到你的 `min` 變數。

```js
assert.equal(min, 1);
```

你應該有一個名為 `max` 的變數。

```js
assert.isNotNull(max);
```

你的 `max` 變數應該是一個數字。

```js
assert.isNumber(max);
```

你應該將數字 `100` 指定到你的 `max` 變數。

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
