---
id: 66ea845f54195c7020cc521b
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

在兩個值之間產生隨機數的公式如下：

```js
Math.random() * (maximum - minimum) + minimum;
```

這將產生介於兩個值之間的浮點數結果。

創建一個名為 `randomNum2` 的變數，並將在 `min` 和 `max` 值之間生成的結果指定給它。

然後，記錄 `randomNum2` 變數以查看結果。嘗試重新整理頁面以查看不同的結果。

# --hints--

你應該有一個名為 `randomNum2` 的變數。

```js
assert.isNotNull(randomNum2);
```

你的 `randomNum2` 變數應該是一個數字。

```js
assert.isNumber(randomNum2);
```

你應該指定在 `min` 和 `max` 值之間產生的結果。請確保不要硬編碼一個值。你應該使用上面提供的公式。

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

你應該將 `randomNum2` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
