---
id: 66ea89c25e1bd871688eba6a
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

到目前為止，你只處理過浮點數和數字 `0`。

在工作坊的下一部分，你將檢視如何使用 `Math.floor()` 方法。

首先為你的程式碼添加一個 `console.log()`，用來記錄訊息 `"The Math.floor() method rounds the value down to the nearest whole integer."`

# --hints--

你應該記錄訊息 `"The Math.floor() method rounds the value down to the nearest whole integer."`

```js
assert.match(code, /console\.log\((["'])The\s+Math\.floor\(\)\s+method\s+rounds\s+the\s+value\s+down\s+to\s+the\s+nearest\s+whole\s+integer\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);
--fcc-editable-region--

--fcc-editable-region--
```
