---
id: 66ea7adae8053065a64f9002
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個 Mathbot 來複習使用不同的 `Math` 物件方法。這個 Mathbot 將會將一些數學操作和訊息記錄到主控台，以幫助你了解 JavaScript 的 `Math` 物件。

在這第一步，先建立一個名為 `botName` 的變數，並將字串值 `"MathBot"` 指定給它。

然後，創建另一個名為 `greeting` 的變數，並將句子 `"Hi there! My name is [botName goes here] and I am here to teach you about the Math object!"` 指定給它。在 `[botName goes here]` 佔位符的位置，使用 `botName` 變數。 

你可以自由使用模板字面值或使用 `+` 運算子進行字串串接來達成此目的。

最後，將 `greeting` 變數記錄到主控台。

# --hints--

你應該有一個名為 `botName` 的變數。

```js
assert.isNotNull(botName);
```

你的 `botName` 變數應該是一個字串。

```js
assert.isString(botName);
```

你應該將字串值 `"MathBot"` 指定給 `botName` 變數。

```js
assert.strictEqual(botName, "MathBot");
```

你應該有一個名為 `greeting` 的變數。

```js
assert.isNotNull(greeting);
```

你的 `greeting` 變數應該是一個字串。

```js
assert.isString(greeting);
```

你應該將以下句子指定給 `greeting` 變數：`"Hi there! My name is [MathBot goes here] and I am here to teach you about the Math object!"`。請確定用 `botName` 變數替換 `[MathBot goes here]`。

```js
assert.strictEqual(greeting, "Hi there! My name is MathBot and I am here to teach you about the Math object!");
```

你應該將 `greeting` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
