---
id: 66eb3ec3b07dfe8f898646c0
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

如你在先前課程中所學，`Math.floor()` 方法會將數值向下取整到最接近的整數。

```js
const price = 10.99;
Math.floor(price); // 10
```

創建一個名為 `numRoundedDown` 的變數，並將將浮點數 `6.7` 向下取整到最接近整數的結果指定給它。

然後，將 `numRoundedDown` 變數記錄到主控台以查看結果。 

# --hints--

你應該有一個 `numRoundedDown` 變數。

```js
assert.isNotNull(numRoundedDown);
```

`numRoundedDown` 變數應該是一個數字。

```js
assert.isNumber(numRoundedDown);
```

你的 `numRoundedDown` 變數應該與將浮點數 `6.7` 向下取整到最接近整數的結果相等。

```js
assert.equal(numRoundedDown, 6);
```

你不應該為 `numRoundedDown` 變數硬編碼值 `6`。

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

你應該將 `numRoundedDown` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
