---
id: 66eb488c0129e192542f0865
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

在之前的課程中，你學會了如何像這樣使用 `Math.ceil()` 方法。

```js
const price = 10.01;
Math.ceil(price); // 11
```

`Math.ceil()` 方法會將數值向上取整到最接近的整數。

創建一個名為 `numRoundedUp` 的變數，並將將浮點數 `3.2` 無條件進位到最接近整數的結果指定給它。

然後，將 `numRoundedUp` 變數記錄到主控台以查看結果。

# --hints--

你應該有一個 `numRoundedUp` 變數。

```js
assert.isNotNull(numRoundedUp);
```

`numRoundedUp` 變數應該是一個數字。

```js
assert.isNumber(numRoundedUp);
```

你的 `numRoundedUp` 變數應該與將浮點數 `3.2` 向下取整到最接近整數的結果相等。

```js
assert.equal(numRoundedUp, 4);
```

你不應該為 `numRoundedUp` 變數硬編碼 `4` 的值。

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

你應該將 `numRoundedUp` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
