---
id: 66eb3ec3b07dfe8f898646c0
title: 步驟7
challengeType: 1
dashedName: step-7
---

# --description--

As you learned in previous lessons, the `Math.floor()` method rounds the value down to the nearest whole integer.

```js
const price = 10.99;
Math.floor(price); // 10
```

創建一個名爲 `numRoundedDown` 的變量，將浮點數 `6.7` 向下取整到最接近的整數後的結果賦值給它。

然後，將 `numRoundedDown` 變量記錄到控制檯以查看結果。

# --hints--

你應該有一個 `numRoundedDown` 變量。

```js
assert.isNotNull(numRoundedDown);
```

`numRoundedDown` 變量應爲數字。

```js
assert.isNumber(numRoundedDown);
```

你的 `numRoundedDown` 變量應等於將浮點數 `6.7` 向下取整到最接近的整數。

```js
assert.equal(numRoundedDown, 6);
```

你不應將 `numRoundedDown` 變量的值硬編碼爲 `6`。

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

你應將 `numRoundedDown` 變量記錄到控制檯。

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
