---
id: 66eb488c0129e192542f0865
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

In previous lessons, you learned how to work with the `Math.ceil()` method like this:

```js
const price = 10.01;
Math.ceil(price); // 11
```

`Math.ceil()` 方法將數值向上取整爲最接近的整數。

創建一個名爲 `numRoundedUp` 的變量，將浮點數 `3.2` 向上取整到最接近的整數賦值給它。

然後，將 `numRoundedUp` 變量記錄到控制檯以查看結果。

# --hints--

你應該有一個 `numRoundedUp` 變量。

```js
assert.isNotNull(numRoundedUp);
```

`numRoundedUp` 變量應爲數字。

```js
assert.isNumber(numRoundedUp);
```

你的 `numRoundedUp` 變量應等於將浮點數 `3.2` 向上取整到最接近的整數。

```js
assert.equal(numRoundedUp, 4);
```

你不應將 `numRoundedUp` 變量的值硬編碼爲 `4`。

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

你應將 `numRoundedUp` 變量記錄到控制檯。

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
