---
id: 66ece147323c91f654bf8dd9
title: 步驟 13
challengeType: 1
dashedName: step-13
---

# --description--

In previous lessons, you learned how to get the maximum and minimum values from a range of numbers like this:

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

創建一個名爲 `maxNum` 的變量，將 `3, 125, 55, 24` 中最大數字賦值給它。 然後，將 `maxNum` 的值記錄到控制檯。

在下面創建一個名爲 `minNum` 的變量，將 `6, 90, 14, 90, 2` 中的最小數賦給它。 然後，將 `minNum` 的值記錄到控制檯。

# --hints--

你應該有一個名爲 `maxNum` 的變量。

```js
assert.isNotNull(maxNum);
```

你的 `maxNum` 變量應賦值爲 `3, 125, 55, 24` 中最大數字。

```js
assert.equal(maxNum, 125);
```

你不應將 `maxNum` 變量的值硬編碼爲 `125`。 請確保使用 `Math.max()` 方法。

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

你應該將 `maxNum` 的值輸出到控制檯。

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

你應該有一個名爲 `minNum` 的變量。

```js
assert.isNotNull(minNum);
```

你的 `minNum` 變量應賦值爲 `6, 90, 14, 90, 2` 中最小數字。

```js
assert.equal(minNum, 2);
```

你不應將 `minNum` 變量的值硬編碼爲 `2`。 請確保使用 `Math.min()` 方法。

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

你應該將 `minNum` 的值輸出到控制檯。

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
