---
id: 66ece147323c91f654bf8dd9
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

在之前的課程中，你學會了如何從一個數字的範圍中取得最大值和最小值，如下所示：

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

創建一個名為 `maxNum` 的變數，並將尋找數字 `3, 125, 55, 24` 中最大數字的結果指定給它。然後，將 `maxNum` 的值記錄到主控台。

在那下面，創建一個名為 `minNum` 的變數，並指定它為尋找數字 `6, 90, 14, 90, 2` 中最小數字的結果。然後，將 `minNum` 的值記錄到主控台。

# --hints--

你應該有一個名為 `maxNum` 的變數。

```js
assert.isNotNull(maxNum);
```

你的 `maxNum` 變數應該被指定為在數字 `3, 125, 55, 24` 中尋找最大數字的結果。

```js
assert.equal(maxNum, 125);
```

你不應該將值 `125` 硬編碼到 `maxNum` 變數。請確保改用 `Math.max()` 方法。

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

你應該將 `maxNum` 的值記錄到主控台。

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

你應該有一個名為 `minNum` 的變數。

```js
assert.isNotNull(minNum);
```

你的 `minNum` 變數應該被指定為尋找數字 `6, 90, 14, 90, 2` 中最小數字的結果。

```js
assert.equal(minNum, 2);
```

你不應該將數值 `2` 硬編碼到 `minNum` 變數。請確保改用 `Math.min()` 方法。

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

你應該將 `minNum` 的值記錄到主控台。

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
