---
id: 69218a15fe536ee1be38aefb
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

工作坊的下一部分是檢視如何在兩個值之間產生隨機整數。 

首先為你的程式碼添加另一個 `console.log()`，將訊息 `"Now, generate a random integer between two values."` 記錄到主控台。

在下一步，你將在你已經創建的 `min` 和 `max` 變數的值之間產生一個隨機整數。

# --hints--

你應該有一個 `console.log()`，用來記錄訊息 `"Now, generate a random integer between two values."`。

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+integer\s+between\s+two\s+values\.\s*\1\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

--fcc-editable-region--

--fcc-editable-region--
```
