---
id: 69218e46f4490b33f2481a1a
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

在兩個值之間產生隨機整數的公式如下：

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

這將產生介於兩個值之間的整數結果。

創建一個名為 `randomInt` 的變數，並將在 `min` 和 `max` 值之間生成的結果指定給它。

然後，將 `randomInt` 變數記錄下來以查看結果。嘗試重新整理頁面以查看不同的結果。

# --hints--

你應該有一個名為 `randomInt` 的變數。

```js
assert.isNotNull(randomInt);
```

你的 `randomInt` 變數應該是一個數字。

```js
assert.isNumber(randomInt);
```

你應該指定在 `min` 和 `max` 值之間產生的結果。請確保不要硬編碼一個值。你應該使用上面提供的公式。

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

你應該將 `randomInt` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
