---
id: 66325a250690a3612c1db0f6
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

CamperBot hat eine neue Variable namens `bgHexCodeSpanElement` erstellt, um die Referenz auf das `span`-Element mit der `id` von `bg-hex-code` zu speichern. Wenn er jedoch versucht, diese Variable in der Konsole auszugeben, erhält er `null`. 

`null` ist ein spezieller Wert in JavaScript, der das Fehlen eines Werts darstellt. Dies kann passieren, wenn Sie versuchen, auf eine Eigenschaft eines Objekts zuzugreifen, die nicht existiert. 

In diesem Fall übergibt CamperBot nicht den korrekten Selektor an die `document.querySelector`-Methode.

Korrigieren Sie die `document.querySelector("bg-hex-code")`-Zeile so, dass sie das Element mit der `id` von `bg-hex-code` korrekt auswählt.

# --hints--

Sie sollten `"bg-hex-code"` so korrigieren, dass es korrekt auf die `id` von `bg-hex-code` zugreift. Denken Sie daran, dass `id`-Selektoren in CSS mit einem Rautezeichen (`#`) beginnen.

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

--fcc-editable-region--
const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("bg-hex-code");

console.log(bgHexCodeSpanElement);
--fcc-editable-region--
```
