---
id: 612e813b3ba67633222cbe54
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Lembre-se, um atributo `class` pode conter múltiplos valores.Para diferenciar entre suas teclas brancas e pretas, adicione um segundo valor de `class` de `black--key` aos seus segundo, terceiro, quinto, sexto e sétimo elementos `.key`.

# --hints--

O segundo elemento `.key` deve ter o atributo `class` com o valor `black--key`.

```js
const key = document.querySelectorAll('.key')[1];
assert.isTrue(key?.className.includes('black--key'));
```

O terceiro elemento `.key` deve ter a `class` `black--key`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert.isTrue(third?.classList?.contains('black--key'));
```

O quinto elemento `.key` deve ter a `class` `black--key`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert.isTrue(fifth?.classList?.contains('black--key'));
```

O sexto elemento `.key` deve ter a `class` `black--key`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert.isTrue(sixth?.classList.contains('black--key'));
```

O sétimo elemento `.key` deve ter a `class` `black--key`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert.isTrue(seventh?.classList?.contains('black--key'));
```

Você deve ter cinco elementos `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 5);
```

Você deve ter sete elementos `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
