---
id: 612e89562043183c86df287c
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Navegadores frequentemente aplicam valores padrão de margem e padding a elementos específicos.Para garantir que o layout do seu piano seja exibido corretamente, você precisa redefinir o box model.

Adicione um seletor de regras `html` ao seu arquivo CSS e defina a propriedade `box-sizing` como `border-box`.

# --hints--

Você deve ter um seletor `html`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

O seletor `html` deve ter a propriedade `box-sizing` com o valor de `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
