---
id: 612e7d1c29fb872d6384379c
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Coloque um segundo elemento `div` dentro da `div` existente e defina `class` como `keys`.

# --hints--

Crie um segundo elemento `div`.

```js
const divDiv = document.querySelectorAll('div');
assert.lengthOf(divDiv, 2);
```

O novo elemento `div` deve estar dentro do elemento `div` existente.

```js
const div = document.querySelector('div');
assert.lengthOf(div?.children, 1);
assert.equal(div?.children?.[0]?.localName, 'div');
```

O novo elemento `div` deve ter `class` definida como `keys`.

```js
const div = document.querySelector('div');
assert.equal(div?.children?.[0]?.className, 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
