---
id: 612e8279827a28352ce83a72
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Agora, copie o conjunto de sete elementos `.key` e cole dois conjuntos adicionais dentro da div `.keys`.

# --hints--

Você deve ter 21 elementos `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

Você deve ter 15 elementos `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

Todos os elementos `.key` devem estar dentro do seu elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

Todos os elementos `.black--key` devem estar dentro do seu elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
