---
id: 5f331e55dfab7a896e53c3a1
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`title`는 웹 페이지에 보이지 않는 추가 정보를 제공하는 여러 요소 중 하나지만, 검색 엔진이나 페이지 표시 방식에 유용합니다.

`head` 요소 안에 `meta` 요소를 중첩하고, `charset`라는 속성을 `utf-8` 값으로 설정하여 브라우저에 페이지 문자를 어떻게 인코딩할지 알려주세요.

`meta` 요소는 종료 태그가 필요 없는 빈 요소임을 기억하세요.

# --hints--

`meta` 태그가 있어야 합니다.

```js
assert.match(code, /<meta.*?\/?>/is);
```

`meta` 요소는 빈 요소이므로 종료 태그 `</meta>`가 없어야 합니다.

```js
assert.notMatch(code, /<\/meta>/i);
```

`meta` 태그에 `charset` 속성이 있어야 합니다.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

`charset` 속성은 `utf-8` 값을 가져야 합니다.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
