---
id: 5f3326b143638ee1a09ff1e3
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

실제 콘텐츠를 만들 준비를 하려면 `body` 요소 아래에 `head` 요소를 추가하세요.

# --hints--

열리는 `<body>` 태그가 있어야 합니다.

```js
assert(code.match(/<body>/i));
```

닫는 `</body>` 태그가 있어야 합니다.

```js
assert(code.match(/<\/body>/i));
```

`head` 요소는 변경하지 마세요. 닫는 태그를 삭제하지 않았는지 확인하세요.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

`body` 요소는 `head` 요소 뒤에 와야 합니다.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
