---
id: 5f332b23c2045fb843337579
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

방문자에게 카페가 2020년에 설립되었음을 알리려면 `p` 요소 아래에 `h1` 요소를 추가하고 텍스트로 `Est. 2020`를 넣으세요.

# --hints--

열리는 `<p>` 태그가 있어야 합니다.

```js
assert(code.match(/<p>/i));
```

닫는 `</p>` 태그가 있어야 합니다.

```js
assert(code.match(/<\/p>/i));
```

기존 `h1` 요소를 변경하지 마세요. 닫는 태그를 삭제하지 않았는지 확인하세요.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

`p` 요소는 `h1` 요소 아래에 있어야 합니다.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

`p` 요소에는 텍스트 `Est. 2020`가 있어야 합니다.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```
