---
id: 5f3313e74582ad9d063e3a38
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`head` 요소 안에 `html` 요소를 추가해서 `title` 요소를 넣을 수 있게 하세요. `title` 요소의 텍스트는 `Cafe Menu` 여야 합니다.

# --hints--

열린 `<head>` 태그가 있어야 합니다.

```js
assert(code.match(/<head>/i));
```

닫힌 `</head>` 태그가 있어야 합니다.

```js
assert(code.match(/<head>/i));
```

열린 `<title>` 태그가 있어야 합니다.

```js
assert(code.match(/<title>/i));
```

닫힌 `</title>` 태그가 있어야 합니다.

```js
assert(code.match(/<\/title>/i));
```

`title` 요소가 `head` 요소 안에 중첩되어 있어야 합니다.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

`title` 요소의 텍스트는 `Cafe Menu` 여야 합니다. 철자를 확인해야 할 수도 있습니다.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
