---
id: 5f33071498eb2472b87ddee4
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 프로젝트에서는 카페 메뉴를 만들면서 CSS(Cascading Style Sheets)의 기본을 배우게 됩니다. CSS는 HTML 문서의 스타일을 지정하는 언어입니다. HTML 요소가 화면에 어떻게 표시되어야 하는지 설명합니다.

Cat Photo App의 마지막 몇 단계에서 배운 것처럼, 웹 페이지를 만들기 위해 필요한 기본 구조가 있습니다. 모든 HTML 문서에는 `DOCTYPE` 선언과 `html` 요소가 있어야 합니다. `DOCTYPE`는 브라우저에 문서가 어떤 HTML 버전인지 알려줍니다. 그리고 `html` 요소는 모든 다른 요소를 포함하는 루트 요소를 나타냅니다.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

`<!DOCTYPE html>` 태그와 `html` 속성이 `lang`인 `en` 요소를 추가하세요.

# --hints--

`<!DOCTYPE html>` 선언이 있어야 합니다.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

언어가 영어로 설정된 `<html lang="en">` 여는 태그가 있어야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

닫는 `</html>` 태그가 있어야 합니다.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
