---
id: 5f33294a6af5e9188dbdb8f3
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

메뉴 내용을 추가할 시간입니다. 기존 `main` 요소 안에 `body` 요소를 추가하세요. 이 요소는 나중에 카페에서 제공하는 커피와 디저트의 가격 정보를 담게 됩니다.

# --hints--

코드에 `<main>` 여는 태그가 있어야 합니다.

```js
assert(code.match(/<main>/i));
```

코드에 `</main>` 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/main>/i));
```

`body` 요소는 변경하지 마세요. 닫는 태그를 실수로 삭제하지 않도록 주의하세요.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

`main` 태그는 `body` 태그 안에 있어야 합니다.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
