---
id: 5f3477cb2e27333b1ab2b955
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

이제 `styles.css` 파일을 연결해야 스타일이 다시 적용됩니다. `head` 요소 안에 `link` 요소를 추가하세요. `rel` 속성을 `"stylesheet"` 값으로, `href` 속성을 `"styles.css"` 값으로 지정하세요.

`link` 요소가 빈 요소라는 점을 참고하세요.

# --hints--

코드에 `link` 요소가 있어야 합니다.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

`link` 요소는 빈 요소이므로 닫는 태그 `</link>`가 없어야 합니다.

```js
assert.notMatch(code, /<\/link>/i);
```

기존 `head` 요소를 변경하지 마세요. 닫는 태그를 삭제하지 않았는지 확인하세요.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

`link` 요소가 하나 있어야 합니다.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

`link` 요소는 `rel` 값을 가진 `stylesheet` 속성을 가져야 합니다.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

`link` 요소는 `href` 값을 가진 `styles.css` 속성을 가져야 합니다.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
