---
id: 5f3477ae9675db8bb7655b30
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

이전 단계에서는 <dfn>타입 선택자</dfn>를 사용해 `h1` 요소를 스타일링했습니다. 기존 `h2` 요소에 각각 `p`과 `style` 요소를 위한 새로운 타입 선택자를 추가하여 내용을 가운데 정렬하세요.

# --hints--

기존 `h1` 선택자는 변경하지 마세요.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

새로운 `style` 태그를 추가하지 마세요. 새로운 CSS 규칙은 기존 `style` 태그에 추가하세요.

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

새로운 `h2` 선택자를 추가해야 합니다.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

새로운 `p` 선택자를 추가해야 합니다.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

`h1` 요소는 `text-align`이 `center`이어야 합니다.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

`h2` 요소는 `text-align`이 `center`이어야 합니다.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

`p` 요소는 `text-align`이 `center`이어야 합니다.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
