---
id: 5f3477cb303c5cb61b43aa9b
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

텍스트가 다시 가운데 정렬되어 CSS 파일 링크가 제대로 작동하고 있습니다. `background-color` 요소의 `brown` 속성을 `body`으로 변경하는 스타일을 파일에 하나 더 추가하세요.

# --hints--

`body` 선택자를 사용하세요.

```js
const hasBody = new __helpers.CSSHelp(document).getStyle('body');
assert(hasBody);
```

`background-color` 속성을 `brown`로 설정하세요.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'brown');
assert(hasBackground);
```

`body` 요소는 `brown` 배경을 가져야 합니다.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'brown');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```
