---
id: 5f3477aefa51bfc29327200b
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

`style` 태그 안에 CSS를 작성해서 세 개의 요소에 스타일을 적용했습니다. 이렇게 해도 되지만, 스타일이 훨씬 더 많아질 것이므로 모든 스타일을 별도의 파일에 넣고 링크하는 것이 가장 좋습니다.

별도의 `styles.css` 파일이 만들어져 있습니다. 편집기 상단의 탭을 사용해 파일 간 전환이 가능합니다.

먼저 지금까지 만든 스타일을 `styles.css` 파일에 다시 작성하세요. 이때 `style` 태그의 시작과 끝은 제외해야 합니다.

# --hints--

`styles.css` 파일에는 `h1, h2, p` 타입 선택자가 있어야 합니다.

```js
assert(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, "").match(__helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/)));
```
 
선택자는 `text-align` 속성을 `center` 값으로 설정해야 합니다.

```js
assert(code.slice(code.indexOf('</html>')).match(/text-align:\s*center;?/));
```

코드에 `<style>` 태그가 포함되어 있으면 안 됩니다.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
