---
id: 5f34a1fd611d003edeafd681
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

갈색 배경 때문에 텍스트를 읽기 어렵습니다. `body` 요소의 배경색을 `burlywood`으로 변경해서 색은 있지만 텍스트를 읽을 수 있게 하세요.

# --hints--

`background-color` 속성을 `burlywood`으로 설정하세요.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'burlywood');
assert(hasBackground);
```

`body` 요소는 `burlywood` 배경을 가져야 합니다.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'burlywood');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: brown;
}
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
```
