---
id: 5f3477ae8466a9a3d2cc953c
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

`styles.css` 파일에 CSS가 있으니, `style` 요소와 그 안의 모든 내용을 제거하세요. 제거하면 가운데 정렬된 텍스트가 다시 왼쪽으로 이동합니다.

# --hints--

코드에 `style` 태그가 없어야 합니다.

```js
assert(!code.match(/style/i));
```

HTML 파일에 CSS 선택자가 없어야 합니다.

```js
assert(!editableContents.includes('style'));
assert(!editableContents.includes('text-align'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
