---
id: 5f344fc1520b6719f2e35605
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

메뉴에는 커피용 섹션과 디저트용 섹션 두 개가 있을 것입니다. 모든 커피를 넣을 공간을 마련하려면 `section` 요소 안에 `main` 요소를 추가하세요.

# --hints--

`<section>` 여는 태그가 있어야 합니다.

```js
assert(code.match(/<section\s*>/i));
```

`</section>` 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/section\s*>/i));
```

기존 `main` 요소를 변경하지 마세요. 닫는 태그를 삭제하지 않았는지 확인하세요.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

`section` 요소는 `main` 요소 안에 있어야 합니다.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
