---
id: 5f356ed656a336993abd9f7c
title: 25단계
challengeType: 0
dashedName: step-25
---

# --description--

다음으로, `#menu`를 수평으로 가운데에 배치하세요. `margin-left`과 `margin-right` 속성을 `auto`로 설정하면 됩니다. 여백은 요소 주변의 보이지 않는 공간이라고 생각하세요. 이 두 여백 속성을 사용해 `#menu` 요소 안에서 `body` 요소를 가운데에 배치하세요.

# --hints--

`margin-left` 속성을 `auto`로 설정하세요.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-left'] === 'auto');
assert(hasMargin);
```

`margin-right` 속성을 `auto`로 설정하세요.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-right'] === 'auto');
assert(hasMargin);
```

`margin-left`과 `margin-right` 속성을 `#menu`로 `auto`에 설정하세요.

```js
const divMarginRight = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-right');
const divMarginLeft = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-left');
assert(divMarginRight === 'auto');
assert(divMarginLeft === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
}
--fcc-editable-region--
```
