---
id: 5f356ed63c7807a4f1e6d054
title: 21단계
challengeType: 0
dashedName: step-21
---

# --description--

지금 목표는 `div`가 페이지 전체 너비를 차지하지 않도록 만드는 것입니다. CSS `width` 속성이 이 목적에 딱 맞습니다.

`id` 선택자를 사용해 특정 요소를 `id` 속성으로 지정할 수 있습니다. `id` 선택자는 해시 기호 `#`를 요소의 `id` 값 바로 앞에 붙여 정의합니다. 예를 들어, 요소에 `id`가 `cat`라면 다음과 같이 해당 요소를 지정할 수 있습니다:

```css 
#cat {
  width: 250px;
}
```

`#menu` 선택자를 사용해 요소의 너비를 `300px`로 지정하세요.

# --hints--

`#menu` 선택자가 있어야 합니다.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

`width` 속성을 `300px`로 설정해야 합니다.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

`div`의 너비는 300px이어야 합니다.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```
