---
id: 5f356ed60785e1f3e9850b6e
title: 24단계
challengeType: 0
dashedName: step-24
---

# --description--

이제 텍스트가 `#menu` 요소 안에서 가운데 정렬된 것을 쉽게 알 수 있습니다. 현재 `#menu` 요소의 너비는 픽셀(`px`) 단위로 지정되어 있습니다.

`width` 속성의 값을 `80%`로 변경하여 부모 요소(`body`) 너비의 80%가 되도록 하세요.

# --hints--

`width` 속성을 `80%`로 설정하세요.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '80%');
assert(hasWidth);
```

`width` 속성을 `300px`로 지정하면 안 됩니다.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(!hasWidth);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--
#menu {
  width: 300px;
  background-color: burlywood;
}
--fcc-editable-region--
```
