---
id: 5f356ed6199b0cdef1d2be8f
title: 26단계
challengeType: 0
dashedName: step-26
---

# --description--

지금까지는 요소를 스타일링할 때 type 선택자와 id 선택자를 사용했습니다. 하지만 요소를 스타일링할 때는 보통 다른 선택자를 더 자주 사용합니다.

<dfn>클래스 선택자</dfn>는 이름 앞에 점(.)이 바로 붙는 형태로 정의됩니다. 예를 들어 다음과 같습니다:

```css
.class-name {
  styles
}
```

기존의 `#menu` 선택자를 `#menu`을 `.menu`라는 이름의 클래스로 바꾸어 클래스 선택자로 변경하세요.

# --hints--

`.menu` 클래스 선택자가 있어야 합니다.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

`#menu` 선택자는 없어야 합니다.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```
