---
id: 5f3cade9fa77275d9f4efe62
title: 40단계
challengeType: 0
dashedName: step-40
---

# --description--

잘 작동했지만 가격 오른쪽에 아직 약간의 공간이 있습니다.

너비에 대해 다양한 백분율을 계속 시도할 수도 있습니다. 대신, 백스페이스 키를 사용해 클래스가 `p`인 `price` 요소를 클래스가 `p`인 `flavor` 요소 옆으로 옮겨 편집기에서 두 요소가 같은 줄에 있도록 하세요. 두 요소 사이에 공간이 없도록 확인하세요.

# --hints--

`p` 요소들 사이에 공간이 없어야 합니다.

```js
assert(code.match(/Vanilla<\/p><p/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
--fcc-editable-region--
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
--fcc-editable-region--
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 49%;
}

.price {
  text-align: right;
  width: 49%;
}
```
