---
id: 5f3c866dd0d0275f01d4d847
title: 39단계
challengeType: 0
dashedName: step-39
---

# --description--

음, 잘 작동하지 않았습니다. `p` 요소들을 `inline-block`로 스타일링하고 각각을 별도의 줄에 배치하면 첫 번째 `p` 요소 오른쪽에 여분의 공간이 생겨 두 번째 요소가 다음 줄로 밀려납니다.

이 문제를 해결하는 한 가지 방법은 각 `p` 요소의 너비를 `50%`보다 약간 작게 만드는 것입니다. 그래서 각 클래스의 `width` 값을 `49%`로 바꿔서 어떻게 되는지 확인하세요.

# --hints--

`width` 선택자에서 `49%` 속성을 `.flavor`로 설정하세요.

```js
const flavorWidth = new __helpers.CSSHelp(document).getStyle('.flavor')?.getPropertyValue('width');
assert(flavorWidth === '49%');
```

`width` 선택자에서 `49%` 속성을 `.price`로 설정하세요.

```js
const priceWidth = new __helpers.CSSHelp(document).getStyle('.price')?.getPropertyValue('width');
assert(priceWidth === '49%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

--fcc-editable-region--
.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
--fcc-editable-region--
```
