---
id: 5f3c866daec9a49519871816
title: 30단계
challengeType: 0
dashedName: step-30
---

# --description--

`article` 요소는 보통 관련 정보를 가진 여러 요소를 포함합니다. 이 경우, 커피 맛과 그 맛의 가격을 포함할 것입니다.

`p` 요소 안에 두 개의 `article` 요소를 중첩하세요. 첫 번째 요소의 텍스트는 `French Vanilla`이어야 하고, 두 번째 요소의 텍스트는 `3.00`이어야 합니다.

# --hints--

기존 `article` 요소는 변경하지 마세요.

```js
assert.lengthOf(document.querySelectorAll('article'),1);
```

`article` 요소 안에 두 개의 `p` 요소가 있어야 합니다.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphChildren,2);
```

첫 번째 `p` 요소의 텍스트는 `French Vanilla`이어야 합니다.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const firstP = paragraphChildren?.[0];
assert.match(firstP?.innerText.trim(), /French Vanilla/i);
```

두 번째 `p` 요소의 텍스트는 `3.00`이어야 합니다.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const secondP = paragraphChildren?.[1];
assert.match(secondP?.innerText.trim(), /3\.00/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
