---
id: 5f769702e6e33127d14aa120
title: 35단계
challengeType: 0
dashedName: step-35
---

# --description--

이제 `right` 클래스를 가진 요소들의 텍스트를 `price`에 맞춰 정렬하세요.

# --hints--

`price` 클래스 선택자가 있어야 합니다.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('.price');
assert.exists(hasSelector);
```

`price` 클래스 선택자는 `text-align` 속성을 `right`로 설정해야 합니다.

```js
const priceSelector = new __helpers.CSSHelp(document).getStyle('.price');
assert.equal(priceSelector?.getPropertyValue('text-align'), 'right');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

--fcc-editable-region--

--fcc-editable-region--

```
