---
id: 681cb1b0dab50c87ddb2e518
title: "Завдання 7: цільова сума"
challengeType: 29
dashedName: challenge-7
---

# --description--

Дано масив чисел і ціле число ціль. Знайдіть два унікальні числа в масиві, які у сумі дорівнюють цільовому значенню. Поверніть масив з індексами цих двох чисел або `"Target not found"`, якщо жодні два числа не дають у сумі ціль.

- Повернутий масив повинен містити індекси у зростаючому порядку.

# --hints--

`find_target([2, 7, 11, 15], 9)` має повертати `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([2, 7, 11, 15], 9), [0, 1])`)
}})
```

`find_target([3, 2, 4, 5], 6)` має повертати `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([3, 2, 4, 5], 6), [1, 2])`)
}})
```

`find_target([1, 3, 5, 6, 7, 8], 15)` має повертати `[4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 6, 7, 8], 15), [4, 5])`)
}})
```

`find_target([1, 3, 5, 7], 14)` має повертати `'Target not found'`.

```js

({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 7], 14), "Target not found")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_target(arr, target):

    return arr
```

# --solutions--

```py
def find_target(arr, target):
    for i in range(len(arr)):
        for j in range(i + 1, len(arr)):
            if arr[i] + arr[j] == target:
                return [i, j]
    return 'Target not found'
```
