---
id: 681cb1b0dab50c87ddb2e51b
title: "Завдання 10: 3 страйки"
challengeType: 29
dashedName: challenge-10
---

# --description--

Дано ціле число від 1 до 10 000. Поверніть кількість чисел від 1 до цього числа, квадрат яких містить принаймні одну цифру 3.

# --hints--

`squares_with_three(1)` має повертати `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(1), 0)`)
}})
```

`squares_with_three(10)` має повертати `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(10), 1)`)
}})
```

`squares_with_three(100)` має повертати `19`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(100), 19)`)
}})
```

`squares_with_three(1000)` має повертати `326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(1000), 326)`)
}})
```

`squares_with_three(10000)` має повертати `4531`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(10000), 4531)`)
}})
```

# --seed--

## --seed-contents--

```py
def squares_with_three(n):

    return n
```

# --solutions--

```py
def squares_with_three(n):
    count = 0
    for i in range(1, n + 1):
        square = i * i
        if '3' in str(square):
            count += 1
    return count
```
