---
id: 681cb1afdab50c87ddb2e516
title: "Завдання 5: jbelmud текст"
challengeType: 29
dashedName: challenge-5
---

# --description--

Дано рядок. Поверніть заплутану версію цього рядка, де кожне слово трансформується за такими правилами:

- Перша і остання букви слова залишаються на своїх місцях
- Всі букви між першою та останньою сортуються в алфавітному порядку.
- Вхідні рядки не міститимуть розділових знаків і будуть повністю в нижньому регістрі.

# --hints--

`jbelmu("hello world")` має повертати `"hello wlord"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("hello world"), "hello wlord")`)
}})
```

`jbelmu("i love jumbled text")` має повертати `"i love jbelmud text"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("i love jumbled text"), "i love jbelmud text")`)
}})
```

`jbelmu("freecodecamp is my favorite place to learn to code")` має повертати `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe")`)
}})
```

`jbelmu("the quick brown fox jumps over the lazy dog")` має повертати `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def jbelmu(text):

    return text
```

# --solutions--

```py
def jbelmu(text):
    words = text.split()
    jumbled = []

    for word in words:
        if len(word) <= 3:
            jumbled.append(word)
        else:
            first = word[0]
            last = word[-1]
            middle = ''.join(sorted(word[1:-1]))
            jumbled.append(first + middle + last)
    
    return ' '.join(jumbled)
```
