---
id: 681cb1afdab50c87ddb2e517
title: "Завдання 6: перевірка анаграми"
challengeType: 29
dashedName: challenge-6
---

# --description--

Дано два рядки, визначте, чи є вони анаграмами один одного (містять ті самі символи в будь-якому порядку).

- Ігноруйте регістр і порожній простір.

# --hints--

`are_anagrams("listen", "silent")` має повертати `true`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("listen", "silent"), True)`)
}})
```

`are_anagrams("School master", "The classroom")` має повертати `true`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("School master", "The classroom"), True)`)
}})
```

`are_anagrams("A gentleman", "Elegant man")` має повертати `true`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("A gentleman", "Elegant man"), True)`)
}})
```

`are_anagrams("Hello", "World")` має повертати `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("Hello", "World"), False)`)
}})
```

`are_anagrams("apple", "banana")` має повертати `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("apple", "banana"), False)`)
}})
```

`are_anagrams("cat", "dog")` має повертати `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("cat", "dog"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def are_anagrams(str1, str2):

    return str1
```

# --solutions--

```py
def are_anagrams(str1, str2):
    def clean(s):
        return sorted(s.replace(" ", "").lower())
    
    return clean(str1) == clean(str2)
```
