---
id: 681cb1a2dab50c87ddb2e514
title: "Завдання 3: послідовність Фібоначчі"
challengeType: 29
dashedName: challenge-3
---

# --description--

Послідовність Фібоначчі — це ряд чисел, де кожне число є сумою двох попередніх. Починаючи з `0` і `1`, перші 10 чисел у послідовності — це `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

Дано масив, що містить перші два числа послідовності Фібоначчі, та ціле число, яке представляє довжину послідовності. Поверніть масив, що містить послідовність заданої довжини.

- Ваша функція має працювати з послідовностями будь-якої довжини, більшої або рівної нулю.
- Якщо довжина дорівнює нулю, поверніть порожній масив.
- Зверніть увагу, що початкові числа є частиною послідовності.

# --hints--


`fibonacci_sequence([0, 1], 20)` має повернути `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181])`)
}})
```

`fibonacci_sequence([21, 32], 1)` має повернути `[21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([21, 32], 1), [21])`)
}})
```

`fibonacci_sequence([0, 1], 0)` має повернути `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 0), [])`)
}})
```

`fibonacci_sequence([10, 20], 2)` має повернути `[10, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([10, 20], 2), [10, 20])`)
}})
```

`fibonacci_sequence([123456789, 987654321], 5)` має повернути `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541])`)
}})
```

# --seed--

## --seed-contents--

```py
def fibonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def fibonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2]
        sequence.append(next_value)

    return sequence
```
