---
id: 681cb1afdab50c87ddb2e515
title: "Завдання 4: S  P  A  C  E  J  A  M"
challengeType: 29
dashedName: challenge-4
---

# --description--

Дано рядок. Видаліть усі пробіли з рядка, вставте по два пробіли між кожним символом, перетворіть усі літери алфавіту на верхній регістр і поверніть результат.

- Неалфавітні символи мають залишатися без змін (окрім пробілів).

# --hints--

`space_jam("freeCodeCamp")` має повернути `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("   free   Code   Camp   ")` має повернути `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("Hello World?!")` має повернути `"H  E  L  L  O  W  O  R  L  D  ?  !"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !")`)
}})
```

`space_jam("C@t$ & D0g$")` має повернути `"C  @  T  $  &  D  0  G  $"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $")`)
}})
```

`space_jam("allyourbase")` має повернути `"A  L  L  Y  O  U  R  B  A  S  E"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E")`)
}})
```

# --seed--

## --seed-contents--

```py
def space_jam(s):

    return s
```

# --solutions--

```py
def space_jam(s):
    s = s.replace(" ", "")
    s = "  ".join(s)
    return s.upper()
```
