---
id: 6821ebce237de8297eaee790
title: "Завдання 14: битва символів"
challengeType: 29
dashedName: challenge-14
---

# --description--

Дано два рядки, які представляють вашу армію та армію супротивника, кожен символ вашої армії бореться з символом на тій самій позиції армії супротивника за такими правилами:

- Символи `a-z` мають відповідно силу `1-26`.
- Символи `A-Z` мають відповідно силу `27-52`.
- Цифри `0-9` мають силу, що дорівнює їхньому значенню.
- Всі інші символи мають значення нуль.
- Кожен символ може брати участь лише в одній битві.

Для кожної битви перемагає сильніший символ. Армія з більшою кількістю перемог виграє війну. Поверніть такі значення:

- `"Opponent retreated"`, якщо у вашій армії більше символів, ніж у армії супротивника.
- `"We retreated"`, якщо у армії супротивника більше символів, ніж у вашій.
- `"We won"`, якщо ваша армія виграла більше битв.
- `"We lost"`, якщо армія супротивника виграла більше битв.
- `"It was a tie"`, якщо обидві армії виграли однакову кількість битв.

# --hints--

`battle("Hello", "World")` має повернути `"We lost"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` має повернути `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` має повернути `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` має повернути `"Opponent retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` має повернути `"We retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` має повернути `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` має повернути `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```
