---
id: 6821ebda237de8297eaee792
title: "Завдання 16: перевернути круглі дужки"
challengeType: 29
dashedName: challenge-16
---

# --description--

Дано рядок, який містить правильно вкладені круглі дужки. Поверніть декодовану версію рядка, використовуючи такі правила:

- Усі символи всередині кожної пари круглих дужок потрібно перевернути.
- Круглі дужки слід видалити з фінального результату.
- Якщо дужки вкладені, спочатку потрібно перевернути найглибшу пару, а потім її результат включити у перевертання зовнішньої пари.
- Припускайте, що всі круглі дужки збалансовані та правильно вкладені.

# --hints--

`decode("(f(b(dc)e)a)")` має повернути `"abcdef"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` має повернути `"Can you read this?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` має повернути `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```
