---
id: 6821ebfd237de8297eaee799
title: "Завдання 23: RGB у hex"
challengeType: 29
dashedName: challenge-23
---

# --description--

Дано рядок кольору CSS `rgb(r, g, b)`, поверніть його еквівалент у шістнадцятковому форматі.

Ось кілька прикладів вихідних даних для заданих вхідних:

| Вхідні дані   | Вихідні дані   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- Зробіть усі літери у нижньому регістрі.
- Поверніть `#`, за яким ідуть шість символів. Не використовуйте скорочені значення.

# --hints--

`rgb_to_hex("rgb(255, 255, 255)")` має повертати `"#ffffff"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(255, 255, 255)"), "#ffffff")`)
}})
```

`rgb_to_hex("rgb(1, 11, 111)")` має повертати `"#010b6f"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(1, 11, 111)"), "#010b6f")`)
}})
```

`rgb_to_hex("rgb(173, 216, 230)")` має повертати `"#add8e6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(173, 216, 230)"), "#add8e6")`)
}})
```

`rgb_to_hex("rgb(79, 123, 201)")` має повертати `"#4f7bc9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(79, 123, 201)"), "#4f7bc9")`)
}})
```

# --seed--

## --seed-contents--

```py
def rgb_to_hex(rgb):

    return rgb
```

# --solutions--

```py
def rgb_to_hex(rgb):
    import re
    match = re.findall(r'\d+', rgb)
    r, g, b = [max(0, min(255, int(x))) for x in match[:3]]
    return f'#{r:02x}{g:02x}{b:02x}'
```
