---
id: 6821ebd4237de8297eaee791
title: "Завдання 15: верблюжий регістр"
challengeType: 29
dashedName: challenge-15
---

# --description--

Дано рядок, поверніть його версію у верблюжому регістрі, використовуючи такі правила:

- Слова в рядковому аргументі розділені одним або кількома символами з наступної множини: пробіл (` `), дефіс (`-`) або підкреслення (`_`). Розглядайте будь-яку послідовність цих символів як розрив слова.
- Перше слово має бути повністю у нижньому регістрі.
- Кожне наступне слово має починатися з великої літери, а решта літер — у нижньому регістрі.
- Всі пробіли та роздільники потрібно видалити.

# --hints--

`to_camel_case("hello world")` має повернути `"helloWorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("hello world"), "helloWorld")`)
}})
```

`to_camel_case("HELLO WORLD")` має повернути `"helloWorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("HELLO WORLD"), "helloWorld")`)
}})
```

`to_camel_case("secret agent-X")` має повернути `"secretAgentX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("secret agent-X"), "secretAgentX")`)
}})
```

`to_camel_case("FREE cODE cAMP")` має повернути `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("FREE cODE cAMP"), "freeCodeCamp")`)
}})
```

`to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` має повернути `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_camel_case(s):

    return s
```

# --solutions--

```py
import re
def to_camel_case(s):
    words = re.split(r'[_\- ]+', s)

    camel = [
        words[0].lower() if words else ''
    ] + [
        word.capitalize() for word in words[1:]
    ]

    return ''.join(camel)
```
