---
id: 6821ebea237de8297eaee795
title: "Завдання 19: Світло свічки"
challengeType: 29
dashedName: challenge-19
---

# --description--

Дано ціле число, що представляє кількість свічок, з якими ви починаєте, і ціле число, що показує, скільки використаних свічок потрібно, щоб створити нову. Поверніть кількість свічок, які ви використаєте після того, як створите і спалите якомога більше.

Наприклад, якщо дано 7 свічок і потрібно 2 використані свічки, щоб зробити нову:

1. Спаліть 7 свічок, щоб отримати 7 залишків,
2. Переробіть 6 залишків у 3 нові свічки (залишився 1 залишок),
3. Спаліть 3 свічки, щоб отримати ще 3 залишки (всього 4),
4. Переробіть 4 залишки у 2 нові свічки,
5. Спаліть 2 свічки, щоб отримати 2 залишки,
6. Переробіть 2 залишки у 1 нову свічку,
7. Спаліть 1 свічку.

У цьому прикладі ви спалите загалом 13 свічок.

# --hints--

`burn_candles(7, 2)` має повертати `13`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(7, 2), 13)`)
}})
```

`burn_candles(10, 5)` має повертати `12`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(10, 5), 12)`)
}})
```

`burn_candles(20, 3)` має повертати `29`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(20, 3), 29)`)
}})
```

`burn_candles(17, 4)` має повертати `22`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(17, 4), 22)`)
}})
```

`burn_candles(2345, 3)` має повертати `3517`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(2345, 3), 3517)`)
}})
```

# --seed--

## --seed-contents--

```py
def burn_candles(candles, leftovers_needed):

    return candles
```

# --solutions--

```py
def burn_candles(candles, leftovers_needed):
    total_burned = 0
    unused_leftovers = 0

    while candles > 0:
        total_burned += candles
        leftovers = candles + unused_leftovers
        candles = leftovers // leftovers_needed
        unused_leftovers = leftovers % leftovers_needed

    return total_burned
```
