---
id: 68216ef80f957572e7c340c5
title: "Завдання 12: декодер повідомлень"
challengeType: 29
dashedName: challenge-12
---

# --description--

Дано рядок секретного повідомлення та ціле число, яке представляє кількість літер, на яку було зсунуто повідомлення для його кодування. Поверніть розкодований рядок.

- Додатне число означає, що повідомлення було зсунуте вперед по алфавіту.
- Від’ємне число означає, що повідомлення було зсунуте назад по алфавіту.
- Регістр має значення, розкодовані символи повинні зберігати регістр своїх закодованих відповідників.
- Неалфавітні символи не слід розкодовувати.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` має повертати `"This is a secret message."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.")`)
}})
```

`decode("Byffi Qilfx!", 20)` має повертати `"Hello World!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Byffi Qilfx!", 20), "Hello World!")`)
}})
```

`decode("Zqd xnt njzx?", -1)` має повертати `"Are you okay?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Zqd xnt njzx?", -1), "Are you okay?")`)
}})
```

`decode("oannLxmnLjvy", 9)` має повертати `"freeCodeCamp"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("oannLxmnLjvy", 9), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message, shift):

    return message
```

# --solutions--

```py
def decode(message, shift):
    decoded_message = []
    for char in message:
        if char.isalpha():
            base = ord('a') if char.islower() else ord('A')
            new_char = chr((ord(char) - base - shift) % 26 + base)
            decoded_message.append(new_char)
        else:
            decoded_message.append(char)
    return ''.join(decoded_message)
```
