---
id: 6821ebdf237de8297eaee793
title: "Завдання 17: порядок операцій без урахування пріоритету"
challengeType: 29
dashedName: challenge-17
---

# --description--

Дано масив цілих чисел і масив рядкових операторів. Застосуйте операції до чисел послідовно зліва направо. Повторюйте операції за потреби, доки не використаєте всі числа. Поверніть фінальний результат.

Наприклад, якщо дано `[1, 2, 3, 4, 5]` і `['+', '*']`, поверніть результат обчислення `1 + 2 * 3 + 4 * 5` зліва направо, ігноруючи стандартний порядок операцій.

- Дійсними операторами є `+`, `-`, `*`, `/` і `%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` має повернути `3`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3)`)
}})
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` має повернути `38`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38)`)
}})
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` має повернути `60`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60)`)
}})
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` має повернути `30`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30)`)
}})
```

`evaluate([33, 11, 29, 13], ['/', '-'])` має повернути `-2`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([33, 11, 29, 13], ['/', '-']), -2)`)
}})
```

# --seed--

## --seed-contents--

```py
def evaluate(numbers, operators):

    return numbers
```

# --solutions--

```py
def do_math(a, b, operator):
    if operator == '+':
        return a + b
    elif operator == '-':
        return a - b
    elif operator == '*':
        return a * b
    elif operator == '/':
        return a / b
    else:
        return a % b

def evaluate(numbers, operators):
    total = numbers[0]

    for i in range(1, len(numbers)):
        operator = operators[(i - 1) % len(operators)]
        total = do_math(total, numbers[i], operator)

    return total
```
