---
id: 68adce01c0e1144d0a902956
title: "Завдання 25: повторювач голосних"
challengeType: 29
dashedName: challenge-25
---

# --description--

Дано рядок, поверніть нову версію рядка, де кожна голосна повторюється на один раз більше, ніж попередня зустрінута голосна. Наприклад, перша голосна у реченні залишається без змін. Друга голосна з’являється двічі поспіль. Третя голосна з’являється тричі поспіль і так далі.

- Літери `a`, `e`, `i`, `o` і `u`, у верхньому чи нижньому регістрі, вважаються голосними.
- Початкова голосна зберігає свій регістр.
- Повторені голосні мають бути в нижньому регістрі.
- Всі символи, що не є голосними, зберігають свій початковий регістр.

# --hints--

`repeat_vowels("hello world")` має повернути `"helloo wooorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` має повернути `"freeeCooodeeeeCaaaaamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` має повернути `"AEeIiiOoooUuuuu"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` має повернути `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
