---
id: 68adce01c0e1144d0a90295c
title: "Завдання 28: аналізатор римських чисел"
challengeType: 29
dashedName: challenge-28
---

# --description--

Дано рядок, що представляє римське число. Поверніть його ціле число.

Римські числа складаються з таких символів і значень:

| Символ | Значення |
|--------|----------|
| I      | 1        |
| V      | 5        |
| X      | 10       |
| L      | 50       |
| C      | 100      |
| D      | 500      |
| M      | 1000     |

- Числа читаються зліва направо. Якщо менше число стоїть перед більшим, його значення віднімається. Інакше значення додаються.

# --hints--

`parse_roman_numeral("III")` має повертати `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("III"), 3)`)
}})
```

`parse_roman_numeral("IV")` має повертати `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("IV"), 4)`)
}})
```

`parse_roman_numeral("XXVI")` має повертати `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("XXVI"), 26)`)
}})
```

`parse_roman_numeral("XCIX")` має повертати `99`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("XCIX"), 99)`)
}})
```

`parse_roman_numeral("CDLX")` має повертати `460`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("CDLX"), 460)`)
}})
```

`parse_roman_numeral("DIV")` має повертати `504`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("DIV"), 504)`)
}})
```

`parse_roman_numeral("MMXXV")` має повертати `2025`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("MMXXV"), 2025)`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_roman_numeral(numeral):

    return numeral
```

# --solutions--

```py
def parse_roman_numeral(numeral):
    roman_map = {
        "I": 1,
        "V": 5,
        "X": 10,
        "L": 50,
        "C": 100,
        "D": 500,
        "M": 1000
    }

    total = 0
    for i in range(len(numeral)):
        current = roman_map[numeral[i]]
        next_val = roman_map[numeral[i + 1]] if i + 1 < len(numeral) else 0

        if current < next_val:
            total -= current
        else:
            total += current

    return total
```
