---
id: 68adce01c0e1144d0a902958
title: "Завдання 26: перевірка IPv4"
challengeType: 29
dashedName: challenge-26
---

# --description--

Дано рядок, визначте, чи є він дійсною IPv4-адресою. Дійсна IPv4-адреса складається з чотирьох цілих чисел, розділених крапками (`.`). Кожне число має відповідати таким умовам:

- Воно знаходиться в діапазоні від 0 до 255 включно.
- Не має провідних нулів (наприклад, 0 дозволено, а 01 — ні).
- Дозволені лише числові символи.

# --hints--

`is_valid_ipv4("192.168.1.1")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.1.1"), True)`)
}})
```

`is_valid_ipv4("0.0.0.0")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("0.0.0.0"), True)`)
}})
```

`is_valid_ipv4("255.01.50.111")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("255.01.50.111"), False)`)
}})
```

`is_valid_ipv4("255.00.50.111")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("255.00.50.111"), False)`)
}})
```

`is_valid_ipv4("256.101.50.115")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("256.101.50.115"), False)`)
}})
```

`is_valid_ipv4("192.168.101.")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.101."), False)`)
}})
```

`is_valid_ipv4("192168145213")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192168145213"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_ipv4(ipv4):

    return ipv4
```

# --solutions--

```py
def is_valid_ipv4(ipv4):
    parts = ipv4.split(".")    

    if len(parts) != 4:
        return False

    for part in parts:
        if not part.isdigit():
            return False

        num = int(part)
        if num < 0 or num > 255:
            return False

        if len(part) > 1 and part.startswith("0"):
            return False

    return True
```
