---
id: 68af0687ef34c76c28ffa54b
title: "Завдання 32: Реверс речення"
challengeType: 29
dashedName: challenge-32
---

# --description--

Дано рядок слів, поверніть новий рядок зі словами у зворотному порядку. Наприклад, перше слово має бути в кінці повернутого рядка, а останнє слово — на початку повернутого рядка.

- У заданому рядку слова можуть бути розділені одним або кількома пробілами.
- Повернутий рядок повинен містити лише один пробіл між словами.

# --hints--

`reverse_sentence("world hello")` має повернути `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("world hello"), "hello world")`)
}})
```

`reverse_sentence("push commit git")` має повернути `"git commit push"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("push commit git"), "git commit push")`)
}})
```

`reverse_sentence("npm  install   apt    sudo")` має повернути `"sudo apt install npm"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("npm  install   apt    sudo"), "sudo apt install npm")`)
}})
```

`reverse_sentence("import    default   function  export")` має повернути `"export function default import"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("import    default   function  export"), "export function default import")`)
}})
```

# --seed--

## --seed-contents--

```py
def reverse_sentence(sentence):

    return sentence
```

# --solutions--

```py
def reverse_sentence(sentence):

   return " ".join(sentence.split()[::-1])
```
