---
id: 68b06e589bf227324381476f
title: "Завдання 35: частота слів"
challengeType: 29
dashedName: challenge-35
---

# --description--

Дано абзац, поверніть масив трьох найчастіше вживаних слів.

- Слова в абзаці будуть розділені пробілами.
- Ігноруйте регістр у заданому абзаці. Наприклад, вважайте `Hello` і `hello` одним і тим самим словом.
- Ігноруйте пунктуацію у заданому абзаці. Пунктуація складається з ком (`,`), крапок (`.`) та знаків оклику (`!`).
- Повернутий масив повинен містити всі слова у нижньому регістрі.
- Повернутий масив повинен бути впорядкований за спаданням, починаючи з найчастіше вживаного слова.

# --hints--

`get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` має повернути `["coding", "python", "in"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

`get_words("I like coding. I like testing. I love debugging!")` має повернути `["i", "like", "coding"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

`get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` має повернути `["debug", "test", "deploy"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```
