---
id: 68b06e589bf2273243814777
title: "Завдання 39: Заповнити бак"
challengeType: 29
dashedName: challenge-39
---

# --description--

Враховуючи розмір паливного бака, поточний рівень палива та ціну за галон, поверніть вартість повного заправлення бака.

- `tankSize` — це загальна місткість бака в галонах.
- `fuelLevel` — це поточна кількість палива в баку в галонах.
- `pricePerGallon` — це вартість одного галона палива.
- Повернене значення має бути округлене до двох десяткових знаків у форматі: `"$d.dd"`.

# --hints--

`cost_to_fill(20, 0, 4.00)` має повертати `"$80.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` має повертати `"$17.50"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` має повертати `"$29.25"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` має повертати `"$0.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` має повертати `"$21.89"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
