---
id: 68b06e589bf2273243814775
title: "Завдання 38: генератор slug"
challengeType: 29
dashedName: challenge-38
---

# --description--

Дано рядок, поверніть URL-дружню версію цього рядка, дотримуючись таких обмежень:

- Усі літери мають бути в нижньому регістрі.
- Усі символи, які не є літерами, числами або пробілами, потрібно видалити.
- Усі пробіли потрібно замінити на URL-кодований пробіл `%20`.
- Послідовні пробіли потрібно замінити на один `%20`.
- Повернутий рядок не повинен починатися або закінчуватися `%20`.

# --hints--

`generate_slug("helloWorld")` має повернути `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("helloWorld"), "helloworld")`)
}})
```

`generate_slug("hello world!")` має повернути `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello world!"), "hello%20world")`)
}})
```

`generate_slug(" hello-world ")` має повернути `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug(" hello-world "), "helloworld")`)
}})
```

`generate_slug("hello  world")` має повернути `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello  world"), "hello%20world")`)
}})
```

`generate_slug("  ?H^3-1*1]0! W[0%R#1]D  ")` має повернути `"h3110%20w0r1d"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_slug(str):

    return str
```

# --solutions--

```py
import re
def generate_slug(s):
    cleaned = re.sub(r'[^a-zA-Z0-9 ]+', '', s)
    cleaned = re.sub(r'\s+', ' ', cleaned).strip()
    return cleaned.lower().replace(' ', '%20')
```
