---
id: 68b1f72371a5ac895ac70a02
title: "Завдання 40: Зберігання фотографій"
challengeType: 29
dashedName: challenge-40
---

# --description--

Дано розмір фотографії в мегабайтах (MB) та обсяг жорсткого диска в гігабайтах (GB). Поверніть кількість фотографій, які жорсткий диск може зберегти, використовуючи такі умови:

- 1 гігабайт дорівнює 1000 мегабайтам.
- Поверніть кількість цілих фотографій, які може зберегти диск.

# --hints--

`number_of_photos(1, 1)` має повертати `1000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(1, 1), 1000)`)
}})
```

`number_of_photos(2, 1)` має повертати `500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(2, 1), 500)`)
}})
```

`number_of_photos(4, 256)` має повертати `64000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(4, 256), 64000)`)
}})
```

`number_of_photos(3.5, 750)` має повертати `214285`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 750), 214285)`)
}})
```

`number_of_photos(3.5, 5.5)` має повертати `1571`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 5.5), 1571)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_photos(photo_size_mb, drive_size_gb):

    return photo_size_mb
```

# --solutions--

```py
def number_of_photos(photo_size_mb, drive_size_gb):
    drive_size_mb = drive_size_gb * 1000
    return drive_size_mb // photo_size_mb
```
