---
id: 68b1f72371a5ac895ac70a04
title: "Завдання 41: зберігання файлів"
challengeType: 29
dashedName: challenge-41
---

# --description--

За заданим розміром файлу, одиницею виміру розміру файлу та обсягом жорсткого диска в гігабайтах (GB) поверніть кількість файлів, які жорсткий диск може зберегти, використовуючи такі обмеження:

- Одиниця виміру розміру файлу може бути байтами (`"B"`), кілобайтами (`"KB"`) або мегабайтами (`"MB"`).
- Поверніть кількість цілих файлів, які помістяться на диску.
- Використовуйте такі перетворення:

| Одиниця | Еквівалент |
|:-------:|:----------:|
| 1 B     |   1 B      |
| 1 KB    |   1000 B   |
| 1 MB    |   1000 KB  |
| 1 GB    |   1000 MB  |

Наприклад, за аргументів `500`, `"KB"` та `1` визначте, скільки файлів розміром 500 KB поміститься на жорсткому диску обсягом 1 GB.

# --hints--

`number_of_files(500, "KB", 1)` має повертати `2000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(500, "KB", 1), 2000)`)
}})
```

`number_of_files(50000, "B", 1)` має повертати `20000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(50000, "B", 1), 20000)`)
}})
```

`number_of_files(5, "MB", 1)` має повертати `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(5, "MB", 1), 200)`)
}})
```

`number_of_files(4096, "B", 1.5)` має повертати `366210`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4096, "B", 1.5), 366210)`)
}})
```

`number_of_files(220.5, "KB", 100)` має повертати `453514`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(220.5, "KB", 100), 453514)`)
}})
```

`number_of_files(4.5, "MB", 750)` має повертати `166666`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4.5, "MB", 750), 166666)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_files(file_size, file_unit, drive_size_gb):

    return file_size
```

# --solutions--

```py
def number_of_files(file_size, file_unit, drive_size_gb):
    drive_size_bytes = drive_size_gb * 1000 * 1000 * 1000

    if file_unit == "B":
        file_size_bytes = file_size
    elif file_unit == "KB":
        file_size_bytes = file_size * 1000
    else:
        file_size_bytes = file_size * 1000 * 1000

    return int(drive_size_bytes // file_size_bytes)
```
