---
id: 68b1f72371a5ac895ac70a08
title: "Завдання 43: цифри проти літер"
challengeType: 29
dashedName: challenge-43
---

# --description--

Дано рядок, поверніть `"digits"`, якщо в рядку більше цифр, ніж літер, `"letters"`, якщо більше літер, ніж цифр, і `"tie"`, якщо кількість цифр і літер однакова.

- Цифри складаються з `0-9`.
- Літери складаються з `a-z` у верхньому або нижньому регістрі.
- Ігноруйте будь-які інші символи.

# --hints--

`digits_or_letters("abc123")` має повертати `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123"), "tie")`)
}})
```

`digits_or_letters("a1b2c3d")` має повертати `"letters"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("a1b2c3d"), "letters")`)
}})
```

`digits_or_letters("1a2b3c4")` має повертати `"digits"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("1a2b3c4"), "digits")`)
}})
```

`digits_or_letters("abc123!@#DEF")` має повертати `"letters"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123!@#DEF"), "letters")`)
}})
```

`digits_or_letters("H3110 W0R1D")` має повертати `"digits"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("H3110 W0R1D"), "digits")`)
}})
```

`digits_or_letters("P455W0RD")` має повертати `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("P455W0RD"), "tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def digits_or_letters(s):

    return s
```

# --solutions--

```py
def digits_or_letters(s):
    digit_count = 0
    letter_count = 0

    for char in s:
        if char.isdigit():
            digit_count += 1
        elif char.isalpha():
            letter_count += 1

    if digit_count > letter_count:
        return "digits"
    if letter_count > digit_count:
        return "letters"
    return "tie"
```
