---
id: 68b1f72371a5ac895ac70a06
title: "Завдання 42: зберігання відео"
challengeType: 29
dashedName: challenge-42
---

# --description--

За заданим розміром відео, одиницею виміру розміру відео, обсягом жорсткого диска та одиницею виміру жорсткого диска поверніть кількість відео, які жорсткий диск може зберегти, використовуючи такі обмеження:

- Одиниця виміру розміру відео може бути байтами (`"B"`), кілобайтами (`"KB"`), мегабайтами (`"MB"`) або гігабайтами (`"GB"`).
- Якщо не вказано одну з одиниць відео, наведених вище, поверніть `"Invalid video unit"`.
- Одиниця виміру обсягу жорсткого диска може бути гігабайтами (`"GB"`) або терабайтами (`"TB"`).
- Якщо не вказано одну з одиниць жорсткого диска, наведених вище, поверніть `"Invalid drive unit"`.
- Поверніть кількість цілих відео, які помістяться на диску.
- Використовуйте такі конверсії:

| Одиниця | Еквівалент |
|:-------:|:----------:|
| 1 B     |   1 B      |
| 1 KB    |   1000 B   |
| 1 MB    |   1000 KB  |
| 1 GB    |   1000 MB  |
| 1 TB    |   1000 GB  |

Наприклад, за аргументами `500`, `"MB"`, `100` та `"GB"` визначте, скільки відео по 500 MB поміститься на жорсткому диску обсягом 100 GB.

# --hints--

`number_of_videos(500, "MB", 100, "GB")` має повернути `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500, "MB", 100, "GB"), 200)`)
}})
```

`number_of_videos(1, "TB", 10, "TB")` має повернути `"Invalid video unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1, "TB", 10, "TB"), "Invalid video unit")`)
}})
```

`number_of_videos(2000, "MB", 100000, "MB")` має повернути `"Invalid drive unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(2000, "MB", 100000, "MB"), "Invalid drive unit")`)
}})
```

`number_of_videos(500000, "KB", 2, "TB")` має повернути `4000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500000, "KB", 2, "TB"), 4000)`)
}})
```

`number_of_videos(1.5, "GB", 2.2, "TB")` має повернути `1466`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1.5, "GB", 2.2, "TB"), 1466)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):

    return video_size
```

# --solutions--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):
    video_units = {"KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000}
    drive_units = {"GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000}

    if video_unit not in video_units:
        return "Invalid video unit"
    if drive_unit not in drive_units:
        return "Invalid drive unit"

    video_bytes = video_size * video_units[video_unit]
    drive_bytes = drive_size * drive_units[drive_unit]

    return int(drive_bytes // video_bytes)
```
