---
id: 68b7687dded630607aceccad
title: "Завдання 46: друге найбільше"
challengeType: 29
dashedName: challenge-46
---

# --description--

Дано масив, поверніть друге найбільше унікальне число.

# --hints--

`second_largest([1, 2, 3, 4])` має повернути `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([1, 2, 3, 4]), 3)`)
}})
```

`second_largest([20, 139, 94, 67, 31])` має повернути `94`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([20, 139, 94, 67, 31]), 94)`)
}})
```

`second_largest([2, 3, 4, 6, 6])` має повернути `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([2, 3, 4, 6, 6]), 4)`)
}})
```

`second_largest([10, -17, 55.5, 44, 91, 0])` має повернути `55.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([10, -17, 55.5, 44, 91, 0]), 55.5)`)
}})
```

`second_largest([1, 0, -1, 0, 1, 0, -1, 1, 0])` має повернути `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([1, 0, -1, 0, 1, 0, -1, 1, 0]), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def second_largest(arr):

    return arr
```

# --solutions--

```py
def second_largest(arr):
    unique = list(set(arr))
    unique.sort(reverse=True)
    return unique[1]
```
