---
id: 68b7cadffed0e75a517da675
title: "Завдання 53: десяткове число в двійкове"
challengeType: 29
dashedName: challenge-53
---

# --description--

Дано невід’ємне ціле число, поверніть його двійкове представлення у вигляді рядка.

Двійкове число використовує лише цифри `0` та `1` для представлення будь-якого числа. Щоб конвертувати десяткове число у двійкове, послідовно діліть число на `2` і записуйте остачу. Повторюйте, доки число не стане нулем. Читайте остачі у зворотному порядку — від останньої записаної до першої. Наприклад, щоб конвертувати `12` у двійкове:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` у двійковому вигляді — це `1100`.

# --hints--

`to_binary(5)` має повернути `"101"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(5), "101")`)
}})
```

`to_binary(12)` має повернути `"1100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(12), "1100")`)
}})
```

`to_binary(50)` має повернути `"110010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(50), "110010")`)
}})
```

`to_binary(99)` має повернути `"1100011"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(99), "1100011")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_binary(decimal):

    return decimal
```

# --solutions--

```py
def to_binary(decimal):
    if decimal == 0:
        return "0"
    binary = ""
    while decimal > 0:
        binary = str(decimal % 2) + binary
        decimal //= 2
    return binary
```
