---
id: 68b7687dded630607aceccb1
title: "Завдання 48: детектор спаму"
challengeType: 29
dashedName: challenge-48
---

# --description--

Дано номер телефону у форматі `"+A (BBB) CCC-DDDD"`, де кожна літера представляє цифру за таким правилом:

- `A` представляє код країни і може містити будь-яку кількість цифр.
- `BBB` представляє код області і завжди містить три цифри.
- `CCC` та `DDDD` представляють місцевий номер і завжди містять відповідно три та чотири цифри.

Визначте, чи є цей номер спамом за такими критеріями:

- Код країни довший за 2 цифри або не починається з нуля (`0`).
- Код області більший за 900 або менший за 200.
- Сума перших трьох цифр місцевого номера зустрічається серед останніх чотирьох цифр місцевого номера.
- У номері одна й та сама цифра повторюється чотири або більше разів поспіль (ігноруючи символи форматування).

# --hints--

`is_spam("+0 (200) 234-0182")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (200) 234-0182"), False)`)
}})
```

`is_spam("+091 (555) 309-1922")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+091 (555) 309-1922"), True)`)
}})
```

`is_spam("+1 (555) 435-4792")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+1 (555) 435-4792"), True)`)
}})
```

`is_spam("+0 (955) 234-4364")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (955) 234-4364"), True)`)
}})
```

`is_spam("+0 (155) 131-6943")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (155) 131-6943"), True)`)
}})
```

`is_spam("+0 (555) 135-0192")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 135-0192"), True)`)
}})
```

`is_spam("+0 (555) 564-1987")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 564-1987"), True)`)
}})
```

`is_spam("+00 (555) 234-0182")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+00 (555) 234-0182"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_spam(number):

    return number
```

# --solutions--

```py
import re
def is_spam(number):
    digits = re.sub(r"\D", "", number)
    match = re.match(r"^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$", number)
    country_code, area_code, ccc, dddd = match.groups()

    if len(country_code) > 2 or not country_code.startswith("0"):
        return True

    area_num = int(area_code)
    if area_num > 900 or area_num < 200:
        return True

    sum_ccc = sum(int(d) for d in ccc)
    if str(sum_ccc) in dddd:
        return True

    if re.search(r"(\d)\1\1\1", digits):
        return True

    return False
```
