---
id: 68b7687dded630607aceccb3
title: "Завдання 49: парсер заголовків CSV"
challengeType: 29
dashedName: challenge-49
---

# --description--

За першою лінією CSV-файлу (значення, розділені комами) поверніть масив, що містить заголовки.

- Перша лінія CSV-файлу містить заголовки, розділені комами.
- Видаліть усі провідні або кінцеві пробіли з кожного заголовка.

# --hints--

`get_headings("name,age,city")` має повернути `["name", "age", "city"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("name,age,city"), ["name", "age", "city"])`)
}})
```

`get_headings("first name,last name,phone")` має повернути `["first name", "last name", "phone"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("first name,last name,phone"), ["first name", "last name", "phone"])`)
}})
```

`get_headings("username , email , signup date ")` має повернути `["username", "email", "signup date"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("username , email , signup date "), ["username", "email", "signup date"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_headings(csv):

    return csv
```

# --solutions--

```py
def get_headings(csv):
    return [h.strip() for h in csv.split(",")]
```
