---
id: 68b7cadffed0e75a517da673
title: "Завдання 52: бінарне число у десяткове"
challengeType: 29
dashedName: challenge-52
---

# --description--

Дано рядок, що представляє бінарне число. Поверніть його десятковий еквівалент у вигляді числа.

Бінарне число використовує лише цифри `0` і `1` для представлення будь-якого числа. Щоб конвертувати бінарне число у десяткове, помножте кожну цифру на ступінь `2` і додайте їх разом. Починайте з множення правої крайньої цифри на `2^0`, наступної зліва — на `2^1` і так далі. Коли всі цифри будуть помножені на ступінь `2`, додайте отримані значення.

Наприклад, бінарне число `101` дорівнює `5` у десятковій системі, тому що:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`to_decimal("101")` має повернути `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("101"), 5)`)
}})
```

`to_decimal("1010")` має повернути `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010"), 10)`)
}})
```

`to_decimal("10010")` має повернути `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("10010"), 18)`)
}})
```

`to_decimal("1010101")` має повернути `85`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010101"), 85)`)
}})
```

# --seed--

## --seed-contents--

```py
def to_decimal(binary):

    return binary
```

# --solutions--

```py
def to_decimal(binary):
    decimal = 0
    for i, bit_char in enumerate(reversed(binary)):
        bit = int(bit_char)
        decimal += bit * (2 ** i)
    return decimal
```
