---
id: 68c1a929005bf54d342aa8d6
title: "Завдання 59: Тиждень космосу, день 5: зона Голділок"
challengeType: 29
dashedName: challenge-59
---

# --description--

На п’ятий день Тижня космосу ви обчислите «зону Голділок» зірки — область навколо зірки, де умови «саме такі», щоб існувала рідка вода.

За заданою масою зірки поверніть масив із початковою та кінцевою відстанями її зони Голділок у астрономічних одиницях.

Щоб обчислити зону Голділок:

1. Знайдіть світність зірки, піднявши її масу у степінь 3.5.
2. Початок зони — це 0.95 від квадратного кореня зі світності.
3. Кінець зони — це 1.37 від квадратного кореня зі світності.

- Поверніть відстані, округлені до двох десяткових знаків.

Наприклад, за маси `1` поверніть `[0.95, 1.37]`.

# --hints--

`goldilocks_zone(1)` має повернути `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` має повернути `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` має повернути `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` має повернути `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` має повернути `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
