---
id: 68c1a929005bf54d342aa8d3
title: "Завдання 56: тиждень космосу, день 2: пошук екзопланет"
challengeType: 29
dashedName: challenge-56
---

# --description--

На другий день тижня космосу вам дають рядок, де кожен символ представляє показник яскравості зірки. Визначте, чи виявили ці показники екзопланету за допомогою методу транзиту. Метод транзиту полягає в тому, що планета проходить перед зіркою, зменшуючи її спостережувану яскравість.

- Показники яскравості складаються лише з символів `0-9` та `A-Z`, де кожен показник відповідає таким числовим значенням:
- Символи `0-9` відповідають рівням яскравості `0-9`.
- Символи `A-Z` відповідають рівням яскравості `10-35`.

Зірка вважається такою, що має екзопланету, якщо будь-який окремий показник менший або дорівнює 80% від середнього значення всіх показників. Наприклад, якщо середня яскравість зірки дорівнює 10, її вважатимуть такою, що має екзопланету, якщо будь-який окремий показник дорівнює 8 або менше.

# --hints--

`has_exoplanet("665544554")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("665544554"), False)`)
}})
```

`has_exoplanet("FGFFCFFGG")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FGFFCFFGG"), True)`)
}})
```

`has_exoplanet("MONOPLONOMONPLNOMPNOMP")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("MONOPLONOMONPLNOMPNOMP"), False)`)
}})
```

`has_exoplanet("FREECODECAMP")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FREECODECAMP"), True)`)
}})
```

`has_exoplanet("9AB98AB9BC98A")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("9AB98AB9BC98A"), False)`)
}})
```

`has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_exoplanet(readings):

    return readings
```

# --solutions--

```py
def has_exoplanet(readings):
    values = [int(c, 36) for c in readings]
    average = sum(values) / len(values)
    threshold = average * 0.8
    for v in values:
        if v <= threshold:
            return True
    return False
```
