---
id: 68c1a929005bf54d342aa8d4
title: "Завдання 57: тиждень космосу, день 3: телефон додому"
challengeType: 29
dashedName: challenge-57
---

# --description--

На третій день тижня космосу вам надано масив чисел, які представляють відстані (у кілометрах) між вами, супутниками та вашою рідною планетою у маршруті зв’язку. Визначте, скільки часу займе повідомлення, надіслане через цей маршрут, щоб досягти планети призначення, враховуючи такі умови:

- Перше значення в масиві — це відстань від вашого місця розташування до першого супутника.
- Кожне наступне значення, крім останнього, — це відстань до наступного супутника.
- Останнє значення в масиві — це відстань від попереднього супутника до вашої рідної планети.
- Повідомлення рухається зі швидкістю 300 000 км/с.
- Кожен супутник, через який проходить повідомлення, додає затримку передачі 0,5 секунди.
- Поверніть число, округлене до 4 знаків після десяткової крапки, без зайвих нулів у кінці.

# --hints--

`send_message([300000, 300000])` має повернути `2.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([300000, 300000]), 2.5)`)
}})
```

`send_message([384400, 384400])` має повернути `3.0627`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([384400, 384400]), 3.0627)`)
}})
```

`send_message([54600000, 54600000])` має повернути `364.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([54600000, 54600000]), 364.5)`)
}})
```

`send_message([1000000, 500000000, 1000000])` має повернути `1674.3333`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([1000000, 500000000, 1000000]), 1674.3333)`)
}})
```

`send_message([10000, 21339, 50000, 31243, 10000])` має повернути `2.4086`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([10000, 21339, 50000, 31243, 10000]), 2.4086)`)
}})
```

`send_message([802101, 725994, 112808, 3625770, 481239])` має повернути `21.1597`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([802101, 725994, 112808, 3625770, 481239]), 21.1597)`)
}})
```

# --seed--

## --seed-contents--

```py
def send_message(route):

    return route
```

# --solutions--

```py
def send_message(route):
    total_distance = sum(route)
    delay = (len(route) - 1) * 0.5
    time = total_distance / 300_000
    total = time + delay
    return round(total, 4)
```
