---
id: 68c1a929005bf54d342aa8d2
title: "Завдання 55: Тиждень космосу, день 1: класифікація зірок"
challengeType: 29
dashedName: challenge-55
---

# --description--

4 жовтня починається Всесвітній тиждень космосу. Наступні сім днів принесуть вам завдання з кодування на астрономічну тематику.

Для сьогоднішнього завдання вам задано температуру поверхні зірки в Кельвінах (K), і потрібно визначити її зоряну класифікацію на основі таких діапазонів:

- `"O"`: 30 000 K або вище
- `"B"`: 10 000 K - 29 999 K
- `"A"`: 7 500 K - 9 999 K
- `"F"`: 6 000 K - 7 499 K
- `"G"`: 5 200 K - 5 999 K
- `"K"`: 3 700 K - 5 199 K
- `"M"`: 0 K - 3 699 K

- Поверніть класифікацію заданої зірки.

# --hints--

`classification(5778)` має повертати `"G"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(5778), "G")`)
}})
```

`classification(2400)` має повертати `"M"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(2400), "M")`)
}})
```

`classification(9999)` має повертати `"A"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(9999), "A")`)
}})
```

`classification(3700)` має повертати `"K"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(3700), "K")`)
}})
```

`classification(3699)` має повертати `"M"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(3699), "M")`)
}})
```

`classification(210000)` має повертати `"O"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(210000), "O")`)
}})
```

`classification(6000)` має повертати `"F"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(6000), "F")`)
}})
```

`classification(11432)` має повертати `"B"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(11432), "B")`)
}})
```

# --seed--

## --seed-contents--

```py
def classification(temp):

    return temp
```

# --solutions--

```py
def classification(temp):
    if temp >= 30000:
        return "O"
    elif temp >= 10000:
        return "B"
    elif temp >= 7500:
        return "A"
    elif temp >= 6000:
        return "F"
    elif temp >= 5200:
        return "G"
    elif temp >= 3700:
        return "K"
    else:
        return "M"
```
