---
id: 68c1a929005bf54d342aa8d5
title: "Завдання 58: тиждень космосу, день 4: місце посадки"
challengeType: 29
dashedName: challenge-58
---

# --description--

У четвертий день тижня космосу вам дають матрицю чисел (масив масивів), що представляє потенційні місця посадки для вашого ровера. Знайдіть найбезпечніше місце посадки за такими правилами:

- Кожна клітинка матриці містить число від `0-9`, включно.
- Будь-яке `0` представляє потенційне місце посадки.
- Будь-яке число, крім `0`, занадто небезпечне для посадки. Чим більше число, тим небезпечніше.
- Найбезпечніше місце визначається як клітинка `0`, навколо якої (до 4 сусідів, діагоналі ігноруємо) сума небезпеки найменша.
- Ігноруйте сусідів за межами матриці (кути та краї мають менше сусідів).
- Поверніть індекси найбезпечнішого місця посадки. Завжди буде лише одне таке місце.

Наприклад, для:

```js
[
  [1, 0],
  [2, 0]
]
```

Поверніть `[0, 1]`, індекси `0` у першому масиві.

# --hints--

`find_landing_spot([[1, 0], [2, 0]])` має повернути `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 0], [2, 0]]), [0, 1])`)
}})
```

`find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` має повернути `[1, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1])`)
}})
```

`find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` має повернути `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2])`)
}})
```

`find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` має повернути `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_landing_spot(matrix):

    return matrix
```

# --solutions--

```py
def find_landing_spot(matrix):
    best_spot = None
    lowest_neighbor_sum = float('inf')

    for i in range(len(matrix)):
        for j in range(len(matrix[i])):
            if matrix[i][j] == 0:
                current_neighbor_sum = 0

                if i > 0:
                    current_neighbor_sum += matrix[i - 1][j]
                if j < len(matrix[i]) - 1:
                    current_neighbor_sum += matrix[i][j + 1]
                if i < len(matrix) - 1:
                    current_neighbor_sum += matrix[i + 1][j]
                if j > 0:
                    current_neighbor_sum += matrix[i][j - 1]

                if current_neighbor_sum < lowest_neighbor_sum:
                    lowest_neighbor_sum = current_neighbor_sum
                    best_spot = [i, j]

    return best_spot
```
