---
id: 68c497f3aaefc9fd9f1b0e26
title: "Завдання 62: з шістнадцяткової у десяткову"
challengeType: 29
dashedName: challenge-62
---

# --description--

Дано рядок, що представляє шістнадцяткове число (основа 16). Поверніть його десяткове (основа 10) значення як ціле число.

Шістнадцяткова система числення використовує 16 цифр:

- `0-9` представляють значення `0` до `9`.
- `A-F` представляють значення `10` до `15`.

Ось часткова таблиця перетворень:

| Шістнадцяткове | Десяткове |
|:--------------:|:---------:|
| 0              | 0         |
| 1              | 1         |
| ...            | ...       |
| 9              | 9         |
| A              | 10        |
| ...            | ...       |
| F              | 15        |
| 10             | 16        |
| ...            | ...       |
| 9F             | 159       |
| A0             | 160       |
| ...            | ...       |
| FF             | 255       |
| 100            | 256       |

- Рядок міститиме лише символи `0–9` та `A–F`.

# --hints--

`hex_to_decimal("A")` має повертати `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A"), 10)`)
}})
```

`hex_to_decimal("15")` має повертати `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("15"), 21)`)
}})
```

`hex_to_decimal("2E")` має повертати `46`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("2E"), 46)`)
}})
```

`hex_to_decimal("FF")` має повертати `255`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("FF"), 255)`)
}})
```

`hex_to_decimal("A3F")` має повертати `2623`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A3F"), 2623)`)
}})
```

# --seed--

## --seed-contents--

```py
def hex_to_decimal(hex):

    return hex
```

# --solutions--

```py
def hex_to_decimal(hex):

    return int(hex, 16)
```
