---
id: 68c497f3aaefc9fd9f1b0e24
title: "Завдання 60: тиждень космосу, день 6: фаза місяця"
challengeType: 29
dashedName: challenge-60
---

# --description--

На шостий день тижня космосу вам буде надано дату у форматі `"YYYY-MM-DD"`, і потрібно визначити фазу місяця на цей день за такими правилами:

Використовуйте спрощений місячний цикл тривалістю 28 днів, поділений на чотири рівні фази:

- `"New"`: дні з 1 по 7
- `"Waxing"`: дні з 8 по 14
- `"Full"`: дні з 15 по 21
- `"Waning"`: дні з 22 по 28

Після 28-го дня цикл повторюється з 1-го дня — новим місяцем.

- Використовуйте `"2000-01-06"` як референс нового місяця (1-й день циклу) для визначення фази заданого дня.
- Вам не будуть надані дати раніше за референсну дату.
- Поверніть правильну фазу у вигляді рядка.

**Примітка:** 1-й день означає день нового місяця, тобто минуло 0 днів від останнього нового місяця.

# --hints--

`moon_phase("2000-01-12")` має повернути `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` має повернути `"Waxing"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` має повернути `"Full"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` має повернути `"Waning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` має повернути `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```
