---
id: 68c497f3aaefc9fd9f1b0e25
title: "Завдання 61: тиждень космосу, день 7: паливо для запуску"
challengeType: 29
dashedName: challenge-61
---

# --description--

Для фінального дня тижня космосу вам буде задано масу корисного навантаження в кілограмах (кг), яке ви хочете відправити на орбіту. Визначте кількість палива, необхідного для відправлення вашого корисного навантаження на орбіту, використовуючи такі правила:

- Ракетам потрібно 1 кг палива на кожні 5 кг маси, яку вони мають підняти.
- Саме паливо має масу. Тож коли ви додаєте паливо, маса для підйому зростає, що вимагає більше палива, що збільшує масу, і так далі.
- Щоб обчислити загальну кількість палива: почніть з маси корисного навантаження, обчисліть необхідне паливо для нього, додайте це паливо до загальної маси і обчисліть знову. Повторюйте цей процес, доки додаткова кількість палива не стане меншою за 1 кг, тоді зупиніться.
- Ігноруйте масу самої ракети. Обчислюйте лише паливо, необхідне для підйому корисного навантаження та самого палива.

Наприклад, якщо маса корисного навантаження 50 кг, вам знадобиться 10 кг палива для його підйому (корисне навантаження / 5), що збільшує загальну масу до 60 кг, для підйому якої потрібно 12 кг палива (додатково 2 кг), що збільшує загальну масу до 62 кг, для підйому якої потрібно 12,4 кг палива — додатково 0,4 кг — що менше за 1 кг, тому зупиняємося тут. Загальна маса для підйому — 62,4 кг, з яких 50 кг — початкове корисне навантаження, а 12,4 кг — паливо.

- Поверніть кількість палива, округлену до одного десяткового знака.

# --hints--

`launch_fuel(50)` має повернути `12.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(50), 12.4)`)
}})
```

`launch_fuel(500)` має повернути `124.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(500), 124.8)`)
}})
```

`launch_fuel(243)` має повернути `60.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(243), 60.7)`)
}})
```

`launch_fuel(11000)` має повернути `2749.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(11000), 2749.8)`)
}})
```

`launch_fuel(6214)` має повернути `1553.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(6214), 1553.4)`)
}})
```

# --seed--

## --seed-contents--

```py
def launch_fuel(payload):

    return payload
```

# --solutions--

```py
def launch_fuel(payload):
    total_mass = payload
    additional_fuel = total_mass / 5
    total_fuel = additional_fuel

    while additional_fuel >= 1:
        total_mass += additional_fuel
        additional_fuel = (total_mass / 5) - total_fuel
        total_fuel += additional_fuel

    return round(total_fuel, 1)
```
