---
id: 68cae5b538ff798bbd4da007
title: "Завдання 69: Відсутні шкарпетки"
challengeType: 29
dashedName: challenge-69
---

# --description--

Дано ціле число, яке представляє кількість пар шкарпеток, з якими ви почали, та інше ціле число, яке представляє, скільки циклів прання ви пройшли. Поверніть кількість повних пар шкарпеток, які у вас зараз є, використовуючи такі умови:

- Кожні 2 цикли прання ви втрачаєте одну шкарпетку.
- Кожні 3 цикли прання ви знаходите одну відсутню шкарпетку.
- Кожні 5 циклів прання одна шкарпетка зношується і її потрібно викинути.
- Кожні 10 циклів прання ви купуєте пару шкарпеток.
- Загальна кількість шкарпеток ніколи не може бути меншою за нуль.
- Правила можуть накладатися. Наприклад, у 10-му циклі прання ви втратите одну шкарпетку, викинете одну шкарпетку і купите нову пару шкарпеток.
- Поверніть кількість повних пар шкарпеток.

# --hints--

`sock_pairs(2, 5)` має повернути `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(2, 5), 1)`)
}})
```

`sock_pairs(1, 2)` має повернути `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 2), 0)`)
}})
```

`sock_pairs(5, 11)` має повернути `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(5, 11), 4)`)
}})
```

`sock_pairs(6, 25)` має повернути `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(6, 25), 3)`)
}})
```

`sock_pairs(1, 8)` має повернути `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 8), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sock_pairs(pairs, cycles):

    return pairs
```

# --solutions--

```py
def sock_pairs(pairs, cycles):
    socks = pairs * 2

    for i in range(1, cycles + 1):
        if i % 2 == 0:
            socks -= 1
        if i % 3 == 0:
            socks += 1
        if i % 5 == 0:
            socks -= 1
        if i % 10 == 0:
            socks += 2

        if socks < 0:
            socks = 0

    return socks // 2
```
