---
id: 68cae5b538ff798bbd4da009
title: "Завдання 71: калькулятор чайових"
challengeType: 29
dashedName: challenge-71
---

# --description--

Враховуючи ціну вашої їжі та кастомний відсоток чайових, поверніть масив із трьома значеннями чайових: 15%, 20% і кастомною сумою.

- Ціни будуть задані у форматі: `"$N.NN"`.
- Кастомні відсотки чайових будуть задані у такому форматі: `"25%"`.
- Поверніть суми у тому ж форматі `"$N.NN"`, округлені до двох десяткових знаків.

Наприклад, якщо ціна їжі `"$10.00"`, а кастомне значення чайових `"25%"`, поверніть `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculate_tips("$10.00", "25%")` має повернути `["$1.50", "$2.00", "$2.50"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"])`)
}})
```

`calculate_tips("$89.67", "26%")` має повернути `["$13.45", "$17.93", "$23.31"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"])`)
}})
```

`calculate_tips("$19.85", "9%")` має повернути `["$2.98", "$3.97", "$1.79"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_tips(meal_price, custom_tip):

    return meal_price
```

# --solutions--

```py
def calculate_tips(meal_price, custom_tip):
    meal = float(meal_price[1:])
    custom_percent = float(custom_tip[:-1])

    tips_percent = [15, 20, custom_percent]
    tips = [f"${meal * p / 100:.2f}" for p in tips_percent]

    return tips
```
