---
id: 68cae5b538ff798bbd4da001
title: "Завдання 63: битва слів"
challengeType: 29
dashedName: challenge-63
---

# --description--

Дано два речення, які представляють вашу команду та команду суперника, де кожне слово вашої команди змагається з відповідним словом команди суперника. Визначте, яка команда перемагає, використовуючи такі правила:

- У заданих реченнях завжди буде однакова кількість слів.
- Слова розділені одним пробілом і містять лише літери.
- Значення кожного слова — це сума значень його літер.
- Літери `a` до `z` відповідають значенням `1` до `26`. Наприклад, `a` має значення `1`, а `z` — `26`.
- Велика літера подвоює значення літери. Наприклад, `A` має значення `2`, а `Z` — `52`.
- Слова змагаються по черзі: перше слово вашої команди змагається з першим словом команди суперника і так далі.
- Слово перемагає, якщо його значення більше за значення слова суперника.
- Команда з більшою кількістю виграшних слів є переможцем.

Поверніть `"We win"`, якщо ваша команда перемагає, `"We lose"`, якщо ваша команда програє, і `"Draw"`, якщо обидві команди мають однакову кількість перемог.

# --hints--

`battle("hello world", "hello word")` має повертати `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` має повертати `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` має повертати `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` має повертати `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` має повертати `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` має повертати `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` має повертати `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
