---
id: 68cae5b538ff798bbd4da004
title: "Завдання 66: видалення HTML-тегів"
challengeType: 29
dashedName: challenge-66
---

# --description--

Дано рядок HTML-коду, видаліть теги та поверніть чистий текстовий вміст.

- Вхідний рядок міститиме лише дійсний HTML.
- HTML-теги можуть бути вкладеними.
- Видаліть теги та будь-які атрибути.

Наприклад, `'<a href="#">Click here</a>'` має повернути `"Click here"`.

# --hints--

`strip_tags('<a href="#">Click here</a>')` має повернути `"Click here"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

`strip_tags('<p class="center">Hello <b>World</b>!</p>')` має повернути `"Hello World!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

`strip_tags('<img src="cat.jpg" alt="Cat">')` має повернути порожній рядок (`""`).

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

`strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` має повернути `sectionsection`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```
