---
id: 68cae5b538ff798bbd4da00a
title: "Завдання 72: регулятор термостата 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

Враховуючи поточну температуру в кімнаті у Фаренгейтах та цільову температуру у Цельсіях, поверніть рядок, який вказує, як регулювати температуру в кімнаті відповідно до цих умов:

- Поверніть `"Heat: X degrees Fahrenheit"`, якщо поточна температура нижча за цільову. `X` — це кількість градусів у Фаренгейтах, на яку потрібно нагріти кімнату, щоб досягти цілі, округлена до 1 десяткового знака.
- Поверніть `"Cool: X degrees Fahrenheit"`, якщо поточна температура вища за цільову. `X` — це кількість градусів у Фаренгейтах, на яку потрібно охолодити кімнату, щоб досягти цілі, округлена до 1 десяткового знака.
- Поверніть `"Hold"`, якщо поточна температура дорівнює цільовій.

Щоб конвертувати температуру з Цельсія у Фаренгейти, помножте температуру в Цельсіях на 1.8 і додайте 32 до результату (`F = (C * 1.8) + 32`).

# --hints--

`adjust_thermostat(32, 0)` має повертати `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` має повертати `"Heat: 7.0 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` має повертати `"Cool: 7.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` має повертати `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` має повертати `"Heat: 12.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
