---
id: 68cae5b538ff798bbd4da008
title: "Завдання 70: витягування атрибутів HTML"
challengeType: 29
dashedName: challenge-70
---

# --description--

Дано рядок із дійсним HTML-елементом. Поверніть атрибути цього елемента за такими критеріями:

- Вам буде надано лише один елемент.
- Атрибути будуть у форматі: `attribute="value"`.
- Поверніть масив рядків, де кожен рядок містить властивість атрибута та його значення, розділені комою, у такому форматі: `["attribute1, value1", "attribute2, value2"]`.
- Поверніть атрибути в тому порядку, в якому вони задані.
- Якщо атрибути не знайдено, поверніть порожній масив.

# --hints--

`extract_attributes('<span class="red"></span>')` має повертати `["class, red"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<span class="red"></span>'), ["class, red"])`)
}})
```

`extract_attributes('<meta charset="UTF-8" />')` має повертати `["charset, UTF-8"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<meta charset="UTF-8" />'), ["charset, UTF-8"])`)
}})
```

`extract_attributes("<p>Lorem ipsum dolor sit amet</p>")` має повертати `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes("<p>Lorem ipsum dolor sit amet</p>"), [])`)
}})
```

`extract_attributes('<input name="email" type="email" required="true" />')` має повертати `["name, email", "type, email", "required, true"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"])`)
}})
```

`extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>')` має повертати `["id, submit", "class, btn btn-primary"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"])`)
}})
```

# --seed--

## --seed-contents--

```py
def extract_attributes(element):

    return element
```

# --solutions--

```py
import re
def extract_attributes(element):
    pattern = r'([\w-]+)="([^"]*)"'
    matches = re.findall(pattern, element)
    return [f"{attr}, {val}" for attr, val in matches]
```
