---
id: 68d2ba1468508398389487cf
title: "Завдання 74: улюблені пісні"
challengeType: 29
dashedName: challenge-74
---

# --description--

Пам’ятаєте iPod? Першу модель випустили рівно 24 роки тому, 23 жовтня 2001 року.

Дано масив об’єктів пісень, які представляють ваш плейлист iPod. Поверніть масив із назвами двох найчастіше прослуховуваних пісень, причому найпопулярніша пісня має бути першою.

- Кожен об’єкт матиме властивість `"title"` (рядок) і властивість `"plays"` (ціле число).

# --hints--

`favorite_songs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` має повертати `["Sync or Swim", "Earbud Blues"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"])`)
}})
```

`favorite_songs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` має повертати `["Clickwheel Love", "99 Downloads"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"])`)
}})
```

`favorite_songs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` має повертати `["Song B", "Song C"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"])`)
}})
```

# --seed--

## --seed-contents--

```py
def favorite_songs(playlist):

    return playlist
```

# --solutions--

```py
def favorite_songs(playlist):
    sorted_songs = sorted(playlist, key=lambda x: x["plays"], reverse=True)
    return [song["title"] for song in sorted_songs[:2]]
```
