---
id: 68d2ba1468508398389487d0
title: "Завдання 75: прихований скарб"
challengeType: 29
dashedName: challenge-75
---

# --description--

Дано двовимірний масив, що представляє карту дна океану з прихованим скарбом, та масив із координатами (`[row, column]`) для наступного занурення у пошуках скарбу. Поверніть `"Empty"`, `"Found"` або `"Recovered"` за такими правилами:

- У заданому двовимірному масиві буде рівно один не знайдений скарб, який займає кілька клітинок.
- Кожна клітинка в двовимірному масиві міститиме одне з наступних значень:
  - `"-"`: немає скарбу.
  - `"O"`: частина скарбу, яку ще не знайшли.
  - `"X"`: частина скарбу, яку вже знайшли.

- Якщо в місці занурення немає скарбу, поверніть `"Empty"`.
- Якщо в місці занурення знайдено скарб, але принаймні одна інша частина скарбу ще не знайдена, поверніть `"Found"`.
- Якщо в місці занурення знайдена остання не знайдена частина скарбу, поверніть `"Recovered"`.

Наприклад, якщо задано:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

І `[2, 1]` як координати місця занурення, поверніть `"Recovered"`, бо занурення знайшло останню не знайдену частину скарбу.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` має повернути `"Recovered"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered")`)
}})
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` має повернути `"Empty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty")`)
}})
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` має повернути `"Found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` має повернути `"Found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` має повернути `"Recovered"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` має повернути `"Empty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty")`)
}})
```

# --seed--

## --seed-contents--

```py
def dive(map, coordinates):

    return map
```

# --solutions--

```py
def dive(map, coordinates):
    row, col = coordinates
    target = map[row][col]

    if target == "-":
        return "Empty"

    if target == "O":
        for r in range(len(map)):
            for c in range(len(map[r])):
                if (r != row or c != col) and map[r][c] == "O":
                    return "Found"
        return "Recovered"

    return "Found"
```
