---
id: 68d2ba1468508398389487ce
title: "Завдання 73: Говори мудро, ти повинен"
challengeType: 29
dashedName: challenge-73
---

# --description--

Дано речення, поверніть його версію, яка звучить як порада мудрого вчителя, використовуючи такі правила:

- Слова розділені одним пробілом.
- Знайдіть перше входження одного з наступних слів у реченні: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Перемістіть усі слова до і включно з цим словом у кінець речення і:
  - Збережіть порядок слів при переміщенні.
  - Зробіть їх усі в нижньому регістрі.
  - Додайте перед ними кому та пробіл.
- Зробіть першу літеру нового першого слова речення великою.
- Всі задані речення закінчуються одним розділовим знаком. Збережіть оригінальну пунктуацію речення і перемістіть її в кінець нового речення.
- Поверніть нове речення, переконайтеся, що між словами є лише один пробіл і немає пробілів на початку чи в кінці речення.

Наприклад, для `"You must speak wisely."` поверніть `"Speak wisely, you must."`

# --hints--

`wise_speak("You must speak wisely.")` має повернути `"Speak wisely, you must."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You must speak wisely."), "Speak wisely, you must.")`)
}})
```

`wise_speak("You can do it!")` має повернути `"Do it, you can!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You can do it!"), "Do it, you can!")`)
}})
```

`wise_speak("Do you think you will complete this?")` має повернути `"Complete this, do you think you will?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("Do you think you will complete this?"), "Complete this, do you think you will?")`)
}})
```

`wise_speak("All your base are belong to us.")` має повернути `"Belong to us, all your base are."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("All your base are belong to us."), "Belong to us, all your base are.")`)
}})
```

`wise_speak("You have much to learn.")` має повернути `"Much to learn, you have."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You have much to learn."), "Much to learn, you have.")`)
}})
```

# --seed--

## --seed-contents--

```py
def wise_speak(sentence):

    return sentence
```

# --solutions--

```py
def wise_speak(sentence):
    triggers = ["have", "must", "are", "will", "can"]
    punctuation = sentence[-1]
    words = sentence[:-1].split(" ")

    index = next(i for i, w in enumerate(words) if w in triggers)

    to_move = [w.lower() for w in words[:index + 1]]
    remaining = words[index + 1:]

    if remaining:
        remaining[0] = remaining[0][0].upper() + remaining[0][1:]

    return " ".join(remaining) + ", " + " ".join(to_move) + punctuation
```
