---
id: 68d30845cc08266018fc46be
title: "Завдання 78: послідовність цілих чисел"
challengeType: 29
dashedName: challenge-78
---

# --description--

Дано додатне ціле число. Поверніть рядок зі всіма цілими числами від `1` до заданого числа включно, у числовому порядку.

Наприклад, для `5` поверніть `"12345"`.

# --hints--

`sequence(5)` має повернути `"12345"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(5), "12345")`)
}})
```

`sequence(10)` має повернути `"12345678910"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(10), "12345678910")`)
}})
```

`sequence(1)` має повернути `"1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(1), "1")`)
}})
```

`sequence(27)` має повернути `"123456789101112131415161718192021222324252627"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(27), "123456789101112131415161718192021222324252627")`)
}})
```

# --seed--

## --seed-contents--

```py
def sequence(n):

    return n
```

# --solutions--

```py
def sequence(n):
    result = ""
    for i in range(1, n + 1):
        result += str(i)
    return result
```
