---
id: 68d30845cc08266018fc46bc
title: "Завдання 76: доповнююча ДНК"
challengeType: 29
dashedName: challenge-76
---

# --description--

Дано рядок, що представляє послідовність ДНК, поверніть її доповнюючу нитку за такими правилами:

- ДНК складається з літер `"A"`, `"C"`, `"G"` та `"T"`.
- Літери `"A"` та `"T"` доповнюють одна одну.
- Літери `"C"` та `"G"` доповнюють одна одну.

Наприклад, для `"ACGT"` поверніть `"TGCA"`.

# --hints--

`complementary_dna("ACGT")` має повернути `"TGCA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ACGT"), "TGCA")`)
}})
```

`complementary_dna("ATGCGTACGTTAGC")` має повернути `"TACGCATGCAATCG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ATGCGTACGTTAGC"), "TACGCATGCAATCG")`)
}})
```

`complementary_dna("GGCTTACGATCGAAG")` має повернути `"CCGAATGCTAGCTTC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC")`)
}})
```

`complementary_dna("GATCTAGCTAGGCTAGCTAG")` має повернути `"CTAGATCGATCCGATCGATC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC")`)
}})
```

# --seed--

## --seed-contents--

```py
def complementary_dna(strand):

    return strand
```

# --solutions--

```py
def complementary_dna(strand):
    complements = { "A": "T", "T": "A", "C": "G", "G": "C" }
    return "".join(complements[n] for n in strand)
```
