---
id: 68d30845cc08266018fc46bd
title: "Завдання 77: форматування тривалості"
challengeType: 29
dashedName: challenge-77
---

# --description--

Дано ціле число секунд, поверніть рядок, який представляє ту саму тривалість у форматі `"H:MM:SS"`, де `"H"` — це кількість годин, `"MM"` — кількість хвилин, а `"SS"` — кількість секунд. Поверніть час, використовуючи такі правила:

- Секунди: завжди мають бути двозначними.
- Хвилини: слід опускати провідні нулі, якщо вони не потрібні. Використовуйте `"0"`, якщо тривалість менша за одну хвилину.
- Години: включайте лише якщо вони більші за нуль.

# --hints--

`format(500)` має повертати `"8:20"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(500), "8:20")`)
}})
```

`format(4000)` має повертати `"1:06:40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(4000), "1:06:40")`)
}})
```

`format(1)` має повертати `"0:01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(1), "0:01")`)
}})
```

`format(5555)` має повертати `"1:32:35"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(5555), "1:32:35")`)
}})
```

`format(99999)` має повертати `"27:46:39"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(99999), "27:46:39")`)
}})
```

# --seed--

## --seed-contents--

```py
def format(seconds):

    return seconds
```

# --solutions--

```py
def format(seconds):
    h = seconds // 3600
    m = (seconds % 3600) // 60
    s = seconds % 60

    seconds_str = f"{s:02d}"
    minutes_str = str(m)

    if h > 0:
        hours_str = str(h)
        return f"{hours_str}:{minutes_str.zfill(2)}:{seconds_str}"
    else:
        return f"{minutes_str}:{seconds_str}"
```
