---
id: 68d30fc57588d97fd3027b30
title: "Завдання 79: Navigator"
challengeType: 29
dashedName: challenge-79
---

# --description--

28 жовтня 1994 року вийшов Netscape Navigator, який допоміг мільйонам досліджувати ранній веб.

Дано масив команд браузера, які ви виконали у Netscape Navigator. Поверніть поточну сторінку, на якій ви опинилися після виконання всіх команд за такими правилами:

- Ви завжди починаєте зі сторінки `"Home"`, яка не входить до масиву команд.
- Дійсні команди:
  - `"Visit Page"`: де `"Page"` — це назва сторінки, яку ви відвідуєте. Наприклад, `"Visit About"` переводить вас на сторінку `"About"`. Коли ви переходите на нову сторінку, обов’язково видаліть будь-яку історію переходів вперед.
  - `"Back"`: повертає вас на попередню сторінку в історії або залишає на поточній, якщо попередньої немає.
  - `"Forward"`: переводить вас вперед в історії на сторінку, з якої ви прийшли, або залишає на поточній, якщо такої немає.

Наприклад, для `["Visit About Us", "Back", "Forward"]` поверніть `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` має повернути `"About Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Back", "Forward"]), "About Us")`)
}})
```

`navigate(["Forward"])` має повернути `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Forward"]), "Home")`)
}})
```

`navigate(["Back"])` має повернути `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Back"]), "Home")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery"])` має повернути `"Gallery"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery"]), "Gallery")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` має повернути `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home")`)
}})
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` має повернути `"Contact"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact")`)
}})
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` має повернути `"Visit Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate(commands):

    return commands
```

# --solutions--

```py
def navigate(commands):
    history = ["Home"]
    current_index = 0

    for command in commands:
        if command.startswith("Visit "):
            history = history[:current_index + 1]
            history.append(command[6:])
            current_index += 1
        elif command == "Back" and current_index > 0:
            current_index -= 1
        elif command == "Forward" and current_index < len(history) - 1:
            current_index += 1

    return history[current_index]
```
