---
id: 68e39ed6106dac2f0a98fd65
title: "Завдання 83: валідація підпису"
challengeType: 29
dashedName: challenge-83
---

# --description--

Дано рядок повідомлення, рядок секретного ключа та число підпису. Визначте, чи є підпис дійсним, використовуючи цей метод кодування:

- Літери в повідомленні та секретному ключі мають такі значення:
  - `a` до `z` мають відповідні значення `1` до `26`.
  - `A` до `Z` мають відповідні значення `27` до `52`.
- Всі інші символи не мають значення.
- Обчисліть підпис, взявши суму значень повідомлення та суму значень секретного ключа.

Наприклад, для повідомлення `"foo"` та секретного ключа `"bar"` підпис буде `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

Перевірте, чи збігається обчислений підпис із наданим підписом.

# --hints--

`verify("foo", "bar", 57)` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("foo", "bar", 57), True)`)
}})
```

`verify("foo", "bar", 54)` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("foo", "bar", 54), False)`)
}})
```

`verify("freeCodeCamp", "Rocks", 238)` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("freeCodeCamp", "Rocks", 238), True)`)
}})
```

`verify("Is this valid?", "No", 210)` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Is this valid?", "No", 210), False)`)
}})
```

`verify("Is this valid?", "Yes", 233)` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Is this valid?", "Yes", 233), True)`)
}})
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def verify(message, key, signature):

    return message
```

# --solutions--

```py
def verify(message, key, signature):
    def charValue(ch) -> int:
        if 'a' <= ch <= 'z':
            return ord(ch) - ord('a') + 1
        elif 'A' <= ch <= 'Z':
            return ord(ch) - ord('A') + 27
        else:
            return 0

    def compute_sum(s):
        return sum(charValue(ch) for ch in s)

    total = compute_sum(message) + compute_sum(key)
    return total == signature
```
