---
id: 68e39ed6106dac2f0a98fd63
title: "Завдання 81: n-те просте число"
challengeType: 29
dashedName: challenge-81
---

# --description--

Просте число — це додатне ціле число більше за 1, яке ділиться лише на 1 та саме на себе. Перші п’ять простих чисел — `2`, `3`, `5`, `7` і `11`.

Дано додатне ціле число `n`, поверніть `n`-те просте число. Наприклад, для `5` поверніть 5-те просте число: `11`.

# --hints--

`nth_prime(5)` має повернути `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(5), 11)`)
}})
```

`nth_prime(10)` має повернути `29`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(10), 29)`)
}})
```

`nth_prime(16)` має повернути `53`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(16), 53)`)
}})
```

`nth_prime(99)` має повернути `523`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(99), 523)`)
}})
```

`nth_prime(1000)` має повернути `7919`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(1000), 7919)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_prime(n):

    return n
```

# --solutions--

```py
def nth_prime(n):
    primes = []
    num = 2

    while len(primes) < n:
        is_prime = True
        for i in range(2, int(num**0.5) + 1):
            if num % i == 0:
                is_prime = False
                break
        if is_prime:
            primes.append(num)
        num += 1

    return primes[n - 1]
```
