---
id: 68e39ed6106dac2f0a98fd66
title: "Завдання 84: Інфіковані"
challengeType: 29
dashedName: challenge-84
---

# --description--

2 листопада 1988 року було випущено перший великий інтернет-черв’як, який за один день інфікував близько 10% комп’ютерів, підключених до інтернету.

У цьому завданні вам дають кількість днів, що минули з моменту випуску інтернет-черв’яка, і потрібно визначити, скільки комп’ютерів інфіковано за такими правилами:

- У день 0 інфіковано перший комп’ютер.
- Кожного наступного дня кількість інфікованих комп’ютерів подвоюється.
- Кожного 3-го дня після поширення вірусу застосовується патч, який зменшує кількість інфікованих комп’ютерів на 20%. Кількість комп’ютерів, що отримали патч, округлюйте вгору до найближчого цілого числа.

Наприклад, у такі дні:

- День 0: інфіковано 1 комп’ютер.
- День 1: інфіковано 2 комп’ютери.
- День 2: інфіковано 4 комп’ютери.
- День 3: інфіковано 8 комп’ютерів. Потім застосовується патч: 8 інфікованих * 20% = 1,6 патчів. Округлюємо 1,6 вгору до 2. 8 інфікованих - 2 патчі = 6 комп’ютерів інфіковано після 3-го дня.

Поверніть кількість інфікованих комп’ютерів після заданої кількості днів.

# --hints--

`infected(1)` має повернути `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(1), 2)`)
}})
```

`infected(3)` має повернути `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(3), 6)`)
}})
```

`infected(8)` має повернути `152`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(8), 152)`)
}})
```

`infected(17)` має повернути `39808`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(17), 39808)`)
}})
```

`infected(25)` має повернути `5217638`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(25), 5217638)`)
}})
```

# --seed--

## --seed-contents--

```py
def infected(days):

    return days
```

# --solutions--

```py
import math
def infected(days):
    infected = 1
    for day in range(1, days + 1):
        infected *= 2
        if day % 3 == 0:
            patched = math.ceil(infected * 0.2)
            infected -= patched

    return infected
```
