---
id: 68e39ed6106dac2f0a98fd64
title: "Завдання 82: SpOoKy~CaSe"
challengeType: 29
dashedName: challenge-82
---

# --description--

Дано рядок, що представляє ім’я змінної. Перетворіть його на «spooky case» за такими правилами:

- Замініть усі підкреслення (`_`) і дефіси (`-`) на тильду (`~`).
- Зробіть першу літеру рядка та кожну другу після неї великою. Ігноруйте символ тильди при підрахунку. Всі інші літери зробіть у нижньому регістрі.

Наприклад, для `hello_world` поверніть `HeLlO~wOrLd`.

# --hints--

`spookify("hello_world")` має повернути `"HeLlO~wOrLd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("hello_world"), "HeLlO~wOrLd")`)
}})
```

`spookify("Spooky_Case")` має повернути `"SpOoKy~CaSe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("Spooky_Case"), "SpOoKy~CaSe")`)
}})
```

`spookify("TRICK-or-TREAT")` має повернути `"TrIcK~oR~tReAt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt")`)
}})
```

`spookify("c_a-n_d-y_-b-o_w_l")` має повернути `"C~a~N~d~Y~~b~O~w~L"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L")`)
}})
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` має повернути `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS")`)
}})
```

# --seed--

## --seed-contents--

```py
def spookify(boo):

    return boo
```

# --solutions--

```py
def spookify(boo):
    replaced = boo.replace("_", "~").replace("-", "~")
    
    result = []
    capitalize = True

    for char in replaced:
        if char == "~":
            result.append(char)
        else:
            result.append(char.upper() if capitalize else char.lower())
            capitalize = not capitalize

    return "".join(result)
```
