---
id: 68ee9e3066cfd4eb2328e8a8
title: "Завдання 89: підрахунок карт"
challengeType: 29
dashedName: challenge-89
---

# --description--

Стандартна колода карт має 13 унікальних карт у кожній масті. Враховуючи ціле число, що представляє кількість карт, які потрібно вибрати з колоди, поверніть кількість унікальних комбінацій карт, які ви можете вибрати.

- Порядок не має значення. Вибір карти A, а потім карти B — це те саме, що вибір карти B, а потім карти A.

Наприклад, враховуючи `52`, поверніть `1`. Існує лише одна комбінація з 52 карт, яку можна вибрати з колоди у 52 карти. А враховуючи `2`, поверніть `1326`. Існує 1326 комбінацій карт, які ви можете отримати, вибираючи 2 карти з колоди.

# --hints--

`combinations(52)` має повернути `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` має повернути `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` має повернути `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` має повернути `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` має повернути `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` має повернути `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```
