---
id: 68ee9e3066cfd4eb2328e8a4
title: "Завдання 85: фішка для слів"
challengeType: 29
dashedName: challenge-85
---

# --description--

Дано рядок із реченням, поверніть число слів, які є в цьому реченні.

- Слова — це будь-яка послідовність символів, окрім пробілів, і вони розділені одним пробілом.

# --hints--

`count_words("Hello world")` має повернути `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("Hello world"), 2)`)
}})
```

`count_words("The quick brown fox jumps over the lazy dog.")` має повернути `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("The quick brown fox jumps over the lazy dog."), 9)`)
}})
```

`count_words("I like coding challenges!")` має повернути `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("I like coding challenges!"), 4)`)
}})
```

`count_words("Complete the challenge in JavaScript and Python.")` має повернути `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("Complete the challenge in JavaScript and Python."), 7)`)
}})
```

`count_words("The missing semi-colon crashed the entire internet.")` має повернути `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("The missing semi-colon crashed the entire internet."), 7)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_words(sentence):

    return sentence
```

# --solutions--

```py
def count_words(sentence):

    return len(sentence.split(' '))
```
