---
id: 68f6587287ad1f4ad39b0c84
title: "Завдання 98: підрахунок прямокутників"
challengeType: 29
dashedName: challenge-98
---

# --description--

Дано два натуральні числа, які представляють ширину та висоту прямокутника. Визначте, скільки прямокутників з цілими шириною та висотою може поміститися у заданому.

- Рахуйте лише прямокутники з цілою шириною та висотою.

Наприклад, для `1` і `3` поверніть `6`. Три прямокутники 1x1, два прямокутники 1x2 і один прямокутник 1x3.

# --hints--

`count_rectangles(1, 3)` має повернути `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 3), 6)`)
}})
```

`count_rectangles(3, 2)` має повернути `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(3, 2), 18)`)
}})
```

`count_rectangles(1, 2)` має повернути `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 2), 3)`)
}})
```

`count_rectangles(5, 4)` має повернути `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(5, 4), 150)`)
}})
```

`count_rectangles(11, 19)` має повернути `12540`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(11, 19), 12540)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_rectangles(width, height):

    return width
```

# --solutions--

```py
def count_rectangles(width, height):
    total = 0
    for w in range(1, width + 1):
        for h in range(1, height + 1):
            total += (width - w + 1) * (height - h + 1)
    return total
```
