---
id: 68f6587287ad1f4ad39b0c7f
title: "Завдання 93: Голосні та приголосні"
challengeType: 29
dashedName: challenge-93
---

# --description--

Дано рядок, поверніть масив із числом голосних і числом приголосних у цьому рядку.

- Голосні складаються з `a`, `e`, `i`, `o`, `u` у будь-якому регістрі.
- Приголосні — це всі інші літери у будь-якому регістрі.
- Ігноруйте всі символи, що не є літерами.

Наприклад, для `"Hello World"` поверніть `[3, 7]`.

# --hints--

`count("Hello World")` має повернути `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("JavaScript")` має повернути `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("JavaScript"), [3, 7])`)
}})
```

`count("Python")` має повернути `[1, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Python"), [1, 5])`)
}})
```

`count("freeCodeCamp")` має повернути `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("freeCodeCamp"), [5, 7])`)
}})
```

`count("Hello, World!")` має повернути `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("The quick brown fox jumps over the lazy dog.")` має повернути `[11, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

# --seed--

## --seed-contents--

```py
def count(s):

    return s
```

# --solutions--

```py
def count(s):
    vowels = "aeiou"
    consonants = "bcdfghjklmnpqrstvwxyz"
    v = 0
    c = 0

    for char in s:
        ch = char.lower()
        if ch in vowels:
            v += 1
        if ch in consonants:
            c += 1

    return [v, c]
```
