---
id: 68f6587287ad1f4ad39b0c80
title: "Завдання 94: генератор підпису для е-пошти"
challengeType: 29
dashedName: challenge-94
---

# --description--

Дано рядки з іменем людини, посадою та компанією. Поверніть підпис для е-пошти у вигляді одного рядка, використовуючи такі правила:

- Ім’я має йти першим, перед яким ставиться префікс, що залежить від першої літери імені. Для імен, що починаються з (незалежно від регістру):
  - `A-I`: Використовуйте `>>` як префікс.
  - `J-R`: Використовуйте `--` як префікс.
  - `S-Z`: Використовуйте `::` як префікс.
- Після імені має йти кома та пробіл (`, `).
- Посада та компанія мають йти після коми та пробілу, розділені `" at "` (з пробілами навколо).

Наприклад, для `"Quinn Waverly"`, `"Founder and CEO"` і `"TechCo"` поверніть `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generate_signature("Quinn Waverly", "Founder and CEO", "TechCo")` має повернути `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo")`)
}})
```

`generate_signature("Alice Reed", "Engineer", "TechCo")` має повернути `">>Alice Reed, Engineer at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo")`)
}})
```

`generate_signature("Tina Vaughn", "Developer", "example.com")` має повернути `"::Tina Vaughn, Developer at example.com"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com")`)
}})
```

`generate_signature("B. B.", "Product Tester", "AcmeCorp")` має повернути `">>B. B., Product Tester at AcmeCorp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp")`)
}})
```

`generate_signature("windstorm", "Cloud Architect", "Atmospheronics")` має повернути `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_signature(name, title, company):

    return name
```

# --solutions--

```py
def generate_signature(name, title, company):
    first_letter = name[0].upper()
    if first_letter in "ABCDEFGHI":
        prefix = ">>"
    elif first_letter in "JKLMNOPQR":
        prefix = "--"
    else:
        prefix = "::"

    return f"{prefix}{name}, {title} at {company}"
```
