---
id: 68f6587287ad1f4ad39b0c81
title: "Завдання 95: зсув масиву"
challengeType: 29
dashedName: challenge-95
---

# --description--

Дано масив і ціле число, яке вказує, на скільки позицій потрібно зсунути масив. Поверніть зсунений масив.

- Додатне ціле число зсуває масив вліво.
- Від’ємне ціле число зсуває масив вправо.
- Зсув відбувається циклічно.

Наприклад, якщо дано `[1, 2, 3]` і `1`, зсуньте масив на 1 позицію вліво, повернувши `[2, 3, 1]`.

# --hints--

`shift_array([1, 2, 3], 1)` має повернути `[2, 3, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([1, 2, 3], 1), [2, 3, 1])`)
}})
```

`shift_array([1, 2, 3], -1)` має повернути `[3, 1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([1, 2, 3], -1), [3, 1, 2])`)
}})
```

`shift_array(["alpha", "bravo", "charlie"], 5)` має повернути `["charlie", "alpha", "bravo"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"])`)
}})
```

`shift_array(["alpha", "bravo", "charlie"], -11)` має повернути `["bravo", "charlie", "alpha"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"])`)
}})
```

`shift_array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` має повернути `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4])`)
}})
```

# --seed--

## --seed-contents--

```py
def shift_array(arr, n):

    return arr
```

# --solutions--

```py
def shift_array(arr, n):
    length = len(arr)
    n = n % length
    if n < 0:
        n += length

    return arr[n:] + arr[:n]
```
