---
id: 68f6587287ad1f4ad39b0c7c
title: "Завдання 90: обмеження символів"
challengeType: 29
dashedName: challenge-90
---

# --description--

У цьому завданні вам дають рядок і потрібно визначити, чи він поміститься у допис у соціальній мережі. Поверніть один із наступних рядків залежно від заданих правил:

- `"short post"`, якщо він поміщається в межах обмеження у 40 символів.
- `"long post"`, якщо він більший за 40 символів, але поміщається в межах обмеження у 80 символів.
- `"invalid post"`, якщо він занадто довгий, щоб поміститися в будь-яке з цих обмежень.

# --hints--

`can_post("Hello world")` має повертати `"short post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("Hello world"), "short post")`)
}})
```

`can_post("This is a longer message but still under eighty characters.")` має повертати `"long post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This is a longer message but still under eighty characters."), "long post")`)
}})
```

`can_post("This message is too long to fit into either of the character limits for a social media post.")` має повертати `"invalid post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This message is too long to fit into either of the character limits for a social media post."), "invalid post")`)
}})
```

# --seed--

## --seed-contents--

```py
def can_post(message):

    return message
```

# --solutions--

```py
def can_post(message):
    if len(message) <= 40:
        return "short post"
    elif len(message) <= 80:
        return "long post"
    else:
        return "invalid post"
```
