---
id: 68ffb91507a5b645769328c6
title: "Завдання 103: НСК"
challengeType: 29
dashedName: challenge-103
---

# --description--

Дано два цілі числа, поверніть найменше спільне кратне (НСК) цих двох чисел.

НСК двох чисел — це найменше додатне ціле число, яке є кратним обох чисел. Наприклад, для `4` і `6` поверніть `12`, тому що:

- Кратні `4` — це `4`, `8`, `12` і так далі.
- Кратні `6` — це `6`, `12`, `18` і так далі.
- `12` — найменше число, яке є кратним обох.

# --hints--

`lcm(4, 6)` має повернути `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(4, 6), 12)`)
}})
```

`lcm(9, 6)` має повернути `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(9, 6), 18)`)
}})
```

`lcm(10, 100)` має повернути `100`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(10, 100), 100)`)
}})
```

`lcm(13, 17)` має повернути `221`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(13, 17), 221)`)
}})
```

`lcm(45, 70)` має повернути `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(45, 70), 630)`)
}})
```

# --seed--

## --seed-contents--

```py
def lcm(a, b):

    return a
```

# --solutions--

```py
def lcm(a, b):
    def gcd(x, y):
        return x if y == 0 else gcd(y, x % y)

    return abs(a * b) // gcd(a, b)
```
