---
id: 68ffb91507a5b645769328c9
title: "Завдання 106: валідатор повідомлень"
challengeType: 29
dashedName: challenge-106
---

# --description--

Дано рядок повідомлення та рядок валідації, визначте, чи є повідомлення дійсним.

- Повідомлення вважається дійсним, якщо кожне слово в повідомленні починається з відповідної літери у рядку валідації, у правильному порядку.
- Літери не чутливі до регістру.
- Слова в повідомленні розділені одним пробілом.

# --hints--

`is_valid_message("hello world", "hw")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("hello world", "hw"), True)`)
}})
```

`is_valid_message("ALL CAPITAL LETTERS", "acl")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("ALL CAPITAL LETTERS", "acl"), True)`)
}})
```

`is_valid_message("Coding challenge are boring.", "cca")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("Coding challenge are boring.", "cca"), False)`)
}})
```

`is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"), True)`)
}})
```

`is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_message(message, validation):

    return message
```

# --solutions--

```py
def is_valid_message(message, validation):
    words = message.split()
    
    if len(words) != len(validation):
        return False

    for word, val_letter in zip(words, validation):
        if word[0].lower() != val_letter.lower():
            return False

    return True
```
