---
id: 68ffb91507a5b645769328c3
title: "Завдання 100: 100 символів"
challengeType: 29
dashedName: challenge-100
---

# --description--

Вітаємо з 100-м щоденним завданням з програмування!

Дано рядок, повторюйте його символи, доки результат не стане рівно 100 символів завдовжки. Якщо повторення перевищують 100 символів, обріжте зайве, щоб отримати рівно 100.

# --hints--

`one_hundred("One hundred ")` має повернути `"One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(one_hundred("One hundred "), "One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One ")`)
}})
```

`one_hundred("freeCodeCamp ")` має повернути `"freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(one_hundred("freeCodeCamp "), "freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC")`)
}})
```

`one_hundred("daily challenges ")` має повернути `"daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(one_hundred("daily challenges "), "daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge")`)
}})
```

`one_hundred("!")` має повернути `"!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(one_hundred("!"), "!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!")`)
}})
```



# --seed--

## --seed-contents--

```py
def one_hundred(chars):

    return chars
```

# --solutions--

```py
def one_hundred(chars):
    result = ""
    i = 0

    while len(result) < 100:
        result += chars[i % len(chars)]
        i += 1

    return result[:100]
```
