---
id: 68ffb91507a5b645769328cc
title: "Завдання 109: Скільки мені років?"
challengeType: 29
dashedName: challenge-109
---

# --description--

За датою народження у форматі `YYYY-MM-DD` поверніть вік людини станом на 27 листопада 2025 року.

- Припускайте, що всі дати народження є дійсними і передують 27 листопада 2025 року.
- Поверніть вік у вигляді цілого числа.
- Обов’язково врахуйте, чи вже була у людини дата народження у 2025 році.

# --hints--

`calculate_age("2000-11-20")` має повернути `25`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-11-20"), 25)`)
}})
```

`calculate_age("2000-12-01")` має повернути `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-12-01"), 24)`)
}})
```

`calculate_age("2014-10-25")` має повернути `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2014-10-25"), 11)`)
}})
```

`calculate_age("1994-01-06")` має повернути `31`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-01-06"), 31)`)
}})
```

`calculate_age("1994-12-14")` має повернути `30`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-12-14"), 30)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_age(birthday):

    return birthday
```

# --solutions--

```py
from datetime import datetime
def calculate_age(birthday):
    today = datetime(2025, 11, 27)
    birth_date = datetime.strptime(birthday, "%Y-%m-%d")

    age = today.year - birth_date.year

    if (today.month, today.day) < (birth_date.month, birth_date.day):
        age -= 1

    return age
```
