---
id: 69162d64f96574d9bb629eff
title: "Завдання 114: camel case у snake case"
challengeType: 29
dashedName: challenge-114
---

# --description--

Дано рядок у camel case. Поверніть версію цього рядка у snake case, використовуючи такі правила:

- Вхідний рядок міститиме лише літери (`A-Z` та `a-z`) і завжди починатиметься з літери нижнього регістру.
- Кожна велика літера у рядку camel case позначає початок нового слова.
- Перетворіть усі літери на нижній регістр.
- Розділяйте слова символом підкреслення (`_`).

# --hints--

`to_snake("helloWorld")` має повернути `"hello_world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` має повернути `"my_variable_name"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` має повернути `"freecodecamp_daily_challenges"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
