---
id: 69162d64f96574d9bb629f01
title: "Завдання 116: підрахунок перестановок"
challengeType: 29
dashedName: challenge-116
---

# --description--

Дано рядок, поверніть число унікальних перестановок, які можна утворити з його символів.

- Перестановка — це будь-яке перенесення символів рядка у новому порядку.
- Не рахуйте однакові перестановки двічі.
- Якщо рядок містить повторювані символи, повторні варіанти слід враховувати лише один раз.
- Рядок міститиме лише літери (`A-Z`, `a-z`).

Наприклад, для `"abb"` поверніть `3`, бо існує три унікальні способи розташувати літери: `"abb"`, `"bab"` та `"bba"`.

# --hints--

`count_permutations("abb")` має повернути `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` має повернути `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` має повернути `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` має повернути `39916800`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
