---
id: 69162d64f96574d9bb629efd
title: "Завдання 112: детектор штучного інтелекту"
challengeType: 29
dashedName: challenge-112
---

# --description--

Сьогоднішнє завдання натхненне виходом ChatGPT 30 листопада 2022 року.

Дано рядок з одного або кількох речень, визначте, чи ймовірно він був згенерований штучним інтелектом, використовуючи такі правила:

- Він містить два або більше дефіси (`-`).
- Він містить два або більше множини дужок (`()`). Текст може бути всередині дужок.
- Він містить три або більше слів з 7 або більше літерами.

- Слова розділені одним пробілом і складаються лише з літер (`A-Z`). Не включайте розділові знаки або інші нелітературні символи як частину слова.

Якщо задане речення відповідає будь-якому з наведених правил, поверніть `"AI"`, інакше поверніть `"Human"`.

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` має повертати `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` має повертати `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` має повертати `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` має повертати `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` має повертати `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```
