---
id: 69162d64f96574d9bb629f00
title: "Завдання 115: конвертер елемента впорядкованого списку Markdown"
challengeType: 29
dashedName: challenge-115
---

# --description--

Дано рядок, який представляє елемент впорядкованого списку в Markdown. Потрібно повернути еквівалентний HTML-рядок.

Дійсний елемент впорядкованого списку в Markdown повинен:

- Починатися з нуля або більше пробілів, за якими йде
- Число (1 або більше) і крапка (`.`), за якими йде
- Принаймні один пробіл, а потім
- Текст елемента списку.

Якщо рядок не має точно такого формату, поверніть `"Invalid format"`. Інакше обгорніть текст елемента списку в теги `li` і поверніть рядок.

Наприклад, для `"1. My item"` поверніть `"<li>My item</li>"`.

Примітка: консоль може не відображати HTML-теги в рядках при виводі повідомлень. Перевірте консоль браузера, щоб побачити логи з тегами. 

# --hints--

`convert_list_item("1. My item")` має повернути `"<li>My item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` має повернути `"<li>Another item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` має повернути `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` має повернути `"<li>list item text</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` має повернути `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` має повернути `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```
