---
id: 69162d64f96574d9bb629efe
title: "Завдання 113: милі в кілометри"
challengeType: 29
dashedName: challenge-113
---

# --description--

Дано відстань у милях у вигляді числа, поверніть еквівалентну відстань у кілометрах.

- Вхідні дані завжди будуть невід’ємним числом.
- 1 миля дорівнює 1,60934 кілометра.
- Округліть результат до двох десяткових знаків.
- Видаліть непотрібні нульові цифри в кінці округленого результату.

# --hints--

`convert_to_km(1)` має повертати `1.61`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(1), 1.61)`)
}})
```

`convert_to_km(21)` має повертати `33.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(21), 33.8)`)
}})
```

`convert_to_km(3.5)` має повертати `5.63`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(3.5), 5.63)`)
}})
```

`convert_to_km(0)` має повертати `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0), 0)`)
}})
```

`convert_to_km(0.621371)` має повертати `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0.621371), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_km(miles):

    return miles
```

# --solutions--

```py
def convert_to_km(miles):
    km = miles * 1.60934
    return round(km, 2)
```
