---
id: 69162d64f96574d9bb629efb
title: "Завдання 110: Відгадувач слів"
challengeType: 29
dashedName: challenge-110
---

# --description--

Дано два рядки однакової довжини: секретне слово та здогадку. Порівняйте здогадку із секретним словом за такими правилами:

- Секретне слово та здогадка складатимуться лише з літер верхнього регістру (`"A"` до `"Z"`);
- Для кожної літери у здогадці замініть її на число відповідно до того, як вона співпадає із секретним словом:
  - `"2"`, якщо літера є у секретному слові і знаходиться на правильній позиції.
  - `"1"`, якщо літера є у секретному слові, але на неправильній позиції.
  - `"0"`, якщо літери немає у секретному слові.
- Кожна літера у секретному слові може бути використана не більше одного разу.
- Спочатку призначаються точні співпадіння (`"2"`), потім часткові співпадіння (`"1"`) призначаються зліва направо для решти літер.
- Якщо літера зустрічається кілька разів у здогадці, вона може співпадати лише стільки разів, скільки вона є у секретному слові.

Наприклад, якщо секретне слово `"APPLE"`, а здогадка `"POPPA"`, поверніть `"10201"`:

Перша `"P"` не на правильній позиції (`"1"`), `"O"` немає у секретному слові (`"0"`), друга `"P"` на правильній позиції (`"2"`), третя `"P"` — це нуль (`"0"`), бо дві `"P"` у секретному слові вже використані, а `"A"` не на правильній позиції (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` має повернути `"10201"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("APPLE", "POPPA"), "10201")`)
}})
```

`compare("REACT", "TRACE")` має повернути `"11221"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("REACT", "TRACE"), "11221")`)
}})
```

`compare("DEBUGS", "PYTHON")` має повернути `"000000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("DEBUGS", "PYTHON"), "000000")`)
}})
```

`compare("JAVASCRIPT", "TYPESCRIPT")` має повернути `"0000222222"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222")`)
}})
```

`compare("ORANGE", "ROUNDS")` має повернути `"110200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("ORANGE", "ROUNDS"), "110200")`)
}})
```

`compare("WIRELESS", "ETHERNET")` має повернути `"10021000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("WIRELESS", "ETHERNET"), "10021000")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare(word, guess):

    return word
```

# --solutions--

```py
def compare(word, guess):
    n = len(word)
    result = ["0"] * n
    used_secret = [False] * n

    for i in range(n):
        if guess[i] == word[i]:
            result[i] = "2"
            used_secret[i] = True

    for i in range(n):
        if result[i] == "0":
            for j in range(n):
                if not used_secret[j] and guess[i] == word[j]:
                    result[i] = "1"
                    used_secret[j] = True
                    break

    return "".join(result)
```
