---
id: 69162d64f96574d9bb629f03
title: "Завдання 118: форматування дати"
challengeType: 29
dashedName: challenge-118
---

# --description--

Дано дату у форматі `"Month day, year"`, поверніть дату у форматі `"YYYY-MM-DD"`.

- Даний місяць буде повною англійською назвою місяця. Наприклад: `"January"`, `"February"` тощо.
- У поверненому значенні додайте провідні нулі до місяця та дня, якщо це необхідно, щоб забезпечити двозначний формат.

Наприклад, для `"December 6, 2025"` поверніть `"2025-12-06"`.

# --hints--

`format_date("December 6, 2025")` має повернути `"2025-12-06"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("December 6, 2025"), "2025-12-06")`)
}})
```

`format_date("January 1, 2000")` має повернути `"2000-01-01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("January 1, 2000"), "2000-01-01")`)
}})
```

`format_date("November 11, 1111")` має повернути `"1111-11-11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("November 11, 1111"), "1111-11-11")`)
}})
```

`format_date("September 7, 512")` має повернути `"512-09-07"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("September 7, 512"), "512-09-07")`)
}})
```

`format_date("May 4, 1950")` має повернути `"1950-05-04"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("May 4, 1950"), "1950-05-04")`)
}})
```

`format_date("February 29, 1992")` має повернути `"1992-02-29"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("February 29, 1992"), "1992-02-29")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_date(date_string):

    return date_string
```

# --solutions--

```py
def format_date(date_string):
    months = {
        "January": "01",
        "February": "02",
        "March": "03",
        "April": "04",
        "May": "05",
        "June": "06",
        "July": "07",
        "August": "08",
        "September": "09",
        "October": "10",
        "November": "11",
        "December": "12"
    }

    month, day, year = date_string.replace(",", "").split(" ")
    day = day.zfill(2)

    return f"{year}-{months[month]}-{day}"
```
