---
id: 691b559495c5cb5a37b9b480
title: "Завдання 120: фунти в кілограми"
challengeType: 29
dashedName: challenge-120
---

# --description--

Дано вагу у фунтах як число, поверніть рядок `"(lbs) pounds equals (kgs) kilograms."`.

- Замініть `"(lbs)"` на вхідне число.
- Замініть `"(kgs)"` на вхідне число, конвертоване в кілограми, округлене до двох десяткових знаків і завжди з двома десятковими знаками у значенні.
- 1 фунт дорівнює 0.453592 кілограма.
- Якщо вхідні дані — це `1`, використовуйте `"pound"` замість `"pounds"`.
- Якщо конвертоване значення — це `1`, використовуйте `"kilogram"` замість `"kilograms"`.

# --hints--

`convert_to_kgs(1)` має повертати `"1 pound equals 0.45 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(1), "1 pound equals 0.45 kilograms.")`)
}})
```

`convert_to_kgs(0)` має повертати `"0 pounds equals 0.00 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(0), "0 pounds equals 0.00 kilograms.")`)
}})
```

`convert_to_kgs(100)` має повертати `"100 pounds equals 45.36 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(100), "100 pounds equals 45.36 kilograms.")`)
}})
```

`convert_to_kgs(2.5)` має повертати `"2.5 pounds equals 1.13 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.5), "2.5 pounds equals 1.13 kilograms.")`)
}})
```

`convert_to_kgs(2.20462)` має повертати `"2.20462 pounds equals 1.00 kilogram."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.20462), "2.20462 pounds equals 1.00 kilogram.")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_kgs(lbs):

    return lbs
```

# --solutions--

```py
def convert_to_kgs(lbs):
    KG_PER_POUND = 0.453592
    kgs = f"{lbs * KG_PER_POUND:.2f}"

    pound_word = "pound" if lbs == 1 else "pounds"
    kilogram_word = "kilogram" if kgs == "1.00" else "kilograms"

    return f"{lbs} {pound_word} equals {kgs} {kilogram_word}."
```
