---
id: 691b559495c5cb5a37b9b486
title: "Завдання 126: зробити першу літеру великою"
challengeType: 29
dashedName: challenge-126
---

# --description--

Дано рядок title, поверніть новий рядок, відформатований у форматі заголовка за такими правилами:

- Зробіть першу літеру кожного слова великою.
- Усі інші літери в кожному слові зробіть у нижньому регістрі.
- Слова завжди розділені одним пробілом.

# --hints--

`title_case("hello world")` має повернути `"Hello World"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("hello world"), "Hello World")`)
}})
```

`title_case("the quick brown fox")` має повернути `"The Quick Brown Fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("the quick brown fox"), "The Quick Brown Fox")`)
}})
```

`title_case("JAVASCRIPT AND PYTHON")` має повернути `"Javascript And Python"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("JAVASCRIPT AND PYTHON"), "Javascript And Python")`)
}})
```

`title_case("AvOcAdO tOAst fOr brEAkfAst")` має повернути `"Avocado Toast For Breakfast"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast")`)
}})
```

# --seed--

## --seed-contents--

```py
def title_case(title):

    return title
```

# --solutions--

```py
def title_case(title):
    return " ".join(
        w[:1].upper() + w[1:].lower()
        for w in title.split(" ")
    )
```
