---
id: 691b559495c5cb5a37b9b487
title: "Завдання 127: перевірка швидкості"
challengeType: 29
dashedName: challenge-127
---

# --description--

Враховуючи швидкість, з якою ви рухаєтеся в милях на годину (MPH), та обмеження швидкості в кілометрах на годину (KPH), визначте, чи перевищуєте ви швидкість, і чи отримаєте попередження або штраф.

- 1 миля дорівнює 1,60934 кілометра.
- Якщо ви рухаєтеся зі швидкістю, меншою або рівною обмеженню, поверніть `"Not Speeding"`.
- Якщо ви рухаєтеся зі швидкістю, що перевищує обмеження на 5 KPH або менше, поверніть `"Warning"`.
- Якщо ви рухаєтеся зі швидкістю, що перевищує обмеження більш ніж на 5 KPH, поверніть `"Ticket"`.

# --hints--

`speed_check(30, 70)` має повертати `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` має повертати `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` має повертати `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` має повертати `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` має повертати `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` має повертати `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```
