---
id: 691b559495c5cb5a37b9b482
title: "Завдання 122: парсер жирного тексту в markdown"
challengeType: 29
dashedName: challenge-122
---

# --description--

Дано рядок, який може містити жирний текст у markdown. Потрібно повернути еквівалентний рядок у HTML.

- Жирний текст у markdown — це будь-який текст, який починається і закінчується двома зірочками (`**`) або двома підкресленнями (`__`).
- Між текстом і зірочками або підкресленнями не може бути пробілів, але пробіли можуть бути всередині самого тексту.
- Перетворіть усі випадки жирного тексту на HTML-теги `b` і поверніть рядок.

Наприклад, для `"**This is bold**"` поверніть `"<b>This is bold</b>"`.

Примітка: консоль може не відображати HTML-теги в рядках при виводі повідомлень. Перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`parse_bold("**This is bold**")` має повернути `"<b>This is bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is bold**"), "<b>This is bold</b>")`)
}})
```

`parse_bold("__This is also bold__")` має повернути `"<b>This is also bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__This is also bold__"), "<b>This is also bold</b>")`)
}})
```

`parse_bold("**This is not bold **")` має повернути `"**This is not bold **"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is not bold **"), "**This is not bold **")`)
}})
```

`parse_bold("__ This is also not bold__")` має повернути `"__ This is also not bold__"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__ This is also not bold__"), "__ This is also not bold__")`)
}})
```

`parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog.")` має повернути `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_bold(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_bold(markdown):
    markdown = re.sub(r'\*\*(\S(?:.*?\S)?)\*\*', r'<b>\1</b>', markdown)
    markdown = re.sub(r'__(\S(?:.*?\S)?)__', r'<b>\1</b>', markdown)

    return markdown
```
