---
id: 691b559495c5cb5a37b9b485
title: "Завдання 125: гра життя"
challengeType: 29
dashedName: challenge-125
---

# --description--

Дано матрицю (масив масивів), що представляє поточний стан у грі життя Конвея. Поверніть наступний стан матриці, використовуючи такі правила:

- Кожна клітинка або `1` (жива), або `0` (мертва).
- Сусідами клітинки є до восьми навколишніх клітинок (вертикально, горизонтально та по діагоналі).
- Клітинки на краях мають менше ніж вісім сусідів.

Правила оновлення кожної клітинки:

- Будь-яка жива клітинка з менше ніж двома живими сусідами вмирає (недостатня кількість).
- Будь-яка жива клітинка з двома або трьома живими сусідами продовжує жити.
- Будь-яка жива клітинка з більш ніж трьома живими сусідами вмирає (надмірна кількість).
- Будь-яка мертва клітинка з рівно трьома живими сусідами стає живою (розмноження).

Наприклад, для:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

поверніть:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Кожна клітинка оновлюється відповідно до кількості живих сусідів. Наприклад, `[0][0]` залишається мертвою (2 живих сусіди), `[0][1]` залишається живою (2 живих сусіди), `[0][2]` вмирає (3 живих сусіди) і так далі.

# --hints--

`game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` має повернути `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]])`)
}})
```

`game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` має повернути `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]])`)
}})
```

`game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` має повернути `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]])`)
}})
```

`game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` має повернути `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def game_of_life(grid):

    return grid
```

# --solutions--

```py
def game_of_life(grid):
    rows, cols = len(grid), len(grid[0])

    def count_live_neighbors(r, c):
        count = 0
        for i in range(r-1, r+2):
            for j in range(c-1, c+2):
                if 0 <= i < rows and 0 <= j < cols and (i != r or j != c):
                    count += grid[i][j]
        return count

    next_state = [row[:] for row in grid]

    for r in range(rows):
        for c in range(cols):
            live_neighbors = count_live_neighbors(r, c)

            if grid[r][c] == 1:
                if live_neighbors < 2 or live_neighbors > 3:
                    next_state[r][c] = 0
            else:
                if live_neighbors == 3:
                    next_state[r][c] = 1

    return next_state
```
