---
id: 691b559495c5cb5a37b9b483
title: "Завдання 123: конструктор римських чисел"
challengeType: 29
dashedName: challenge-123
---

# --description--

Дано ціле число, поверніть його еквівалентне значення римськими цифрами.

Римські цифри використовують такі символи:

| Символ | Значення |
| ------ | -------- |
| I      | 1        |
| V      | 5        |
| X      | 10       |
| L      | 50       |
| C      | 100      |
| D      | 500      |
| M      | 1000     |

Римські цифри записуються від найбільшого до найменшого, зліва направо, за такими правилами:

- Додавання використовується, коли символ слідує за символом з таким самим або меншим значенням. Наприклад, `18` записується як `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- Віднімання використовується, коли менший символ стоїть перед більшим, щоб представити 4 або 9 у будь-якому розряді. Наприклад, 19 записується як `XIX` (10 + (10 - 1)).
- Жоден символ не може повторюватися більше трьох разів поспіль. Віднімання використовується, коли інакше довелося б записати символ більше трьох разів поспіль. Тож найбільше число, яке можна записати, — 3999.

Ось ще один приклад: для `1464` поверніть `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`to_roman(18)` має повернути `"XVIII"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(18), "XVIII")`)
}})
```

`to_roman(19)` має повернути `"XIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(19), "XIX")`)
}})
```

`to_roman(1464)` має повернути `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(1464), "MCDLXIV")`)
}})
```

`to_roman(2025)` має повернути `"MMXXV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(2025), "MMXXV")`)
}})
```

`to_roman(3999)` має повернути `"MMMCMXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(3999), "MMMCMXCIX")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_roman(num):

    return num
```

# --solutions--

```py
def to_roman(num):
    symbols = [
        ("M", 1000),
        ("CM", 900),
        ("D", 500),
        ("CD", 400),
        ("C", 100),
        ("XC", 90),
        ("L", 50),
        ("XL", 40),
        ("X", 10),
        ("IX", 9),
        ("V", 5),
        ("IV", 4),
        ("I", 1)
    ]

    result = ""
    for sym, val in symbols:
        while num >= val:
            result += sym
            num -= val
    return result
```
