---
id: 6925e2068081f40f549ced1c
title: "Завдання 138: сума дільників"
challengeType: 29
dashedName: challenge-138
---

# --description--

Дано додатне ціле число, поверніть суму всіх його дільників.

- Дільник — це будь-яке ціле число, яке ділить задане число без остачі (остача дорівнює `0`).
- Рахуйте кожен дільник лише один раз.

Наприклад, для `6` поверніть `12`, тому що дільниками `6` є `1`, `2`, `3` і `6`, а їх сума дорівнює `12`.

# --hints--

`sum_divisors(6)` має повернути `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` має повернути `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` має повернути `56`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` має повернути `224`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` має повернути `806`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` має повернути `23520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```
