---
id: 6925e2068081f40f549ced1a
title: "Завдання 136: парсер зображень у markdown"
challengeType: 29
dashedName: challenge-136
---

# --description--

Дано рядок із зображенням у форматі Markdown, поверніть еквівалентний HTML-рядок.

Зображення в Markdown має такий формат: `"![alt text](image_url)"`. Де:

- `alt text` — це опис зображення (значення атрибута `alt`).
- `image_url` — це URL-адреса джерела зображення (значення атрибута `src`).

Поверніть рядок HTML-тега `img` з множиною атрибутів, де `src` дорівнює URL зображення, а `alt` — альтернативному тексту.

Наприклад, для `"![Cute cat](cat.png)"` поверніть `'<img src="cat.png" alt="Cute cat">'`;

- Переконайтеся, що тег, порядок атрибутів, пробіли та використання лапок збігаються з наведеним вище.

Примітка: консоль може не відображати HTML-теги у рядках при виводі повідомлень — перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`parse_image("![Cute cat](cat.png)")` має повернути `'<img src="cat.png" alt="Cute cat">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` має повернути `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` має повернути `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```
