---
id: 6925e2068081f40f549ced1d
title: "Завдання 139: камінь, ножиці, папір"
challengeType: 29
dashedName: challenge-139
---

# --description--

Дано два рядки: перший представляє гравця 1, а другий — гравця 2. Визначте переможця у грі камінь, ножиці, папір.

- Вхідні рядки завжди будуть `"Rock"`, `"Paper"` або `"Scissors"`.
- `"Rock"` перемагає `"Scissors"`.
- `"Paper"` перемагає `"Rock"`.
- `"Scissors"` перемагає `"Paper"`.

Поверніть:

- `"Player 1 wins"`, якщо виграє гравець 1.
- `"Player 2 wins"`, якщо виграє гравець 2.
- `"Tie"`, якщо обидва гравці вибрали однаковий варіант.

# --hints--

`rock_paper_scissors("Rock", "Rock")` має повернути `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Rock"), "Tie")`)
}})
```

`rock_paper_scissors("Rock", "Paper")` має повернути `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Paper"), "Player 2 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Paper")` має повернути `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Paper"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Rock", "Scissors")` має повернути `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Scissors"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Scissors")` має повернути `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Scissors"), "Tie")`)
}})
```

`rock_paper_scissors("Scissors", "Rock")` має повернути `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Rock"), "Player 2 wins")`)
}})
```

# --seed--

## --seed-contents--

```py
def rock_paper_scissors(player1, player2):

    return player1
```

# --solutions--

```py
def rock_paper_scissors(player1, player2):
    if player1 == player2:
        return "Tie"
    
    if (player1 == "Rock" and player2 == "Scissors") or \
       (player1 == "Paper" and player2 == "Rock") or \
       (player1 == "Scissors" and player2 == "Paper"):
        return "Player 1 wins"
    else:
        return "Player 2 wins"
```
