---
id: 6925e2068081f40f549ced1b
title: "Завдання 137: генератор сніжинок"
challengeType: 29
dashedName: challenge-137
---

# --description--

Дано багаторядковий рядок, який використовує символи нового рядка (`\n`) для представлення розриву рядка. Поверніть новий рядок, у якому кожен рядок віддзеркалений по горизонталі та приєднаний до кінця оригінального рядка.

- Віддзеркальте рядок, перевернувши всі його символи, включно з пробілами.

Наприклад, для `"* \n *\n* "`, який виводить у консоль:

```sh
* 
 *
* 
```

Поверніть `"*  *\n ** \n*  *"`, який виводить у консоль:

```sh
*  *
 ** 
*  *
```

Уважно зверніть увагу на пробіли у заданих і повернутих рядках. Обов’язково не обрізайте жодного з них.

# --hints--

`generate_snowflake("* \n *\n* ")` має повернути `"*  *\n ** \n*  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` має повернути `"X=~~=X"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` має повернути `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` має повернути `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` має повернути `"*  --  *\n * -- * \n*  --  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```
