---
id: 69272dcf1c24b44fd79137c4
title: "Завдання 141: пальне для зльоту"
challengeType: 29
dashedName: challenge-141
---

# --description--

Враховуючи кількість галонів пального, що наразі є у вашому літаку, та необхідну кількість літрів пального для досягнення пункту призначення, визначте, скільки додаткових галонів пального слід додати.

- 1 галон дорівнює 3.78541 літрів.
- Якщо в літаку вже достатньо пального, поверніть `0`.
- Ви можете додавати лише цілі галони.
- Не включайте десяткові числа у повернене значення.

# --hints--

`fuel_to_add(0, 1)` має повертати `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(0, 1), 1)`)
}})
```

`fuel_to_add(5, 40)` має повертати `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(5, 40), 6)`)
}})
```

`fuel_to_add(10, 30)` має повертати `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(10, 30), 0)`)
}})
```

`fuel_to_add(896, 20500)` має повертати `4520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(896, 20500), 4520)`)
}})
```

`fuel_to_add(1000, 50000)` має повертати `12209`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(1000, 50000), 12209)`)
}})
```

# --seed--

## --seed-contents--

```py
def fuel_to_add(current_gallons, required_liters):

    return current_gallons
```

# --solutions--

```py
import math
def fuel_to_add(current_gallons, required_liters):
    liters_per_gallon = 3.78541
    current_liters = current_gallons * liters_per_gallon
    if current_liters >= required_liters:
        return 0
    liters_needed = required_liters - current_liters
    return math.ceil(liters_needed / liters_per_gallon)
```
