---
id: 69272dcf1c24b44fd79137c5
title: "Завдання 142: Сума рядка"
challengeType: 29
dashedName: challenge-142
---

# --description--

Дано рядок, що містить цифри та інші символи. Поверніть суму всіх чисел у цьому рядку.

- Розглядайте послідовні цифри як одне число. Наприклад, `"13"` рахується як 13, а не 1 + 3.
- Ігноруйте всі символи, які не є цифрами.

# --hints--

`string_sum("3apples2bananas")` має повернути `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("3apples2bananas"), 5)`)
}})
```

`string_sum("10cats5dogs2birds")` має повернути `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("10cats5dogs2birds"), 17)`)
}})
```

`string_sum("125344")` має повернути `125344`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("125344"), 125344)`)
}})
```

`string_sum("a1b20c300")` має повернути `321`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a1b20c300"), 321)`)
}})
```

`string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` має повернути `1653`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653)`)
}})
```

# --seed--

## --seed-contents--

```py
def string_sum(s):

    return s
```

# --solutions--

```py
import re
def string_sum(s):
    numbers = re.findall(r'\d+', s)
    return sum(int(num) for num in numbers)
```
