---
id: 69272dcf1c24b44fd79137c3
title: "Завдання 140: SCREAMING_SNAKE_CASE"
challengeType: 29
dashedName: challenge-140
---

# --description--

Дано рядок, що представляє ім’я змінної. Поверніть ім’я змінної, конвертоване у SCREAMING_SNAKE_CASE.

Дані імена змінних будуть записані в одному з наступних форматів:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

У наведених форматах слова розділені підкресленням (`_`), дефісом (`-`) або нове слово починається з великої літери.

Щоб конвертувати у SCREAMING_SNAKE_CASE:

- Зробіть усі літери у верхньому регістрі
- Розділіть слова підкресленням (`_`)

# --hints--

`to_screaming_snake_case("userEmail")` має повернути `"USER_EMAIL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("userEmail"), "USER_EMAIL")`)
}})
```

`to_screaming_snake_case("UserPassword")` має повернути `"USER_PASSWORD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("UserPassword"), "USER_PASSWORD")`)
}})
```

`to_screaming_snake_case("user_id")` має повернути `"USER_ID"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user_id"), "USER_ID")`)
}})
```

`to_screaming_snake_case("user-address")` має повернути `"USER_ADDRESS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user-address"), "USER_ADDRESS")`)
}})
```

`to_screaming_snake_case("username")` має повернути `"USERNAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("username"), "USERNAME")`)
}})
```

`to_screaming_snake_case("my_variable_name")` має повернути `"MY_VARIABLE_NAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("my_variable_name"), "MY_VARIABLE_NAME")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_screaming_snake_case(variable_name):

    return variable_name
```

# --solutions--

```py
import re
def to_screaming_snake_case(variable_name):
    temp = re.sub(r'[-_]+', ' ', variable_name)
    temp = re.sub(r'([a-z0-9])([A-Z])', r'\1 \2', temp)
    words = temp.strip().split()
    return '_'.join(words).upper()
```
