---
id: 69306364df283fcaff2e1ad7
title: "Завдання 146: місце для лівші за столом"
challengeType: 29
dashedName: challenge-146
---

# --description--

Дано матрицю 4x2 (масив масивів), що представляє розсадження за столом під час їжі. Визначте, скільки місць підходять для лівші.

- Лівша не може сидіти там, де правша сидить на місці безпосередньо зліва від нього.

У заданій матриці:

- `"R"` — це місце, зайняте правшею.
- `"L"` — це місце, зайняте лівшею.
- `"U"` — це вільне місце.
- Сідати можна лише на вільні місця.
- Місця у верхньому ряду повернені "вниз", а місця в нижньому ряду — "вгору" (як за столом), тому ліво і право відносні до орієнтації місця.
- Кутові місця мають лише одне сусіднє місце.

Наприклад, у заданій матриці:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

Місце у верхньому лівому куті зайняти не можна, бо зліва сидить правша. Інші два вільні місця можна зайняти, бо зліва немає правші.

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` має повернути `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` має повернути `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` має повернути `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` має повернути `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` має повернути `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```
