---
id: 69306364df283fcaff2e1ad6
title: "Завдання 145: N-те число Фібоначчі"
challengeType: 29
dashedName: challenge-145
---

# --description--

Дано ціле число `n`. Поверніть `n`-те число в послідовності Фібоначчі.

Послідовність Фібоначчі — це ряд чисел, де кожне число є сумою двох попередніх. Перші 10 чисел у послідовності: `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

# --hints--

`nth_fibonacci(4)` має повернути `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(4), 2)`)
}})
```

`nth_fibonacci(10)` має повернути `34`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(10), 34)`)
}})
```

`nth_fibonacci(15)` має повернути `377`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(15), 377)`)
}})
```

`nth_fibonacci(40)` має повернути `63245986`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(40), 63245986)`)
}})
```

`nth_fibonacci(75)` має повернути `1304969544928657`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(75), 1304969544928657)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_fibonacci(n):

    return n
```

# --solutions--

```py
def nth_fibonacci(n):
    if n == 1:
        return 0
    if n == 2:
        return 1

    a, b = 0, 1
    for _ in range(3, n + 1):
        a, b = b, a + b
    return b
```
