---
id: 69373793f5a867f769cde13a
title: "Завдання 155: Посадіть урожай"
challengeType: 29
dashedName: challenge-155
---

# --description--

Дано ціле число, що представляє розмір вашого сільськогосподарського поля, та `"acres"` або `"hectares"`, які позначають одиницю виміру розміру вашого поля, а також тип культури. Визначте, скільки рослин цього типу ви можете розмістити на своєму полі.

- 1 акр дорівнює 4046,86 квадратних метрів.
- 1 гектар дорівнює 10 000 квадратних метрів.

Ось список культур, які будуть задані у вхідних даних, і скільки простору займає одна рослина:

|Культура|Простір на одну рослину|
|-|-|
|`"corn"`|1 квадратний метр|
|`"wheat"`|0,1 квадратного метра|
|`"soybeans"`|0,5 квадратного метра|
|`"tomatoes"`|0,25 квадратного метра|
|`"lettuce"`|0,2 квадратного метра|

Поверніть кількість рослин, які поміщаються на полі, округлену вниз до найближчого цілого числа рослин.

# --hints--

`get_number_of_plants(1, "acres", "corn")` має повернути `4046`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(1, "acres", "corn"), 4046)`)
}})
```

`get_number_of_plants(2, "hectares", "lettuce")` має повернути `100000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(2, "hectares", "lettuce"), 100000)`)
}})
```

`get_number_of_plants(20, "acres", "soybeans")` має повернути `161874`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(20, "acres", "soybeans"), 161874)`)
}})
```

`get_number_of_plants(3.75, "hectares", "tomatoes")` має повернути `150000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(3.75, "hectares", "tomatoes"), 150000)`)
}})
```

`get_number_of_plants(16.75, "acres", "tomatoes")` має повернути `271139`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(16.75, "acres", "tomatoes"), 271139)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_of_plants(field_size, unit, crop):

    return field_size
```

# --solutions--

```py
import math
def get_number_of_plants(field_size, unit, crop):
    crop_space = {
        "corn": 1,
        "wheat": 0.1,
        "soybeans": 0.5,
        "tomatoes": 0.25,
        "lettuce": 0.2
    }

    field_in_sqm = field_size * 4046.86 if unit == "acres" else field_size * 10000

    return math.floor(field_in_sqm / crop_space[crop.lower()])
```
