---
id: 69373793f5a867f769cde139
title: "Завдання 154: Пар для лунки"
challengeType: 29
dashedName: challenge-154
---

# --description--

Дано два цілі числа: пар для гольф-лунки та кількість ударів, які гравець зробив на цій лунці. Поверніть результат гравця, використовуючи терміни гольфу.

Поверніть:

- `"Hole in one!"`, якщо знадобився один удар.
- `"Eagle"`, якщо знадобилося на два удари менше за пар.
- `"Birdie"`, якщо знадобився на один удар менше за пар.
- `"Par"`, якщо кількість ударів дорівнює пар.
- `"Bogey"`, якщо знадобився на один удар більше за пар.
- `"Double bogey"`, якщо знадобилося на два удари більше за пар.

# --hints--

`golf_score(3, 3)` має повертати `"Par"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 3), "Par")`)
}})
```

`golf_score(4, 3)` має повертати `"Birdie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 3), "Birdie")`)
}})
```

`golf_score(3, 1)` має повертати `"Hole in one!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 1), "Hole in one!")`)
}})
```

`golf_score(5, 7)` має повертати `"Double bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 7), "Double bogey")`)
}})
```

`golf_score(4, 5)` має повертати `"Bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 5), "Bogey")`)
}})
```

`golf_score(5, 3)` має повертати `"Eagle"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 3), "Eagle")`)
}})
```

# --seed--

## --seed-contents--

```py
def golf_score(par, strokes):

    return par
```

# --solutions--

```py
def golf_score(par, strokes):
    if strokes == 1:
        return "Hole in one!"

    diff = strokes - par

    if diff == -2:
        return "Eagle"
    if diff == -1:
        return "Birdie"
    if diff == 0:
        return "Par"
    if diff == 1:
        return "Bogey"
    if diff == 2:
        return "Double bogey"
```
