---
id: 6939b873185d8e00d453563b
title: "Завдання 156: непарне чи парне?"
challengeType: 29
dashedName: challenge-156
---

# --description--

Дано додатне ціле число. Поверніть `"Odd"`, якщо це непарне число, і `"Even"`, якщо парне.

# --hints--

`odd_or_even(1)` має повертати `"Odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(1), "Odd")`)
}})
```

`odd_or_even(2)` має повертати `"Even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(2), "Even")`)
}})
```

`odd_or_even(13)` має повертати `"Odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(13), "Odd")`)
}})
```

`odd_or_even(196)` має повертати `"Even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(196), "Even")`)
}})
```

`odd_or_even(123456789)` має повертати `"Odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(123456789), "Odd")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even(n):

    return n
```

# --solutions--

```py
def odd_or_even(n):
    return "Even" if n % 2 == 0 else "Odd"
```
