---
id: 6939b873185d8e00d453563c
title: "Завдання 157: парсер посилань у markdown"
challengeType: 29
dashedName: challenge-157
---

# --description--

Дано рядок посилання у форматі Markdown, поверніть еквівалентний рядок у форматі HTML.

Зображення у Markdown має такий формат: `"[link_text](link_url)"`. Поверніть рядок HTML-тегу `a` з множиною `href`, встановленою у `link_url`, а `link_text` як вміст тегу.

Наприклад, для `"[freeCodeCamp](https://freecodecamp.org/)"` поверніть `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

Примітка: консоль може не відображати HTML-теги у рядках під час виведення повідомлень — перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`parse_link("[freeCodeCamp](https://freecodecamp.org/)")` має повертати `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>')`)
}})
```

`parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` має повертати `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>')`)
}})
```

`parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` має повертати `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_link(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_link(markdown):
    match = re.match(r'\[(.*?)\]\((.*?)\)', markdown)
    text, url = match.groups()
    return f'<a href="{url}">{text}</a>'
```
