---
id: 694596b0585c11170ac7c7fa
title: "Завдання 162: споживання енергії"
challengeType: 29
dashedName: challenge-162
---

# --description--

Враховуючи число калорій, спалених під час тренування, та число ват-годин, використаних вашими електронними пристроями під час цього тренування, визначте, що використало більше енергії.

Щоб порівняти їх, перетворіть обидва значення в джоулі, використовуючи такі конверсії:

- 1 калорія дорівнює 4184 джоулям.
- 1 ват-година дорівнює 3600 джоулям.

Поверніть:

- `"Workout"`, якщо тренування використало більше енергії.
- `"Devices"`, якщо пристрій використав більше енергії.
- `"Equal"`, якщо обидва використали однакову кількість енергії.

# --hints--

`compare_energy(250, 50)` має повертати `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` має повертати `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` має повертати `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` має повертати `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` має повертати `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` має повертати `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```
