---
id: 694596b0585c11170ac7c7fb
title: "Завдання 163: приголосний регістр"
challengeType: 29
dashedName: challenge-163
---

# --description--

Дано рядок, що представляє ім’я змінної. Перетворіть його на приголосний регістр за такими правилами:

- Усі приголосні мають бути перетворені на верхній регістр.
- Усі голосні (`a`, `e`, `i`, `o`, `u` у будь-якому регістрі) мають бути перетворені на нижній регістр.
- Усі дефіси (`-`) мають бути замінені на підкреслення (`_`).

# --hints--

`to_consonant_case("helloworld")` має повернути `"HeLLoWoRLD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("helloworld"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("HELLOWORLD")` має повернути `"HeLLoWoRLD"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("HELLOWORLD"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("_hElLO-WOrlD-")` має повернути `"_HeLLo_WoRLD_"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_hElLO-WOrlD-"), "_HeLLo_WoRLD_")`)
}})
```

`to_consonant_case("_~-generic_~-variable_~-name_~-here-~_")` має повернути `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_consonant_case(s):

    return s
```

# --solutions--

```py
def to_consonant_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char == "-":
            result += "_"
        elif char.isalpha():
            if char in vowels:
                result += char.lower()
            else:
                result += char.upper()
        else:
            result += char

    return result
```
