---
id: 696655d24b614176d4c9b78c
title: "Завдання 169: міні FizzBuzz"
challengeType: 29
dashedName: challenge-169
---

# --description--

Дано ціле число, поверніть рядок відповідно до таких правил:

- Якщо число ділиться на 3, поверніть `"Fizz"`.
- Якщо число ділиться на 5, поверніть `"Buzz"`.
- Якщо число ділиться і на 3, і на 5, поверніть `"FizzBuzz"`.
- Інакше поверніть задане число у вигляді рядка.

# --hints--

`fizz_buzz_mini(3)` має повертати `"Fizz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(3), "Fizz")`)
}})
```

`fizz_buzz_mini(4)` має повертати `"4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(4), "4")`)
}})
```

`fizz_buzz_mini(35)` має повертати `"Buzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(35), "Buzz")`)
}})
```

`fizz_buzz_mini(75)` має повертати `"FizzBuzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(75), "FizzBuzz")`)
}})
```

`fizz_buzz_mini(98)` має повертати `"98"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(98), "98")`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_mini(n):

    return n
```

# --solutions--

```py
def fizz_buzz_mini(n):
    if n % 3 == 0 and n % 5 == 0:
        return "FizzBuzz"
    if n % 3 == 0:
        return "Fizz"
    if n % 5 == 0:
        return "Buzz"
    return str(n)
```
