---
id: 696655d24b614176d4c9b78d
title: "Завдання 170: непарний чи парний день"
challengeType: 29
dashedName: challenge-170
---

# --description--

За заданим часовим штампом (число мілісекунд від початку епохи Unix) поверніть:

- `"odd"`, якщо число дня місяця для цього часовогo штампа непарне.
- `"even"`, якщо число дня місяця для цього часовогo штампа парне.

Наприклад, для `1769472000000`, часовогo штампа 27 січня 2026 року, поверніть `"odd"`, бо число дня (`27`) є непарним числом.

Примітка: часовий штамп заданий у мілісекундах, і слід використовувати дату в часовому поясі UTC, а не у вашому локальному часі.

# --hints--

`odd_or_even_day(1769472000000)` має повертати `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` має повертати `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` має повертати `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` має повертати `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` має повертати `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```
