---
id: 696655d24b614176d4c9b78b
title: "Завдання 168: масштабоване зображення"
challengeType: 29
dashedName: challenge-168
---

# --description--

Дано рядок, який представляє ширину та висоту зображення, а також число для масштабування зображення. Потрібно повернути масштабовані ширину та висоту.

- Вхідний рядок має формат `"WxH"`. Наприклад, `"800x600"`.
- Масштаб — це число, на яке потрібно помножити ширину та висоту.

Поверніть масштабовані розміри у тому ж форматі `"WxH"`.

# --hints--

`scale_image("800x600", 2)` має повернути `"1600x1200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("800x600", 2), "1600x1200")`)
}})
```

`scale_image("100x100", 10)` має повернути `"1000x1000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("100x100", 10), "1000x1000")`)
}})
```

`scale_image("1024x768", 0.5)` має повернути `"512x384"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("1024x768", 0.5), "512x384")`)
}})
```

`scale_image("300x200", 1.5)` має повернути `"450x300"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("300x200", 1.5), "450x300")`)
}})
```

# --seed--

## --seed-contents--

```py
def scale_image(size, scale):

    return size
```

# --solutions--

```py
def scale_image(size, scale):
    width, height = map(int, size.split("x"))

    new_width = round(width * scale)
    new_height = round(height * scale)

    return f"{new_width}x{new_height}"
```
