---
id: 69738771fb5a7b8b24cca2a3
title: "Завдання 177: дзеркальний рядок"
challengeType: 29
dashedName: challenge-177
---

# --description--

Дано рядок. Поверніть новий рядок, який складається з початкового рядка з доданою в кінці перевернутою копією самого себе.

# --hints--

`mirror("freeCodeCamp")` має повернути `"freeCodeCamppmaCedoCeerf"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf")`)
}})
```

`mirror("RaceCar")` має повернути `"RaceCarraCecaR"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("RaceCar"), "RaceCarraCecaR")`)
}})
```

`mirror("helloworld")` має повернути `"helloworlddlrowolleh"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("helloworld"), "helloworlddlrowolleh")`)
}})
```

`mirror("The quick brown fox...")` має повернути `"The quick brown fox......xof nworb kciuq ehT"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT")`)
}})
```

# --seed--

## --seed-contents--

```py
def mirror(s):

    return s
```

# --solutions--

```py
def mirror(s):
    reversed_s = s[::-1]
    return s + reversed_s
```
