---
id: 69738771fb5a7b8b24cca2a2
title: "Завдання 176: День бабака"
challengeType: 29
dashedName: challenge-176
---

# --description--

Сьогодні День бабака, коли бабак передбачає погоду, спостерігаючи, чи бачить він свою тінь.

Враховуючи значення, що представляє появу бабака, поверніть правильне передбачення:

- Якщо задане значення — булевий `true` (бабак побачив свою тінь), поверніть `"Looks like we'll have six more weeks of winter."`.
- Якщо значення — булевий `false` (бабак не побачив свою тінь), поверніть `"It's going to be an early spring."`.
- Якщо значення інше (бабак не з’явився), поверніть `"No prediction this year."`.

# --hints--

`groundhog_day_prediction(True)` має повертати `"Looks like we'll have six more weeks of winter."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(True), "Looks like we'll have six more weeks of winter.")`)
}})
```

`groundhog_day_prediction(False)` має повертати `"It's going to be an early spring."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(False), "It's going to be an early spring.")`)
}})
```

`groundhog_day_prediction(None)` має повертати `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(None), "No prediction this year.")`)
}})
```

`groundhog_day_prediction(" ")` має повертати `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(" "), "No prediction this year.")`)
}})
```

`groundhog_day_prediction("True")` має повертати `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction("True"), "No prediction this year.")`)
}})
```

# --seed--

## --seed-contents--

```py
def groundhog_day_prediction(appearance):

    return appearance
```

# --solutions--

```py
def groundhog_day_prediction(appearance):
    if appearance is True:
        return "Looks like we'll have six more weeks of winter."
    if appearance is False:
        return "It's going to be an early spring."
    return "No prediction this year."
```
