---
id: 69738771fb5a7b8b24cca2a4
title: "Завдання 178: Обрізати текст"
challengeType: 29
dashedName: challenge-178
---

# --description--

Дано рядок. Поверніть його без змін, якщо він містить 20 символів або менше. Якщо рядок довший за 20 символів, обріжте його до перших 17 символів, додайте `"..."` наприкінці (щоб загальна довжина була 20 символів) і поверніть результат.

# --hints--

`truncate_text("Hello, world!")` має повертати `"Hello, world!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Hello, world!"), "Hello, world!")`)
}})
```

`truncate_text("This string should get truncated.")` має повертати `"This string shoul..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("This string should get truncated."), "This string shoul...")`)
}})
```

`truncate_text("Exactly twenty chars")` має повертати `"Exactly twenty chars"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Exactly twenty chars"), "Exactly twenty chars")`)
}})
```

`truncate_text(".....................")` має повертати `"...................."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("....................."), "....................")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(text):

    return text
```

# --solutions--

```py
def truncate_text(text):
    if len(text) <= 20:
        return text
    return text[:17] + "..."
```
