---
id: 69738771fb5a7b8b24cca2a0
title: "Завдання 174: пошук знаку зодіаку"
challengeType: 29
dashedName: challenge-174
---

# --description--

Дано рядок дати у форматі `"YYYY-MM-DD"`. Поверніть знак зодіаку для цієї дати, використовуючи таку діаграму:

| Діапазон дат | Знак зодіаку |
| - | - |
| 21 березня - 19 квітня | `"Aries"` |
| 20 квітня - 20 травня | `"Taurus"` |
| 21 травня - 20 червня | `"Gemini"` |
| 21 червня - 22 липня | `"Cancer"` |
| 23 липня - 22 серпня | `"Leo"` |
| 23 серпня - 22 вересня | `"Virgo"` |
| 23 вересня - 22 жовтня | `"Libra"` |
| 23 жовтня - 21 листопада | `"Scorpio"` |
| 22 листопада - 21 грудня | `"Sagittarius"` |
| 22 грудня - 19 січня | `"Capricorn"` |
| 20 січня - 18 лютого | `"Aquarius"` |
| 19 лютого - 20 березня | `"Pisces"` |

- Знаки зодіаку визначаються лише за місяцем і днем, рік можна ігнорувати.

# --hints--

`get_sign("2026-01-31")` має повернути `"Aquarius"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2026-01-31"), "Aquarius")`)
}})
```

`get_sign("2001-06-10")` має повернути `"Gemini"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2001-06-10"), "Gemini")`)
}})
```

`get_sign("1985-09-07")` має повернути `"Virgo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1985-09-07"), "Virgo")`)
}})
```

`get_sign("2023-03-19")` має повернути `"Pisces"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2023-03-19"), "Pisces")`)
}})
```

`get_sign("2045-11-05")` має повернути `"Scorpio"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2045-11-05"), "Scorpio")`)
}})
```


`get_sign("1985-12-06")` має повернути `"Sagittarius"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1985-12-06"), "Sagittarius")`)
}})
```

`get_sign("2025-12-30")` має повернути `"Capricorn"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2025-12-30"), "Capricorn")`)
}})
```

`get_sign("2018-10-08")` має повернути `"Libra"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2018-10-08"), "Libra")`)
}})
```

`get_sign("1958-05-04")` має повернути `"Taurus"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1958-05-04"), "Taurus")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_sign(date_str):

    return date_str
```

# --solutions--

```py
def get_sign(date_str):
    _, month_str, day_str = date_str.split("-")
    month = int(month_str)
    day = int(day_str)

    if (month == 3 and day >= 21) or (month == 4 and day <= 19):
        return "Aries"
    if (month == 4 and day >= 20) or (month == 5 and day <= 20):
        return "Taurus"
    if (month == 5 and day >= 21) or (month == 6 and day <= 20):
        return "Gemini"
    if (month == 6 and day >= 21) or (month == 7 and day <= 22):
        return "Cancer"
    if (month == 7 and day >= 23) or (month == 8 and day <= 22):
        return "Leo"
    if (month == 8 and day >= 23) or (month == 9 and day <= 22):
        return "Virgo"
    if (month == 9 and day >= 23) or (month == 10 and day <= 22):
        return "Libra"
    if (month == 10 and day >= 23) or (month == 11 and day <= 21):
        return "Scorpio"
    if (month == 11 and day >= 22) or (month == 12 and day <= 21):
        return "Sagittarius"
    if (month == 12 and day >= 22) or (month == 1 and day <= 19):
        return "Capricorn"
    if (month == 1 and day >= 20) or (month == 2 and day <= 18):
        return "Aquarius"
    if (month == 2 and day >= 19) or (month == 3 and day <= 20):
        return "Pisces"
```
