---
id: 69738771fb5a7b8b24cca29d
title: "Завдання 171: Розгортання масиву"
challengeType: 29
dashedName: challenge-171
---

# --description--

Дано масив, який містить вкладені масиви. Поверніть новий масив, у якому всі значення розгорнуті в один одномірний масив. Збережіть початковий порядок елементів у масивах.

# --hints--

`flatten([1, [2, 3], 4])` має повернути `[1, 2, 3, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([1, [2, 3], 4]), [1, 2, 3, 4])`)
}})
```

`flatten([5, [4, [3, 2]], 1])` має повернути `[5, 4, 3, 2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([5, [4, [3, 2]], 1]), [5, 4, 3, 2, 1])`)
}})
```

`flatten(["A", [[[["B"]]]], "C"])` має повернути `["A", "B", "C"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten(["A", [[[["B"]]]], "C"]), ["A", "B", "C"])`)
}})
```

`flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]])` має повернути `["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]]), ["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"])`)
}})
```

`flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]])` має повернути `["red","blue","green","yellow","purple","orange","pink","brown"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]]), ["red","blue","green","yellow","purple","orange","pink","brown"])`)
}})
```

# --seed--

## --seed-contents--

```py
def flatten(arr):

    return arr
```

# --solutions--

```py
def flatten(arr):
    result = []
    for item in arr:
        if isinstance(item, list):
            result.extend(flatten(item))
        else:
            result.append(item)
    return result
```
