---
id: 697a49e6ff50d756c9b69365
title: "Завдання 188: зимові ігри 2026, день 9: скелетон"
challengeType: 29
dashedName: challenge-188
---

# --description--

Дано рядок, який представляє вигини на трасі для скелетону. Визначте складність траси.

- Даний рядок міститиме лише літери:
  - `"L"` для лівого повороту
  - `"R"` для правого повороту
  - `"S"` для прямої ділянки

- Кожна зміна напрямку додає 15 балів (`"L"`, за яким слідує `"R"` або навпаки).
- Усі інші вигини додають по 5 балів кожен (усі інші символи `"L"` або `"R"`).
- Прямі ділянки додають 0 балів.

Складність траси базується на загальній сумі балів. Поверніть:

- `"Easy"`, якщо сума від 0 до 100
- `"Medium"`, якщо сума від 101 до 200
- `"Hard"`, якщо сума понад 200

# --hints--

`get_difficulty("SLSLLSRRLSRLRL")` має повернути `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLSLLSRRLSRLRL"), "Easy")`)
}})
```

`get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` має повернути `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard")`)
}})
```

`get_difficulty("SRRRRLSLLRLRSSRLSRL")` має повернути `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRRRRLSLLRLRSSRLSRL"), "Medium")`)
}})
```

`get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` має повернути `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard")`)
}})
```

`get_difficulty("SLLSSLRLSLSLRSLSSLRL")` має повернути `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium")`)
}})
```

`get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` має повернути `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_difficulty(track):

    return track
```

# --solutions--

```py
def get_difficulty(track):
  score = 0

  for i in range(len(track)):
    current = track[i]
    previous = track[i - 1] if i > 0 else None

    if current == 'S':
      score += 0
    elif current == 'L' or current == 'R':
      if previous and previous != 'S' and previous != current:
        score += 15
      else:
        score += 5

  if score <= 100:
    return 'Easy'
  elif score <= 200:
    return 'Medium'
  else:
    return 'Hard'
```
