---
id: 697a49e6ff50d756c9b6935f
title: "Завдання 182: зимові ігри 2026, день 3: біатлон"
challengeType: 29
dashedName: challenge-182
---

# --description--

Дано масив цілих чисел, де кожне значення представляє кількість влучень у ціль за один раунд біатлону. Поверніть загальну штрафну дистанцію, яку спортсмен має пробігти.

- Кожен раунд складається з 5 мішеней.
- За кожну пропущену мішень спортсмен має пробігти штрафний цикл довжиною 150 метрів.

# --hints--

`calculate_penalty_distance([4, 4])` має повернути `300`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 4]), 300)`)
}})
```

`calculate_penalty_distance([5, 5])` має повернути `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 5]), 0)`)
}})
```

`calculate_penalty_distance([4, 5, 3, 5])` має повернути `450`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 5, 3, 5]), 450)`)
}})
```

`calculate_penalty_distance([5, 4, 5, 5])` має повернути `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 4, 5, 5]), 150)`)
}})
```

`calculate_penalty_distance([4, 3, 0, 3])` має повернути `1500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 3, 0, 3]), 1500)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_penalty_distance(rounds):

    return rounds
```

# --solutions--

```py
def calculate_penalty_distance(rounds):
    total_penalty = 0

    for hits in rounds:
        misses = 5 - hits
        total_penalty += misses * 150

    return total_penalty
```
