---
id: 697a49e9860d24853adef67d
title: "Завдання 192: зимові ігри 2026, день 13: лижне двоборство"
challengeType: 29
dashedName: challenge-192
---

# --description--

Дано масив балів за стрибки спортсменів, обчисліть час затримки їхнього старту для крос-кантрі у лижному двоборстві.

Спортсмен із найвищим балом за стрибок стартує першим (затримка 0 секунд). Всі інші спортсмени стартують пізніше залежно від того, наскільки їхній бал за стрибок відстає від найкращого.

Щоб обчислити затримку для кожного спортсмена, відніміть бал спортсмена за стрибок від найкращого загального балу і помножте результат на 1.5. Округліть затримку вгору до найближчого цілого числа.

# --hints--

`calculate_start_delays([120, 110, 125])` має повертати `[8, 23, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([120, 110, 125]), [8, 23, 0])`)
}})
```

`calculate_start_delays([118, 125, 122, 120])` має повертати `[11, 0, 5, 8]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([118, 125, 122, 120]), [11, 0, 5, 8])`)
}})
```

`calculate_start_delays([100, 105, 95, 110, 120, 115, 108])` має повертати `[30, 23, 38, 15, 0, 8, 18]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18])`)
}})
```

`calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` має повертати `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_start_delays(jump_scores):

    return jump_scores
```

# --solutions--

```py
import math
def calculate_start_delays(jump_scores):
    best_jump = max(jump_scores)
    return [math.ceil((best_jump - score) * 1.5) for score in jump_scores]
```
