---
id: 697a49e6ff50d756c9b6935d
title: "袟邪胁写邪薪薪褟 180: 蟹懈屑芯胁褨 褨谐褉懈 2026, 写械薪褜 1: 写械薪褜 胁褨写泻褉懈褌褌褟"
challengeType: 29
dashedName: challenge-180
---

# --description--

小褜芯谐芯写薪褨 锌芯褔懈薪邪褞褌褜褋褟 蟹懈屑芯胁褨 褨谐褉懈 2026 褉芯泻褍. 袧邪褋褌褍锌薪褨 17 写薪褨胁 锌褉懈薪械褋褍褌褜 胁邪屑 蟹邪胁写邪薪薪褟 蟹 泻芯写褍胁邪薪薪褟, 薪邪褌褏薪械薪薪褨 褑懈屑懈 褨谐褉邪屑懈.

袛谢褟 锌械褉褕芯谐芯 蟹邪胁写邪薪薪褟 胁邪屑 写邪褞褌褜 写胁芯谢褨褌械褉薪懈泄 泻芯写 泻褉邪褩薪懈, 褨 锌芯褌褉褨斜薪芯 锌芯胁械褉薪褍褌懈 械屑芯写蟹褨 锌褉邪锌芯褉褑褟 褑褨褦褩 泻褉邪褩薪懈.

袙懈泻芯褉懈褋褌芯胁褍泄褌械 褑械泄 褋锌懈褋芯泻:

| 袣褉邪褩薪邪                    | 袣芯写            | 袩褉邪锌芯褉械褑褜 |
| -                         | -:             | :--:      |
| Albania                   |            `"AL"` | `"馃嚘馃嚤"` |
| Andorra                   |            `"AD"` | `"馃嚘馃嚛"` |
| Argentina                 |            `"AR"` | `"馃嚘馃嚪"` |
| Armenia                   |            `"AM"` | `"馃嚘馃嚥"` |
| Australia                 |            `"AU"` | `"馃嚘馃嚭"` |
| Austria                   |            `"AT"` | `"馃嚘馃嚬"` |
| Azerbaijan                |            `"AZ"` | `"馃嚘馃嚳"` |
| Belgium                   |            `"BE"` | `"馃嚙馃嚜"` |
| Benin                     |            `"BJ"` | `"馃嚙馃嚡"` |
| Bolivia                   |            `"BO"` | `"馃嚙馃嚧"` |
| Bosnia and Herzegovina    |            `"BA"` | `"馃嚙馃嚘"` |
| Brazil                    |            `"BR"` | `"馃嚙馃嚪"` |
| Bulgaria                  |            `"BG"` | `"馃嚙馃嚞"` |
| Canada                    |            `"CA"` | `"馃嚚馃嚘"` |
| Chile                     |            `"CL"` | `"馃嚚馃嚤"` |
| China                     |            `"CN"` | `"馃嚚馃嚦"` |
| Colombia                  |            `"CO"` | `"馃嚚馃嚧"` |
| Croatia                   |            `"HR"` | `"馃嚟馃嚪"` |
| Cyprus                    |            `"CY"` | `"馃嚚馃嚲"` |
| Czech Republic            |            `"CZ"` | `"馃嚚馃嚳"` |
| Denmark                   |            `"DK"` | `"馃嚛馃嚢"` |
| Ecuador                   |            `"EC"` | `"馃嚜馃嚚"` |
| Eritrea                   |            `"ER"` | `"馃嚜馃嚪"` |
| Estonia                   |            `"EE"` | `"馃嚜馃嚜"` |
| Finland                   |            `"FI"` | `"馃嚝馃嚠"` |
| France                    |            `"FR"` | `"馃嚝馃嚪"` |
| Georgia                   |            `"GE"` | `"馃嚞馃嚜"` |
| Germany                   |            `"DE"` | `"馃嚛馃嚜"` |
| Great Britain             |            `"GB"` | `"馃嚞馃嚙"` |
| Greece                    |            `"GR"` | `"馃嚞馃嚪"` |
| Guinea-Bissau             |            `"GW"` | `"馃嚞馃嚰"` |
| Haiti                     |            `"HT"` | `"馃嚟馃嚬"` |
| Hong Kong                 |            `"HK"` | `"馃嚟馃嚢"` |
| Hungary                   |            `"HU"` | `"馃嚟馃嚭"` |
| Iceland                   |            `"IS"` | `"馃嚠馃嚫"` |
| India                     |            `"IN"` | `"馃嚠馃嚦"` |
| Iran                      |            `"IR"` | `"馃嚠馃嚪"` |
| Ireland                   |            `"IE"` | `"馃嚠馃嚜"` |
| Israel                    |            `"IL"` | `"馃嚠馃嚤"` |
| Italy                     |            `"IT"` | `"馃嚠馃嚬"` |
| Jamaica                   |            `"JM"` | `"馃嚡馃嚥"` |
| Japan                     |            `"JP"` | `"馃嚡馃嚨"` |
| Kazakhstan                |            `"KZ"` | `"馃嚢馃嚳"` |
| Kenya                     |            `"KE"` | `"馃嚢馃嚜"` |
| Kosovo                    |            `"XK"` | `"馃嚱馃嚢"` |
| Kyrgyzstan                |            `"KG"` | `"馃嚢馃嚞"` |
| Latvia                    |            `"LV"` | `"馃嚤馃嚮"` |
| Lebanon                   |            `"LB"` | `"馃嚤馃嚙"` |
| Liechtenstein             |            `"LI"` | `"馃嚤馃嚠"` |
| Lithuania                 |            `"LT"` | `"馃嚤馃嚬"` |
| Luxembourg                |            `"LU"` | `"馃嚤馃嚭"` |
| Madagascar                |            `"MG"` | `"馃嚥馃嚞"` |
| Malaysia                  |            `"MY"` | `"馃嚥馃嚲"` |
| Malta                     |            `"MT"` | `"馃嚥馃嚬"` |
| Mexico                    |            `"MX"` | `"馃嚥馃嚱"` |
| Moldova                   |            `"MD"` | `"馃嚥馃嚛"` |
| Monaco                    |            `"MC"` | `"馃嚥馃嚚"` |
| Mongolia                  |            `"MN"` | `"馃嚥馃嚦"` |
| Montenegro                |            `"ME"` | `"馃嚥馃嚜"` |
| Morocco                   |            `"MA"` | `"馃嚥馃嚘"` |
| Netherlands               |            `"NL"` | `"馃嚦馃嚤"` |
| New Zealand               |            `"NZ"` | `"馃嚦馃嚳"` |
| Nigeria                   |            `"NG"` | `"馃嚦馃嚞"` |
| North Macedonia           |            `"MK"` | `"馃嚥馃嚢"` |
| Norway                    |            `"NO"` | `"馃嚦馃嚧"` |
| Pakistan                  |            `"PK"` | `"馃嚨馃嚢"` |
| Philippines               |            `"PH"` | `"馃嚨馃嚟"` |
| Poland                    |            `"PL"` | `"馃嚨馃嚤"` |
| Portugal                  |            `"PT"` | `"馃嚨馃嚬"` |
| Puerto Rico               |            `"PR"` | `"馃嚨馃嚪"` |
| Romania                   |            `"RO"` | `"馃嚪馃嚧"` |
| San Marino                |            `"SM"` | `"馃嚫馃嚥"` |
| Saudi Arabia              |            `"SA"` | `"馃嚫馃嚘"` |
| Serbia                    |            `"RS"` | `"馃嚪馃嚫"` |
| Singapore                 |            `"SG"` | `"馃嚫馃嚞"` |
| Slovakia                  |            `"SK"` | `"馃嚫馃嚢"` |
| Slovenia                  |            `"SI"` | `"馃嚫馃嚠"` |
| South Africa              |            `"ZA"` | `"馃嚳馃嚘"` |
| South Korea               |            `"KR"` | `"馃嚢馃嚪"` |
| Spain                     |            `"ES"` | `"馃嚜馃嚫"` |
| Sweden                    |            `"SE"` | `"馃嚫馃嚜"` |
| Switzerland               |            `"CH"` | `"馃嚚馃嚟"` |
| Thailand                  |            `"TH"` | `"馃嚬馃嚟"` |
| Trinidad & Tobago         |            `"TT"` | `"馃嚬馃嚬"` |
| Turkey                    |            `"TR"` | `"馃嚬馃嚪"` |
| Ukraine                   |            `"UA"` | `"馃嚭馃嚘"` |
| United Arab Emirates      |            `"AE"` | `"馃嚘馃嚜"` |
| United States             |            `"US"` | `"馃嚭馃嚫"` |
| Uruguay                   |            `"UY"` | `"馃嚭馃嚲"` |
| Uzbekistan                |            `"UZ"` | `"馃嚭馃嚳"` |
| Venezuela                 |            `"VE"` | `"馃嚮馃嚜"` |

# --hints--

`get_flag("AL")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚤"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AL"), "馃嚘馃嚤")`)
}})
```

`get_flag("AD")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚛"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AD"), "馃嚘馃嚛")`)
}})
```

`get_flag("AR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AR"), "馃嚘馃嚪")`)
}})
```

`get_flag("AM")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚥"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AM"), "馃嚘馃嚥")`)
}})
```

`get_flag("AU")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AU"), "馃嚘馃嚭")`)
}})
```

`get_flag("AT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AT"), "馃嚘馃嚬")`)
}})
```

`get_flag("AZ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AZ"), "馃嚘馃嚳")`)
}})
```

`get_flag("BE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚜"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BE"), "馃嚙馃嚜")`)
}})
```

`get_flag("BJ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚡"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BJ"), "馃嚙馃嚡")`)
}})
```

`get_flag("BO")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BO"), "馃嚙馃嚧")`)
}})
```

`get_flag("BA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚘"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BA"), "馃嚙馃嚘")`)
}})
```

`get_flag("BR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BR"), "馃嚙馃嚪")`)
}})
```

`get_flag("BG")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚞"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BG"), "馃嚙馃嚞")`)
}})
```

`get_flag("CA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚘"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CA"), "馃嚚馃嚘")`)
}})
```

`get_flag("CL")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚤"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CL"), "馃嚚馃嚤")`)
}})
```

`get_flag("CN")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CN"), "馃嚚馃嚦")`)
}})
```

`get_flag("CO")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CO"), "馃嚚馃嚧")`)
}})
```

`get_flag("HR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚟馃嚪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HR"), "馃嚟馃嚪")`)
}})
```

`get_flag("CY")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CY"), "馃嚚馃嚲")`)
}})
```

`get_flag("CZ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CZ"), "馃嚚馃嚳")`)
}})
```

`get_flag("DK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚛馃嚢"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DK"), "馃嚛馃嚢")`)
}})
```

`get_flag("EC")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚜馃嚚"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EC"), "馃嚜馃嚚")`)
}})
```

`get_flag("ER")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚜馃嚪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ER"), "馃嚜馃嚪")`)
}})
```

`get_flag("EE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚜馃嚜"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EE"), "馃嚜馃嚜")`)
}})
```

`get_flag("FI")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚝馃嚠"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FI"), "馃嚝馃嚠")`)
}})
```

`get_flag("FR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚝馃嚪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FR"), "馃嚝馃嚪")`)
}})
```

`get_flag("GE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚞馃嚜"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GE"), "馃嚞馃嚜")`)
}})
```

`get_flag("DE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚛馃嚜"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DE"), "馃嚛馃嚜")`)
}})
```

`get_flag("GB")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚞馃嚙"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GB"), "馃嚞馃嚙")`)
}})
```

`get_flag("GR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚞馃嚪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GR"), "馃嚞馃嚪")`)
}})
```

`get_flag("GW")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚞馃嚰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GW"), "馃嚞馃嚰")`)
}})
```

`get_flag("HT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚟馃嚬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HT"), "馃嚟馃嚬")`)
}})
```

`get_flag("HK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚟馃嚢"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HK"), "馃嚟馃嚢")`)
}})
```

`get_flag("HU")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚟馃嚭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HU"), "馃嚟馃嚭")`)
}})
```

`get_flag("IS")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚫"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IS"), "馃嚠馃嚫")`)
}})
```

`get_flag("IN")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IN"), "馃嚠馃嚦")`)
}})
```

`get_flag("IR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IR"), "馃嚠馃嚪")`)
}})
```

`get_flag("IE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚜"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IE"), "馃嚠馃嚜")`)
}})
```

`get_flag("IL")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚤"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IL"), "馃嚠馃嚤")`)
}})
```

`get_flag("IT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IT"), "馃嚠馃嚬")`)
}})
```

`get_flag("JM")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚡馃嚥"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JM"), "馃嚡馃嚥")`)
}})
```

`get_flag("JP")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚡馃嚨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JP"), "馃嚡馃嚨")`)
}})
```

`get_flag("KZ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚢馃嚳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KZ"), "馃嚢馃嚳")`)
}})
```

`get_flag("KE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚢馃嚜"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KE"), "馃嚢馃嚜")`)
}})
```

`get_flag("XK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚱馃嚢"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("XK"), "馃嚱馃嚢")`)
}})
```

`get_flag("KG")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚢馃嚞"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KG"), "馃嚢馃嚞")`)
}})
```

`get_flag("LV")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚤馃嚮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LV"), "馃嚤馃嚮")`)
}})
```

`get_flag("LB")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚤馃嚙"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LB"), "馃嚤馃嚙")`)
}})
```

`get_flag("LI")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚤馃嚠"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LI"), "馃嚤馃嚠")`)
}})
```

`get_flag("LT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚤馃嚬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LT"), "馃嚤馃嚬")`)
}})
```

`get_flag("LU")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚤馃嚭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LU"), "馃嚤馃嚭")`)
}})
```

`get_flag("MG")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚞"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MG"), "馃嚥馃嚞")`)
}})
```

`get_flag("MY")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MY"), "馃嚥馃嚲")`)
}})
```

`get_flag("MT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MT"), "馃嚥馃嚬")`)
}})
```

`get_flag("MX")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MX"), "馃嚥馃嚱")`)
}})
```

`get_flag("MD")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚛"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MD"), "馃嚥馃嚛")`)
}})
```

`get_flag("MC")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚚"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MC"), "馃嚥馃嚚")`)
}})
```

`get_flag("MN")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MN"), "馃嚥馃嚦")`)
}})
```

`get_flag("ME")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚜"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ME"), "馃嚥馃嚜")`)
}})
```

`get_flag("MA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚘"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MA"), "馃嚥馃嚘")`)
}})
```

`get_flag("NL")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚦馃嚤"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NL"), "馃嚦馃嚤")`)
}})
```

`get_flag("NZ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚦馃嚳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NZ"), "馃嚦馃嚳")`)
}})
```

`get_flag("NG")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚦馃嚞"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NG"), "馃嚦馃嚞")`)
}})
```

`get_flag("MK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚢"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MK"), "馃嚥馃嚢")`)
}})
```

`get_flag("NO")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚦馃嚧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NO"), "馃嚦馃嚧")`)
}})
```

`get_flag("PK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚨馃嚢"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PK"), "馃嚨馃嚢")`)
}})
```

`get_flag("PH")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚨馃嚟"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PH"), "馃嚨馃嚟")`)
}})
```

`get_flag("PL")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚨馃嚤"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PL"), "馃嚨馃嚤")`)
}})
```

`get_flag("PT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚨馃嚬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PT"), "馃嚨馃嚬")`)
}})
```

`get_flag("PR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚨馃嚪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PR"), "馃嚨馃嚪")`)
}})
```

`get_flag("RO")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚪馃嚧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RO"), "馃嚪馃嚧")`)
}})
```

`get_flag("SM")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚥"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SM"), "馃嚫馃嚥")`)
}})
```

`get_flag("SA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚘"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SA"), "馃嚫馃嚘")`)
}})
```

`get_flag("RS")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚪馃嚫"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RS"), "馃嚪馃嚫")`)
}})
```

`get_flag("SG")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚞"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SG"), "馃嚫馃嚞")`)
}})
```

`get_flag("SK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚢"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SK"), "馃嚫馃嚢")`)
}})
```

`get_flag("SI")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚠"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SI"), "馃嚫馃嚠")`)
}})
```

`get_flag("ZA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚳馃嚘"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ZA"), "馃嚳馃嚘")`)
}})
```

`get_flag("KR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚢馃嚪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KR"), "馃嚢馃嚪")`)
}})
```

`get_flag("ES")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚜馃嚫"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ES"), "馃嚜馃嚫")`)
}})
```

`get_flag("SE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚜"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SE"), "馃嚫馃嚜")`)
}})
```

`get_flag("CH")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚟"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CH"), "馃嚚馃嚟")`)
}})
```

`get_flag("TH")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚬馃嚟"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TH"), "馃嚬馃嚟")`)
}})
```

`get_flag("TT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚬馃嚬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TT"), "馃嚬馃嚬")`)
}})
```

`get_flag("TR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚬馃嚪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TR"), "馃嚬馃嚪")`)
}})
```

`get_flag("UA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚭馃嚘"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UA"), "馃嚭馃嚘")`)
}})
```

`get_flag("AE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚜"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AE"), "馃嚘馃嚜")`)
}})
```

`get_flag("US")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚭馃嚫"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("US"), "馃嚭馃嚫")`)
}})
```

`get_flag("UY")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚭馃嚲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UY"), "馃嚭馃嚲")`)
}})
```

`get_flag("UZ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚭馃嚳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UZ"), "馃嚭馃嚳")`)
}})
```

`get_flag("VE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚮馃嚜"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("VE"), "馃嚮馃嚜")`)
}})
```



# --seed--

## --seed-contents--

```py
def get_flag(code):

    return code
```

# --solutions--

```py
def get_flag(code):
    return "".join(chr(127397 + ord(char)) for char in code)
```
