---
id: 697a49e9860d24853adef67c
title: "Завдання 191: зимові ігри 2026, день 12: бобслей"
challengeType: 29
dashedName: challenge-191
---

# --description--

Дано масив, який представляє ваги спортсменів у команді з бобслею, та число, що представляє вагу бобслея. Визначте, чи команда має право брати участь у змаганні.

- Довжина масиву визначає розмір команди: 1, 2 або 4 особи.
- Усі задані значення ваги вказані в кілограмах (кг).

Бобслей (сані без спортсменів) повинен мати мінімальну вагу:

- 162 кг для команди з 1 особи
- 170 кг для команди з 2 осіб
- 210 кг для команди з 4 осіб

Загальна вага бобслея (спортсмени плюс сани) не повинна перевищувати:

- 247 кг для команди з 1 особи
- 390 кг для команди з 2 осіб
- 630 кг для команди з 4 осіб

Поверніть `"Eligible"`, якщо команда відповідає всім вимогам, або `"Not Eligible"`, якщо команда не виконує одну або кілька вимог.

# --hints--

`check_eligibility([78], 165)` має повертати `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([78], 165), "Eligible")`)
}})
```

`check_eligibility([80], 160)` має повертати `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 160), "Not Eligible")`)
}})
```

`check_eligibility([80], 170)` має повертати `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 170), "Not Eligible")`)
}})
```

`check_eligibility([85, 90], 170)` має повертати `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 90], 170), "Eligible")`)
}})
```

`check_eligibility([85, 95], 168)` має повертати `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 95], 168), "Not Eligible")`)
}})
```

`check_eligibility([112, 97], 185)` має повертати `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([112, 97], 185), "Not Eligible")`)
}})
```

`check_eligibility([110, 102, 90, 106], 222)` має повертати `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([110, 102, 90, 106], 222), "Eligible")`)
}})
```

`check_eligibility([106, 99, 90, 88], 205)` має повертати `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 90, 88], 205), "Not Eligible")`)
}})
```

`check_eligibility([106, 99, 103, 96], 227)` має повертати `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 103, 96], 227), "Not Eligible")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_eligibility(athlete_weights, sled_weight):

    return athlete_weights
```

# --solutions--

```py
def check_eligibility(athlete_weights, sled_weight):
    team_size = len(athlete_weights)

    rules = {
        1: {"min_sled": 162, "max_total": 247},
        2: {"min_sled": 170, "max_total": 390},
        4: {"min_sled": 210, "max_total": 630}
    }

    min_sled = rules[team_size]["min_sled"]
    max_total = rules[team_size]["max_total"]

    total_weight = sum(athlete_weights) + sled_weight

    if sled_weight < min_sled:
        return "Not Eligible"

    if total_weight > max_total:
        return "Not Eligible"

    return "Eligible"
```
