---
id: 697a49e6ff50d756c9b6935d
title: "Завдання 180: Зимові ігри, день 1: день відкриття"
challengeType: 29
dashedName: challenge-180
---

# --description--

Наступні 17 днів принесуть вам завдання з програмування, натхненні Зимовими іграми.

Для першого завдання вам дають дволітерний код країни, і потрібно повернути емодзі прапорця цієї країни.

Використовуйте цей список:

| Країна                    | Код            | Прапорець |
| -                         | -:             | :--:      |
| Albania                   |            `"AL"` | `"🇦🇱"` |
| Andorra                   |            `"AD"` | `"🇦🇩"` |
| Argentina                 |            `"AR"` | `"🇦🇷"` |
| Armenia                   |            `"AM"` | `"🇦🇲"` |
| Australia                 |            `"AU"` | `"🇦🇺"` |
| Austria                   |            `"AT"` | `"🇦🇹"` |
| Azerbaijan                |            `"AZ"` | `"🇦🇿"` |
| Belgium                   |            `"BE"` | `"🇧🇪"` |
| Benin                     |            `"BJ"` | `"🇧🇯"` |
| Bolivia                   |            `"BO"` | `"🇧🇴"` |
| Bosnia and Herzegovina    |            `"BA"` | `"🇧🇦"` |
| Brazil                    |            `"BR"` | `"🇧🇷"` |
| Bulgaria                  |            `"BG"` | `"🇧🇬"` |
| Canada                    |            `"CA"` | `"🇨🇦"` |
| Chile                     |            `"CL"` | `"🇨🇱"` |
| China                     |            `"CN"` | `"🇨🇳"` |
| Colombia                  |            `"CO"` | `"🇨🇴"` |
| Croatia                   |            `"HR"` | `"🇭🇷"` |
| Cyprus                    |            `"CY"` | `"🇨🇾"` |
| Czech Republic            |            `"CZ"` | `"🇨🇿"` |
| Denmark                   |            `"DK"` | `"🇩🇰"` |
| Ecuador                   |            `"EC"` | `"🇪🇨"` |
| Eritrea                   |            `"ER"` | `"🇪🇷"` |
| Estonia                   |            `"EE"` | `"🇪🇪"` |
| Finland                   |            `"FI"` | `"🇫🇮"` |
| France                    |            `"FR"` | `"🇫🇷"` |
| Georgia                   |            `"GE"` | `"🇬🇪"` |
| Germany                   |            `"DE"` | `"🇩🇪"` |
| Great Britain             |            `"GB"` | `"🇬🇧"` |
| Greece                    |            `"GR"` | `"🇬🇷"` |
| Guinea-Bissau             |            `"GW"` | `"🇬🇼"` |
| Haiti                     |            `"HT"` | `"🇭🇹"` |
| Hong Kong                 |            `"HK"` | `"🇭🇰"` |
| Hungary                   |            `"HU"` | `"🇭🇺"` |
| Iceland                   |            `"IS"` | `"🇮🇸"` |
| India                     |            `"IN"` | `"🇮🇳"` |
| Iran                      |            `"IR"` | `"🇮🇷"` |
| Ireland                   |            `"IE"` | `"🇮🇪"` |
| Israel                    |            `"IL"` | `"🇮🇱"` |
| Italy                     |            `"IT"` | `"🇮🇹"` |
| Jamaica                   |            `"JM"` | `"🇯🇲"` |
| Japan                     |            `"JP"` | `"🇯🇵"` |
| Kazakhstan                |            `"KZ"` | `"🇰🇿"` |
| Kenya                     |            `"KE"` | `"🇰🇪"` |
| Kosovo                    |            `"XK"` | `"🇽🇰"` |
| Kyrgyzstan                |            `"KG"` | `"🇰🇬"` |
| Latvia                    |            `"LV"` | `"🇱🇻"` |
| Lebanon                   |            `"LB"` | `"🇱🇧"` |
| Liechtenstein             |            `"LI"` | `"🇱🇮"` |
| Lithuania                 |            `"LT"` | `"🇱🇹"` |
| Luxembourg                |            `"LU"` | `"🇱🇺"` |
| Madagascar                |            `"MG"` | `"🇲🇬"` |
| Malaysia                  |            `"MY"` | `"🇲🇾"` |
| Malta                     |            `"MT"` | `"🇲🇹"` |
| Mexico                    |            `"MX"` | `"🇲🇽"` |
| Moldova                   |            `"MD"` | `"🇲🇩"` |
| Monaco                    |            `"MC"` | `"🇲🇨"` |
| Mongolia                  |            `"MN"` | `"🇲🇳"` |
| Montenegro                |            `"ME"` | `"🇲🇪"` |
| Morocco                   |            `"MA"` | `"🇲🇦"` |
| Netherlands               |            `"NL"` | `"🇳🇱"` |
| New Zealand               |            `"NZ"` | `"🇳🇿"` |
| Nigeria                   |            `"NG"` | `"🇳🇬"` |
| North Macedonia           |            `"MK"` | `"🇲🇰"` |
| Norway                    |            `"NO"` | `"🇳🇴"` |
| Pakistan                  |            `"PK"` | `"🇵🇰"` |
| Philippines               |            `"PH"` | `"🇵🇭"` |
| Poland                    |            `"PL"` | `"🇵🇱"` |
| Portugal                  |            `"PT"` | `"🇵🇹"` |
| Puerto Rico               |            `"PR"` | `"🇵🇷"` |
| Romania                   |            `"RO"` | `"🇷🇴"` |
| San Marino                |            `"SM"` | `"🇸🇲"` |
| Saudi Arabia              |            `"SA"` | `"🇸🇦"` |
| Serbia                    |            `"RS"` | `"🇷🇸"` |
| Singapore                 |            `"SG"` | `"🇸🇬"` |
| Slovakia                  |            `"SK"` | `"🇸🇰"` |
| Slovenia                  |            `"SI"` | `"🇸🇮"` |
| South Africa              |            `"ZA"` | `"🇿🇦"` |
| South Korea               |            `"KR"` | `"🇰🇷"` |
| Spain                     |            `"ES"` | `"🇪🇸"` |
| Sweden                    |            `"SE"` | `"🇸🇪"` |
| Switzerland               |            `"CH"` | `"🇨🇭"` |
| Thailand                  |            `"TH"` | `"🇹🇭"` |
| Trinidad & Tobago         |            `"TT"` | `"🇹🇹"` |
| Turkey                    |            `"TR"` | `"🇹🇷"` |
| Ukraine                   |            `"UA"` | `"🇺🇦"` |
| United Arab Emirates      |            `"AE"` | `"🇦🇪"` |
| United States             |            `"US"` | `"🇺🇸"` |
| Uruguay                   |            `"UY"` | `"🇺🇾"` |
| Uzbekistan                |            `"UZ"` | `"🇺🇿"` |
| Venezuela                 |            `"VE"` | `"🇻🇪"` |

# --hints--

`get_flag("AL")` має повертати `"🇦🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AL"), "🇦🇱")`)
}})
```

`get_flag("AD")` має повертати `"🇦🇩"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AD"), "🇦🇩")`)
}})
```

`get_flag("AR")` має повертати `"🇦🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AR"), "🇦🇷")`)
}})
```

`get_flag("AM")` має повертати `"🇦🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AM"), "🇦🇲")`)
}})
```

`get_flag("AU")` має повертати `"🇦🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AU"), "🇦🇺")`)
}})
```

`get_flag("AT")` має повертати `"🇦🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AT"), "🇦🇹")`)
}})
```

`get_flag("AZ")` має повертати `"🇦🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AZ"), "🇦🇿")`)
}})
```

`get_flag("BE")` має повертати `"🇧🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BE"), "🇧🇪")`)
}})
```

`get_flag("BJ")` має повертати `"🇧🇯"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BJ"), "🇧🇯")`)
}})
```

`get_flag("BO")` має повертати `"🇧🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BO"), "🇧🇴")`)
}})
```

`get_flag("BA")` має повертати `"🇧🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BA"), "🇧🇦")`)
}})
```

`get_flag("BR")` має повертати `"🇧🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BR"), "🇧🇷")`)
}})
```

`get_flag("BG")` має повертати `"🇧🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BG"), "🇧🇬")`)
}})
```

`get_flag("CA")` має повертати `"🇨🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CA"), "🇨🇦")`)
}})
```

`get_flag("CL")` має повертати `"🇨🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CL"), "🇨🇱")`)
}})
```

`get_flag("CN")` має повертати `"🇨🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CN"), "🇨🇳")`)
}})
```

`get_flag("CO")` має повертати `"🇨🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CO"), "🇨🇴")`)
}})
```

`get_flag("HR")` має повертати `"🇭🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HR"), "🇭🇷")`)
}})
```

`get_flag("CY")` має повертати `"🇨🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CY"), "🇨🇾")`)
}})
```

`get_flag("CZ")` має повертати `"🇨🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CZ"), "🇨🇿")`)
}})
```

`get_flag("DK")` має повертати `"🇩🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DK"), "🇩🇰")`)
}})
```

`get_flag("EC")` має повертати `"🇪🇨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EC"), "🇪🇨")`)
}})
```

`get_flag("ER")` має повертати `"🇪🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ER"), "🇪🇷")`)
}})
```

`get_flag("EE")` має повертати `"🇪🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EE"), "🇪🇪")`)
}})
```

`get_flag("FI")` має повертати `"🇫🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FI"), "🇫🇮")`)
}})
```

`get_flag("FR")` має повертати `"🇫🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FR"), "🇫🇷")`)
}})
```

`get_flag("GE")` має повертати `"🇬🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GE"), "🇬🇪")`)
}})
```

`get_flag("DE")` має повертати `"🇩🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DE"), "🇩🇪")`)
}})
```

`get_flag("GB")` має повертати `"🇬🇧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GB"), "🇬🇧")`)
}})
```

`get_flag("GR")` має повертати `"🇬🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GR"), "🇬🇷")`)
}})
```

`get_flag("GW")` має повертати `"🇬🇼"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GW"), "🇬🇼")`)
}})
```

`get_flag("HT")` має повертати `"🇭🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HT"), "🇭🇹")`)
}})
```

`get_flag("HK")` має повертати `"🇭🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HK"), "🇭🇰")`)
}})
```

`get_flag("HU")` має повертати `"🇭🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HU"), "🇭🇺")`)
}})
```

`get_flag("IS")` має повертати `"🇮🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IS"), "🇮🇸")`)
}})
```

`get_flag("IN")` має повертати `"🇮🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IN"), "🇮🇳")`)
}})
```

`get_flag("IR")` має повертати `"🇮🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IR"), "🇮🇷")`)
}})
```

`get_flag("IE")` має повертати `"🇮🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IE"), "🇮🇪")`)
}})
```

`get_flag("IL")` має повертати `"🇮🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IL"), "🇮🇱")`)
}})
```

`get_flag("IT")` має повертати `"🇮🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IT"), "🇮🇹")`)
}})
```

`get_flag("JM")` має повертати `"🇯🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JM"), "🇯🇲")`)
}})
```

`get_flag("JP")` має повертати `"🇯🇵"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JP"), "🇯🇵")`)
}})
```

`get_flag("KZ")` має повертати `"🇰🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KZ"), "🇰🇿")`)
}})
```

`get_flag("KE")` має повертати `"🇰🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KE"), "🇰🇪")`)
}})
```

`get_flag("XK")` має повертати `"🇽🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("XK"), "🇽🇰")`)
}})
```

`get_flag("KG")` має повертати `"🇰🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KG"), "🇰🇬")`)
}})
```

`get_flag("LV")` має повертати `"🇱🇻"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LV"), "🇱🇻")`)
}})
```

`get_flag("LB")` має повертати `"🇱🇧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LB"), "🇱🇧")`)
}})
```

`get_flag("LI")` має повертати `"🇱🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LI"), "🇱🇮")`)
}})
```

`get_flag("LT")` має повертати `"🇱🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LT"), "🇱🇹")`)
}})
```

`get_flag("LU")` має повертати `"🇱🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LU"), "🇱🇺")`)
}})
```

`get_flag("MG")` має повертати `"🇲🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MG"), "🇲🇬")`)
}})
```

`get_flag("MY")` має повертати `"🇲🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MY"), "🇲🇾")`)
}})
```

`get_flag("MT")` має повертати `"🇲🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MT"), "🇲🇹")`)
}})
```

`get_flag("MX")` має повертати `"🇲🇽"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MX"), "🇲🇽")`)
}})
```

`get_flag("MD")` має повертати `"🇲🇩"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MD"), "🇲🇩")`)
}})
```

`get_flag("MC")` має повертати `"🇲🇨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MC"), "🇲🇨")`)
}})
```

`get_flag("MN")` має повертати `"🇲🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MN"), "🇲🇳")`)
}})
```

`get_flag("ME")` має повертати `"🇲🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ME"), "🇲🇪")`)
}})
```

`get_flag("MA")` має повертати `"🇲🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MA"), "🇲🇦")`)
}})
```

`get_flag("NL")` має повертати `"🇳🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NL"), "🇳🇱")`)
}})
```

`get_flag("NZ")` має повертати `"🇳🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NZ"), "🇳🇿")`)
}})
```

`get_flag("NG")` має повертати `"🇳🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NG"), "🇳🇬")`)
}})
```

`get_flag("MK")` має повертати `"🇲🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MK"), "🇲🇰")`)
}})
```

`get_flag("NO")` має повертати `"🇳🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NO"), "🇳🇴")`)
}})
```

`get_flag("PK")` має повертати `"🇵🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PK"), "🇵🇰")`)
}})
```

`get_flag("PH")` має повертати `"🇵🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PH"), "🇵🇭")`)
}})
```

`get_flag("PL")` має повертати `"🇵🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PL"), "🇵🇱")`)
}})
```

`get_flag("PT")` має повертати `"🇵🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PT"), "🇵🇹")`)
}})
```

`get_flag("PR")` має повертати `"🇵🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PR"), "🇵🇷")`)
}})
```

`get_flag("RO")` має повертати `"🇷🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RO"), "🇷🇴")`)
}})
```

`get_flag("SM")` має повертати `"🇸🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SM"), "🇸🇲")`)
}})
```

`get_flag("SA")` має повертати `"🇸🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SA"), "🇸🇦")`)
}})
```

`get_flag("RS")` має повертати `"🇷🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RS"), "🇷🇸")`)
}})
```

`get_flag("SG")` має повертати `"🇸🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SG"), "🇸🇬")`)
}})
```

`get_flag("SK")` має повертати `"🇸🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SK"), "🇸🇰")`)
}})
```

`get_flag("SI")` має повертати `"🇸🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SI"), "🇸🇮")`)
}})
```

`get_flag("ZA")` має повертати `"🇿🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ZA"), "🇿🇦")`)
}})
```

`get_flag("KR")` має повертати `"🇰🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KR"), "🇰🇷")`)
}})
```

`get_flag("ES")` має повертати `"🇪🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ES"), "🇪🇸")`)
}})
```

`get_flag("SE")` має повертати `"🇸🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SE"), "🇸🇪")`)
}})
```

`get_flag("CH")` має повертати `"🇨🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CH"), "🇨🇭")`)
}})
```

`get_flag("TH")` має повертати `"🇹🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TH"), "🇹🇭")`)
}})
```

`get_flag("TT")` має повертати `"🇹🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TT"), "🇹🇹")`)
}})
```

`get_flag("TR")` має повертати `"🇹🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TR"), "🇹🇷")`)
}})
```

`get_flag("UA")` має повертати `"🇺🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UA"), "🇺🇦")`)
}})
```

`get_flag("AE")` має повертати `"🇦🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AE"), "🇦🇪")`)
}})
```

`get_flag("US")` має повертати `"🇺🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("US"), "🇺🇸")`)
}})
```

`get_flag("UY")` має повертати `"🇺🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UY"), "🇺🇾")`)
}})
```

`get_flag("UZ")` має повертати `"🇺🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UZ"), "🇺🇿")`)
}})
```

`get_flag("VE")` має повертати `"🇻🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("VE"), "🇻🇪")`)
}})
```



# --seed--

## --seed-contents--

```py
def get_flag(code):

    return code
```

# --solutions--

```py
def get_flag(code):
    return "".join(chr(127397 + ord(char)) for char in code)
```
