---
id: 697a49e6ff50d756c9b69364
title: "Завдання 187: Зимові ігри 2026, день 8: санний спорт"
challengeType: 29
dashedName: challenge-187
---

# --description--

Дано масив із п’яти чисел, кожне з яких представляє час (у секундах), за який саночник пройшов сегмент траси. Визначте, у якому сегменті була найвища швидкість і якою вона була.

Траса поділена на такі сегменти:

- Сегмент 1: 320 метрів
- Сегмент 2: 280 метрів
- Сегмент 3: 350 метрів
- Сегмент 4: 300 метрів
- Сегмент 5: 250 метрів

Перше значення в масиві відповідає часу для сегмента 1, друге — для сегмента 2 і так далі.

Щоб обчислити швидкість (у метрах за секунду) для сегмента, поділіть відстань на час.

Поверніть `"The luger's fastest speed was X m/s on segment Y."`, де `X` — найвища швидкість, округлена до двох знаків після коми, а `Y` — номер сегмента, у якому була найвища швидкість.

# --hints--

`get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128])` має повертати `"The luger's fastest speed was 35.07 m/s on segment 5."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.")`)
}})
```

`get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284])` має повертати `"The luger's fastest speed was 37.75 m/s on segment 2."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.")`)
}})
```

`get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912])` має повертати `"The luger's fastest speed was 38.49 m/s on segment 3."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.")`)
}})
```

`get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840])` має повертати `"The luger's fastest speed was 37.69 m/s on segment 1."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.")`)
}})
```

`get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508])` має повертати `"The luger's fastest speed was 39.24 m/s on segment 4."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_fastest_speed(times):

    return times
```

# --solutions--

```py
def get_fastest_speed(times):
    distances = [320, 280, 350, 300, 250]
    fastest_speed = 0
    fastest_segment = 0

    for i, t in enumerate(times):
        speed = distances[i] / t
        if speed > fastest_speed:
            fastest_speed = speed
            fastest_segment = i + 1  # 1-indexed

    return f"The luger's fastest speed was {fastest_speed:.2f} m/s on segment {fastest_segment}."
```
