---
id: 697a49e9860d24853adef67f
title: "Завдання 194: зимові ігри 2026, день 15: фристайл лижний спорт"
challengeType: 29
dashedName: challenge-194
---

# --description--

Дано назву трюку, що складається з двох слів. Визначте, чи є вона дійсною назвою трюку у фристайл лижному спорті.

Трюк вважається дійсним, якщо перше слово є у списку дійсних перших слів, а друге слово — у списку дійсних других слів.

- Два слова будуть розділені одним пробілом.

Дійсні перші слова:

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

Дійсні другі слова:

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`is_valid_trick("Polar Vortex")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Polar Vortex"), True)`)
}})
```

`is_valid_trick("Solar Icequake")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Solar Icequake"), True)`)
}})
```

`is_valid_trick("Thunder Blizzard")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Thunder Blizzard"), True)`)
}})
```

`is_valid_trick("Phantom Frostbite")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Phantom Frostbite"), True)`)
}})
```

`is_valid_trick("Ghost Avalanche")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Ghost Avalanche"), True)`)
}})
```

`is_valid_trick("Snowstorm Shadow")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Snowstorm Shadow"), False)`)
}})
```

`is_valid_trick("Solar Sky")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Solar Sky"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_trick(trick_name):

    return trick_name
```

# --solutions--

```py
def is_valid_trick(trick_name):
    valid_first = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"]
    valid_second = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"]

    words = trick_name.split(" ")
    first, second = words
    return first in valid_first and second in valid_second
```
