---
id: 697a49e6ff50d756c9b69360
title: "Завдання 183: зимові ігри 2026, день 4: стрибки на лижах"
challengeType: 29
dashedName: challenge-183
---

# --description--

За заданими балами за дистанцію, стилем, компенсацією вітру та бонусом за K-точку обчисліть свій результат у стрибках на лижах і визначте, чи отримали ви медаль.

- Ваш результат обчислюється шляхом додавання цих чотирьох значень.

Поточні загальні бали інших стрибунів:

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- Якщо ваш результат найкращий, поверніть `"Gold"`
- Якщо другий за величиною, поверніть `"Silver"`
- Якщо третій за величиною, поверніть `"Bronze"`
- Інакше поверніть `"No Medal"`

# --hints--

`ski_jump_medal(125.0, 58.0, 0.0, 6.0)` має повертати `"Gold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(125.0, 58.0, 0.0, 6.0), "Gold")`)
}})
```

`ski_jump_medal(119.0, 50.0, 1.0, 4.0)` має повертати `"Bronze"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(119.0, 50.0, 1.0, 4.0), "Bronze")`)
}})
```

`ski_jump_medal(122.0, 52.0, -1.0, 4.0)` має повертати `"Silver"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(122.0, 52.0, -1.0, 4.0), "Silver")`)
}})
```

`ski_jump_medal(118.0, 50.5, -1.5, 4.0)` має повертати `"No Medal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(118.0, 50.5, -1.5, 4.0), "No Medal")`)
}})
```

`ski_jump_medal(124.0, 50.5, 2.0, 5.0)` має повертати `"Gold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(124.0, 50.5, 2.0, 5.0), "Gold")`)
}})
```

`ski_jump_medal(119.0, 49.5, 0.0, 3.0)` має повертати `"No Medal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(119.0, 49.5, 0.0, 3.0), "No Medal")`)
}})
```

# --seed--

## --seed-contents--

```py
def ski_jump_medal(distance_points, style_points, wind_comp, k_point_bonus):

    return distance_points
```

# --solutions--

```py
def ski_jump_medal(distance_points, style_points, wind_comp, k_point_bonus):
    my_score = distance_points + style_points + wind_comp + k_point_bonus
    other_scores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0]
    all_scores = other_scores + [my_score]
    all_scores.sort(reverse=True)
    rank = all_scores.index(my_score) + 1

    if rank == 1:
        return "Gold"
    elif rank == 2:
        return "Silver"
    elif rank == 3:
        return "Bronze"
    else:
        return "No Medal"
```
