---
id: 697a49e6ff50d756c9b6935e
title: "Завдання 181: зимові ігри 2026, день 2: сноубординг"
challengeType: 29
dashedName: challenge-181
---

# --description--

Визначте позу сноубордиста при приземленні, враховуючи його початкову позу та обертання в градусах.

- Поза сноубордиста може бути або `"Regular"`, або `"Goofy"`.
- Обертання трюків — кратні 90 градусів. Позитивне значення означає обертання за годинниковою стрілкою, від’ємне — проти годинникової стрілки.
- Поза при приземленні змінюється на протилежну кожні 180 градусів обертання.

Наприклад, якщо задано `"Regular"` і `90`, поверніть `"Regular"`. Якщо задано `"Regular"` і `180` градусів, поверніть `"Goofy"`.

# --hints--

`get_landing_stance("Regular", 90)` має повернути `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` має повернути `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` має повернути `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` має повернути `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` має повернути `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` має повернути `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
