---
id: 698a1a73ade5ac0e19180fa8
title: "Завдання 205: підрахунок ідеальних кубів"
challengeType: 29
dashedName: challenge-205
---

# --description--

Дано два цілі числа, визначте, скільки ідеальних кубів існує в діапазоні між цими двома числами включно.

- Нижче число не обов’язково є першим аргументом.
- Число є ідеальним кубом, якщо існує ціле число (`n`), для якого `n * n * n = number`. Наприклад, 27 є ідеальним кубом, тому що `3 * 3 * 3 = 27`.

# --hints--

`count_perfect_cubes(3, 30)` має повертати `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(3, 30), 2)`)
}})
```

`count_perfect_cubes(1, 30)` має повертати `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(1, 30), 3)`)
}})
```

`count_perfect_cubes(30, 0)` має повертати `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(30, 0), 4)`)
}})
```

`count_perfect_cubes(-64, 64)` має повертати `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(-64, 64), 9)`)
}})
```

`count_perfect_cubes(9214, -8127)` має повертати `41`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(9214, -8127), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_perfect_cubes(a, b):

    return a
```

# --solutions--

```py
def count_perfect_cubes(a, b):
    start, end = min(a, b), max(a, b)
    count = 0

    n = 0
    while n ** 3 <= end:
        if n ** 3 >= start:
            count += 1
        n += 1

    n = -1
    while n ** 3 >= start:
        if n ** 3 <= end:
            count += 1
        n -= 1

    return count
```
