---
id: 698a1a73ade5ac0e19180fa7
title: "Завдання 204: Сума літер"
challengeType: 29
dashedName: challenge-204
---

# --description--

Дано рядок, поверніть суму його літер.

- Літери — це A-Z у верхньому або нижньому регістрі
- Значення літер: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26
- Літери у верхньому та нижньому регістрі мають однакове значення.
- Ігноруйте всі символи, що не є літерами.

# --hints--

`sum_letters("Hello")` має повернути `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("Hello"), 52)`)
}})
```

`sum_letters("freeCodeCamp")` має повернути `94`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("freeCodeCamp"), 94)`)
}})
```

`sum_letters("The quick brown fox jumps over the lazy dog.")` має повернути `473`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("The quick brown fox jumps over the lazy dog."), 473)`)
}})
```

`sum_letters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` має повернути `1681`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681)`)
}})
```

`sum_letters("</404>")` має повернути `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("</404>"), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_letters(s):

    return s
```

# --solutions--

```py
def sum_letters(s):
    total = 0
    for char in s:
        upper = char.upper()
        if 'A' <= upper <= 'Z':
            total += ord(upper) - ord('A') + 1
    return total
```
