---
id: 698a1a73ade5ac0e19180fa4
title: "Завдання 201: зсув матриці"
challengeType: 29
dashedName: challenge-201
---

# --description--

Дано матрицю (масив масивів) чисел і ціле число. Зсунути всі значення в матриці на задану величину.

- Позитивний зсув переміщує значення вправо.
- Від’ємний зсув переміщує значення вліво.

Значення повинні обертатися:

- Розглядайте матрицю як одну безперервну послідовність значень.
- Коли значення виходить за межі рядка, воно продовжує рух з початку наступного рядка.
- Коли значення виходить за останню позицію в матриці, воно обертається на першу позицію.
- Те саме застосовується у зворотному напрямку при зсуві вліво.

Наприклад, для:

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

зі зсувом `1` перемістіть усі числа вправо на один:

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shift_matrix([[1, 2, 3], [4, 5, 6]], 1)` має повернути `[[6, 1, 2], [3, 4, 5]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6]], -1)` має повернути `[[2, 3, 4], [5, 6, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)` має повернути `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)` має повернути `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]])`)
}})
```

`shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)` має повернути `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]])`)
}})
```

`shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)` має повернути `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]])`)
}})
```

# --seed--

## --seed-contents--

```py
def shift_matrix(matrix, shift):

    return matrix
```

# --solutions--

```py
def shift_matrix(matrix, shift):
    rows = len(matrix)
    cols = len(matrix[0])
    flat = [num for row in matrix for num in row]
    n = len(flat)
    s = shift % n
    shifted = flat[-s:] + flat[:-s] if s != 0 else flat[:]
    result = []
    for i in range(rows):
        result.append(shifted[i*cols:(i+1)*cols])

    return result
```
