---
id: 698a1a73ade5ac0e19180fa0
title: "Завдання 197: сумісність груп крові"
challengeType: 29
dashedName: challenge-197
---

# --description--

За заданими групою крові донора та групою крові реципієнта визначте, чи може донор здати кров реципієнту.

Кожна група крові складається з:

- Літери: `"A"`, `"B"`, `"AB"` або `"O"`
- І резус-фактора: `"+"` або `"-"`

Групи крові будуть однією з дійсних літер, за якою слідує резус-фактор. Наприклад, `"AB+"` і `"O-"` — дійсні групи крові.

Правила для літер:

- `"O"` може здавати кров іншим типам літер.
- `"A"` може здавати кров `"A"` і `"AB"`.
- `"B"` може здавати кров `"B"` і `"AB"`.
- `"AB"` може здавати кров лише `"AB"`.

Правила для резус-фактора:

- Від’ємний (`"-"`) може здавати кров як `"-"`, так і `"+"`.
- Позитивний (`"+"`) може здавати кров лише `"+"`.

Для того, щоб донор міг здати кров реципієнту, мають виконуватися обидва правила — і для літери, і для резус-фактора.

# --hints--

`can_donate("B+", "B+")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B+", "B+"), True)`)
}})
```

`can_donate("O-", "AB-")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "AB-"), True)`)
}})
```

`can_donate("O+", "A-")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "A-"), False)`)
}})
```

`can_donate("A+", "AB+")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A+", "AB+"), True)`)
}})
```

`can_donate("A-", "B-")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A-", "B-"), False)`)
}})
```

`can_donate("B-", "AB+")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "AB+"), True)`)
}})
```

`can_donate("B-", "A+")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "A+"), False)`)
}})
```

`can_donate("O-", "O+")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "O+"), True)`)
}})
```

`can_donate("O+", "O-")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "O-"), False)`)
}})
```

`can_donate("AB+", "AB-")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("AB+", "AB-"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_donate(donor, recipient):

    return donor
```

# --solutions--

```py
def can_donate(donor, recipient):
    donor_type = donor[:-1]
    donor_rh = donor[-1]

    recipient_type = recipient[:-1]
    recipient_rh = recipient[-1]

    abo_compatibility = {
        "O": ["A", "B", "AB", "O"],
        "A": ["A", "AB"],
        "B": ["B", "AB"],
        "AB": ["AB"]
    }

    abo_match = recipient_type in abo_compatibility[donor_type]
    rh_match = (
        donor_rh == "-" or
        (donor_rh == "+" and recipient_rh == "+")
    )

    return abo_match and rh_match
```
