---
id: 698a1a73ade5ac0e19180fa9
title: "Завдання 206: значення ігрових карт"
challengeType: 29
dashedName: challenge-206
---

# --description--

Дано масив ігрових карт, поверніть новий масив із числовим значенням кожної карти.

Значення карт:

- Туз (`"A"`) має значення 1.
- Картки з числами (`"2"` - `"10"`) мають своє номінальне значення: відповідно від 2 до 10.
- Фігурні карти: Валет (`"J"`), Дама (`"Q"`) і Король (`"K"`) оцінюються в 10.

Масти:

- Кожна карта має масть: Піки (`"S"`), Трефи (`"C"`), Бубни (`"D"`) або Черви (`"H"`).

Формат карти:

- Кожна карта представлена у вигляді рядка: `"valueSuit"`. Наприклад: `"AS"` — це туз пік, `"10H"` — десятка черв, а `"QC"` — дама треф.

# --hints--

`card_values(["3H", "4D", "5S"])` має повернути `[3, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["3H", "4D", "5S"]), [3, 4, 5])`)
}})
```

`card_values(["AS", "10S", "10H", "6D", "7D"])` має повернути `[1, 10, 10, 6, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7])`)
}})
```

`card_values(["8D", "QS", "2H", "JC", "9C"])` має повернути `[8, 10, 2, 10, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9])`)
}})
```

`card_values(["AS", "KS"])` має повернути `[1, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "KS"]), [1, 10])`)
}})
```

`card_values(["10H", "JH", "QH", "KH", "AH"])` має повернути `[10, 10, 10, 10, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def card_values(cards):

    return cards
```

# --solutions--

```py
def card_values(cards):
    values = []
    for card in cards:
        value_str = card[:-1]

        if value_str == "A":
            value = 1
        elif value_str in ["J", "Q", "K"]:
            value = 10
        else:
            value = int(value_str)

        values.append(value)
    return values
```
