---
id: 698a1a863194f1f4e63f645e
title: "Завдання 207: найменший горизонтальний розрив"
challengeType: 29
dashedName: challenge-207
---

# --description--

Дано рядок, поверніть підрядок між двома ідентичними символами, між якими найменша кількість символів (найменший горизонтальний розрив).

- Завжди буде принаймні одна пара однакових символів.
- Повернутий підрядок не повинен включати однакові символи.
- Якщо два або більше горизонтальних розривів однакової довжини, поверніть символи з першого.

Наприклад, для `"ABCDAC"` поверніть `"DA"`, тому що:

- Лише `"A"` і `"C"` повторюються в рядку.
- Кількість символів між двома `"A"` символами дорівнює 3, а між `"C"` символами — 2.
- Тож поверніть рядок між двома `"C"` символами.

# --hints--

`smallest_gap("ABCDAC")` має повернути `"DA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("ABCDAC"), "DA")`)
}})
```

`smallest_gap("racecar")` має повернути `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("racecar"), "e")`)
}})
```

`smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` має повернути `"#q6e&rkf(p"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p")`)
}})
```

`smallest_gap("Hello World")` має повернути `""`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("Hello World"), "")`)
}})
```

`smallest_gap("The quick brown fox jumps over the lazy dog.")` має повернути `"fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("The quick brown fox jumps over the lazy dog."), "fox")`)
}})
```

# --seed--

## --seed-contents--

```py
def smallest_gap(s):

    return s
```

# --solutions--

```py
def smallest_gap(s):
    min_gap = float("inf")
    result = ""

    for i in range(len(s)):
        for j in range(i + 1, len(s)):
            if s[i] == s[j]:
                gap = j - i - 1

                if gap < min_gap:
                    min_gap = gap
                    result = s[i + 1:j]

                break

    return result
```
