---
id: 6994cff2290543b3aec9f50f
title: "Завдання 210: валідатор HSL"
challengeType: 29
dashedName: challenge-210
---

# --description--

Дано рядок, визначте, чи є він дійсним CSS-значенням кольору `hsl()` у форматі HSL.

- Дійсне значення HSL має бути у форматі `"hsl(h, s%, l%)"`, де:
  - `h` (відтінок) має бути числом від 0 до 360 (включно).
  - `s` (насиченість) має бути відсотком від 0% до 100%.
  - `l` (світлість) має бути відсотком від 0% до 100%.

- Пробіли дозволені лише:
  - Після відкриваючої дужки
  - Перед і/або після ком
  - Перед і/або після закриваючої дужки

- За бажанням, значення може закінчуватися крапкою з комою (`";"`).

Наприклад, `"hsl(240, 50%, 50%)"` є дійсним значенням HSL.

# --hints--

`is_valid_hsl("hsl(240, 50%, 50%)")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(240, 50%, 50%)"), True)`)
}})
```

`is_valid_hsl("hsl( 200 , 50% , 75% )")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl( 200 , 50% , 75% )"), True)`)
}})
```

`is_valid_hsl("hsl(99,60%,80%);")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(99,60%,80%);"), True)`)
}})
```

`is_valid_hsl("hsl(0, 0%, 0%) ;")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(0, 0%, 0%) ;"), True)`)
}})
```

`is_valid_hsl("hsl(  10  ,  20%   ,  30%   )    ;")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(  10  ,  20%   ,  30%   )    ;"), True)`)
}})
```

`is_valid_hsl("hsl(361, 50%, 80%)")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(361, 50%, 80%)"), False)`)
}})
```

`is_valid_hsl("hsl(300, 101%, 70%)")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(300, 101%, 70%)"), False)`)
}})
```

`is_valid_hsl("hsl(200, 55%, 75)")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(200, 55%, 75)"), False)`)
}})
```

`is_valid_hsl("hsl (80, 20%, 10%)")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl (80, 20%, 10%)"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_hsl(hsl):

    return hsl
```

# --solutions--

```py
import re
def is_valid_hsl(hsl):
    pattern = r"^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$"
    
    match = re.match(pattern, hsl)
    if not match:
        return False

    hue = int(match.group(1))
    sat = int(match.group(2))
    light = int(match.group(3))

    if not (0 <= hue <= 360):
        return False
    if not (0 <= sat <= 100):
        return False
    if not (0 <= light <= 100):
        return False

    return True
```
