---
id: 6994cff2290543b3aec9f511
title: "Завдання 212: вставлення в масив"
challengeType: 29
dashedName: challenge-212
---

# --description--

Дано масив, значення для вставлення в масив і індекс, за яким потрібно вставити це значення. Поверніть новий масив зі вставленим значенням на вказаному індексі.

# --hints--

`insert_into_array([2, 4, 8, 10], 6, 2)` має повертати `[2, 4, 6, 8, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(insert_into_array([2, 4, 8, 10], 6, 2), [2, 4, 6, 8, 10])`)
}})
```

`insert_into_array(["the", "quick", "fox"], "brown", 2)` має повертати `["the", "quick", "brown", "fox"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(insert_into_array(["the", "quick", "fox"], "brown", 2), ["the", "quick", "brown", "fox"])`)
}})
```

`insert_into_array([], 0, 0)` має повертати `[0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(insert_into_array([], 0, 0), [0])`)
}})
```

`insert_into_array([0, 1, 1, 2, 3, 8, 13], 5, 5)` має повертати `[0, 1, 1, 2, 3, 5, 8, 13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(insert_into_array([0, 1, 1, 2, 3, 8, 13], 5, 5), [0, 1, 1, 2, 3, 5, 8, 13])`)
}})
```

# --seed--

## --seed-contents--

```py
def insert_into_array(arr, value, index):

    return arr
```

# --solutions--

```py
def insert_into_array(arr, value, index):
    return arr[:index] + [value] + arr[index:]
```
