---
id: 699c8e045ee7cb94ed2322da
title: "Завдання 219: ювілейні віхи"
challengeType: 29
dashedName: challenge-219
---

# --description--

Дано ціле число, що представляє кількість років, які пара прожила у шлюбі. Поверніть їхню найновішу ювілейну віху відповідно до цієї діаграми:

| Роки шлюбу   | Віха         |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- Якщо вони не досягли першої віхи, поверніть `"Newlyweds"`.

# --hints--

`get_milestone(0)` має повертати `"Newlyweds"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(0), "Newlyweds")`)
}})
```

`get_milestone(1)` має повертати `"Paper"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(1), "Paper")`)
}})
```

`get_milestone(8)` має повертати `"Wood"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(8), "Wood")`)
}})
```

`get_milestone(10)` має повертати `"Tin"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(10), "Tin")`)
}})
```

`get_milestone(26)` має повертати `"Silver"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(26), "Silver")`)
}})
```

`get_milestone(45)` має повертати `"Ruby"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(45), "Ruby")`)
}})
```

`get_milestone(50)` має повертати `"Gold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(50), "Gold")`)
}})
```

`get_milestone(64)` має повертати `"Diamond"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(64), "Diamond")`)
}})
```

`get_milestone(71)` має повертати `"Platinum"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(71), "Platinum")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_milestone(years):

    return years
```

# --solutions--

```py
def get_milestone(years):
    if years < 1:
        return "Newlyweds"

    milestones = [
        (1, "Paper"),
        (5, "Wood"),
        (10, "Tin"),
        (25, "Silver"),
        (40, "Ruby"),
        (50, "Gold"),
        (60, "Diamond"),
        (70, "Platinum")
    ]

    for year, milestone in reversed(milestones):
        if years >= year:
            return milestone
```
