---
id: 699c8e045ee7cb94ed2322d8
title: "Завдання 217: захоплені шахові фігури"
challengeType: 29
dashedName: challenge-217
---

# --description--

Дано масив рядків, які представляють шахові фігури, що у вас ще залишилися на дошці. Обчисліть загальну вартість фігур, які захопив ваш суперник.

У шахах ви починаєте з 16 фігур:

| Фігура | Скорочення    | Кількість | Вартість |
| ------ | ------------- | --------- | -------- |
| Пішак  | `"P"`     | 8         | 1        |
| Тура   | `"R"`     | 2         | 5        |
| Кінь   | `"N"`     | 2         | 3        |
| Слон   | `"B"`     | 2         | 3        |
| Ферзь  | `"Q"`     | 1         | 9        |
| Король | `"K"`     | 1         | 0        |

- Даний масив міститиме лише наведені вище скорочення.
- Будь-які з 16 фігур, яких немає в масиві, були захоплені.
- Поверніть загальну вартість усіх захоплених фігур, якщо...
- Якщо Король був захоплений, поверніть `"Checkmate"`.

# --hints--

`get_captured_value(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` має повертати `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8)`)
}})
```

`get_captured_value(["P", "P", "P", "P", "P", "R", "B", "K"])` має повертати `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "P", "P", "P", "P", "R", "B", "K"]), 26)`)
}})
```

`get_captured_value(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` має повертати `16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16)`)
}})
```

`get_captured_value(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` має повертати `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4)`)
}})
```

`get_captured_value(["P", "K"])` має повертати `38`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "K"]), 38)`)
}})
```

`get_captured_value(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` має повертати `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0)`)
}})
```

`get_captured_value(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` має повертати `"Checkmate"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_captured_value(pieces):

    return pieces
```

# --solutions--

```py
def get_captured_value(pieces):
    if "K" not in pieces:
        return "Checkmate"

    values = {"P":1, "R":5, "N":3, "B":3, "Q":9, "K":0}
    remaining = "PPPPPPPPRRNNBBQK"

    for p in pieces:
        remaining = remaining.replace(p, "", 1)

    return sum(values[p] for p in remaining)
```
