---
id: 699c8e045ee7cb94ed2322d6
title: "Завдання 215: калькулятор плати за паркування"
challengeType: 29
dashedName: challenge-215
---

# --description--

Дано два рядки, які представляють час, коли ви припаркували свій автомобіль, і час, коли ви його забрали. Обчисліть плату за паркування.

- Дано рядки у форматі `"HH:MM"` з використанням 24-годинного формату. Наприклад, `"14:00"` означає 14:00.
- Перший рядок — це час, коли ви припаркували автомобіль, а другий — час, коли його забрали.
- Якщо час забору раніше за час паркування, це означає, що паркування тривало після опівночі до наступного дня.

Правила оплати:

- Кожна година паркування коштує 3 долари.
- Часткові години округлюються до наступної повної години.
- Якщо паркування триває вночі (після опівночі), додається додаткова плата за ніч у розмірі 10 доларів.
- Мінімальна плата становить 5 доларів (використовується лише якщо загальна сума менша за 5).

Поверніть загальну вартість у форматі `"$cost"`, наприклад `"$5"`.

# --hints--

`calculate_parking_fee("09:00", "11:00")` має повертати `"$6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("09:00", "11:00"), "$6")`)
}})
```

`calculate_parking_fee("10:00", "10:30")` має повертати `"$5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("10:00", "10:30"), "$5")`)
}})
```

`calculate_parking_fee("08:10", "10:45")` має повертати `"$9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("08:10", "10:45"), "$9")`)
}})
```

`calculate_parking_fee("14:40", "23:10")` має повертати `"$27"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("14:40", "23:10"), "$27")`)
}})
```

`calculate_parking_fee("18:15", "01:30")` має повертати `"$34"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("18:15", "01:30"), "$34")`)
}})
```

`calculate_parking_fee("11:11", "11:10")` має повертати `"$82"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("11:11", "11:10"), "$82")`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_parking_fee(park_time, pickup_time):

    return park_time
```

# --solutions--

```py
import math

def calculate_parking_fee(park_time, pickup_time):
    def to_minutes(time_str):
        hours, minutes = map(int, time_str.split(":"))
        return hours * 60 + minutes

    entry_minutes = to_minutes(park_time)
    exit_minutes = to_minutes(pickup_time)

    overnight = False

    if exit_minutes < entry_minutes:
        overnight = True
        total_minutes = (24 * 60 - entry_minutes) + exit_minutes
    else:
        total_minutes = exit_minutes - entry_minutes

    hours = math.ceil(total_minutes / 60)

    cost = hours * 3

    if overnight:
        cost += 10

    if cost < 5:
        cost = 5

    return f"${cost}"
```
