---
id: 699c8e045ee7cb94ed2322dd
title: "Завдання 222: тіні рівнодення"
challengeType: 29
dashedName: challenge-222
---

# --description--

Сьогодні рівнодення, коли сонце знаходиться безпосередньо над екватором і точно в зеніті опівдні. За заданим часом визначте тінь, яку відкидає вертикальний стовп висотою 4 фути.

- Час буде рядком у `"HH:MM"` 24-годинному форматі (наприклад, `"15:00"` — це 15:00).
- Час задається лише з кроком у 30 хвилин.

Правила:

- Сонце сходить о 6:00 прямо `"east"` і заходить о 18:00 прямо `"west"`.
- Тінь завжди спрямована в протилежний бік від сонця.
- Довжина тіні (у футах) — це число годин від опівдня, піднесене до куба.
- Тіні немає до сходу сонця (до 6:00), після заходу (18:00 або пізніше) та в опівдень.

Поверніть:

- Якщо тінь існує, поверніть `"(length)ft (direction)"`. Наприклад, `"8ft west"`.
- Інакше поверніть `"No shadow"`.

Наприклад, для `"10:00"` поверніть `"8ft west"`, бо 10:00 — це 2 години від опівдня, отже 2<sup>3</sup> = 8 футів, а тінь спрямована на захід, бо сонце о 10:00 на сході.

# --hints--

`get_shadow("10:00")` має повернути `"8ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` має повернути `"27ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` має повернути `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` має повернути `"166.375ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` має повернути `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` має повернути `"216ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` має повернути `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` має повернути `"91.125ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` має повернути `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
