---
id: 699c8e045ee7cb94ed2322d4
title: "Завдання 213: конвертер довжини слова"
challengeType: 29
dashedName: challenge-213
---

# --description--

Дано рядок слів, поверніть новий рядок, у якому кожне слово замінене на його довжину.

- Слова у заданому рядку будуть розділені одним пробілом
- Збережіть пробіли у поверненому рядку.

Наприклад, для `"hello world"` поверніть `"5 5"`.

# --hints--

`convert_words("hello world")` має повернути `"5 5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("hello world"), "5 5")`)
}})
```

`convert_words("Thanks and happy coding")` має повернути `"6 3 5 6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("Thanks and happy coding"), "6 3 5 6")`)
}})
```

`convert_words("The quick brown fox jumps over the lazy dog")` має повернути `"3 5 5 3 5 4 3 4 3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3")`)
}})
```

`convert_words("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` має повернути `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_words(s):

    return s
```

# --solutions--

```py
def convert_words(s):

    return " ".join(str(len(word)) for word in s.split(" "))
```
