---
id: 69a890af247de743333bd4cc
title: "Завдання 223: декодер QR"
challengeType: 29
dashedName: challenge-223
---

# --description--

Дано матрицю 6x6 (масив масивів), що представляє QR-код. Поверніть рядок бінарних даних у коді.

- QR-код може бути повернутий на будь-який кут, кратний 90 градусам.
- Правильно орієнтований код має 2x2 групу `1` (маркерів орієнтації) у нижньому лівому, верхньому лівому та верхньому правому кутах.
- Три 2x2 маркери орієнтації не є частиною бінарних даних.
- Бінарні дані читаються зліва направо, зверху вниз (як у книзі), коли QR-код правильно орієнтований.
- Код завжди має рівно одну правильну орієнтацію.

Наприклад, якщо дано:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

або той самий код з іншою орієнтацією:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

Поверніть `"000000000000000000000001"`, усі бінарні дані, виключаючи три 2x2 маркери орієнтації.

# --hints--

`decode_qr(["110011", "110011", "000000", "000000", "110000", "110001"])` має повертати `"000000000000000000000001"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001")`)
}})
```

`decode_qr(["100011", "000011", "000000", "000000", "110011", "110011"])` має повертати `"000000000000000000000001"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001")`)
}})
```

`decode_qr(["110011", "111111", "010000", "110000", "110011", "110100"])` має повертати `"001101000011000000110100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100")`)
}})
```

`decode_qr(["011011", "101011", "101000", "100010", "110011", "111011"])` має повертати `"010001000100010101010110"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110")`)
}})
```

`decode_qr(["111100", "110001", "100011", "001101", "110011", "110011"])` має повертати `"010000100100100101001110"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode_qr(qr_code):

    return qr_code
```

# --solutions--

```py
def is_oriented(code):
    marker_indices = [
        (0, 0),
        (0, 1),
        (0, 4),
        (0, 5),
        (1, 0),
        (1, 1),
        (1, 4),
        (1, 5),
        (4, 0),
        (4, 1),
        (5, 0),
        (5, 1)
    ]

    for row, col in marker_indices:
        if code[row][col] != '1':
            return False

    return True


def rotate(code):
    new_code = ['', '', '', '', '', '']
    for i in range(6):
        for j in range(5, -1, -1):
            new_code[i] += code[j][i]
    return new_code


def get_data(code):
    data = ''
    data += code[0][2:4]
    data += code[1][2:4]
    data += code[2]
    data += code[3]
    data += code[4][2:]
    data += code[5][2:]
    return data


def decode_qr(qr_code):
    code = qr_code[:]

    while not is_oriented(code):
        code = rotate(code)

    return get_data(code)
```
