---
id: 69a890af247de743333bd4cd
title: "Завдання 224: детектор ступеня обсмаження кави"
challengeType: 29
dashedName: challenge-224
---

# --description--

Дано рядок, який представляє зерна, використані для приготування чашки кави. Визначте ступінь обсмаження цієї чашки.

- Даний рядок міститиме такі символи, кожен з яких представляє тип зерна:

  - Апостроф (`'`) — це зерно світлого обсмаження, яке оцінюється в 1 бал.
  - Дефіс (`-`) — це зерно середнього обсмаження, яке оцінюється в 2 бали.
  - Крапка (`.`) — це зерно темного обсмаження, яке оцінюється в 3 бали.

- Рівень обсмаження визначається середнім значенням усіх зерен.

Поверніть:

- `"Light"`, якщо середнє менше за 1.75.
- `"Medium"`, якщо середнє від 1.75 до 2.5.
- `"Dark"`, якщо середнє більше за 2.5.

# --hints--

`detect_roast("''-''''''-'-''--''''")` має повертати `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("''-''''''-'-''--''''"), "Light")`)
}})
```

`detect_roast(".'-''-''..'''.-.-''-")` має повертати `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".'-''-''..'''.-.-''-"), "Medium")`)
}})
```

`detect_roast("--.''--'-''.--..-.--")` має повертати `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("--.''--'-''.--..-.--"), "Medium")`)
}})
```

`detect_roast("-...'-......-..-...-")` має повертати `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-...'-......-..-...-"), "Dark")`)
}})
```

`detect_roast(".--.-..-......----.'")` має повертати `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".--.-..-......----.'"), "Medium")`)
}})
```

`detect_roast("..-..-..-..-....-.-.")` має повертати `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("..-..-..-..-....-.-."), "Dark")`)
}})
```

`detect_roast("-'-''''''..-'.''-'.'")` має повертати `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-'-''''''..-'.''-'.'"), "Light")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_roast(beans):

    return beans
```

# --solutions--

```py
def detect_roast(beans):
    total = 0

    for bean in beans:
        if bean == "'":
            total += 1
        elif bean == "-":
            total += 2
        elif bean == ".":
            total += 3

    avg = total / len(beans)

    if avg < 1.75:
        return "Light"
    elif avg <= 2.5:
        return "Medium"
    else:
        return "Dark"
```
