---
id: 69a890af247de743333bd4d0
title: "Завдання 227: час охолодження"
challengeType: 29
dashedName: challenge-227
---

# --description--

Дано два часові позначки: перша позначає, коли користувач закінчив екзамен, а друга — поточний час. Визначте, чи може користувач знову складати екзамен.

- Обидві часові позначки будуть у форматі `"YYYY-MM-DDTHH:MM:SS"`, наприклад `"2026-03-25T14:00:00"`. Зверніть увагу, що час подано у 24-годинному форматі.
- Користувач повинен почекати щонайменше 48 годин перед повторним складанням екзамену.

# --hints--

`can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00"), True)`)
}})
```

`can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00"), False)`)
}})
```

`can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00"), True)`)
}})
```

`can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_retake(finish_time, current_time):

    return finish_time
```

# --solutions--

```py
from datetime import datetime, timedelta

def can_retake(finish_time, current_time):
    fmt = "%Y-%m-%dT%H:%M:%S"
    finished_dt = datetime.strptime(finish_time, fmt)
    current_dt = datetime.strptime(current_time, fmt)
    diff = current_dt - finished_dt
    cooldown = timedelta(hours=48)

    return diff >= cooldown
```
