---
id: 69b5b2be76ec8135a7fbe973
title: "Завдання 244: атака тура і слона"
challengeType: 29
dashedName: challenge-244
---

# --description--

Дано рядок із позицією тура на шахівниці та інший рядок із позицією слона. Визначте, чи може одна фігура атакувати іншу.

Стандартна шахівниця має розмір 8x8, з колонками, позначеними `A` до `H` (зліва направо), і рядами, позначеними `1` до `8` (знизу вгору). Вона виглядає так:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- Тури можуть рухатися на будь-яку кількість клітинок по горизонталі або вертикалі.
- Слони можуть рухатися на будь-яку кількість клітинок по діагоналі.
- Одна фігура може атакувати іншу, якщо вона може переміститися на позицію цієї фігури.

Поверніть:

- `"rook"`, якщо тур може атакувати слона.
- `"bishop"`, якщо слон може атакувати тура.
- `"neither"`, якщо жодна з фігур не може атакувати іншу.

# --hints--

`rook_bishop_attack("A1", "A5")` має повертати `"rook"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("A1", "A5"), "rook")`)
}})
```

`rook_bishop_attack("C3", "F6")` має повертати `"bishop"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("C3", "F6"), "bishop")`)
}})
```

`rook_bishop_attack("D4", "D7")` має повертати `"rook"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("D4", "D7"), "rook")`)
}})
```

`rook_bishop_attack("B7", "H1")` має повертати `"bishop"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B7", "H1"), "bishop")`)
}})
```

`rook_bishop_attack("B3", "C5")` має повертати `"neither"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B3", "C5"), "neither")`)
}})
```

`rook_bishop_attack("G3", "E8")` має повертати `"neither"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("G3", "E8"), "neither")`)
}})
```

# --seed--

## --seed-contents--

```py
def rook_bishop_attack(rook, bishop):

    return rook
```

# --solutions--

```py
def rook_bishop_attack(rook, bishop):
    col_diff = abs(ord(rook[0]) - ord(bishop[0]))
    row_diff = abs(int(rook[1]) - int(bishop[1]))

    if rook[0] == bishop[0] or rook[1] == bishop[1]:
        return "rook"
    if col_diff == row_diff:
        return "bishop"
    return "neither"
```
