---
id: 69b559d2903b9e4afe9075f8
title: "Завдання 238: екранування обертання цифр"
challengeType: 29
dashedName: challenge-238
---

# --description--

Дано додатне ціле число. Визначте, чи воно або будь-яке з його обертань ділиться націло на кількість його цифр.

Обертання означає перемістити першу цифру в кінець. Наприклад, після одного обертання 123 стає 231.

- Спочатку перевірте обертання `0` (дане число).
- У заданих числах не буде нулів.
- Поверніть перше число обертання, якщо таке знайдеться, або `"none"`, якщо ні.

# --hints--

`get_rotation(123)` має повернути `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(123), 0)`)
}})
```

`get_rotation(13579)` має повернути `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(13579), 3)`)
}})
```

`get_rotation(24681)` має повернути `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(24681), "none")`)
}})
```

`get_rotation(84138789345)` має повернути `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(84138789345), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rotation(n):

    return n
```

# --solutions--

```py
def get_rotation(n):
    s = str(n)
    digit_count = len(s)
    current = s

    for i in range(digit_count):
        if int(current) % digit_count == 0:
            return i
        current = current[1:] + current[0]

    return "none"
```
