---
id: 69b559d2903b9e4afe9075f7
title: "Завдання 237: валідація рівняння"
challengeType: 29
dashedName: challenge-237
---

# --description--

Дано рядок, що представляє математичне рівняння. Визначте, чи воно правильне.

- Ліва частина може містити до трьох натуральних чисел і оператори `+`, `-`, `*` та `/`.
- Рівняння буде задане у форматі: `"number operator number = number"` (з двома або трьома числами зліва). Наприклад: `"2 + 2 = 4"` або `"2 + 3 - 1 = 4"`.
- Права частина завжди буде одним цілим числом.

Дотримуйтесь стандартного порядку операцій: множення та ділення виконуються перед додаванням і відніманням, зліва направо.

# --hints--

`is_valid_equation("2 + 2 = 4")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 2 = 4"), True)`)
}})
```

`is_valid_equation("2 + 3 - 1 = 4")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 3 - 1 = 4"), True)`)
}})
```

`is_valid_equation("8 / 2 = 4")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("8 / 2 = 4"), True)`)
}})
```

`is_valid_equation("10 * 5 = 50")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("10 * 5 = 50"), True)`)
}})
```

`is_valid_equation("2 - 2 = 0")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 - 2 = 0"), True)`)
}})
```

`is_valid_equation("2 + 9 / 3 = 5")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 9 / 3 = 5"), True)`)
}})
```

`is_valid_equation("20 - 2 * 3 = 14")` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("20 - 2 * 3 = 14"), True)`)
}})
```

`is_valid_equation("2 + 5 = 6")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 5 = 6"), False)`)
}})
```

`is_valid_equation("10 - 2 * 3 = 24")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("10 - 2 * 3 = 24"), False)`)
}})
```

`is_valid_equation("3 + 9 / 3 = 4")` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("3 + 9 / 3 = 4"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_equation(equation):

    return equation
```

# --solutions--

```py
def is_valid_equation(equation):
    left, right = equation.split(" = ")
    tokens = left.split(" ")

    i = 1
    while i < len(tokens) - 1:
        op = tokens[i]
        if op in ("*", "/"):
            a = int(tokens[i - 1])
            b = int(tokens[i + 1])
            computed = a * b if op == "*" else a // b
            tokens[i - 1:i + 2] = [str(computed)]
        else:
            i += 2

    result = int(tokens[0])
    i = 1
    while i < len(tokens) - 1:
        op = tokens[i]
        val = int(tokens[i + 1])
        result = result + val if op == "+" else result - val
        i += 2

    return result == int(right)
```
