---
id: 69bc6cb30c1d112a2e110a07
title: "Завдання 250: прихований ключ"
challengeType: 29
dashedName: challenge-250
---

# --description--

Вітаємо з 250-м щоденним завданням!

Дано зашифрований рядок, розшифруйте його за допомогою ключа шифрування та поверніть результат.

Щоб знайти ключ:

- Подивіться на всі щоденні завдання до сьогоднішнього дня, номер яких кратний 25 (включно з цим).
- Візьміть першу літеру з кожної назви цих завдань і об’єднайте їх у рядок. Якщо назва починається з не літери, знайдіть її першу літеру.

Щоб розшифрувати повідомлення, пройдіться по кожній літері зашифрованого рядка і:

  1. Подивіться на відповідну літеру в ключі (повторюйте ключ, якщо повідомлення довше за ключ).
  2. Перетворіть літеру ключа на відповідне число: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
  3. Зсуньте зашифровану літеру назад в алфавіті на це число.
  4. Якщо зсув виходить за межі `"A"`, оберніться до `"Z"`.

Наприклад, якщо зашифроване повідомлення починається з `"Y"`, а перша літера ключа — `"V"` (22), зсуньте `"Y"` назад на 22 позиції, щоб отримати `"C"`. Повторюйте цей процес для кожної літери, щоб розшифрувати повне повідомлення.

- Зсуваються лише літери, пробіли повертаються без змін.
- Всі задані та повернені літери — у верхньому регістрі.

# --hints--

`decode("YAVJYNXE")` має повернути `"CONGRATS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")` має повернути `"CODING LEGEND"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")` має повернути `"YOU ARE AWESOME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")` має повернути `"KEEP GOING"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` має повернути `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```
