---
id: 69c5f3d787b1725d5f00c8bd
title: "Завдання 259: вибух FizzBuzz"
challengeType: 29
dashedName: challenge-259
---

# --description--

Дано ціле число, поверніть кількість кроків, необхідних, щоб перетворити слово `"fizzbuzz"` на рядок з принаймні заданою кількістю `"z"` за такими правилами:

- Починайте з рядка `"fizzbuzz"`.
- Кожного кроку застосовуйте стандартні правила FizzBuzz, використовуючи позицію літери в рядку (перша `"f"` має позицію 1).
  - Якщо позиція літери ділиться на 3, замініть літеру на `"fizz"`
  - Якщо вона ділиться на 5, замініть літеру на `"buzz"`
  - Якщо вона ділиться і на 3, і на 5, замініть літеру на `"fizzbuzz"`

Отже, після 1 кроку `"fizzbuzz"` перетворюється на `"fifizzzbuzzfizzzz"`, який містить 9 `"z"`.

# --hints--

`explode_fizzbuzz(9)` має повернути `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

`explode_fizzbuzz(15)` має повернути `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

`explode_fizzbuzz(51)` має повернути `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

`explode_fizzbuzz(52)` має повернути `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

`explode_fizzbuzz(359)` має повернути `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

`explode_fizzbuzz(789)` має повернути `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

`explode_fizzbuzz(54482)` має повернути `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

`explode_fizzbuzz(1000000)` має повернути `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```
