---
id: 69c5f3d787b1725d5f00c8ba
title: "Завдання 256: найближчий напрямок часу"
challengeType: 29
dashedName: challenge-256
---

# --description--

Дано два часи, визначте, чи можна дістатися від першого до другого швидше, рухаючись вперед чи назад.

- Часи задані у 24-годинному форматі (`"HH:MM"`)
- Годинник циклічний (23:59 переходить у 00:00 при русі вперед, а 00:00 переходить у 23:59 при русі назад)

Поверніть:

- `"forward"`, якщо рух вперед коротший
- `"backward"`, якщо рух назад коротший
- `"equal"`, якщо обидва напрямки займають однаковий час

# --hints--

`get_direction("10:00", "12:00")` має повертати `"forward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("10:00", "12:00"), "forward")`)
}})
```

`get_direction("11:00", "05:00")` має повертати `"backward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("11:00", "05:00"), "backward")`)
}})
```

`get_direction("00:00", "12:00")` має повертати `"equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("00:00", "12:00"), "equal")`)
}})
```

`get_direction("15:45", "01:10")` має повертати `"forward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("15:45", "01:10"), "forward")`)
}})
```

`get_direction("03:30", "19:50")` має повертати `"backward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("03:30", "19:50"), "backward")`)
}})
```

`get_direction("06:30", "18:30")` має повертати `"equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("06:30", "18:30"), "equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_direction(time1, time2):

    return time1
```

# --solutions--

```py
def get_direction(time1, time2):
    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    start = to_minutes(time1)
    end = to_minutes(time2)
    forward = (end - start + 1440) % 1440
    backward = (start - end + 1440) % 1440

    if forward < backward:
        return 'forward'
    if backward < forward:
        return 'backward'
    return 'equal'
```
