---
id: 69c6ff713a52713463aa7928
title: "Challenge 262: URL Query Parser"
challengeType: 29
dashedName: challenge-262
---

# --description--

Дано URL, який містить рядок запиту. Проаналізуйте цей рядок запиту та перетворіть його на об’єкт (або словник) пар ключ-значення.

- Рядок запиту починається після `"?"`,
- кожен параметр відокремлений `"&"`,
- кожна пара ключ/значення відокремлена `"="`

Наприклад, для `"https://example.com/search?name=Alice&age=30"` поверніть:

```json
{
  "name": "Alice",
  "age": "30"
}
```

Усі значення мають бути повернені у вигляді рядків.

# --hints--

`parse_url_query("https://example.com/search?name=Alice&age=30")` should return `{"name": "Alice", "age": "30"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"})`)
}})
```

`parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python")` should return `{"skill": "programming", "language": "python"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"})`)
}})
```

`parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2")` should return `{"category": "books", "sort": "asc", "page": "2"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"})`)
}})
```

`parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now")` should return `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_url_query(url):

    return url
```

# --solutions--

```py
def parse_url_query(url):
    query_string = url.split('?')[1]
    return dict(param.split('=') for param in query_string.split('&'))
```
