---
id: 69c6ff713a52713463aa7927
title: "Завдання 261: числові слова"
challengeType: 29
dashedName: challenge-261
---

# --description--

Дано ціле число від 0 до 99, поверніть його англійське словесне представлення.

- `0` повертає `"zero"`.
- Числа від 1 до 19 мають унікальні назви (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- Кратні 10 від 20 до 90 мають власні назви (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- Числа від 21 до 99, які не є кратними 10, записуються як два слова, з’єднані дефісом. Наприклад `"forty-two"` і `"fifty-three"`.

# --hints--

`get_number_words(0)` має повертати `"zero"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(0), "zero")`)
}})
```

`get_number_words(10)` має повертати `"ten"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(10), "ten")`)
}})
```

`get_number_words(19)` має повертати `"nineteen"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(19), "nineteen")`)
}})
```

`get_number_words(30)` має повертати `"thirty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(30), "thirty")`)
}})
```

`get_number_words(53)` має повертати `"fifty-three"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(53), "fifty-three")`)
}})
```

`get_number_words(7)` має повертати `"seven"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(7), "seven")`)
}})
```

`get_number_words(12)` має повертати `"twelve"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(12), "twelve")`)
}})
```

`get_number_words(60)` має повертати `"sixty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(60), "sixty")`)
}})
```

`get_number_words(67)` має повертати `"sixty-seven"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(67), "sixty-seven")`)
}})
```

`get_number_words(98)` має повертати `"ninety-eight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(98), "ninety-eight")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_words(n):

    return n
```

# --solutions--

```py
def get_number_words(n):
    ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
            "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
            "seventeen","eighteen","nineteen"]
    tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"]

    if n < 20:
        return ones[n]
    if n % 10 == 0:
        return tens[n // 10]
    return tens[n // 10] + "-" + ones[n % 10]
```
