---
id: 69cfca90e8a0a6d4d6871c51
title: "Завдання 267: конвертер парсеків"
challengeType: 29
dashedName: challenge-267
---

# --description--

У далекій галактиці парсеки використовують для вимірювання як часу, так і відстані. За заданим цілим числом парсеків поверніть його еквівалент у часі або відстані.

- Якщо задане ціле число непарне, воно представляє час. Якщо парне — відстань.

Використовуйте такі коефіцієнти конвертації:

| Парсеки | Час/Відстань |
| - | - |
| 1 | 2 години |
| 2 | 6 світлових років |

Поверніть конвертоване значення як ціле число.

# --hints--

`convert_parsecs(1)` має повернути `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(1), 2)`)
}})
```

`convert_parsecs(2)` має повернути `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(2), 6)`)
}})
```

`convert_parsecs(31)` має повернути `62`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(31), 62)`)
}})
```

`convert_parsecs(88)` має повернути `264`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(88), 264)`)
}})
```

`convert_parsecs(17)` має повернути `34`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(17), 34)`)
}})
```

`convert_parsecs(14)` має повернути `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(14), 42)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_parsecs(parsecs):

    return parsecs
```

# --solutions--

```py
def convert_parsecs(parsecs):
    if parsecs % 2 != 0:
        return parsecs * 2
    else:
        return parsecs * 3
```
