---
id: 69cfca90e8a0a6d4d6871c53
title: "Challenge 269: Allergen Friendly Meals"
challengeType: 29
dashedName: challenge-269
---

# --description--

Дано масив страв і масив алергенів, яких слід уникати. Поверніть назви всіх страв, які не містять жодного з цих алергенів.

- Кожна страва має формат `[meal, allergens]`, де `meal` — назва страви, а `allergens` — масив алергенів, які містить страва. Наприклад, `["pasta", ["wheat", "milk"]]`.
- Алергени, яких слід уникати, будуть масивом рядків.

Поверніть назви безпечних страв у тому ж порядку, що й задано. Якщо безпечних страв немає, поверніть порожній масив.

# --hints--

`get_allergen_friendly_meals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"])` should return `["salad"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"]), ["salad"])`)
}})
```

`get_allergen_friendly_meals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"])` should return `["fried rice", "chicken parmesan"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"]), ["fried rice", "chicken parmesan"])`)
}})
```

`get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"])` should return `["oatmeal", "granola", "toast"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"]), ["oatmeal", "granola", "toast"])`)
}})
```

`get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"])` should return `["granola", "yogurt", "eggs"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"]), ["granola", "yogurt", "eggs"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_allergen_friendly_meals(meals, allergens):

    return meals
```

# --solutions--

```py
def get_allergen_friendly_meals(meals, allergens):
    return [name for name, meal_allergens in meals if not any(a in meal_allergens for a in allergens)]
```
