---
id: 69cfca90e8a0a6d4d6871c50
title: "Challenge 266: Good Day"
challengeType: 29
dashedName: challenge-266
---

# --description--

Дано рядок часу у форматі `"HH:MM"` (24-годинний формат), поверніть:

- `"Good morning"` для часу від `05:00` до `11:59`
- `"Good afternoon"` для часу від `12:00` до `17:59`
- `"Good evening"` для часу від `18:00` до `21:59`
- `"Good night"` для часу від `22:00` до `04:59`

# --hints--

`get_greeting("06:30")` should return `"Good morning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("06:30"), "Good morning")`)
}})
```

`get_greeting("12:00")` should return `"Good afternoon"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("12:00"), "Good afternoon")`)
}})
```

`get_greeting("21:59")` should return `"Good evening"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("21:59"), "Good evening")`)
}})
```

`get_greeting("00:01")` should return `"Good night"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("00:01"), "Good night")`)
}})
```

`get_greeting("11:30")` should return `"Good morning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("11:30"), "Good morning")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_greeting(s):

    return s
```

# --solutions--

```py
def get_greeting(s):
    hours, minutes = map(int, s.split(':'))
    total = hours * 60 + minutes

    if 300 <= total < 720:
        return 'Good morning'
    if 720 <= total < 1080:
        return 'Good afternoon'
    if 1080 <= total < 1320:
        return 'Good evening'
    return 'Good night'
```
