---
id: 69e2383af7832c8032603b93
title: "Завдання 277: дзеркальне відображення"
challengeType: 29
dashedName: challenge-277
---

# --description--

Дано два рядки, визначте, чи є другий рядок дзеркальним відображенням першого.

Дзеркальне відображення утворюється шляхом обертання рядка в зворотному напрямку та заміни кожного символу на його дзеркальний еквівалент.

- Симетричні символи виглядають у дзеркалі так само, як і самі:

`W`, `T`, `Y`, `U`, `I`, `O`, `H`, `A`, `X`, `V`, `M`, `w`, `o`, `x`, `v`, `0`, `8`, `=`, `+`, `:`, `|`, `-`, `_`, `*`, `^`, `!`, `.`, і пробіл (` `).

- Дзеркальні пари міняються місцями у дзеркалі:

| Символ | Міняється місцями з |
| - | - |
| `[` | `]` |
| `{` | `}` |
| `<` | `>` |
| `b` | `d` |
| `p` | `q` |
| `(` | `)` |

Якщо будь-який із рядків містить символ, якого немає у наведених вище списках, то дзеркальне відображення для нього створити неможливо.

Наприклад, дзеркальне відображення `"[HOW]"` — це `"[WOH]"`.

# --hints--

`is_mirror_image("[HOW]", "[WOH]")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("[HOW]", "[WOH]"), True)`)
}})
```

`is_mirror_image("MOM", "MOM")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("MOM", "MOM"), True)`)
}})
```

`is_mirror_image("vow", "wov")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("vow", "wov"), True)`)
}})
```

`is_mirror_image("TIM", "TIM")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("TIM", "TIM"), False)`)
}})
```

`is_mirror_image("{WOW}", "}WOW{")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("{WOW}", "}WOW{"), False)`)
}})
```

`is_mirror_image("XXVII", "IIV%X")` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("XXVII", "IIV%X"), False)`)
}})
```

`is_mirror_image("><(((*>", "<*)))><")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("><(((*>", "<*)))><"), True)`)
}})
```

`is_mirror_image("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW")` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_mirror_image(s1, s2):

    return s1
```

# --solutions--

```py
def is_mirror_image(s1, s2):
    symmetric = set('WTYUIOHA XVMwoxv08=+:|-_*^!.')
    pairs = {'[': ']', ']': '[', '{': '}', '}': '{', '<': '>', '>': '<',
             'b': 'd', 'd': 'b', 'p': 'q', 'q': 'p', '(': ')', ')': '('}
    mirrored = []
    for c in reversed(s1):
        if c in symmetric:
            mirrored.append(c)
        elif c in pairs:
            mirrored.append(pairs[c])
        else:
            return False
    return ''.join(mirrored) == s2
```
