---
id: 69e2383af7832c8032603b8f
title: "Challenge 273: ISBN-13 Validator"
challengeType: 29
dashedName: challenge-273
---

# --description--

Дано рядок, визначте, чи є він дійсним номером ISBN-13.

A valid ISBN-13:

- Містить лише цифри та дефіси
- Містить рівно 13 цифр після видалення дефісів
- Проходить таку перевірку:
  1. Помножте кожну цифру на 1 або 3, чергуючи (помножте першу цифру на 1, другу на 3, третю на 1 і так далі).
  2. Сума результатів має бути кратною 10.

# --hints--

`is_valid_isbn_13("9780306406157")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9780306406157"), True)`)
}})
```

`is_valid_isbn_13("97803064061570")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("97803064061570"), False)`)
}})
```

`is_valid_isbn_13("978-0-13-595705-9")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-0-13-595705-9"), True)`)
}})
```

`is_valid_isbn_13("978-030-64061A-4")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-030-64061A-4"), False)`)
}})
```

`is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_isbn_13(s):

    return s
```

# --solutions--

```py
def is_valid_isbn_13(s):
    digits = s.replace("-", "")
    if not digits.isdigit() or len(digits) != 13:
        return False
    total = sum(int(d) * (1 if i % 2 == 0 else 3) for i, d in enumerate(digits))
    return total % 10 == 0
```
