---
id: 69f8c998d78ad3171a0713ba
title: "Завдання 287: виправлення римських чисел"
challengeType: 29
dashedName: challenge-287
---

# --description--

Дано рядок неправильно записаних римських чисел, поверніть значення у стандартній римській нотації.

У вхідних даних використовується лише адитивна нотація, тому кожен символ додає своє значення до загальної суми. Нагадаємо символи та їхні значення:

| Символ | Значення |
| ------ | -------- |
| `"I"`  | 1        |
| `"V"`  | 5        |
| `"X"`  | 10       |
| `"L"`  | 50       |
| `"C"`  | 100      |
| `"D"`  | 500      |
| `"M"`  | 1000     |

При повторному кодуванні використовуйте найбільший можливий символ на кожному кроці, застосовуючи віднімаючі пари (`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`) там, де це потрібно.

# --hints--

`fix_numerals("XIIIII")` має повернути `"XV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("XIIIII"), "XV")`)
}})
```

`fix_numerals("IIIILX")` має повернути `"LXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("IIIILX"), "LXIV")`)
}})
```

`fix_numerals("XXVVVIIIII")` має повернути `"XL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("XXVVVIIIII"), "XL")`)
}})
```

`fix_numerals("MDCCLXXXXVIIII")` має повернути `"MDCCXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("MDCCLXXXXVIIII"), "MDCCXCIX")`)
}})
```

`fix_numerals("IIIIVVVVXXXXLLLLCCDD")` має повернути `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV")`)
}})
```

`fix_numerals("ILCDMIVDIIXLCVCXDL")` має повернути `"MMCMLXXXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV")`)
}})
```

# --seed--

## --seed-contents--

```py
def fix_numerals(s):

    return s
```

# --solutions--

```py
def fix_numerals(s):
    values = {"I": 1, "V": 5, "X": 10, "L": 50, "C": 100, "D": 500, "M": 1000}
    encoding = [
        (1000, "M"), (900, "CM"), (500, "D"), (400, "CD"),
        (100, "C"), (90, "XC"), (50, "L"), (40, "XL"),
        (10, "X"), (9, "IX"), (5, "V"), (4, "IV"), (1, "I")
    ]

    total = sum(values[c] for c in s)

    result = ""
    for value, symbol in encoding:
        while total >= value:
            result += symbol
            total -= value
    return result
```
