---
id: 69f8c998d78ad3171a0713b9
title: "Завдання 286: відкриті проблеми"
challengeType: 29
dashedName: challenge-286
---

# --description--

Дано масив номерів проблем і ще один масив номерів запитів на злиття (PR). Поверніть масив проблем, які залишаються відкритими після того, як усі PR були об’єднані.

- PR закриває проблему, якщо їхні цифри є обертанням одна одної. Наприклад, проблему 123 закриє PR 231 або 312.
- PR не закриває проблему з точно таким самим числом. Наприклад, PR 123 не закриває проблему 123. Тож проблему з усіма однаковими цифрами не можна закрити.
- Будь-яке число може мати видалені провідні нулі. Наприклад, PR 201 закриє проблему 12 (012, обертання 201). Аналогічно, проблему 201 закриє PR 12.

Поверніть залишкові відкриті проблеми в тому порядку, в якому вони були задані.

# --hints--

`get_open_issues([123, 234], [231])` має повертати `[234]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 234], [231]), [234])`)
}})
```

`get_open_issues([123, 345, 16], [345, 231])` має повертати `[345, 16]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 345, 16], [345, 231]), [345, 16])`)
}})
```

`get_open_issues([456, 332, 12, 15], [201, 945, 180])` має повертати `[456, 332, 15]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15])`)
}})
```

`get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` має повертати `[115, 296, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24])`)
}})
```

`get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` має повертати `[95, 395, 754, 296, 709, 237, 1802]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_open_issues(issues, prs):

    return issues
```

# --solutions--

```py
def get_open_issues(issues, prs):
    def get_rotations(n):
        s = str(n)
        rotations = set()
        for i in range(len(s)):
            rotated = s[i:] + s[:i]
            rotations.add(int(rotated))
        return rotations

    def are_rotations(a, b):
        if a == b:
            return False
        return b in get_rotations(a)

    return [issue for issue in issues if not any(are_rotations(pr, issue) or are_rotations(issue, pr) for pr in prs)]
```
