---
id: 69f8c998d78ad3171a0713be
title: "Challenge 291: FizzBuzz Count"
challengeType: 29
dashedName: challenge-291
---

# --description--

За заданими початковим і кінцевим числом порахуйте кількість появ fizz і buzz у діапазоні (включно).

- Числа, кратні 3, рахуйте як fizz.
- Числа, кратні 5, рахуйте як buzz.
- Числа, кратні і 3, і 5, рахуйте як одночасно fizz і buzz.

Поверніть об’єкт або словник із підрахунками у форматі: `{ fizz, buzz }`.

# --hints--

`fizz_buzz_count(1, 11)` should return `{"fizz": 3, "buzz": 2}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(1, 11), {"fizz": 3, "buzz": 2})`)
}})
```

`fizz_buzz_count(14, 41)` should return `{"fizz": 9, "buzz": 6}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(14, 41), {"fizz": 9, "buzz": 6})`)
}})
```

`fizz_buzz_count(24, 100)` should return `{"fizz": 26, "buzz": 16}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(24, 100), {"fizz": 26, "buzz": 16})`)
}})
```

`fizz_buzz_count(-635, -14)` should return `{"fizz": 207, "buzz": 125}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-635, -14), {"fizz": 207, "buzz": 125})`)
}})
```

`fizz_buzz_count(-5432, 6789)` should return `{"fizz": 4074, "buzz": 2444}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-5432, 6789), {"fizz": 4074, "buzz": 2444})`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_count(start, end):

    return start
```

# --solutions--

```py
def fizz_buzz_count(start, end):
    fizz = sum(1 for i in range(start, end + 1) if i % 3 == 0)
    buzz = sum(1 for i in range(start, end + 1) if i % 5 == 0)
    return {"fizz": fizz, "buzz": buzz}
```
