---
id: 69f8c998d78ad3171a0713bf
title: "Завдання 292: ширше співвідношення сторін"
challengeType: 29
dashedName: challenge-292
---

# --description--

Дано два рядки з різними розмірами зображень, поверніть співвідношення сторін зображення з більшим співвідношенням ширини до висоти.

- Дано рядки у форматі `"WxH"`, наприклад, `"1920x1080"`.
- Співвідношення сторін — це відношення ширини до висоти, скорочене до найменших цілих чисел. Наприклад, `"1920x1080"` скорочується до `"16:9"`.
- Поверніть рядок у форматі `"W:H"`, наприклад, `"16:9"`.

# --hints--

`get_wider_aspect_ratio("1920x1080", "800x600")` має повернути `"16:9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1920x1080", "800x600"), "16:9")`)
}})
```

`get_wider_aspect_ratio("1080x1350", "2048x1536")` має повернути `"4:3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1080x1350", "2048x1536"), "4:3")`)
}})
```

`get_wider_aspect_ratio("640x480", "2440x1220")` має повернути `"2:1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("640x480", "2440x1220"), "2:1")`)
}})
```

`get_wider_aspect_ratio("360x640", "1080x1920")` має повернути `"9:16"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("360x640", "1080x1920"), "9:16")`)
}})
```

`get_wider_aspect_ratio("3440x1440", "2048x858")` має повернути `"43:18"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("3440x1440", "2048x858"), "43:18")`)
}})
```

`get_wider_aspect_ratio("12345x61234", "12534x51234")` має повернути `"2089:8539"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("12345x61234", "12534x51234"), "2089:8539")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_wider_aspect_ratio(a, b):

    return a
```

# --solutions--

```py
from math import gcd

def get_wider_aspect_ratio(a, b):
    def parse(s):
        w, h = map(int, s.split("x"))
        g = gcd(w, h)
        return w / h, f"{w // g}:{h // g}"

    ratio_a, simplified_a = parse(a)
    ratio_b, simplified_b = parse(b)

    return simplified_a if ratio_a >= ratio_b else simplified_b
```
