---
id: 69f90c1329a94b37e2a2086d
title: "Завдання 294: комбінації круглих дужок"
challengeType: 29
dashedName: challenge-294
---

# --description--

Дано ціле число `n`, поверніть кількість дійсних комбінацій `n` пар круглих дужок.

- Дійсна комбінація — це рядок, у якому кожна відкрита кругла дужка має відповідну закриту круглу дужку, і жодна закрита кругла дужка не з’являється раніше за свою відповідну відкриту.

Наприклад, для `2` існує 2 дійсні комбінації:

```json
(())
()()
```

# --hints--

`get_combinations(2)` має повернути `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(2), 2)`)
}})
```

`get_combinations(3)` має повернути `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(3), 5)`)
}})
```

`get_combinations(5)` має повернути `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(5), 42)`)
}})
```

`get_combinations(8)` має повернути `1430`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(8), 1430)`)
}})
```

`get_combinations(13)` має повернути `742900`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(13), 742900)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_combinations(n):

    return n
```

# --solutions--

```py
def get_combinations(n):
    count = 0

    def generate(op, cl):
        nonlocal count
        if op == 0 and cl == 0:
            count += 1
            return
        if op > 0:
            generate(op - 1, cl)
        if cl > op:
            generate(op, cl - 1)

    generate(n, n)
    return count
```
