---
id: 6a0dcc730cb92a616f86f0c5
title: "Завдання 301: останнє завантаження"
challengeType: 29
dashedName: challenge-301
---

# --description--

Враховуючи кількість мірок прального порошку, що залишилися, та масив із кількістю мірок, які ви використовували в кожен із попередніх днів, поверніть кількість повних днів, на які вистачить порошку.

Обчисліть середнє щоденне використання за історією використання і припустіть, що саме таку кількість ви використовуватимете щодня надалі.

# --hints--

`last_load_date(10, [2, 2, 2, 2, 2, 2, 2])` має повернути `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(10, [2, 2, 2, 2, 2, 2, 2]), 5)`)
}})
```

`last_load_date(16, [2, 3, 0, 3, 4, 2, 1])` має повернути `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(16, [2, 3, 0, 3, 4, 2, 1]), 7)`)
}})
```

`last_load_date(33, [5, 0, 4, 3, 3, 2])` має повернути `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(33, [5, 0, 4, 3, 3, 2]), 11)`)
}})
```

`last_load_date(50, [2, 0, 2, 9, 12, 0, 2])` має повернути `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(50, [2, 0, 2, 9, 12, 0, 2]), 12)`)
}})
```

`last_load_date(20, [13, 9, 12, 10, 8])` має повернути `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(20, [13, 9, 12, 10, 8]), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def last_load_date(scoops, usage):

    return scoops
```

# --solutions--

```py
def last_load_date(scoops, usage):
    avg = sum(usage) / len(usage)
    return int(scoops / avg)
```
