---
id: 6a0dcc730cb92a616f86f0c1
title: "Завдання 297: Валідатор схеми Частина 3"
challengeType: 29
dashedName: challenge-297
---

# --description--

Дано об’єкт (JavaScript) або словник (Python), визначте, чи відповідає він наступній схемі:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles
}
```

- Символ вертикальної риски (`|`) означає "або". `role` має бути одним із перелічених значень `Roles`.
- Додаткові ключі дозволені

# --hints--

`is_valid_schema({"username": "henry", "posts": 0, "verified": True, "role": "staff"})` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "henry", "posts": 0, "verified": True, "role": "staff"}), True)`)
}})
```

`is_valid_schema({"username": "sara", "posts": 45, "verified": False, "role": "creator", "followers": 70})` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "sara", "posts": 45, "verified": False, "role": "creator", "followers": 70}), True)`)
}})
```

`is_valid_schema({"username": "penelope", "posts": 20, "verified": True, "role": "admin"})` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "penelope", "posts": 20, "verified": True, "role": "admin"}), True)`)
}})
```

`is_valid_schema({"username": "kevin", "posts": 0, "verified": False, "role": "user"})` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "kevin", "posts": 0, "verified": False, "role": "user"}), True)`)
}})
```

`is_valid_schema({"username": "george", "posts": 15, "verified": True, "role": "moderator"})` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "george", "posts": 15, "verified": True, "role": "moderator"}), True)`)
}})
```

`is_valid_schema({"username": "david", "posts": 0, "verified": False, "role": "guest"})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "david", "posts": 0, "verified": False, "role": "guest"}), False)`)
}})
```

`is_valid_schema({"username": "wendy", "posts": 10, "verified": True})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "wendy", "posts": 10, "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "fabian", "posts": 1, "verified": True, "role": True})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "fabian", "posts": 1, "verified": True, "role": True}), False)`)
}})
```

`is_valid_schema({"username": 8, "posts": 1, "verified": True, "role": "user"})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 8, "posts": 1, "verified": True, "role": "user"}), False)`)
}})
```

`is_valid_schema({"username": "penny", "posts": "10", "verified": True, "role": "staff"})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "penny", "posts": "10", "verified": True, "role": "staff"}), False)`)
}})
```

`is_valid_schema({"username": "john", "posts": "1", "verified": "true", "role": "admin"})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "john", "posts": "1", "verified": "true", "role": "admin"}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    roles = ["user", "creator", "moderator", "staff", "admin"]
    return (
        isinstance(obj.get("username"), str) and
        isinstance(obj.get("posts"), int) and not isinstance(obj.get("posts"), bool) and
        isinstance(obj.get("verified"), bool) and
        obj.get("role") in roles
    )
```
