---
id: 6a0dcc730cb92a616f86f0c0
title: "Завдання 296: Валідатор схеми. Частина 2"
challengeType: 29
dashedName: challenge-296
---

# --description--

Дано об’єкт (JavaScript) або словник (Python). Визначте, чи відповідає він наступній схемі:

```json
{
  username: string,
  posts: number,
  verified: boolean
}
```

- Додаткові ключі дозволені

# --hints--

`is_valid_schema({"username": "alice", "posts": 10, "verified": False})` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "alice", "posts": 10, "verified": False}), True)`)
}})
```

`is_valid_schema({"username": "carol", "posts": 15, "verified": True, "followers": 25})` має повертати `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "carol", "posts": 15, "verified": True, "followers": 25}), True)`)
}})
```

`is_valid_schema({"username": "frank", "posts": "21", "verified": True})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "frank", "posts": "21", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "sam", "posts": 17, "verified": "false"})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "sam", "posts": 17, "verified": "false"}), False)`)
}})
```

`is_valid_schema({"username": "bill", "verified": True})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bill", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "fred", "verified": True})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "fred", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": 5, "posts": 10, "verified": True})` має повертати `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 5, "posts": 10, "verified": True}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    return (
        isinstance(obj.get("username"), str) and
        isinstance(obj.get("posts"), int) and not isinstance(obj.get("posts"), bool) and
        isinstance(obj.get("verified"), bool)
    )
```
