---
id: 6a0dcd03ee4e68698080ef6a
title: "Завдання 306: витягування вмісту HTML"
challengeType: 29
dashedName: challenge-306
---

# --description--

Дано рядок HTML, поверніть звичайний текстовий вміст із видаленими всіма тегами.

# --hints--

`extract_content('<p>hello world</p>')` має повернути `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_content('<p>hello world</p>'), "hello world")`)
}})
```

`extract_content('<p>hello <span>world</span></p>')` має повернути `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_content('<p>hello <span>world</span></p>'), "hello world")`)
}})
```

`extract_content('<a href="example.com">Click me</a>')` має повернути `"Click me"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_content('<a href="example.com">Click me</a>'), "Click me")`)
}})
```

`extract_content('<p><button onClick="learnToCode()">Learn</button> to <code>code<code> <br/>for <strong>free</strong> <br/>on <a href="https://freecodecamp.org/" target="_blank"><span class="highlight">freecodecamp</span>.org</a>')` має повернути `"Learn to code for free on freecodecamp.org"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_content('<p><button onClick="learnToCode()">Learn</button> to <code>code<code> <br/>for <strong>free</strong> <br/>on <a href="https://freecodecamp.org/" target="_blank"><span class="highlight">freecodecamp</span>.org</a>'), "Learn to code for free on freecodecamp.org")`)
}})
```

`extract_content('<div class="container"><h1 id="title">Welcome to <strong>My</strong> Website.</h1><p>This is a <a href="https://example.com" target="_blank">link</a> to something <em>really</em> <span class="highlight">important</span>.</p><ul><li>Item <strong>one</strong></li><li>Item <em>two</em></li><li>Item three</li></ul><img src="pic.jpg" alt="A picture"/><p class="footer">Contact us at <a href="mailto:hello@example.com">hello@example.com</a> for <span>more <strong>info</strong></span>.</p></div>')` має повернути `"Welcome to My Website.This is a link to something really important.Item oneItem twoItem threeContact us at hello@example.com for more info."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_content('<div class="container"><h1 id="title">Welcome to <strong>My</strong> Website.</h1><p>This is a <a href="https://example.com" target="_blank">link</a> to something <em>really</em> <span class="highlight">important</span>.</p><ul><li>Item <strong>one</strong></li><li>Item <em>two</em></li><li>Item three</li></ul><img src="pic.jpg" alt="A picture"/><p class="footer">Contact us at <a href="mailto:hello@example.com">hello@example.com</a> for <span>more <strong>info</strong></span>.</p></div>'), "Welcome to My Website.This is a link to something really important.Item oneItem twoItem threeContact us at hello@example.com for more info.")`)
}})
```

# --seed--

## --seed-contents--

```py
def extract_content(html):

    return html
```

# --solutions--

```py
import re

def extract_content(html):
    return re.sub(r'<[^>]*>', '', html).strip()
```
