---
id: 6a0dcd03ee4e68698080ef6b
title: "Завдання 307: правила зонування"
challengeType: 29
dashedName: challenge-307
---

# --description--

Дано двовимірну сітку (масив масивів), що представляє макет будівель міста. Поверніть координати всіх будівель, які порушують правила зонування.

Кожна клітинка сітки містить одну з міток із таблиці нижче. Будівля порушує правила, якщо будь-який із її (до) 4 сусідів по горизонталі або вертикалі є типом, з яким вона не може бути поруч.

| Мітка | Тип | Не може бути поруч із |
| - | - | - |
| `"i"` | industrial | `"R"`, `"I"` |
| `"A"` | Agricultural | `"C"` |
| `"R"` | Residential | `"i"`, `"C"` |
| `"I"` | Institutional | `"i"` |
| `"C"` | Commercial | `"R"`, `"A"` |
| `""` (порожній рядок) | undeveloped | без обмежень |

Поверніть координати всіх клітинок, що порушують правила, у вигляді масиву пар `[row, col]`, у будь-якому порядку. Якщо порушень немає, поверніть порожній масив.

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` має повертати `[[0, 0], [0, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` має повертати `[[0, 1], [1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` має повертати `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` має повертати `[[0, 1], [0, 2], [0, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` має повертати `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
