---
id: 6a0dcd03ee4e68698080ef68
title: "Завдання 304: Організація маршруту"
challengeType: 29
dashedName: challenge-304
---

# --description--

Дано масив щонайменше з двома необов’язковими зупинками для одноденної поїздки. Потрібно повернути кількість дійсних варіантів організації маршруту.

Маршрут завжди включає `"breakfast"`, `"lunch"` і `"dinner"`, їх не передають як аргументи. Необов’язкові зупинки можна розміщувати будь-де в маршруті за такими правилами:

- `"breakfast"` завжди перша, перед `"lunch"` має бути щонайменше одна зупинка.
- `"lunch"` має з’являтися перед `"dinner"`, між ними має бути щонайменше одна зупинка.
- Після `"dinner"` може бути не більше однієї необов’язкової зупинки.

Поверніть кількість дійсних варіантів організації.

# --hints--

`get_itinerary_count(["library", "park"])` має повертати `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` має повертати `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` має повертати `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` має повертати `840`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` має повертати `55440`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
