---
id: 6a151d32d772271dd2448c2d
title: "Завдання 311: Заклинатель"
challengeType: 29
dashedName: challenge-311
---

# --description--

Дано рядок кодів заклять, які ви накладаєте. Обчисліть загальну кількість балів.

Кожен символ у рядку представляє закляття:

| Код | Закляття | Категорія | Базовий бал |
| - | - | - | - |
| `"f"` | Вогонь | Руйнування | 3 |
| `"l"` | Блискавка | Руйнування | 3 |
| `"i"` | Лід | Контроль | 2 |
| `"w"` | Вітер | Контроль | 2 |
| `"h"` | Лікування | Відновлення | 1 |
| `"s"` | Щит | Відновлення | 1 |

Комбінований множник застосовується залежно від того, скільки заклять підряд було накладено з різних категорій:

- Перше закляття завжди оцінюється за базовим значенням.
- Кожне наступне закляття з іншої категорії, ніж попереднє, збільшує множник на 1.
- Накладення закляття з тієї ж категорії, що й попереднє, скидає множник назад до 1.
- Бал за кожне закляття — це його базовий бал, помножений на поточний множник.

Поверніть загальну кількість балів за послідовність заклять.

# --hints--

`cast("fihwl")` має повернути `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` має повернути `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` має повернути `37`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` має повернути `50`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` має повернути `101`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
