---
id: 6a15cadf5f240d05a264955b
title: "Завдання 317: калькулятор BMI"
challengeType: 29
dashedName: challenge-317
---

# --description--

За заданою вагою в фунтах і зростом в дюймах поверніть BMI (індекс маси тіла), округлений до одного десяткового знака.

Щоб отримати BMI: поділіть вагу на квадрат зросту, а потім помножте результат на 703.

# --hints--

`calculate_bmi(180, 70)` має повернути `25.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(180, 70), 25.8)`)
}})
```

`calculate_bmi(140, 64)` має повернути `24.0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(140, 64), 24.0)`)
}})
```

`calculate_bmi(160, 76)` має повернути `19.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(160, 76), 19.5)`)
}})
```

`calculate_bmi(200, 60)` має повернути `39.1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(200, 60), 39.1)`)
}})
```

`calculate_bmi(150, 68)` має повернути `22.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(150, 68), 22.8)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_bmi(weight, height):

    return weight
```

# --solutions--

```py
def calculate_bmi(weight, height):
    return round((weight / (height * height)) * 703, 1)
```
