---
id: 6a15cadf5f240d05a2649557
title: "Завдання 313: Вартість оренди"
challengeType: 29
dashedName: challenge-313
---

# --description--

За заданими часовими позначками оренди, повернення та рівнем оренди поверніть загальну вартість оренди з урахуванням будь-яких штрафів за прострочення.

- Задані часові позначки — це UTC ISO рядки, наприклад: `"2026-06-18T18:30:00Z"`.
- Рівень оренди — це кількість днів до дати повернення: `1`, `3` або `7`.
- Оренду потрібно повернути до 12:00 за UTC або раніше в останній день періоду оренди. Наприклад, оренда на 1 день, оформлена в будь-який час 15 березня, має бути повернена до 12:00 за UTC 16 березня.
- За кожен день після дати та часу повернення нараховується штраф за прострочення.

Ціни наступні:

| Рівень | Базова вартість | Штраф за прострочення за день |
| - | - | - |
| 1 день | $4.99 | $3.99 |
| 3 дні | $3.99 | $2.99 |
| 7 днів | $2.99 | $0.99 |

Поверніть загальну вартість, округлену до двох десяткових знаків у форматі `"$D.CC"`.

# --hints--

`get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` має повернути `"$4.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99")`)
}})
```

`get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` має повернути `"$12.97"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97")`)
}})
```

`get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` має повернути `"$3.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99")`)
}})
```

`get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` має повернути `"$9.97"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97")`)
}})
```

`get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` має повернути `"$2.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99")`)
}})
```

`get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` має повернути `"$358.40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rental_cost(rented, returned, tier):

    return rented
```

# --solutions--

```py
from datetime import datetime, timedelta
import math

def get_rental_cost(rented, returned, tier):
    pricing = {
        1: {"base": 4.99, "late": 3.99},
        3: {"base": 3.99, "late": 2.99},
        7: {"base": 2.99, "late": 0.99},
    }
    rented_dt = datetime.fromisoformat(rented.replace("Z", "+00:00"))
    returned_dt = datetime.fromisoformat(returned.replace("Z", "+00:00"))
    due = rented_dt + timedelta(days=tier)
    due = due.replace(hour=12, minute=0, second=0, microsecond=0)
    base = pricing[tier]["base"]
    late = pricing[tier]["late"]
    if returned_dt <= due:
        return f"${base:.2f}"
    days_late = math.ceil((returned_dt - due).total_seconds() / 86400)
    return f"${base + days_late * late:.2f}"
```
