---
id: 6a19b1062d1b153d8ac76d72
title: "Завдання 324: підрахунок дублікатів символів"
challengeType: 29
dashedName: challenge-324
---

# --description--

Дано два рядки, поверніть кількість символів із другого рядка, які можна знайти в першому.

- Повторювані символи у другому рядку рахуйте окремо.

# --hints--

`duplicate_character_count("aloha", "hei")` має повернути `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("aloha", "hei"), 1)`)
}})
```

`duplicate_character_count("jambo", "bonjour")` має повернути `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("jambo", "bonjour"), 4)`)
}})
```

`duplicate_character_count("hello", "hola")` має повернути `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("hello", "hola"), 3)`)
}})
```

`duplicate_character_count("ola", "hej")` має повернути `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("ola", "hej"), 0)`)
}})
```

`duplicate_character_count("ciao", "konnichiwa")` має повернути `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("ciao", "konnichiwa"), 5)`)
}})
```

`duplicate_character_count("merhaba", "xin chao")` має повернути `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("merhaba", "xin chao"), 2)`)
}})
```

`duplicate_character_count("hello world", "hello to everyone around the world")` має повернути `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("hello world", "hello to everyone around the world"), 26)`)
}})
```

# --seed--

## --seed-contents--

```py
def duplicate_character_count(str1, str2):

    return str1
```

# --solutions--

```py
def duplicate_character_count(str1, str2):
    return sum(1 for char in str2 if char in str1)
```
