---
id: 6a1d9f98e819ed70a0e994de
title: "Завдання 332: сортування проблем"
challengeType: 29
dashedName: challenge-332
---

# --description--

Враховуючи кількість мілісекунд від останнього повідомлення у проблемі та текст останнього повідомлення, визначте, що слід зробити з проблемою за цими правилами:

- Якщо останнє повідомлення було менше ніж 7 днів тому, поверніть `"leave it"`
- Якщо останнє повідомлення було 7 або більше днів тому і його вміст містить `"bump"` (без урахування регістру), поверніть `"close it"`
- Інакше поверніть `"bump it"`

# --hints--

`triage_issue(86400000, "Lets fix it")` має повернути `"leave it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(86400000, "Lets fix it"), "leave it")`)
}})
```

`triage_issue(1209600000, "still waiting")` має повернути `"bump it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(1209600000, "still waiting"), "bump it")`)
}})
```

`triage_issue(864000000, "bump")` має повернути `"close it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(864000000, "bump"), "close it")`)
}})
```

`triage_issue(604800000, "Do we still want this?")` має повернути `"bump it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Do we still want this?"), "bump it")`)
}})
```

`triage_issue(604800000, "Bumping this")` має повернути `"close it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Bumping this"), "close it")`)
}})
```

`triage_issue(345600000, "I'll make a PR")` має повернути `"leave it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(345600000, "I'll make a PR"), "leave it")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(ms, message):

    return ms
```

# --solutions--

```py
def triage_issue(ms, message):
    seven_days = 7 * 24 * 60 * 60 * 1000
    if ms < seven_days:
        return "leave it"
    if "bump" in message.lower():
        return "close it"
    return "bump it"
```
