---
id: 6a1d9f98e819ed70a0e994dd
title: "Завдання 331: найближчий кратний"
challengeType: 29
dashedName: challenge-331
---

# --description--

Дано два цілі числа, округліть перше до найближчого кратного другого.

# --hints--

`round_to_nearest_multiple(5, 3)` має повернути `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(5, 3), 6)`)
}})
```

`round_to_nearest_multiple(17, 4)` має повернути `16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(17, 4), 16)`)
}})
```

`round_to_nearest_multiple(43, 5)` має повернути `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(43, 5), 45)`)
}})
```

`round_to_nearest_multiple(38, 11)` має повернути `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(38, 11), 33)`)
}})
```

`round_to_nearest_multiple(93, 12)` має повернути `96`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(93, 12), 96)`)
}})
```

# --seed--

## --seed-contents--

```py
def round_to_nearest_multiple(num, multiple):

    return num
```

# --solutions--

```py
def round_to_nearest_multiple(num, multiple):
    return round(num / multiple) * multiple
```
