---
id: 6a1d9f98e819ed70a0e994e0
title: "Завдання 334: Точна здача"
challengeType: 29
dashedName: challenge-334
---

# --description--

Дано ціле число в центах, поверніть кількість різних способів дати точну здачу, використовуючи пенні (1 цент), нікелі (5 центів), дайми (10 центів) і квартери (25 центів).

# --hints--

`exact_change(3)` має повернути `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(3), 1)`)
}})
```

`exact_change(9)` має повернути `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(9), 2)`)
}})
```

`exact_change(17)` має повернути `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(17), 6)`)
}})
```

`exact_change(39)` має повернути `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(39), 24)`)
}})
```

`exact_change(61)` має повернути `73`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(61), 73)`)
}})
```

`exact_change(99)` має повернути `213`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(99), 213)`)
}})
```

# --seed--

## --seed-contents--

```py
def exact_change(amount):

    return amount
```

# --solutions--

```py
def exact_change(amount):
    coins = [1, 5, 10, 25]
    dp = [0] * (amount + 1)
    dp[0] = 1
    for coin in coins:
        for i in range(coin, amount + 1):
            dp[i] += dp[i - coin]
    return dp[amount]
```
