---
id: 6a1d9f98e819ed70a0e994da
title: "Завдання 328: алгоритм Капрекара"
challengeType: 29
dashedName: challenge-328
---

# --description--

Дано 4-цифрове число, поверніть кількість разів, які потрібно застосувати алгоритм Капрекара, щоб дійти до 6174.

Алгоритм Капрекара працює так:

- Впорядкуйте цифри за спаданням, щоб утворити найбільше число
- Впорядкуйте цифри за зростанням, щоб утворити найменше число (за потреби додайте провідні нулі)
- Відніміть менше число від більшого
- Повторюйте з новим числом

# --hints--

`kaprekar(1234)` має повернути `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(1234), 3)`)
}})
```

`kaprekar(2025)` має повернути `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(2025), 6)`)
}})
```

`kaprekar(7173)` має повернути `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(7173), 4)`)
}})
```

`kaprekar(3164)` має повернути `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(3164), 7)`)
}})
```

`kaprekar(8082)` має повернути `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(8082), 2)`)
}})
```

# --seed--

## --seed-contents--

```py
def kaprekar(n):

    return n
```

# --solutions--

```py
def kaprekar(n):
    steps = 0
    current = n
    while current != 6174:
        digits = list(str(current).zfill(4))
        desc = int("".join(sorted(digits, reverse=True)))
        asc = int("".join(sorted(digits)))
        current = desc - asc
        steps += 1
    return steps
```
