---
id: 6a1d9f98e819ed70a0e994e2
title: "Завдання 336: Відповідність гороскопів"
challengeType: 29
dashedName: challenge-336
---

# --description--

Дано два рядки зі знаками зодіаку, поверніть їхній відсоток сумісності.

Знаки розташовані у колі з 12 позицій у такому порядку: `"Aries"`, `"Taurus"`, `"Gemini"`, `"Cancer"`, `"Leo"`, `"Virgo"`, `"Libra"`, `"Scorpio"`, `"Sagittarius"`, `"Capricorn"`, `"Aquarius"`, `"Pisces"`, після `"Pisces"` коло знову повертається до `"Aries"`. Знайдіть найкоротшу відстань між двома знаками і поверніть сумісність:

| Відстань | Сумісність |
|----------|------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscope_match("Libra", "Sagittarius")` має повернути `"80%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Libra", "Sagittarius"), "80%")`)
}})
```

`horoscope_match("Gemini", "Scorpio")` має повернути `"20%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Gemini", "Scorpio"), "20%")`)
}})
```

`horoscope_match("Pisces", "Aries")` має повернути `"40%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Pisces", "Aries"), "40%")`)
}})
```

`horoscope_match("Capricorn", "Cancer")` має повернути `"50%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Capricorn", "Cancer"), "50%")`)
}})
```

`horoscope_match("Aquarius", "Aquarius")` має повернути `"100%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Aquarius", "Aquarius"), "100%")`)
}})
```

`horoscope_match("Virgo", "Taurus")` має повернути `"90%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Virgo", "Taurus"), "90%")`)
}})
```

`horoscope_match("Leo", "Scorpio")` має повернути `"30%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Leo", "Scorpio"), "30%")`)
}})
```

# --seed--

## --seed-contents--

```py
def horoscope_match(sign1, sign2):

    return sign1
```

# --solutions--

```py
def horoscope_match(sign1, sign2):
    signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"]
    a = signs.index(sign1)
    b = signs.index(sign2)
    diff = abs(a - b)
    distance = min(diff, 12 - diff)
    compatibility = {0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%"}
    return compatibility[distance]
```
