---
id: 6a2037a68a0bc2aef0006004
title: "Завдання 342: ймовірність кидка кубиків"
challengeType: 29
dashedName: challenge-342
---

# --description--

За заданою кількістю шестигранних кубиків для кидка та цільовою сумою поверніть ймовірність отримати цю суму у вигляді рядка у форматі `"1 in X"`.

- Кількість кубиків буде від 1 до 6.
- Цільова сума завжди досяжна з заданою кількістю кубиків.
- Округліть `"X"` до найближчого цілого числа.

# --hints--

`get_odds(1, 5)` має повернути `"1 in 6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(1, 5), "1 in 6")`)
}})
```

`get_odds(2, 4)` має повернути `"1 in 12"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(2, 4), "1 in 12")`)
}})
```

`get_odds(3, 10)` має повернути `"1 in 8"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(3, 10), "1 in 8")`)
}})
```

`get_odds(4, 7)` має повернути `"1 in 65"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(4, 7), "1 in 65")`)
}})
```

`get_odds(5, 26)` має повернути `"1 in 111"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(5, 26), "1 in 111")`)
}})
```

`get_odds(6, 35)` має повернути `"1 in 7776"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(6, 35), "1 in 7776")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_odds(dice, target):

    return dice
```

# --solutions--

```py
def get_odds(dice, target):
    valid_combos = 0
    total_combos = 6 ** dice

    def roll(dice_left, total):
        nonlocal valid_combos
        if dice_left == 0:
            if total == target:
                valid_combos += 1
            return
        for face in range(1, 7):
            roll(dice_left - 1, total + face)

    roll(dice, 0)
    return f"1 in {round(total_combos / valid_combos)}"
```
