---
id: 6a2037a68a0bc2aef0006007
title: "Завдання 345: Змішувач слів"
challengeType: 29
dashedName: challenge-345
---

# --description--

Дано два слова, поверніть нове слово, поєднавши першу половину першого слова з другою половиною другого слова.

- Для слів непарної довжини перша половина є коротшою половиною.

# --hints--

`blend_words("turtle", "toucan")` має повернути `"turcan"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("turtle", "toucan"), "turcan")`)
}})
```

`blend_words("chipmunk", "flamingo")` має повернути `"chipingo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("chipmunk", "flamingo"), "chipingo")`)
}})
```

`blend_words("falcon", "pelican")` має повернути `"falican"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("falcon", "pelican"), "falican")`)
}})
```

`blend_words("hyena", "iguana")` має повернути `"hyana"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("hyena", "iguana"), "hyana")`)
}})
```

`blend_words("scorpion", "gorilla")` має повернути `"scorilla"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("scorpion", "gorilla"), "scorilla")`)
}})
```

`blend_words("platypus", "wolverine")` має повернути `"platerine"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("platypus", "wolverine"), "platerine")`)
}})
```

# --seed--

## --seed-contents--

```py
def blend_words(word1, word2):

    return word1
```

# --solutions--

```py
def blend_words(word1, word2):
    first_half = word1[:len(word1) // 2]
    second_half = word2[len(word2) // 2:]
    return first_half + second_half
```
