---
id: 6a22d77ddf034bc4e35b1d54
title: "Завдання 347: теорія ігор"
challengeType: 29
dashedName: challenge-347
---

# --description--

Дано два рядки однакової довжини, які представляють стратегії двох гравців у грі, поверніть рахунки у вигляді масиву `[player1, player2]`.

- Дано рядки міститимуть лише одну з двох літер: `"C"` (співпрацювати) або `"D"` (зраджувати).
- Кожен символ представляє один раунд, оцінений так:
  - Якщо обидва гравці співпрацюють, кожен отримує 3 очки.
  - Якщо обидва гравці зраджують, кожен отримує 1 очко.
  - Якщо один гравець зраджує, а інший співпрацює, зрадник отримує 5 очок, а співпрацюючий — 0.

# --hints--

`play_game("CCCC", "CCCC")` має повернути `[12, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCC", "CCCC"), [12, 12])`)
}})
```

`play_game("DDDD", "DDDD")` має повернути `[4, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDDD", "DDDD"), [4, 4])`)
}})
```

`play_game("CCDD", "CDDD")` має повернути `[5, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCDD", "CDDD"), [5, 10])`)
}})
```

`play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD")` має повернути `[24, 34]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34])`)
}})
```

`play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` має повернути `[66, 21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21])`)
}})
```

# --seed--

## --seed-contents--

```py
def play_game(p1, p2):

    return p1
```

# --solutions--

```py
def play_game(p1, p2):
    scores = [0, 0]
    for i in range(len(p1)):
        if p1[i] == 'C' and p2[i] == 'C':
            scores[0] += 3
            scores[1] += 3
        elif p1[i] == 'D' and p2[i] == 'D':
            scores[0] += 1
            scores[1] += 1
        elif p1[i] == 'D' and p2[i] == 'C':
            scores[0] += 5
        else:
            scores[1] += 5
    return scores
```
