---
id: 6a22d77ddf034bc4e35b1d5b
title: "Завдання 354: рейтинг контрастності 3"
challengeType: 29
dashedName: challenge-354
---

# --description--

Дано два масиви, які представляють значення RGB, і булевий параметр, що вказує, чи є текст великим, поверніть рейтинг контрастності WCAG, використовуючи наступний метод:

Спочатку перетворіть кожне значення RGB у відносну яскравість:

- Поділіть кожен канал `[R, G, B]` на 255, щоб отримати значення від `0` до `1`
- Застосуйте формулу гамма-корекції до кожного каналу:
  - Якщо значення каналу менше або дорівнює `0.04045`: `channel / 12.92`
  - Інакше: `((channel + 0.055) / 1.055) ^ 2.4`
- Обчисліть яскравість: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Потім обчисліть коефіцієнт контрастності, додавши 0.05 до кожного значення яскравості, а потім поділивши світліше на темніше. Світліше завжди буде першим аргументом.

Поверніть рейтинг на основі коефіцієнта контрастності, використовуючи таку таблицю:

| Рейтинг | Звичайний текст | Великий текст |
|--------|-----------------|---------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | нижче 4.5 | нижче 3.0 |

# --hints--

`get_contrast_rating([255, 255, 255], [0, 0, 0], False)` має повернути `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([255, 255, 255], [0, 0, 0], False), "AAA")`)
}})
```

`get_contrast_rating([215, 188, 188], [55, 55, 55], False)` має повернути `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([215, 188, 188], [55, 55, 55], False), "AA")`)
}})
```

`get_contrast_rating([143, 144, 210], [46, 47, 61], False)` має повернути `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([143, 144, 210], [46, 47, 61], False), "Fail")`)
}})
```

`get_contrast_rating([167, 167, 210], [53, 10, 53], True)` має повернути `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([167, 167, 210], [53, 10, 53], True), "AAA")`)
}})
```

`get_contrast_rating([135, 147, 155], [60, 70, 90], True)` має повернути `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([135, 147, 155], [60, 70, 90], True), "AA")`)
}})
```

`get_contrast_rating([125, 210, 195], [105, 130, 90], True)` має повернути `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([125, 210, 195], [105, 130, 90], True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):

    return rgb1
```

# --solutions--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):
    def to_luminance(rgb):
        channels = []
        for c in rgb:
            c = c / 255
            if c <= 0.04045:
                channels.append(c / 12.92)
            else:
                channels.append(((c + 0.055) / 1.055) ** 2.4)
        return 0.2126 * channels[0] + 0.7152 * channels[1] + 0.0722 * channels[2]

    l1 = to_luminance(rgb1)
    l2 = to_luminance(rgb2)
    ratio = (l1 + 0.05) / (l2 + 0.05)

    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
