---
id: 6a22d77ddf034bc4e35b1d58
title: "Завдання 351: пронічне число"
challengeType: 29
dashedName: challenge-351
---

# --description--

Дано число, визначте, чи є воно пронічним числом.

Пронічне число — це добуток двох послідовних цілих чисел. Наприклад, 6 є пронічним, бо 2 * 3 = 6.

# --hints--

`is_pronic(6)` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(6), True)`)
}})
```

`is_pronic(15)` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(15), False)`)
}})
```

`is_pronic(12)` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(12), True)`)
}})
```

`is_pronic(132)` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(132), True)`)
}})
```

`is_pronic(80)` має повернути `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(80), False)`)
}})
```

`is_pronic(0)` має повернути `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(0), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_pronic(n):

    return n
```

# --solutions--

```py
def is_pronic(n):
    k = int(n ** 0.5)
    return k * (k + 1) == n
```
