---
id: 6a22d77ddf034bc4e35b1d59
title: "Завдання 352: рейтинг коефіцієнта контрастності 1"
challengeType: 29
dashedName: challenge-352
---

# --description--

За заданим коефіцієнтом контрастності та булевим значенням, яке вказує, чи є текст великим, поверніть рейтинг WCAG, використовуючи таку таблицю:

| Рейтинг | Звичайний текст | Великий текст |
|--------|-----------------|---------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | нижче 4.5 | нижче 3.0 |

# --hints--

`get_contrast_rating("7.5", False)` має повернути `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("7.5", False), "AAA")`)
}})
```

`get_contrast_rating("4.8", False)` має повернути `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.8", False), "AA")`)
}})
```

`get_contrast_rating("4.2", False)` має повернути `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.2", False), "Fail")`)
}})
```

`get_contrast_rating("4.5", True)` має повернути `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.5", True), "AAA")`)
}})
```

`get_contrast_rating("3.0", True)` має повернути `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("3.0", True), "AA")`)
}})
```

`get_contrast_rating("2.7", False)` має повернути `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("2.7", False), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(ratio, is_large_text):

    return ratio
```

# --solutions--

```py
def get_contrast_rating(ratio, is_large_text):
    ratio = float(ratio)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
