---
id: 6a26df95efa55a2524399745
title: "Завдання 364: між двома відрами"
challengeType: 29
dashedName: challenge-364
---

# --description--

Дано два відра фарби, кожне з кольором у форматі RGB і рівнем заповнення. Поверніть змішаний колір RGB у вигляді масиву з трьох цілих чисел.

- Кожне відро — це об’єкт (JavaScript) або словник (Python) з властивістю `color` (масив з трьох цілих чисел `[r, g, b]`) і властивістю `fullness` (від 0 до 100).
- Змішаний колір — це зважений середній кожного каналу двох кольорів на основі рівня заповнення, при цьому кожен канал округлюється до найближчого цілого числа.

# --hints--

`mix_paint({"color": [250, 250, 250], "fullness": 50}, {"color": [0, 0, 0], "fullness": 50})` має повернути `[125, 125, 125]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [250, 250, 250], "fullness": 50}, {"color": [0, 0, 0], "fullness": 50}), [125, 125, 125])`)
}})
```

`mix_paint({"color": [250, 250, 250], "fullness": 80}, {"color": [0, 0, 0], "fullness": 20})` має повернути `[200, 200, 200]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [250, 250, 250], "fullness": 80}, {"color": [0, 0, 0], "fullness": 20}), [200, 200, 200])`)
}})
```

`mix_paint({"color": [100, 150, 200], "fullness": 30}, {"color": [100, 150, 200], "fullness": 70})` має повернути `[100, 150, 200]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [100, 150, 200], "fullness": 30}, {"color": [100, 150, 200], "fullness": 70}), [100, 150, 200])`)
}})
```

`mix_paint({"color": [143, 143, 101], "fullness": 45}, {"color": [100, 204, 204], "fullness": 90})` має повернути `[114, 184, 170]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [143, 143, 101], "fullness": 45}, {"color": [100, 204, 204], "fullness": 90}), [114, 184, 170])`)
}})
```

`mix_paint({"color": [15, 134, 249], "fullness": 29}, {"color": [97, 178, 55], "fullness": 54})` має повернути `[68, 163, 123]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [15, 134, 249], "fullness": 29}, {"color": [97, 178, 55], "fullness": 54}), [68, 163, 123])`)
}})
```

# --seed--

## --seed-contents--

```py
def mix_paint(bucket1, bucket2):

    return bucket1
```

# --solutions--

```py
def mix_paint(bucket1, bucket2):
    total = bucket1['fullness'] + bucket2['fullness']
    return [
        round((c1 * bucket1['fullness'] + c2 * bucket2['fullness']) / total)
        for c1, c2 in zip(bucket1['color'], bucket2['color'])
    ]
```
