---
id: 6a26df3e2988bcdded204896
title: "Завдання 360: тривалість у розмовній формі"
challengeType: 29
dashedName: challenge-360
---

# --description--

За заданою кількістю секунд поверніть тривалість у розмовній англійській мові.

- Розбийте тривалість на години, хвилини та секунди.
- Пропускайте будь-які нульові значення.
- Використовуйте однину або множину за потреби (`"1 hour"`, `"2 hours"`).
- Якщо є, з'єднайте останні два елементи словом `"and"`, а другий і третій з кінця — комою (`"1 hour, 2 minutes and 3 seconds"`).

# --hints--

`get_spoken_duration(3723)` має повернути `"1 hour, 2 minutes and 3 seconds"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(3723), "1 hour, 2 minutes and 3 seconds")`)
}})
```

`get_spoken_duration(7295)` має повернути `"2 hours, 1 minute and 35 seconds"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(7295), "2 hours, 1 minute and 35 seconds")`)
}})
```

`get_spoken_duration(8521)` має повернути `"2 hours, 22 minutes and 1 second"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(8521), "2 hours, 22 minutes and 1 second")`)
}})
```

`get_spoken_duration(435)` має повернути `"7 minutes and 15 seconds"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(435), "7 minutes and 15 seconds")`)
}})
```

`get_spoken_duration(14455)` має повернути `"4 hours and 55 seconds"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(14455), "4 hours and 55 seconds")`)
}})
```

`get_spoken_duration(72000)` має повернути `"20 hours"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(72000), "20 hours")`)
}})
```

`get_spoken_duration(1)` має повернути `"1 second"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(1), "1 second")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_duration(seconds):

    return seconds
```

# --solutions--

```py
def get_spoken_duration(seconds):
    hours = seconds // 3600
    minutes = (seconds % 3600) // 60
    secs = seconds % 60

    parts = []
    if hours:
        parts.append(f"{hours} {'hour' if hours == 1 else 'hours'}")
    if minutes:
        parts.append(f"{minutes} {'minute' if minutes == 1 else 'minutes'}")
    if secs:
        parts.append(f"{secs} {'second' if secs == 1 else 'seconds'}")

    if len(parts) == 1:
        return parts[0]
    if len(parts) == 2:
        return f"{parts[0]} and {parts[1]}"
    return f"{parts[0]}, {parts[1]} and {parts[2]}"
```
