---
id: 6909ed80e1c195f49564edac
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、条件文と比較演算子の使い方を学びます。

まず、`truthyOrFalsy` という名前の変数を宣言し、ブール値 `true` を代入してください。

# --hints--

`truthyOrFalsy` という名前の変数があるはずです。

```js
assert.exists(truthyOrFalsy);
```

`true` を `truthyOrFalsy` 変数に代入してください。

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
