---
id: 690a562fd611e8155599ce98
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

`truthyOrFalsy` の代入を変更して、その値を `"freeCodeCamp"` に設定してください。空でない文字列は真値なので、`true` がコンソールに表示され続けます。

# --hints--

`truthyOrFalsy` 変数は `"freeCodeCamp"` の値を持つべきです。

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
