---
id: 690a57095c4afb27bbe751b3
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

偽値とは、ブール値の文脈で `false` と評価される値のことです。偽値の例は以下の通りです：

- `""`（空の文字列）
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

このカリキュラムを進める中で、これらすべてに出会います。

まずは、`truthyOrFalsy` の代入を空の文字列に変更してください。その後、`false` がコンソールに表示されます。

# --hints--

`truthyOrFalsy` に空の文字列を代入してください。

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
