---
id: 690a51806fb743bc896c3af7
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

前のレッスンで学んだように、ブール値は `true` と `false` の値だけを持つデータ型です。

真値とは、ブール値の文脈で `true` と評価される値のことです。真値の例は次の通りです：

- 空でない文字列
- ゼロでない数値
- ブール値の `true`

そして、後で出会う他の値もあります。

`Boolean()` 関数を使うと、値の真値性を確認できます。例えば、`Boolean("Hello World!")` は `true` が真値なので `"Hello World!"` を返します。

では、`Boolean(truthyOrFalsy)` を呼び出して、その結果をコンソールにログしてください。

# --hints--

`Boolean(truthyOrFalsy)` をコンソールにログする必要があります。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
