---
id: 690a0df27ee9ee36d935d530
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

`isTimmyAGamer` という名前の変数を宣言し、ブール値 `false` を代入してください。

# --hints--

`isTimmyAGamer` という名前の変数を宣言してください。

```js
assert.exists(isTimmyAGamer);
```

`false` 変数にブール値 `isTimmyAGamer` を代入してください。

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
