---
id: 690a06833d3e83b5226ade78
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

条件文は、条件に応じてコードを実行するために使います。これは `if` キーワードと、丸括弧で囲まれた条件、波括弧で囲まれた本体から構成されます。以下は `if` 文の例です。

```js
if (condition) {
  console.log("condition is truthy");
}
```

`condition` が真値（空でない文字列、ゼロ以外の数値、`true` など）の場合、`if` 文の本体内のコードが実行されます。

変数の下に、`if` を条件に使った `hasDeveloperJob` 文を作成してください。`if` 文の本体内で、`"Timmy is employed as a developer."` をコンソールに出力してください。

# --hints--

`if` 文が必要です。

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

`if` 文は `hasDeveloperJob` を条件に使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

`"Timmy is employed as a developer."` を `if` 文の中でコンソールに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
