---
id: 690a58cc4f828a49eab1b53a
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

`hasDeveloperJob` という名前の変数を宣言し、ブール値 `true` を代入してください。

# --hints--

`hasDeveloperJob` という名前の変数があるはずです。

```js
assert.exists(hasDeveloperJob);
```

`true` を `hasDeveloperJob` 変数に代入してください。

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
