---
id: 690a57fef66aac39dffc61a2
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

truthy値と偽値が何かを理解したので、変数の宣言と`console.log`文の両方をコードから削除してください。

# --hints--

`truthyOrFalsy`という変数を持ってはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s+truthyOrFalsy\s*=/);
```

コードに`console.log(Boolean(truthyOrFalsy))`があってはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
