---
id: 690a3ad44874a928a2fe3077
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

比較演算子を使うと値を比較して、より複雑な条件文を書くことができます。

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

例えば、上のコードは変数 `"Happy Birthday!"` と `today` の値が等しい場合にのみ `birthday` をコンソールに出力します。

3つ目の `if` 文を作成してください。その条件には `>=` 演算子を使って、`timmyAge` が `16` 以上かどうかをチェックしてください。

新しい `if` 文の中で、`"Timmy is old enough to drive."` をコンソールに出力してください。

# --hints--

コードには3つ目の `if` 文があるはずです。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

3つ目の `if` 文は `timmyAge >= 16` を条件に使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

3つ目の `"Timmy is old enough to drive."` 文の中で `if` をコンソールに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
