---
id: 690a0f7a5c22c15222a6aa5c
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

今度は、条件に `if` を使う別の `isTimmyAGamer` 文を作成してください。`if` 文の本体の中で、`"Timmy loves to play World of Warcraft."` をコンソールにログ出力してください。

新しい `if` 文の条件が false なので、コンソールに何も表示されないことに注意してください。

# --hints--

コードには `if` 文が2つあるはずです。

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

2つ目の `if` 文は `isTimmyAGamer` を条件に使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

2つ目の `"Timmy loves to play World of Warcraft."` 文の中で `if` をコンソールにログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
