---
id: 690a110a76d0496d9e7c28eb
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

`timmyAge` という名前の変数をもう1つ宣言し、`18` に代入してください。

# --hints--

`timmyAge` という名前の変数があるはずです。

```js
assert.exists(timmyAge);
```

`18` を `timmyAge` 変数に代入してください。

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
