---
id: 66c06fad3475cd92421b9ac2
title: Побудова калькулятора високосного року
challengeType: 26
dashedName: build-a-leap-year-calculator
---

# --description--

Високосний рік — це рік, який ділиться націло на `4`, за винятком років, які діляться націло на `100` і не діляться націло на `400`. Наприклад, `2000` — високосний рік, а `1900` — ні. Також у високосному році у лютому є додатковий день — 29-й день місяця.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Оголосіть функцію з назвою `isLeapYear`, яка приймає число як аргумент.
2. Поза межами функції оголосіть змінну `year`, яка зберігає значення року, який ви хочете перевірити.
3. Усередині функції використайте інструкцію `if/ else` або тернарний оператор, щоб перевірити, чи є рік високосним.
4. Щоб перевірити, чи є рік високосним, виконайте такі умови:

   - Якщо рік ділиться націло на `4`, тоді це високосний рік.
   - Якщо рік також ділиться націло на `100`, тоді це не високосний рік.
   - Якщо рік також ділиться націло на `400`, тоді це високосний рік.

5. Якщо рік високосний, поверніть `[year] is a leap year.`. Інакше поверніть `[year] is not a leap year.`. Ви заміните `[year]` на параметр, оголошений у функції `isLeapYear`.
6. Ви повинні викликати функцію `isLeapYear` з `year` як аргументом і призначити результат змінній з назвою `result`.
7. Ви повинні вивести змінну `result` у консоль за допомогою `console.log()`.

# --hints--

Ви повинні оголосити функцію з назвою `isLeapYear`.

```js
assert.isFunction(isLeapYear);
```

Функція `isLeapYear` має приймати число як аргумент.

```js
assert.match(isLeapYear.toString(), /\s*isLeapYear\(\s*\w+\s*\)/);
```

Ви повинні оголосити змінну `year` і призначити їй значення для перевірки, чи є рік високосним.

```js
assert.isDefined(year);
assert.isNumber(year);
assert.isAtLeast(year, 0);
assert.isTrue(Number.isInteger(year));
```

Змінна `year` не повинна бути порожньою.

```js
assert.isNotNull(year);
```

Якщо `2024` є значенням змінної `year`, тоді `result` має бути `2024 is a leap year.`

```js
assert.strictEqual(isLeapYear(2024), '2024 is a leap year.');
```

Якщо `2000` є значенням змінної `year`, тоді `result` має бути `2000 is a leap year.`

```js
assert.strictEqual(isLeapYear(2000), '2000 is a leap year.');
```

Якщо `1900` є значенням змінної `year`, тоді `result` має бути `1900 is not a leap year.`

```js
assert.strictEqual(isLeapYear(1900), '1900 is not a leap year.');

```

Ви повинні викликати функцію `isLeapYear` і передати `year` як параметр.

```js
assert.match(__helpers.removeJSComments(code), /isLeapYear\(\s*year\s*\)/);
```

Ви повинні оголосити змінну `result`.

```js
assert.isDefined(result);
```

Ви повинні зберегти результат виклику функції `isLeapYear` у змінній з назвою `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*isLeapYear\(\s*year\s*\)/);
```

Ви повинні вивести `result` у консоль за допомогою `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function isLeapYear(year) {
  if ((year % 4 === 0 && year % 100 !== 0) || year % 400 === 0) {
    return(year + " is a leap year.");
  } else {
    return(year + " is not a leap year.");
  }
}

const year = 2024;
const result = isLeapYear(year);
console.log(result);
```
