---
id: a97fd23d9b809dac9921074f
title: オプションの引数を持つ合計関数を作成する
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

このラボでは、最大2つの引数を受け取り、それらを合計する関数を作成します。ただし、引数が1つだけの場合は、2つ目の数値を待つ関数を返します。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. 関数`addTogether`を作成してください。
2. 関数が2つの引数を受け取った場合、2つの引数の合計を返す必要があります。
3. 引数が1つだけの場合、`addTogether`は関数を返す必要があります。
   - 返された関数が1つの引数で呼び出されたとき、その合計を返す必要があります。

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. どちらかの引数が有効な数値でない場合は、`undefined`を返してください。

# --hints--

`addTogether`関数を用意してください。

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)`は5を返すべきです。

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)`は53.4を返すべきです。

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)`は`undefined`を返すべきです。

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)`は`undefined`を返すべきです。

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")`は`undefined`を返すべきです。

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)`は関数を返すべきです。

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)`は12を返すべきです。

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])`は`undefined`を返すべきです。

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")`は`undefined`を返すべきです。

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
