---
id: a97fd23d9b809dac9921074f
title: Construir uma Função de Soma com Argumentos Opcionais
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

Neste laboratório você construirá uma função que aceita até dois argumentos, e os soma, mas se houver apenas um argumento retorna uma função que espera pelo segundo número para somar.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Crie uma função `addTogether`.
2. Se a função receber dois argumentos, a função deve retornar a soma dos dois argumentos.
3. Se apenas um argumento for fornecido, `addTogether` deve retornar uma função.
   - Quando a função retornada for chamada com um único argumento, ela deve retornar a soma.

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. Se qualquer um dos argumentos não for um número válido, retorne `undefined`.

# --hints--

Você deve ter a função `addTogether`.

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` deve retornar 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` deve retornar 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` deve retornar `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` deve retornar `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` deve retornar `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` deve retornar uma função.

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` deve retornar 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` deve retornar `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` deve retornar `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
