---
id: a3566b1109230028080c9345
title: 모든 숫자 합산 알고리즘 설계
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

이 실습에서는 모든 숫자를 합산하는 알고리즘을 설계해야 합니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 두 개의 숫자를 요소로 가진 배열을 인수로 받는 `sumAll`라는 이름의 함수가 있어야 합니다.
1. `sumAll([n, m])`은 `n`와 `m`의 합과 그 사이에 있는 모든 숫자의 합을 반환해야 합니다. 가장 작은 숫자가 항상 먼저 오지는 않습니다. 예를 들어, `sumAll([4,1])`는 `10`를 반환해야 하는데, `1`과 `4` (둘 다 포함) 사이에 있는 모든 숫자의 합이 `10`이기 때문입니다.

# --hints--

`sumAll`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])`은 숫자를 반환해야 합니다.

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])`은 `10`를 반환해야 합니다.

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])`은 `10`를 반환해야 합니다.

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])`은 `45`를 반환해야 합니다.

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])`은 `45`를 반환해야 합니다.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
