---
id: 60a3e3396c7b40068ad6997f
title: Etapa 22
challengeType: 0
dashedName: step-22
---

# --description--

Agregar 1 píxel de espacio de relleno en la parte superior, inferior, izquierda y derecha del lienzo cambió sus dimensiones a 502 píxeles x 602 píxeles.

Reemplazar la propiedad `padding` con `overflow` establecida en `hidden` - cambiando el lienzo de nuevo a sus dimensiones originales.

# --hints--

Debe eliminar la propiedad `padding` del selector `.canvas`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

Debe establecer la propiedad `overflow` en `hidden`.

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

El elemento `.canvas` debería tener un valor de `overflow` de `hidden`.

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
