---
id: 60a3e3396c7b40068ad69971
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Tiempo para CSS.

Even though your `<div>` has no text, it's still treated as a box with content. Escribe una regla CSS que utilice el selector de clase `.canvas` y asigna su `width` a 500 píxeles. Aquí tienes una regla CSS que establece el ancho de la clase `card` a 300 píxeles:

```css
.card {
  width: 300px;
}
```

# --hints--

Tu código debe tener un selector `.canvas`.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

Debe establecer la propiedad `width` a `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

Tu elemento `.canvas` debe tener un `width` de `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
