---
id: a103376db3ba46b2d50db289
title: Tekeleza mabadiliko ya mfuatano wa herufi kuwa spinal case
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

Spinal case ni muundo wa mfuatano wa herufi ambapo maneno yote yameandikwa kwa herufi ndogo na yamegawanywa kwa alama za msalaba (-). `"this-is-spinal-tap"` ni mfano wa mfuatano wa herufi katika muundo wa spinal case.

Katika maabara hii, utaunda kitendakazi kinachobadilisha mfuatano wa herufi uliotolewa kuwa spinal case.

**Lengo**: Timiza hadithi za mtumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za mtumizi:**

1. Unapaswa kuunda kitendakazi chenye jina `spinalCase`.
2. Kitendakazi `spinalCase` kinapaswa kupokea hoja moja, mfuatano wa herufi.
3. Kitendakazi `spinalCase` kinapaswa kurudisha mfuatano wa herufi katika muundo wa spinal case. Kwa mfano, ikiwa hoja ni `ProductLanding page`, kitendakazi kinapaswa kurudisha `product-landing-page`.

# --hints--

Unapaswa kuunda kitendakazi chenye jina `spinalCase`.

```js
assert.isFunction(spinalCase);
```

`spinalCase` inapaswa kupokea hoja moja.

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` inapaswa kurudisha mfuatano wa herufi `this-is-spinal-tap`.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` inapaswa kurudisha mfuatano wa herufi `this-is-spinal-tap`.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` inapaswa kurudisha mfuatano wa herufi `the-andy-griffith-show`.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` inapaswa kurudisha mfuatano wa herufi `teletubbies-say-eh-oh`.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` inapaswa kurudisha mfuatano wa herufi `all-the-small-things`.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
