---
id: 641cdf57c3f7ee276e1d9b32
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

이벤트 리스너로 돌아가서 `result` 요소의 텍스트를 업데이트해야 합니다. 이 작업을 위해 삼항 연산자를 사용할 수 있습니다.

다음은 삼항 연산자의 결과를 요소의 텍스트 콘텐츠에 할당하는 예시입니다:

```js
el.textContent = condition ? "Use this text if the condition is true" : "Use this text if the condition is false";
```

`if` 문장 뒤에, `result`를 인수로 하여 `isSpam()`을 호출한 결과의 진릿값에 따라 `messageInput.value` 요소의 텍스트를 업데이트하세요. 참이면 `textContent` 요소의 `result` 속성을 `"Oh no! This looks like a spam message."`로 설정하세요. 그렇지 않으면 `"This message does not seem to contain any spam."`로 설정하세요.

그런 다음 `messageInput` 요소의 `value` 속성을 빈 문자열로 설정하세요.

# --hints--

`checkMessageButton`을 클릭하고 `isSpam(messageInput.value)`가 참일 때, `textContent`의 `result` 속성을 `"Oh no! This looks like a spam message."`로 설정해야 합니다.

```js
const temp = isSpam;
isSpam = (msg) => true;
messageInput.value = "spam";
try {
  checkMessageButton.dispatchEvent(new Event("click"));
  assert.equal(document.querySelector("#result-message").textContent, "Oh no! This looks like a spam message.");
} finally {
  isSpam = temp;
}
```

`checkMessageButton`을 클릭하고 `isSpam(messageInput.value)`가 거짓일 때, `textContent`의 `result` 속성을 `"This message does not seem to contain any spam."`로 설정해야 합니다.

```js
messageInput.value = "spam";
checkMessageButton.dispatchEvent(new Event("click"));
assert.equal(document.querySelector("#result-message").textContent, "This message does not seem to contain any spam.");
```

`textContent`의 `result`을 설정한 후, `value` 요소의 `messageInput` 속성을 빈 문자열로 재설정해야 합니다.

```js
messageInput.value = "spam";
checkMessageButton.dispatchEvent(new Event("click"));
assert.isEmpty(messageInput.value);

messageInput.value = "spam";
const temp = isSpam;
isSpam = (msg) => true;
try {
  checkMessageButton.dispatchEvent(new Event("click"));
  assert.isEmpty(messageInput.value);
} finally {
  isSpam = temp;
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const isSpam = (msg) => false;

--fcc-editable-region--
checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }


});
--fcc-editable-region--
```
