---
id: 6421f98f4999d1179ce37cb4
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

배열에는 `some()` 메서드가 있습니다. `filter()` 메서드처럼 `some()`는 배열의 요소를 인수로 받는 콜백 함수를 인수로 받습니다. `some()` 메서드는 콜백 함수가 배열의 적어도 한 요소에 대해 `true`를 반환하면 `true`를 반환합니다.

다음은 배열에 대문자가 하나라도 있는지 확인하는 `.some()` 메서드 호출 예시입니다.

```js
const arr = ["A", "b", "C"];
arr.some(letter => letter === letter.toUpperCase());
```

`isSpam()` 함수를 `some()` 메서드를 사용하도록 업데이트하세요. 이 메서드는 `msg`를 `denyList`에 있는 정규 표현식 중 하나에 대해 테스트했을 때 `true`를 반환하는지 확인합니다.

# --hints--

`isSpam` 함수는 `.some()` 메서드를 사용해야 합니다.

```js
let flag = false;
const temp = Array.prototype.some;
Array.prototype.some = (arg) => {flag = true};
try {
  isSpam("");
  assert.isTrue(flag);
} finally {
  Array.prototype.some = temp;
}
```

`isSpam` 함수는 `denyList.some()`를 호출해야 합니다.

```js
let flag = false;
const temp = denyList.some;
denyList.some = (arg) => {flag = true};
try {
  isSpam("");
  assert.isTrue(flag);
} finally {
  denyList.some = temp;
} 
```

`isSpam` 함수는 `denyList.some()`의 결과를 반환해야 하며, 콜백 함수는 `.test(msg)`의 각 항목에 대해 `denyList`을 평가해야 합니다.

```js
const help = "plEase hElp";
const assist = "AssIst ME";
const notSpam = "not spam";
const mockSpam = (msg) => denyList.some((regex) => regex.test(msg));
denyList = [helpRegex, /test1/, /test2/]
assert.strictEqual(isSpam(help), mockSpam(help));
assert.strictEqual(isSpam(assist), mockSpam(assist));
assert.strictEqual(isSpam(notSpam), mockSpam(notSpam));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;

const denyList = [helpRegex];

--fcc-editable-region--
const isSpam = (msg) => helpRegex.test(msg);
--fcc-editable-region--

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
