---
id: 642213bf8d38b0227ed6ab0b
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

다음 정규 표현식은 `free money` 같은 문자열을 찾습니다. `freeRegex` 변수를 선언하고 `free money` 문자열과 일치하는 정규 표현식을 할당하세요. 대소문자를 구분하지 않도록 만드세요.

그런 다음, 새 정규 표현식을 `denyList` 배열에 추가하여 테스트할 수 있게 하세요.

# --hints--

`freeRegex` 변수를 선언해야 합니다.

```js
assert.exists(freeRegex);
```

`freeRegex` 변수는 정규 표현식이어야 합니다.

```js
assert.instanceOf(freeRegex, RegExp);
```

`freeRegex` 는 `free money` 와 일치해야 합니다.

```js
const spam = "free money";
assert.deepEqual(spam.match(freeRegex), spam.match(/free money/));
```

`freeRegex` 는 대소문자를 구분하지 않아야 합니다.

```js
assert.include(freeRegex.flags, 'i');
```

`denyList` 배열에 `freeRegex` 변수를 포함해야 합니다.

```js
assert.deepInclude(denyList, freeRegex);
```

`denyList` 배열에 `dollarRegex` 변수를 포함해야 합니다.

```js
assert.deepInclude(denyList, dollarRegex);
```

`denyList` 배열에 `helpRegex` 변수를 포함해야 합니다.

```js
assert.deepInclude(denyList, helpRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;
const dollarRegex = /[0-9]+\s*(?:hundred|thousand|million|billion)?\s+dollars/i;
--fcc-editable-region--

const denyList = [helpRegex, dollarRegex];
--fcc-editable-region--
const isSpam = (msg) => denyList.some((regex) => regex.test(msg));

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
