---
id: 642206e054eef81b5e3092ed
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

텍스트 `dollars` 앞에 숫자를 매치해야 합니다. `0|1|2`처럼 일일이 작성할 수도 있지만, 정규 표현식에는 이를 더 쉽게 해주는 기능이 있습니다.

<dfn>문자 클래스</dfn>는 대괄호로 정의하며, 괄호 안에 있는 문자 중 하나와 매치됩니다. 예를 들어, `[aeiou]`는 목록 `aeiou`에 있는 문자 중 하나와 매치됩니다. 하이픈을 사용해 문자 범위를 정의할 수도 있습니다. 예를 들어, `[a-z]`는 `a`부터 `z`까지의 문자 중 하나와 매치됩니다.

숫자 `0`부터 `9`까지를 매치하는 문자 클래스를 `dollarRegex` 표현식에 추가하세요 — 숫자는 단어 `dollars` 앞에 와야 하며, 숫자와 단어 사이에는 공백이 있어야 합니다.

# --hints--

`dollarRegex`에는 문자 클래스를 사용해야 합니다.

```js
assert.match(dollarRegex.source, /\[.*\]/);
```

문자 클래스는 `0-9`여야 합니다.

```js
assert.match(dollarRegex.source, /\[0-9\]/);
```

`dollarRegex`은 여전히 `"1 dollars"`와 매치되어야 합니다.

```js
const oneDollar = '1 dollars'
assert.deepEqual(oneDollar.match(dollarRegex), oneDollar.match(/[0-9] dollars/i));
```

`dollarRegex`는 대소문자를 구분하지 않아야 합니다.

```js
assert.include(dollarRegex.flags, 'i');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;
--fcc-editable-region--
const dollarRegex = /dollars/i;
--fcc-editable-region--

const denyList = [helpRegex, dollarRegex];

const isSpam = (msg) => denyList.some((regex) => regex.test(msg));

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
