---
id: 66da326c02141df538f29ba5
title: Construir una página de pago
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Instrucciones:**

1. Debes tener un elemento `h1` con el texto `Checkout`.
1. Debes tener dos elementos `section` inmediatamente después del elemento `h1`.
1. Debes tener un elemento `h2` con el texto `Your Cart` dentro de la primera sección.
1. Debes tener una imagen de un artículo en la primera sección con texto alternativo adecuado. Puedes usar esta imagen si lo deseas: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. Debes tener un elemento `h2` con el texto `Payment Information` dentro de la segunda sección.
1. Debes tener un elemento `form` dentro de la segunda sección.
1. Debes tener un input con un `id` y `name` de `card-name` dentro de tu formulario y un `label` asociado con él.
1. Debes tener un input con un `id` y `name` de `card-number` dentro de tu formulario y un `label` asociado con él.
1. Al menos dos de tus elementos `input` deben ser requeridos y tener un atributo `aria-required` establecido en `true`.

# --hints--

Debes tener un elemento `h1` con el texto `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

Debes tener solo un elemento `h1` en tu página.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Debes tener al menos dos elementos `section` después del elemento `h1`.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

Debes tener un elemento `h2` con el texto `Your Cart` dentro de la primera sección.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

Debes tener al menos una imagen con texto alternativo en tu primera sección.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

Debes tener un elemento `h2` con el texto `Payment Information` dentro de la segunda sección.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

Debes tener un elemento `form` dentro del segundo elemento `section`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

Debes tener un elemento `input` con un `id` y `name` de `card-name` dentro de tu formulario.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

Debes tener un elemento `input` con un `id` y `name` de `card-number` dentro de tu formulario.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

Todos tus elementos `input` que no son de `type` `submit` deben tener un elemento `label` asociado con ellos.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => assert.exists(document.querySelector(`label[for="${input.id}"]`)));
```

Al menos dos de tus elementos `input` deben ser requeridos y tener un atributo `aria-required` establecido en `true`.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required && input.getAttribute('aria-required') === 'true');
assert.isAtLeast(requiredInputs?.length, 2)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name</label>
                    <input type="text" id="card-name" name="card-name" required aria-required="true">
                </div>
                <div>
                    <label for="card-number">Card Number</label>
                    <input type="text" id="card-number" name="card-number" required aria-required="true">
                </div>
                <div>
                    <label for="expiry-date">Expiry Date</label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required aria-required="true">
                </div>
                <div>
                    <label for="cvv">CVV</label>
                    <input type="text" id="cvv" name="cvv" required aria-required="true" aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
