---
id: 579e2a2c335b9d72dd32e05c
title: 슬라이스와 스플라이스 알고리즘 구현하기
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

이 실습에서는 두 배열을 병합하는 알고리즘을 만들어야 합니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 두 배열과 인덱스를 받는 `frankenSplice` 함수를 만드세요.
1. 첫 번째 배열의 각 요소를 주어진 인덱스부터 두 번째 배열에 순서대로 복사하고, 결과 배열을 반환하세요.
1. 함수 실행 후 입력 배열들은 변경되지 않아야 합니다.

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)`은 `[4, 1, 2, 3, 5]`를 반환해야 합니다.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)`은 `["a", 1, 2, "b"]`를 반환해야 합니다.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)`은 `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`를 반환해야 합니다.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

첫 번째 배열의 모든 요소는 원래 순서대로 두 번째 배열에 추가되어야 합니다. `frankenSplice([1, 2, 3, 4], [], 0)`은 `[1, 2, 3, 4]`를 반환해야 합니다.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

함수 실행 후 첫 번째 배열은 변경되지 않아야 합니다.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

함수 실행 후 두 번째 배열은 변경되지 않아야 합니다.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
