---
id: 587d774c367417b2b2512a9d
title: 知道何時應該將 Alt Text 留空
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9P4t2'
forumTopicId: 301019
dashedName: know-when-alt-text-should-be-left-blank
---

# --description--

在上一個挑戰中，你學到使用 `img` 頁籤時必須包含 `alt` 屬性。然而，有時影像會與已描述它們的標題群組在一起，或僅用於裝飾。在這些情況下，`alt` 文字可能顯得多餘或不必要。

當影像已經用文字內容說明或不為頁面增加意義時，`img` 仍然需要一個 `alt` 屬性，但可以設定為空字串。以下是一個範例：

```html
<img src="visualDecoration.jpeg" alt="">
```

背景影像通常也屬於「裝飾性」標籤。然而，它們通常是透過 CSS 規則套用，因此不屬於螢幕閱讀器所處理的標記部分。

**注意：** 對於帶有標題的影像，你仍然可能想要包含 `alt` 文字，因為它有助於搜尋引擎編目影像的內容。

# --instructions--

Camper Cat 已經為他網站的部落格部分編寫了一個骨架頁面。他計劃在兩篇文章之間加入一個視覺分隔，使用一張武士刀的裝飾圖片。為 `img` 頁籤添加一個 `alt` 屬性，並將其設為空字串。（請注意，圖片的 `src` 並未連結到實際的檔案－不用擔心畫面中沒有顯示出任何刀劍。）

# --hints--

你的 `img` 頁籤應該有一個 `alt` 屬性。

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

`alt` 屬性應該設為空字串。

```js
assert.isEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg" alt="">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```
