---
id: 587d774e367417b2b2512aa0
title: 將內容包裹在 `article` 元素中
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp79S3'
forumTopicId: 301029
dashedName: wrap-content-in-the-article-element
---

# --description--

`article` 是另一個新的 HTML5 元素，能為你的標記添加語意。`article` 是一個區段元素，用來包裹獨立且自包含的內容。該標籤非常適合用於部落格文章、論壇貼文或新聞報導。

判斷內容是否能獨立存在通常是一種判斷呼叫，但你可以使用幾個簡單的測試。問問你自己，如果你移除所有週遭環境，該內容是否仍然有意義？同樣地，對於文字，如果內容出現在 RSS 訂閱中，是否仍然成立？

請記住，使用輔助技術的人士依賴有組織且語意明確的標記來更好地理解你的作品。

**注意：** `section` 元素也是 HTML5 的新元素，且其語意與 `article` 稍有不同。`article` 用於獨立內容，而 `section` 用於群組主題相關的內容。它們可以根據需要互相使用。 例如，若一本書是 `article`，那麼每一章節就是一個 `section`。當內容群組之間沒有關聯時，則使用 `div`。

`<div>` - 群組內容  
`<section>` - 群組相關內容  
`<article>` - 群組獨立且自包含的內容

# --instructions--

Camper Cat 使用 `article` 頁籤包裹他部落格頁面上的文章，但他忘記在最上面那篇文章周圍使用它們。將 `div` 頁籤改為使用 `article` 頁籤。

# --hints--

你的程式碼應該有三個 `article` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('article') ,3);
```

你的程式碼不應該有任何 `div` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <div>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </div>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```
