---
id: 587d7790367417b2b2512ab0
title: 使用 tabindex 為元素添加鍵盤焦點
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMDHW'
forumTopicId: 301027
dashedName: use-tabindex-to-add-keyboard-focus-to-an-element
---

# --description--

HTML 的 `tabindex` 屬性有三個與元素鍵盤焦點相關的不同函式（程式）。當它出現在頁籤上時，表示該元素可以被聚焦。其值（一個正數、負數或零的整數）決定行為。

某些元素，例如聯結和表單控制元件，當使用者使用 Tab 鍵瀏覽頁面時會自動獲得鍵盤焦點。焦點的順序與元素在 HTML 原始碼中的出現順序相同。透過在其他元素（例如 `div`、`span` 和 `p`）上放置 `tabindex="0"` 屬性，也可以賦予相同的功能。以下是一個範例：

```html
<div tabindex="0">I need keyboard focus!</div>
```

**注意：** 負的 `tabindex` 值（通常為 -1）顯示元素是可聚焦的，但無法透過鍵盤到達。這種方法通常用於以程式方式將焦點帶到內容（例如當用於快顯視窗的 `div` 被作用時），且超出這些挑戰的範圍。

# --instructions--

Camper Cat 創建了一個新的調查來收集他的使用者資訊。他知道輸入欄位會自動獲得鍵盤焦點，但他想確保他的鍵盤使用者在使用 Tab 鍵瀏覽條款時會暫停在說明上。為 `p` 頁籤添加一個 `tabindex` 屬性並將其值設為 `0`。額外提示－使用 `tabindex` 也能啟用 CSS 偽類 `:focus` 在 `p` 頁籤上的作用。

# --hints--

你的程式碼應該為包含表單說明的 `p` 頁籤添加 `tabindex` 屬性。

```js
assert.isNotNull(document.querySelector('p')?.getAttribute('tabindex'));
```

你的程式碼應該將 `p` 標記上的 `tabindex` 屬性設為 0。

```js
assert.equal(document.querySelector('p')?.getAttribute('tabindex'), '0');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p>Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p tabindex="0">Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
