---
id: 587d778f367417b2b2512aad
title: >-
  Avoid Colorblindness Issues by Carefully Choosing Colors that Convey
  Information
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437as3'
forumTopicId: 301011
dashedName: >-
  avoid-colorblindness-issues-by-carefully-choosing-colors-that-convey-information
---

# --description--

色盲有各種不同的形式。這些形式可以從對某一波長光的敏感度降低，到完全無法辨識顏色不等。最常見的形式是對綠色的敏感度降低。

例如，如果兩個相似的綠色是你內容的前景色和背景色，色盲使用者可能無法分辨它們。相近的顏色可以被視為色輪上的鄰居，當傳達重要資訊時應避免這些組合。

**注意：** 有些線上取色工具包含不同色盲型別顏色顯示的視覺模擬。這些是除了線上對比度檢查計算器之外的絕佳資源。

# --instructions--

Camper Cat 正在測試一個重要按鈕的不同樣式，但黃色（`#FFFF33`）的 `background-color` 和綠色（`#33FF33`）的文字 `color` 是色輪上相鄰的色調，對某些色盲使用者來說幾乎無法區分。（它們相似的亮度也未通過對比度比率檢查）。將文字 `color` 改為深藍色（`#003366`）即可解決這兩個問題。

# --hints--

你的程式碼應該將 `button` 的文字 `color` 更改為深藍色。

```js
const button = document.querySelector('button');
const buttonColor = window.getComputedStyle(button).color; 
assert.equal(buttonColor, 'rgb(0, 51, 102)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  button {
    color: #33FF33;
    background-color: #FFFF33;
    font-size: 14px;
    padding: 10px;
  }
  </style>
</head>
<body>
  <header>
    <h1>Danger!</h1>
  </header>
  <button>Delete Internet</button>
</body>
```

# --solutions--

```html
<head>
  <style>
    button {
      color: #003366;
      background-color: #FFFF33;
      font-size: 14px;
      padding: 10px;
    }
  </style>
</head>
<body>
  <header>
    <h1>Danger!</h1>
  </header>
  <button>Delete Internet</button>
</body>
```
