---
id: 587d7790367417b2b2512aaf
title: 使用 HTML 存取鍵使聯結可導覽
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQvmaTp'
forumTopicId: 301021
dashedName: make-links-navigable-with-html-access-keys
---

# --description--

HTML 提供 `accesskey` 屬性來指定快捷鍵，以作用或將焦點帶到元素。為你的元素添加 `accesskey` 屬性可以讓僅使用鍵盤的使用者更高效地導航。

HTML5 允許在任何元素上使用這個屬性，但當它用於互動式元素時特別有用。這包括聯結、按鈕和表單控制元件。

這裡有一個範例：

```html
<button accesskey="b">Important Button</button>
```

# --instructions--

Camper Cat 想要在兩個部落格文章標題周圍的聯結設置鍵盤快速鍵，讓他網站的使用者可以快速存取完整故事。為兩個聯結添加 `accesskey` 屬性，並將第一個設為 `g`（代表 Garfield），第二個設為 `c`（代表 Chuck Norris）。

# --hints--

你的程式碼應該為 `id` 為 `first` 的 `a` 頁籤添加 `accesskey` 屬性。

```js
assert.isTrue(document.querySelector('#first')?.hasAttribute('accesskey'));
```

你的程式碼應該為具有 `id` 為 `second` 的 `a` 頁籤添加 `accesskey` 屬性。

```js
assert.isTrue(document.querySelector('#second')?.hasAttribute('accesskey'));
```

你的程式碼應該在 `id` 為 `first` 的 `a` 頁籤上設定 `accesskey` 屬性為 `g`。請注意大小寫有區別。

```js
assert.equal(document.querySelector('#first')?.getAttribute('accesskey'), 'g');
```

你的程式碼應該在 `id` 為 `second` 的 `a` 頁籤上設定 `accesskey` 屬性為 `c`。請注意大小寫有區分。

```js
assert.equal(document.querySelector('#second')?.getAttribute('accesskey'), 'c');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" accesskey="g" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" accesskey="c" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
