---
id: 587d778a367417b2b2512aa5
title: 使用 figure 元素提升圖表的存取性
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMqtE'
forumTopicId: 301015
dashedName: improve-chart-accessibility-with-the-figure-element
---

# --description--

HTML5 介紹了 `figure` 元素和相關的 `figcaption`。這些項目一起使用時，會包裹一個視覺表現（例如影像、圖表或圖示）以及其說明文字。將這些元素一起包裹，透過語意上群組相關內容並提供解釋 `figure` 的文字替代，帶來雙重的存取性提升。

對於像圖表這類的資料視覺化，說明文字可以用來簡要註明趨勢或結論，供視覺障礙使用者參考。另一個挑戰涵蓋如何將圖表資料的表格版本（使用 CSS）移出螢幕，供螢幕閱讀器使用者使用。

這裡有一個範例－請注意 `figcaption` 位於 `figure` 頁籤內，且可以與其他元素結合：

```html
<figure>
  <img src="roundhouseDestruction.jpeg" alt="Photo of Camper Cat executing a roundhouse kick">
  <br>
  <figcaption>
    Master Camper Cat demonstrates proper form of a roundhouse kick.
  </figcaption>
</figure>
```

# --instructions--

Camper Cat 正在努力製作一個堆疊長條圖，顯示每週花在潛行、戰鬥和武器訓練的時間。幫助他更好地結構他的頁面，將他使用的 `div` 標記改為 `figure` 標記，並將包圍說明文字的 `p` 標記改為 `figcaption` 標記。

# --hints--

你的程式碼應該只有一個 `figure` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('figure') , 1);
```

你的程式碼應該只有一個 `figcaption` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('figcaption') , 1);
```

你的程式碼不應該有任何 `div` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

你的程式碼中不應該有任何 `p` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('p') , 0);
```

`figcaption` 應該是 `figure` 頁籤的子元素。

```js
const figure = document.querySelector('figure');
const children = figure?.querySelectorAll(`:scope ${'figcaption'}`);
assert.lengthOf(children, 1);
```

你的 `figure` 元素應該有一個結束標籤。

```js
assert.isTrue(
    code.match(/<\/figure>/g)?.length === code.match(/<figure>/g)?.length
);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>

      <!-- Only change code below this line -->
      <div>
        <!-- Stacked bar chart will go here -->
        <br>
        <p>Breakdown per week of time to spend training in stealth, combat, and weapons.</p>
      </div>
      <!-- Only change code above this line -->

    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>
      <figure>
        <!-- Stacked bar chart will go here -->
        <br>
        <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
      </figure>
    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
