---
id: 587d774c367417b2b2512a9c
title: 為視覺障礙者的無障礙性，為影像添加文字替代方案
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 16628
dashedName: add-a-text-alternative-to-images-for-visually-impaired-accessibility
---

# --description--

你很可能在其他挑戰中看過 `img` 頁籤上的 `alt` 屬性。`alt` 文字描述影像的內容，並提供其文字替代。當影像無法載入或使用者無法看到時，`alt` 屬性會有所幫助。搜尋引擎也會使用它來了解影像包含的內容，以便將其納入搜尋結果。以下是一個範例：

```html
<img src="importantLogo.jpeg" alt="Company logo">
```

視覺障礙者依賴螢幕閱讀器將網頁內容轉換為音訊介面。如果資訊僅以視覺方式呈現，他們將無法取得該資訊。對於圖片，螢幕閱讀器可以存取 `alt` 屬性並朗讀其內容，以傳達關鍵資訊。

良好的 `alt` 文字為讀者提供影像的簡短描述。你應該始終在你的影像上包含 `alt` 屬性。根據 HTML5 規範，這現在被視為強制性的。

# --instructions--

Camper Cat 剛好既是程式設計忍者也是真正的忍者，他正在建置一個網站來分享他的知識。他想使用的評測圖片展示了他的技能，應該讓所有網站訪客欣賞。為 `img` 頁籤添加一個 `alt` 屬性，說明 Camper Cat 正在練習空手道。（影像的 `src` 並未連結到實際的檔案，因此你應該會在顯示中看到 `alt` 文字。）

# --hints--

你的 `img` 頁籤應該有一個 `alt` 屬性，且該屬性不應該是空的。

```js
assert.isNotEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<img src="doingKarateWow.jpeg">
```

# --solutions--

```html
<img src="doingKarateWow.jpeg" alt="Someone doing karate">
```
