---
id: 587d7790367417b2b2512ab1
title: 使用 tabindex 指定多個元素的鍵盤焦點順序
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzRRcb'
forumTopicId: 301028
dashedName: use-tabindex-to-specify-the-order-of-keyboard-focus-for-several-elements
---

# --description--

`tabindex` 屬性也指定元素的確切切換順序。當該屬性的值設定為 1 或更高的正整數時，即可達成此目的。

設定 `tabindex="1"` 將會首先將鍵盤焦點帶到該元素。然後它會依序循環指定的 `tabindex` 值（2、3 等），接著才移動到預設值和 `tabindex="0"` 項目。

重要的是要注意，當以這種方式設定標籤順序時，它會覆寫預設順序（使用 HTML 原始碼）。這可能會讓期望從頁面頂部開始導覽的使用者感到困惑。在某些情況下，這種技術可能是必要的，但就存取性而言，請在應用前謹慎考量。

這裡有一個範例：

```html
<div tabindex="1">I get keyboard focus, and I get it first!</div>
```

```html
<div tabindex="2">I get keyboard focus, and I get it second!</div>
```

# --instructions--

Camper Cat 在他的 Inspirational Quotes 頁面上有一個搜尋欄，他計劃用 CSS 將它定位在右上角。他希望搜尋 `input` 和送出 `input` 控制元件成為分頁順序中的前兩個項目。為 `search` `input` 添加一個設定為 `1` 的 `tabindex` 屬性，並為 `submit` `input` 添加一個設定為 `2` 的 `tabindex` 屬性。

另一件需要注意的是，有些瀏覽器在點擊元素時可能會將你置於分頁鍵順序的中間。頁面中已新增一個元素，確保你將始終從分頁鍵順序的開始處開始。

# --hints--

你的程式碼應該為 `search` `input` 頁籤添加 `tabindex` 屬性。

```js
assert.isNotNull(document.querySelector('#search')?.getAttribute('tabindex'));
```

你的程式碼應該為 `submit` `input` 頁籤添加 `tabindex` 屬性。

```js
assert.isNotNull(document.querySelector('#submit')?.getAttribute('tabindex'));
```

你的程式碼應該將 `tabindex` 屬性設置在 `search` `input` 頁籤上，值為 1。

```js
assert.equal(document.querySelector('#search')?.getAttribute('tabindex'), '1');
```

你的程式碼應該將 `submit` `input` 頁籤的 `tabindex` 屬性設為 2。

```js
assert.equal(document.querySelector('#submit')?.getAttribute('tabindex'), '2');
```

# --seed--

## --seed-contents--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input type="search" name="search" id="search">
    <input type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```

# --solutions--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input tabindex="1" type="search" name="search" id="search">
    <input tabindex="2" type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```
