---
id: 587d7789367417b2b2512aa4
title: 使用 `audio` 元素提升音訊內容的可存取性
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVkcZ'
forumTopicId: 301014
dashedName: improve-accessibility-of-audio-content-with-the-audio-element
---

# --description--

HTML5 的 `audio` 元素在你的標記中包裹聲音或音訊資料流內容時，賦予語意意義。音訊內容也需要文字替代方案，以便讓聾啞或聽力受損的人能夠存取。這可以透過頁面上的鄰近文字或指向文字稿的連結來完成。

`audio` 頁籤支援 `controls` 屬性。這會顯示瀏覽器預設的播放、暫停及其他控制元件，並支援鍵盤功能。這是一個布林值屬性，表示不需要設定值，只要在頁籤中出現即啟用該設定。

這裡有一個範例：

```html
<audio id="meowClip" controls>
  <source src="audio/meow.mp3" type="audio/mpeg">
  <source src="audio/meow.ogg" type="audio/ogg">
</audio>
```

**注意：** 多媒體內容通常同時具有視覺和聽覺組件。它需要同步字幕和文字稿，以便有視覺和／或聽覺障礙的使用者能夠存取。一般來說，網頁開發人員不負責製作字幕或文字稿，但需要知道要將它們包含在內。

# --instructions--

是時候暫時離開 Camper Cat，來認識同為露營者的 Zersiax（@zersiax），他是無障礙的推動者，也是螢幕閱讀器的使用者。若要聽他的螢幕閱讀器實際運作的片段，請在 `p` 元素後為你的內容添加一個 `audio` 元素。包含 `controls` 屬性。然後在 `audio` 標籤內放置一個 `source` 元素，並將 `src` 屬性設為 `https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3`，`type` 屬性設為 `"audio/mpeg"`。

**注意：** 音訊片段可能聽起來很快且難以理解，但這是螢幕閱讀器使用者的正常速度。

# --hints--

你的程式碼應該只有一個 `audio` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('audio'),1);
```

你的 `audio` 元素應該有一個結束標籤。

```js
assert.match(code,/<audio.*>[\s\S]*<\/audio>/g);
assert.lengthOf(code.match(/<\/audio>/g),1);
```

`audio` 頁籤應該有 `controls` 屬性。

```js
assert.exists(document.querySelector('audio')?.getAttribute('controls'));
```

你的程式碼應該有一個 `source` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('source'), 1);
```

你的 `source` 頁籤應該在 `audio` 頁籤內。

```js
const audio = document.querySelector('audio');
const children = audio.querySelectorAll(`:scope ${'source'}`);
assert.lengthOf(children,1);
```

`source` 標記上 `src` 屬性的值應該與說明中的連結完全相符。

```js
assert.equal(
  document.querySelector('source')?.getAttribute('src'),
  'https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3'
);
```

你的程式碼應該在 `source` 頁籤上包含一個 `type` 屬性，其值為 audio/mpeg。

```js
assert.equal(document.querySelector('source')?.getAttribute('type'), 'audio/mpeg');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>



  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>
    <audio controls>
      <source src="https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3" type="audio/mpeg"/>
    </audio>
  </main>
</body>
```
