---
id: 587d778e367417b2b2512aab
title: 使用高對比度文字提升可讀性
challengeType: 0
videoUrl: 'https://scrimba.com/c/cKb3nCq'
forumTopicId: 301017
dashedName: improve-readability-with-high-contrast-text
---

# --description--

前景色與背景色之間的低對比度會使文字難以閱讀。足夠的對比度能提升你的內容可讀性，但「足夠」到底是什麼意思？

網頁內容存取指引（WCAG）建議一般文字的對比度至少為 4.5 比 1。此比率是透過比較兩種顏色的相對亮度值來計算。範圍從相同顏色、無對比的 1 比 1，到白色對黑色、最高對比的 21 比 1。線上有許多對比度檢查工具可幫你計算此比率。

# --instructions--

Camper Cat 最近部落格文章中選擇在白色背景上使用淺灰色文字，其對比度為 1.5:1，導致難以閱讀。將文字的 `color` 從目前的灰色（`#D3D3D3`）更改為較深的灰色（`#636363`），以將對比度提升至 6:1。

# --hints--

你的程式碼應該將 `body` 的文字 `color` 更改為較深的灰色。

```js
const body = document.querySelector('body');
const bodyColor = window.getComputedStyle(body).color; 
assert(bodyColor == 'rgb(99, 99, 99)');
```

你的程式碼不應該更改 `body` 的 `background-color`。

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body).backgroundColor; 
assert.equal(backgroundColor , 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: #D3D3D3;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: #636363;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
