---
id: 587d778f367417b2b2512aac
title: 避免色盲問題，請使用足夠的對比度
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMEUw'
forumTopicId: 301012
dashedName: avoid-colorblindness-issues-by-using-sufficient-contrast
---

# --description--

顏色是視覺設計的重要零件，但其使用會引入兩個存取性問題。首先，不應僅使用顏色作為傳達重要資訊的唯一方式，因為螢幕閱讀器使用者無法看到它。其次，前景色和背景色需要有足夠的對比度，讓色盲使用者能夠分辨它們。

先前的挑戰涵蓋了為第一個問題提供文字替代方案。最後一個挑戰介紹了對比度檢查工具，以協助解決第二個問題。WCAG 建議的 4.5:1 對比度比率適用於顏色使用以及灰階組合。

色盲使用者難以區分某些顏色－通常是在色相上，但有時也包括明度。你可能還記得對比度比率是使用前景和背景顏色的相對亮度（或明度）值計算出來的。

在實務上，4.5：1 的對比度可以透過將較深色調暗（加入黑色）和將較淺色調淡（加入白色）來達成。色輪上較深的色調被視為藍色、紫羅蘭色、洋紅色和紅色的色調，而較淺的淡色則是橘色、黃色、綠色和藍綠色。

# --instructions--

Camper Cat 正在嘗試為他的部落格文字和背景使用顏色，但他目前綠色調的 `background-color` 與栗色文字 `color` 的對比度為 2.5:1。你可以輕鬆調整顏色的亮度，因為他使用 CSS 的 `hsl()` 屬性（代表色相、飽和度、亮度）宣告這些顏色，只需更改第三個引數。將 `background-color` 的亮度值從 35% 提高到 55%，並將 `color` 的亮度值從 20% 降低到 15%。這樣可以將對比度提升到 5.9:1。

# --hints--

你的程式碼應該只更改文字 `color` 屬性的亮度值為 15%。

```js
assert.match(code ,/color:\s*?hsl\(0,\s*?55%,\s*?15%\)/gi);
```

你的程式碼應該只更改 `background-color` 屬性的亮度值為 55%。

```js
assert.match(code ,/background-color:\s*?hsl\(120,\s*?25%,\s*?55%\)/gi);
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 20%);
    background-color: hsl(120, 25%, 35%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 15%);
    background-color: hsl(120, 25%, 55%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
