---
id: 587d7787367417b2b2512aa1
title: 使用標題 Landmark 讓螢幕閱讀器導航更簡單
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB76vtv'
forumTopicId: 301023
dashedName: make-screen-reader-navigation-easier-with-the-header-landmark
---

# --description--

下一個增加語意意義並提升無障礙性的 HTML5 元素是 `header` 頁籤。它用來包裹其父標籤的介紹性資訊或導覽連結，並且非常適合用於多個頁面頂部重複出現的內容。

`header` 共享你在 `main` 中看到的內嵌地標特性，讓輔助技術能快速導覽至該內容。

**注意：** `header` 是用於你的 HTML 文件中 `body` 標記的。它不同於包含頁面標題、meta 資訊等的 `head` 元素。

# --instructions--

Camper Cat 正在撰寫一些關於忍者訓練的精彩文章，並且想為它們在他的网站上添加一個頁面。將目前包含 `h1` 的最上層 `div` 改成 `header` 標籤。

# --hints--

你的程式碼應該有一個 `header` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('header'),1);
```

你的 `header` 頁籤應該包裹住 `h1`。

```js
const header = document.querySelector('header');
const children = header?.querySelectorAll(`:scope ${'h1'}`);
assert.lengthOf(children , 1);
```

你的程式碼不應該有任何 `div` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

你的 `header` 元素應該有一個結束標籤。

```js
assert.isTrue(code.match(/<\/header>/g)?.length === code.match(/<header>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>

  <div>
    <h1>Training with Camper Cat</h1>
  </div>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>

  <header>
    <h1>Training with Camper Cat</h1>
  </header>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
