---
id: 587d778f367417b2b2512aae
title: 使用描述性聯結文字賦予聯結意義
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437DcV'
forumTopicId: 301013
dashedName: give-links-meaning-by-using-descriptive-link-text
---

# --description--

螢幕閱讀器使用者有各種選項，可選擇他們的裝置閱讀何種型別的內容。這些選項包括跳過（或略過）地標元素、跳轉到主要內容，或從標題取得頁面摘要。另一個選項是只聽取頁面上可用的聯結。

螢幕閱讀器透過朗讀聯結文字，或是位於錨點（`a`）標籤之間的內容來達成此目的。擁有一個「點擊這裡」或「閱讀更多」的聯結列表並沒有幫助。相反地，請在 `a` 標籤內使用簡短但具描述性的文字，以提供這些使用者更多意義。

# --instructions--

Camper Cat 使用的聯結文字在沒有上下文的情況下描述性不強。將錨點（`a`）標記移動，使其包覆 `information about batteries` 的文字，而不是 `Click here`。

# --hints--

你的程式碼應該將錨點 `a` 標記從包圍字詞 `Click here` 移動到包圍字詞 `information about batteries`。

```js
assert.match(document.querySelector('a')?.textContent, /^(information about batteries)$/g);
```

`a` 元素應該有一個 `href` 屬性，其值為空字串 `""`。

```js
assert.isEmpty(document.querySelector('a')?.getAttribute('href'));
```

`a` 元素應該有一個結束標籤。

```js
assert.isTrue(code.match(/<\/a>/g)?.length === code.match(/<a href=(''|"")>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. <a href="">Click here</a> for information about batteries</p>
  </article>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. Click here for <a href="">information about batteries</a></p>
  </article>
</body>
```
