---
id: 587d774d367417b2b2512a9e
title: 使用標題來顯示內容的階層關係
challengeType: 0
videoUrl: 'https://scrimba.com/c/cqVEktm'
forumTopicId: 301026
dashedName: use-headings-to-show-hierarchical-relationships-of-content
---

# --description--

標題（`h1` 到 `h6` 元素）是主要的標記，有助於為你的內容提供結構和標籤。螢幕閱讀器可以設定為只讀取頁面上的標題，讓使用者獲得摘要。這表示你的標記中的標題標籤必須具有語意意義並彼此相關，而不是僅僅根據大小值來選擇。

*語意意義* 是指你用於內容周圍的標記顯示它所包含的資訊型別。

如果你正在撰寫一篇包含引言、主體和結論的論文，將結論放在大綱中主體的子章節下是不太合理的。結論應該是獨立的章節。同樣地，網頁中的標題頁籤需要依序排列，並顯示你內容的階層關係。

標題具有相等（或更高）等級時，會開始新的隱含章節，標題具有較低等級時，會開始前一個標題的子章節。

舉例來說，一個頁面中有一個 `h2` 元素，接著是幾個以 `h4` 元素標記的子章節，會讓螢幕閱讀器使用者感到困惑。當有六個選項時，因為在瀏覽器中看起來比較好看，會很想使用標記，但你可以使用 CSS 來調整相對大小。

最後一點，每個頁面應該始終只有一個（且只能有一個）`h1` 元素，這是你內容的主要主題。這個以及其他標題部分用於搜尋引擎理解頁面的主題。

# --instructions--

Camper Cat 想在他網站上建立一個專門介紹成為忍者的頁面。幫助他修正標題，使他的標記能為內容賦予語意意義，並顯示各區段的正確父子關係。將所有的 `h5` 頁籤改成適當的標題層級，以顯示它們是 `h2` 頁籤的子區段。為此目的使用 `h3` 頁籤。

# --hints--

你的程式碼應該有 6 個 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll('h3') , 6);
```

你的程式碼應該有 6 個 `h3` 結束標籤。

```js
assert.lengthOf((code.match(/\/h3/g) || []) ,6);
```

你的程式碼中不應該有任何 `h5` 元素。

```js
assert.isEmpty(document.querySelectorAll('h5'));
```

你的程式碼不應該有任何 `h5` 結束標籤。

```js
assert.notMatch(code, /\/h5/);
```

# --seed--

## --seed-contents--

```html
<h1>How to Become a Ninja</h1>
<main>
  <h2>Learn the Art of Moving Stealthily</h2>
  <h5>How to Hide in Plain Sight</h5>
  <h5>How to Climb a Wall</h5>

  <h2>Learn the Art of Battle</h2>
  <h5>How to Strengthen your Body</h5>
  <h5>How to Fight like a Ninja</h5>

  <h2>Learn the Art of Living with Honor</h2>
  <h5>How to Breathe Properly</h5>
  <h5>How to Simplify your Life</h5>
</main>
```

# --solutions--

```html
<h1>How to Become a Ninja</h1>
<main>
  <h2>Learn the Art of Moving Stealthily</h2>
  <h3>How to Hide in Plain Sight</h3>
  <h3>How to Climb a Wall</h3>

  <h2>Learn the Art of Battle</h2>
  <h3>How to Strengthen your Body</h3>
  <h3>How to Fight like a Ninja</h3>

  <h2>Learn the Art of Living with Honor</h2>
  <h3>How to Breathe Properly</h3>
  <h3>How to Simplify your Life</h3>
</main>
```
