---
id: 587d778b367417b2b2512aa8
title: 為你的介面添加可存取的日期選擇器
challengeType: 0
videoUrl: 'https://scrimba.com/c/cR3bRbCV'
forumTopicId: 301008
dashedName: add-an-accessible-date-picker
---

# --description--

表單通常包含 `input` 欄位，可用於創建多種不同的表單控制元件。此元素上的 `type` 屬性指出將會創建哪種類型的 `input` 元素。

你可能已經注意到在之前的挑戰中有 `text` 和 `submit` 輸入型別，而 HTML5 介紹了一個可指定 `date` 欄位的選項。根據瀏覽器的支援情況，當 `input` 欄位獲得焦點時會顯示日期選擇器，這讓所有使用者填寫表單更為容易。

對於較舊的瀏覽器，型別將預設為 `text`，因此在 `label` 或 `placeholder` 文字中顯示用戶預期的日期格式會有所幫助，以防萬一。

這裡有一個範例：

```html
<label for="input1">Enter a date:</label>
<input type="date" id="input1" name="input1">
```

# --instructions--

Camper Cat 正在設置一個 Mortal Kombat 錦標賽，並想詢問他的競爭者哪一天最合適。為你的 `input` 頁籤添加一個 `type` 屬性為 `date`、`id` 屬性為 `pickdate` 和 `name` 屬性為 `date`。

# --hints--

你的程式碼應該為日期選擇欄添加一個 `input` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('input'), 2);
```

你的 `input` 頁籤應該有一個 `type` 屬性，值為 `date`。

```js
assert.equal(document.querySelector('input')?.getAttribute('type'), 'date');
```

你的 `input` 頁籤應該有一個 `id` 屬性，值為 `pickdate`。

```js
assert.equal(document.querySelector('input')?.getAttribute('id'),'pickdate');
```

你的 `input` 頁籤應該有一個 `name` 屬性，其值為 `date`。

```js
assert.equal(document.querySelector('input')?.getAttribute('name'), 'date');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>

        <!-- Only change code below this line -->



        <!-- Only change code above this line -->

        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>
        <input type="date" id="pickdate" name="date">
        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
