---
id: 587d778a367417b2b2512aa6
title: 使用標籤元素改善表單欄位的存取性
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMMAN'
forumTopicId: 301016
dashedName: improve-form-field-accessibility-with-the-label-element
---

# --description--

使用語意化的 HTML 標記來改善存取性，適用於同時使用適當的頁籤名稱和屬性。接下來的幾個挑戰涵蓋了在表單中使用屬性的一些重要情境。

`label` 頁籤包裹特定表單控制元件項目的文字，通常是選項的名稱或標籤。這會將意義連結到該項目，並使表單更易讀。`label` 頁籤上的 `for` 屬性明確地將該 `label` 與表單控制元件關聯，並由螢幕閱讀器使用。

你在基礎 HTML 章節的一課中學習了圓鈕和它們的標籤。在那一課中，我們將圓鈕輸入元素包裹在一個 `label` 元素內，並搭配標籤文字，以使文字可點擊。另一種達成此目的的方法是使用 `for` 屬性，如本課所說明。

`for` 屬性的值必須與表單控制元件的 `id` 屬性值相同。以下是一個範例：

```html
<form>
  <label for="name">Name:</label>
  <input type="text" id="name" name="name">
</form>
```

# --instructions--

Camper Cat 預期他的深思熟慮的部落格文章會引起大量關注，並且想要包含一個電子郵件註冊表單。為電子郵件 `label` 添加一個與其 `input` 欄位的 `id` 相符的 `for` 屬性。

# --hints--

你的程式碼中，`label` 頁籤應該有一個不為空的 `for` 屬性。

```js
assert.isNotEmpty(document.querySelector('label')?.getAttribute('for'));
```

你的 `for` 屬性值應該與電子郵件 `input` 上的 `id` 值相符。

```js
assert.equal(document.querySelector('label')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label>Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
