---
id: 587d774e367417b2b2512a9f
title: 使用 main 元素直接跳至內容
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7zuE'
forumTopicId: 301018
dashedName: jump-straight-to-the-content-using-the-main-element
---

# --description--

HTML5 介紹了幾個新的元素，讓開發者擁有更多的選項，同時也結合了存取功能。這些頁籤包括 `main`、`header`、`footer`、`nav`、`article` 和 `section` 等。

預設情況下，瀏覽器算繪這些元素類似於簡單的 `div`。然而，在適當的地方使用它們會為你的標記添加額外的意義。僅憑頁籤名稱就能顯示它所包含的資訊型別，這為該內容增加語意意義。輔助技術可以存取這些資訊，以提供更好的頁面摘要或導覽選項給它們的使用者。

`main` 元素用來包裹（你猜對了）主要內容，每個頁面應該只有一個。它用來圍繞與你頁面核心主題相關的資訊。不應該包含跨頁重複出現的條款，例如導覽聯結或橫幅。

`main` 頁籤也具有內嵌的地標特性，輔助技術可以使用該特性快速導覽到主要內容。如果你曾經在頁面頂端看到「跳至主要內容」的連結，使用 `main` 頁籤會自動為輔助裝置提供該功能。

# --instructions--

Camper Cat 對他的忍者武器頁籤有一些大構想。幫助他在 `header` 和 `footer`（涵蓋於其他挑戰中）之間為他的標記添加開啟和關閉的 `main` 頁籤。暫時保持 `main` 頁籤為空。

# --hints--

你的程式碼應該只有一個 `main` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

`main` 頁籤應該位於關閉的 `header` 頁籤與開啟的 `footer` 頁籤之間。

```js
assert.match(code,/<\/header>\s*?<main>\s*?<\/main>/gi);
```

# --seed--

## --seed-contents--

```html
<header>
  <h1>Weapons of the Ninja</h1>
</header>



<footer></footer>
```

# --solutions--

```html
<header>
  <h1>Weapons of the Ninja</h1>
</header>
<main>

</main>
<footer></footer>
```
