---
id: 66e45c8140f9fda5c105ae26
title: イベントフライヤーページを作成する
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. 本文内に`header`要素があること。
1. `header`要素にはイベントの画像があり、その中にイベントタイトルの`h1`が順番に含まれていること。必要であればこの画像を使えます：`https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. 本文内に`main`要素があること。
1. `main`要素には少なくとも2つの`section`要素があり、イベントの機能を紹介していること。
1. `section`要素それぞれに`h2`が含まれていること。
1. 本文には上下に`50px`のパディングがあること。
1. 本文には上下に`0`のマージンがあり、左右のマージンは中央に配置されていること。
1. 本文の幅はビューポートの幅に対して相対的に設定されていること。
1. 本文は`calc`関数を使って、本文の上下に適用されたすべてのパディングを差し引いたビューポートの高さの100％に`min-height`プロパティを設定していること。
1. フライヤー内に少なくとも1つの`hr`があること。
1. すべての`width`および`hr`要素の`section`は親要素に対するパーセント値で設定されていること。

**注意：** HTMLにスタイルシートをリンクし、CSSを適用してください。

# --hints--

`header`要素内に`body`要素があるはずです。

```js
assert.exists(document.querySelector('body > header'));
```

`img`の最初の子要素として`header`要素があるはずです。

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

`h1`の2番目の子要素として`header`要素があるはずです。

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

`h1`要素は1つだけにしてください。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`h1`にはイベントのタイトルがあるはずです。

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

`main`要素内に`body`要素があるはずです。

```js
assert.exists(document.querySelector('body > main'));
```

`section`要素内に少なくとも2つの`main`要素があるはずです。

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

`section`要素それぞれに`h2`が含まれているはずです。

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

`h2`要素は空であってはいけません。

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

`body`要素には上下に`50px`のパディングがあるはずです。

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

`body`要素には上下に`0`のマージンがあるはずです。

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

`body`要素には左右に`auto`のマージンがあるはずです。

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

`body`要素は`vw`を使って幅を設定しているはずです。

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

`body`は`calc`を使って`min-height`を`100vh - 100px`に設定しているはずです。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

少なくとも1つの`hr`要素があるはずです。

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

`hr`要素の幅はパーセント値で設定されているはずです。

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

`section`要素の幅はパーセント値で設定されているはずです。

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
