---
id: 646c5e727e56e743c9aed4a1
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

CSS 定位讓你設定元素在瀏覽器中的定位方式。它有一個 `position` 屬性，你可以將其設為 `static`、`absolute`、`relative`、`sticky` 或 `fixed`。

一旦你設定元素的 `position` 屬性，就可以透過為一個或多個 `top`、`right`、`left` 或 `bottom` 屬性設定像素或百分比值來移動該元素。

`static` 是所有元素的預設定位。如果你將它指定給某個元素，你將無法使用 `top`、`right`、`left` 或 `bottom` 移動該元素。

給 `.cat-head` 一個 `position` 屬性為 `static`，然後將 `top` 和 `left` 屬性分別設為 `100px`。


# --hints--

你的 `.cat-head` 選擇器應該有一個 `position` 屬性設定為 `static`。確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'static')
```

你的 `.cat-head` 選擇器應該有一個設定為 `100px` 的 `top` 屬性。請確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px')
```

你的 `.cat-head` 選擇器應該有一個 `left` 屬性設定為 `100px`。確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
