---
id: 646c5ace05e4be4211407935
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

要看到 `.cat-head` 元素，為它添加一個線性漸層背景，`#5e5e5e` 在 `85%`，`#45454f` 在 `100%`。

你可能不會注意到這兩種顏色之間的差異，但它們確實存在。

# --hints--

你的 `.cat-head` 選擇器應該有一個 `background` 屬性。

```js
assert.match(code, /background:/)
```

你的 `background` 屬性應該使用 `linear-gradient` 函式（程式）。

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

你的 `linear-gradient` 函式（程式）應該將第一個顏色設定為 `#5e5e5e`，位置在 `85%`。

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

你的 `linear-gradient` 函式（程式）應該將第二個顏色設為 `#45454f`，位置在 `100%`。

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
