---
id: 646cf48d8f8e1f535a1821d3
title: 步驟 27
challengeType: 0
dashedName: step-27
---

# --description--

將右耳移動到位置，並將 `position` 屬性設為 `absolute`，`top` 設為 `-26px`，`left` 設為 `163px`。

# --hints--

你的 `.cat-right-ear` 選擇器應該有一個設定為 `absolute` 的 `position` 屬性。別忘了添加分號。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.position, 'absolute')
```

你的 `.cat-right-ear` 選擇器應該有一個 `top` 屬性設定為 `-26px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.top, '-26px')
```

你的 `.cat-right-ear` 選擇器應該有一個設定為 `163px` 的 `left` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.left, '163px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  --fcc-editable-region--
  
  --fcc-editable-region--
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
