---
id: 646c48df8674cf2b91020ecc
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

你已經獲得一個基本的樣板。請先在你的 `head` 元素中添加一個 `link` 元素。對該 `link` 元素，將 `rel` 屬性設為 `stylesheet`，並將 `href` 設為 `./styles.css`。

# --hints--

你應該有一個 `link` 元素。

```js
assert.exists(document.querySelector("head>link"));
```

你的 `link` 元素應該設定 `rel` 和 `href` 屬性為 `stylesheet` 和 `./styles.css`。

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
