---
id: 646ce9d790d2a44de5f99e04
title: 第 19 步
challengeType: 0
dashedName: step-19
---

# --description--

創建兩個 `div` 元素，第一個在 `.cat-left-ear` 元素內，並帶有 `cat-left-inner-ear` 類別，第二個在 `.cat-right-ear` 元素內，並帶有 `cat-right-inner-ear` 類別。

# --hints--

你不應該更改具有 `cat-left-ear` 類別的現有 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div.cat-left-ear'), 1);
```

你不應該更改具有 `cat-right-ear` 類別的現有 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div.cat-right-ear'), 1);
```

你的 `.cat-left-ear` 元素內應該有一個 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('.cat-left-ear div'), 1);
```

你的 `.cat-right-ear` 元素內應該有一個 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('.cat-right-ear div'), 1);
```

`.cat-left-ear` 內的新 `div` 元素應該有 `cat-left-inner-ear` 類別。

```js
assert.isTrue(document.querySelectorAll('.cat-left-ear div')[0]?.classList.contains('cat-left-inner-ear'));
```

`.cat-right-ear` 內的新 `div` 元素應該有 `cat-right-inner-ear` 類別。

```js
assert.isTrue(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          --fcc-editable-region--
          <div class="cat-left-ear">

          </div>
          <div class="cat-right-ear">
            
          </div>
          --fcc-editable-region--
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
