---
id: 646c5d7057c45f432fcdd46c
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

使用類別選擇器，為 `.cat-head` 元素設定寬度為 `205px`、高度為 `180px`。同時，為它添加 `1px solid #000` 的邊框和 `46%` 的 `border-radius`。 

# --hints--

你應該有一個 `.cat-head` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

你的 `.cat-head` 選擇器應該設定 `width` 為 `205px`。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width, '205px');
```

你的 `.cat-head` 選擇器應該設定 `height` 為 `180px`。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height,'180px')
```

你的 `.cat-head` 選擇器應該設定 `border` 為 `1px solid #000`。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border, '1px solid rgb(0, 0, 0)')
```

你的 `.cat-head` 選擇器應該設定 `border-radius` 為 `46%`。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius, '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
