---
id: 646cf1206cac5f51804f49cf
title: 第 25 步
challengeType: 0
dashedName: step-25
---

# --description--

要正確定位左耳，你可以使用 CSS transform 將它旋轉一定角度。

`transform` 屬性允許你修改元素的形狀、位置和大小，而不改變版面配置或影響周圍的元素。它具有 `translate()`、`rotate()`、`scale()`、`skew()` 和 `matrix()` 等函式（程式）。

將 `transform` 屬性設定為 `rotate(-45deg)`，然後看看會發生什麼。

# --hints--

你應該將 `.cat-left-ear` 元素的 `transform` 屬性設為 `rotate(-45deg)`。別忘了添加分號。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.transform, 'rotate(-45deg)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  --fcc-editable-region--
  
  --fcc-editable-region--
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
