---
id: 646ddd3f9f97a0667b964bdb
title: 步驟 40
challengeType: 0
dashedName: step-40
---

# --description--

創建兩個 `div` 元素，一個帶有類別 `cat-left-inner-eye` 並置於 `.cat-left-eye` 元素內，另一個帶有類別 `cat-right-inner-eye` 並置於 `.cat-right-eye` 元素內。

# --hints--

你不應該更改具有 `cat-left-eye` 類別的現有 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div.cat-left-eye'), 1);
```

你不應該更改具有 `cat-right-eye` 類別的現有 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div.cat-right-eye'), 1);
```

你應該在你的 `.cat-left-eye` 元素內有一個 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('.cat-left-eye div'), 1);
```

你應該在你的 `.cat-right-eye` 元素內有一個 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('.cat-right-eye div'), 1);
```

`.cat-left-eye` 內的新 `div` 元素應該有 `cat-left-inner-eye` 類別。

```js
assert.isTrue(document.querySelectorAll('.cat-left-eye div')[0]?.classList.contains('cat-left-inner-eye'));
```

`.cat-right-eye` 內的新 `div` 元素應該有 `cat-right-inner-eye` 類別。

```js
assert.isTrue(document.querySelectorAll('.cat-right-eye div')[0]?.classList.contains('cat-right-inner-eye'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>

        <div class="cat-eyes">
          --fcc-editable-region--
          <div class="cat-left-eye">
            
          </div>
          <div class="cat-right-eye">
            
          </div>
          --fcc-editable-region--
        </div>        
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-left-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}

.cat-right-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}
```
