---
id: 646dd556d524bc61c0139bd6
title: 步驟 33
challengeType: 0
dashedName: step-33
---

# --description--

大多數貓的耳朵和同一隻耳朵的內部零件顏色不同。你也可以實作相同的效果。這就是為什麼你之前為右耳和左耳的內部都定義了一個 `div` 元素。

使用類別選擇器，為你的 `.cat-left-inner-ear` 元素的左右邊框各設定 `20px solid transparent`。同時為它設定一個底部邊框為 `40px solid #3b3b4f`。

# --hints--

你應該有一個 `.cat-left-inner-ear` 選擇器。

```js 
assert.exists(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear'))
```

你的 `.cat-left-inner-ear` 選擇器應該有一個設定為 `20px solid transparent` 的 `border-left` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderLeft, '20px solid transparent')
```

你的 `.cat-left-inner-ear` 選擇器應該有一個 `border-right` 屬性設定為 `20px solid transparent`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderRight, '20px solid transparent')
```

你的 `.cat-left-inner-ear` 選擇器應該有一個設定為 `40px solid #3b3b4f` 的 `border-bottom` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderBottom, '40px solid rgb(59, 59, 79)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

--fcc-editable-region--

--fcc-editable-region--
```
