---
id: 6476f5c17f99146071ee884c
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

你會看到什麼都沒有發生。即使設定了 `top` 和 `left` 各為 `100px`，`.cat-head` 元素也沒有移動。但使用 `relative` 定位就不是這種情況。

當你使用 `relative` 值時，元素仍然根據文件的正常流程定位，但 `top`、`left`、`bottom` 和 `right` 值變為作用中。

將 `.cat-head` 的位置設為 `relative`，而非 `static`，並保持 `top` 和 `left` 屬性不變。

# --hints--

你的 `.cat-head` 選擇器應該有一個設定為 `relative` 的 `position` 屬性。確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'relative');
```

你的 `.cat-head` 選擇器應該有一個設定為 `100px` 的 `top` 屬性。請確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

你的 `.cat-head` 選擇器應該有一個 `left` 屬性設定為 `100px`。確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
