---
id: 6477062778c85972eb648030
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

你現在應該將貓頭置中。 

為 `.cat-head` 元素賦予 `position` 屬性並設為 `absolute`。將 `right`、`left`、`top`、`bottom` 屬性的值設為 `0`，然後將其所有邊的 `margin` 屬性設為 `auto`。這是一種使用 CSS 定位將元素垂直且水平置中的方法。

# --hints--

你的 `.cat-head` 選擇器應該有一個設定為 `absolute` 的 `position` 屬性。請確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

你的 `.cat-head` 選擇器應該有一個設定為 `0` 的 `top` 屬性。請確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

你的 `.cat-head` 選擇器應該有一個設定為 `0` 的 `left` 屬性。請確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '0px');
```

你的 `.cat-head` 選擇器應該有一個設定為 `0` 的 `right` 屬性。請確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right, '0px');
```

你的 `.cat-head` 選擇器應該有一個設定為 `0` 的 `bottom` 屬性。請確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom, '0px');
```

你的 `.cat-head` 選擇器應該有一個設定為 `auto` 的 `margin` 屬性。請確保你添加分號。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
