---
id: 647d855ac12cd436059acd39
title: 步驟2
challengeType: 0
dashedName: step-2
---

# --description--

使用通用選擇器在 CSS 中添加 `box-sizing: border-box;`。 這可確保元素在其指定的寬度和高度內包含內邊距和邊框。

# --hints--

你應該有通用選擇器（`*`）。

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

你的通用選擇器應該有一個 `box-sizing` 屬性，設置爲 `border-box`。

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
