---
id: 683ecb010f4c6bafd428802c
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, validarás un conjunto de datos médicos para asegurarte de que cumple con un conjunto de reglas.

Declara una variable llamada `medical_records` y asígnale una lista vacía. En los siguientes pasos, la vas a usar para almacenar tus datos médicos.

# --hints--

Debes declarar una variable llamada `medical_records`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

Debes asignar una lista vacía a tu variable `medical_records`.

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

