---
id: 6842c3eee327042f1da17b89
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Crea una estructura `if`. Para su condición, usa el operador `not` para negar `is_sequence`. Dentro de la estructura `if`, imprime `Invalid format: expected a list or tuple.` y retorna `False`.

# --hints--

Debes crear una estructura `if` dentro de tu función `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_ifs()[0]
`) })
```

Tu sentencia `if` debe tener `not is_sequence` como condición.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_ifs()[0].find_conditions()[0].is_equivalent("not is_sequence")
`) })
```

Debes imprimir `Invalid format: expected a list or tuple.` dentro de tu sentencia `if`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_ifs()[0].find_bodies()[0].has_call("print('Invalid format: expected a list or tuple.')")
`) })
```

Debes devolver `False` al final de tu estructura `if`.

```js
({ test: () => runPython(`
_if_body = _Node(_code).find_function("validate").find_ifs()[0].find_bodies()[0]
_p_statement = "print('Invalid format: expected a list or tuple.')"

if not _if_body.has_call(_p_statement):
  assert _if_body[-1].is_equivalent("return False")
else:
  assert _if_body.is_ordered(_p_statement, "return False")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]
--fcc-editable-region--
def validate(data):
    is_sequence = isinstance(data, (list, tuple))
    
--fcc-editable-region--
```
