---
id: 6842d5b537c64266420faf62
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Agrega una clave `age` con el valor del entero `34` y una clave `gender` con el valor de la cadena `Female` a tu diccionario. No olvides la coma entre los pares clave-valor.

# --hints--

Tu lista `medical_records` debe contener un solo diccionario.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

Debes tener una clave `age` con el valor del entero `34` en tu diccionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

Debes tener una clave `gender` con el valor de la cadena `Female` en tu diccionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

Debes tener una clave `patient_id` con el valor de la cadena `P1001` en tu diccionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
