---
id: 6842ba0086201089ea268c92
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Siguiendo la misma estructura que usaste en los pasos anteriores, la lista `medical_records` ha sido completada para ti con los datos de otros pacientes. Siéntete libre de echarle un vistazo.

A continuación comenzarás a escribir la función para validar el conjunto de datos. Crea una función llamada `validate` con un solo parámetro `data`.

Quieres asegurarte de que tus datos sean una lista o una tupla. Por lo tanto, dentro de la función `validate`, declara una variable llamada `is_sequence` y asígnale una llamada a `isinstance`. Pasa `data` como el primer argumento y una tupla que contiene `list` y `tuple` como el segundo argumento.

# --hints--

Debes tener una función llamada `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

Tu función `validate` debe tener un solo parámetro `data`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

Debes tener una variable llamada `is_sequence` dentro de tu función `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

Debes asignar `isinstance(data, (list, tuple))` a tu variable `is_sequence`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]
--fcc-editable-region--

--fcc-editable-region--
```
