---
id: 6842d659a4ae7973dabab3b0
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Completa el diccionario agregando los siguientes tres pares clave-valor:

- Una clave `diagnosis` con el valor de `Hypertension`
- Una clave `medications` con el valor de la lista `['Lisinopril']`
- Una clave `last_visit_id` con el valor de `V2301`.

# --hints--

Debes tener una clave `diagnosis` con el valor de la cadena `Hypertension` en tu diccionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

Debes tener una clave `medications` con el valor de la lista `['Lisinopril']` en tu diccionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

Debes tener una clave `last_visit_id` con el valor de la cadena `V2301` en tu diccionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

Debes tener una clave `patient_id` con el valor de la cadena `P1001` en tu diccionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

Debes tener una clave `age` con el valor del entero `34` en tu diccionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

Debes tener una clave `gender` con el valor de la cadena `Female` en tu diccionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
