---
id: 6842aee5bcaf4ab8883f8191
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Como aprendiste en lecciones anteriores, un diccionario es una estructura de datos que contiene pares clave-valor, donde las claves son tipos de datos inmutables (como cadenas) y únicas dentro de un diccionario:

```py
person = {
    'name': 'John',
    'age': 33
}
```

La lista `medical_records` almacenará diccionarios, cada uno representando a un paciente. Agrega un diccionario con una clave `patient_id` y un valor de la cadena `P1001` a la lista `medical_records`.

# --hints--

Tu lista `medical_records` debe contener un solo diccionario.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

Debes tener una clave `patient_id` con el valor de la cadena `P1001` en tu diccionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
medical_records = [
    
]
--fcc-editable-region--
```
