---
id: 6846ac424529900e7327d326
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

Como aprendiste en una lección anterior, un conjunto es una colección desordenada de elementos únicos:

```py
integers = set([3, 5, 1, 2, 1, 3, 4])
print(integers) # {1, 2, 3, 4, 5}
```

Vas a usar un conjunto para asegurarte de que cada diccionario no contenga claves adicionales o mal escritas.

Usa la estructura `set()` para crear un conjunto a partir de la siguiente lista de claves que cada diccionario debe tener: `['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']`. Asigna el conjunto a una variable llamada `key_set`.

# --hints--

Debes tener una variable llamada `key_set`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("key_set")
`) })
```

Debes asignar `set(['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id'])` a tu variable `key_set`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_variable("key_set").is_equivalent("key_set = set(['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id'])")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }   
]


def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
--fcc-editable-region--
    
--fcc-editable-region--
    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)

```
