---
id: 6846a1f2805ac73ce671a5d8
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Como aprendiste en una lección anterior, la función `enumerate` permite llevar un registro del índice mientras recorres un iterable:

```py
person = {'name': 'John', 'age': 33}

for index, item in enumerate(person):
    print(index, item)

# 0 name
# 1 age
```

Crea un ciclo `for` que itere sobre `data`. Usa la función `enumerate` para obtener tanto el índice como el elemento en `data` en cada iteración. Usa `index` y `dictionary` como variables de iteración.

Por ahora usa `pass` para llenar el cuerpo del ciclo.

# --hints--

Debes tener un ciclo `for` en tu función `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0]
`) })
```

Tu ciclo `for` debe iterar sobre `enumerate(data)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_for_iter().is_equivalent("enumerate(data)")
`) })
```

Debes usar `index` y `dictionary` como las variables de iteración en tu ciclo `for`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_for_vars().is_equivalent("index, dictionary")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

--fcc-editable-region--
    
--fcc-editable-region--
```
