---
id: 65577739f57ecca6c39bb4e9
title: Schritt 33
challengeType: 20
dashedName: step-33
---

# --description--

Der `list()`-Typ-Konstruktor ermöglicht es Ihnen, eine Liste aus einem iterierbaren Objekt zu erstellen.

Ändern Sie die Zuweisung Ihrer `unvisited`-Variablen so, dass `list()` verwendet wird, und übergeben Sie `graph` als iterierbares Objekt.

# --hints--

Sie sollten `list()` verwenden, um aus dem `graph`-Wörterbuch eine Liste zu erzeugen.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/list\s*\(\s*graph\s*\)/));
  }
})
```

Sie sollten `list(graph)` `unvisited` zuweisen.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\s*=\s*list\s*\(\s*graph\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
