---
id: 6557709b0aee699a6a00528c
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Zum Beispiel kann ein Graph verwendet werden, um zwei Punkte im Raum, `A` und `B`, darzustellen, die durch einen Pfad verbunden sind. Ein solcher Graph besteht aus zwei Knoten, die durch eine Kante verbunden sind.

Ersetzen Sie den vorhandenen `'species'`-Schlüssel durch die Strings `'A'`. Ersetzen Sie dann den entsprechenden Wert durch den String `'B'`, um die Verbindung zwischen den Knoten `'A'` und `'B'` darzustellen.

# --hints--

Ihr Wörterbuch sollte einen `'A'`-Schlüssel enthalten.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Ihr `'A'`-Schlüssel sollte `'B'` als Wert haben.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Ihr Wörterbuch sollte zwei Schlüssel enthalten.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
