---
id: 65577a17564ce8a8e06c1460
title: Schritt 35
challengeType: 20
dashedName: step-35
---

# --description--

Dictionary-Comprehensions unterstützen auch die bedingte `if`/`else`-Syntax:

```py
{key: val_1 if condition else val_2 for key in dict}
```

Im obigen Beispiel ist `dict` das bestehende Dictionary. Wenn `condition` zu `True` ausgewertet wird, erhält `key` den Wert `val_1`, andernfalls `val_2`.

Verwenden Sie eine Dictionary-Comprehension, um ein Dictionary basierend auf `graph` zu erstellen, und weisen Sie es der Variablen `distances` zu. Geben Sie dem Schlüssel den Wert null, wenn der Knoten dem Startknoten entspricht, und unendlich sonst. Verwenden Sie `float('inf')`, um Letzteres zu erreichen.

# --hints--

Sie sollten die Syntax der Dictionary-Comprehension verwenden, um Ihrer Variablen `distances` einen Wert zuzuweisen.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}distances\s*=\s*\{\s*(\w+)\s*:\s*0\s+if\s+(?:\1\s*==\s*start|start\s*==\s*\1)\s+else\s+float\s*\(\s*("|')inf\2\s*\)\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    paths = {node: [] for node in graph}
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
