---
id: 655773b0591c5f9f4045883e
title: Schritt 24
challengeType: 20
dashedName: step-24
---

# --description--

Der Algorithmus startet an einem angegebenen Knoten. Dann durchsucht er den Graphen, um den kürzesten Weg zwischen dem Startknoten, also der *source*, und allen anderen Knoten zu finden.

Dafür benötigt Ihre Funktion zwei Parameter: `graph` und `start`. Fügen Sie diese Ihrer Funktionsdeklaration hinzu.

# --hints--

Ihre Funktion sollte `graph` und `start` als Parameter in dieser Reihenfolge annehmen.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
