---
id: 655774955b097ea14897db12
title: Schritt 28
challengeType: 20
dashedName: step-28
---

# --description--

Die Entfernung vom Startknoten ist null, da der Algorithmus seine Bewertung genau dort beginnt.

Nachdem Sie `node` an `unvisited` in Ihrer Schleife angehängt haben, erstellen Sie eine `if`-Anweisung, die ausgelöst wird, wenn der Knoten gleich dem Startknoten ist. Weisen Sie dann `0` diesem Knoten innerhalb des `distances`-Wörterbuchs zu.

# --hints--

Sie sollten eine `if`-Anweisung erstellen, die ausgeführt wird, wenn `node` gleich `start` ist.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

Innerhalb Ihrer neuen `if`-Anweisung sollten Sie `0` dem Knoten im `distances`-Wörterbuch zuweisen.

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
