---
id: 65577236b056379d5dbc7000
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Jetzt beginnen Sie damit, den Algorithmus zu entwickeln, der den kürzesten Weg zwischen jedem Knoten in Ihrem neuen Graphen berechnet.

Deklarieren Sie eine leere Funktion namens `shortest_path`. Verwenden Sie das Schlüsselwort `pass`, um den Funktionskörper zu füllen.

# --hints--

Sie sollten eine Funktion mit dem Namen `shortest_path` definieren.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
