---
id: 655775221059f5a20493d5d7
title: Schritt 30
challengeType: 20
dashedName: step-30
---

# --description--

Fügen Sie nach Ihrer `for`-Schleife einen `print()`-Aufruf hinzu und übergeben Sie den folgenden String, um die Werte der von Ihnen erstellten Variablen zu sehen: `f'Unvisited: {unvisited}\nDistances: {distances}'`.

# --hints--

Sie sollten `f'Unvisited: {unvisited}\nDistances: {distances}'` nach Ihrer `for`-Schleife ausgeben.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(^\s*)for.*:.*^\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
--fcc-editable-region--
```
