---
id: 6557746aad2844a0cd864e12
title: Schritt 27
challengeType: 20
dashedName: step-27
---

# --description--

Während der Algorithmus den Graphen durchsucht, sollte er die aktuell bekannte kürzeste Entfernung zwischen dem Startknoten und den anderen Knoten im Blick behalten.

Erstellen Sie vor Ihrer `for`-Schleife eine neue Variable namens `distances` und weisen Sie ihr ein leeres Wörterbuch zu.

# --hints--

Sie sollten eine Variable namens `distances` haben.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=.*(?=^\1for.*:)/ms));
  }
})
```

Ihre `distances`-Variable sollte ein leeres Wörterbuch sein.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=\s*\{\s*\}.*(?=^\1for.*:)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
