---
id: 65578d0f6c78a0b868a43b9c
title: Schritt 40
challengeType: 20
dashedName: step-40
---

# --description--

Innerhalb der `while`-Schleife ist der erste Schritt, den aktuellen Knoten zu definieren, den Sie besuchen möchten. Dafür können Sie die `min()`-Funktion verwenden. Sie gibt das kleinste Element aus dem als Argument übergebenen iterierbaren Objekt zurück.

Löschen Sie `pass` und erstellen Sie dann eine Variable namens `current`, der Sie `min(unvisited)` zuweisen.

# --hints--

Sie sollten eine `current`-Variable in Ihrer `while`-Schleife erstellen.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

Sie sollten `min(unvisited)` Ihrer `current`-Variable zuweisen. Denken Sie daran, `pass` zu löschen.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
