---
id: 655773f8b8b5db9fc6d0ae76
title: Schritt 25
challengeType: 20
dashedName: step-25
---

# --description--

Um die besuchten Knoten zu verfolgen, benötigen Sie eine Liste aller Knoten im Graphen. Sobald ein Knoten besucht wurde, wird er aus dieser Liste entfernt.

Ersetzen Sie nun das Schlüsselwort `pass` durch eine Variable namens `unvisited` und weisen Sie ihr eine leere Liste zu.

# --hints--

Sie sollten eine Variable namens `unvisited` innerhalb der `shortest_path`-Funktion haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

Sie sollten Ihrer `unvisited`-Variable eine leere Liste zuweisen. Denken Sie daran, `pass` zu löschen.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
