---
id: 65578cb031cd93b77a285db2
title: Schritt 38
challengeType: 20
dashedName: step-38
---

# --description--

Ihre Funktion wird alle Knoten erkunden, die mit dem Startknoten verbunden sind. Sie berechnet die kürzesten Pfade für alle diese Knoten. Anschließend entfernt sie den Startknoten aus den nicht besuchten Knoten.

Danach wird der nächstgelegene Nachbarknoten besucht, und der Vorgang wird wiederholt, bis alle Knoten besucht wurden.

Ab jetzt arbeiten Sie an der Hauptschleife, die die Knoten im Graphen erkundet. Um Probleme mit einer unendlichen Schleife während der Entwicklung des Algorithmus zu vermeiden, kommentieren Sie den Funktionsaufruf aus.

# --hints--

Sie sollten den Funktionsaufruf in einen Kommentar umwandeln.

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
