---
id: 655771d889132f9ccd341060
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Ein Graph wird als *gewichteter* Graph bezeichnet, wenn seine Kanten mit Gewichten versehen sind, die eine Entfernung, Zeit oder einen anderen quantitativen Wert darstellen.

In Ihrem Fall sind diese Gewichte die Entfernungen zwischen jedem Knoten oder Punkt im Raum. Um einen gewichteten Graphen darzustellen, können Sie Ihr Wörterbuch so ändern, dass für jeden Wert eine Liste von Tuples verwendet wird.

Das erste Element im Tuple ist der verbundene Knoten, und das zweite Element ist eine ganze Zahl, die die Entfernung angibt.

Ändern Sie `my_graph['A']` in eine Liste von Tuples, wobei die Entfernung von `A-B` `3` ist und die Entfernung von `A-D` `1` ist.

# --hints--

`my_graph["A"]` sollte eine Liste sein, die die Tuples `('B', 3)` und `('D', 1)` enthält.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph` sollte 4 Schlüssel mit den Namen `'A'`, `'B'`, `'C'` und `'D'` haben.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
