---
id: 65578d4fc3afc3b8f554c882
title: Schritt 41
challengeType: 20
dashedName: step-41
---

# --description--

`min()` nimmt auch ein nur-Keyword-Argument entgegen. Wenn Sie eine Funktion als zusätzliches Argument an `min()` übergeben, können Sie die Art und Weise ändern, wie die Listenelemente verglichen werden.

Das Ergebnis der Zeile, die Sie im vorherigen Schritt geschrieben haben, ist der Knoten, der alphabetisch zuerst kommt. Stattdessen möchten Sie den unbesuchten Knoten auswählen, der den kleinsten Abstand vom Startknoten hat.

Übergeben Sie `key=distances.get` als zweites Argument an Ihren `min()`-Aufruf. So erfolgt der Vergleich abhängig vom Wert, den jedes `unvisited`-Listenelement im `distances`-Wörterbuch hat.

# --hints--

Sie sollten `key=distances.get` als zweites Argument an Ihren `min()`-Aufruf übergeben.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
