---
id: 65578fcf00322dbad5dee05b
title: Schritt 43
challengeType: 20
dashedName: step-43
---

# --description--

Erstellen Sie eine `if`-Anweisung, um zu prüfen, ob die Distanz des Nachbarknotens (das zweite Element im verarbeiteten Tuple) plus die Distanz von `current` kleiner ist als die aktuell bekannte Distanz des Nachbarknotens (das erste Element im verarbeiteten Tuple).

Verwenden Sie das Schlüsselwort `pass`, um den Körper der `if` vorübergehend zu füllen.

# --hints--

Sie sollten eine `if`-Anweisung haben, um zu prüfen, ob `distance + distances[current]` kleiner ist als `distances[node]`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)for\s+\w+\s*,\s*\w+\s+in\s+graph\s*\[\s*current\s*\]\s*:\s*^\1(\s{4})if\s+distance\s*\+\s*distances\s*\[\s*current\s*\]\s*<\s*distances\s*\[\s*node\s*\]\s*:/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
--fcc-editable-region--
        for node, distance in graph[current]:
            pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
