---
id: 65578f895f2a65ba7a916804
title: Schritt 42
challengeType: 20
dashedName: step-42
---

# --description--

Nach der `current` Variablenzuweisung erstellen Sie eine `for` Schleife, um die Tuple in der `graph[current]` Liste zu durchlaufen. Dafür benötigen Sie zwei iterierende Variablen. Denken Sie daran, `pass` zu verwenden, um den Schleifenkörper zu füllen.

# --hints--

Sie sollten eine `for` Schleife erstellen, um die Tuple-Elemente in der `graph[current]` Liste zu durchlaufen. Verwenden Sie zwei iterierende Variablen und vergessen Sie nicht, das Schlüsselwort `pass` hinzuzufügen.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)while\s+unvisited\s*:.*^\1\1for\s+\w+\s*,\s*\w+\s+in\s+graph\s*\[\s*current\s*\]\s*:\s*^\1\1\1pass/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited, key=distances.get)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
