---
id: 65577791ad8c26a7705e2919
title: Schritt 34
challengeType: 20
dashedName: step-34
---

# --description--

Mit einer Dictionary-Comprehension können Sie ein Dictionary aus einem bestehenden Dictionary erstellen:

```py
{key: val for key in dict}
```

Im obigen Beispiel ist `val` der Wert, den `key` im neuen Dictionary haben wird, und `dict` ist das bestehende Dictionary.

Sie möchten die Pfade zwischen dem Startknoten und jedem anderen Knoten verfolgen.

Erstellen Sie nach der `distances`-Variablen eine `paths`-Variable und weisen Sie ihr ein Dictionary mit allen Schlüsseln von `graph` zu. Weisen Sie jedem Schlüssel eine leere Liste zu und verwenden Sie eine Dictionary-Comprehension, um Ihr Dictionary zu erstellen.

# --hints--

Sie sollten eine Variable namens `paths` haben.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

Ihre `paths`-Variable sollte die Syntax der Dictionary-Comprehension verwenden, um jedem Knoten im graph eine leere Liste zuzuweisen.

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
