---
id: 6557712d77ce2d9bd7e63afd
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Fügen Sie Ihrem Graphen einen weiteren Knoten hinzu, der mit `B` verbunden ist, und nennen Sie ihn `C`.

Ändern Sie Ihr bestehendes Dictionary so, dass es diese Anordnung darstellt: Fügen Sie einen weiteren Schlüssel `'C'` zu `my_graph` hinzu und geben Sie ihm den Wert des Strings `'B'`.

Ändern Sie außerdem den Wert des bestehenden Schlüssels `'B'` in die Liste `['A', 'C']`, um die mehreren Verbindungen Ihres `'B'`-Knotens darzustellen.

# --hints--

Ihr Dictionary sollte 3 Schlüssel haben — `'A'`, `'B'` und `'C'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

Der Wert von `my_graph['A']` sollte der String `'B'` sein.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` sollte eine Liste sein.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

Der Wert von `my_graph['B']` sollte eine Liste sein, die `'A'` und `'C'` enthält.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

Der Wert von `my_graph['C']` sollte der String `'B'` sein.

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
