---
id: 65576ff7888f9e96f52a4be1
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Dictionaries speichern Daten in Form von *Schlüssel*-*Wert*-Paaren. Ein Schlüssel wird durch einen Doppelpunkt vom entsprechenden Wert getrennt. Und jedes Schlüssel-Wert-Paar wird durch ein Komma vom nächsten Paar getrennt:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

Fügen Sie Ihrem Dictionary ein neues Schlüssel-Wert-Paar hinzu. Verwenden Sie den String `'species'` als Schlüssel und den String `'guinea pig'` als Wert.

# --hints--

Sie sollten einen `'species'`-Schlüssel mit dem Wert `'guinea pig'` in Ihrem `copper`-Dictionary haben.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

Ihr `copper`-Dictionary sollte genau ein Schlüssel-Wert-Paar enthalten.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
