---
id: 6557716aadbd2d9c42c0e69a
title: Schritt 19
challengeType: 20
dashedName: step-19
---

# --description--

Fügen Sie einen letzten Knoten `'D'` hinzu, der mit `'A'` und `'C'` verbunden ist.

Ändern Sie Ihr Wörterbuch so, dass diese Struktur dargestellt wird. Verwenden Sie erneut eine Liste, um mehrere Verbindungen darzustellen.

# --hints--

Ihr Wörterbuch sollte 4 Schlüssel mit den Namen `'A'`, `'B'`, `'C'` und `'D'` enthalten.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` sollte eine Liste sein.

```js
({ test: () => assert(runPython(`
    type(my_graph["A"]) is list
  `))
})
```

`my_graph['A']` sollte eine Liste sein, die `'B'` und `'D'` enthält.

```js
({ test: () => assert(runPython(`
    len(my_graph["A"]) == 2 and "B" in my_graph["A"] and "D" in my_graph["A"]
  `))
})
```

`my_graph['B']` sollte eine Liste sein.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']` sollte eine Liste sein, die `'A'` und `'C'` enthält.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']` sollte eine Liste sein.

```js
({ test: () => assert(runPython(`
    type(my_graph["C"]) is list
  `))
})
```

`my_graph['C']` sollte eine Liste sein, die `'B'` und `'D'` enthält.

```js
({ test: () => assert(runPython(`
    len(my_graph["C"]) == 2 and "B" in my_graph["C"] and "D" in my_graph["C"]
  `))
})
```

`my_graph['D']` sollte eine Liste sein.

```js
({ test: () => assert(runPython(`
    type(my_graph["D"]) is list
  `))
})
```

`my_graph['D']` sollte eine Liste sein, die `'A'` und `'C'` enthält.

```js
({ test: () => assert(runPython(`
    len(my_graph["D"]) == 2 and "A" in my_graph["D"] and "C" in my_graph["D"]
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': ['A', 'C'],
    'C': 'B'
}
--fcc-editable-region--
```
