---
id: 655774d01daeeaa1978b99d5
title: Schritt 29
challengeType: 20
dashedName: step-29
---

# --description--

Am Anfang werden alle anderen Knoten im Graphen als unendlich weit vom Startknoten entfernt betrachtet, da die Entfernung noch nicht bestimmt wurde.

Erstellen Sie eine `else`-Klausel und weisen Sie dem Knoten im `distances`-Wörterbuch einen unendlichen Wert zu. Verwenden Sie dazu die `float()`-Funktion mit dem String `'inf'` als Argument, um eine Fließkommazahl zu erzeugen, die die positive Unendlichkeit darstellt.

# --hints--

Sie sollten eine `else`-Klausel haben.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

Sie sollten `float('inf')` innerhalb Ihrer neuen `distances[node]`-Klausel `else` zuweisen.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
