---
id: 655777060d8ddea6741be1b1
title: Schritt 32
challengeType: 20
dashedName: step-32
---

# --description--

Alle Distanzen in `distances` sind auf unendlich gesetzt, außer für den Startknoten. Die `unvisited`-Liste enthält alle Knoten in Ihrem Graphen. Aber eigentlich brauchen Sie diese `for`-Schleife nicht, um dieses Ergebnis zu erzielen.

Entfernen Sie Ihre `for`-Schleife mit ihrem gesamten Inhalt.

# --hints--

Sie sollten Ihre `for`-Schleife und den gesamten verschachtelten Code entfernen.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(^\s*)distances\s*=\s*\{\s*\}\s*\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
