---
id: 65661b72d6745ebec6a96923
title: Schritt 22
challengeType: 20
dashedName: step-22
---

# --description--

Ändern Sie auf die gleiche Weise die beiden übrigen Listen, wobei Sie berücksichtigen, dass die `C-D` Entfernung `7` ist.

# --hints--

`my_graph['C']` sollte eine Liste enthalten, die die Tuples `('B', 4)` und `('D', 7)` enthält.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 4), ("D", 7)]
    len(my_graph["C"]) == 2 and all(t in my_graph["C"] for t in tuples) 
  `))
})
```

`my_graph['D']` sollte eine Liste enthalten, die die Tuples `('A', 1)` und `('C', 7)` enthält.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 1), ("C", 7)]
    len(my_graph["D"]) == 2 and all(t in my_graph["D"] for t in tuples)
  `))
})
```

`my_graph` sollte 4 Schlüssel mit den Namen `'A'`, `'B'`, `'C'` und `'D'` haben.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` sollte eine Liste enthalten, die die Tuples `('B', 3)` und `('D', 1)` enthält.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph['B']` sollte eine Liste enthalten, die die Tuples `('A', 3)` und `('C', 4)` enthält.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
