---
id: 6566195b0a021bb660b2b4b1
title: Schritt 21
challengeType: 20
dashedName: step-21
---

# --description--

Ändern Sie nun `my_graph['B']` in eine Liste von Tuples, wobei das erste Element im Tuple der verbundene Knoten ist und das zweite Element die Entfernung. Die `B-C`-Entfernung ist `4`.

# --hints--

`my_graph['B']` sollte eine Liste sein, die die Tuples `('A', 3)` und `('C', 4)` enthält.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

`my_graph` sollte 4 Schlüssel mit den Namen `'A'`, `'B'`, `'C'` und `'D'` haben.

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
