---
id: 6578b13757611e2825beb8a5
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Schlüssel müssen innerhalb eines Dictionaries eindeutig sein und können nur unveränderliche Datentypen sein. Das bedeutet, Sie können keine Liste oder ein anderes Dictionary als Schlüssel verwenden.

Fügen Sie Ihrem Dictionary einen weiteren Schlüssel `'age'` hinzu und weisen Sie ihm die Ganzzahl `2` als Wert zu.

# --hints--

Sie sollten einen `'age'`-Schlüssel mit dem Wert `2` direkt nach `'species': 'guinea pig'` in Ihrem `copper`-Dictionary haben. Vergessen Sie das Komma nicht.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Sie sollten zwei Schlüssel-Wert-Paare in Ihrem `copper`-Dictionary haben.

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
