---
id: 6578b57361f2f132a02e2a18
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Sie können auf die in einem Dictionary gespeicherten Daten über dessen Schlüssel zugreifen:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

Folgen Sie nach Ihrem Dictionary dem obigen Beispiel, um auf den Schlüssel `'species'` von `copper` zuzugreifen und das Ergebnis auszugeben.

# --hints--

Sie sollten die Zuweisung Ihres Dictionary nicht verändern.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Sie sollten `copper['species']` verwenden, um auf den Wert des Schlüssels `'species'` zuzugreifen.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

Sie sollten `print()` aufrufen und `copper['species']` als Argument übergeben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
