---
id: 65789506b30453080f77470c
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

Bisher sind Ihnen bereits verschiedene Datentypen begegnet:

- Unveränderliche Datentypen wie ganze Zahlen, Strings, Tuple und Booleans.
- Veränderliche Datentypen wie Listen und Dictionaries.

Ein Dictionary erkennen Sie an einem Paar geschweifter Klammern, `{}`.

Beginnen Sie damit, eine Variable namens `copper` zu erstellen und weisen Sie ihr ein leeres Dictionary mit einem Paar geschweifter Klammern zu, so wie Sie eine leere Liste mit einem Paar eckiger Klammern erstellen würden.

# --hints--

Sie sollten eine Variable namens `copper` haben.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

Ihre Variable `copper` sollte den Wert eines leeren Dictionaries haben. Verwenden Sie dafür ein Paar geschweifter Klammern.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
