---
id: 6579ca0923cfa7162089d2f0
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Die gleiche Syntax kann verwendet werden, um den Wert eines bestehenden Schlüssels zu ändern.

Greifen Sie direkt vor dem `print()`-Aufruf auf den `'species'`-Schlüssel zu und weisen Sie seinen Wert neu zu `'Cavia porcellus'`.

# --hints--

Sie sollten den `copper['species']` vor dem `'Cavia porcellus'`-Aufruf auf den String `print()` neu zuweisen.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
