---
id: 6579cd5f6dd62c189e53ddbb
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Nachdem Sie die grundlegenden Aspekte von Dictionaries überprüft haben, können Sie mit dem Aufbau des Algorithmus für den kürzesten Pfad fortfahren.

Löschen Sie jede Codezeile nach der Deklaration des `copper`-Dictionaries.

# --hints--

Sie sollten alle Codezeilen nach der Deklaration Ihres Dictionaries löschen.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
