---
id: 6579cbab9825b8170974c69a
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Sie können ein Schlüssel-Wert-Paar aus einem Dictionary entfernen, indem Sie das Schlüsselwort `del` verwenden:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

Direkt vor Ihrer `for`-Schleife verwenden Sie das Schlüsselwort `del`, um den Schlüssel `'age'` und seinen Wert aus `copper` zu löschen.

# --hints--

Sie sollten das Schlüsselwort `del` verwenden, um `copper['age']` vor der `for`-Schleife zu löschen.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
