---
id: 6579dd49fa8a8e1fd06b85a9
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Graphen sind Datenstrukturen, die Beziehungen zwischen Paaren von Elementen darstellen.
Diese Elemente, genannt *Knoten*, können reale Objekte, Entitäten, Punkte im Raum oder andere sein. Die Verbindungen zwischen den Knoten werden *Kanten* genannt.

Hier ist eine visuelle Darstellung eines Graphen:

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


Benennen Sie das `copper`-Dictionary in `my_graph` um. Dieses wird den Graphen darstellen, um Ihren Algorithmus zu testen.

# --hints--

Sie sollten Ihr `copper`-Dictionary in `my_graph` umbenennen.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

Ihre `my_graph`-Variable sollte ein Dictionary sein.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
