---
id: 65796fac81f983127558f3f4
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Ändern Sie nun Ihren bestehenden `print()`-Aufruf, um den Wert des `age`-Schlüssels auszugeben.


# --hints--

Sie sollten `copper['age']` verwenden, um auf den Wert des `age`-Schlüssels zuzugreifen.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

Sie sollten `print()` aufrufen und `copper['age']` als Argument übergeben.

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
