---
id: 6579717f0920131304286804
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Um nach der Deklaration eines Dictionaries ein neues Schlüssel-Wert-Paar hinzuzufügen, können Sie den Schlüssel auf dieselbe Weise angeben, wie Sie auf einen vorhandenen Schlüssel zugreifen würden, und den Wert des neuen Schlüssels mit dem Zuweisungsoperator festlegen:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

Löschen Sie Ihren `print()`-Aufruf. Fügen Sie dann nach der `copper`-Deklaration den Schlüssel `'food'` zu Ihrem Dictionary hinzu und weisen Sie ihm den Wert `'hay'` zu.

# --hints--

Sie sollten `print(copper['age'])` nicht in Ihrem Code haben.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

Sie sollten den Schlüssel `'food'` zu `copper` hinzufügen, nachdem Sie das Dictionary deklariert haben.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

Sie sollten `copper['food']` nach der Deklaration des Dictionaries auf `'hay'` setzen.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
