---
id: 65affe602a2a828b0382a058
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Ersetzen Sie den Schlüssel `'age'` durch den String `'B'` und setzen Sie dessen Wert auf den String `'A'`, um die Verbindung zwischen den Knoten in beide Richtungen darzustellen.

# --hints--

Ihr Wörterbuch sollte einen Schlüssel `'B'` enthalten.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

Ihr Schlüssel `'B'` sollte den Wert `'A'` haben.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

Ihr Wörterbuch sollte einen Schlüssel `'A'` enthalten.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Ihr Schlüssel `'A'` sollte den Wert `'B'` haben.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Ihr Wörterbuch sollte zwei Schlüssel enthalten.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
