---
id: 65b7d25992879698180e6a71
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Schließlich, wenn Sie die Schlüssel-Wert-Paare durchlaufen möchten, können Sie die `.items()`-Methode verwenden.

Ändern Sie Ihre `for`-Schleife so, dass sie `copper.items()` anstelle von `copper.values()` durchläuft.

# --hints--

Sie sollten Ihre `for`-Schleife so ändern, dass sie `copper.items()` anstelle von `copper.values()` durchläuft.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.items()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.values():
    print(i)
--fcc-editable-region--
```
