---
id: 65b7cf140d34058d7ea8935f
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Wenn Sie die Werte der Wörterbuchschlüssel durchlaufen möchten, ist eine Möglichkeit, die `.values()`-Methode zu verwenden.

Ändern Sie Ihre `for`-Schleife so, dass sie `copper.values()` anstelle von `copper` durchläuft, und sehen Sie sich die Ausgabe an.

# --hints--

Sie sollten Ihre `for`-Schleife so ändern, dass sie `copper.values()` anstelle von `copper` durchläuft.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
