---
id: 65b7d3319ebec69b983fb91b
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Wie Sie an der Ausgabe sehen können, erstellt `.items()` eine Datenstruktur, die jedes Schlüssel-Wert-Paar in einem eigenen Tuple speichert. Um die Elemente in diesen Tuples zu durchlaufen, können Sie eine zweite Schleifenvariable hinzufügen:

```py
for i, j in dict.items():
    print(i, j)
```

Ändern Sie Ihre `for`-Schleife so, dass sie zwei Schleifenvariablen verwendet, und geben Sie beide innerhalb des Schleifenblocks aus.

# --hints--

Sie sollten Ihre `for`-Schleife so ändern, dass sie zwei Variablen verwendet.

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

Sie sollten beide Schleifenvariablen innerhalb Ihrer `for`-Schleife ausgeben.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
