---
id: 65b7cd2b7bd9a684ccf1dc16
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Um die Schlüssel eines Dictionaries zu durchlaufen, können Sie einfach das Dictionary in eine `for` Schleife setzen. Der folgende Code würde jeden Schlüssel im Dictionary `dict` ausgeben:

```py
for i in dict:
   print(i)
```

Ersetzen Sie den `print()` Aufruf durch eine `for` Schleife, die über `copper` iteriert und jeden Schlüssel ausgibt.

# --hints--

Sie sollten keinen `print(copper)` in Ihrem Code haben.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

Sie sollten eine `for` Schleife erstellen, um das `copper` Dictionary zu durchlaufen.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

Sie sollten jeden Schlüssel des `copper` innerhalb Ihrer `for` Schleife ausgeben.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
