---
id: 612e8279827a28352ce83a72
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Ahora, copie el conjunto de siete elementos `.key` y pegue dos conjuntos adicionales en el div `.keys`.

# --hints--

Debe tener 21 elementos `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

Debe tener 15 elementos `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

Todos los elementos `.key` deben estar dentro de su elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

Todos los elementos `.black--key` deben estar dentro de su elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
