---
id: 612e804c54d5e7308d7ebe56
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Dentro de tu elemento `.keys`, agrega siete elementos `div`, cada uno con la clase `key`.

# --hints--

Debes crear siete nuevos elementos `div`.

```js
const divDivDiv = document.querySelectorAll('div');
assert.lengthOf(divDivDiv, 9);
```

Tus siete nuevos elementos `div` deben estar dentro de tu elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf([...keys?.children], 7);
[...keys?.children].forEach(
  child => assert.equal(child.tagName, 'DIV')
)
```

Tus siete nuevos elementos `div` deben tener todos la `class` establecida como `key`.

```js
const keys = document.querySelector('.keys');
[...keys?.children].forEach(child => assert.isTrue(child?.classList.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
