---
id: 612e78af05201622d4bab8aa
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Crea un elemento `div` dentro de tu elemento `body` con el `id` establecido en `piano`.

# --hints--

Debes crear un nuevo elemento `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Su elemento `div` debe estar dentro del elemento `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

Su elemento `div` debe tener el `id` configurado como `piano`.

```js
const div = document.querySelector('div');
assert.equal(div?.getAttribute('id'), 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
