---
id: 65c4f22498d22ed775ef8efb
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Jetzt müssen Sie überprüfen, ob der `node` Parameter `None` ist. Wenn ja, bedeutet das, dass die Methode einen Blattknoten oder eine freie Stelle im Baum erreicht hat, an der der neue Knoten eingefügt werden soll.

Ersetzen Sie im Rumpf der `_insert` Methode `pass` durch eine `if` Anweisung, die überprüft, ob `node is None`.

Geben Sie im neuen `if` Block `TreeNode(key)` zurück, um eine neue `TreeNode` Instanz mit dem übergebenen Schlüssel zu erstellen. Dies wird der neue Blattknoten und fügt den Schlüssel effektiv in den Baum ein.

# --hints--

Sie sollten in Ihrer `pass` Methode kein `_insert` haben.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

Sie sollten eine `if` Anweisung schreiben, um zu überprüfen, ob `node is None`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

Sie sollten aus dem `TreeNode(key)` Block `if` zurückgeben.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
