---
id: 65c4f0d32ca17ad4b1636b0e
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Löschen Sie innerhalb der `__init__`-Methode `pass` und initialisieren Sie `root` mit dem Wert `None`.

Das `root`-Attribut stellt den Wurzelknoten des binären Suchbaums dar. Da dies der Konstruktor ist, beginnt ein neues `BinarySearchTree`-Objekt mit einem leeren Baum, weshalb das `root`-Attribut auf `None` gesetzt wird.

# --hints--

Sie sollten die `pass`-Anweisung aus der `__init__`-Methode entfernen.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Sie sollten das `root`-Attribut mit `None` unter Verwendung von `self.root` initialisieren.

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
