---
id: 65c4f02dbd3557d23f12db62
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Löschen Sie innerhalb der `__init__`-Methode `pass` und weisen Sie den Wert des `key`-Parameters dem `key`-Attribut des Knotens mit `self.key` zu.

Das bedeutet, dass das `key`-Attribut der `TreeNode`-Instanz auf den Wert gesetzt wird, der bei der Erstellung des Objekts übergeben wurde.

# --hints--

Sie sollten die `pass`-Anweisung aus der `__init__`-Methode löschen.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

Sie sollten den Wert des `key`-Parameters dem `key`-Attribut des Knotens mit `self.key` zuweisen.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
