---
id: 65c4f2d9fd872fd99ac659dd
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Fügen Sie eine `elif`-Bedingung hinzu, die überprüft, ob `key` > `node.key` ist.

# --hints--

Sie sollten eine `elif`-Bedingung hinzufügen, um zu prüfen, ob `key > node.key` zutrifft. Denken Sie daran, das Schlüsselwort `pass` zu verwenden.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[1].is_equivalent("key > node.key")`))})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)

--fcc-editable-region--    

```

