---
id: 65c4f2b7178afed8e88f782f
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Wenn `key` < `node.key` zu `True` ausgewertet wird, sollte der neue Knoten im linken Teilbaum platziert werden.

Löschen Sie `pass` und rufen Sie rekursiv die `_insert`-Methode mit dem linken Kind als erstem Argument und `key` als zweitem Argument auf. Weisen Sie das Ergebnis dem `left`-Attribut des aktuellen Knotens zu.

# --hints--

Sie sollten das Schlüsselwort `pass` aus dem `if`-Block entfernen.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

Sie sollten die `self._insert`-Methode aufrufen und `node.left` sowie `key` als Argumente übergeben.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

Sie sollten das Ergebnis Ihres `self._insert()`-Aufrufs dem `left`-Attribut des aktuellen Knotens zuweisen.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
