---
id: 65c4ef8463f869d0eaf87c67
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Ersetzen Sie innerhalb der `TreeNode`-Klasse `pass` durch eine `__init__`-Methode, damit Sie die Attribute des Objekts initialisieren können. Fügen Sie vorerst keine Parameter hinzu.

# --hints--

Sie sollten das Schlüsselwort `pass` aus der `TreeNode`-Klasse entfernen und es in die `__init__`-Methode verschieben.

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

Sie sollten eine `__init__`-Methode innerhalb Ihrer `TreeNode`-Klasse definieren. Denken Sie daran, das Schlüsselwort `pass` zu verwenden.

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
