---
id: 65c4f013851cefd1a4fe4c96
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Die `__init__`-Methode nimmt zwei Parameter entgegen: `self` (der die Instanz der zu erstellenden Klasse darstellt) und `key` (den Wert, der im Knoten gespeichert werden soll). Fügen Sie diese beiden Parameter der `__init__()`-Methode hinzu.

# --hints--

Sie sollten den `self`-Parameter zur Methode hinzufügen.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

Sie sollten den `key`-Parameter zur Methode hinzufügen.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
