---
id: 65c4f300da28d8da361bfa93
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Innerhalb Ihrer `elif`-Klausel ersetzen Sie `pass` durch einen Aufruf der `_insert`-Methode mit dem rechten Kind des aktuellen Knotens als erstem Argument und `key` als zweitem Argument. Weisen Sie das Ergebnis dem `right`-Attribut des aktuellen Knotens zu.

# --hints--

Sie sollten das Schlüsselwort `pass` aus dem `elif`-Block entfernen.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Sie sollten die `self._insert()`-Methode aufrufen und `node.right` sowie `key` als Argumente übergeben.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

Sie sollten das Ergebnis Ihres `self._insert()`-Aufrufs dem richtigen Attribut (`node.right`) des aktuellen Knotens zuweisen.

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
