---
id: 65c4f0b82db314d3fcc0b8e1
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Ersetzen Sie innerhalb der `BinarySearchTree`-Klasse `pass` durch eine `__init__`-Methode und fügen Sie dieser Methode einen `self`-Parameter hinzu.

# --hints--

Sie sollten das Schlüsselwort `pass` aus der `BinarySearchTree`-Klasse entfernen.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


Sie sollten eine `__init__`-Methode definieren und dieser Methode einen `self`-Parameter hinzufügen. Denken Sie daran, das Schlüsselwort `pass` im Methodenrumpf zu verwenden.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
