---
id: 65c4eb814cc977c95cd7df0e
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Projekt erstellen Sie einen binären Suchbaum (BST). Ein BST ist eine Datenstruktur, bei der jeder Knoten höchstens zwei Kinder hat, wobei das linke Kind Werte enthält, die kleiner als der Elternknoten sind, und das rechte Kind Werte, die größer als der Elternknoten sind. Dadurch sind effiziente Such- und Sortieroperationen möglich.

So sieht ein binärer Suchbaum aus:

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Beginnen Sie damit, eine leere `TreeNode`-Klasse zu definieren. Die `TreeNode`-Klasse stellt einen Knoten in einem binären Suchbaum dar. Verwenden Sie das Schlüsselwort `pass`, um den Klassenkörper zu füllen und einen Fehler zu vermeiden.

# --hints--

Sie sollten das Schlüsselwort `class` verwenden, um eine leere Klasse mit dem Namen `TreeNode` zu deklarieren.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
