---
id: 65c4f195de7e2ad5932be717
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Als Nächstes müssen Sie einen Mechanismus definieren, um Knoten in den Baum einzufügen. Dafür müssen Sie eine `_insert`-Methode definieren, die eine Hilfsfunktion ist und später von der eigentlichen `insert`-Methode verwendet wird.

Diese Methode ist rekursiv, das heißt, sie ruft sich selbst auf, um den Baum zu durchlaufen, bis der passende Ort für den neuen Knoten gefunden ist.

Definieren Sie eine `_insert`-Methode mit den Parametern `self`, `node` und `key`.

# --hints--

Sie sollten eine `_insert`-Methode innerhalb der `BinarySearchTree`-Klasse definieren. Denken Sie daran, `pass` zu verwenden.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

Ihre `_insert`-Methode sollte drei Parameter annehmen: `self`, `node` und `key`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
