---
id: 65c4f07da204f4d2f325fbd3
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Initialisieren Sie innerhalb der `__init__`-Methode die `left`- und `right`-Attribute des Knotens mit `None`. Das liegt daran, dass ein Knoten beim ersten Erstellen keine linken oder rechten Kinder hat. Denken Sie daran, das Schlüsselwort `self` zu verwenden.

# --hints--

Sie sollten `None` dem `left`-Attribut des Knotens mit `self.left` zuweisen.

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

Sie sollten `None` dem `right`-Attribut des Knotens mit `self.right` zuweisen.

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
