---
id: 65c4f2851645e8d84d14f5d1
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Jetzt müssen Sie den Baum rekursiv durchlaufen und die Werte nach dem Prinzip für binäre Bäume einfügen:

- Werte, die kleiner als der Schlüssel sind, werden im linken Teilbaum platziert
- Werte, die größer als der Schlüssel sind, werden im rechten Teilbaum platziert

Schreiben Sie nach Ihrer bestehenden bedingten Anweisung eine weitere `if`-Anweisung, um zu prüfen, ob `key` kleiner als `node.key` ist.

# --hints--

Sie sollten eine weitere if-Anweisung schreiben, um zu prüfen, ob `key` kleiner als `node.key` ist. Denken Sie daran, das Schlüsselwort `pass` zu verwenden.

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```

