---
id: 65c4f09e074dd8d37830ea00
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Erstellen Sie eine weitere leere Klasse namens `BinarySearchTree`, die einen binären Suchbaum darstellt.

# --hints--

Sie sollten eine Klasse mit dem Namen `BinarySearchTree` haben. Denken Sie daran, das Schlüsselwort `pass` zu verwenden.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
