---
id: 65c4f3258d2e4cdacc919dfd
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Am Ende Ihrer `_insert`-Methode, nachdem der Einfügevorgang abgeschlossen ist, geben Sie den aktuellen Knoten zurück, um die Baumstruktur auf den höheren Ebenen des rekursiven Aufrufstapels zu aktualisieren.


# --hints--

Sie sollten den aktuellen Knoten außerhalb der bedingten Blöcke zurückgeben.

```js
({ test: () => assert.match(code, /^\s{8}return\s+node/m) })
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)

--fcc-editable-region--

```

