---
id: 65c4f37ba6cbcfdc77e24165
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Die `insert`-Methode wird vom Benutzer aufgerufen. Zusätzlich zum `self`-Parameter benötigt sie auch einen `key`-Parameter. Dieser Parameter ist der Schlüsselwert, der in den binären Suchbaum eingefügt wird.

Fügen Sie `key` als zweiten Parameter zur Funktionsdefinition hinzu.

# --hints--

Die `insert`-Methode sollte den `self`-Parameter enthalten.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

Die `insert`-Methode sollte den `key`-Parameter enthalten.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
