---
id: 65c644829cfb63acf3479d09
title: Schritt 28
challengeType: 20
dashedName: step-28
---

# --description--

Fügen Sie die Knoten in der Liste `nodes` in die `bst`-Instanz ein, indem Sie über die Liste iterieren und die `insert`-Methode für jeden Knoten in der Liste aufrufen. Verwenden Sie `node` als Iterationsvariable.

# --hints--

Sie sollten eine `for node in nodes`-Schleife haben.

```js
({ test: () => (runPython(`assert _Node(_code).find_for_loops()[0].find_for_vars().is_equivalent("node") and _Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("nodes")`)) });
```

Sie sollten `bst.insert(node)` innerhalb der `for`-Schleife haben.

```js
({ test: () => (runPython(`assert _Node(_code).find_for("node", "nodes").find_body().is_equivalent("bst.insert(node)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

--fcc-editable-region--
```
