---
id: 65c63e6962159fa372ecdbec
title: Schritt 22
challengeType: 20
dashedName: step-22
---

# --description--

Wenn die zweite `if`-Anweisung nicht `True` ist, bedeutet das, dass der Zielschlüssel größer oder gleich dem Schlüssel des aktuellen Knotens ist.

In einem binären Suchbaum, wenn der Zielschlüssel größer als der Schlüssel des aktuellen Knotens ist, setzt sich die Suche im rechten Teilbaum fort.

Nach dem `if`-Block geben Sie das Ergebnis der `_search`-Methodenaufrufs mit dem rechten Kind des aktuellen Knotens und dem Schlüssel als Argumente zurück.

# --hints--

Sie sollten die `_search`-Methode rekursiv mit `self._search()` aufrufen.

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

Sie sollten zwei Argumente an den Aufruf der `_search`-Methode übergeben: `node.right` und `key`.

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

Sie sollten das Ergebnis des `_search`-Methodenaufrufs zurückgeben.

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
