---
id: 65c63eb01b9563a4b8a046b6
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Definieren Sie als Nächstes die Methode mit dem Namen `search` innerhalb der `BinarySearchTree`-Klasse. Geben Sie der `search`-Methode zwei Parameter: `self` und `key`.

# --hints--

Sie sollten eine Methode mit dem Namen `search` definieren. Denken Sie an das Schlüsselwort `pass`.

```js
({ test: () => assert.match(code, /def\s+search/) });
```

Die `search`-Methode sollte zwei Parameter annehmen: `self` und `key`.

```js
({ test: () => assert.match(code, /def\s+search\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
        
--fcc-editable-region--

--fcc-editable-region--
```
