---
id: 65c63fd9b387f0a6c5bf6a72
title: Schritt 24
challengeType: 20
dashedName: step-24
---

# --description--

Löschen Sie innerhalb der `search`-Methode `pass` und rufen Sie die Hilfsmethode `_search` mit den folgenden Argumenten auf.

- `self.root`: Dies ist die Wurzel des binären Suchbaums. Die Suche beginnt bei der Wurzel.
- `key`: Dies ist der Wert, den der Benutzer im binären Suchbaum finden möchte.

Intern delegiert `search` die eigentliche Suchlogik an die Hilfsmethode `_search`, die die tatsächliche rekursive Suche im binären Suchbaum durchführt.

# --hints--

Sie sollten das Schlüsselwort `pass` aus der `search`-Methode entfernen.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

Sie sollten die `_search`-Methode innerhalb der `search`-Methode aufrufen.

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

Sie sollten die `_search`-Methode aufrufen und `self.root` sowie `key` als Argumente übergeben.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
