---
id: 65c6409418806da8d0636ffc
title: Schritt 26
challengeType: 20
dashedName: step-26
---

# --description--

Die `insert` und `search` Funktionen sind fertiggestellt, jetzt ist es Zeit, sie zu testen.

Beachten Sie, dass die Knoten zu diesem Zeitpunkt nicht sortiert sind, sondern nur eingefügt werden. Das Sortieren mittels Inorder-Traversierung wird später behandelt.

Sie können eine Instanz einer Klasse in Python so erstellen:

```python
object_name = ClassName()
```

Außerhalb der Klassendefinitionen erstellen Sie eine Instanz der `BinarySearchTree`-Klasse und weisen Sie der Variablen `bst` zu.


# --hints--

Sie sollten eine Instanz der `BinarySearchTree`-Klasse erstellen.

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

Sie sollten die neue Instanz von `BinarySearchTree` der Variablen `bst` zuweisen.

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
