---
id: 65c63a4c4da62e9ae18e321a
title: Schritt 19
challengeType: 20
dashedName: step-19
---

# --description--

Es ist Zeit, an der Suchfunktion zu arbeiten. So wie Sie eine Hilfsmethode `_insert` für die `insert`-Methode erstellt haben, müssen Sie eine Hilfsmethode `_search` für die `search`-Methode erstellen.

Definieren Sie die `_search`-Methode mit drei Parametern, nämlich `self`, `node` und `key`.

# --hints--

Sie sollten eine `_search`-Methode deklarieren. Beachten Sie den führenden Unterstrich und denken Sie daran, `pass` zu verwenden.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

Ihre `_search`-Methode sollte drei Parameter annehmen: `self`, `node` und `key`.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
