---
id: 65c645b838e7deb080fc25e0
title: Schritt 29
challengeType: 20
dashedName: step-29
---

# --description--

Jetzt wurden die Knoten eingefügt. Um zu überprüfen, ob sie korrekt eingefügt wurden, können Sie im Baum nach einem Knoten suchen.

Außerhalb der for-Schleife suchen Sie nach dem Knoten `80` in der `bst`-Instanz und fügen ihn einem `print`-Aufruf hinzu. Das erste Argument der `print`-Funktion sollte ebenfalls `'Search for 80:'` sein.

# --hints--

Sie sollten das Ergebnis des Aufrufs von `bst.search(80)` ausgeben, und Ihre `print`-Anweisung sollte als erstes Argument `'Search for 80:'` haben.

```js
assert.match(code, /^print\(('|")Search for 80:('|"),\s*bst\.search\(80\)/gm);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

--fcc-editable-region--
```
