---
id: 65c63da2ed6769a10e141341
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Jetzt definieren Sie einen Basisfall für die rekursive Suche. Entfernen Sie das aktuelle `pass` und schreiben Sie eine `if`-Anweisung, die zwei Bedingungen überprüft:

- Wenn `node` `None` ist: Dies zeigt an, dass die Suche das Ende eines Zweigs erreicht hat, ohne den Schlüssel zu finden.
- Wenn `node.key == key` ist: Das bedeutet, dass der Schlüssel im aktuellen Knoten gefunden wurde.

Kombinieren Sie die beiden Bedingungen mit dem `or`-Operator und geben Sie den aktuellen Knoten innerhalb des `if`-Blocks zurück.

# --hints--

Sie sollten das Schlüsselwort `pass` aus der `_search`-Methode entfernen.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

Sie sollten eine `if`-Anweisung schreiben, die überprüft, ob `node` `None` ist oder ob `node.key` gleich `key` ist.

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

Sie sollten `node` aus dem `if`-Block zurückgeben.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
