---
id: 65c63df529bd15a24c187c62
title: Schritt 21
challengeType: 20
dashedName: step-21
---

# --description--

Schreiben Sie eine weitere `if`-Anweisung, die überprüft, ob der Zielschlüssel kleiner ist als der Schlüssel des aktuellen Knotens.

Innerhalb des `if`-Blocks geben Sie das Ergebnis des Aufrufs der `_search`-Methode mit dem linken Kind des aktuellen Knotens und `key` als Argumente zurück.

# --hints--

Sie sollten eine weitere `if`-Anweisung schreiben, um zu überprüfen, ob `key` < `node.key` ist.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

Sie sollten `self._search(node.left, key)` aus Ihrem neuen `if`-Block zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_bodies()[0].is_equivalent("return self._search(node.left, key)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node

--fcc-editable-region--
```
