---
id: 65ca03bd8eb5faf24b250c56
title: Schritt 33
challengeType: 20
dashedName: step-33
---

# --description--

Der nächste Schritt besteht darin, an der Löschung von Knoten zu arbeiten.  
Dafür gehen Sie genauso vor: Zuerst definieren Sie eine Hilfsmethode und dann die eigentliche Methode.

Definieren Sie innerhalb der `BinarySearchTree`-Klasse eine neue Hilfsmethode namens `_delete`, die drei Parameter entgegennimmt: `self`, `node` und `key`.

# --hints--

Sie sollten die `_delete`-Methode innerhalb der `BinarySearchTree`-Klasse mit den Parametern `self`, `node` und `key` definieren. Denken Sie daran, das Schlüsselwort `pass` zu verwenden.

```js
({ test: () => assert.match(code, /def\s+_delete\(\s*self\s*,\s*node\s*,\s*key\s*\)/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))

```
