---
id: 65ca0d5adf39c410cd1177cc
title: Schritt 50
challengeType: 20
dashedName: step-50
---

# --description--

Jetzt arbeiten Sie daran, den Baum basierend auf der in Reihenfolge Traversierungsmethode zu durchlaufen. Die in Reihenfolge Traversierung ist ein Tiefensuche-Algorithmus für binäre Bäume, der zuerst den linken Teilbaum, dann den aktuellen Knoten und anschließend den rechten Teilbaum besucht.

Definieren Sie die `_inorder_traversal`-Methode innerhalb der `BinarySearchTree`-Klasse und geben Sie ihr drei Parameter: `self`, `node` und `result`. Dabei ist `node` der aktuell betrachtete Knoten während der Traversierung und `result` die Liste, an die die Schlüssel in sortierter Reihenfolge angehängt werden.

# --hints--

Sie sollten eine Methode `_inorder_traversal` definieren, die drei Parameter annimmt: `self`, `node` und `result`. Denken Sie daran, `pass` zu verwenden.

```js
assert.match(code, /def\s+_inorder_traversal\(\s*self\s*,\s*node\s*,\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
            
            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)   
        
        return node

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key

--fcc-editable-region--

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
