---
id: 655cd899f8de09431eabb40c
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在这个项目中，你将通过构建合并排序算法学习数据结构。 

这是一种使用分而治之原理对数据集合进行排序的排序算法。 也就是说，它将一个集合“划分”为更小的子部分，然后通过对子部分进行独立排序来“征服”它们，然后合并已排序的子部分。

创建一个名为 `merge_sort` 的函数。 此函数将处理排序数字列表的任务。 

在函数体中使用 `pass` 关键字。

# --hints--

您应该声明一个名为 `merge_sort` 的函数。 别忘记在函数体中使用 `pass` 关键字。

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
