---
id: 6564a9fe51964c229d5b7f4c
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

你需要一个表示要排序的数据集合的参数。 在 `merge_sort` 函数中创建一个 `array` 参数。

# --hints--

`merge_sort` 函数应该有一个参数：`array`。

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
