---
id: 656627b47bd2d2a4afbc945d
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

在上一步中，你得到了中间点。 你可以使用它将 `array` 分成两部分，并将每个部分分配给新变量。 

使用切片语法提取`array`的左半部分并将其分配给名为`left_part`的变量。


# --hints--

`merge_sort` 函数中应该有一个名为 `left_part` 的变量。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

你应该将 `array[:middle_point]` 分配给 `left_part` 变量。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
