---
id: 656680b0fc79f2c38a34d90e
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

现在是时候对列表（`left_part` 和 `right_part`）进行排序并合并到原始 `array` 中了。 

你可以通过比较两个列表中的元素并将较小的元素合并到主列表中来实现。 你将对 `left_part` 和 `right_part` 中的所有索引进行此比较。

创建三个变量：`left_array_index`、`right_array_index` 和 `sorted_index`，并将它们的值设置为 `0`。 这些变量将帮助您在排序过程中跟踪每个索引。

# --hints--

你应该在 `merge_sort` 函数中有一个名为 `left_array_index` 的变量

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index/);
  }
})
```

`left_array_index` 的值应为 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index\s*=\s*0/);
  }
})
```

你应该在 `merge_sort` 函数中有一个名为 `right_array_index` 的变量

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index/);
  }
})
```

`right_array_index` 的值应为 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index\s*=\s*0/);
  }
})

```

你应该在 `merge_sort` 函数中有一个名为 `sorted_index` 的变量

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index/);
  }
})
```

`sorted_index` 的值应为 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index\s*=\s*0/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)

--fcc-editable-region--

--fcc-editable-region--
```
