---
id: 656639771fed09b5c0e8fe71
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

现在你已将 `array` 列表分成两个单独的列表，你将继续划分每个列表，直到每个元素都独立于其自己的列表中。 具有单个数字的列表始终是有序的。

为此，请在函数内部递归调用 `merge_sort`。 现在不要传入任何参数。

# --hints--

你应该在函数体底部调用 `merge_sort`。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
