---
id: 65670d1ef177e7e2b76d9528
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

你在前面的步骤中创建的 `if` 和 `else` 语句将元素分配给已排序的数组。 

分配给排序数组的每个元素都会占用列表中的一个索引。 因此你必须移动到下一个没有元素的索引。

在 `if`/`else` 块下方，但仍在 `while` 循环内，将 `sorted_index` 增加 1。 这不应该位于 `if` 或 `else` 语句的主体中

# --hints--

你应该将 `sorted_index` 的当前值增加 1。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{8}sorted_index\s*(\+=\s*1|=\s*sorted_index\s*\+\s*1)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
--fcc-editable-region--
```
