---
id: 656702f8b4cbd8cbf0a433c6
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

在你的函数里写一个`while`循环。此循环会从 `left_part`和`right_part`里个拿一个变量然后比较。最小的变量会合并到`array`串列。 

为循环创建两个条件：一个检查 `left_array_index` 是否小于 `left_part` 的长度，另一个检查 `right_array_index` 是否小于 `right_part` 的长度。

# --hints--

你的 `while` 循环应满足这些条件：`left_array_index < len(left_part) and right_array_index < len(right_part)`。不要忘记在循环体中使用 `pass` 关键字。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

在循环体中使用 `pass` 关键字

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
