---
id: 656706afd65547d22bee0b68
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

在 `while` 循环中，将 `pass` 替换为 `if` 语句，该语句检查 `left_part` 的索引是否小于 `right_part` 的索引。

在 `if` 语句中使用 `pass` 关键字。

# --hints--

你应该将 `pass` 关键字替换为带有此条件的 `if` 语句：`left_part[left_array_index] < right_part[right_array_index]`。为 `if` 语句的体添加 `pass` 关键字。

```js
({
 test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_body().is_equivalent("if left_part[left_array_index] < right_part[right_array_index]:\\n  pass")`))
 })
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        pass
--fcc-editable-region--
```
