---
id: 6569b743630ee592a65a7e2f
title: 步骤 26
challengeType: 20
dashedName: step-26
---

# --description--

在测试 `merge_sort()` 函数之前，你需要创建一个基本案例，当 `array` 的长度小于或等于 `1` 时停止函数执行。

此基本案例将停止递归调用。 如果没有它，即使列表已经排序或其中没有元素，合并排序操作仍将继续运行。

紧接着函数声明，创建一个带有此条件的 `if` 语句：`len(array) <= 1`。在函数体中使用 `pass` 关键字。

# --hints--

你应该创建一个 `if` 语句来查看 `len(array) <= 1`。不要忘记在 `if` 语句的体中使用 `pass` 关键字。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):

--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
