---
id: 6577320da0d4c2e594d418e2
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

把 `right_part` 作为参数传递给你在上一步中添加的 `merge_sort()` 调用。

# --hints--

你应该将 `right_part` 作为参数传递给最后一个 `merge_sort()` 调用。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*right_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort()
--fcc-editable-region--
```
