---
id: 65772ef923f922cd720e5008
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

将 `left_part` 作为参数传入在最后一步添加的 `merge_sort()` 调用。

# --hints--

你应该将 `left_part` 作为参数传递给 `merge_sort()` 调用。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*left_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort()
--fcc-editable-region--
```
