---
id: 657f59751d5a5c9b069ae0f3
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

你可以使用 `__name__` 变量来确定 Python 脚本是否作为主程序运行，或者是否作为模块（在另一个 Python 文件中编写的代码）导入。

如果 `__name__` 的值设置为 `'__main__'`，则表示当前脚本是主程序，而不是模块。

在这个项目中，你将使用当前脚本作为主程序。

创建一个 `if` 语句来判断 `__name__` 的值是否是 `'__main__'`。

在 `if` 语句中使用 `pass` 关键字。

# --hints--

`if` 语句中应该检查 `__name__` 的值是否为 `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1

    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
