---
id: 657f3a661730891aa64f3568
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

使用切片语法提取 `array` 的右一半，并将其分配给变量 `right_part` 。

# --hints--

你应该有一个变量 `right_part`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_part\s*(?!=)/m);
  }
})
```

你应该将 `array[middle_point:]` 分配给 `right_part` 变量

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_part\s*=\s*array\s*\[\s*middle_point\s*:\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
