---
id: 686e64be3cd9db25282a956d
title: Design einer Fotoausstellung
challengeType: 25
dashedName: lab-photography-exhibit
demoType: onClick
saveSubmissionToDB: true
---

# --description--

In diesem Labor erstellen Sie ein Layout für eine Fotoausstellung mit Tailwind CSS Utility-Klassen.

Für die Bilder können Sie bei Bedarf die folgenden URLs verwenden:

- https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg
- https://cdn.freecodecamp.org/curriculum/labs/alps.jpg
- https://cdn.freecodecamp.org/curriculum/labs/sea.jpg

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein Element mit einer Klasse namens `main-container` haben.
2. Ihr `.main-container`-Element sollte außerdem die folgenden Tailwind CSS Utility-Klassen haben:
   - die korrekte Utility-Klasse zum Erstellen eines CSS-Grid-Containers.
   - eine Utility-Klasse zur Definition der Anzahl der Spalten im Grid. Sie sollten die Utility-Klasse verwenden, die eine feste Anzahl von Spalten nutzt. Zum Beispiel `grid-cols-<number>`.
   - eine Utility-Klasse zum Festlegen des Zeilen- und Spaltenabstands für die Grid-Elemente.
3. Innerhalb Ihres `.main-container`-Elements sollten sich mindestens drei Elemente befinden, die jeweils eine Klasse namens `card` haben.
4. Jedes Ihrer `.card`-Elemente sollte außerdem die folgenden Tailwind CSS Utility-Klassen haben:
   - eine Utility-Klasse zum Festlegen eines abgerundeten Rahmens Ihrer Wahl.
   - eine Utility-Klasse zum Festlegen des Innenabstands Ihrer Wahl. Sie sollten dieses Format `p-<number>` verwenden.
5. Innerhalb jedes Ihrer `.card`-Elemente sollten sich folgende Elemente befinden:
   - ein Bildelement mit den Attributen `alt` und `src`.
   - ein Element mit der Klasse namens `subheading`.
   - ein Element mit der Klasse namens `description`.
6. Jedes Ihrer Bildelemente sollte eine Utility-Klasse zum Festlegen eines abgerundeten Rahmens Ihrer Wahl haben.
7. Jedes Ihrer `.subheading`-Elemente sollte Utility-Klassen verwenden, um die Schriftstärke und -größe Ihrer Wahl festzulegen.
8. Jedes Ihrer `.description`-Elemente sollte eine Utility-Klasse verwenden, um die Schriftgröße Ihrer Wahl festzulegen.

# --hints--

Sie sollten ein Element mit einer Klasse namens `main-container` haben.

```js
assert.exists(document.querySelector(".main-container"));
```

Ihr `.main-container`-Element sollte die korrekte Tailwind CSS Utility-Klasse zum Erstellen eines CSS-Grid-Containers verwenden.

```js
const mainContainer = document.querySelector(".main-container");
assert.isTrue(mainContainer.classList.contains("grid"));
```

Ihr `.main-container`-Element sollte eine Utility-Klasse zur Definition der Anzahl der Spalten im Grid verwenden.

```js
const mainContainer = document.querySelector(".main-container");
// Grid classes can vary, so this check is for the presence of any grid-cols class
const hasGridCols = [...mainContainer.classList].some(className => className.startsWith("grid-cols-"));
assert.isTrue(hasGridCols);
```

Ihr `.main-container`-Element sollte eine Utility-Klasse zum Festlegen des Zeilen- und Spaltenabstands für die Grid-Elemente verwenden.

```js
const mainContainer = document.querySelector(".main-container");
// Check for any gap class including gap-x and gap-y

const hasGap = [...mainContainer.classList].some(c => c.startsWith("gap-")) ||
  (classList.some(c => c.startsWith("gap-x-")) &&
   classList.some(c => c.startsWith("gap-y-")));

assert.isTrue(hasGap);
```

Ihr `.main-container`-Element sollte mindestens drei Elemente enthalten, die jeweils eine Klasse namens `card` haben.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isAtLeast(cards.length, 3);
```

Jedes Ihrer `.card`-Elemente sollte eine Utility-Klasse zum Festlegen eines abgerundeten Rahmens haben.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);

const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

cards.forEach(card => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => card.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
})
```

Jedes Ihrer `.card`-Elemente sollte eine Utility-Klasse zum Festlegen des Innenabstands haben.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);
const hasPadding = [...cards].every(card => {
  return [...card.classList].some(c => c.startsWith("p-"));
});
assert.isTrue(hasPadding);
```

Jedes Ihrer `.card`-Elemente sollte ein Bildelement enthalten.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
```

Jedes Ihrer `img`-Elemente sollte ein `src`-Attribut haben.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("src")));
images.forEach(img => assert.isTrue(img.getAttribute("src").length > 0));
```

Jedes Ihrer `img`-Elemente sollte ein `alt`-Attribut haben.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("alt")));
images.forEach(img => assert.isTrue(img.getAttribute("alt").length > 0));
```

Jedes Ihrer `.card`-Elemente sollte ein Element mit der Klasse namens `subheading` enthalten.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isAtLeast(subheadings.length, 3);
```

Jedes Ihrer `.subheading`-Elemente sollte Textinhalt haben.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(subheadings);

subheadings.forEach(subheading => {
  assert.isNotEmpty(subheading.textContent.trim());
});
```

Jedes Ihrer `.card`-Elemente sollte ein Element mit der Klasse namens `description` enthalten.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
assert.isAtLeast(descriptions.length, 3);
```

Jedes Ihrer `.description`-Elemente sollte Textinhalt haben.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(descriptions);

descriptions.forEach(description => {
  assert.isNotEmpty(description.textContent.trim());
});
```

Jedes Ihrer `img`-Elemente sollte eine Utility-Klasse zum Festlegen eines abgerundeten Rahmens haben.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isNotEmpty(images);
const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

images.forEach(img => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => img.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
});
```

Jedes Ihrer `.subheading`-Elemente sollte Utility-Klassen verwenden, um die Schriftstärke festzulegen.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedFontWeights = [
  "font-thin",
  "font-extralight",
  "font-light",
  "font-normal",
  "font-medium",
  "font-semibold",
  "font-bold",
  "font-extrabold",
  "font-black"
];

subheadings.forEach(subheading => {
  const hasFontWeight = predefinedFontWeights.some(weight => subheading.classList.contains(weight));
  assert.isTrue(hasFontWeight);
});
```

Jedes Ihrer `.subheading`-Elemente sollte Utility-Klassen verwenden, um die Schriftgröße festzulegen.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

subheadings.forEach(subheading => {
  const hasFontSize = predefinedSizes.some(size => subheading.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

Jedes Ihrer `.description`-Elemente sollte eine Utility-Klasse verwenden, um die Schriftgröße festzulegen.

```js
const descriptions = document.querySelectorAll(".description");
assert.isNotEmpty(descriptions);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

descriptions.forEach(description => {
  const hasFontSize = predefinedSizes.some(size => description.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
 
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-purple-800 text-yellow-400 min-h-screen">
    <h1 class="text-center text-4xl font-bold my-12">Photography Exhibit</h1>

    <main class="main-container grid grid-cols-1 md:grid-cols-3 gap-8 px-4">
      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="Ancient Colosseum"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Colosseum, Rome</h2>
          <p class="description text-sm mb-4">
            There is so much history and beauty. The light hits it just right.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="Alpine Mountains"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">The Alps</h2>
          <p class="description text-sm mb-4">
            The mountains go on forever and the view is breathtaking.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="Calm Sea"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Seascape</h2>
          <p class="description text-sm mb-4">
            The water was still, the air was cool and the feeling is pure peace.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>
    </main>
  </body>
</html>
```
