---
id: 66ebd4ae2812430bb883c787
title: Створіть хаб для подій
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

У цій лабораторній роботі ви будете використовувати семантичні HTML-елементи, щоб створити структуру вебсторінки. Ви додасте вміст і зображення, щоб сторінка виглядала як справжній хаб для подій.

Виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати елемент `header`.
1. В елементі `header` мають бути елемент `h1` з текстом `Event Hub` та елемент `nav`.
1. В елементі `nav` має бути невпорядкований список з двох елементів, що містять посилання на різні розділи сторінки. Перший елемент повинен мати текст `Upcoming Events`, а другий — `Past Events`.
1. Кожне посилання має бути представлене елементом `a` з атрибутом `href`, що направляє на відповідний розділ сторінки: `#upcoming-events` або `#past-events`.
1. Ви повинні мати елемент `main`, який містить різні розділи сторінки.
1. Елемент `main` має містити два елементи `section`.
1. Перший елемент `section` повинен мати атрибут `id` зі значенням `upcoming-events`.
1. Розділ `#upcoming-events` повинен містити:

   - елемент `h2` з текстом `Upcoming Events`;
   - два елементи `article`. Кожен з елементів повинен представляти подію і мати:
      - елемент `h3` з назвою події;
      - елемент `p` з описом події. Внизу можна додати дату, якщо хочете.

1. Другий елемент `section` повинен мати атрибут `id` зі значенням `past-events`.
1. Розділ `#past-events` повинен містити:

   - елемент `h2` з текстом `Past Events`;
   - два елементи `article`. Кожен з елементів повинен представляти минулу подію і мати:
     - елемент `h3` з назвою події;
     - елемент `p` з описом події. Внизу можна додати дату, якщо хочете;
     - елемент-зображення з атрибутом `src`, що вказує на файл зображення, та атрибутом `alt`, що описує зображення.

**Зверніть увагу:** для опису і дати подій можна використовувати будь-який текст. Якщо хочете, можете використати ці посилання для зображень:

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

Ви повинні мати елемент `header`.

```js
assert.isNotNull(document.querySelector("header"));
```

Елемент `header` має бути після початкового тега `body`.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

В елементі `header` має бути елемент `h1` з текстом `Event Hub`.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText, "Event Hub");
```

В елементі `header` має бути елемент `nav` (після елемента `h1`).

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

Елемент `nav` має містити невпорядкований список з двох елементів.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

Першим елементом у невпорядкованому списку має бути посилання.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

Другим елементом у невпорядкованому списку має бути посилання.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

Текстом першого елемента в невпорядкованому списку має бути `Upcoming Events`.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText, "Upcoming Events");
```

Перший елемент у невпорядкованому списку повинен мати атрибут `href` зі значенням `#upcoming-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

Текстом другого елемента в невпорядкованому списку має бути `Past Events`.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText, "Past Events");
```

Другий елемент у невпорядкованому списку повинен мати атрибут `href` зі значенням `#past-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

Після кінцевого тегу елемента `header` має бути елемент `main`.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

Елемент `main` має містити два елементи `section`.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

Перший елемент `section` повинен мати атрибут `id` зі значенням `upcoming-events`.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

Другий елемент `section` повинен мати атрибут `id` зі значенням `past-events`.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

У розділі `#upcoming-events` має бути елемент `h2` з текстом `Upcoming Events`.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText, "Upcoming Events");
```

У розділі `#upcoming-events` має бути два елементи `article` (під елементом `h2`).

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Обидва елементи `article` в розділі `#upcoming-events` повинні мати елемент `h3` з назвою події.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Обидва елементи `article` в розділі `#upcoming-events` повинні мати елемент-абзац з описом події.

```js
const pElement = document.querySelectorAll('#upcoming-events article h3+p').length;
assert.strictEqual(pElement, 2);
```

У розділі `#past-events` має бути елемент `h2` з текстом `Past Events`.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText, "Past Events");
```

У розділі `#past-events` має бути два елементи `article` (під елементом `h2`).

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Обидва елементи `article` в розділі `#past-events` повинні мати елемент `h3` з назвою події.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Обидва елементи `article` в розділі `#past-events` повинні мати елемент-абзац з описом події.

```js
const pElement = document.querySelectorAll('#past-events article h3+p').length;
assert.strictEqual(pElement, 2);
```

Обидва елементи `article` в розділі `#past-events` повинні мати елемент-зображення.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

Обидва елементи-зображення в розділі `#past-events` повинні мати атрибут `src`, що вказує на файл зображення.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

Обидва елементи-зображення в розділі `#past-events` повинні мати атрибут `alt`, що описує зображення.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

Всі елементи `h3` повинні містити назву події.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText)));
```

Всі елементи `p` повинні містити опис події.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText)));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```

