---
id: 68e6aff5325410118e76e7be
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop üben Sie den Umgang mit ARIA, indem Sie einen barrierefreien Audiocontroller erstellen.

Fügen Sie zunächst ein `h1`-Element mit dem Text `Audio Controls` hinzu.

# --hints--

Sie sollten ein `h1`-Element haben.

```js
assert.exists(document.querySelector('h1'));
```

Ihr `h1`-Element sollte sich innerhalb Ihres `body`-Elements befinden.

```js
assert.exists(document.querySelector('body > h1'));
```

Ihr `h1`-Element sollte den Text `Audio Controls` enthalten.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

--fcc-editable-region--
    
--fcc-editable-region--

  </body>
</html>
```
