---
id: 68eac57fadcc662ce917e75e
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie als Nächstes ein `button`-Element hinzu, dessen `type`-Attribut auf den Wert `button` gesetzt ist.

Der Text des `button`-Elements sollte `Play` sein.

# --hints--

Sie sollten ein `button`-Element haben.

```js
assert.exists(document.querySelector('button'));
```

Ihr `button`-Element sollte ein `type`-Attribut haben.

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

Ihr `button`-Element sollte das `type`-Attribut mit dem Wert `button` haben.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

Ihr `button`-Element sollte den Text `Play` enthalten.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

--fcc-editable-region--
    
--fcc-editable-region--

  </body>
</html>
```
