---
id: 68eb7f2397da251a40e04aa2
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Fügen Sie als Nächstes ein `div`-Element hinzu, das alle Elemente für die Lautstärkeregelung enthält.

Ein `span`-Element ist ein generischer inline Container für Inhalte. Es wird häufig verwendet, um einen Textabschnitt anders zu gestalten als den umgebenden Text oder um Text einzuschließen und ihm eine `id` zu geben, auf die andere Elemente (wie ARIA-Attribute) verweisen können.

Verschachteln Sie innerhalb des `div`-Elements ein `span`-Element mit einem `id`-Attribut, das auf `volume-label` gesetzt ist.
Der Text für das `span`-Element sollte `Volume` sein.

Das `id` wird später von ARIA-Attributen referenziert, damit unterstützende Technologien diesen Text als Bezeichnung für die Lautstärkeregelung erkennen können.

# --hints--

Sie sollten ein `div`-Element unter Ihrem `button`-Element haben.

```js
assert.exists(document.querySelector('button + div'));
```

Sie sollten ein `span`-Element innerhalb Ihres `div`-Elements haben.

```js
assert.exists(document.querySelector('button + div > span'));
```

Ihr `span`-Element sollte ein `id`-Attribut haben.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

Ihr `span`-Element sollte ein `id`-Attribut mit dem Wert `volume-label` haben.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

Ihr `span`-Element sollte den Text `Volume` haben.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
