---
id: 68ec0a48c9d6db3594d768c3
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Jetzt ist es Zeit, ein weiteres `span`-Element hinzuzufügen.

Dieses `span`-Element sollte ein `id`-Attribut mit dem Wert `volume-description` haben.
Außerdem sollte der Text für dieses `span`-Element `Adjust the sound level` sein.

Ähnlich wie beim vorherigen `span`-Element wird die `id` dieses `span` später durch ARIA-Attribute referenziert, damit unterstützende Technologien diesen Text als Beschreibung für die Lautstärkeregelung erkennen können.

# --hints--

Sie sollten zwei `span`-Elemente innerhalb des `div`-Elements haben.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

Ihr zweites `span`-Element sollte ein `id`-Attribut haben.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

Ihr zweites `span`-Element sollte ein `id`-Attribut mit dem Wert `volume-description` haben.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

Ihr zweites `span`-Element sollte den Text `Adjust the sound level` enthalten.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
