---
id: 68ed488821845990c911a223
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Ihr barrierefreier Audiocontroller ist fast fertig, aber es gibt noch eine letzte Sache hinzuzufügen.

Fügen Sie ein `button`-Element mit dem `type`-Attribut, das auf `button` gesetzt ist, und mit dem Text `Mute` unter dem `div`-Element hinzu.

Und mit dieser Änderung haben Sie diesen Workshop abgeschlossen!

# --hints--

Sie sollten ein `button`-Element unter dem `div`-Element haben.

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

Ihr `button`-Element sollte ein `type`-Attribut haben.

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

Ihr `button`-Element sollte ein `type`-Attribut mit dem Wert `button` haben.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

Ihr `button`-Element sollte den Text `Mute` enthalten.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
