---
id: 68ed4204ab7edc66b6cb2383
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Wie Sie aus vorherigen Lektionen wissen, wird das `aria-labelledby`-Attribut verwendet, wenn bereits Text auf der Seite vorhanden ist, der als Bezeichnung dienen kann. In diesem Fall haben Sie das `id`-Attribut zu Ihren `span`-Elementen hinzugefügt, die hier nun verwendet werden.

Fügen Sie dem `input`-Element ein `aria-labelledby`-Attribut hinzu, das auf `volume-label volume-description` gesetzt ist.

# --hints--

Ihr `input`-Element sollte ein `aria-labelledby`-Attribut haben.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

Ihr `input`-Element sollte ein `aria-labelledby`-Attribut mit dem Wert `volume-label volume-description` haben.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
      <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
