---
id: 68ed3a96e0ebcc36565b39e5
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Als Nächstes müssen Sie ein `input`-Element mit dem `type`-Attribut, das auf `range` gesetzt ist, hinzufügen, um den Lautstärkeregler zu erstellen. Setzen Sie das `min`-Attribut auf `0`, das `max`-Attribut auf `100` und das `value`-Attribut auf `50`, um den Standardlautstärkepegel festzulegen.

# --hints--

Sie sollten ein `input`-Element unter Ihrem zweiten `span`-Element haben.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

Ihr `input`-Element sollte ein `type`-Attribut besitzen.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

Ihr `input`-Element sollte ein `type`-Attribut mit dem Wert `range` besitzen.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

Ihr `input`-Element sollte ein `min`-Attribut besitzen.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

Ihr `input`-Element sollte ein `min`-Attribut mit dem Wert `0` besitzen.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

Ihr `input`-Element sollte ein `max`-Attribut besitzen.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

Ihr `input`-Element sollte ein `max`-Attribut mit dem Wert `100` besitzen.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

Ihr `input`-Element sollte ein `value`-Attribut besitzen.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

Ihr `input`-Element sollte ein `value`-Attribut mit dem Wert `50` besitzen.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
