---
id: 66b205e6eacba4c4e54ea434
title: 創建一個電子郵件掩碼器
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

在本實驗室中，您將用星號遮蓋電子郵件地址中的用戶名部分。 掩碼是一個用星號或其他字符隱藏或替換敏感信息的術語。

例如，如果電子郵件地址是 `myEmail@email.com`，那麼掩碼的電子郵件地址將是 `m*****l@email.com` 。

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶故事：**

1. 創建一個名爲 `maskEmail` 的函數， 將 `email` 作爲參數。
2. 在函數中，你應對 `email` 進行掩碼處理，並將域名附於其後。  記住你可以使用 `slice`、`repeat`、`indexOf` 或 `replace` 等字符串方法來幫忙實現。
3. 在函數外部，聲明一個名爲 `email` 的變量，用於存儲要掩碼的電子郵件地址。
4. 調用 `maskEmail` 函數，將 `email` 變量作爲參數，並將結果輸出到控制檯。
5. `maskEmail("apple.pie@example.com")` 應返回 `"a*******e@example.com"`。
6. `maskEmail("freecodecamp@example.com")` 應返回 `"f**********p@example.com"`。
7. `maskEmail("info@test.dev")` 應返回 `"i**o@test.dev"`。
8. `maskEmail("user@domain.org")` 應返回 `"u**r@domain.org"`。

# --hints--

你應定義一個名爲 `maskEmail` 的函數。

```js
assert.isFunction(maskEmail);
```

`maskEmail` 函數應接收一個字符串 `email` 作爲參數。

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

在函數外部，你應該有一個 `email` 變量。

```js
assert.isDefined(email);
```

你應爲 `email` 變量指定一個有效的電子郵件地址。

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` 應返回 `"a*******e@example.com"`。

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` 應返回 `"f**********p@example.com"`。

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` 應返回 `"i**o@test.dev"`。

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` 應返回 `"u**r@domain.org"`。

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

你的 `maskEmail` 應產生正確的結果。

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

你應記錄調用 `maskEmail` 的輸出結果，其參數爲 `email`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
