---
id: 66b205e6eacba4c4e54ea434
title: 建置一個 Email 掩碼器
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

在本實驗中，你將使用星號遮蔽電子郵件位址中的使用者名稱部分。遮蔽是一個用來隱藏或以星號或其他字元取代敏感資訊的術語。

例如，如果電子郵件位址是 `myEmail@email.com`，那麼遮罩後的電子郵件位址將是 `m*****l@email.com`。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 創建一個名為 `maskEmail` 的函式（程式），該函式（程式）以 `email` 作為引數。
2. 在函式中，你應該遮蔽 `email` 並將網域名稱附加到它。請記得你可以使用像是 `slice`、`repeat`、`indexOf` 或甚至 `replace` 這些方法來幫助你。
3. 在函式外，宣告一個名為 `email` 的變數來儲存你想要遮蔽的電子郵件位址。
4. 使用 `email` 變數呼叫 `maskEmail` 函式，並將結果輸出到主控台。
5. `maskEmail("apple.pie@example.com")` 應該傳回 `"a*******e@example.com"`。
6. `maskEmail("freecodecamp@example.com")` 應該傳回 `"f**********p@example.com"`。
7. `maskEmail("info@test.dev")` 應該傳回 `"i**o@test.dev"`。
8. `maskEmail("user@domain.org")` 應該傳回 `"u**r@domain.org"`。

# --hints--

你應該定義一個名為 `maskEmail` 的函式（程式）。

```js
assert.isFunction(maskEmail);
```

`maskEmail` 函式（程式）應該接受一個字串，`email` 作為引數。

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

在函式外，你應該有一個 `email` 變數。

```js
assert.isDefined(email);
```

你應該為你的 `email` 變數指定一個有效的電子郵件位址。

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` 應該傳回 `"a*******e@example.com"`。

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` 應該傳回 `"f**********p@example.com"`。

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` 應該傳回 `"i**o@test.dev"`。

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` 應該傳回 `"u**r@domain.org"`。

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

你的 `maskEmail` 應該產生正確的結果。

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

你應該記錄使用 `email` 作為引數呼叫 `maskEmail` 的輸出。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
