---
id: 68e529226b0defdff26c4a8b
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

现在你将开始创建折扣系统接口。Python 提供了一种定义抽象类的方法，这些抽象类作为其他类的蓝图。这是通过使用 `abc` 模块的 `ABC` 类（抽象基类）来完成的。

在你的代码顶部，添加以下导入语句：

```py
from abc import ABC, abstractmethod
```

这将导入 `ABC` 类和 `abstractmethod` 装饰器，你将用它们来定义必须由继承你的抽象类的任何类实现的方法。

# --hints--

你应该在你的代码顶部有提到的导入语句。

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
