---
id: 68e5293bd00d2fe134f58986
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

现在在你的 `PercentageDiscount` 中实现 `apply_discount` 方法。

此方法应计算并返回折扣价。

公式是：`product.price * (1 - self.percent / 100)`

# --hints--

你的 `PercentageDiscount` 类应该有一个 `apply_discount` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("apply_discount")`)
})
```

你的 `apply_discount` 方法应有参数 `self` 和 `product: Product`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_args("self, product: Product")`)
})
```

你的 `apply_discount` 方法应有一个 `float` 的返回类型提示。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_returns("float")`)
})
```

你的 `apply_discount` 方法应使用公式 `product.price * (1 - self.percent / 100)` 计算折后价格。

```js
({
  test: () => runPython(`
    test_product = Product("Test", 100.0)
    discount = PercentageDiscount(20)
    result = discount.apply_discount(test_product)
    assert result == 80.0
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
