---
id: 68e5293bd00d2fe134f58981
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

创建一个名为 `DiscountStrategy` 的类，该类继承自 `ABC`。这将是一个抽象的基类，定义所有折扣策略的接口。

目前，只需在类体中使用 `pass` 来创建一个空类。

# --hints--

你应该创建一个名为 `DiscountStrategy` 的类。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

你的 `DiscountStrategy` 类应继承自 `ABC`。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
