---
id: 68e5293bd00d2fe134f5898a
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

让我们测试你的新折扣策略。在现有代码之后，创建一个金额为 `5` 的 `FixedAmountDiscount` 实例，并将其赋值给名为 `fixed_discount` 的变量。然后将此折扣应用到产品上并打印结果。

# --hints--

你应该创建一个名为 `fixed_discount` 的变量，并赋值为一个 `FixedAmountDiscount` 实例。

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("fixed_discount")
    assert isinstance(fixed_discount, FixedAmountDiscount)
  `)
})
```

你的 `FixedAmountDiscount` 实例应初始化为金额 `5`。

```js
({
  test: () => runPython(`assert fixed_discount.amount == 5`)
})
```

你应该调用 `fixed_discount.apply_discount(product)` 并打印结果。

```js
({
  test: () => runPython(`
    assert _Node(_code).has_call("print(fixed_discount.apply_discount(product))")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
--fcc-editable-region--

--fcc-editable-region--
```
