---
id: 68e529226b0defdff26c4a8a
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

现在，创建一个名为 `Wireless Mouse` 且价格为 `50.0` 的 `Product` 实例。将其赋值给名为 `product` 的变量。然后，将 `product` 打印到控制台以查看其外观。

# --hints--

你应该有一个名为 `product` 的变量。

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

你应该将一个 `Product` 实例赋值给 `product` 变量。

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

你的 `Product` 实例应具有名称 `Wireless Mouse` 和价格 `50.0`。

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

你应该将 `product` 打印到控制台。

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

当你打印一个名称为 `'Wireless Mouse'` 且价格为 `50.0` 的 `Product` 时，它应显示为 `Wireless Mouse - $50.0`。

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
```
