---
id: 68e5293bd00d2fe134f58984
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

现在是创建你的抽象类的具体实现的时候了。创建一个名为 `PercentageDiscount` 的类，该类继承自 `DiscountStrategy`。

为它提供一个带有 `self` 和 `percent` 参数的 `__init__` 方法。为 `percent` 添加 `int` 类型提示，并将该方法的返回类型设置为 `None`。在 `__init__` 方法内部，将 `percent` 作为实例属性保存。

# --hints--

你应该创建一个名为 `PercentageDiscount` 的类。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

你的 `PercentageDiscount` 类应继承自 `DiscountStrategy`。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

你的 `PercentageDiscount` 类应该有一个 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

你的 `__init__` 方法应有参数 `self` 和 `percent: int`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

你的 `__init__` 方法应有 `None` 的返回类型提示。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

你应该在你的 `__init__` 方法中将 `percent` 赋值给 `self.percent`。

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
