---
id: 68e524c94e2e4ab7b4bfc7e2
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次研讨会中，你将构建一个折扣计算器，可以为产品应用不同的折扣策略。系统将根据多个折扣规则为客户确定最佳价格。

首先定义一个名为 `Product` 的类。为你的类添加一个带有以下参数的 `__init__` 方法：`self`、`name` 和 `price`。在 `__init__` 方法内部，将 `name` 和 `price` 参数赋值给同名的实例属性。

# --hints--

你应该创建一个名为 `Product` 的类。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

你的 `Product` 类应该有一个 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

你的 `__init__` 方法应包含以下参数：`self`、`name` 和 `price`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

你应该在你的 `__init__` 方法内将 `name` 赋值给 `self.name`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

你应该在你的 `__init__` 方法内将 `price` 赋值给 `self.price`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
