---
id: 68e529226b0defdff26c4a89
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

现在，如果你创建一个 `Product` 的实例并打印它，你会看到一个默认的对象表现，这并不是很有信息量。例如，它可能看起来像 `<__main__.Product object at 0x7f9b8c2d1d30>`。

要将默认对象表现转换为可读的字串，在你的 `Product` 类中创建一个 `__str__` 方法，该方法返回格式为 `name - $price` 的字串（其中 `name` 和 `price` 应替换为对应属性的值）。

# --hints--

你的 `Product` 类应该有一个 `__str__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

你的 `__str__` 方法应返回格式为 `name - $price` 的字串。

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
--fcc-editable-region--
    
--fcc-editable-region--
```
