---
id: 68e5293bd00d2fe134f58982
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

如前一课所述，抽象方法是任何继承自抽象基类的类必须实现的方法。它们作为一种契约，确保所有折扣策略都具有某些必需的方法。

在你的 `DiscountStrategy` 类中，定义一个名为 `is_applicable` 的方法。在方法定义上方使用 `@abstractmethod` 装饰器。该方法应具有以下参数：`self`、`product` 和 `user_tier`。在方法体中使用 `pass` 语句。

为你的方法签名添加类型提示以使其清晰：`product` 应为 `Product` 类型，`user_tier` 应为 `str` 类型。


# --hints--

你的 `DiscountStrategy` 类应有一个名为 `is_applicable` 的方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("is_applicable")`)
})
```

你的 `is_applicable` 方法应有 `@abstractmethod` 装饰器。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_decorators("abstractmethod")`)
})
```

你的 `is_applicable` 方法应包含参数 `self`、`product: Product` 和 `user_tier: str`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

--fcc-editable-region--
class DiscountStrategy(ABC):
    pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
