---
id: 68e5293bd00d2fe134f58987
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

现在你将创建另一个折扣策略。创建一个名为 `FixedAmountDiscount` 的类，该类继承自 `DiscountStrategy`。

为它提供一个带有 `self` 和 `amount` 参数的 `__init__` 方法。为 `amount` 参数添加 `int` 类型提示，为该方法添加 `None` 返回类型提示。在 `__init__` 方法内部，将 `amount` 作为实例属性保存。

# --hints--

你应该创建一个名为 `FixedAmountDiscount` 的类。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("FixedAmountDiscount")`)
})
```

你的 `FixedAmountDiscount` 类应继承自 `DiscountStrategy`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").inherits_from("DiscountStrategy")`)
})
```

你的 `FixedAmountDiscount` 类应有一个 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("__init__")`)
})
```

你的 `__init__` 方法应有参数 `self` 和 `amount: int`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_args("self, amount: int")`)
})
```

你的 `__init__` 方法的返回类型应为 `None`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_returns("None")`)
})
```

你应该在你的 `__init__` 方法内将 `amount` 赋值给 `self.amount`。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_stmt("self.amount = amount")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
