---
id: 68e52994bc7ea2e3dec5309d
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

现在让我们为高级用户创建第三个折扣策略。创建一个名为 `PremiumUserDiscount` 的类，该类继承自 `DiscountStrategy`。

与之前的策略不同，这个策略不需要 `__init__` 方法，因为它不保存任何配置。它将对所有高级用户应用固定的 20% 折扣。

目前，只需创建一个继承自 `DiscountStrategy` 的空类，使用 `pass`。

# --hints--

你应该创建一个名为 `PremiumUserDiscount` 的类。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PremiumUserDiscount")`)
})
```

你的 `PremiumUserDiscount` 类应继承自 `DiscountStrategy`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").inherits_from("DiscountStrategy")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
