---
id: 68e52994bc7ea2e3dec5309e
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

为 `PremiumUserDiscount` 实现 `is_applicable` 方法。此方法应查看用户是否为高级成员。

如果 `user_tier.lower()` 等于 `'premium'`，则返回 `True`，否则返回 `False`。使用 `.lower()` 确保检查不区分分支。

# --hints--

你的 `PremiumUserDiscount` 类应有一个 `is_applicable` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").has_function("is_applicable")`)
})
```

你的 `is_applicable` 方法应包含参数 `self`、`product: Product` 和 `user_tier: str`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").find_function("is_applicable").has_args("self, product:Product, user_tier:str")`)
})
```

你的 `is_applicable` 方法应具有 `-> bool` 的返回类型提示。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

当 `user_tier.lower()` 与 `'premium'` 相等时，你的 `is_applicable` 方法应返回 `True`。

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPremiumUserDiscount(PremiumUserDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPremiumUserDiscount()
    assert discount.is_applicable(test_product, "Premium") == True
    assert discount.is_applicable(test_product, "premium") == True
    assert discount.is_applicable(test_product, "PREMIUM") == True
  `)
})
```

当 `user_tier` 不是 `'premium'` 时，你的 `is_applicable` 方法应返回 `False`。

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPremiumUserDiscount(PremiumUserDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPremiumUserDiscount()
    assert discount.is_applicable(test_product, "regular") == False
    assert discount.is_applicable(test_product, "standard") == False
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

--fcc-editable-region--
class PremiumUserDiscount(DiscountStrategy):
    pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
