---
id: 68e79d91daf8fb4808f93c47
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

在实现 `apply_discount` 方法后，测试你的折扣策略。创建一个 `PercentageDiscount` 的实例，折扣为 10％，并赋值给名为 `discount` 的变量。然后使用 `discount.apply_discount(product)` 应用折扣，并打印结果以查看折扣效果。

# --hints--

你应该创建一个名为 `discount` 的变量，其值为一个 `PercentageDiscount` 实例。

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

你的 `PercentageDiscount` 实例应初始化为 `10` 百分比。

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

你应该调用 `discount.apply_discount(product)` 并打印结果。

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
    
    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--
```
