---
id: 68e9207ddcba053c5cb24413
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

在 Python 中，类型提示是可选的信号，用于告诉其他开发者变量或函数的预期数据类型。下面是一个为 `name` 参数添加类型提示 `str` 的示例。

```py
def greet(name: str):
    return 'Hello, ' + name
```

对于 `__init__` 方法参数，将 `name` 参数更新为类型提示 `str`，将 `price` 参数更新为类型提示 `float`。

# --hints--

你应该为 `Product` 类的 `__init__` 方法中的 `name` 参数添加类型提示 `str`，并为 `price` 参数添加类型提示 `float`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
