---
id: 699f49669ea51a31d2ba89d8
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

이전 수업에서 기억하듯이, 함수의 반환 타입을 명시적으로 지정할 수 있습니다:

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

이 예제에서 `circleArea` 함수의 반환 타입은 `number`입니다. 다른 타입을 반환하려 하면 TypeScript가 오류를 표시합니다.

`getAverage` 함수를 업데이트하여 반환 타입을 명시적으로 `number`로 지정하세요.

# --hints--

`getAverage` 함수는 반환 타입이 `number`여야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
