---
id: 699f458e47fb88c756fdba81
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

현재 `numbers` 매개변수는 암묵적으로 `any` 타입을 가지고 있습니다. 이를 고치려면, 매개변수를 숫자 배열로 명시적으로 타입 지정할 수 있습니다.

이전 수업에서 배운 대로, 숫자 배열을 이렇게 타입 지정할 수 있습니다:

```ts
const numbers: number[] = [1, 2, 3];
```

`getAverage` 함수를 업데이트하여 `numbers` 매개변수를 숫자 배열로 명시적으로 타입 지정하세요.

# --hints--

`getAverage` 함수는 숫자 배열로 명시적으로 타입 지정된 `numbers`라는 매개변수를 가져야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
