---
id: 699f4861dd5b5f164e1bda4a
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

TypeScript를 사용하면 반환 타입이 강제되고 예상한 대로인지 확인할 수 있다는 또 다른 장점이 있습니다. 때로는 TypeScript가 함수의 반환 타입을 추론할 수 있지만, 명시적으로 지정할 수도 있습니다.

이를 보여주기 위해 `getAverage` 함수의 반환 타입을 문자열 `"something"`로 변경하세요.

# --hints--

`getAverage` 함수는 문자열을 반환해야 합니다.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

`getAverage` 함수는 문자열 `"something"`를 반환해야 합니다.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
