---
id: 699f3cb8d057f29954205d7d
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이전 수업에서 기억하셨겠지만, 함수 매개변수에 다음과 같이 타입 주석을 추가할 수 있습니다:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

이 예제에서 TypeScript는 숫자가 아닌 인수로 `add` 함수를 호출하려 하면 오류를 표시합니다.

`square` 함수에 `num` 매개변수에 대한 타입 주석을 추가하세요. 타입은 `number`이어야 합니다.

# --hints--

`num` 함수의 `square` 매개변수에 타입 주석을 추가해야 합니다. 타입은 `number`이어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
