---
id: 699f412114c87221f16a99e3
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

TypeScript를 사용하면 함수 호출에 너무 많거나 너무 적은 인수를 전달할 때 오류를 잡아내는 또 다른 장점이 있습니다.

이를 보여주기 위해 `square` 함수 호출을 `5`과 `10` 두 개의 인수를 받도록 업데이트하세요.

# --hints--

`square` 함수 호출에는 `5`와 `10` 두 개의 인수가 있어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
