---
id: 699f4a3e7b4c6bfa538cc4f7
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

콘솔을 열면 다음과 같은 오류 메시지가 표시됩니다:

```md
Type 'string' is not assignable to type 'number'.
```

이 오류를 해결하려면 반환값을 `sum / numbers.length`로 변경해야 합니다.

# --hints--

`getAverage` 함수는 숫자를 반환해야 합니다.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

`getAverage` 함수는 배열에 있는 숫자들의 평균을 반환해야 합니다.

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```
