---
id: 699f41b8da20f816b5c7dece
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

콘솔을 열면 다음 오류 메시지가 표시될 것입니다:

```md
Expected 1 arguments, but got 2.
```

오류를 해결하려면 `square` 함수 호출이 `5` 하나의 인수만 받도록 수정하세요.

# --hints--

`square` 함수 호출은 `5` 하나의 인수만 가져야 합니다.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
