---
id: 699f3738921fb79a651ded56
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이제 함수를 호출할 차례입니다.

`result`라는 변수를 만들고 `square` 함수를 `5` 인수와 함께 호출한 결과를 할당하세요.

그 아래에 `result` 변수의 값을 콘솔에 출력하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`result`라는 변수가 있어야 합니다.

```js
assert.isDefined(result);
```

`result` 변수는 숫자여야 합니다.

```js
assert.isNumber(result);
```

`result` 변수는 `5`의 제곱이어야 합니다.

```js
assert.strictEqual(result, 25);
```

`result` 변수를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
