---
id: 699f513e01eb1768cc5df726
title: 17단계
challengeType: 1
dashedName: step-17
---

# --description--

이제 `exponent` 매개변수가 선택 사항으로 표시되었으니, 함수에 인수 하나만 전달해 호출할 수 있습니다.

`raiseTo` 함수 호출을 업데이트하여 숫자 `3`만 전달하세요.

# --hints--

`raiseTo` 함수는 `3` 인수 하나만 전달해 호출해야 합니다.

```js
assert.match(
  code,
  /raiseTo\s*\(\s*3\s*\)/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}

--fcc-editable-region--
const raisedResult = raiseTo(2, 3);
--fcc-editable-region--
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
