---
id: 699f3f636aa840b7201aa37c
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

콘솔을 열면 다음과 같은 TypeScript 오류 메시지가 보일 것입니다:

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

이제 매개변수가 명시적으로 숫자형으로 지정되어 있으므로, TypeScript가 `square` 함수에 문자열 인수를 전달하려는 오류를 잡아낼 수 있습니다.

오류를 고치고 콘솔에서 올바른 출력을 보려면 `square` 함수에 전달된 인수를 숫자 `5`로 변경하세요.

# --hints--

`square` 함수 호출에 `5` 대신 `"something"` 인수가 있어야 합니다.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
