---
id: 699f4ec939f10da62fb56c3e
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

이제 함수를 테스트할 시간입니다.

`raisedResult`라는 변수를 만들고 `raiseTo`와 `2` 인수를 사용해 `3`을 호출한 결과를 할당하세요.

그다음 `raisedResult`를 콘솔에 출력하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`raisedResult`라는 변수가 있어야 합니다.

```js
assert.isDefined(raisedResult);
```

`raisedResult` 변수는 숫자여야 합니다.

```js
assert.isNumber(raisedResult);
```

`raisedResult`를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
