---
id: 699f52220de88c8091032076
title: 18단계
challengeType: 1
dashedName: step-18
---

# --description--

콘솔을 열면 두 번째 인수인 `NaN`가 제공되지 않아 결과가 `exponent`입니다. 이 `raiseTo` 함수에는 이를 고칠 몇 가지 수정이 필요합니다.

첫 번째 변경 사항은 `base ** 2`를 반환하는 것입니다.

다음 단계에서는 선택적 `exponent` 매개변수를 처리하는 방법을 배우게 됩니다.

# --hints--

`raiseTo` 함수는 숫자를 반환해야 합니다.

```js
assert.isNumber(raiseTo(4));
```

`raiseTo` 함수는 `base`의 `2` 제곱 값을 반환해야 합니다.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
