---
id: 699f434e294a9e11cf4f65f0
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

이제 툴킷 프로젝트에 또 다른 수학 도우미를 추가할 시간입니다.

하나의 매개변수를 받는 `getAverage`라는 함수를 만드는 것부터 시작하세요. 이 함수는 숫자 배열의 평균을 반환해야 합니다.

다음 몇 단계에서 이 함수에 대한 타입 주석을 추가할 것입니다.

# --hints--

`getAverage`라는 함수가 있어야 합니다.

```js
assert.isFunction(getAverage);
```

함수는 하나의 매개변수를 받아야 합니다.

```js
assert.strictEqual(getAverage.length, 1);
```

함수는 숫자를 반환해야 합니다.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

함수는 어떤 숫자 배열에 대해서도 올바른 결과를 반환해야 합니다.

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
