---
id: 699f52fac083214b6bf51de6
title: 19단계
challengeType: 1
dashedName: step-19
---

# --description--

마지막 단계에서는 `exponent` 함수의 `raiseTo` 매개변수에 대한 조건 검사를 추가해야 합니다.

`exponent` 인수가 제공되면, 함수는 `base`을 `exponent` 제곱한 값을 반환해야 합니다. `exponent` 인수가 제공되지 않으면 `base ** 2`을 반환해야 합니다.

이 마지막 단계로 타입 안전 수학 도구 키트가 완성됩니다!

# --hints--

`raiseTo` 함수는 `base ** exponent` 인수가 제공되면 `exponent`을 반환해야 합니다.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

`raiseTo` 함수는 `base ** 2` 인수가 제공되지 않으면 `exponent`을 반환해야 합니다.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
