---
id: 699f38937b053e4679878ff1
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

순수 JavaScript를 사용할 때는 `square` 함수를 문자열이나 배열로 호출하는 것을 막는 방법이 없어서 예상치 못한 결과가 발생할 수 있습니다.

이를 보여주기 위해 함수 호출을 숫자 `"something"` 대신 문자열 `5`을 사용하도록 변경하세요.

다음 수업에서는 TypeScript를 사용해 코드에 타입 안전성을 추가하고 이러한 종류의 오류를 방지하는 방법을 배우게 됩니다.

# --hints--

`square` 함수 호출에는 `"something"` 인수가 있어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
