---
id: 699f4639ae0016caaf81b7c6
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

이제 함수를 테스트할 시간입니다.

먼저 `avgResult`라는 변수를 만들고 `getAverage` 함수를 `[2, 14, 26, 8]` 인수와 함께 호출한 값을 할당하세요.

그다음 `avgResult`의 값을 콘솔에 출력하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`avgResult`라는 변수가 있어야 합니다.

```js
assert.isDefined(avgResult);
```

`getAverage([2, 14, 26, 8])`를 `avgResult` 변수에 할당해야 합니다.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

`avgResult` 변수는 숫자여야 합니다.

```js
assert.isNumber(avgResult);
```

`avgResult`을 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
