---
id: 699f1d72e906e4616155290d
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 타입 주석과 함수와 함께 작동하는 방식을 계속 배우면서 타입 안전 수학 도구 키트를 만들어 보세요.

먼저 매개변수를 받아 숫자의 제곱을 반환하는 `square`라는 함수를 만드세요.

# --hints--

`square`라는 함수가 있어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

`square` 함수는 매개변수 하나를 가져야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

`square` 함수는 숫자의 제곱을 반환해야 합니다.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
