---
id: 699f4b903bf5427a599f791d
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

마지막 몇 단계에서는 함수에 대한 타입 주석 연습을 계속하기 위해 하나의 함수를 더 작성할 것입니다.

먼저 `raiseTo`라는 함수를 만들고 `base`과 `exponent`라는 두 매개변수를 받으세요. 두 매개변수 모두 `number` 타입이어야 합니다. 이 함수는 `base`를 `exponent` 거듭제곱한 결과를 반환해야 합니다. 함수의 반환 타입도 `number`로 주석을 달아야 합니다.

# --hints--

`raiseTo` 함수가 있어야 합니다.

```js
assert.isFunction(raiseTo);
```

`raiseTo` 함수는 두 개의 매개변수 `base`와 `exponent`을 가져야 합니다.

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*=>/
);
```

`base`과 `exponent` 매개변수는 모두 `number` 타입이어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

`raiseTo` 함수는 숫자를 반환해야 합니다.

```js
assert.isNumber(raiseTo(2, 3));
```

`raiseTo` 함수는 `base`를 `exponent` 거듭제곱한 결과를 반환해야 합니다.

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

`raiseTo` 함수는 `number` 반환 타입 주석이 있어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
