---
id: 699f5063d2393c674b1f45dc
title: 16단계
challengeType: 1
dashedName: step-16
---

# --description--

이전 단계에서 기억하겠지만, 함수에 잘못된 개수의 인수를 전달하면 TypeScript가 오류를 발생시킵니다.

매개변수 이름 뒤에 `?`를 추가하면 매개변수를 선택 사항으로 표시할 수 있습니다. 이렇게 하면 매개변수보다 적은 인수로 함수를 호출할 수 있습니다.

예를 들어 다음과 같습니다:

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

이 예제에서 `greeting` 함수를 호출할 때 `greet` 매개변수가 제공되지 않으면 `undefined`가 되고, 함수는 기본 인사말을 반환합니다.

`raiseTo` 함수를 업데이트하여 `exponent` 매개변수를 선택 사항으로 만드세요.

# --hints--

`raiseTo` 함수에서 `exponent` 매개변수가 선택 사항으로 표시되어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
