---
id: 66db57ad34c7089b9b41bfd6
title: Build a Lightbox Viewer
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

ライトボックスは、ウェブサイトでマルチメディアコンテンツをページのメインコンテンツの上にポップアップやモーダルウィンドウとして表示し、より魅力的に見せるために使われます。

このラボでは、サムネイルをクリックするとフルサイズの画像を表示するライトボックスギャラリーを作成します。各画像には、サムネイルとフルサイズの2つのバージョンが用意されています。フルサイズの画像はサムネイルと同じですが、`-thumbnail` のサフィックスがありません。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了させてください。

**User Stories:**

1. `div` の中に `gallery` があり、そのクラスは `body` である必要があります。
1. `.gallery` 要素の中に、クラスが `gallery-item` のサムネイル画像が3つある必要があります。サムネイル画像には以下のリンクを使用してください：
   
    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. `div` の中にクラスが `lightbox` の `body` がある必要があります。
1. `button` 要素の中に `id` があり、その `close-btn` は `.lightbox` である必要があります。テキストは `&times;` を使ってもかまいません。
1. `img` 要素の中に `id` があり、その `lightbox-image` は `.lightbox` である必要があります。
1. `.lightbox` 要素は固定位置に設定されていて、プレビューが現在の画像の上に開くようにしてください。
1. `.lightbox` 要素はコンテナの高さと幅を100%に設定してビューポート全体を覆うようにしてください。`.lightbox` 要素はコンテナの左上隅から始まるようにしてください。
1. `.lightbox` は背景色を持ち、最初は `display` プロパティが `none` に設定されて非表示にしてください。
1. `.gallery-item` のいずれかをクリックすると、`.lightbox` 要素の `display` プロパティが `flex` に設定され、`.lightbox` 要素とその中の2つの要素が表示されるようにしてください。
1. `.gallery-item` のいずれかをクリックすると、`#lightbox-image` 要素の `src` がクリックされた画像の `-thumbnail` 属性から `src` を取り除いたフルサイズの画像に設定されます。フルサイズ画像は以下のリンクにあります：

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`
  
1. `.lightbox` 要素が表示されているときに `#close-btn` または `.lightbox` 要素をクリックすると、`.lightbox` 要素の `display` が `none` に戻されます。

**注意:** スタイルシートとJavaScriptファイルをHTMLにリンクするのを忘れないでください。

# --hints--

`div` 要素の中にクラスが `gallery` の `body` があるはずです。

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

`.gallery` 要素の中にクラスが `img` の `gallery-item` 要素が3つあるはずです。

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

`.gallery` 要素の中に `img` 要素があり、その `src` は `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg` に設定されているはずです。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

`.gallery` 要素の中に `img` 要素があり、その `src` は `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg` に設定されているはずです。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

`.gallery` 要素の中に `img` 要素があり、その `src` は `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg` に設定されているはずです。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

`div` 要素の中にクラスが `lightbox` の `body` 要素があるはずです。

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

`.lightbox` 要素の中に `button` 要素があり、その `id` は `close-btn` に設定されているはずです。

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

`.lightbox` 要素の中に `img` 要素があり、その `id` は `lightbox-image` に設定されているはずです。

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

`.lightbox` 要素は固定位置に設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

`.lightbox` 要素はビューポート全体を覆うはずです。

```js
const style = new __helpers.CSSHelp(document).getStyle(".lightbox");

const width = style?.width;
const height = style?.height;

assert.oneOf(
  width,
  ["100%", "100vw"]
);

assert.oneOf(
  height,
  ["100%", "100vh"]
);
```

`.lightbox` 要素はコンテナの左上隅に揃っているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

`.lightbox` 要素は背景色を持っているはずです。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

`.lightbox` 要素は最初は非表示になっているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

`.gallery-item` のいずれかをクリックすると、`.lightbox` 要素の `display` プロパティが `flex` に設定され、`.lightbox` とその中の2つの要素が表示されるはずです。

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

`.gallery-item` のいずれかをクリックすると、`#lightbox-image` 要素の `src` がクリックされた画像の `-thumbnail` 属性から `src` を取り除いたフルサイズの画像に設定されるはずです。

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

`.lightbox` 要素が表示されているときに `#close-btn` ボタンをクリックすると、`.lightbox` 要素の `display` が `none` に戻されるはずです。

```js
const lightbox = document.querySelector(".lightbox");
const closeBtn = document.getElementById("close-btn"); 
const galleryItem = document.querySelector(".gallery-item");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

closeBtn.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

`.lightbox` 要素が表示されているときに `.lightbox` 要素をクリックすると、`.lightbox` 要素の `display` が `none` に戻されるはずです。

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

lightbox.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```
