---
id: 657efce98e958b75df86b305
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

`if` 문장 본문 안에서, 변환된 문자를 이전에 만든 리스트에 추가할 것입니다.

이를 위해 `.append()` 메서드를 사용할 것입니다. 이 메서드는 호출된 리스트의 끝에 주어진 객체를 추가합니다.

`.append()`에서 `snake_cased_char_list` 메서드를 사용해 `converted_character`를 리스트에 추가하세요.

# --hints--

변환된 문자를 `.append()`에 추가하려면 `snake_cased_char_list` 메서드를 사용하세요.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\s*\(\s*converted_character\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
