---
id: 657e928716b77b2277980276
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 프로젝트에서는 `camelCase` 또는 `PascalCase` 형식의 문자열을 입력받아 `snake_case` 형식의 문자열로 변환하는 프로그램을 두 가지 방법으로 만들 것입니다. 먼저 for 루프를 사용하고, 그다음에는 리스트 내포를 사용하여 같은 결과를 얻을 것입니다. 리스트 내포가 코드를 더 간결하게 만드는 방법을 확인할 수 있습니다.

`convert_to_snake_case()`라는 문자열을 입력으로 받는 `pascal_or_camel_cased_string`라는 새 함수를 정의하기 시작하세요. 우선 함수 안에 `pass` 문장을 추가하세요.

# --hints--

`convert_to_snake_case()`라는 새 함수를 정의해야 합니다. 끝에 콜론을 잊지 말고 함수 본문을 채우기 위해 `pass`를 사용하세요.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

함수는 `pascal_or_camel_cased_string`라는 매개변수를 받아야 합니다.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
