---
id: 657effaa2a5e0277d71f9cbe
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

이 시점에서 변수 `snake_cased_char_list`는 변환된 문자들의 리스트를 담고 있습니다. 이 문자들을 하나의 문자열로 합치려면 `.join()` 메서드를 사용할 수 있습니다.

`join` 메서드는 리스트의 각 요소를 지정된 문자열인 구분자를 사이에 두고 연결하여 하나의 문자열로 만듭니다.

```py
result_string = ''.join(characters)
```

위 예제는 `characters` 리스트의 요소들을 빈 문자열을 구분자로 사용해 하나의 문자열로 합칩니다.

이제 `for` 반복문 바로 다음에 `.join()` 메서드를 사용해 `snake_cased_char_list`의 요소들을 빈 문자열을 구분자로 합치세요. 결과를 `snake_cased_string`라는 새 변수에 할당하세요.


# --hints--

`snake_cased_char_list` 리스트의 문자들을 `.join()` 메서드를 사용해 문자열로 합치고, `snake_cased_string` 반복문 뒤에 `for`라는 변수에 할당하세요.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
