---
id: 657ef2a86d4e545cec9a85fb
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

camel case와 pascal case 모두에서 대문자는 새 단어의 시작을 표시합니다. 입력 문자열을 snake case로 변환하려면 입력 문자열의 문자가 대문자인지 확인해야 합니다.

문자가 대문자인지 확인하려면 `.isupper()` 문자열 메서드를 사용할 수 있습니다. 이 메서드는 문자가 대문자면 `True`를, 그렇지 않으면 `False`를 반환합니다.

`for` 반복문 안에 현재 문자가 대문자인지 확인하는 `if` 문을 추가하세요. `pass` 문을 새 `if` 문 안으로 옮기세요.

# --hints--

조건으로 char.isupper()를 사용하는 새로운 if 문을 작성해야 합니다. 끝에 콜론을 추가하고 if 문 본문을 채우기 위해 pass 키워드를 사용하세요.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

반복문 본문에 있는 `pass` 문을 `if` 문으로 바꿔야 합니다.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
