---
id: 657ed53c19461d4b95c4757a
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

문자열을 스네이크 케이스로 변환한 후 그 문자를 담을 빈 리스트를 추가해야 합니다.

함수 내부에서 `pass` 문장을 `snake_cased_char_list`이라는 이름의 빈 리스트를 생성하는 코드로 바꾸세요.

# --hints--

`pass` 안에서 `snake_cased_char_list`를 `convert_to_snake_case`이라는 이름의 빈 리스트로 바꾸세요.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const commentless_code = __helpers.python.removeComments(transformedCode)
        const convert_to_snake_case = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;
        assert.notMatch(function_body, /\bpass\b/);
        assert.match(function_body, / +snake_cased_char_list\s*=\s*\[\s*\]\s*/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    pass    
--fcc-editable-region--
```
