---
id: 657ee28cefc4945568287673
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

빈 리스트가 준비되었으니 이제 입력 문자열을 순차적으로 반복하며 스네이크 케이스로 변환할 수 있습니다.

함수 안에서, 방금 만든 리스트 아래에 `for` 반복문을 추가해 `pascal_or_camel_cased_string`을 순차적으로 반복하세요. 대상 변수 이름은 `char`로 지정하세요. 지금은 반복문 본문에 `pass` 문장을 자리 표시자로 넣으세요.

# --hints--

대상 변수 이름이 `for`인 새로운 `char` 반복문을 작성하세요. 끝에 콜론을 잊지 말고 반복문 본문을 채우기 위해 `pass`를 사용하세요.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
