---
id: 657efa642593c5746acc5c81
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

`if` 문장 본문 안에서 대문자를 소문자로 변환하고 이 소문자 앞에 밑줄을 붙여야 합니다.

대문자를 소문자로 변환하려면 `.lower()` 문자열 메서드를 사용하세요. 그런 다음 문자 앞에 밑줄을 붙이세요. 결과를 `converted_character`라는 변수에 할당하세요.

# --hints--

수정한 문자를 `converted_character`라는 변수에 할당해야 합니다.

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

`pass` 문장 안에 `if`가 없어야 합니다.

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
