---
id: 657f43d341a0dd17120cdb08
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

대문자를 소문자로 변환하고 그 앞에 밑줄을 추가해야 합니다.

리스트 내포 작업을 진행하기 전에, 함수에 반환값을 지정하세요. 이렇게 하면 출력 결과를 확인할 수 있습니다.

`return` 문장을 사용해 함수에서 `snake_cased_char_list` 리스트를 반환하세요.

# --hints--

`snake_cased_char_list` 리스트를 반환해야 합니다. 들여쓰기가 올바른지 확인하세요.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*snake_cased_char_list/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--

--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
