---
id: 657f425dbab54e11993c80f0
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

지금까지 이 프로젝트에서 `for` 반복문을 사용해 입력 문자열을 순회하며 원하는 출력으로 변환했습니다. 이제 `for` 반복문에서 리스트 내포로 전환을 시작할 것입니다.

먼저 `convert_to_snake_case()` 함수 안의 모든 코드를 주석 처리하세요. 리스트 내포를 사용해 로직을 구현할 때 도움이 되니 삭제하지 마세요.

테스트 중 코드가 실패하지 않도록 함수 본문에 `pass` 키워드를 추가하는 것도 잊지 마세요.

# --hints--

`convert_to_snake_case()` 함수 안의 10줄 코드를 모두 주석 처리하고 다른 부분은 건드리지 마세요. 오류를 피하려면 함수 본문에 임시로 `pass`를 추가하세요.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
