---
id: 657f4add33ea4b1f61ba3dc8
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

마지막으로 프로그램을 실행하여 이 새로운 구현을 시험해 보세요. 입력 문자열을 `'IAmAPascalCasedString'`로 바꾸고 `'i_am_a_pascal_cased_string'`로 출력되는지 확인하세요. 비록 그게 거짓말이라 해도요.

모든 단계를 올바르게 수행했다면, 이전과 같이 입력 문자열이 스네이크 케이스로 변환된 것을 볼 수 있을 겁니다.

축하합니다! 이제 `convert_to_snake_case()` 함수가 준비되었습니다.

# --hints--

입력 문자열을 `'aLongAndComplexString'`에서 `'IAmAPascalCasedString'`로 `main()` 함수 안에서 바꿔야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
