---
id: 657f04ed0035f47ed04d0f1f
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

`convert_to_snake_case()` 함수의 출력을 표시하려면 `main()` 함수를 호출해야 합니다.

기존 두 함수와 같은 수준에 `main()` 함수 호출을 추가하세요. 실행하면 주어진 카멜 또는 파스칼 표기 문자열이 스네이크 케이스로 변환된 것을 볼 수 있습니다.

# --hints--

`main()` 함수 호출을 추가해야 합니다.

```js
({
  test: () => {
    const pythonCode = `_Node(_code.replace('\\\\r', '')).has_call("main()")`;
    const result = runPython(pythonCode);
    assert(result);
  }
});
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

--fcc-editable-region--
```
