---
id: 657f01ae9aea647b27402d3e
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

함수를 마무리하려면 `clean_snake_cased_string`를 반환하세요. 이렇게 하면 함수가 완성되어 pascal 또는 camel case 문자열을 snake case로 변환하는 데 사용할 수 있습니다.

함수 끝에 `return` 문을 추가하여 `clean_snake_cased_string`를 반환하세요.


# --hints--

`return clean_snake_cased_string` 함수 끝에 `convert_to_snake_case`를 추가해야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +return\s+clean_snake_cased_string/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')
--fcc-editable-region--


--fcc-editable-region--
```
