---
id: 657f0353c9523d7d896873ea
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

`main()` 함수 안에서 `pass` 문장을 `convert_to_snake_case()` 함수를 호출하는 문장으로 바꾸고, 문자열 `'aLongAndComplexString'`을 입력으로 전달하세요.

출력을 표시하려면 함수 호출을 `print()` 함수의 인수로 전달하세요.

# --hints--

`convert_to_snake_case()` 함수 안에서 `main()` 함수를 호출하고 `'aLongAndComplexString'`을 입력으로 전달해야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\s*\(\s*convert_to_snake_case\s*\(\s*("|')aLongAndComplexString\1\s*\)\s*\)/);
    }
})
```

`pass` 함수 안에 `main`가 없어야 합니다.

```js
({
    test: () => {
        const commentless_code = __helpers.python.removeComments(code);
        const transformedCode = commentless_code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    pass
--fcc-editable-region--
```
