---
id: 657f025ec86c3d7c4177b6be
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

함수가 완성되었으니 이제 다른 함수 안에서 사용할 수 있습니다.

`main()`라는 새 함수를 만들고 `pass`을 함수 본문으로 작성하세요.

# --hints--

`main()`라는 새 함수를 정의해야 합니다. 끝에 콜론을 잊지 말고 함수 본문을 채우기 위해 `pass`을 사용하세요.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(main)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--


--fcc-editable-region--
```
