---
id: 657f456223b8c1187b461987
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

Python에서 <dfn>리스트 내포</dfn>는 기존 반복 가능한 항목 각각에 표현식을 적용하고 선택적으로 조건으로 항목을 필터링하여 새 리스트를 생성할 수 있는 구문입니다. 간결할 뿐만 아니라 리스트 내포는 종종 더 빠르게 실행됩니다.

기본 리스트 내포는 표현식 다음에 `for` 절이 옵니다:

```py
spam = [i * 2 for i in iterable]
```

위 예제는 변수 `i`를 사용해 `iterable`를 순회합니다. 결과 리스트의 각 요소는 현재 반복에서 표현식 `i * 2`를 평가해 얻습니다.

이번 단계에서는 빈 리스트 `snake_cased_char_list`를 리스트 내포 구문으로 채우세요.

빈 리스트를 리스트 내포로 바꿔 `pascal_or_camel_cased_string`의 각 문자를 소문자로 변환하고 앞에 밑줄을 붙이도록 하세요(이전에 주석 처리한 코드가 표현식을 작성하는 데 도움이 될 수 있습니다). `char`를 사용해 `pascal_or_camel_cased_string`를 순회하세요.

# --hints--

`snake_cased_char_list`를 `pascal_or_camel_cased_string`를 순회하는 리스트 내포로 바꾸세요.

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

리스트 내포는 `char`를 사용해 `pascal_or_camel_cased_string`를 순회해야 합니다.

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

리스트 내포는 `'_' + char.lower()`의 각 `char`에 대해 `pascal_or_camel_cased_string`를 평가해야 합니다.

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
