---
id: 657f465f8e718b19c5105ae5
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

리스트 내포는 조건문을 받아 특정 조건이 충족될 때만 주어진 표현식을 평가합니다:

```py
spam = [i * 2 for i in iterable if i > 0]
```

출력에서 볼 수 있듯이, `pascal_or_camel_cased_string`에서 생성된 문자 목록이 연결되었습니다. 리스트 내포 안의 표현식이 각 문자마다 평가되기 때문에, 결과는 모든 문자가 밑줄로 구분된 소문자 문자열입니다.

위 예제를 따라 리스트 내포에 `if` 절을 추가하여 문자가 대문자인 경우에만 표현식이 실행되도록 하세요.

# --hints--

리스트 내포에 `if` 절을 `char.isupper()` 조건과 함께 추가하세요.

```js
({ 
    test: () => assert(runPython(`
    ifs = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_ifs()
    len(ifs) == 1 and ifs[0].is_equivalent("char.isupper()")
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
