---
id: 657f47b12c51e41b3149e584
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

그래도 최종 결과가 정확히 원하는 대로 나오지 않습니다. `if` 절에서 걸러진 문자에 대해 다른 표현식을 실행해야 합니다. 이를 위해 `else` 절을 사용할 것입니다:

```py
spam = [i * 2 if i > 0 else -1 for i in iterable]
```

`if` 절과 달리 `if`/`else` 구문은 표현식과 `for` 키워드 사이에 위치해야 한다는 점을 참고하세요.

대문자가 아닌 문자는 변경하지 않고 그대로 두도록 리스트 내포를 수정하세요.

# --hints--

리스트 내포를 수정하여 `'_' + char.lower()`일 때 `char.isupper()` 표현식을 평가하고, 그렇지 않으면 char를 반환하도록 하세요.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_expr().is_equivalent("'_' + char.lower() if char.isupper() else char")
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string if char.isupper()]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
