---
id: 657f0044be09db790b1eb1c5
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

파스칼 케이스에서는 문자열이 대문자로 시작합니다. 모든 문자를 소문자로 변환하고 밑줄을 추가한 후에는 문자열 시작 부분에 불필요한 밑줄이 생길 수 있습니다.

이 문제를 가장 쉽게 고치는 방법은 `.strip()` 문자열 메서드를 사용하는 것입니다. 이 메서드는 인수로 전달된 문자 집합 중 문자열의 앞이나 뒤에 있는 문자를 제거합니다. 예를 들어:

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

`strip()` 메서드는 `original_string`에 적용됩니다. 이 메서드는 문자열의 앞뒤에 있는 밑줄을 제거합니다. 위 예시의 결과는 `'example_string'` 문자열이 됩니다.

`clean_snake_cased_string`라는 새 변수를 선언하고, `.strip()`에 `snake_cased_string` 메서드를 적용한 결과를 할당하세요. 이때 메서드의 인수로 `'_'`를 전달합니다.

# --hints--

`snake_cased_string.strip('_')`의 결과 문자열을 `clean_snake_cased_string`라는 변수에 할당해야 합니다.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
