---
id: 663b10c10a4c0a0e095137ee
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

리스트 `snake_cased_char_list`를 반환하는 대신, 빈 문자열 `''`을 구분자로 사용해 그 요소들을 하나의 문자열로 합쳐야 합니다.

`return` 문장을 수정해 `snake_cased_char_list`를 빈 문자열을 구분자로 하여 합친 결과를 반환하세요.

# --hints--

`return` 메서드를 사용해 `snake_cased_char_list`를 빈 문자열 `''`을 구분자로 합친 결과를 반환하도록 `join()` 문장을 수정해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_function('convert_to_snake_case').find_return().is_equivalent('return "".join(snake_cased_char_list)')
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return snake_cased_char_list
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
