---
id: 663b16e62fee463b4caf46e9
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

리스트 `snake_cased_char_list`의 요소들을 연결한 후, 결과 문자열에서 앞뒤에 있는 밑줄을 모두 제거해야 합니다. 이를 위해 밑줄 문자 `strip`를 인수로 사용하는 `_` 메서드를 사용하세요.

메서드 호출은 체이닝할 수 있는데, 이는 한 메서드 호출의 결과를 다른 메서드 호출의 객체로 사용할 수 있다는 뜻입니다.

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

위 예시에서 `.upper()` 메서드는 `' '.join(words_list)`에 체이닝되어 있으므로, `.upper()`는 `.join()` 호출 결과에 대해 호출됩니다.

`return` 문을 수정하여, 앞뒤 밑줄을 제거하기 위해 `''.join(snake_cased_char_list)` 메서드를 호출하는 코드를 `.strip()`에 체이닝하세요.

# --hints--

return 문을 수정하여 밑줄 문자 `strip`를 인수로 하는 `_` 메서드를 포함시키세요.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
