---
id: 646c59e942f35541923104bf
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

回到你的 HTML 中，创建一个 `main` 元素。 在 `main` 元素中，添加一个 `div` 元素，类为 `cat-head`。

# --hints--

你应该有一个 `main` 元素。

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

你应该有一个 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

你的 `div` 元素应该有类 `cat-head`。

```js
assert.equal(document.querySelector('div')?.getAttribute('class'), 'cat-head');
```

你的 `div` 元素应该在 `main` 标签内。

```js
assert.lengthOf(document.querySelectorAll('main div'), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
