---
id: 646c48df8674cf2b91020ecc
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

基本的模板已为你提供。 首先在你的 `head` 元素中添加一个 `link` 元素。 对于该 `link` 元素，将 `rel` 属性设置为 `stylesheet`，`href` 设置为 `./styles.css`。

# --hints--

你应该有一个 `link` 元素。

```js
assert.exists(document.querySelector("head>link"));
```

你的 `link` 元素应该有一个 `rel` 和 `href` 属性，分别设置为 `stylesheet` 和 `./styles.css`。

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
