---
id: 6476f7a4827bcc61682f2347
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

当你将 `position` 属性设置为 `absolute`时，元素就会脱离文件的正常流，它的位置由 `top`、`right`、`bottom` 和 `left` 属性决定。

将你的 `.cat-head` 元素的 `position` 属性设置为 `absolute`。

# --hints--

你的 `.cat-head` 选择器应将 `position` 属性设置为 `absolute`。 确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

你的 `.cat-head` 选择器应将 `top` 属性设置为以 `px` 结尾的任意数字。 确保你添加了分号。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/);
```

你的 `.cat-head` 选择器应将 `left` 属性设置为以 `px` 结尾的任意数字。 确保你添加了分号。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>fCC Cat Painting</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class="cat-head"></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
--fcc-editable-region--
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
