---
id: 69708d59833ceded7d00541f
title: Operationen für verkettete Listen implementieren
challengeType: 26
dashedName: implement-linked-list-operations
---

# --description--

In diesem Labor implementieren Sie zusätzliche Operationen für eine verkettete Liste, basierend auf der Grundfunktionalität zum Hinzufügen und Entfernen von Knoten.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `contains`-Funktion haben, die eine verkettete Liste und ein Element entgegennimmt. Sie soll `true` zurückgeben, wenn das angegebene Element in der verketteten Liste existiert, und `false` sonst.
1. Sie sollten eine `getAt`-Funktion haben, die eine verkettete Liste und einen Index entgegennimmt. Sie soll das Element an dem angegebenen Index in der verketteten Liste zurückgeben. Wenn der Index außerhalb des gültigen Bereichs liegt, soll sie `undefined` zurückgeben.
1. Sie sollten eine `insertAt`-Funktion haben, die eine verkettete Liste, einen Index und ein Element entgegennimmt. Sie soll das angegebene Element an der spezifizierten Position in der verketteten Liste einfügen. Wenn der Index außerhalb des gültigen Bereichs liegt, soll die Liste nicht verändert werden.
1. Sie sollten eine `removeAt`-Funktion haben, die eine verkettete Liste und einen Index entgegennimmt. Sie soll den Knoten an dem angegebenen Index in der verketteten Liste entfernen. Wenn der Index außerhalb des gültigen Bereichs liegt, soll die Liste nicht verändert werden.
1. Sie sollten eine `clear`-Funktion haben, die eine verkettete Liste entgegennimmt. Sie soll alle Elemente aus der verketteten Liste entfernen und sie effektiv in einen leeren Zustand zurücksetzen.

**Hinweis:** Einige spätere Tests bauen auf früheren Methoden auf. Wenn zum Beispiel `getAt` nicht korrekt implementiert ist, können Tests für Funktionen wie `insertAt` und `removeAt` fehlschlagen, selbst wenn diese Funktionen fast korrekt sind.

# --hints--

`initList` sollte ein Objekt zurückgeben, bei dem `head` auf `null` gesetzt ist und `length` auf `0`.

```js
const list = initList();
assert.isNull(list.head);
assert.strictEqual(list.length, 0);
```

`add` sollte die Länge der Liste um eins erhöhen.

```js
const list = initList();
add(list, 'apple');
assert.strictEqual(list.length, 1);
```

`remove` sollte die Länge der Liste um eins verringern.

```js
const list = initList();  
add(list, 'apple');  
add(list, 'banana');  
remove(list, 'apple');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(list.head.element, 'banana');  
assert.isNull(list.head.next);
```

Sie sollten eine `contains`-Funktion haben.

```js
assert.isFunction(contains);
```

`contains` sollte `true` zurückgeben, wenn das Element in der Liste existiert.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

`contains` sollte `false` zurückgeben, wenn das Element nicht in der Liste existiert.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

`contains` sollte `false` für eine leere Liste zurückgeben.

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

Sie sollten eine `getAt`-Funktion haben.

```js
assert.isFunction(getAt);
```

`getAt` sollte das Element an einem gegebenen Index zurückgeben.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

`getAt` sollte `undefined` für ungültige Indizes zurückgeben.

```js
const list = initList();  
add(list, 'apple');  
assert.isUndefined(getAt(list, -1));  
assert.isUndefined(getAt(list, 1));  
assert.isUndefined(getAt(list, 5));
```

`getAt` sollte `undefined` für eine leere Liste zurückgeben.

```js
const list = initList();  
assert.isUndefined(getAt(list, 0));
```

Sie sollten eine `insertAt`-Funktion haben.

```js
assert.isFunction(insertAt);
```

`insertAt` sollte ein Element am Anfang der Liste einfügen.

```js
const list = initList();  
add(list, 'second');  
add(list, 'third');  
insertAt(list, 0, 'first');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` sollte ein Element in der Mitte der Liste einfügen.

```js
const list = initList();  
add(list, 'first');  
add(list, 'third');  
insertAt(list, 1, 'second');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` sollte ein Element am Ende der Liste einfügen.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` sollte in eine leere Liste an Index `0` einfügen.

```js  
const list = initList();  
insertAt(list, 0, 'first');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.isNull(list.head.next); 
```

`insertAt` sollte die Liste bei ungültigen Indizes nicht verändern.

```js
const list = initList();  
add(list, 'apple');  
insertAt(list, -1, 'invalid');  
insertAt(list, 10, 'invalid');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
assert.isUndefined(getAt(list, 1));
```

Sie sollten eine `removeAt`-Funktion haben.

```js
assert.isFunction(removeAt);
```

`removeAt` sollte das Element an einem angegebenen Index entfernen.

```js
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'third');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` sollte das erste Element korrekt entfernen.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
assert.isNull(list.head.next);
```

`removeAt` sollte das letzte Element korrekt entfernen.

```js  
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 2);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` sollte die Liste zurücksetzen, wenn das einzige Element entfernt wird.

```js
const list = initList();  
add(list, 'only');  
removeAt(list, 0);  
assert.strictEqual(list.length, 0);  
assert.isNull(list.head);  
assert.isTrue(isEmpty(list));
```

`removeAt` sollte die Liste bei ungültigen Indizes nicht verändern.

```js
const list = initList();  
add(list, 'apple');  
removeAt(list, -1);  
removeAt(list, 10);  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
```

Sie sollten eine `clear`-Funktion haben.

```js
assert.isFunction(clear);
```

`clear` sollte alle Elemente aus der Liste entfernen.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

`clear` sollte die Wiederverwendung der Liste ermöglichen.

```js  
const list = initList();  
add(list, 'first');  
clear(list);  
clear(list);  
add(list, 'second');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'second');  
assert.isNull(list.head.next); 
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  
}

function getAt(list, index) {
 
}

function insertAt(list, index, element) {
  
}

function removeAt(list, index) {
 
}

function clear(list) {

}
```

# --solutions--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  let current = list.head;
  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }
  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;
  let current = list.head;
  let i = 0;
  while (i < index) {
    current = current.next;
    i++;
  }
  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;
  const node = { element, next: null };
  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;
  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
