---
id: 615f34ecc1091b4fd5a8a484
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Within your `head` element, add a `link` element with the `rel` attribute set to `stylesheet` and the `href` attribute set to `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Isto importará a família de fontes `Open Sans`, com os valores de peso da fonte `400`, `700`, e `800`.

Adicione também um elemento `link` para vincular o arquivo `styles.css`.

# --hints--

O código deve ter dois elementos `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Seus elementos `link` devem ser elementos vazios, ou seja, não devem ter uma tag de fechamento `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Os dois elementos `link` devem estar dentro do elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Both of your `link` elements should have the `rel` attribute set to `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Um dos elementos `link` deve ter o atributo `href` com o valor de `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

One of your `link` elements should have an `href` attribute set to `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
