---
id: 5cd9a70215d3c4e65518328f
title: 카운트다운 만들기
challengeType: 26
dashedName: build-a-countdown
---

# --description--

이 실습에서는 시작 숫자를 받아 카운트다운 목록을 생성하는 `countdown` 함수를 만드세요.

`countdown(n)`을 호출하면 `n`에서 시작해 `1`씩 감소하며 `1`에 도달할 때까지 카운트다운하는 배열을 반환해야 합니다.

예를 들어:

- `countdown(5)`는 `[5, 4, 3, 2, 1]`를 반환해야 합니다.
- `countdown(1)`는 `[1]`를 반환해야 합니다.

시작 숫자가 `1`보다 작으면 카운트다운할 것이 없으므로 빈 배열을 반환하세요:

- `countdown(0)`는 `[]`를 반환해야 합니다.
- `countdown(-3)`는 `[]`를 반환해야 합니다.

실습을 완료하려면 재귀를 사용해 결과를 만들어야 하며, 루프(`for`, `while`), `Array.from()`, 또는 배열 반복 메서드(`forEach`, `map`, `filter`, `reduce`)를 사용하면 안 됩니다. 각 호출은 자체 결과 배열을 생성해야 하며(전역 변수를 사용해 카운트다운을 저장하지 마세요).

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `countdown`라는 이름의 함수를 만들어야 합니다.
2. `countdown` 함수는 하나의 인수 `n` (숫자)을 받아야 합니다.
3. `n`가 `1`보다 작으면 `countdown`는 빈 배열을 반환해야 합니다.
4. 그렇지 않으면 `countdown`는 `n`부터 `1`까지 내림차순으로 정수 배열을 반환해야 합니다.
5. 재귀를 사용해야 하며(함수가 자기 자신을 호출해야 함), 어떤 종류의 루프(`for`, `while`, `Array.from()`)나 고차 메서드(`forEach`, `map`, `filter`, `reduce`)를 사용하면 안 됩니다.
6. 각 재귀 호출은 이전 호출보다 작은 값을 사용해야 하며(예: `n - 1`), 함수가 기본 사례에 도달하도록 해야 합니다.
7. 결과 배열을 저장하거나 캐시하기 위해 전역 변수를 사용하면 안 됩니다.
8. `countdown`를 여러 번 다른 입력으로 호출해도 항상 올바르고 독립적인 결과를 반환해야 합니다.

# --hints--

`countdown`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(countdown);
```

`countdown(-1)`는 빈 배열을 반환해야 합니다.

```js
assert.isEmpty(countdown(-1));
```

`countdown(10)`는 `[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`를 반환해야 합니다.

```js
assert.deepStrictEqual(countdown(10), [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

`countdown(5)`는 `[5, 4, 3, 2, 1]`를 반환해야 합니다.

```js
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

코드는 어떤 종류의 루프(`for`, `while`, `Array.from()`)나 고차 함수(`forEach`, `map`, `filter`, `reduce`)에 의존하면 안 됩니다.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

이 문제를 해결하려면 재귀를 사용하세요.

```js
assert.isAtLeast(
  __helpers.removeJSComments(countdown.toString()).match(/countdown\s*\(.+\)/g)?.length,
  2
);
```

전역 변수는 배열을 캐시하는 데 사용하면 안 됩니다.

```js
countdown(1)
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function countdown(n) {
  if (n < 1) {
    return [];
  }

  return [n, ...countdown(n - 1)];
}
```
