---
id: 66c8ba41a77db20a93f9d7a1
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you'll review JavaScript conditionals by building a loan qualification checker app.

The app will check whether the user is eligible for a duplex, car, and condo loan based on their annual income and credit score. A credit score is a number that represents how good you are at managing borrowed money.

To get started, create the following variables and values.

| Variable Name             | Value   |
| ------------------------- | ------- |
| `minIncomeForDuplex`      | `60000` |
| `minCreditScoreForDuplex` | `700`   |
| `minIncomeForCondo`       | `45000` |
| `minCreditScoreForCondo`  | `680`   |
| `minIncomeForCar`         | `30000` |
| `minCreditScoreForCar`    | `650`   |

# --hints--

Створіть змінну `minIncomeForDuplex` зі значенням `60000`.

```js
assert.equal(minIncomeForDuplex, 60000);
```

Створіть `minCreditScoreForDuplex` зі значенням `700`.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

Створіть змінну `minIncomeForCar` зі значенням `30000`.

```js
assert.equal(minIncomeForCar, 30000);
```

Створіть змінну `minCreditScoreForCar` зі значенням `650`.

```js
assert.equal(minCreditScoreForCar, 650);
```

Створіть змінну `minIncomeForCondo` зі значенням `45000`.

```js
assert.equal(minIncomeForCondo, 45000);
```

Створіть змінну `minCreditScoreForCondo` зі значенням `680`.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
