---
id: 646ce8bb4b35544d501c7184
title: Schritt 20
challengeType: 0
dashedName: step-20
---

# --description--

Erstellen Sie innerhalb Ihres `.cat-ears`-Elements zwei `div`-Elemente mit den Klassen `cat-left-ear` und `cat-right-ear`.

# --hints--

Sie sollten das vorhandene `div`-Element mit der Klasse `cat-ears` nicht verändern.

```js
assert(document.querySelectorAll('div.cat-ears').length === 1);
```

Sie sollten zwei `div`-Elemente innerhalb Ihres `.cat-ears`-Elements erstellen.

```js
assert(document.querySelectorAll('.cat-ears div').length === 2);
```

Ihr erstes `div`-Element sollte die Klasse `cat-left-ear` haben.

```js
assert(document.querySelectorAll('.cat-ears div')[0]?.classList.contains('cat-left-ear'));
```

Ihr zweites `div`-Element sollte die Klasse `cat-right-ear` haben.

```js
assert(document.querySelectorAll('.cat-ears div')[1]?.classList.contains('cat-right-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        --fcc-editable-region--
        <div class="cat-ears">
          
        </div>
        --fcc-editable-region--
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
