---
id: 646c48df8674cf2b91020ecb
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie innerhalb Ihres `head`-Elements ein `meta`-Tag mit dem `charset`-Attribut von `utf-8` hinzu. Fügen Sie außerdem ein `title`-Element mit dem Text `fCC Cat Painting` hinzu.

# --hints--

Sie sollten ein einzelnes `meta`-Element hinzufügen.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Ihr `meta`-Element sollte ein `charset`-Attribut mit dem Wert `utf-8` haben.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Sie sollten ein einzelnes `title`-Element hinzufügen.

```js
assert(document.querySelectorAll('title').length === 1);
```

Ihr `title`-Element sollte den Text `fCC Cat Painting` enthalten. Beachten Sie, dass Rechtschreibung und Groß-/Kleinschreibung wichtig sind.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
