---
id: 646c5ace05e4be4211407935
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Um das `.cat-head`-Element zu sehen, geben Sie ihm einen linearen Farbverlauf als Hintergrund mit `#5e5e5e` bei `85%` und `#45454f` bei `100%`.

Vielleicht bemerken Sie den Unterschied zwischen diesen beiden Farben nicht, aber sie sind vorhanden.

# --hints--

Ihr `.cat-head`-Selektor sollte eine `background`-Eigenschaft haben.

```js
assert.match(code, /background:/)
```

Ihre `background`-Eigenschaft sollte die `linear-gradient`-Funktion verwenden.

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

Ihre `linear-gradient`-Funktion sollte die erste Farbe auf `#5e5e5e` bei `85%` setzen.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

Ihre `linear-gradient`-Funktion sollte die zweite Farbe auf `#45454f` bei `100%` setzen.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
