---
id: 646c48df8674cf2b91020ecc
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Fügen Sie ein `link`-Element innerhalb Ihres `head`-Elements hinzu. Für dieses `link`-Element setzen Sie das `rel`-Attribut auf `stylesheet` und das `href`-Attribut auf `./styles.css`.

# --hints--

Sie sollten ein `link`-Element haben.

```js
assert.match(code, /<link/)
```

Ihr `link`-Element sollte ein `rel`- und ein `href`-Attribut haben, die auf `stylesheet` und `./styles.css` gesetzt sind.

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
