---
id: 646c47867800472a4ed5d2ea
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Beginnen Sie mit der grundlegenden HTML-Struktur. Fügen Sie eine `DOCTYPE` Referenz von `html` und ein `html` Element mit seinem `lang` Attribut, das auf `en` gesetzt ist, hinzu. Fügen Sie außerdem ein `head` und ein `body` Element innerhalb des `html` Elements hinzu.

# --hints--

Sie sollten die `DOCTYPE` Deklaration von `html` haben.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Ihre `DOCTYPE` Deklaration sollte am Anfang Ihres HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Sie sollten ein öffnendes `html` Tag mit einem `lang` Attribut von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Sie sollten ein schließendes `html` Tag haben.

```js
assert(code.match(/<\/html>/i));
```

Sie sollten ein öffnendes `head` Tag haben.

```js
assert(code.match(/<head>/i));
```

Sie sollten ein schließendes `head` Tag haben.

```js
assert(code.match(/<\/head>/i));
```

Sie sollten ein öffnendes `body` Tag haben.

```js
assert(code.match(/<body>/i));
```

Sie sollten ein schließendes `body` Tag haben.

```js
assert(code.match(/<\/body>/i));
```

Ihr `body` Element sollte nach dem `head` Element kommen.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

Ihre `head` und `body` Elemente sollten sich innerhalb des `html` Elements befinden.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
