---
id: 6476fc5cf14b276e6d04e82a
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Verwenden Sie einen Klassenselektor, um Ihrem `.box`-Element eine Breite von `200px`, eine Höhe von `600px` und eine Hintergrundfarbe von `#000` zu geben. Geben Sie ihm außerdem eine `position` von `absolute`, eine `top` von `800px` und eine `left` von `650px`.

# --hints--

Sie sollten einen `.box`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

Ihr `.box`-Selektor sollte eine `width`-Eigenschaft mit dem Wert `200px` haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

Ihr `.box`-Selektor sollte eine `height`-Eigenschaft mit dem Wert `600px` haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

Ihr `.box`-Selektor sollte eine `background-color`-Eigenschaft mit dem Wert `#000` haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

Ihr `.box`-Selektor sollte eine `position`-Eigenschaft mit dem Wert `absolute` haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

Ihr `.box`-Selektor sollte eine `left`-Eigenschaft mit dem Wert `650px` haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

Ihr `.box`-Selektor sollte eine `top`-Eigenschaft mit dem Wert `800px` haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

