---
id: 6476fd4213318f6ee211028a
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Ersetzen Sie nun den Wert der `position` Eigenschaft Ihres `.cat-head` durch `fixed`. Lassen Sie sowohl `top` als auch `left` unverändert.

Scrollen Sie danach nach oben und unten, um zu sehen, wie der `fixed` Wert funktioniert.

# --hints--

Ihr `.cat-head` Selektor sollte eine `position` Eigenschaft mit dem Wert `fixed` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'fixed')
```

Ihr `.cat-head` Selektor sollte eine `top` Eigenschaft mit dem Wert `100px` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Ihr `.cat-head` Selektor sollte eine `left` Eigenschaft mit dem Wert `100px` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
