---
id: 6477062778c85972eb648030
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Sie sollten nun den Katzenkopf zentrieren.

Geben Sie dem `.cat-head`-Element eine `position`-Eigenschaft mit dem Wert `absolute`. Setzen Sie für die `0`, `right`, `left`, `top`-Eigenschaften den Wert `bottom` und anschließend die `margin`-Eigenschaft auf allen Seiten auf `auto`. Das ist eine Möglichkeit, ein Element mit CSS-Positionierung vertikal und horizontal zu zentrieren.

# --hints--

Ihr `.cat-head`-Selektor sollte eine `position`-Eigenschaft mit dem Wert `absolute` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Ihr `.cat-head`-Selektor sollte eine `top`-Eigenschaft mit dem Wert `0` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Ihr `.cat-head`-Selektor sollte eine `left`-Eigenschaft mit dem Wert `0` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

Ihr `.cat-head`-Selektor sollte eine `right`-Eigenschaft mit dem Wert `0` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

Ihr `.cat-head`-Selektor sollte eine `bottom`-Eigenschaft mit dem Wert `0` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

Ihr `.cat-head`-Selektor sollte eine `margin`-Eigenschaft mit dem Wert `auto` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
