---
id: 68e44519c769bf545bb3f22a
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

要取得 `"I love dogs and dogs are so much fun!"` 的預期輸出，你需要使用 `replaceAll()` 方法，而不是 `replace()` 方法。此方法會為所有符合子字串的匹配傳回一個新的字串。 

將 `exampleSentence.replace("cats", "dogs")` 更新為使用 `replaceAll()` 方法。現在你應該會在主控台看到正確的文字。

# --hints--

你應該更新 `exampleSentence.replace("cats", "dogs")` 這行程式碼以使用 `replaceAll()` 方法。

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
