---
id: 68e41084bb275539f72de9fe
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個字串轉換器應用程式來練習使用 `replace()`、`replaceAll()` 和 `repeat()` 方法。 

首先，創建一個名為 `originalString` 的變數，並將其指定為 `"I love cats."`。

# --hints--

你應該有一個名為 `originalString` 的變數。

```js
assert.exists(originalString);
```

你的 `originalString` 變數應該是一個字串。

```js
assert.isString(originalString);
```

你應該將字串 `"I love cats."` 指定給 `originalString` 變數。

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
