---
id: 68e44c305f1f935d949d76bb
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

如前一步驟所述，句子 `"I love love love  learning."` 中最後一個 `"love"` 後面有一個多餘的空格。若要移除該多餘空格，你可以使用 `trimEnd()` 方法。 

在 `"love ".repeat(3)` 方法結束時，串鏈 `trimEnd()` 方法。你可以這樣串鏈方法：

```js
.firstMethod().secondMethod()
```

現在當你檢查主控台時，你應該會看到多餘的空格已被移除。

隨著最後一次變更，你的工作坊已完成！ 

# --hints--

你應該將 `trimEnd()` 方法串鏈到 `"love ".repeat(3)` 方法。

```js
assert.equal(repeatedLove, "love ".repeat(3).trimEnd());
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--
const repeatedLove = "love ".repeat(3)
--fcc-editable-region--
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```

# --solutions--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3).trimEnd();
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```
