---
id: 68e42de6396f5c49d90d1951
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

在先前的課程中，你學習了用於在字串中尋找子字串並將其替換為另一個值的 `replace()` 方法。

這裡有一個範例：

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

請記住，字串是不可變的，這表示在該範例中原始的 `text` 不會被修改，而是會創建一個新的字串。 

創建一個名為 `replacedString` 的新變數，並指定它為 `originalString.replace("cats", "dogs")`。 

# --hints--

你應該有一個名為 `replacedString` 的變數。

```js
assert.exists(replacedString);
```

你的 `replacedString` 變數應該是一個字串。

```js
assert.isString(replacedString);
```

你應該將 `originalString.replace("cats", "dogs")` 指定給 `replacedString` 變數。

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
