---
id: 68e42b62d7093345c6975f22
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

在你的 `originalString` 變數下方，添加一個帶有字串 `"Original string:"` 的 `console.log()`。 

然後在那個 `console.log()` 下面，為你的 `console.log()` 添加另一個帶有 `originalString` 變數的 `console.log()`。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該將字串 `"Original string:"` 記錄到主控台。

```js
assert.equal(getLogs()[0], "Original string:");
```

你應該將 `originalString` 變數記錄到主控台。 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
