---
id: 68e44abea678a65b877864dc
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

看到那個重複的字出現在句子中會很不錯。創建一個名為 `newSentence` 的新變數，並將將字串 `"I "` 與變數 `repeatedLove` 串接後，再串接字串 `" learning."` 的結果指定給它。然後將變數 `newSentence` 輸出到主控台。你可以選擇使用 `+` 運算子或模板字串來進行字串串接。 

現在你應該會看到文字 `"I love love love  learning."`。 

*注意*：在最後一個 `"love"` 後面有一個多餘的空格，這是正確的。在最後一步，你會修剪掉那個多餘的空格。 

# --hints--

你應該有一個名為 `newSentence` 的變數。

```js
assert.exists(newSentence);
```

你的 `newSentence` 變數應該是一個字串。

```js
assert.isString(newSentence);
```

你應該將字串 `"I "` 與 `repeatedLove` 變數串接，接著再串接字串 `" learning."`。這個結果應該指定給 `newSentence` 變數。請注意空格。 

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

你不應該將字面字串 `"I love love love  learning."` 指定給 `newSentence` 變數。請重新閱讀說明以取得正確的方法。

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
