---
id: 68e43e884ef09d4d2ab05631
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

現在是時候查看使用 `replace()` 方法後的結果。

首先為你的程式碼添加一個 `console.log()`，其值為 `"After using the replace() method:"`。

然後在那個 `console.log()` 下面，為你的 `console.log()` 添加另一個帶有 `replacedString` 變數的 `console.log()`。

查看主控台，你應該會看到新句子顯示為 `"I love dogs."`，而不是原本的 `"I love cats."`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該將 `"After using the replace() method:"` 記錄到主控台。

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

你應該將 `replacedString` 變數記錄到主控台。 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
