---
id: 68e4465fba0bd256868b7deb
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

現在是時候使用新的字串方法。

首先建立一個名為 `learningSentence` 的變數，並指定字串 `"I love learning!"`。 

然後在該變數下方，向主控台記錄 `"Original learning sentence:"`。在該 `console.log()` 下方，為 `learningSentence` 變數添加另一個 `console.log()`。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該有一個 `learningSentence` 變數。

```js
assert.exists(learningSentence);
```

你的 `learningSentence` 變數應該是一個字串。

```js
assert.isString(learningSentence);
```

你應該將字串 `"I love learning!"` 指定給 `learningSentence` 變數。

```js
assert.equal(learningSentence, "I love learning!");
```

你應該將 `"Original learning sentence:"` 記錄到主控台。

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

你應該將 `learningSentence` 變數記錄到主控台。

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
