---
id: 68e43f94365a964f26c88f78
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

現在是時候使用新的字串方法。 

首先建立一個名為 `exampleSentence` 的新變數，並將字串 `"I love cats and cats are so much fun!"` 指定給它。

然後在該變數下方，向主控台記錄 `"Original sentence:"`。接著在該 `console.log()` 下方，使用 `exampleSentence` 變數再新增一個 `console.log()`。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該擁有一個 `exampleSentence` 變數。

```js
assert.exists(exampleSentence);
```

你的 `exampleSentence` 變數應該是一個字串。

```js
assert.isString(exampleSentence);
```

你應該將字串 `"I love cats and cats are so much fun!"` 指定到 `exampleSentence` 變數。

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

你應該將 `"Original sentence:"` 記錄到主控台。

```js
assert.equal(getLogs()[4], "Original sentence:");
```

你應該將 `exampleSentence` 變數記錄到主控台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
