---
id: 68e4417b7317835208145323
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

接著，創建一個名為 `dogsOnlySentence` 的新變數，並指定它為 `exampleSentence.replace("cats", "dogs")`。

然後在該變數下方，將 `"Replacing all occurrences of cats with dogs:"` 記錄到主控台。在該 `console.log()` 下方，再新增另一個 `console.log()`，用於 `dogsOnlySentence` 變數。 

你現在應該會在主控台看到記錄的文字 `"I love dogs and cats are so much fun!"`。但那不是期望的結果。在下一步，你將學習如何修正它，使文字顯示為 `"I love dogs and dogs are so much fun!"`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該有一個 `dogsOnlySentence` 變數。

```js
assert.exists(dogsOnlySentence);
```

你的 `dogsOnlySentence` 變數應該是一個字串。

```js
assert.isString(dogsOnlySentence);
```

你應該將 `exampleSentence.replace("cats", "dogs")` 指定給你的 `dogsOnlySentence` 變數。

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

你應該將 `"Replacing all occurrences of cats with dogs:"` 記錄到主控台。

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

你應該將 `dogsOnlySentence` 變數記錄到主控台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
