---
id: 68e448ec4ac4205968c0a805
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

如果在句子 `"I love learning!"` 中重複三次字 `"love"`，那會很好。 

如你在先前的課程中所回想，你可以使用 `repeat()` 方法將字串重複特定次數。以下是一個範例：

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

由於字串是不可變的，這個方法不會修改原始字串。它會傳回一個包含重複內容的新字串。 

創建一個名為 `repeatedLove` 的變數並指定它為 `"love ".repeat(3)`。然後將 `repeatedLove` 變數輸出到主控台，以便你可以看到結果。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你應該有一個名為 `repeatedLove` 的變數。

```js
assert.exists(repeatedLove);
```

你的 `repeatedLove` 變數應該是一個字串。

```js
assert.isString(repeatedLove);
```

你應該將 `"love ".repeat(3)` 指定給 `repeatedLove` 變數。

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

你應該將 `repeatedLove` 變數記錄到主控台。

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
