---
id: 587d7fad367417b2b2512be1
title: 使用 JavaScript 和 `onclick` 屬性控制點擊事件
challengeType: 6
forumTopicId: 301503
dashedName: handle-click-events-with-javascript-using-the-onclick-property
---

# --description--

你希望你的程式碼只在頁面完成載入後執行。為此，你可以將一個名為 `DOMContentLoaded` 的 JavaScript 事件附加到文件。以下是執行此操作的程式碼：

```js
document.addEventListener('DOMContentLoaded', function () {});
```

你可以實作放在 `DOMContentLoaded` 函式內的事件處理常式。你可以實作一個 `onclick` 事件處理常式，當使用者點擊 `#getMessage` 元素時觸發，方法是為你的程式碼添加以下程式碼：

```js
document.getElementById('getMessage').onclick = function () {};
```

# --instructions--

在 `DOMContentLoaded` 函式（程式）內，為 id 為 `getMessage` 的元素添加點擊事件處理常式。

# --hints--

你的程式碼應該使用 `document.getElementById` 方法來選取 id 為 `getMessage` 的元素。

```js
assert.match(code, /document\s*\.getElementById\(\s*('|")getMessage\1\s*\)/g);
```

你的程式碼應該為你的 `onclick` 事件添加處理常式。

```js
assert.isFunction(document.getElementById('getMessage').onclick);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    document.getElementById('getMessage').onclick = function () {};
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
