---
id: 587d7fae367417b2b2512be4
title: 從 API 存取 JSON 資料
challengeType: 6
forumTopicId: 301499
dashedName: access-the-json-data-from-an-api
---

# --description--

在前一個挑戰中，你看到了如何從 freeCodeCamp Cat Photo API 取得 JSON 資料。

現在你將更仔細地查看傳回的資料，以更好地理解 JSON 格式。回想 JavaScript 中的一些符號表示法：

<blockquote>[ ] -> Square brackets represent an array.<br>{ } -> Curly brackets represent an object.<br>" " -> Double quotes represent a string. They are also used for key names in JSON.</blockquote>

了解 API 傳回的資料結構很重要，因為它會影響你如何取得所需的值。

在右側，點擊 `Get Message` 按鈕以將 freeCodeCamp Cat Photo API JSON 載入到 HTML 中。

你在 JSON 資料中看到的第一個和最後一個字元是方括號 `[ ]`。這表示傳回的資料是一個陣列。JSON 資料中的第二個字元是大括號 `{`，它代表一個物件的開始。仔細看，你可以看到有三個獨立的物件。JSON 資料是一個包含三個物件的陣列，每個物件都包含有關貓咪照片的資訊。

你之前學到物件包含以逗號分隔的「鍵-值對組」。在 Cat Photo 範例中，第一個物件有 `"id":0`，其中 `id` 是鍵，`0` 是其對應的值。類似地，還有 `imageLink`、`altText` 和 `codeNames` 的鍵。每個貓咪照片物件都有這些相同的鍵，但值不同。

第一個物件中另一個有趣的「鍵值對」是 `"codeNames":["Juggernaut","Mrs. Wallace","ButterCup"]`。這裡 `codeNames` 是鍵，其值是一個包含三個字串的陣列。也可以有物件陣列，以及鍵的值為陣列的情況。

記住如何存取陣列和物件中的資料。陣列使用括號表示法來存取特定索引的項目。物件則使用括號或點表示法來存取特定屬性的值。以下是一個範例，列印第一張貓咪照片的 `altText` 屬性－請注意，編輯器中解析的 JSON 資料已存檔在名為 `json` 的變數中：

```js
console.log(json[0].altText);
```

主控台會顯示字串 `A white cat wearing a green helmet shaped melon on its head.`。

# --instructions--

對於 `id` 為 2 的貓，將 `codeNames` 陣列中的第二個值列印到主控台。你應該使用括號和點記法來存取儲存在變數 `json` 中的物件的該值。

# --hints--

你的程式碼應該使用括號和點記法來存取正確的程式碼名稱，並在主控台列印 `Loki`。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*json\s*\[2\]\s*(\.\s*codeNames|\[\s*('|`|")codeNames\2\s*\])\s*\[\s*1\s*\]\s*\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line

        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line
        console.log(json[2].codeNames[1]);
        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
