---
id: 587d7fae367417b2b2512be3
title: 使用 JavaScript 的 XMLHttpRequest 方法取得 JSON
challengeType: 6
forumTopicId: 301502
dashedName: get-json-with-the-javascript-xmlhttprequest-method
---

# --description--

你也可以從外部原始碼請求資料。這就是 API 派上用場的地方。

請記住，API－或稱應用程式介面－是電腦用來彼此溝通的工具。你將學習如何使用一種稱為 AJAX 的技術，利用從 API 獲取的資料來更新 HTML。

大多數網路 API 以稱為 JSON 的格式傳輸資料。JSON 代表 JavaScript 物件表示法。

JSON 語法看起來非常類似 JavaScript 物件字面值表示法。JSON 有物件屬性及其當前值，夾在 `{` 和 `}` 之間。

這些屬性及其值通常被稱為「鍵值配對」。

然而，API 傳送的 JSON 是以 `bytes` 形式傳送，而你的應用程式接收到的是 `string`。這些可以轉換成 JavaScript 物件，但預設並不是 JavaScript 物件。`JSON.parse` 方法會解析該字串並構建其描述的 JavaScript 物件。

你可以從 freeCodeCamp 的 Cat Photo API 請求 JSON。以下是你可以放在點擊 `event` 中的程式碼來執行此操作：

```js
const req = new XMLHttpRequest();
req.open('GET', '/json/cats.json', true);
req.send();
req.onload = function () {
  const json = JSON.parse(req.responseText);
  document.getElementsByClassName('message')[0].innerHTML =
    JSON.stringify(json);
};
```

以下是每個部分的功能說明。JavaScript 的 `XMLHttpRequest` 物件有許多用於傳輸資料的屬性和方法。首先，創建一個 `XMLHttpRequest` 物件的實體並將其存入 `req` 變數。接著，`open` 方法初始化一個請求－此範例是從 API 請求資料，因此是 `GET` 請求。`open` 的第二個引數是你要請求資料的 API URL。第三個引數是布林值，`true` 表示非同步請求。`send` 方法送出請求。最後，`onload` 事件處理常式解析返回的資料，並套用 `JSON.stringify` 方法將 JavaScript 物件轉換成字串。此字串隨後被插入為訊息文字。

# --instructions--

更新程式碼以創建並發送 `GET` 請求到 freeCodeCamp Cat Photo API。然後點擊 `Get Message` 按鈕。你的 AJAX 函式將會用 API 回傳的未經處理的 JSON 輸出取代 `The message will go here` 文字。

# --hints--

你的程式碼應該創建一個新的 `XMLHttpRequest`。

```js
assert.match(code, /new\s+?XMLHttpRequest\(\s*?\)/g);
```

你的程式碼應該使用 `open` 方法來初始化對 freeCodeCamp Cat Photo API 的 `GET` 請求。

```js
assert.match(
  code,
  /\.open\(\s*('|")GET\1\s*,\s*('|")\/json\/cats\.json\2\s*,\s*true\s*\)/g
);
```

你的程式碼應該使用 `send` 方法來發送請求。

```js
assert.match(code, /\.send\(\s*\)/g);
```

你的程式碼應該有一個 `onload` 事件處理常式設定為一個函式。

```js
assert.match(
  code,
  /\.onload\s*=\s*(function|\(\s*\))\s*(\(\s*\)|\=\>)\s*{/g
);
```

你的程式碼應該使用 `JSON.parse` 方法來解析 `responseText`。

```js
assert.match(code, /JSON\s*\.parse\(\s*.*\.responseText\s*\)/g);
```

你的程式碼應該取得具有類別 `message` 的元素，並將其內部 HTML 更改為 JSON 資料的字串。

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.innerHTML\s*?=\s*?JSON\.stringify\(.+?\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = () => {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
