---
id: 587d7faf367417b2b2512be8
title: 取得地理位置資料以尋找使用者的 GPS 座標
challengeType: 6
forumTopicId: 18188
dashedName: get-geolocation-data-to-find-a-users-gps-coordinates
---

# --description--

另一個你可以做的很酷的事情是存取你使用者的當前位置。每個瀏覽器都有內建的 navigator，可以提供你這些資訊。

導航器將取得使用者當前的經度和緯度。

你會看到提示，詢問是否允許或阻擋此網站知道你目前的位置。只要程式碼正確，挑戰無論哪種方式都能完成。

選擇允許後，你將會看到輸出電話上的文字變成你的緯度和經度。

這是完成此操作的程式碼：

```js
if (navigator.geolocation) {
  navigator.geolocation.getCurrentPosition(function (position) {
    document.getElementById('data').innerHTML =
      'latitude: ' +
      position.coords.latitude +
      '<br>longitude: ' +
      position.coords.longitude;
  });
}
```

首先，它會檢查 `navigator.geolocation` 物件是否存在。如果存在，則會呼叫該物件上的 `getCurrentPosition` 方法，該方法會啟動非同步請求以取得使用者的位置。如果請求成功，方法中的回呼函式會被執行。此函式會使用點記法存取 `position` 物件的緯度和經度值，並更新 HTML。

# --instructions--

將範例程式碼添加到 `script` 頁籤中，以檢查使用者的當前位置並將其插入到 HTML 中。

# --hints--

你的程式碼應該使用 `navigator.geolocation` 來存取使用者的當前位置。

```js
assert.match(code, /navigator\.geolocation\.getCurrentPosition/g);
```

你的程式碼應該使用 `position.coords.latitude` 來顯示使用者的緯度位置。

```js
assert.match(code, /position\.coords\.latitude/g);
```

你的程式碼應該使用 `position.coords.longitude` 來顯示使用者的經度位置。

```js
assert.match(code, /position\.coords\.longitude/g);
```

你應該在具有 `id="data"` 的 `div` 元素中顯示使用者的位置。

```js
assert.match(
  code,
  /document\.getElementById\(\s*('|")data\1\s*\)\.innerHTML/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  // Add your code below this line

  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data"></div>
```

# --solutions--

```html
<script>
  // Add your code below this line
  if (navigator.geolocation) {
    navigator.geolocation.getCurrentPosition(function(position) {
      document.getElementById('data').innerHTML = "latitude: " + position.coords.latitude + "<br>longitude: " + position.coords.longitude;
    });
  }
  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data">

</div>

</section>
```
