---
id: 587d7faf367417b2b2512be9
title: 使用 JavaScript 的 XMLHttpRequest 方法傳送資料(Post Data)
challengeType: 6
forumTopicId: 301504
dashedName: post-data-with-the-javascript-xmlhttprequest-method
---

# --description--

在前面的範例中，你從外部資源接收資料。只要該資源支援 AJAX 請求且你知道 URL，你也可以將資料傳送到外部資源。

JavaScript 的 `XMLHttpRequest` 方法也用於將資料發送到伺服器。以下是一個範例：

```js
const xhr = new XMLHttpRequest();
xhr.open('POST', url, true);
xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
xhr.onreadystatechange = function () {
  if (xhr.readyState === 4 && xhr.status === 201) {
    const serverResponse = JSON.parse(xhr.response);
    document.getElementsByClassName('message')[0].textContent =
      serverResponse.userName + serverResponse.suffix;
  }
};
const body = JSON.stringify({ userName: userName, suffix: ' loves cats!' });
xhr.send(body);
```

你之前已經看過幾個這些方法。這裡的 `open` 方法將請求初始化為對外部資源給定 URL 的 `POST`，並將 `true` 作為第三個參數傳入－表示要非同步執行該操作。

`setRequestHeader` 方法設定 HTTP 請求標頭的值，該標頭包含有關發送者和請求的資訊。它必須在 `open` 方法之後，但在 `send` 方法之前被呼叫。這兩個參數分別是標頭的名稱和要設定為該標頭主體的值。

接著，`onreadystatechange` 事件監聽器控制代碼請求狀態的變化。`readyState` 為 `4` 表示操作已完成，而 `status` 為 `201` 表示請求成功。因此，可以更新文件的 HTML。

最後，`send` 方法會使用 `body` 值發送請求。`body` 包含 `userName` 和 `suffix` 鍵。

# --instructions--

更新程式碼，使其對 API 端點發出 `POST` 請求。然後在輸入欄位中輸入你的名字並點擊 `Send Message`。你的 AJAX 函式應該用來自伺服端的資料取代 `Reply from Server will be here.`。格式化回應以顯示附加了文字 ` loves cats` 的你的名字。

# --hints--

你的程式碼應該創建一個新的 `XMLHttpRequest`。

```js
assert.match(code, /new\s*XMLHttpRequest\(\s*\)/g);
```

你的程式碼應該使用 `open` 方法來初始化對伺服端的 `POST` 請求。

```js
assert.match(code, /\.open\(\s*('|")POST\1\s*,\s*url\s*,\s*true\s*\)/g);
```

你的程式碼應該使用 `setRequestHeader` 方法。

```js
assert.match(
  code,
  /\.setRequestHeader\(\s*('|")Content-Type\1\s*,\s*('|")application\/json;\s*charset=UTF-8\2\s*\)/g
);
```

你的程式碼應該有一個 `onreadystatechange` 事件處理常式設定為一個函式。

```js
assert.match(code, /\.onreadystatechange\s*=/g);
```

你的程式碼應該取得具有類別 `message` 的元素，並將其 `textContent` 更改為 `userName loves cats`。

```js
assert.match(
  code,
  /document\.getElementsByClassName\(\s*('|")message\1\s*\)\[0\]\.textContent\s*=\s*.+?\.userName\s*\+\s*.+?\.suffix/g
);
```

你的程式碼應該使用 `send` 方法。

```js
assert.match(code, /\.send\(\s*body\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line

      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message box">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line
      const xhr = new XMLHttpRequest();
      xhr.open('POST', url, true);
      xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
      xhr.onreadystatechange = function () {
        if (xhr.readyState === 4 && xhr.status === 201) {
          const serverResponse = JSON.parse(xhr.response);
          document.getElementsByClassName('message')[0].textContent =
            serverResponse.userName + serverResponse.suffix;
        }
      };
      const body = JSON.stringify({
        userName: userName,
        suffix: ' loves cats!'
      });
      xhr.send(body);
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```
