---
id: 5ccfad82bb2dc6c965a848e5
title: 使用 JavaScript 的 fetch 方法取得 JSON
challengeType: 6
forumTopicId: 301501
dashedName: get-json-with-the-javascript-fetch-method
---

# --description--

另一種請求外部資料的方法是使用 JavaScript 的 `fetch()` 行為。它等同於 `XMLHttpRequest`，但語法被認為較容易理解。

這是對 `/json/cats.json` 發出 GET 請求的程式碼。

```js
fetch('/json/cats.json')
  .then(response => response.json())
  .then(data => {
    document.getElementById('message').innerHTML = JSON.stringify(data);
  });
```

注意：`fetch()` 方法使用 `GET` 作為預設的 `HTTP` 方法。這表示你不需要為基本資料擷取明確指定它。

看看這段 `code` 的每一部分。

第一行是發出請求的那一行。所以，`fetch(URL)` 對指定的 URL 發出 `GET` 請求。該方法傳回一個 Promise。

當 Promise 被傳回後，如果請求成功，則會執行 `then` 方法，該方法會取得回應並將其轉換為 JSON 格式。

`then` 方法也會傳回一個 Promise，該 Promise 由下一個 `then` 方法處理。第二個 `then` 中的引數就是你要尋找的 JSON 物件！

現在，它使用 `document.getElementById()` 選取將接收資料的元素。然後，它透過插入從請求傳回的 JSON 物件生成的字串來修改該元素的 HTML 程式碼。

# --instructions--

更新程式碼以創建並發送 `GET` 請求到 freeCodeCamp Cat Photo API。但這次使用 `fetch` 方法而非 `XMLHttpRequest`。

# --hints--

你的程式碼應該使用擷取的資料來取代內部的 HTML。

```js
const catData = 'dummy data';
const ref = fetch;
fetch = () => Promise.resolve({ json: () => catData });

  try {
    document.getElementById('getMessage').click();
    await new Promise((resolve, reject) => setTimeout(() => resolve(), 250));
  } catch (error) {
    console.log(error);
  } finally {
    fetch = ref;
    assert.equal(
      document.getElementById('message').textContent,
      JSON.stringify(catData)
    );
  }

```

你的程式碼應該使用 `fetch` 發出 `GET` 請求。

```js
assert.match(code, /fetch\s*\(\s*('|")\/json\/cats\.json\1\s*\)/g);
```

你的程式碼應該使用 `then` 將回應轉換為 JSON。

```js
assert.match(
  code,
  /\.then\s*\(\s*\(?(?<var>\w+)\)?\s*=>\s*\k<var>\s*\.json\s*\(\s*\)\s*\)/g
);
```

你的程式碼應該使用 `then` 來控制代碼由另一個 `then` 轉換成 JSON 的資料。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.then\(\(?\w+\)?=>{[^}]*}\)/g);
```

你的程式碼應該取得 id 為 `message` 的元素，並將其內部 HTML 更改為 JSON 資料的字串。

```js
assert.match(
  __helpers.removeWhiteSpace(code),
  /document\.getElementById\(('|")message\1\)\.innerHTML=JSON\.stringify\(?\w+\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      fetch('/json/cats.json')
        .then(response => response.json())
        .then(data => {
          document.getElementById('message').innerHTML = JSON.stringify(data);
        });
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
