---
id: 669e81368e52b3a5c35a2dc5
title: Crea una página de compilación de vídeo
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener un elemento `main` como el único hijo del elemento `body`.
1. Debes tener un elemento `h1` con el tema de tu página.
1. Debes tener un párrafo que introduzca el tema de tu página debajo del elemento `h1`.
1. Debes tener tres elementos `section` debajo de tu primer párrafo.
1. Cada sección debe contener un elemento `h2`, un párrafo y un elemento `iframe`, en este orden.
1. Los tres elementos `iframe` deben tener un atributo `src` establecido en un video válido.
1. Cada elemento `iframe` también debe tener un atributo `title` para describir el contenido incrustado, y un atributo `height` y un atributo `width` para establecer el tamaño adecuado del elemento.

# --hints--

Debes tener un elemento `main` dentro de tu elemento `body`.

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

Tu elemento `main` debe ser el único hijo del elemento `body`.

```js
assert.equal(document.querySelector('body').children.length, 1);
```

Debes tener un elemento `h1` con el tema de tu página dentro del elemento `main`.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Debes tener un párrafo que introduzca el tema de tu página debajo del elemento `h1`.

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

Debes tener tres elementos `section` debajo de tu primer elemento `p`.

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

Cada elemento `section` debe comenzar con un elemento `h2` que sirva como título para esa sección.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

Cada elemento `section` debe contener un elemento `p` para introducir el contenido del video como su segundo hijo.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

Cada elemento `section` debe contener un elemento `iframe` como su tercer hijo.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

Tu primer elemento `iframe` debe tener un atributo `src` establecido en un video válido.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

Tu primer elemento `iframe` debe tener un atributo `title` para describir el contenido incrustado.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

Tu primer elemento `iframe` debe tener un atributo `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

Tu primer elemento `iframe` debe tener un atributo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

Tu segundo elemento `iframe` debe tener un atributo `src` establecido en un video válido.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

Tu segundo elemento `iframe` debe tener un atributo `title` para describir el contenido incrustado.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

Tu segundo elemento `iframe` debe tener un atributo `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

Tu segundo elemento `iframe` debe tener un atributo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

Tu tercer elemento `iframe` debe tener un atributo `src` establecido en un video válido.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

Tu tercer elemento `iframe` debe tener un atributo `title` para describir el contenido incrustado.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

Tu tercer elemento `iframe` debe tener un atributo `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

Tu tercer elemento `iframe` debe tener un atributo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Frontend Web Development</h1>
        <p>Frontend Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Frontend
            developers
            ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
