---
id: a77dbc43c33f39daa4429b4f
title: Tengeneza kitendakazi cha ukaguzi wa thamani ya kweli au siyo kweli
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

Katika maabara hii, utatengeneza kitendakazi kinachokagua kama thamani imeainishwa kama thamani ya kweli au siyo kweli ya msingi.

Thamani za kweli au siyo kweli za msingi ni `true` na `false`.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi kinachoitwa `booWho` kinachopokea hoja moja.
1. Ikiwa hoja inayopokelewa ni thamani ya kweli au siyo kweli ya msingi, kitendakazi kinapaswa kurudisha `true`.
1. Ikiwa hoja ni thamani nyingine yoyote, kitendakazi kinapaswa kurudisha `false`.

# --hints--

Unapaswa kuwa na kitendakazi cha `booWho`.

```js
assert.isFunction(booWho);
```

`booWho(true)` inapaswa kurudisha `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` inapaswa kurudisha `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
