---
id: 67fccc2463253b213a000142
title: Einen Multimedia-Player erstellen
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

In den vorherigen Lektionen haben Sie den Umgang mit `audio` und `video` Elementen kennengelernt. In diesem Labor werden Sie einen Multimedia-Player erstellen, der eine `audio` Spur und `video` mit einem Transkript anzeigt.

Für das `audio` Element müssen Sie ein `source` Element einfügen, das verwendet wird, um die verwendeten Medien anzugeben.

Hier ist ein Beispiel:

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

Das `source` Element kann auch im `video` Element wie folgt verwendet werden:

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `h1` Element für den Haupttitel der Seite haben.
2. Sie sollten drei `section` Elemente haben.
3. Im ersten `section` Element sollten Sie ein `h2` Element für den Titel des abgespielten Songs haben.
4. Unter dem `h2` Element sollten Sie ein `audio` Element mit dem `controls` Attribut und einem `aria-label` Attribut haben.
5. Im `audio` Element sollten Sie ein `source` Element mit einem `src` Attribut, das auf eine Audiodatei zeigt, und einem `type` Attribut haben. Sie können gerne diese Audio-URL verwenden: `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. Im zweiten `section` Element sollten Sie ein `h2` Element für den Titel des abgespielten Videos haben.
7. Unter dem `h2` Element sollten Sie ein `video` Element mit den `controls`, `width` Attributen und einem `aria-label` Attribut haben.
8. Im `video` Element sollten Sie ein `source` Element mit einem `src` Attribut, das auf eine Videodatei zeigt, und einem `type` Attribut haben. Sie können gerne diese Video-URL verwenden: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. Unter dem `source` Element sollten Sie ein `track` Element mit einem `src` Attribut, das auf eine Untertiteldatei zeigt, sowie einem `kind`, einem `srclang` und einem `label` Attribut haben. Sie können gerne diese Untertitel-URL verwenden: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.vtt`
10. Im dritten `section` Element sollten Sie ein `h2` Element für den Titel des Abschnitts, z. B. „Transcript“, haben.
11. Unter dem `h2` Element sollten Sie ein `p` Element mit dem Transkript des Videos haben.

# --hints--

Sie sollten ein `h1` Element haben.

```js
assert.isNotNull(document.querySelector("h1"));
```

Ihr `h1` sollte einen Text enthalten, der den Haupttitel der Seite darstellt.

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```

Sie sollten drei `section` Elemente haben.

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

Das erste `section` sollte ein `h2` Element enthalten.

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

Ihr `h2` sollte einen Text enthalten, der den Titel des abgespielten Songs darstellt.

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

Das erste `section` sollte ein `audio` Element enthalten.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

Ihr `audio` Element sollte ein `controls` Attribut haben.

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

Ihr `audio` Element sollte ein `aria-label` Attribut mit einem Text haben, der das Audio beschreibt.

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

Ihr `audio` Element sollte ein `source` Element enthalten.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

Ihr `source` Element sollte ein `src` Attribut haben, das auf eine Audiodatei zeigt.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

Ihr `source` Element sollte ein `type` Attribut mit einem Wert haben, der den Typ des Audios angibt.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

Das zweite `section` sollte ein `h2` Element enthalten.

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

Ihr `h2` sollte einen Text enthalten, der den Titel des abgespielten Videos darstellt.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

Das zweite `section` sollte ein `video` Element enthalten.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

Ihr `video` Element sollte ein `controls` Attribut haben.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

Ihr `video` Element sollte ein `width` Attribut mit einem Wert haben, der die Breite des Videos angibt.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

Das `video` Element sollte ein `source` Element enthalten.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

Das `source` Element sollte ein `src` Attribut haben, das auf eine Videodatei zeigt.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

Das `source` Element sollte ein `type` Attribut mit einem Wert haben, der den Typ des Videos angibt.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

Ihr `video` Element sollte ein `track` Element enthalten.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

Ihr `track` Element sollte ein `src` Attribut haben, das auf eine Untertiteldatei zeigt.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

Ihr `track` Element sollte ein `kind` Attribut mit einem Text haben, der die Art der Datei beschreibt.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

Ihr `track` Element sollte ein `srclang` Attribut mit einem Text haben, der die Sprache der Untertitel beschreibt.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

Ihr `track` Element sollte ein `label` Attribut mit einem Text haben, der die Sprache der Untertitel beschreibt.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

Das dritte `section` sollte ein `h2` Element enthalten.

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

Ihr `h2` sollte einen Text enthalten, der den Titel des Abschnitts darstellt.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

Das dritte `section` sollte ein `p` Element für das Transkript enthalten.

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
