---
id: 671fa769540d152e53327f4b
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

选择 `#scary-btn`、`#funny-btn` 和 `#adventure-btn` 按钮，并将它们存储在名为 `scaryStoryBtn`、`funnyStoryBtn` 和 `adventureStoryBtn` 的变量中。

# --hints--

你应该使用 `getElementById` 方法选择 `#scary-btn` 按钮并将其存储在 `scaryStoryBtn` 变量中。

```js
assert.isDefined(scaryStoryBtn);
assert.deepEqual(scaryStoryBtn, document.getElementById("scary-btn"));
```

你应该使用 `getElementById` 方法选择 `#funny-btn` 按钮并将其存储在 `funnyStoryBtn` 变量中。

```js
assert.isDefined(funnyStoryBtn);
assert.deepEqual(funnyStoryBtn, document.getElementById("funny-btn"));
```

你应该使用 `getElementById` 方法选择 `#adventure-btn` 按钮并将其存储在 `adventureStoryBtn` 变量中。

```js
assert.isDefined(adventureStoryBtn);
assert.deepEqual(adventureStoryBtn, document.getElementById("adventure-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Storytelling App</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>

    <h1>Want to hear a short story?</h1>

    <main class="story-container">
        <div class="btn-container">
            <button class="btn" id="scary-btn">Scary Story</button>
            <button class="btn" id="funny-btn">Funny Story</button>
            <button class="btn" id="adventure-btn">Adventure Story</button>
        </div>
        <p id="result"></p>
    </main>
    <script src="./script.js"></script>
</body>
</html>
```

```js
--fcc-editable-region--
const storyContainer = document.querySelector(".story-container");

--fcc-editable-region--
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --dark-grey: #1b1b32;
    --black: #000;
    --white: #fff;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
}

body {
    background-color: var(--dark-grey);
    color: var(--white);
}

h1,
#result {
    text-align: center;
}

h1 {
    margin: 10px 0 15px;
}

.story-container {
    margin: auto;
    padding: 10px;
    width: 80%;
    border-style: double;
    border-width: 14px;
    border-color: var(--white);
}

.btn-container {
    display: flex;
    flex-direction: column;
    justify-content: center;
    align-items: center;
}

@media (min-width: 760px) {
    .btn-container {
        flex-direction: row;
    }
}

#result {
    margin-top: 15px;
    font-size: 1.2rem;
    line-height: 30px;
}

.btn {
    cursor: pointer;
    width: 200px;
    margin: 10px 0 10px 0.5rem;
    color: var(--black);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}
```
