---
id: 68e4c81577243e4db1676491
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將透過建置一個包含 Quincy Larson 求職技巧的網頁來練習使用語意化 HTML。基本的 HTML 範本已為你準備好。

首先建立一個 `h1` 元素，內容為 `Quincy's Tips for Getting a Developer Job`。

# --hints--

你應該有一個 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素應該有文字 `Quincy's Tips for Getting a Developer Job`。

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
