---
id: 68e964b84f954c322376e592
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

如果引用的來源是網站，你可以使用 `cite` 屬性來引用它。此屬性的值應該是有效的 URL。雖然此屬性不會改變區塊引用的呈現方式，但它對於提供螢幕閱讀器和搜尋引擎更多關於引用的資訊非常有幫助。

這是一個帶有 `cite` 屬性的行內引用元素範例：

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

為你的行內引用元素添加帶有此 URL 的 `cite` 屬性：

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

你的 `q` 元素應該有一個 `cite` 屬性。

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

你的 `q` 元素的 `cite` 屬性應該具有值 `https://www.freecodecamp.org/news/learn-to-code-book/`。

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
