---
id: 68e90335bd53a901c41dfc13
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在之前的課程中，你學到 `q` 元素用於區分引用的文字與周圍的內容。

這裡有一個範例：

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

大多數現代瀏覽器在你使用 `q` 元素時，會自動為行內引用加上引號。

在段落內將 `You can become a developer.` 包裹在行內引用元素中，其餘段落保持不變。

# --hints--

你應該有一個 `q` 元素。

```js
assert.exists(document.querySelector('q'));
```

你的 `q` 元素應該巢狀在位於 `h1` 元素下方的段落內。

```js
assert.exists(document.querySelector('h1 + p > q'));
```

你的 `q` 元素應該包含文字 `You can become a developer.`

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

在你的 `q` 元素內，文字不得有前置或後置空白（第一個字元前或最後一個字元後不得有空格、製表符或換行）。

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

你應該只在句子的適當零件周圍添加 `q` 頁籤，而不更改原始文字。`p` 元素仍應顯示：`Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
