---
id: 68ea9d23318b6481ac3de777
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

就像行內引用元素一樣，你也可以為區塊引用元素添加 `cite` 屬性。

以下是一個帶有 `cite` 屬性的區塊引用元素範例：

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

現在，為區塊引用元素添加一個 `cite` 屬性，並設置 URL 為 `https://www.freecodecamp.org/news/learn-to-code-book/`。

# --hints--

你的 `blockquote` 元素應該有一個 `cite` 屬性。

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.exists(firstBlockquoteEl?.getAttribute('cite'));
```

`cite` 屬性應該具有值 `https://www.freecodecamp.org/news/learn-to-code-book/`。

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
      --fcc-editable-region--
        <blockquote>
      --fcc-editable-region--
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
