---
id: 68eaa9ef8a0867071f195f59
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

到目前為止，你一直在使用 `cite` 屬性來標示引用的來源，但該屬性並不會真正向使用者顯示來源。

如果你想要以視覺方式標示來源，你可以在區塊引用元素外部添加一個引用元素 `cite`。引用元素是一個 HTML 元素，你可以用來標記被引用創作作品的標題，例如書籍、文章、歌曲、電影、網站或研究論文。

這裡有一個範例：

```html
<div>
  <blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
    The first thing you should consider about education is this is an economic decision.
  </blockquote>
  <p>&mdash;Quincy Larson, <cite>Is College Still Worth it? Tips from my 20 Years in Adult Education</cite></p>
</div>
```

在區塊引用元素下方的 `p` 元素內，將 `How to Learn to Code and Get a Developer Job [Full Book]` 包裹在 `cite` 元素中。

# --hints--

你應該在 `blockquote` 元素下方的 `p` 元素內包含一個 `cite` 元素。

```js
assert.exists(document.querySelector('blockquote + p > cite'));
```

你的 `cite` 元素應該包含文字 `How to Learn to Code and Get a Developer Job [Full Book]`。

```js
const citeEl = document.querySelector('blockquote + p > cite');
assert.equal(citeEl?.innerText.trim(), 'How to Learn to Code and Get a Developer Job [Full Book]');
```

忽略你剛剛添加的 `cite` 元素，`p` 元素的文字仍應該是 `—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]`。

```js
const pElCleanedText = document.querySelector('blockquote + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
      --fcc-editable-region--
          &mdash;Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]
      --fcc-editable-region--
        </p>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
