---
id: 68ea7e7b916e8270483689b6
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

在 `h2` 元素下方，你將添加另一則由 Quincy 所說的引言。這次，引言不會是較大段落的一部分。相反地，整個段落將會是一則引言。為了區分像這樣的引用文字，你應該使用區塊引用元素：`blockquote`。在瀏覽器中，你會看到文字稍微有縮排。

這是一個帶有引用文字的區塊引用元素範例：

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

現在，在第一個區段內，於 `h2` 元素下方添加一個區塊引用元素，內容為 `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`。

# --hints--

在第一個 `section` 元素內，你應該在 `h2` 元素下方放置一個 `blockquote` 元素。

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

你的 `blockquote` 元素應該包含文字 `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`。

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
