---
id: 587d7788367417b2b2512aa2
title: جعل التنقل لقارئ الشاشة أسهل باستخدام معلم nav
challengeType: 0
videoUrl: 'https://scrimba.com/c/czVwWSv'
forumTopicId: 301024
dashedName: make-screen-reader-navigation-easier-with-the-nav-landmark
---

# --description--

عنصر `nav` هو عنصر آخر في HTML5 يحتوي على ميزة المعلم المضمّن لتسهيل التنقل لقارئ الشاشة. يُستخدم هذا الوسم لتغليف روابط التنقل الرئيسية في صفحتك.

إذا كانت هناك روابط متكررة للموقع في أسفل الصفحة، فلا حاجة لترميزها باستخدام وسم `nav` أيضًا. يكفي استخدام وسم `footer` (الذي ستتعلمه في التحدٍّ التالي).

# --instructions--

أدرج Camper Cat روابط التنقل في أعلى صفحة التدريب الخاصة به، لكنه غلفها داخل وسم `div`. غيّر وسم `div` إلى وسم `nav` لتحسين إمكانية الوصول في صفحته.

# --hints--

يجب أن يحتوي شفرتك على وسم `nav` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('nav') , 1);
```

يجب أن تغلف وسوم `nav` عنصر `ul` وعناصر القائمة التابعة له.

```js
const nav = document.querySelector('nav');
const children = nav?.querySelectorAll(`:scope ${'ul'}`);
assert.lengthOf(children,1);
```

يجب ألا تحتوي شفرتك على أي وسوم `div`.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

يجب أن يحتوي عنصر `nav` على وسم إغلاق.

```js
assert.isTrue(code.match(/<\/nav>/g)?.length === code.match(/<nav>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training with Camper Cat</h1>

    <div>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </div>

  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training with Camper Cat</h1>

    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>

  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
