---
id: 587d7787367417b2b2512aa1
title: جعل التنقل باستخدام قارئ الشاشة أسهل مع معلم القسم العلوي
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB76vtv'
forumTopicId: 301023
dashedName: make-screen-reader-navigation-easier-with-the-header-landmark
---

# --description--

العنصر التالي في HTML5 الذي يضيف المعنى الدلالي ويحسن إمكانية الوصول هو وسم `header`. يُستخدم لتغليف المعلومات التمهيدية أو روابط التنقل للوسم الأب، ويعمل جيدًا حول المحتوى الذي يتكرر في الأعلى عبر صفحات متعددة.

يشارك وسم `header` ميزة المعلم المضمنة التي رأيتها مع `main`، مما يسمح لتقنيات المساعدة بالتنقل السريع إلى ذلك المحتوى.

**ملاحظة:** وسم `header` مخصص للاستخدام داخل وسم `body` في وثيقة HTML الخاصة بك. وهو مختلف عن عنصر `head`، الذي يحتوي على عنوان الصفحة والمعلومات الوصفية وغيرها.

# --instructions--

يكتب Camper Cat بعض المقالات الرائعة عن تدريب النينجا، ويريد إضافة صفحة لها في موقعه. غيّر وسم `div` العلوي الذي يحتوي حاليًا على `h1` إلى وسم `header` بدلاً منه.

# --hints--

يجب أن يحتوي شفرتك على وسم `header` واحد.

```js
assert.lengthOf(document.querySelectorAll('header'),1);
```

يجب أن تغلف وسوم `header` وسم `h1`.

```js
const header = document.querySelector('header');
const children = header?.querySelectorAll(`:scope ${'h1'}`);
assert.lengthOf(children , 1);
```

يجب ألا تحتوي شفرتك على أي وسوم `div`.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

يجب أن يحتوي وسم `header` على وسم إغلاق.

```js
assert.isTrue(code.match(/<\/header>/g)?.length === code.match(/<header>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>

  <div>
    <h1>Training with Camper Cat</h1>
  </div>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>

  <header>
    <h1>Training with Camper Cat</h1>
  </header>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
