---
id: 587d778d367417b2b2512aaa
title: اجعل العناصر مرئية فقط لقارئ الشاشة باستخدام CSS مخصص
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJ8QGkhJ'
forumTopicId: 301020
dashedName: make-elements-only-visible-to-a-screen-reader-by-using-custom-css
---

# --description--

هل لاحظت أن جميع تحديات إمكانية الوصول التي طبقتها حتى الآن لم تستخدم أي CSS؟ هذا يوضح أهمية استخدام مخطط منطقي للوثيقة وأوسمة ذات معنى دلالي حول المحتوى قبل إدخال جانب التصميم البصري.

مع ذلك، يمكن لسحر CSS تحسين إمكانية الوصول في صفحتك عندما تريد إخفاء المحتوى بصريًا المخصص فقط لقارئي الشاشة. يحدث هذا عندما تكون المعلومات في شكل بصري (مثل مخطط)، لكن مستخدمي قارئ الشاشة يحتاجون إلى عرض بديل (مثل جدول) للوصول إلى البيانات. يُستخدم CSS لوضع العناصر المخصصة لقارئ الشاشة خارج مساحة العرض في نافذة المتصفح.

إليك مثالًا على قواعد CSS التي تحقق ذلك:

```css
.sr-only {
  position: absolute;
  left: -10000px;
  width: 1px;
  height: 1px;
  top: auto;
  overflow: hidden;
}
```

**ملاحظة:** الطرق التالية في CSS لن تفعل الشيء نفسه:

<ul>
<li><code>display: none;</code> أو <code>visibility: hidden;</code> يخفي المحتوى عن الجميع، بما في ذلك مستخدمي قارئ الشاشة</li>
<li>القيم الصفرية لأحجام البيكسل، مثل <code>width: 0px; height: 0px;</code> تزيل ذلك العنصر من تدفق الوثيقة، مما يعني أن قارئي الشاشة سيتجاهلونه</li>
</ul>

# --instructions--

أنشأ Camper Cat مخططًا بيانيًا مكدسًا رائعًا لصفحة تدريبه، ووضع البيانات في جدول لمستخدميه من ضعاف البصر. الجدول يحتوي بالفعل على صنف `sr-only`، لكن قواعد CSS لم تُملأ بعد. اجعل خاصية `position` بقيمة `absolute`، و`left` بقيمة `-10000px`، و`width` و`height` كلاهما بقيمة `1px`.

# --hints--

يجب أن يضبط شفرتك خاصية `position` للصنف `sr-only` إلى القيمة `absolute`.

```js
const srOnly = document.querySelector('.sr-only');
const position = window.getComputedStyle(srOnly).position; 
assert.equal(position, 'absolute');
```

يجب أن يضبط شفرتك خاصية `left` للصنف `sr-only` إلى القيمة `-10000px`.

```js
const srOnly = document.querySelector('.sr-only');
const left = window.getComputedStyle(srOnly).left;
assert.equal(left, '-10000px');
```

يجب أن يضبط شفرتك خاصية `width` للصنف `sr-only` إلى قيمة `1` بيكسل.

```js
assert.match(code , /width:\s*?1px/gi);
```

يجب أن يضبط شفرتك خاصية `height` للصنف `sr-only` إلى قيمة `1` بيكسل.

```js
assert.match(code , /height:\s*?1px/gi);
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  .sr-only {
    position: ;
    left: ;
    width: ;
    height: ;
    top: auto;
    overflow: hidden;
  }
  </style>
</head>
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <section>
    <h2>Master Camper Cat's Beginner Three Week Training Program</h2>
    <figure>
      <!-- Stacked bar chart of weekly training -->
      <p>[Stacked bar chart]</p>
      <br />
      <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
    </figure>
    <table class="sr-only">
      <caption>Hours of Weekly Training in Stealth, Combat, and Weapons</caption>
      <thead>
        <tr>
          <th></th>
          <th scope="col">Stealth &amp; Agility</th>
          <th scope="col">Combat</th>
          <th scope="col">Weapons</th>
          <th scope="col">Total</th>
        </tr>
      </thead>
      <tbody>
        <tr>
          <th scope="row">Week One</th>
          <td>3</td>
          <td>5</td>
          <td>2</td>
          <td>10</td>
        </tr>
        <tr>
          <th scope="row">Week Two</th>
          <td>4</td>
          <td>5</td>
          <td>3</td>
          <td>12</td>
        </tr>
        <tr>
          <th scope="row">Week Three</th>
          <td>4</td>
          <td>6</td>
          <td>3</td>
          <td>13</td>
        </tr>
      </tbody>
    </table>
  </section>
  <section id="stealth">
    <h2>Stealth &amp; Agility Training</h2>
    <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
    <article><h3>No training is NP-complete without parkour</h3></article>
  </section>
  <section id="combat">
    <h2>Combat Training</h2>
    <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
    <article><h3>Goodbye, world: 5 proven ways to knock out an opponent</h3></article>
  </section>
  <section id="weapons">
    <h2>Weapons Training</h2>
    <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
    <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  .sr-only {
    position: absolute;
    left: -10000px;
    width: 1px;
    height: 1px;
    top: auto;
    overflow: hidden;
  }
  </style>
</head>
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <section>
    <h2>Master Camper Cat's Beginner Three Week Training Program</h2>
    <figure>
      <!-- Stacked bar chart of weekly training -->
      <p>[Stacked bar chart]</p>
      <br />
      <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
    </figure>
    <table class="sr-only">
      <caption>Hours of Weekly Training in Stealth, Combat, and Weapons</caption>
      <thead>
        <tr>
          <th></th>
          <th scope="col">Stealth &amp; Agility</th>
          <th scope="col">Combat</th>
          <th scope="col">Weapons</th>
          <th scope="col">Total</th>
        </tr>
      </thead>
      <tbody>
        <tr>
          <th scope="row">Week One</th>
          <td>3</td>
          <td>5</td>
          <td>2</td>
          <td>10</td>
        </tr>
        <tr>
          <th scope="row">Week Two</th>
          <td>4</td>
          <td>5</td>
          <td>3</td>
          <td>12</td>
        </tr>
        <tr>
          <th scope="row">Week Three</th>
          <td>4</td>
          <td>6</td>
          <td>3</td>
          <td>13</td>
        </tr>
      </tbody>
    </table>
  </section>
  <section id="stealth">
    <h2>Stealth &amp; Agility Training</h2>
    <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
    <article><h3>No training is NP-complete without parkour</h3></article>
  </section>
  <section id="combat">
    <h2>Combat Training</h2>
    <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
    <article><h3>Goodbye, world: 5 proven ways to knock out an opponent</h3></article>
  </section>
  <section id="weapons">
    <h2>Weapons Training</h2>
    <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
    <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
