---
id: 587d778e367417b2b2512aab
title: تحسين قابلية القراءة باستخدام نص عالي التباين
challengeType: 0
videoUrl: 'https://scrimba.com/c/cKb3nCq'
forumTopicId: 301017
dashedName: improve-readability-with-high-contrast-text
---

# --description--

يمكن أن يجعل التباين المنخفض بين ألوان المقدمة والخلفية النص صعب القراءة. التباين الكافي يحسن من قابلية قراءة المحتوى، لكن ماذا يعني "كافي" بالضبط؟

توصي إرشادات إمكانية الوصول لمحتوى الويب (WCAG) بنسبة تباين لا تقل عن 4.5 إلى 1 للنص العادي. تُحسب النسبة بمقارنة قيم الإضاءة النسبية بين لونين. تتراوح هذه النسبة من 1:1 لنفس اللون، أي بدون تباين، إلى 21:1 للأبيض مقابل الأسود، وهو أعلى تباين ممكن. هناك العديد من أدوات فحص التباين المتاحة على الإنترنت التي تحسب هذه النسبة لك.

# --instructions--

اختيار Camper Cat لنص رمادي فاتح على خلفية بيضاء في منشوره الأخير على المدونة له نسبة تباين 1.5:1، مما يجعل القراءة صعبة. غيّر `color` النص من الرمادي الحالي (`#D3D3D3`) إلى رمادي أغمق (`#636363`) لتحسين نسبة التباين إلى 6:1.

# --hints--

يجب أن يغيّر شفرتك `color` النص لعنصر `body` إلى الرمادي الأغمق.

```js
const body = document.querySelector('body');
const bodyColor = window.getComputedStyle(body).color; 
assert(bodyColor == 'rgb(99, 99, 99)');
```

يجب ألا تغيّر شفرتك `background-color` لعنصر `body`.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body).backgroundColor; 
assert.equal(backgroundColor , 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: #D3D3D3;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: #636363;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
