---
id: 587d774e367417b2b2512aa0
title: لف المحتوى داخل عنصر article
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp79S3'
forumTopicId: 301029
dashedName: wrap-content-in-the-article-element
---

# --description--

`article` هو أحد عناصر HTML5 الجديدة التي تضيف المعنى الدلالي إلى ترميزك. `article` هو عنصر تقسيم يُستخدم لاحتواء محتوى مستقل ومكتفٍ بذاته. يعمل هذا الوسم جيدًا مع مدخلات المدونات، منشورات المنتديات، أو المقالات الإخبارية.

تحديد ما إذا كان المحتوى يمكن أن يقف بمفرده عادةً ما يكون قرارًا تقديريًا، لكن يمكنك استخدام بعض الاختبارات البسيطة. اسأل نفسك إذا أزلت كل السياق المحيط، هل سيظل هذا المحتوى منطقيًا؟ وبالمثل، بالنسبة للنص، هل سيظل المحتوى صالحًا إذا وُضع في خلاصة RSS؟

تذكر أن الأشخاص الذين يستخدمون تقنيات المساعدة يعتمدون على ترميز منظم وذو معنى دلالي لفهم عملك بشكل أفضل.

**ملاحظة:** عنصر `section` هو أيضًا جديد في HTML5، وله معنى دلالي مختلف قليلاً عن `article`. يستخدم `article` للمحتوى المستقل، بينما يستخدم `section` لتجميع المحتوى ذي الصلة الموضوعية. يمكن استخدامهما داخل بعضهما حسب الحاجة. على سبيل المثال، إذا كان الكتاب هو `article`، فكل فصل هو `section`. وعندما لا يكون هناك علاقة بين مجموعات المحتوى، فاستخدم `div`.

`<div>` - يجمع المحتوى
`<section>` - يجمع المحتوى ذي الصلة
`<article>` - يجمع المحتوى المستقل والمكتفٍ بذاته

# --instructions--

استخدم Camper Cat وسم `article` لاحتواء المنشورات في صفحة مدونته، لكنه نسي استخدامه حول المنشور الأعلى. غيّر وسم `div` إلى وسم `article` بدلاً منه.

# --hints--

يجب أن يحتوي شفرتك على ثلاثة وسوم `article`.

```js
assert.lengthOf(document.querySelectorAll('article') ,3);
```

يجب ألا تحتوي شفرتك على أي وسم `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <div>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </div>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```
