---
id: 587d774d367417b2b2512a9e
title: استخدام العناوين لإظهار العلاقات الهرمية للمحتوى
challengeType: 0
videoUrl: 'https://scrimba.com/c/cqVEktm'
forumTopicId: 301026
dashedName: use-headings-to-show-hierarchical-relationships-of-content
---

# --description--

العناوين (`h1` إلى `h6`) هي وُسوم أساسية تساعد في توفير هيكل وتسميات لمحتواك. يمكن ضبط قارئات الشاشة لقراءة العناوين فقط في الصفحة ليحصل المستخدم على ملخص. هذا يعني أهمية أن تحمل وُسوم العناوين في ترميزك معنى دلاليًا وأن ترتبط ببعضها البعض، لا أن تُختار فقط بناءً على قيم حجمها.

*المعنى الدلالي* يعني أن الوسم الذي تستخدمه حول المحتوى يشير إلى نوع المعلومات التي يحتويها.

إذا كنت تكتب ورقة تحتوي على مقدمة وجسم وخاتمة، فلن يكون من المنطقي أن تجعل الخاتمة فرعًا فرعيًا من الجسم في مخططك. يجب أن تكون قسمًا مستقلًا. وبالمثل، يجب أن تتبع وُسوم العناوين في صفحة الويب ترتيبًا معينًا وتُظهر العلاقات الهرمية لمحتواك.

العناوين ذات الرتبة المتساوية (أو الأعلى) تبدأ أقسامًا جديدة ضمنية، والعناوين ذات الرتبة الأدنى تبدأ فروعًا فرعية من القسم السابق.

كمثال، صفحة تحتوي على وسم `h2` يتبعه عدة فروع فرعية معنونة بـ `h4` ستربك مستخدم قارئ الشاشة. مع وجود ستة خيارات، قد تميل لاستخدام وسم لأنه يبدو أفضل في المتصفح، لكن يمكنك استخدام CSS لتعديل الحجم النسبي.

نقطة أخيرة، يجب أن تحتوي كل صفحة دائمًا على وسم `h1` واحد فقط، وهو الموضوع الرئيسي لمحتواك. هذا الوسم والعناوين الأخرى تُستخدم جزئيًا من قبل محركات البحث لفهم موضوع الصفحة.

# --instructions--

يريد Camper Cat صفحة على موقعه مخصصة لتعلم أن تصبح نينجا. ساعده في تصحيح العناوين بحيث يعطي ترميزه معنى دلاليًا للمحتوى، ويُظهر العلاقات الصحيحة بين الأقسام الرئيسية والفرعية. غيّر كل وُسوم `h5` إلى المستوى الصحيح من العناوين لتدل على أنها فروع فرعية من وُسوم `h2`. استخدم وُسوم `h3` لهذا الغرض.

# --hints--

يجب أن يحتوي شفرتك على 6 وُسوم `h3`.

```js
assert.lengthOf(document.querySelectorAll('h3') , 6);
```

يجب أن يحتوي شفرتك على 6 وُسوم إغلاق `h3`.

```js
assert.lengthOf((code.match(/\/h3/g) || []) ,6);
```

يجب ألا يحتوي شفرتك على أي وُسوم `h5`.

```js
assert.isEmpty(document.querySelectorAll('h5'));
```

يجب ألا يحتوي شفرتك على أي وُسوم إغلاق `h5`.

```js
assert.notMatch(code, /\/h5/);
```

# --seed--

## --seed-contents--

```html
<h1>How to Become a Ninja</h1>
<main>
  <h2>Learn the Art of Moving Stealthily</h2>
  <h5>How to Hide in Plain Sight</h5>
  <h5>How to Climb a Wall</h5>

  <h2>Learn the Art of Battle</h2>
  <h5>How to Strengthen your Body</h5>
  <h5>How to Fight like a Ninja</h5>

  <h2>Learn the Art of Living with Honor</h2>
  <h5>How to Breathe Properly</h5>
  <h5>How to Simplify your Life</h5>
</main>
```

# --solutions--

```html
<h1>How to Become a Ninja</h1>
<main>
  <h2>Learn the Art of Moving Stealthily</h2>
  <h3>How to Hide in Plain Sight</h3>
  <h3>How to Climb a Wall</h3>

  <h2>Learn the Art of Battle</h2>
  <h3>How to Strengthen your Body</h3>
  <h3>How to Fight like a Ninja</h3>

  <h2>Learn the Art of Living with Honor</h2>
  <h3>How to Breathe Properly</h3>
  <h3>How to Simplify your Life</h3>
</main>
```
