---
id: 587d778f367417b2b2512aac
title: تجنب مشاكل عمى الألوان باستخدام نسبة تباين كافية
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMEUw'
forumTopicId: 301012
dashedName: avoid-colorblindness-issues-by-using-sufficient-contrast
---

# --description--

اللون جزء كبير من التصميم البصري، لكن استخدامه يطرح مشكلتين في إمكانية الوصول. أولًا، لا يجب استخدام اللون وحده كطريقة وحيدة لنقل معلومات مهمة لأن مستخدمي قارئ الشاشة لن يروا اللون. ثانيًا، يجب أن يكون هناك نسبة تباين كافية بين ألوان المقدمة والخلفية حتى يتمكن المصابون بعمى الألوان من التمييز بينها.

التحديات السابقة تناولت وجود بدائل نصية لمعالجة المشكلة الأولى. أما التحدي الأخير فقدم أدوات لفحص نسبة التباين لمساعدتك في المشكلة الثانية. نسبة التباين الموصى بها من WCAG وهي 4.5:1 تنطبق على استخدام الألوان وكذلك على التركيبات ذات التدرج الرمادي.

يواجه المصابون بعمى الألوان صعوبة في التمييز بين بعض الألوان - عادة في درجة اللون (hue) وأحيانًا في فاتح اللون أيضًا. قد تتذكر أن نسبة التباين تُحسب باستخدام قيم اللمعان النسبي (أو فاتح اللون) لألوان المقدمة والخلفية.

عمليًا، يمكن الوصول إلى نسبة التباين 4.5:1 عن طريق التظليل (إضافة الأسود) للون الأغمق والتلوين (إضافة الأبيض) للون الأفتح. الظلال الأغمق على عجلة الألوان تُعتبر ظلالًا من الأزرق والبنفسجي والماجنتا والأحمر، بينما الألوان الملوّنة بالأبيض الأفتح هي البرتقالي والأصفر والأخضر والأزرق-الأخضر.

# --instructions--

يقوم Camper Cat بتجربة استخدام اللون لنص مدونته وخلفيتها، لكن تركيبة اللون الحالية التي تجمع بين `background-color` أخضر مائل و`color` نص خمري لها نسبة تباين 2.5:1. يمكنك تعديل فاتح اللون بسهولة لأنهما معلنان باستخدام خاصية CSS `hsl()` (التي تعني درجة اللون، التشبع، فاتح اللون) بتغيير المعلمة الثالثة. زد قيمة فاتح اللون في `background-color` من 35% إلى 55%، وقلل قيمة فاتح اللون في `color` من 20% إلى 15%. هذا يحسن نسبة التباين إلى 5.9:1.

# --hints--

يجب أن يغير شفرتك فقط قيمة فاتح اللون لخاصية نص `color` إلى 15%.

```js
assert.match(code ,/color:\s*?hsl\(0,\s*?55%,\s*?15%\)/gi);
```

يجب أن يغير شفرتك فقط قيمة فاتح اللون لخاصية `background-color` إلى 55%.

```js
assert.match(code ,/background-color:\s*?hsl\(120,\s*?25%,\s*?55%\)/gi);
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 20%);
    background-color: hsl(120, 25%, 35%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 15%);
    background-color: hsl(120, 25%, 55%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
