---
id: 587d778f367417b2b2512aad
title: تجنب مشاكل عمى الألوان باختيار الألوان التي تنقل المعلومات بعناية
  Avoid Colorblindness Issues by Carefully Choosing Colors that Convey
  Information
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437as3'
forumTopicId: 301011
dashedName: >-
  avoid-colorblindness-issues-by-carefully-choosing-colors-that-convey-information
---

# --description--

هناك أشكال مختلفة من عمى الألوان. تتراوح هذه من انخفاض الحساسية لطول موجي معين من الضوء إلى عدم القدرة على رؤية اللون على الإطلاق. الشكل الأكثر شيوعًا هو انخفاض الحساسية لاكتشاف اللون الأخضر.

على سبيل المثال، إذا كان لونان أخضران متشابهان هما لون المقدمة والخلفية لمحتواك، فقد لا يتمكن المستخدم المصاب بعمى الألوان من التمييز بينهما. يمكن اعتبار الألوان المتقاربة جيرانًا على عجلة الألوان، ويجب تجنب تلك التركيبات عند نقل معلومات مهمة.

**ملاحظة:** تتضمن بعض أدوات اختيار الألوان عبر الإنترنت محاكاة بصرية لكيفية ظهور الألوان لأنواع مختلفة من عمى الألوان. هذه موارد رائعة بالإضافة إلى حاسبات التحقق من نسبة التباين عبر الإنترنت.

# --instructions--

يقوم Camper Cat باختبار أنماط مختلفة لزر مهم، لكن اللون الأصفر (`#FFFF33`) في `background-color` واللون الأخضر (`#33FF33`) في نص `color` هما درجتان متجاورتان على عجلة الألوان ولا يمكن تمييزهما عمليًا لبعض المستخدمين المصابين بعمى الألوان. (تشابه فاتح اللون لديهما يفشل أيضًا في اختبار نسبة التباين). غيّر نص `color` إلى أزرق داكن (`#003366`) لحل المشكلتين.

# --hints--

يجب أن يغيّر شفرتك نص `color` للزر `button` إلى الأزرق الداكن.

```js
const button = document.querySelector('button');
const buttonColor = window.getComputedStyle(button).color; 
assert.equal(buttonColor, 'rgb(0, 51, 102)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  button {
    color: #33FF33;
    background-color: #FFFF33;
    font-size: 14px;
    padding: 10px;
  }
  </style>
</head>
<body>
  <header>
    <h1>Danger!</h1>
  </header>
  <button>Delete Internet</button>
</body>
```

# --solutions--

```html
<head>
  <style>
    button {
      color: #003366;
      background-color: #FFFF33;
      font-size: 14px;
      padding: 10px;
    }
  </style>
</head>
<body>
  <header>
    <h1>Danger!</h1>
  </header>
  <button>Delete Internet</button>
</body>
```
