---
id: 587d778a367417b2b2512aa6
title: تحسين إمكانية الوصول لحقول النموذج باستخدام عنصر label
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMMAN'
forumTopicId: 301016
dashedName: improve-form-field-accessibility-with-the-label-element
---

# --description--

تحسين إمكانية الوصول باستخدام ترميز HTML الدلالي يشمل استخدام أسماء الوسوم والخصائص المناسبة. التحديات التالية تغطي بعض السيناريوهات المهمة لاستخدام الخصائص في النماذج.

وسم `label` يحيط بالنص الخاص بعنصر تحكم معين في النموذج، وغالبًا ما يكون اسم الخيار أو مسماه. هذا يربط المعنى بالعنصر ويجعل النموذج أسهل في القراءة. الخاصية `for` في وسم `label` تربط هذا الـ `label` صراحةً بعنصر التحكم في النموذج ويستخدمها قارئ الشاشة.

تعلمت عن أزرار الاختيار ومسامياتها في درس ضمن قسم HTML الأساسي. في ذلك الدرس، وضعنا عنصر الإدخال الخاص بزر الاختيار داخل وسم `label` مع نص المسمى لجعل النص قابلًا للنقر. طريقة أخرى لتحقيق ذلك هي باستخدام الخاصية `for`، كما هو موضح في هذا الدرس.

يجب أن تكون قيمة الخاصية `for` هي نفسها قيمة الخاصية `id` لعنصر التحكم في النموذج. إليك مثالًا:

```html
<form>
  <label for="name">Name:</label>
  <input type="text" id="name" name="name">
</form>
```

# --instructions--

يتوقع Camper Cat اهتمامًا كبيرًا بمنشوراته المدونة المدروسة ويريد تضمين نموذج تسجيل بريد إلكتروني. أضف خاصية `for` على مسمى البريد الإلكتروني `label` بحيث تطابق قيمة `id` في حقل الإدخال `input` الخاص به.

# --hints--

يجب أن يحتوي شفرتك على خاصية `for` على وسم `label` لا تكون فارغة.

```js
assert.isNotEmpty(document.querySelector('label')?.getAttribute('for'));
```

يجب أن تطابق قيمة الخاصية `for` قيمة الخاصية `id` في حقل الإدخال الخاص بالبريد الإلكتروني `input`.

```js
assert.equal(document.querySelector('label')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label>Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
