---
id: 587d778b367417b2b2512aa8
title: إضافة منتقي تاريخ يمكن الوصول إليه
challengeType: 0
videoUrl: 'https://scrimba.com/c/cR3bRbCV'
forumTopicId: 301008
dashedName: add-an-accessible-date-picker
---

# --description--

غالبًا ما تتضمن النماذج حقل `input`، الذي يمكن استخدامه لإنشاء عدة عناصر تحكم في النموذج مختلفة. الخاصية `type` على هذا العنصر تحدد نوع عنصر `input` الذي سيتم إنشاؤه.

ربما لاحظت أنواع الإدخال `text` و `submit` في التحديات السابقة، وقد قدم HTML5 خيارًا لتحديد حقل `date`. اعتمادًا على دعم المتصفح، يظهر منتقي التاريخ في حقل `input` عند تركيزه، مما يسهل ملء النموذج لجميع المستخدمين.

في المتصفحات القديمة، سيكون النوع الافتراضي `text`، لذا من المفيد عرض تنسيق التاريخ المتوقع للمستخدمين في نص الـ `label` أو `placeholder` تحسبًا لذلك.

إليك مثالًا:

```html
<label for="input1">Enter a date:</label>
<input type="date" id="input1" name="input1">
```

# --instructions--

يقوم Camper Cat بإعداد بطولة Mortal Kombat ويريد سؤال منافسيه لمعرفة التاريخ الأنسب. أضف وسم `input` مع خاصية `type` بقيمة `date`، وخاصية `id` بقيمة `pickdate`، وخاصية `name` بقيمة `date`.

# --hints--

يجب أن تضيف شفرتك وسم `input` واحد لحقل محدد التاريخ.

```js
assert.lengthOf(document.querySelectorAll('input'), 2);
```

يجب أن يحتوي وسم `input` على خاصية `type` بقيمة `date`.

```js
assert.equal(document.querySelector('input')?.getAttribute('type'), 'date');
```

يجب أن يحتوي وسم `input` على خاصية `id` بقيمة `pickdate`.

```js
assert.equal(document.querySelector('input')?.getAttribute('id'),'pickdate');
```

يجب أن يحتوي وسم `input` على خاصية `name` بقيمة `date`.

```js
assert.equal(document.querySelector('input')?.getAttribute('name'), 'date');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>

        <!-- Only change code below this line -->



        <!-- Only change code above this line -->

        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>
        <input type="date" id="pickdate" name="date">
        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
