---
id: 587d778f367417b2b2512aae
title: أعطِ الروابط معنى باستخدام نص رابط وصفي
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437DcV'
forumTopicId: 301013
dashedName: give-links-meaning-by-using-descriptive-link-text
---

# --description--

لدى مستخدمي قارئ الشاشة خيارات متعددة لنوع المحتوى الذي يقرأه جهازهم. تشمل هذه الخيارات تخطي (أو تجاوز) عناصر العلامات المميزة، القفز إلى المحتوى الرئيسي، أو الحصول على ملخص للصفحة من العناوين. خيار آخر هو سماع الروابط المتاحة فقط على الصفحة.

يقوم قارئو الشاشة بذلك عن طريق قراءة نص الرابط، أو ما بين وسم الرابط (`a`). وجود قائمة بروابط مثل "نقر هنا" أو "اقرأ المزيد" ليس مفيدًا. بدلاً من ذلك، استخدم نصًا موجزًا لكنه وصفي داخل وسم `a` ليعطي معنى أكبر لهؤلاء المستخدمين.

# --instructions--

نص الرابط الذي يستخدمه Camper Cat ليس وصفيًا جدًا بدون السياق المحيط. حرّك وسم الرابط (`a`) ليغلف النص `information about batteries` بدلًا من `Click here`.

# --hints--

يجب أن يحرك شفرتك وسم الرابط `a` من حول الكلمات `Click here` ليغلف الكلمات `information about batteries`.

```js
assert.match(document.querySelector('a')?.textContent, /^(information about batteries)$/g);
```

يجب أن يحتوي عنصر `a` على خاصية `href` بقيمة سلسلة نصية فارغة `""`.

```js
assert.isEmpty(document.querySelector('a')?.getAttribute('href'));
```

يجب أن يحتوي عنصر `a` على وسم إغلاق.

```js
assert.isTrue(code.match(/<\/a>/g)?.length === code.match(/<a href=(''|"")>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. <a href="">Click here</a> for information about batteries</p>
  </article>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. Click here for <a href="">information about batteries</a></p>
  </article>
</body>
```
