---
id: 587d7790367417b2b2512aaf
title: جعل الروابط قابلة للتنقل باستخدام مفاتيح الوصول في HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQvmaTp'
forumTopicId: 301021
dashedName: make-links-navigable-with-html-access-keys
---

# --description--

تقدم HTML الخاصية `accesskey` لتحديد مفتاح اختصار لتنشيط عنصر أو جلب التركيز إليه. إضافة خاصية `accesskey` يمكن أن تجعل التنقل أكثر كفاءة لمستخدمي لوحة المفاتيح فقط.

يسمح HTML5 باستخدام هذه الخاصية على أي عنصر، لكنها مفيدة بشكل خاص عند استخدامها مع العناصر التفاعلية. يشمل ذلك الروابط، الأزرار، وعناصر التحكم في الموذج.

إليك مثالًا:

```html
<button accesskey="b">Important Button</button>
```

# --instructions--

يريد Camper Cat أن تحتوي الروابط حول عنواني المقالتين في المدونة على اختصارات لوحة مفاتيح حتى يتمكن مستخدمو موقعه من التنقل بسرعة إلى القصة الكاملة. أضف خاصية `accesskey` إلى كلا الرابطين واضبط الأول على `g` (لـ Garfield) والثاني على `c` (لـ Chuck Norris).

# --hints--

يجب أن تضيف شفرتك خاصية `accesskey` إلى وسم `a` الذي يحمل الخاصية `id` بقيمة `first`.

```js
assert.isTrue(document.querySelector('#first')?.hasAttribute('accesskey'));
```

يجب أن تضيف شفرتك خاصية `accesskey` إلى وسم `a` الذي يحمل الخاصية `id` بقيمة `second`.

```js
assert.isTrue(document.querySelector('#second')?.hasAttribute('accesskey'));
```

يجب أن يضبط شفرتك الخاصية `accesskey` على وسم `a` الذي يحمل `id` بقيمة `first` إلى `g`. لاحظ أن حالة الأحرف مهمة.

```js
assert.equal(document.querySelector('#first')?.getAttribute('accesskey'), 'g');
```

يجب أن يضبط شفرتك الخاصية `accesskey` على الوسم `a` الذي يحمل الخاصية `id` بقيمة `second` إلى `c`. لاحظ أن حالة الأحرف مهمة.

```js
assert.equal(document.querySelector('#second')?.getAttribute('accesskey'), 'c');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" accesskey="g" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" accesskey="c" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
