---
id: 587d7790367417b2b2512ab0
title: استخدام tabindex لإضافة تركيز لوحة المفاتيح إلى عنصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMDHW'
forumTopicId: 301027
dashedName: use-tabindex-to-add-keyboard-focus-to-an-element
---

# --description--

خاصية `tabindex` في HTML لها ثلاث وظائف مميزة تتعلق بتركيز لوحة المفاتيح على العنصر. عندما تكون على وسم، فهي تشير إلى أن العنصر يمكن تركيزه. القيمة (عدد صحيح موجب أو سالب أو صفر) تحدد السلوك.

بعض العناصر، مثل الروابط وعناصر التحكم في الموذج، تحصل تلقائيًا على تركيز لوحة المفاتيح عندما ينتقل المستخدم عبر الصفحة باستخدام التبويب. يكون ذلك بنفس ترتيب العناصر في ترميز المصدر HTML. يمكن إعطاء هذه الوظيفة نفسها لعناصر أخرى، مثل `div` و`span` و`p`، بوضع خاصية `tabindex="0"` عليها. إليك مثالًا:

```html
<div tabindex="0">I need keyboard focus!</div>
```

**ملاحظة:** قيمة `tabindex` السالبة (عادة -1) تشير إلى أن العنصر قابل للتركيز، لكنه غير قابل للوصول عبر لوحة المفاتيح. تُستخدم هذه الطريقة عادةً لجلب التركيز إلى المحتوى برمجيًا (مثل تفعيل `div` المستخدم في نافذة منبثقة)، وهي خارج نطاق هذه التحديات.

# --instructions--

أنشأ Camper Cat استبيانًا جديدًا لجمع معلومات عن مستخدميه. يعلم أن حقول الإدخال تحصل تلقائيًا على تركيز لوحة المفاتيح، لكنه يريد التأكد من أن مستخدمي لوحة المفاتيح يتوقفون عند التعليمات أثناء التنقل عبر العناصر باستخدام التبويب. أضف خاصية `tabindex` إلى وسم `p` واضبط قيمتها على `0`. مكافأة - استخدام `tabindex` يتيح أيضًا لفئة CSS المزيفة `:focus` العمل على وسم `p`.

# --hints--

يجب أن تضيف شفرتك خاصية `tabindex` إلى وسم `p` الذي يحتوي على تعليمات النموذج.

```js
assert.isNotNull(document.querySelector('p')?.getAttribute('tabindex'));
```

يجب أن تضبط شفرتك خاصية `tabindex` على وسم `p` بقيمة 0.

```js
assert.equal(document.querySelector('p')?.getAttribute('tabindex'), '0');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p>Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p tabindex="0">Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
