---
id: 587d778b367417b2b2512aa7
title: لف أزرار الاختيار في عنصر مجموعة الحقول لتحسين إمكانية الوصول
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVefw'
forumTopicId: 301030
dashedName: wrap-radio-buttons-in-a-fieldset-element-for-better-accessibility
---

# --description--

الموضوع التالي في النماذج يغطي إمكانية الوصول لأزرار الاختيار. يُعطى كل خيار `label` مع خاصية `for` تربطها بـ `id` للعنصر المقابل كما تناولناه في التحدٍّ السابق. بما أن أزرار الاختيار غالبًا ما تأتي في مجموعة يجب على المستخدم اختيار واحد منها، فهناك طريقة لعرض الخيارات كجزء من مجموعة بشكل دلالي.

وسم `fieldset` يحيط بالمجموعة الكاملة من أزرار الاختيار لتحقيق ذلك. غالبًا ما يستخدم وسم `legend` لتوفير وصف للمجموعة، يقرأه قارئ الشاشة لكل خيار داخل عنصر `fieldset`.

لا يكون غلاف `fieldset` ووسم `legend` ضروريين عندما تكون الخيارات واضحة بذاتها، مثل اختيار الجنس. استخدام `label` مع خاصية `for` لكل زر اختيار يكون كافيًا.

إليك مثالًا:

```html
<form>
  <fieldset>
    <legend>Choose one of these three items:</legend>
    <input id="one" type="radio" name="items" value="one">
    <label for="one">Choice One</label><br>
    <input id="two" type="radio" name="items" value="two">
    <label for="two">Choice Two</label><br>
    <input id="three" type="radio" name="items" value="three">
    <label for="three">Choice Three</label>
  </fieldset>
</form>
```

# --instructions--

يريد Camper Cat معلومات عن مستوى النينجا لمستخدميه عند تسجيلهم في قائمة بريده الإلكتروني. أضاف مجموعة من أزرار الاختيار وتعلم من درسنا السابق استخدام وسوم `label` مع خاصية `for` لكل خيار. هيا يا Camper Cat! مع ذلك، يحتاج شفرة إلى بعض التعديل. غيّر وسم `div` المحيط بأزرار الاختيار إلى وسم `fieldset`، وغيّر وسم `p` بداخله إلى وسم `legend`.

# --hints--

يجب أن يحتوي شفرتك على وسم `fieldset` حول مجموعة أزرار الاختيار.

```js
assert.lengthOf(document.querySelectorAll('fieldset'),1);
```

يجب أن يحتوي عنصر `fieldset` على وسم إغلاق.

```js
assert.isTrue(code.match(/<\/fieldset>/g)?.length === code.match(/<fieldset>/g)?.length);
```

يجب أن يحتوي شفرتك على وسم `legend` حول النص الذي يسأل عن مستوى النينجا للمستخدم.

```js
assert.lengthOf(document.querySelectorAll('legend') ,1);
```

يجب ألا يحتوي شفرتك على أي وسوم `div`.

```js
assert.lengthOf(document.querySelectorAll('div') ,0);
```

يجب ألا يحتوي شفرتك على وسم `p` حول النص الذي يسأل عن مستوى النينجا للمستخدم.

```js
assert.lengthOf(document.querySelectorAll('p') ,4);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>
      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <!-- Only change code below this line -->
      <div>
        <p>What level ninja are you?</p>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">Master</label>
      </div>
      <!-- Only change code above this line -->


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>
      <label for="email">Email:</label>
      <input type="text" id="email" name="email">

      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">Master</label>
      </fieldset>

      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
