---
id: 587d7789367417b2b2512aa4
title: تحسين إمكانية الوصول لمحتوى الصوت باستخدام عنصر audio
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVkcZ'
forumTopicId: 301014
dashedName: improve-accessibility-of-audio-content-with-the-audio-element
---

# --description--

عنصر `audio` في HTML5 يمنح المعنى الدلالي عند تغليفه لمحتوى الصوت أو تدفق الصوت في ترميزك. يحتاج محتوى الصوت أيضًا إلى بديل نصي ليكون متاحًا للأشخاص الصم أو ضعاف السمع. يمكن تحقيق ذلك بنص قريب على الصفحة أو برابط إلى نص مكتوب.

وسم `audio` يدعم الخاصية `controls`. هذا يعرض عناصر التحكم الافتراضية في المتصفح مثل التشغيل والإيقاف المؤقت وغيرها، ويدعم وظائف لوحة المفاتيح. هذه خاصية منطقية، بمعنى أنها لا تحتاج إلى قيمة، فوجودها على الوسم يُفعّل الإعداد.

إليك مثالًا:

```html
<audio id="meowClip" controls>
  <source src="audio/meow.mp3" type="audio/mpeg">
  <source src="audio/meow.ogg" type="audio/ogg">
</audio>
```

**ملاحظة:** عادةً ما يحتوي المحتوى المتعدد الوسائط على مكونات بصرية وسمعية معًا. يحتاج إلى تسميات توضيحية متزامنة ونص مكتوب حتى يتمكن المستخدمون ذوو الإعاقات البصرية و/أو السمعية من الوصول إليه. عمومًا، لا يكون مطوِّر الويب مسؤولًا عن إنشاء التسميات التوضيحية أو النص المكتوب، لكنه يحتاج إلى معرفة تضمينها.

# --instructions--

حان الوقت لأخذ استراحة من Camper Cat والتعرف على الزميل Zersiax (@zersiax)، بطل إمكانية الوصول ومستخدم قارئ الشاشة. لسماع مقطع من قارئ الشاشة الخاص به أثناء العمل، أضف عنصر `audio` بعد عنصر `p`. أدرج الخاصية `controls`. ثم ضع عنصر `source` داخل وسم `audio` مع تعيين الخاصية `src` إلى `https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3` والخاصية `type` إلى `"audio/mpeg"`.

**ملاحظة:** قد يبدو المقطع الصوتي سريعًا وصعب الفهم، لكن هذه سرعة طبيعية لمستخدمي قارئ الشاشة.

# --hints--

يجب أن يحتوي شفرتك على وسم `audio` واحد.

```js
assert.lengthOf(document.querySelectorAll('audio'),1);
```

يجب أن يحتوي عنصر `audio` على وسم إغلاق.

```js
assert.match(code,/<audio.*>[\s\S]*<\/audio>/g);
assert.lengthOf(code.match(/<\/audio>/g),1);
```

يجب أن يحتوي وسم `audio` على الخاصية `controls`.

```js
assert.exists(document.querySelector('audio')?.getAttribute('controls'));
```

يجب أن يحتوي شفرتك على وسم `source` واحد.

```js
assert.lengthOf(document.querySelectorAll('source'), 1);
```

يجب أن يكون وسم `source` داخل وسوم `audio`.

```js
const audio = document.querySelector('audio');
const children = audio.querySelectorAll(`:scope ${'source'}`);
assert.lengthOf(children,1);
```

يجب أن تتطابق قيمة الخاصية `src` في وسم `source` مع الرابط في التعليمات تمامًا.

```js
assert.equal(
  document.querySelector('source')?.getAttribute('src'),
  'https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3'
);
```

يجب أن تتضمن شفرتك خاصية `type` في وسم `source` بقيمة audio/mpeg.

```js
assert.equal(document.querySelector('source')?.getAttribute('type'), 'audio/mpeg');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>



  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>
    <audio controls>
      <source src="https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3" type="audio/mpeg"/>
    </audio>
  </main>
</body>
```
