---
id: 587d774c367417b2b2512a9c
title: أضف بديلاً نصيًا للصور من أجل إمكانية الوصول لذوي الإعاقة البصرية
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 16628
dashedName: add-a-text-alternative-to-images-for-visually-impaired-accessibility
---

# --description--

ربما رأيت خاصية `alt` على وسم `img` في تحديات أخرى. يصف نص `alt` محتوى الصورة ويوفر بديلاً نصيًا لها. تساعد خاصية `alt` في الحالات التي تفشل فيها الصورة في التحميل أو لا يمكن للمستخدم رؤيتها. كما تستخدم محركات البحث هذه الخاصية لفهم محتوى الصورة لتضمينها في نتائج البحث. إليك مثالًا:

```html
<img src="importantLogo.jpeg" alt="Company logo">
```

يعتمد الأشخاص ذوو الإعاقة البصرية على قارئات الشاشة لتحويل محتوى الويب إلى واجهة صوتية. لن يحصلوا على المعلومات إذا عُرضت فقط بصريًا. بالنسبة للصور، يمكن لقارئات الشاشة الوصول إلى خاصية `alt` وقراءة محتواها لتقديم المعلومات الأساسية.

يوفر نص `alt` الجيد للقارئ وصفًا موجزًا للصورة. يجب أن تتضمن دائمًا خاصية `alt` على صورتك. وفقًا لمواصفات HTML5، أصبح هذا الآن إلزاميًا.

# --instructions--

يُعد Camper Cat نينجا في البرمجة ونينجا حقيقيًا أيضًا، وهو يبني موقعًا لمشاركة معرفته. صورة الملف الشخصي التي يريد استخدامها تُظهر مهاراته ويجب أن يقدرها جميع زوار الموقع. أضف خاصية `alt` في وسم `img` تشرح أن Camper Cat يمارس الكاراتيه. (مصدر الصورة `src` لا يربط بملف فعلي، لذا يجب أن ترى نص `alt` في العرض.)

# --hints--

يجب أن يحتوي وسم `img` على خاصية `alt` وألا تكون فارغة.

```js
assert.isNotEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<img src="doingKarateWow.jpeg">
```

# --solutions--

```html
<img src="doingKarateWow.jpeg" alt="Someone doing karate">
```
