---
id: 587d774c367417b2b2512a9d
title: اعرف متى يجب ترك نص alt فارغًا
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9P4t2'
forumTopicId: 301019
dashedName: know-when-alt-text-should-be-left-blank
---

# --description--

في التحدٍّ السابق، تعلمت أن تضمين خاصية `alt` عند استخدام وسم `img` أمر إلزامي. مع ذلك، أحيانًا تُجمَع الصور مع تسمية توضيحية تشرحها بالفعل، أو تُستخدم للزينة فقط. في هذه الحالات، قد يبدو نص `alt` زائدًا أو غير ضروري.

عندما تُشرح الصورة بالفعل بمحتوى نصي أو لا تضيف معنى إلى الصفحة، لا يزال وسم `img` يحتاج إلى خاصية `alt`، لكن يمكن تعيينها إلى سلسلة نصية فارغة. إليك مثالًا:

```html
<img src="visualDecoration.jpeg" alt="">
```

عادةً ما تُصنّف الصور الخلفية أيضًا تحت تسمية "زخرفية". لكنها تُطبّق عادةً بقواعد CSS، ولذلك ليست جزءًا من الترميز الذي يعالجه قارئ الشاشة.

**ملاحظة:** بالنسبة للصور التي تحتوي على تسمية توضيحية، قد ترغب في تضمين نص `alt` لأنه يساعد محركات البحث على فهرسة محتوى الصورة.

# --instructions--

قام Camper Cat ببرمجة صفحة هيكلية لجزء المدونة في موقعه. يخطط لإضافة فاصل بصري بين مقالتين باستخدام صورة زخرفية لسيف ساموراي. أضف خاصية `alt` إلى وسم `img` واجعلها سلسلة نصية فارغة. (لاحظ أن مصدر الصورة `src` لا يربط بملف فعلي - لا تقلق من عدم ظهور السيوف في العرض.)

# --hints--

يجب أن يحتوي وسم `img` على خاصية `alt`.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

يجب تعيين خاصية `alt` إلى سلسلة نصية فارغة.

```js
assert.isEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg" alt="">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```
