---
id: 587d7790367417b2b2512ab1
title: استخدام tabindex لتحديد ترتيب تركيز لوحة المفاتيح لعدة عناصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzRRcb'
forumTopicId: 301028
dashedName: use-tabindex-to-specify-the-order-of-keyboard-focus-for-several-elements
---

# --description--

تحدد الخاصية `tabindex` أيضًا الترتيب الدقيق لعناصر التبويب. يتحقق ذلك عندما تُعيّن قيمة الخاصية إلى عدد موجب يساوي 1 أو أكثر.

تعيين `tabindex="1"` يجلب تركيز لوحة المفاتيح إلى ذلك العنصر أولًا. ثم يتبع تسلسل قيم `tabindex` المحددة (2، 3، إلخ)، قبل الانتقال إلى العناصر الافتراضية وذات `tabindex="0"`.

من المهم ملاحظة أنه عند تعيين ترتيب التبويب بهذه الطريقة، فإنه يتجاوز الترتيب الافتراضي (الذي يستخدم شفرة HTML المصدر). قد يربك هذا المستخدمين الذين يتوقعون بدء التنقل من أعلى الصفحة. قد تكون هذه التقنية ضرورية في بعض الحالات، لكن من حيث إمكانية الوصول، احرص على الحذر قبل تطبيقها.

إليك مثالًا:

```html
<div tabindex="1">I get keyboard focus, and I get it first!</div>
```

```html
<div tabindex="2">I get keyboard focus, and I get it second!</div>
```

# --instructions--

لدى Camper Cat حقل بحث في صفحة الاقتباسات الملهمة التي يخطط لوضعها في الزاوية العلوية اليمنى باستخدام CSS. يريد أن يكون عنصر التحكم في الموذج `input` الخاص بالبحث وعنصر التحكم في الموذج `input` الخاص بالإرسال أول عنصرين في ترتيب التبويب. أضف خاصية `tabindex` بقيمة `1` إلى عنصر التحكم في الموذج `search` `input`، وخاصية `tabindex` بقيمة `2` إلى عنصر التحكم في الموذج `submit` `input`.

شيء آخر يجب ملاحظته هو أن بعض المتصفحات قد تضعك في منتصف ترتيب التبويب عند نقر عنصر ما. تم إضافة عنصر إلى الصفحة يضمن أن تبدأ دائمًا من بداية ترتيب التبويب.

# --hints--

يجب أن تضيف شفرتك خاصية `tabindex` إلى وسم `input` الخاص بـ `search`.

```js
assert.isNotNull(document.querySelector('#search')?.getAttribute('tabindex'));
```

يجب أن تضيف شفرتك خاصية `tabindex` إلى وسم `input` الخاص بـ `submit`.

```js
assert.isNotNull(document.querySelector('#submit')?.getAttribute('tabindex'));
```

يجب أن يضبط شفرتك الخاصية `tabindex` على وسم `input` الخاص بـ `search` إلى القيمة 1.

```js
assert.equal(document.querySelector('#search')?.getAttribute('tabindex'), '1');
```

يجب أن يضبط شفرتك الخاصية `tabindex` على وسم `input` الخاص بـ `submit` إلى القيمة 2.

```js
assert.equal(document.querySelector('#submit')?.getAttribute('tabindex'), '2');
```

# --seed--

## --seed-contents--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input type="search" name="search" id="search">
    <input type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```

# --solutions--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input tabindex="1" type="search" name="search" id="search">
    <input tabindex="2" type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```
