---
id: 587d778a367417b2b2512aa5
title: تحسين إمكانية الوصول للمخططات باستخدام عنصر figure
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMqtE'
forumTopicId: 301015
dashedName: improve-chart-accessibility-with-the-figure-element
---

# --description--

قدم HTML5 عنصر `figure` والعنصر المرتبط به `figcaption`. عند استخدامهما معًا، يلتفان حول تمثيل بصري (مثل صورة، رسم بياني، أو مخطط) مع التسمية التوضيحية الخاصة به. يجمع التفاف هذه العناصر معًا بين المحتوى ذي الصلة دلاليًا ويوفر بديلًا نصيًا يشرح `figure`، مما يعزز إمكانية الوصول بشكل مزدوج.

بالنسبة لتمثيل البيانات مثل المخططات، يمكن استخدام التسمية التوضيحية لتلخيص الاتجاهات أو الاستنتاجات بإيجاز للمستخدمين ذوي الإعاقة البصرية. يغطي تحدٍّ آخر كيفية نقل نسخة الجدول من بيانات المخطط خارج الشاشة (باستخدام CSS) لمستخدمي قارئ الشاشة.

إليك مثالًا - لاحظ أن `figcaption` يكون داخل وسم `figure` ويمكن دمجه مع عناصر أخرى:

```html
<figure>
  <img src="roundhouseDestruction.jpeg" alt="Photo of Camper Cat executing a roundhouse kick">
  <br>
  <figcaption>
    Master Camper Cat demonstrates proper form of a roundhouse kick.
  </figcaption>
</figure>
```

# --instructions--

يعمل Camper Cat بجد لإنشاء مخطط عمودي مكدس يظهر مقدار الوقت الأسبوعي المخصص للتدريب على التخفي والقتال والأسلحة. ساعده على تنظيم صفحته بشكل أفضل بتغيير وسم `div` الذي استخدمه إلى وسم `figure`، ووسم `p` الذي يحيط بالتسمية التوضيحية إلى وسم `figcaption`.

# --hints--

يجب أن يحتوي شفرتك على وسم `figure` واحد.

```js
assert.lengthOf(document.querySelectorAll('figure') , 1);
```

يجب أن يحتوي شفرتك على وسم `figcaption` واحد.

```js
assert.lengthOf(document.querySelectorAll('figcaption') , 1);
```

يجب ألا تحتوي شفرتك على أي وسوم `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

يجب ألا تحتوي شفرتك على أي وسوم `p`.

```js
assert.lengthOf(document.querySelectorAll('p') , 0);
```

يجب أن يكون `figcaption` عنصرًا فرعيًا داخل وسم `figure`.

```js
const figure = document.querySelector('figure');
const children = figure?.querySelectorAll(`:scope ${'figcaption'}`);
assert.lengthOf(children, 1);
```

يجب أن يحتوي عنصر `figure` على وسم إغلاق.

```js
assert.isTrue(
    code.match(/<\/figure>/g)?.length === code.match(/<figure>/g)?.length
);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>

      <!-- Only change code below this line -->
      <div>
        <!-- Stacked bar chart will go here -->
        <br>
        <p>Breakdown per week of time to spend training in stealth, combat, and weapons.</p>
      </div>
      <!-- Only change code above this line -->

    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>
      <figure>
        <!-- Stacked bar chart will go here -->
        <br>
        <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
      </figure>
    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
