---
id: 587d7da9367417b2b2512b67
title: push 대신 concat을 사용해 배열 끝에 요소 추가하기
challengeType: 1
forumTopicId: 301226
dashedName: add-elements-to-the-end-of-an-array-using-concat-instead-of-push
---

# --description--

함수형 프로그래밍은 변경하지 않는 함수를 만들고 사용하는 것에 관한 것입니다.

이전 도전 과제에서는 원본 배열을 변경하지 않고 배열을 새 배열로 병합하는 방법으로 `concat` 메서드를 소개했습니다. `concat`과 `push` 메서드를 비교해 보세요. `push`은 호출된 같은 배열 끝에 항목을 추가하여 그 배열을 변경합니다. 예를 들어 다음과 같습니다:

```js
const arr = [1, 2, 3];
arr.push(4, 5, 6);
```

`arr`는 `[1, 2, 3, 4, 5, 6]`로 값이 변경되는데, 이는 함수형 프로그래밍 방식이 아닙니다.

`concat`은 배열 끝에 새 항목을 변경 없이 병합하는 방법을 제공합니다.

# --instructions--

`nonMutatingPush` 함수를 변경해 `concat`을 사용하여 `newItem`를 `original` 끝에 병합하되 `original`과 `newItem` 배열을 변경하지 않도록 하세요. 함수는 배열을 반환해야 합니다.

# --hints--

코드는 `concat` 메서드를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

코드는 `push` 메서드를 사용하지 않아야 합니다.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?push/g));
```

`first` 배열은 변경되지 않아야 합니다.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

`second` 배열은 변경되지 않아야 합니다.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingPush([1, 2, 3], [4, 5])`는 `[1, 2, 3, 4, 5]`을 반환해야 합니다.

```js
assert(
  JSON.stringify(nonMutatingPush([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingPush(original, newItem) {
  // Only change code below this line
  return original.push(newItem);

  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingPush(first, second);
```

# --solutions--

```js
function nonMutatingPush(original, newItem) {
  return original.concat(newItem);
}
const first = [1, 2, 3];
const second = [4, 5];
```
