---
id: 587d7b8e367417b2b2512b5e
title: 함수형 프로그래밍을 사용해 변이와 부작용 피하기
challengeType: 1
forumTopicId: 301228
dashedName: avoid-mutations-and-side-effects-using-functional-programming
---

# --description--

아직 눈치채지 못했다면, 이전 도전 과제의 문제는 `splice` 함수 내 `tabClose()` 호출에 있었습니다. 불행히도 `splice`는 호출된 원래 배열을 변경하기 때문에, 두 번째 호출은 수정된 배열을 사용했고 예상치 못한 결과를 낳았습니다.

이것은 훨씬 더 큰 패턴의 작은 예입니다 - 변수, 배열, 또는 객체에 함수를 호출하면, 그 함수가 변수나 객체 내 무언가를 변경합니다.

함수형 프로그래밍의 핵심 원칙 중 하나는 변화를 주지 않는 것입니다. 변화는 버그를 초래합니다. 함수가 인수나 전역 변수를 포함해 아무것도 변경하지 않는다는 것을 알면 버그를 예방하기가 더 쉽습니다.

이전 예제는 복잡한 연산이 없었지만 `splice` 메서드가 원래 배열을 변경했고, 그 결과 버그가 발생했습니다.

함수형 프로그래밍에서 무언가를 변경하거나 바꾸는 것을 <dfn>변이</dfn>라고 하며, 그 결과를 <dfn>부작용</dfn>이라고 합니다. 함수는 이상적으로 <dfn>순수한 함수</dfn>여야 하며, 부작용을 일으키지 않아야 합니다.

이 원칙을 익혀 코드 내 어떤 변수나 객체도 변경하지 않도록 합시다.

# --instructions--

함수 `incrementer`의 코드를 작성해 전역 변수 `fixedValue`의 값을 1만큼 증가시킨 값을 반환하세요.

# --hints--

함수 `incrementer`는 `fixedValue` (`4`)의 값을 변경하면 안 됩니다.

```js
incrementer();
assert(fixedValue === 4);
```

함수 `incrementer`는 `fixedValue` 값보다 1 큰 값을 반환해야 합니다.

```js
const __newValue = incrementer();
assert(__newValue === 5);
```

함수 `incrementer`는 전역 `fixedValue` 변수 값을 기반으로 값을 반환해야 합니다.

```js
(function () {
  fixedValue = 10;
  const newValue = incrementer();
  assert(fixedValue === 10 && newValue === 11);
  fixedValue = 4;
})();
```

# --seed--

## --seed-contents--

```js
// The global variable
let fixedValue = 4;

function incrementer() {
  // Only change code below this line


  // Only change code above this line
}
```

# --solutions--

```js
let fixedValue = 4

function incrementer() {
  return fixedValue + 1
}
```
