---
id: 587d7da9367417b2b2512b69
title: sort 메서드를 사용해 배열을 알파벳순으로 정렬하기
challengeType: 1
forumTopicId: 18303
dashedName: sort-an-array-alphabetically-using-the-sort-method
---

# --description--

`sort` 메서드는 콜백 함수에 따라 배열의 요소를 정렬합니다.

예를 들어:

```js
function ascendingOrder(arr) {
  return arr.sort(function(a, b) {
    return a - b;
  });
}

ascendingOrder([1, 5, 2, 3, 4]);
```

이 경우 `[1, 2, 3, 4, 5]` 값을 반환합니다.

```js
function reverseAlpha(arr) {
  return arr.sort(function(a, b) {
    return a === b ? 0 : a < b ? 1 : -1;
  });
}

reverseAlpha(['l', 'h', 'z', 'b', 's']);
```

이 경우 `['z', 's', 'l', 'h', 'b']` 값을 반환합니다.

JavaScript의 기본 정렬 방식은 문자열 유니코드 포인트 값에 따른 것으로, 예상치 못한 결과가 나올 수 있습니다. 따라서 배열 항목을 어떻게 정렬할지 지정하는 콜백 함수, 보통 `compareFunction`라고 하는 함수를 제공하는 것이 권장됩니다. 이 콜백 함수가 제공되면 배열 요소는 `compareFunction`의 반환값에 따라 정렬됩니다: 만약 `compareFunction(a,b)`가 두 요소 `a`와 `b`에 대해 0보다 작은 값을 반환하면 `a`가 `b`보다 앞에 옵니다. 만약 `compareFunction(a,b)`가 두 요소 `a`와 `b`에 대해 0보다 큰 값을 반환하면 `b`가 `a`보다 앞에 옵니다. 만약 `compareFunction(a,b)`가 두 요소 `a`와 `b`에 대해 0과 같은 값을 반환하면 `a`와 `b`는 변경되지 않습니다.

# --instructions--

`sort` 메서드를 `alphabeticalOrder` 함수 안에서 사용해 `arr`의 요소를 알파벳순으로 정렬하세요. 함수는 정렬된 배열을 반환해야 합니다.

# --hints--

코드에서 `sort` 메서드를 사용하세요.

```js
assert(__helpers.removeJSComments(code).match(/\.sort/g));
```

`alphabeticalOrder(["a", "d", "c", "a", "z", "g"])`는 `["a", "a", "c", "d", "g", "z"]`를 반환해야 합니다.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'd', 'c', 'a', 'z', 'g'])) ===
    JSON.stringify(['a', 'a', 'c', 'd', 'g', 'z'])
);
```

`alphabeticalOrder(["x", "h", "a", "m", "n", "m"])`는 `["a", "h", "m", "m", "n", "x"]`를 반환해야 합니다.

```js
assert(
  JSON.stringify(alphabeticalOrder(['x', 'h', 'a', 'm', 'n', 'm'])) ===
    JSON.stringify(['a', 'h', 'm', 'm', 'n', 'x'])
);
```

`alphabeticalOrder(["a", "a", "a", "a", "x", "t"])`는 `["a", "a", "a", "a", "t", "x"]`를 반환해야 합니다.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'a', 'a', 'a', 'x', 't'])) ===
    JSON.stringify(['a', 'a', 'a', 'a', 't', 'x'])
);
```

# --seed--

## --seed-contents--

```js
function alphabeticalOrder(arr) {
  // Only change code below this line

  return arr
  // Only change code above this line
}

alphabeticalOrder(["a", "d", "c", "a", "z", "g"]);
```

# --solutions--

```js
function alphabeticalOrder(arr) {
  return arr.sort();
}
```
