---
id: 587d7b8e367417b2b2512b5c
title: 함수형 프로그래밍 용어 이해하기
challengeType: 1
forumTopicId: 301240
dashedName: understand-functional-programming-terminology
---

# --description--

FCC 팀의 기분이 변해서 이제 두 가지 종류의 차를 원합니다: 녹차와 홍차. 일반적인 사실: 클라이언트의 기분 변화는 꽤 흔합니다.

이 정보를 바탕으로, 지난 도전 과제에서 사용한 `getTea` 함수를 다시 살펴서 다양한 차 요청을 처리해야 합니다. `getTea`을 매개변수로 함수를 받도록 수정하면 준비하는 차 종류를 바꿀 수 있습니다. 이렇게 하면 `getTea`가 더 유연해지고, 클라이언트 요청이 바뀔 때 프로그래머가 더 많은 제어권을 가질 수 있습니다.

하지만 먼저, 함수형 용어를 살펴보겠습니다:

<dfn>콜백</dfn>은 다른 함수에 전달되어 그 함수의 호출 방식을 결정하는 함수입니다. 예를 들어 `filter`에서 콜백 함수는 JavaScript에 배열을 어떻게 필터링할지 기준을 알려줍니다.

변수에 할당하거나, 다른 함수에 전달하거나, 다른 함수에서 반환할 수 있는 함수는 <dfn>일급</dfn> 함수라고 합니다. JavaScript에서는 모든 함수가 일급 함수입니다.

함수를 인수로 받거나 반환값으로 함수를 반환하는 함수는 <dfn>고차</dfn> 함수라고 합니다.

함수가 다른 함수에 전달되거나 반환될 때, 전달되거나 반환된 함수를 <dfn>람다</dfn>라고 부릅니다.

# --instructions--

녹차 27잔과 홍차 13잔을 각각 `tea4GreenTeamFCC`와 `tea4BlackTeamFCC` 변수에 저장하세요. `getTea` 함수가 수정되어 첫 번째 인수로 함수를 받는다는 점에 유의하세요.

참고: 차 잔 수 데이터는 마지막 인수로 전달됩니다. 이 부분은 이후 수업에서 더 자세히 다룹니다.

# --hints--

`tea4GreenTeamFCC` 변수에는 팀을 위한 녹차 27잔이 들어 있어야 합니다.

```js
assert(tea4GreenTeamFCC.length === 27);
```

`tea4GreenTeamFCC` 변수에는 녹차 잔이 들어 있어야 합니다.

```js
assert(tea4GreenTeamFCC[0] === 'greenTea');
```

`tea4BlackTeamFCC` 변수에는 홍차 13잔이 들어 있어야 합니다.

```js
assert(tea4BlackTeamFCC.length === 13);
```

`tea4BlackTeamFCC` 변수에는 홍차 잔이 들어 있어야 합니다.

```js
assert(tea4BlackTeamFCC[0] === 'blackTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareGreenTea = () => 'greenTea';

// Function that returns a string representing a cup of black tea
const prepareBlackTea = () => 'blackTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (prepareTea, numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4GreenTeamFCC = null;
const tea4BlackTeamFCC = null;
// Only change code above this line

console.log(
  tea4GreenTeamFCC,
  tea4BlackTeamFCC
);
```

# --solutions--

```js
const prepareGreenTea = () => 'greenTea';
const prepareBlackTea = () => 'blackTea';

const getTea = (prepareTea, numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

const tea4BlackTeamFCC = getTea(prepareBlackTea, 13);
const tea4GreenTeamFCC = getTea(prepareGreenTea, 27);
```
