---
id: 587d7dab367417b2b2512b70
title: 커링과 부분 적용 소개
challengeType: 1
forumTopicId: 301232
dashedName: introduction-to-currying-and-partial-application
---

# --description--

함수의 <dfn>arity</dfn>는 필요한 인수의 개수를 의미합니다. <dfn>커링</dfn>은 N개의 인수를 받는 함수를 arity가 1인 N개의 함수로 변환하는 것을 뜻합니다.

즉, 함수를 하나의 인수만 받도록 재구성하고, 그 다음 인수를 받는 또 다른 함수를 반환하는 식으로 이어집니다.

예를 들어 다음과 같습니다:

```js
function unCurried(x, y) {
  return x + y;
}

function curried(x) {
  return function(y) {
    return x + y;
  }
}

const curried = x => y => x + y

curried(1)(2)
```

`curried(1)(2)`는 `3`를 반환합니다.

이 방법은 한 번에 모든 인수를 함수에 전달할 수 없을 때 유용합니다. 각 함수 호출을 변수에 저장해두면, 다음 인수를 받을 때 반환된 함수 참조를 사용할 수 있습니다. 위 예제의 커링된 함수를 사용한 예는 다음과 같습니다:

```js
const funcForY = curried(1);
console.log(funcForY(2)); // 3
```

마찬가지로, <dfn>부분 적용</dfn>은 한 번에 몇 개의 인수를 함수에 적용하고, 더 많은 인수를 적용할 수 있는 또 다른 함수를 반환하는 것으로 설명할 수 있습니다. 예시는 다음과 같습니다:

```js
function impartial(x, y, z) {
  return x + y + z;
}

const partialFn = impartial.bind(this, 1, 2);
partialFn(10); // 13
```

# --instructions--

`add` 함수 본문을 채워 `x`, `y`, `z` 매개변수를 커링을 사용해 더하세요.

# --hints--

`add(10)(20)(30)`는 `60`를 반환해야 합니다.

```js
assert(add(10)(20)(30) === 60);
```

`add(1)(2)(3)`는 `6`를 반환해야 합니다.

```js
assert(add(1)(2)(3) === 6);
```

`add(11)(22)(33)`는 `66`를 반환해야 합니다.

```js
assert(add(11)(22)(33) === 66);
```

코드에는 마지막에 `x + y + z`를 반환하는 문장이 포함되어야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/[xyz]\s*?\+\s*?[xyz]\s*?\+\s*?[xyz]/g));
```

# --seed--

## --seed-contents--

```js
function add(x) {
  // Only change code below this line


  // Only change code above this line
}

add(10)(20)(30);
```

# --solutions--

```js
const add = x => y => z => x + y + z
```
