---
id: 587d7b88367417b2b2512b45
title: '고차 함수 map, filter, reduce를 사용해 복잡한 문제 해결하기'
challengeType: 1
forumTopicId: 301311
dashedName: use-higher-order-functions-map-filter-or-reduce-to-solve-a-complex-problem
---

# --description--

`map()`, `filter()`, `reduce()` 같은 고차 함수를 사용한 몇 가지 도전 과제를 완료했으니, 이제 이 함수들을 적용해 더 복잡한 도전 과제를 해결해 보세요.

# --instructions--

`squareList` 함수를 `map()`, `filter()`, `reduce()` 중 어떤 조합이든 사용해 완성하세요. 이 함수는 실수 배열이 주어졌을 때 *양의 정수*만 제곱한 값을 포함하는 새 배열을 반환해야 합니다(소수는 정수가 아닙니다). 실수 배열의 예는 `[-3, 4.8, 5, 3, -3.2]`입니다.

**참고:** 함수 내에서 어떤 종류의 `for` 또는 `while` 반복문이나 `forEach()` 함수는 사용하지 마세요.

# --hints--

`squareList`은 `function`여야 합니다.

```js
assert.typeOf(squareList, 'function'),
  '<code>squareList</code> should be a <code>function</code>';
```

`for`, `while`, `forEach`는 사용하지 마세요.

```js
assert(!__helpers.removeJSComments(code).match(/for|while|forEach/g));
```

`map`, `filter`, `reduce` 중 하나를 사용하세요.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/\.(map|filter|reduce)\(/g)
);
```

함수는 `array`를 반환해야 합니다.

```js
assert(Array.isArray(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])));
```

`squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])`는 `[16, 1764, 36]`를 반환해야 합니다.

```js
assert.deepStrictEqual(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2]), [
  16,
  1764,
  36
]);
```

`squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3])`는 `[9, 100, 49]`를 반환해야 합니다.

```js
assert.deepStrictEqual(squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3]), [
  9,
  100,
  49
]);
```

# --seed--

## --seed-contents--

```js
const squareList = arr => {
  // Only change code below this line
  return arr;
  // Only change code above this line
};

const squaredIntegers = squareList([-3, 4.8, 5, 3, -3.2]);
console.log(squaredIntegers);
```

# --solutions--

```js
const squareList = arr => {
  const positiveIntegers = arr.filter(num => {
    return num >= 0 && Number.isInteger(num);
  });
  const squaredIntegers = positiveIntegers.map(num => {
    return num ** 2;
  });
  return squaredIntegers;
};
```
