---
id: 587d7daa367417b2b2512b6c
title: join 메서드를 사용해 배열을 문자열로 결합하기
challengeType: 1
forumTopicId: 18221
dashedName: combine-an-array-into-a-string-using-the-join-method
---

# --description--

`join` 메서드는 배열의 요소들을 연결해 문자열을 만드는 데 사용됩니다. 이 메서드는 문자열에서 배열 요소를 구분하는 데 쓰이는 구분자를 인수로 받습니다.

예를 들어 보겠습니다:

```js
const arr = ["Hello", "World"];
const str = arr.join(" ");
```

`str`은 문자열 `Hello World` 값을 갖게 됩니다.
# --instructions--

`join` 메서드(그 외 다른 메서드도 가능)를 `sentensify` 함수 안에서 사용해 문자열 `str`에 있는 단어들로 문장을 만드세요. 함수는 문자열을 반환해야 합니다. 예를 들어 `I-like-Star-Wars`은 `I like Star Wars`로 변환됩니다. 이번 도전 과제에서는 `replace` 메서드를 사용하지 마세요.

# --hints--

코드에서 `join` 메서드를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\.join/g));
```

코드에서 `replace` 메서드를 사용하지 않아야 합니다.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?replace/g));
```

`sentensify("May-the-force-be-with-you")`은 문자열을 반환해야 합니다.

```js
assert(typeof sentensify('May-the-force-be-with-you') === 'string');
```

`sentensify("May-the-force-be-with-you")`은 문자열 `May the force be with you`을 반환해야 합니다.

```js
assert(sentensify('May-the-force-be-with-you') === 'May the force be with you');
```

`sentensify("The.force.is.strong.with.this.one")`은 문자열 `The force is strong with this one`을 반환해야 합니다.

```js
assert(
  sentensify('The.force.is.strong.with.this.one') ===
    'The force is strong with this one'
);
```

`sentensify("There,has,been,an,awakening")`은 문자열 `There has been an awakening`을 반환해야 합니다.

```js
assert(
  sentensify('There,has,been,an,awakening') === 'There has been an awakening'
);
```

# --seed--

## --seed-contents--

```js
function sentensify(str) {
  // Only change code below this line


  // Only change code above this line
}

sentensify("May-the-force-be-with-you");
```

# --solutions--

```js
function sentensify(str) {
  return str.split(/\W/).join(' ');
}
```
