---
id: 587d7b90367417b2b2512b65
title: slice 메서드를 사용해 배열의 일부를 반환하기
challengeType: 1
forumTopicId: 301239
dashedName: return-part-of-an-array-using-the-slice-method
---

# --description--

`slice` 메서드는 배열의 특정 요소들을 복사해 반환합니다. 두 개의 인수를 받을 수 있는데, 첫 번째는 슬라이스를 시작할 인덱스이고, 두 번째는 슬라이스를 끝낼 인덱스(포함하지 않음)입니다. 인수가 제공되지 않으면 기본값은 배열의 처음부터 끝까지이며, 이는 배열 전체를 복사하는 쉬운 방법입니다. `slice` 메서드는 원본 배열을 변경하지 않고 새 배열을 반환합니다.

예시는 다음과 같습니다:

```js
const arr = ["Cat", "Dog", "Tiger", "Zebra"];
const newArray = arr.slice(1, 3);
```

`newArray`는 `["Dog", "Tiger"]` 값을 갖게 됩니다.

# --instructions--

`slice` 함수에서 `sliceArray` 메서드를 사용해 주어진 `anim`과 `beginSlice` 인덱스에 따라 `endSlice` 배열의 일부를 반환하세요. 함수는 배열을 반환해야 합니다.

# --hints--

코드는 `slice` 메서드를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\.slice/g));
```

`inputAnim` 변수는 변경되지 않아야 합니다.

```js
assert(
  JSON.stringify(inputAnim) ===
    JSON.stringify(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 3)`은 `["Dog", "Tiger"]`를 반환해야 합니다.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 3)) ===
    JSON.stringify(['Dog', 'Tiger'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 0, 1)`은 `["Cat"]`를 반환해야 합니다.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 0, 1)) ===
    JSON.stringify(['Cat'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 4)`는 `["Dog", "Tiger", "Zebra"]`를 반환해야 합니다.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 4)) ===
    JSON.stringify(['Dog', 'Tiger', 'Zebra'])
);
```

# --seed--

## --seed-contents--

```js
function sliceArray(anim, beginSlice, endSlice) {
  // Only change code below this line


  // Only change code above this line
}

const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
sliceArray(inputAnim, 1, 3);
```

# --solutions--

```js
function sliceArray(anim, beginSlice, endSlice) {
  return anim.slice(beginSlice, endSlice);
}
const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
```
