---
id: 587d7dab367417b2b2512b6d
title: 문자열을 URL 슬러그로 변환하는 함수형 프로그래밍 적용하기
challengeType: 1
forumTopicId: 301227
dashedName: apply-functional-programming-to-convert-strings-to-url-slugs
---

# --description--

최근 도전 과제들에서는 함수형 프로그래밍 원칙을 따르는 유용한 배열과 문자열 메서드를 다뤘습니다. 또한 문제를 더 단순한 형태로 줄이는 데 사용되는 강력한 메서드인 `reduce`에 대해서도 배웠습니다. 평균 계산부터 정렬까지, 모든 배열 작업은 이 메서드를 적용해 수행할 수 있습니다. `map`과 `filter`가 `reduce`의 특수한 경우임을 기억하세요.

이제 배운 내용을 결합해 실용적인 문제를 해결해 봅시다.

많은 콘텐츠 관리 사이트(CMS)는 게시물 제목을 URL의 일부로 추가해 간단한 북마크 용도로 사용합니다. 예를 들어, Medium에 `Stop Using Reduce`라는 제목의 게시물을 작성하면 URL에 제목 문자열의 어떤 형태가 포함될 가능성이 큽니다(`.../stop-using-reduce`). freeCodeCamp 사이트에서도 이미 이런 것을 보셨을 겁니다.

# --instructions--

`urlSlug` 함수를 완성해 문자열 `title`을 변환하고 URL용 하이픈으로 연결된 버전을 반환하세요. 이 섹션에서 다룬 메서드를 자유롭게 사용하되 `replace`는 사용하지 마세요. 요구사항은 다음과 같습니다:

입력은 공백과 제목 대문자가 포함된 문자열입니다

출력은 단어 사이 공백이 하이픈(`-`)으로 교체된 문자열이어야 합니다

출력은 모두 소문자여야 합니다

출력에 공백이 없어야 합니다

# --hints--

이번 도전 과제에서는 `replace` 메서드를 사용하지 마세요.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?replace/g));
```

`urlSlug("Winter Is Coming")`은 문자열 `winter-is-coming`를 반환해야 합니다.

```js
assert(urlSlug('Winter Is Coming') === 'winter-is-coming');
```

`urlSlug(" Winter Is  Coming")`은 문자열 `winter-is-coming`를 반환해야 합니다.

```js
assert(urlSlug(' Winter Is  Coming') === 'winter-is-coming');
```

`urlSlug("A Mind Needs Books Like A Sword Needs A Whetstone")`은 문자열 `a-mind-needs-books-like-a-sword-needs-a-whetstone`를 반환해야 합니다.

```js
assert(
  urlSlug('A Mind Needs Books Like A Sword Needs A Whetstone') ===
    'a-mind-needs-books-like-a-sword-needs-a-whetstone'
);
```

`urlSlug("Hold The Door")`은 문자열 `hold-the-door`를 반환해야 합니다.

```js
assert(urlSlug('Hold The Door') === 'hold-the-door');
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
function urlSlug(title) {


}
// Only change code above this line
urlSlug("A Mind Needs Books Like A Sword Needs A Whetstone");
```

# --solutions--

```js
function urlSlug(title) {
  return title.trim().split(/\s+/).join("-").toLowerCase();
}
```
