---
id: 587d7dab367417b2b2512b6f
title: 배열의 일부 요소가 조건을 충족하는지 확인하기 위해 some 메서드 사용하기
challengeType: 1
forumTopicId: 301314
dashedName: use-the-some-method-to-check-that-any-elements-in-an-array-meet-a-criteria
---

# --description--

`some` 메서드는 배열과 함께 작동하여 *어떤* 요소가 특정 테스트를 통과하는지 확인합니다. 조건을 충족하는 값이 하나라도 있으면 불리언 값 `true`을 반환하고, 그렇지 않으면 `false`를 반환합니다.

예를 들어, 다음 코드는 `numbers` 배열의 어떤 요소가 10보다 작은지 확인합니다:

```js
const numbers = [10, 50, 8, 220, 110, 11];

numbers.some(function(currentValue) {
  return currentValue < 10;
});
```

`some` 메서드는 `true`를 반환합니다.

# --instructions--

`some` 함수 안에서 `checkPositive` 메서드를 사용하여 `arr`의 어떤 요소가 양수인지 확인하세요. 함수는 불리언 값을 반환해야 합니다.

# --hints--

코드는 `some` 메서드를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\.some/g));
```

`checkPositive([1, 2, 3, -4, 5])`은 `true`을 반환해야 합니다.

```js
assert(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])`은 `true`을 반환해야 합니다.

```js
assert(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([-1, -2, -3, -4, -5])`는 `false`를 반환해야 합니다.

```js
assert(!checkPositive([-1, -2, -3, -4, -5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.some(elem => elem > 0);
}
```
