---
id: 587d7b8e367417b2b2512b5d
title: 명령형 코드 사용의 위험성 이해하기
challengeType: 1
forumTopicId: 301241
dashedName: understand-the-hazards-of-using-imperative-code
---

# --description--

함수형 프로그래밍은 좋은 습관입니다. 코드를 관리하기 쉽게 유지하고, 교묘한 버그를 피할 수 있게 도와줍니다. 하지만 그 전에, 문제점이 발생할 수 있는 부분을 강조하기 위해 명령형 프로그래밍 방식을 살펴보겠습니다.

영어(및 다른 여러 언어)에서 명령형 시제는 명령을 내릴 때 사용합니다. 마찬가지로, 프로그래밍에서 명령형 스타일은 컴퓨터에게 작업을 수행할 일련의 명령문을 주는 방식입니다.

종종 이 명령문들은 전역 변수 같은 프로그램 상태를 변경합니다. 고전적인 예는 배열의 인덱스를 순회하는 정확한 지시를 주는 `for` 루프를 작성하는 것입니다.

반면에 함수형 프로그래밍은 선언형 프로그래밍의 한 형태입니다. 메서드나 함수를 호출하여 컴퓨터에게 원하는 작업을 알려줍니다.

JavaScript는 일반적인 작업을 처리하는 미리 정의된 메서드를 많이 제공하므로, 컴퓨터가 어떻게 수행해야 하는지 일일이 작성할 필요가 없습니다. 예를 들어, 위에서 언급한 `for` 루프 대신 배열을 순회하는 세부 사항을 처리하는 `map` 메서드를 호출할 수 있습니다. 이는 디버깅 섹션에서 다룬 "Off By One Errors" 같은 의미론적 오류를 피하는 데 도움이 됩니다.

상황을 생각해 보세요: 브라우저에서 웹을 탐색하며 열린 탭을 추적하고 싶습니다. 간단한 객체지향 코드를 사용해 이를 모델링해 보겠습니다.

Window 객체는 탭들로 구성되며, 보통 여러 개의 Window를 엽니다. 각 Window 객체에 열린 사이트의 제목은 배열에 저장됩니다. 브라우저에서 작업(새 탭 열기, 창 병합, 탭 닫기)을 한 후, 아직 열린 탭을 출력하고 싶습니다. 닫힌 탭은 배열에서 제거되고, 새 탭은(간단히 하기 위해) 배열 끝에 추가됩니다.

코드 편집기에는 `tabOpen()`, `tabClose()`, `join()` 함수로 이 기능을 구현한 코드가 나와 있습니다. `tabs` 배열은 열린 페이지 이름을 저장하는 Window 객체의 일부입니다.

# --instructions--

편집기의 코드를 살펴보세요. 프로그램에 부작용을 일으켜 잘못된 동작을 하는 메서드를 사용하고 있습니다. `finalTabs.tabs`에 저장된 최종 열린 탭 목록은 `['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']`이어야 하지만, 코드가 생성하는 목록은 약간 다릅니다.

`Window.prototype.tabClose`을 변경하여 올바른 탭을 제거하세요.

# --hints--

`finalTabs.tabs`은 `['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']`이어야 합니다.

```js
assert.deepEqual(finalTabs.tabs, [
  'FB',
  'Gitter',
  'Reddit',
  'Twitter',
  'Medium',
  'new tab',
  'Netflix',
  'YouTube',
  'Vine',
  'GMail',
  'Work mail',
  'Docs',
  'freeCodeCamp',
  'new tab'
]);
```

# --seed--

## --seed-contents--

```js
// tabs is an array of titles of each site open within the window
const Window = function(tabs) {
  this.tabs = tabs; // We keep a record of the array inside the object
};

// When you join two windows into one window
Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

// When you open a new tab at the end
Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab'); // Let's open a new tab for now
  return this;
};

// When you close a tab
Window.prototype.tabClose = function(index) {

  // Only change code below this line

  const tabsBeforeIndex = this.tabs.splice(0, index); // Get the tabs before the tab
  const tabsAfterIndex = this.tabs.splice(index + 1); // Get the tabs after the tab

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex); // Join them together

  // Only change code above this line

  return this;
 };

// Let's create three browser windows
const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']); // Your mailbox, drive, and other work sites
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']); // Social sites
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']); // Entertainment sites

// Now perform the tab opening, closing, and other operations
const finalTabs = socialWindow
  .tabOpen() // Open a new tab for cat memes
  .join(videoWindow.tabClose(2)) // Close third tab in video window, and join
  .join(workWindow.tabClose(1).tabOpen());
console.log(finalTabs.tabs);
```

# --solutions--

```js
const Window = function(tabs) {
  this.tabs = tabs;
};

Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab');
  return this;
};

Window.prototype.tabClose = function(index) {
  const tabsBeforeIndex = this.tabs.slice(0, index);
  const tabsAfterIndex = this.tabs.slice(index + 1);

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex);
  return this;
 };

const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']);
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']);
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']);

const finalTabs = socialWindow
  .tabOpen()
  .join(videoWindow.tabClose(2))
  .join(workWindow.tabClose(1).tabOpen());
```
