---
id: 587d7b8f367417b2b2512b62
title: 프로토타입에 map 구현하기
challengeType: 1
forumTopicId: 301230
dashedName: implement-map-on-a-prototype
---

# --description--

이전에 `Array.prototype.map()` 또는 간단히 `map()`을 적용해 본 것처럼, `map` 메서드는 호출된 배열과 같은 길이의 배열을 반환합니다. 또한 콜백 함수가 원본 배열을 변경하지 않는 한 원본 배열도 변경하지 않습니다.

즉, `map`은 순수한 함수이며 출력은 오직 입력에만 의존합니다. 게다가 다른 함수를 인수로 받습니다.

자신만의 `map` 메서드를 구현하면 `map` 메서드에 대해 많은 것을 배울 수 있습니다. `for` 반복문이나 `Array.prototype.forEach()`를 사용하는 것이 좋습니다.

# --instructions--

내장 `map` 메서드를 사용하지 않고, `Array.prototype.map()`과 정확히 같은 동작을 하는 자신만의 `Array.prototype.myMap()`을 작성하세요. `Array` 인스턴스는 `myMap` 메서드 내에서 `this`를 사용해 접근할 수 있습니다.

# --hints--

`[23, 65, 98, 5, 13].myMap(item => item * 2)`은 `[46, 130, 196, 10, 26]`와 같아야 합니다.

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item * 2;
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`["naomi", "quincy", "camperbot"].myMap(element => element.toUpperCase())`은 `["NAOMI", "QUINCY", "CAMPERBOT"]`을 반환해야 합니다.

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element.toUpperCase();
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`[1, 1, 2, 5, 2].myMap((element, index, array) => array[index + 1] || array[0])`은 `[1, 2, 5, 2, 1]`을 반환해야 합니다.

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array[index + 1] || array[0];
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

코드에서 `map` 메서드를 사용하지 않아야 합니다.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?map/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    newArray.push(callback(this[i], i, this));
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const doubled_s = s.myMap(item => item * 2);
```
