---
id: 587d7b8f367417b2b2512b60
title: 함수에서 전역 변수 리팩터링하기
challengeType: 1
forumTopicId: 301235
dashedName: refactor-global-variables-out-of-functions
---

# --description--

지금까지 함수형 프로그래밍의 두 가지 뚜렷한 원칙을 살펴봤습니다:

1) 변수나 객체를 변경하지 말고, 필요하다면 새로운 변수와 객체를 만들어 함수에서 반환하세요. 힌트: `const newArr = arrVar`처럼 `arrVar`이 배열일 경우 기존 변수를 복사하지 않고 참조만 생성합니다. 따라서 `newArr`에서 값을 변경하면 `arrVar`의 값도 변경됩니다.

2) 함수 매개변수를 선언하세요 - 함수 내부의 모든 계산은 함수에 전달된 인수에만 의존하며, 전역 객체나 변수에는 의존하지 않습니다.

숫자에 1을 더하는 것은 별로 흥미롭지 않지만, 배열이나 더 복잡한 객체를 다룰 때 이 원칙들을 적용할 수 있습니다.

# --instructions--

전역 배열 `bookList`가 두 함수 내부에서 변경되지 않도록 코드를 다시 작성하세요. `add` 함수는 전달받은 배열 끝에 주어진 `bookName`를 추가하고 새로운 배열을 반환해야 합니다. `remove` 함수는 전달받은 배열에서 주어진 `bookName`를 제거해야 합니다.

**참고:** 두 함수 모두 배열을 반환해야 하며, 새로운 매개변수는 `bookName` 매개변수 앞에 추가하세요.

# --hints--

`bookList`은 변경되지 않고 여전히 `["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"]`과 같아야 합니다.

```js
add(bookList, "Test");
remove(bookList, "The Hound of the Baskervilles");
assert(
  JSON.stringify(bookList) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'On The Electrodynamics of Moving Bodies',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae'
    ])
);
```

`add(bookList, "A Brief History of Time")`은 `["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae", "A Brief History of Time"]`을 반환해야 합니다.

```js
assert(
  JSON.stringify(add(bookList, "A Brief History of Time")) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'On The Electrodynamics of Moving Bodies',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae',
      'A Brief History of Time'
    ])
);
```

`remove(bookList, "On The Electrodynamics of Moving Bodies")`은 `["The Hound of the Baskervilles", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"]`를 반환해야 합니다.

```js
assert(
  JSON.stringify(remove(bookList, 'On The Electrodynamics of Moving Bodies')) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae'
    ])
);
```

`remove(add(bookList, "A Brief History of Time"), "On The Electrodynamics of Moving Bodies");`은 `["The Hound of the Baskervilles", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae", "A Brief History of Time"]`과 같아야 합니다.

```js
assert(
  JSON.stringify(remove(add(bookList, 'A Brief History of Time'), 'On The Electrodynamics of Moving Bodies')) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae',
      'A Brief History of Time'
    ])
);
```

# --seed--

## --seed-contents--

```js
// The global variable
const bookList = ["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"];

// Change code below this line
function add(bookName) {

  bookList.push(bookName);
  return bookList;
  
  // Change code above this line
}

// Change code below this line
function remove(bookName) {
  const book_index = bookList.indexOf(bookName);
  if (book_index >= 0) {

    bookList.splice(book_index, 1);
    return bookList;

    // Change code above this line
    }
}
```

# --solutions--

```js
// The global variable
const bookList = ["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"];

function add(bookList, bookName) {
  return [...bookList, bookName];
}

function remove(bookList, bookName) {
  const bookListCopy = [...bookList];
  const bookNameIndex = bookList.indexOf(bookName);
  if (bookNameIndex >= 0) {
    bookListCopy.splice(bookNameIndex, 1);
  }
  return bookListCopy;
}
```
