---
id: 587d7daa367417b2b2512b6b
title: split 메서드를 사용해 문자열을 배열로 나누기
challengeType: 1
forumTopicId: 18305
dashedName: split-a-string-into-an-array-using-the-split-method
---

# --description--

`split` 메서드는 문자열을 문자열 배열로 나눕니다. 구분자로 사용할 수 있는 문자를 인수로 받으며, 정규 표현식도 가능합니다. 예를 들어 구분자가 공백이면 단어 배열을 얻고, 구분자가 빈 문자열이면 문자열의 각 문자를 요소로 하는 배열을 얻습니다.

다음은 하나의 문자열을 공백으로 나누고, 다른 문자열을 정규 표현식으로 숫자를 기준으로 나누는 두 가지 예입니다:

```js
const str = "Hello World";
const bySpace = str.split(" ");

const otherString = "How9are7you2today";
const byDigits = otherString.split(/\d/);
```

`bySpace`의 값은 `["Hello", "World"]`이고 `byDigits`의 값은 `["How", "are", "you", "today"]`입니다.

문자열은 불변이므로 `split` 메서드를 사용하면 문자열 작업이 더 쉬워집니다.

# --instructions--

`split` 함수 안에서 `splitify` 메서드를 사용해 `str`을 단어 배열로 나누세요. 함수는 배열을 반환해야 합니다. 단어는 항상 공백으로 구분되지 않으며, 배열에 구두점이 포함되지 않아야 합니다.

# --hints--

코드는 `split` 메서드를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\.split/g));
```

`splitify("Hello World,I-am code")`은 `["Hello", "World", "I", "am", "code"]`을 반환해야 합니다.

```js
assert(
  JSON.stringify(splitify('Hello World,I-am code')) ===
    JSON.stringify(['Hello', 'World', 'I', 'am', 'code'])
);
```

`splitify("Earth-is-our home")`은 `["Earth", "is", "our", "home"]`을 반환해야 합니다.

```js
assert(
  JSON.stringify(splitify('Earth-is-our home')) ===
    JSON.stringify(['Earth', 'is', 'our', 'home'])
);
```

`splitify("This.is.a-sentence")`는 `["This", "is", "a", "sentence"]`을 반환해야 합니다.

```js
assert(
  JSON.stringify(splitify('This.is.a-sentence')) ===
    JSON.stringify(['This', 'is', 'a', 'sentence'])
);
```

# --seed--

## --seed-contents--

```js
function splitify(str) {
  // Only change code below this line


  // Only change code above this line
}

splitify("Hello World,I-am code");
```

# --solutions--

```js
function splitify(str) {
  return str.split(/\W/);
}
```
