---
id: 587d7dab367417b2b2512b6e
title: 배열의 모든 요소가 조건을 충족하는지 확인하려면 every 메서드를 사용하세요
challengeType: 1
forumTopicId: 301312
dashedName: use-the-every-method-to-check-that-every-element-in-an-array-meets-a-criteria
---

# --description--

`every` 메서드는 배열과 함께 작동하여 *모든* 요소가 특정 테스트를 통과하는지 확인합니다. 모든 값이 조건을 충족하면 불리언 값 `true`을 반환하고, 그렇지 않으면 `false`를 반환합니다.

예를 들어, 다음 코드는 `numbers` 배열의 모든 요소가 10보다 작은지 확인합니다:

```js
const numbers = [1, 5, 8, 0, 10, 11];

numbers.every(function(currentValue) {
  return currentValue < 10;
});
```

여기서 `every` 메서드는 `false`를 반환합니다.

# --instructions--

`every` 함수 안에서 `checkPositive` 메서드를 사용하여 `arr`의 모든 요소가 양수인지 확인하세요. 함수는 불리언 값을 반환해야 합니다.

# --hints--

코드는 `every` 메서드를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\.every/g));
```

`checkPositive([1, 2, 3, -4, 5])`은 `false`을 반환해야 합니다.

```js
assert.isFalse(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])`은 `true`을 반환해야 합니다.

```js
assert.isTrue(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([1, -2, 3, -4, 5])`는 `false`를 반환해야 합니다.

```js
assert.isFalse(checkPositive([1, -2, 3, -4, 5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.every(num => num > 0);
}
```
