---
id: 587d7da9367417b2b2512b6a
title: 원본 배열을 변경하지 않고 정렬된 배열 반환하기
challengeType: 1
forumTopicId: 301237
dashedName: return-a-sorted-array-without-changing-the-original-array
---

# --description--

`sort` 메서드의 부작용은 원본 배열의 요소 순서를 변경한다는 점입니다. 즉, 배열을 제자리에서 변경합니다. 이를 피하는 한 가지 방법은 먼저 정렬할 배열에 빈 배열을 연결하는 것인데(`slice`과 `concat`가 새 배열을 반환한다는 점을 기억하세요), 그다음 `sort` 메서드를 실행하는 것입니다.

# --instructions--

`sort` 함수에서 `nonMutatingSort` 메서드를 사용해 배열 요소를 오름차순으로 정렬하세요. 함수는 새 배열을 반환해야 하며 `globalArray` 변수를 변경하면 안 됩니다.

# --hints--

코드에서 `sort` 메서드를 사용하세요.

```js
assert(nonMutatingSort.toString().match(/\.sort/g));
```

`globalArray` 변수는 변경되지 않아야 합니다.

```js
assert(JSON.stringify(globalArray) === JSON.stringify([5, 6, 3, 2, 9]));
```

`nonMutatingSort(globalArray)`는 `[2, 3, 5, 6, 9]`을 반환해야 합니다.

```js
assert(
  JSON.stringify(nonMutatingSort(globalArray)) ===
    JSON.stringify([2, 3, 5, 6, 9])
);
```

`nonMutatingSort(globalArray)`는 하드코딩하지 마세요.

```js
assert(!nonMutatingSort.toString().match(/\[.*?[23569].*?\]/gs));
```

함수는 전달받은 배열이 아닌 새 배열을 반환해야 합니다.

```js
assert(nonMutatingSort(globalArray) !== globalArray);
```

`nonMutatingSort([1, 30, 4, 21, 100000])`는 `[1, 4, 21, 30, 100000]`을 반환해야 합니다.

```js
assert(JSON.stringify(nonMutatingSort([1, 30, 4, 21, 100000])) ===
    JSON.stringify([1, 4, 21, 30, 100000]))
```

`nonMutatingSort([140000, 104, 99])`는 `[99, 104, 140000]`을 반환해야 합니다.

```js
assert(JSON.stringify(nonMutatingSort([140000, 104, 99])) ===
    JSON.stringify([99, 104, 140000]))
```

# --seed--

## --seed-contents--

```js
const globalArray = [5, 6, 3, 2, 9];

function nonMutatingSort(arr) {
  // Only change code below this line


  // Only change code above this line
}

nonMutatingSort(globalArray);
```

# --solutions--

```js
const globalArray = [5, 6, 3, 2, 9];
function nonMutatingSort(arr) {
  return [].concat(arr).sort((a,b) => a-b);
}
```
