---
id: 587d7da9367417b2b2512b66
title: concat 메서드를 사용해 두 배열 합치기
challengeType: 1
forumTopicId: 301229
dashedName: combine-two-arrays-using-the-concat-method
---

# --description--

<dfn>연결(concatenation)</dfn>은 항목을 끝에서 끝으로 이어 붙이는 것을 의미합니다. JavaScript는 문자열과 배열 모두에 동일하게 작동하는 `concat` 메서드를 제공합니다. 배열의 경우, 이 메서드는 하나의 배열에서 호출되고, 다른 배열이 `concat`의 인수로 제공되어 첫 번째 배열 끝에 추가됩니다. 이 메서드는 새로운 배열을 반환하며 원래 배열 둘 다를 변경하지 않습니다. 예시는 다음과 같습니다:

```js
[1, 2, 3].concat([4, 5, 6]);
```

반환된 배열은 `[1, 2, 3, 4, 5, 6]`가 됩니다.

# --instructions--

`concat` 함수에서 `nonMutatingConcat` 메서드를 사용해 `attach` 끝에 `original`를 연결하세요. 함수는 연결된 배열을 반환해야 합니다.

# --hints--

코드는 `concat` 메서드를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

`first` 배열은 변경되지 않아야 합니다.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

`second` 배열은 변경되지 않아야 합니다.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingConcat([1, 2, 3], [4, 5])`은 `[1, 2, 3, 4, 5]`을 반환해야 합니다.

```js
assert(
  JSON.stringify(nonMutatingConcat([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingConcat(original, attach) {
  // Only change code below this line


  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingConcat(first, second);
```

# --solutions--

```js
function nonMutatingConcat(original, attach) {
  return original.concat(attach);
}
const first = [1, 2, 3];
const second = [4, 5];
```
