---
id: 587d7b8e367417b2b2512b5f
title: 함수에서 외부 의존성을 피하려면 인수를 전달하세요
challengeType: 1
forumTopicId: 301234
dashedName: pass-arguments-to-avoid-external-dependence-in-a-function
---

# --description--

지난 도전 과제는 함수형 프로그래밍 원칙에 한 걸음 더 다가갔지만, 아직 부족한 점이 있습니다.

전역 변수 값을 변경하지는 않았지만, `incrementer` 함수는 전역 변수 `fixedValue`이 없으면 작동하지 않습니다.

함수형 프로그래밍의 또 다른 원칙은 의존성을 항상 명확히 선언하는 것입니다. 즉, 함수가 변수나 객체에 의존한다면, 그 변수나 객체를 인수로 직접 함수에 전달해야 합니다.

이 원칙에는 여러 가지 좋은 결과가 있습니다. 함수 테스트가 더 쉬워지고, 어떤 입력을 받는지 정확히 알 수 있으며, 프로그램 내 다른 부분에 의존하지 않게 됩니다.

이로 인해 코드를 수정, 삭제 또는 추가할 때 더 자신감을 가질 수 있습니다. 무엇을 변경할 수 있고 무엇을 변경할 수 없는지 알 수 있으며, 잠재적인 함정이 어디에 있는지도 파악할 수 있습니다.

마지막으로, 함수는 코드의 어느 부분에서 실행되든 동일한 입력에 대해 항상 같은 출력을 생성합니다.

# --instructions--

`incrementer` 함수를 업데이트하여 의존성을 명확히 선언하세요.

`incrementer` 함수를 작성하여 인수를 받고, 그 값을 1만큼 증가시킨 후 결과를 반환하세요.

# --hints--

`incrementer` 함수는 `fixedValue` (즉, `4`) 값을 변경하지 않아야 합니다.

```js
assert(fixedValue === 4);
```

`incrementer` 함수는 인수를 받아야 합니다.

```js
assert(incrementer.length === 1);
```

`incrementer` 함수는 `fixedValue` 값보다 1 큰 값을 반환해야 합니다.

```js
const __newValue = incrementer(fixedValue);
assert(__newValue === 5);
```

# --seed--

## --seed-contents--

```js
// The global variable
let fixedValue = 4;

// Only change code below this line
function incrementer() {


  // Only change code above this line
}
```

# --solutions--

```js
let fixedValue = 4;

function incrementer(fixedValue) {
  return fixedValue + 1;
}
```
