---
id: 587d7b8d367417b2b2512b5b
title: 함수형 프로그래밍 배우기
challengeType: 1
forumTopicId: 301233
dashedName: learn-about-functional-programming
---

# --description--

함수형 프로그래밍은 함수 범위 밖에 부작용이 없는 단순하고 독립적인 함수들로 해결책을 만드는 프로그래밍 스타일입니다: `INPUT -> PROCESS -> OUTPUT`

함수형 프로그래밍은 다음과 같습니다:

1) 독립적인 함수 - 전역 변수와 같이 변할 수 있는 프로그램 상태에 의존하지 않습니다

2) 순수한 함수 - 같은 입력은 항상 같은 출력을 반환합니다

3) 제한된 부작용을 가진 함수 - 함수 외부의 프로그램 상태에 대한 모든 변경이나 변이는 신중하게 제어됩니다

# --instructions--

freeCodeCamp 구성원들은 차를 정말 좋아합니다.

코드 편집기에서 `prepareTea` 함수와 `getTea` 함수가 이미 정의되어 있습니다. `getTea` 함수를 호출하여 팀을 위해 40잔의 차를 받고, `tea4TeamFCC` 변수에 저장하세요.

# --hints--

`tea4TeamFCC` 변수에는 팀을 위한 40잔의 차가 저장되어야 합니다.

```js
assert(tea4TeamFCC.length === 40);
```

`tea4TeamFCC` 변수에는 녹차 잔이 저장되어야 합니다.

```js
assert(tea4TeamFCC[0] === 'greenTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareTea = () => 'greenTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4TeamFCC = null;
// Only change code above this line
```

# --solutions--

```js
const prepareTea = () => 'greenTea';

const getTea = (numOfCups) => {
  const teaCups = [];
  
  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }

  return teaCups;
};

const tea4TeamFCC = getTea(40); 
```
