---
id: 63ee3f71381756f9716727ef
title: Вправи з основ CSS А
challengeType: 25
dashedName: css-foundations-exercise-a
---

# --description--

**Мета:**  
У цьому завданні ви практикуватимете додавання CSS до HTML-файлу, використовуючи всі три методи: зовнішній CSS, внутрішній CSS і рядковий CSS. Для додавання стилів за допомогою зовнішнього та внутрішнього методів слід використовувати лише селектори типу. Також слід використовувати ключові слова для кольорів (наприклад, "синій") замість використання значень `RGB` або `HEX`.

## Історія користувача

- Ви повинні бачити елемент `div` з деяким текстом.
  - Він має мати фон `red`, текст `white`, розмір шрифту `32px`, текст вирівняний по центру та `bold`.
  - CSS для елемента `div` має бути доданий зовнішньо, використовуючи селектор типу.
- Ви повинні бачити елемент `p` з деяким текстом.
  - Він має мати фон `green`, текст `white` і розмір шрифту `18px`.
  - CSS для елемента `p` має бути доданий внутрішньо, використовуючи селектор типу.
- Ви повинні бачити елемент `button` з деяким текстом.
  - Елемент `button` має мати фон `orange` і розмір шрифту `18px`.
  - CSS для елемента `button` має бути доданий за допомогою рядкових стилів.

# --hints--

Ви повинні мати один елемент `div`, що містить деякий текст.

```js
const divElement = document.querySelector('div');

assert.isNotNull(divElement);
assert.isAtLeast(divElement?.innerText.length, 1);
```

Ви повинні мати зовнішній файл стилів, що містить стилі для елемента `div`.

```js
const styleSheet = new __helpers.CSSHelp(document).getStyleSheet();
const isExternal = styleSheet?.ownerNode.classList.contains('fcc-injected-styles');
const divStyle = new __helpers.CSSHelp(document).getStyle('div');

assert.isTrue(isExternal);
assert.isNotNull(divStyle);
```

Ваш елемент `div` не повинен мати CSS, доданий внутрішньо або рядковими стилями.

```js
const styleElement = document.querySelector('style:not([class])');

assert.isNotTrue(styleElement?.innerText.includes('div'));
assert.isNotTrue(document.querySelector('div')?.hasAttribute('style'));
```

Ваш елемент `div` має мати `background-color` зі значенням `red` і `color` зі значенням `white`.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const divBGColor = divStyle?.getPropertyValue('background-color');
const divColor = divStyle?.getPropertyValue('color');

assert.equal(divBGColor, 'red');
assert.equal(divColor, 'white');
```

Ваш елемент `div` має мати `font-weight`, встановлений у `bold`, `font-size`, встановлений у `32px`, і `text-align`, встановлений у `center`.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const textAlign = divStyle?.getPropertyValue('text-align');
const fontSize = divStyle?.getPropertyValue('font-size');
const fontWeight = divStyle?.getPropertyValue('font-weight');

assert.equal(textAlign, 'center');
assert.equal(fontSize, '32px');
assert.equal(fontWeight,'bold');
```

Ви повинні мати один елемент `p`, і він має містити деякий текст.

```js
const pElement = document.querySelector('p');

assert.isNotNull(pElement);
assert.isAtLeast(pElement?.innerText.length, 1)
```

Ваш елемент `p` має мати стилі, додані внутрішньо за допомогою елемента `style`.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let isStyled = false;

if (rules && rules.selectorText === 'p') {
  isStyled = true;
}

assert.isTrue(isStyled);
```

Ваш елемент `p` має мати `font-size` зі значенням `18px` і `color`, встановлений у `white`.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let fontSize, color;

if (rules && rules.selectorText === 'p') {
  fontSize = rules.style.fontSize;
  color = rules.style.color;
}

assert.equal(fontSize, "18px");
assert.equal(color, 'white');
```

Ви повинні мати один елемент `button`, що містить деякий текст.

```js
const btnElement = document.querySelector('button');

assert.isNotNull(btnElement);
assert.isAtLeast(btnElement?.innerText.length, 1);
```

Ваш елемент `button` має мати рядковий стиль.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('style'));
```
 
Ваш елемент `button` має мати `background-color`, встановлений у `orange`.

```js
assert.equal(document.querySelector('button')?.style.backgroundColor, 'orange')
```

Ваш елемент `button` має мати `font-size`, встановлений у `18px`.

```js
assert.equal(document.querySelector('button')?.style.fontSize, '18px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>

  </body>
</html>
```

```css
/* styles.css */
```

# --solutions--

```html
<!DOCTYPE html>
<html>
<head>
  <title>My Styling Example</title>
  <link rel="stylesheet" type="text/css" href="styles.css">
  <style>
    p {
      background-color: green;
      color: white;
      font-size: 18px;
    }
  </style>
</head>
<body>
  <div>Hello World!</div>
  <p>This is a paragraph.</p>
  <button style="background-color: orange; font-size: 18px;">Click Me</button>
</body>
</html>
```

```css
div {
  background-color: red;
  color: white;
  font-size: 32px;
  text-align: center;
  font-weight: bold;
}
```
