---
id: 66c8ba975ee7230e29f6c4b0
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Ahora, debe verificar si el/la solicitante está calificado/a sólo para un préstamo de coche. Si lo son, devuelve la cadena `"You qualify for a car loan."`.

# --hints--

Su función `getLoanMessage` debe devolver una cadena si el/la solicitante califica para un préstamo de coche.

```js
assert.isString(getLoanMessage(30000, 650));
```

Su función `getLoanMessage` debe devolver la cadena `"You qualify for a car loan."`.

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

Su función `getLoanMessage` debe devolver `undefined` si los ingresos anuales y el puntaje de crédito del/la solicitante no cumplen con los requisitos para un préstamo de coche.

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 

--fcc-editable-region--
}
```
