---
id: 66c8ba41a77db20a93f9d7a1
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, revisarás los condicionales de JavaScript construyendo una aplicación verificada de calificación de préstamos.

La aplicación verificará si el usuario es elegible para un préstamo de dúplex, coche y apartamento en base a sus ingresos anuales y puntaje de crédito. Un puntaje de crédito es un número que representa qué tan bueno eres administrando dinero prestado.

Para empezar, crea las siguientes variables y valores.

| Nombre de la variable | Valor   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

Debes crear una variable `minIncomeForDuplex` establecida en `60000`.

```js
assert.equal(minIncomeForDuplex, 60000);
```

Debes crear un `minCreditScoreForDuplex` establecido en `700`.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

Debes crear una variable `minIncomeForCar` establecida en `30000`.

```js
assert.equal(minIncomeForCar, 30000);
```

Debes crear una variable `minCreditScoreForCar` establecida en `650`.

```js
assert.equal(minCreditScoreForCar, 650);
```

Debes crear una variable `minIncomeForCondo` establecida en `45000`.

```js
assert.equal(minIncomeForCondo, 45000);
```

Debes crear una variable `minCreditScoreForCondo` establecida en `680`.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
