---
id: 66c8ba975ee7230e29f6c4ad
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Para verificar qué préstamo califica un usuario según el `annualIncome` y `creditScore`, debes usar una instrucción `if/else if` o un ternario directamente dentro de la función `getLoanMessage`. Luego, devolverás el mensaje apropiado en el bloque de cada condición.

Comenzando con el préstamo para dúplex, verifica si `annualIncome` es mayor o igual a `minIncomeForDuplex` Y si `creditScore` es mayor o igual a `minCreditScoreForDuplex`.

Si esa condición es verdadera, entonces el solicitante es elegible para un préstamo para dúplex y los otros préstamos. Entonces, dentro de la verificación, devuelve la cadena `"You qualify for a duplex, condo, and car loan."`

# --hints--

Tu función `getLoanMessage` debe devolver una cadena.

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

Tu función `getLoanMessage` debe devolver la cadena `"You qualify for a duplex, condo, and car loan."`.

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

Tu función `getLoanMessage` debe devolver `undefined` si los ingresos anuales y la puntuación de crédito del solicitante no cumplen con los requisitos para un préstamo para dúplex.

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--
function getLoanMessage(annualIncome, creditScore) {

}
--fcc-editable-region--
```
