---
id: 66c8ba975ee7230e29f6c4b2
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Ahora, es hora de probar tu función `getLoanMessage`.

Usa la tabla a continuación para crear `4` variables y sus valores:

| Variable Name | Value   |
| ----------- | ------- |
| `duplexLoanMsg`       | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`   | `getLoanMessage(65000, 690)`   |
| `carLoanMsg`   | `getLoanMessage(45000, 660)`   |
| `noLoanMsg`   | `getLoanMessage(25000, 550)`   |

Después de eso, registra cada variable en la consola para ver los mensajes.

Con eso, ¡tu proyecto de verificación de calificación de préstamos está completo!

# --hints--

Deberías tener una variable llamada `duplexLoanMsg`.

```js
assert.isNotNull(duplexLoanMsg);
```

Tu variable `duplexLoanMsg` debería estar configurada con el resultado de `getLoanMessage(85000, 850)`.

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

Deberías registrar la variable `duplexLoanMsg` en la consola.

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

Deberías tener una variable llamada `condoLoanMsg`.

```js
assert.isNotNull(condoLoanMsg);
```

Tu variable `condoLoanMsg` debería estar configurada con el resultado de `getLoanMessage(65000, 690)`.

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

Deberías registrar la variable `condoLoanMsg` en la consola.

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

Deberías tener una variable llamada `carLoanMsg`.

```js
assert.isNotNull(carLoanMsg);
```

Tu variable `carLoanMsg` debería estar configurada con el resultado de `getLoanMessage(45000, 660)`.

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

Deberías registrar la variable `carLoanMsg` en la consola.

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

Deberías tener una variable llamada `noLoanMsg`.

```js
assert.isNotNull(noLoanMsg);
```

Tu variable `noLoanMsg` debería estar configurada con el resultado de `getLoanMessage(25000, 550)`.

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

Deberías registrar la variable `noLoanMsg` en la consola.

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
