---
id: 66c8ba975ee7230e29f6c4af
title: Pasos 4
challengeType: 1
dashedName: step-4
---

# --description--

Si el ingreso anual del solicitante es mayor o igual a `minIncomeForCondo`, Y si su puntaje de crédito es mayor o igual a `minCreditScoreForCondo`, entonces califica para un préstamo de condominio y carro.

Verifica si eso es cierto en la función `getLoanMessage`. Si lo es, devuelve la cadena `"You qualify for a condo and car loan."`

# --hints--

La función `getLoanMessage` debe retornar una cadena si el solicitante califica para un préstamo de condominio y carro.

```js
assert.isString(getLoanMessage(45000, 680));
```

La función `getLoanMessage` debe retornar la cadena `"You qualify for a condo and car loan."`.

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

La función `getLoanMessage` debe retornar `undefined` si los ingresos anuales y el puntaje de crédito del solicitante no cumplen con los requisitos para un préstamo de condominio.

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }

--fcc-editable-region--
}
```
