---
id: 6899b1d3bb656e56bd6937e2
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Erstellen Sie unter dem `p`-Element ein `div`-Element mit der Klasse `btn-container`.

Innerhalb des `div`-Elements erstellen Sie ein `button`-Element.

Dieses `button` sollte eine `id` von `happy-btn`, eine Klasse von `emoji-btn` und eine `aria-label` von `Happy face emoji` haben.

# --hints--

Erstellen Sie ein neues `div`-Element unter dem `p`-Element.

```js
assert.exists(document.querySelector("p + div"));
```

Das neue `div`-Element sollte die Klasse `btn-container` haben.

```js
assert.equal(document.querySelector("p + div")?.className, "btn-container");
```

Es sollte ein `button`-Element innerhalb des `div`-Elements geben.

```js
assert.exists(document.querySelector("div > button"));
```

Das `button`-Element sollte eine `id` von `happy-btn` haben.

```js
assert.equal(document.querySelector("div > button")?.id, "happy-btn");
```

Das `button`-Element sollte eine `class` von `emoji-btn` haben.

```js
assert.equal(document.querySelector("div > button")?.className, "emoji-btn");
```

Das `button`-Element sollte eine `aria-label` von `Happy face emoji` haben.

```js
assert.equal(document.querySelector("div > button")?.getAttribute("aria-label"), "Happy face emoji");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <main>
    <h1 class="title">How are you feeling today?</h1>
    <p class="description">
      Click on the buttons below to rate your emotions.
    </p>
--fcc-editable-region--
    
--fcc-editable-region--
  </main>
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```
