---
id: 587d78b1367417b2b2512b09
title: "Rendere un'immagine responsiva"
challengeType: 0
forumTopicId: 301140
dashedName: make-an-image-responsive
---

# --description--

Rendere le immagini responsivi con CSS è in realtà molto semplice. Devi solo aggiungere queste proprietà a un'immagine:

```css
img {
  max-width: 100%;
  height: auto;
}
```

L'`max-width` di `100%` farà in modo che l'immagine non sia mai più larga del contenitore in cui si trova, e l'`height` di `auto` farà mantenere all'immagine il suo rapporto d'aspetto originale.

# --instructions--

Aggiungi le regole di definire lo stile (verb) alla classe `responsive-img` per renderla responsiva. Non dovrebbe mai essere più larga del suo contenitore (in questo caso, è la finestra di anteprima) e dovrebbe mantenere il suo rapporto d'aspetto originale. Dopo aver aggiunto il tuo codice, ridimensiona l'anteprima per vedere come si comportano le tue immagini.

# --hints--

La tua classe `responsive-img` dovrebbe avere un insieme `max-width` impostato su `100%`.

```js
assert.strictEqual(getComputedStyle(document.querySelector('.responsive-img')).maxWidth, '100%');
```

La tua classe `responsive-img` dovrebbe avere un insieme `height` impostato su `auto`.

```js
assert.match(code, /height:\s*?auto;/g);
```

# --seed--

## --seed-contents--

```html
<style>
.responsive-img {


}

img {
  width: 600px;
}
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```

# --solutions--

```html
<style>
  .responsive-img {
    max-width: 100%;
    height: auto;
  }

  img {
    width: 600px;
  }
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```
