---
id: 587d78b1367417b2b2512b0a
title: "Usa un'immagine retina per display ad alta risoluzione"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cVZ4Rfp'
forumTopicId: 301142
dashedName: use-a-retina-image-for-higher-resolution-displays
---

# --description--

Con l'aumento dei dispositivi connessi a internet, le loro dimensioni e specifiche variano, e i display che usano possono essere diversi sia esternamente che internamente. La densità di pixel è un aspetto che può differire da un dispositivo all'altro e questa densità è nota come Pixel Per Inch (PPI) o Dots Per Inch (DPI). Il display più famoso di questo tipo è quello noto come "Retina Display" sugli ultimi notebook Apple MacBook Pro e recentemente sui computer iMac. A causa della differenza nella densità di pixel tra i display "Retina" e "Non-Retina", alcune immagini non create pensando a un display ad alta risoluzione potrebbero apparire "pixelate" quando vengono renderizzate su un display ad alta risoluzione.

Il modo più semplice per far apparire correttamente le tue immagini su display ad alta risoluzione, come il "retina display" dei MacBook Pro, è definire i valori `width` e `height` come la metà di quelli del file originale. Ecco un esempio di un'immagine che usa solo metà dell'altezza e della larghezza originali:

```html
<style>
  img {
    height: 250px;
    width: 250px;
  }
</style>
<img src="coolPic500x500" alt="A most excellent picture" />
```

# --instructions--

Imposta i valori `width` e `height` del tag `img` a metà dei loro valori originali. In questo caso, sia il valore originale `height` che quello originale `width` sono `200px`.

# --hints--

Il tuo tag `img` dovrebbe avere un `width` di 100 pixel.

```js
assert.strictEqual(document.querySelector('img').width, 100);
```

Il tuo tag `img` dovrebbe avere un `height` di 100 pixel.

```js
assert.strictEqual(document.querySelector('img').height, 100);
```

# --seed--

## --seed-contents--

```html
<style></style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```

# --solutions--

```html
<style>
  img {
    height: 100px;
    width: 100px;
  }
</style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```
