---
id: 587d78b1367417b2b2512b0c
title: Rendere responsiva la tipografia
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/crzN7T8'
forumTopicId: 301141
dashedName: make-typography-responsive
---

# --description--

Invece di usare `em` o `px` per dimensionare il testo, puoi usare le unità viewport per una tipografia responsiva. Le unità viewport, come le percentuali, sono unità relative, ma si basano su elementi diversi. Le unità viewport sono relative alle dimensioni del viewport (larghezza o altezza) di un dispositivo, mentre le percentuali sono relative alla dimensione del contenitore genitore.

Le quattro diverse unità viewport sono:

<ul><li><code>vw</code> (viewport width): <code>10vw</code> corrisponderebbe al 10% della larghezza del viewport.</li><li><code>vh</code> (viewport height): <code>3vh</code> corrisponderebbe al 3% dell'altezza del viewport.</li><li><code>vmin</code> (viewport minimum): <code>70vmin</code> corrisponderebbe al 70% della dimensione minore del viewport (altezza o larghezza).</li><li><code>vmax</code> (viewport maximum): <code>100vmax</code> corrisponderebbe al 100% della dimensione maggiore del viewport (altezza o larghezza).</li></ul>

Ecco un esempio che imposta un tag `body` al 30% della larghezza del viewport.

```css
body {
  width: 30vw;
}
```

# --instructions--

Imposta la `width` del tag `h2` all'80% della larghezza del viewport e la `width` del paragrafo al 75% della dimensione minore del viewport.

# --hints--

Il tuo tag `h2` dovrebbe avere una `width` di 80vw.

```js
assert.match(
  __helpers.removeCssComments(code),
  /h2\s*?{\s*?width:\s*?80vw;\s*?}/g
);
```

Il tuo tag `p` dovrebbe avere una `width` di 75vmin.

```js
assert.match(
  __helpers.removeCssComments(code),
  /p\s*?{\s*?width:\s*?75vmin;\s*?}/g
);
```

# --seed--

## --seed-contents--

```html
<style>

</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  h2 {
    width: 80vw;
  }
  p {
    width: 75vmin;
  }
</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
