---
id: a9bd25c716030ec90084d8a1
title: 實作 Chunky Monkey 演算法
challengeType: 26
dashedName: implement-the-chunky-monkey-algorithm
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 撰寫一個名為 `chunkArrayInGroups` 的函式（程式），該函式以陣列作為第一個引數，以數字作為第二個引數。該函式應將陣列拆分成長度等於第二個引數的較小陣列，並將它們傳回為二維陣列。

# --hints--

`chunkArrayInGroups(["a", "b", "c", "d"], 2)` 應該傳回 `[["a", "b"], ["c", "d"]]`。

```js
assert.deepEqual(chunkArrayInGroups(['a', 'b', 'c', 'd'], 2), [
  ['a', 'b'],
  ['c', 'd']
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 3)` 應該傳回 `[[0, 1, 2], [3, 4, 5]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 3), [
  [0, 1, 2],
  [3, 4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 2)` 應該傳回 `[[0, 1], [2, 3], [4, 5]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 2), [
  [0, 1],
  [2, 3],
  [4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 4)` 應該傳回 `[[0, 1, 2, 3], [4, 5]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 4), [
  [0, 1, 2, 3],
  [4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6], 3)` 應該傳回 `[[0, 1, 2], [3, 4, 5], [6]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6], 3), [
  [0, 1, 2],
  [3, 4, 5],
  [6]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 4)` 應該傳回 `[[0, 1, 2, 3], [4, 5, 6, 7], [8]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 4), [
  [0, 1, 2, 3],
  [4, 5, 6, 7],
  [8]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 2)` 應該傳回 `[[0, 1], [2, 3], [4, 5], [6, 7], [8]]`。

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 2), [
  [0, 1],
  [2, 3],
  [4, 5],
  [6, 7],
  [8]
]);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function chunkArrayInGroups(arr, size) {
  let out = [];

  for (let i = 0; i < arr.length; i += size) {
    out.push(arr.slice(i, i + size));
  }

  return out;
}

chunkArrayInGroups(['a', 'b', 'c', 'd'], 2);
```
