---
id: 69757cc0faae0152c1418aad
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, practicarás trabajar con números y operaciones matemáticas para construir un divisor de cuentas. Esta herramienta calculará cuánto debe cada persona después de sumar los costos de la comida y la propina.

Para comenzar, necesitas una forma de llevar un registro del monto total a medida que se agregan los costos. En Python, puedes usar una variable para almacenar un entero (un número entero) que cambia con el tiempo.

Por ejemplo, podrías escribir:

```py
my_number = 2
```

Crea una variable llamada `running_total` y asígnale el valor `0`.

# --hints--

Debes tener una variable llamada `running_total`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

Debes asignar el entero `0` a tu variable `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
