---
id: 6976cc0b0c135686a4620b94
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

A continuación, necesitas tener en cuenta el número de personas que comparten la cuenta. Guarda este valor en una variable, como hiciste en el paso anterior.

Crea una variable llamada `num_of_friends` y asígnale el valor de `4`. Esto se usará más adelante en el taller para calcular la división final.

# --hints--

Debes tener una variable llamada `num_of_friends`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

Debes asignar el entero `4` a tu variable `num_of_friends`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
