---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Cada curso tiene un costo. Necesitas almacenar estos montos en variables para usarlos después. Dado que estos montos incluyen centavos, usarás el tipo float, que se usa para representar números decimales. Aquí tienes un ejemplo de una variable con un valor float:

```py
change = 2.35
```

Crea cuatro variables: `appetizers` con valor `37.89`, `main_courses` con valor `57.34`, `desserts` con valor `39.39` y `drinks` con valor `64.21`. 

# --hints--

Debes definir una variable llamada `appetizers` y asignarle el valor `37.89`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

Debes definir una variable llamada `main_courses` y asignarle el valor `57.34`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

Debes definir una variable llamada `desserts` y asignarle el valor `39.39`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

Debes definir una variable llamada `drinks` y asignarle el valor `64.21`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
