---
id: 69779dcb44016eccc05c15a4
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

El servicio fue excelente, así que el grupo decide dejar una propina del 25%. Para calcular un porcentaje en Python, puedes multiplicar el total por el equivalente decimal del porcentaje.

Por ejemplo, para encontrar el 10% de un valor, lo multiplicarías por `0.10` usando el operador `*`:

```py
tax = total * 0.10
```

Crea una variable llamada `tip` y asígnale el resultado de multiplicar `running_total` por `0.25`.

Finalmente, usa `print()` para mostrar la cadena `Tip amount:` seguida de un espacio y el valor de tu variable `tip`.

# --hints--

Debes tener una variable llamada `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

Tu variable `tip` debe ser el resultado de `running_total * 0.25`.

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

Debes imprimir la cadena `Tip amount:` seguida de un espacio y la variable `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tip amount:', tip)") or _Node(_code).has_call("print(f'Tip amount: {tip}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total = appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
