---
id: 6977a8bbcbdd557886bd0718
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Con la propina ya incluida, tienes el monto final para todo el grupo. Tienes que determinar cuánto debe cada persona dividiendo la cuenta total entre el número de amigos. 

En Python, usas la barra diagonal `/` para realizar divisiones. Por ejemplo:

```py
half = 10 / 2
```

Crea una variable llamada `final_bill` y asígnale el resultado de dividir `running_total` entre `num_of_friends`. 

Finalmente, usa la función `print()` para mostrar la cadena `Bill per person:` seguida de un espacio y el valor de `final_bill`.

# --hints--

Debes definir una variable llamada `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

Debes usar el operador `/` para dividir `running_total` entre `num_of_friends` y asignar el resultado a tu variable `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

Debes imprimir la cadena `Bill per person:` seguida de un espacio y la variable `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
