---
id: 6977a4306bb3e856690a4890
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Ahora que has calculado la propina, necesitas añadirla a tu `running_total` para encontrar el monto final de la cuenta.

En Python, puedes usar el operador de asignación aumentada `+=` para sumar un valor a una variable y actualizar esa variable al mismo tiempo. Por ejemplo, `total += 5` es una forma abreviada de escribir `total = total + 5`.

Usa el operador `+=` para sumar el valor de `tip` a tu `running_total`. Finalmente, usa `print()` para mostrar la cadena `Total with tip:` seguida de un espacio y el valor de `running_total`.

# --hints--

Debes usar el operador `+=` para sumar el valor de `tip` a tu `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

Debes imprimir la cadena `Total with tip:` seguida de un espacio y la variable `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Total with tip:', running_total)") or _Node(_code).has_call("print(f'Total with tip: {running_total}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
