---
id: 6982684f3a25f379e195a5fc
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora que has almacenado los costos individuales, puedes calcular el total. En Python, usas el operador de suma `+` para sumar valores. 

Actualiza la variable `running_total` sumando `appetizers`, `main_courses`, `desserts` y `drinks`. 

Finalmente, usa `print()` para mostrar la cadena `Total bill so far:` seguida de un espacio y el valor de `running_total`.

**Nota:** Puede que notes que la salida tiene más dígitos decimales de los esperados. Como aprendiste en una lección anterior, esto sucede porque los números se almacenan en binario, y muchos valores decimales no pueden representarse exactamente en este formato, lo que provoca errores de redondeo.

# --hints--

Debes actualizar `running_total` usando el operador `+` para sumar el costo de los cuatro cursos. 

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total = {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

Debes imprimir la cadena `Total bill so far:` seguida de un espacio y el valor de `running_total`.

```js
({
    test: () => runPython(`
    sol1 = "print('Total bill so far:', running_total)"
    sol2 = "print(f'Total bill so far: {running_total}')"
    assert _Node(_code).has_call(sol1) or _Node(_code).has_call(sol2)`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
