---
id: 6982684f3a25f379e195a5fc
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora que has almacenado los costos individuales, puedes calcular el total.

Recuerda que el operador `+=` suma un valor a una variable existente y la actualiza al mismo tiempo. Por ejemplo:

```py
total = 10
total += 2 + 2 + 1
print(total)  # total is now 15
```

Usa el operador `+=` una vez para agregar `appetizers`, `main_courses`, `desserts` y `drinks` a `running_total`.

Finalmente, usa `print()` para mostrar la cadena `Total bill so far:` seguida de un espacio y el valor de `running_total`.

**Nota:** Puede que notes que la salida tiene más dígitos decimales de los esperados. Como aprendiste en una lección anterior, esto sucede porque los números se almacenan en binario, y muchos valores decimales no pueden representarse exactamente en este formato, lo que provoca errores de redondeo.

# --hints--

Debes actualizar `running_total` usando el operador `+=` una vez para sumar los cuatro costos del curso.

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total += {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

Debes imprimir la cadena `Total bill so far:` seguida de un espacio y el valor de `running_total`.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Total bill so far: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - (37.89 + 57.34 + 39.39 + 64.21)) < 1e-6
`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
