---
id: 68e707257ea91c2e12d489fb
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop üben Sie, wie Sie benutzerdefinierte Styles für Optionsschaltflächen hinzufügen, indem Sie ein Formular für Elternsprechtage erstellen. Der HTML-Boilerplate-Code wurde für Sie bereitgestellt.

Beginnen Sie damit, ein `main`-Element mit einer Klasse namens `container` hinzuzufügen.

# --hints--

Sie sollten ein `main`-Element haben.

```js
assert.exists(document.querySelector("main"));
```

Ihr `main`-Element sollte eine Klasse namens `container` haben.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
