---
id: 68e7e729f6d14f69a23786ef
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie als Nächstes innerhalb Ihres `main`-Elements ein `h1`-Element mit dem Text `Parent Teacher Conference Form` und den Klassen `title` und `center` hinzu.

# --hints--

Sie sollten ein `h1`-Element haben.

```js
assert.exists(document.querySelector("h1"));
```

Ihr `h1`-Element sollte den Text `Parent Teacher Conference Form` enthalten.

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

Ihr `h1`-Element sollte die Klassen `title` und `center` haben.

```js
assert.exists(document.querySelector("h1.title.center"));
```

Ihr `h1`-Element sollte sich innerhalb des `main`-Elements befinden.

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
--fcc-editable-region--
    
--fcc-editable-region--
  </main>
</body>

</html>
```
