---
id: 68ea8fdf2c9b167307e7f671
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Jetzt ist es Zeit, die Formularelemente hinzuzufügen, um die Informationen der Eltern zu erfassen.

Beginnen Sie damit, ein `label`-Element mit dem Text `Parent/Guardian Name: ` und dem `for`-Attribut, das auf `"parent-name"` gesetzt ist, hinzuzufügen.

Fügen Sie dann unter Ihrem `label`-Element ein `input`-Element mit dem `type`-Attribut hinzu, das auf `"text"` gesetzt ist. Die `name`- und `id`-Attribute sollten auf `"parent-name"` gesetzt sein. Das `placeholder`-Attribut sollte auf `"E.g., Nancy Doe"` gesetzt sein. Schließlich sollte Ihr `input` erforderlich sein.

# --hints--

Sie sollten insgesamt drei `label`-Elemente auf der Seite haben.

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

Ihr drittes `label`-Element sollte den Text `Parent/Guardian Name: ` haben. Vergessen Sie nicht das Leerzeichen nach dem Doppelpunkt.

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

Ihr drittes `label`-Element sollte ein `for`-Attribut haben.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

Das `label`-Attribut Ihres dritten `for`-Elements sollte den Wert `"parent-name"` haben.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

Sie sollten drei `input`-Elemente auf der Seite haben.

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

Ihr drittes `input`-Element sollte unter Ihrem dritten `label`-Element sein.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Ihr drittes `input`-Element sollte ein `type`-Attribut haben, das auf `"text"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Ihr drittes `input`-Element sollte ein `name`-Attribut haben, das auf `"parent-name"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

Ihr drittes `input`-Element sollte ein `id`-Attribut haben, das auf `"parent-name"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

Ihr drittes `input`-Element sollte ein `placeholder`-Attribut haben, das auf `"E.g., Nancy Doe"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

Ihr drittes `input`-Element sollte ein `required`-Attribut haben.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
