---
id: 68ea975c7cbbb894a386e402
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Der nächste Schritt ist, die `label` und `input` Elemente für die E-Mail-Kontaktmethode hinzuzufügen.

Beginnen Sie damit, eine Bezeichnung mit der Klasse `"contact-method"` und einem `for`-Attribut, das auf `"email"` gesetzt ist, hinzuzufügen. Der Text der Bezeichnung sollte `Email: ` sein.

Fügen Sie unter Ihrem `label`-Element eine Optionsschaltfläche mit `id`- und `value`-Attributen hinzu, die auf `"email"` gesetzt sind. Das `name`-Attribut sollte auf `"contact-method"` gesetzt sein und die Klasse sollte `"contact-method-radio-btn"` sein.

Stellen Sie zuletzt sicher, dass diese Optionsschaltfläche standardmäßig `checked` ist.

# --hints--


Ihr letztes Fieldset sollte ein `label`-Element enthalten.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

Dieses neue `label`-Element (für E-Mail) sollte die Klasse `"contact-method"` haben.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

Das `label`-Element sollte ein `for`-Attribut haben, das auf `"email"` gesetzt ist.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

Das `label`-Element sollte den Text `Email: ` enthalten. Achten Sie auf das Leerzeichen nach dem Doppelpunkt.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

Sie sollten ein `input`-Element innerhalb des letzten `fieldset` haben.

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

Ihr `input`-Element sollte ein `type`-Attribut haben, das auf `radio` gesetzt ist.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

Das `input`-Element sollte ein `id`-Attribut haben, das auf `"email"` gesetzt ist.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

Das `input`-Element sollte ein `value`-Attribut haben, das auf `"email"` gesetzt ist.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

Das `input`-Element sollte ein `name`-Attribut haben, das auf `"contact-method"` gesetzt ist.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

Das `input`-Element sollte die Klasse `"contact-method-radio-btn"` haben.

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

Ihr `radio`-Button sollte standardmäßig `checked` sein.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
