---
id: 68ea9f01c51b66b937793611
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Fügen Sie als Nächstes eine weitere Bezeichnung mit einer Klasse von `"contact-method"` und einem `for`-Attribut hinzu, das auf `"phone"` gesetzt ist. Der Text der Bezeichnung sollte `Phone: ` sein.

Fügen Sie unter Ihrem `label`-Element eine Optionsschaltfläche mit `id`- und `value`-Attributen hinzu, die auf `"phone"` gesetzt sind. Das `name`-Attribut sollte auf `"contact-method"` gesetzt werden und die Klasse sollte auf `"contact-method-radio-btn"` gesetzt sein.

# --hints--

Sie sollten zwei `label`-Elemente im letzten `fieldset` haben.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type label"), 2);
```

Das `label`-Element sollte eine Klasse von `"contact-method"` haben.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

Das `label`-Element sollte ein `for`-Attribut haben, das auf `"phone"` gesetzt ist.

```js
assert.exists(document.querySelector("fieldset:last-of-type label[for='phone']"));
```

Das `label`-Element sollte den Text `Phone: ` haben.

```js
const label = document.querySelector("fieldset:last-of-type label[for='phone']");
assert.match(label?.innerText, /Phone:\s?/);
```

Sie sollten zwei `input`-Elemente innerhalb des letzten `fieldset` haben.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type input"), 2);
```

Ihr `input`-Element sollte ein `type`-Attribut haben, das auf `radio` gesetzt ist.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

Das `input`-Element sollte ein `id`-Attribut haben, das auf `"phone"` gesetzt ist.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("id"), "phone");
```

Das `input`-Element sollte ein `value`-Attribut haben, das auf `"phone"` gesetzt ist.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("value"), "phone");
```

Das `input`-Element sollte ein `name`-Attribut haben, das auf `"contact-method"` gesetzt ist.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

Das `input`-Element sollte eine Klasse von `"contact-method-radio-btn"` haben.

```js
assert.exists(document.querySelector("fieldset:last-of-type input[id='phone'].contact-method-radio-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
