---
id: 68ea7c41f2555e16174e7cfa
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Fügen Sie als Nächstes ein Absatz-Element mit dem Text `Please fill out the form below to help schedule your parent-teacher conference.` hinzu. Ihr Absatz-Element sollte außerdem die Klassen `description` und `center` haben.

# --hints--

Sie sollten ein `p`-Element haben.

```js
assert.exists(document.querySelector("p"));
```

Ihr `p`-Element sollte den Text `Please fill out the form below to help schedule your parent-teacher conference.` enthalten.

```js
assert.strictEqual(document.querySelector("p")?.innerText, "Please fill out the form below to help schedule your parent-teacher conference.");
```

Ihr `p`-Element sollte die Klassen `description` und `center` haben.

```js
assert.exists(document.querySelector("p.description.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
