---
id: 68eaa668d4d9fbe529a0e16b
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Fügen Sie nun ein `label`-Element mit dem Text `Any specific concerns or topics you'd like to discuss?` und einem `for`-Attribut hinzu, das auf `notes` gesetzt ist.

Fügen Sie unter diesem `label` ein `textarea`-Element hinzu. Setzen Sie dessen `id` auf `notes`, `name` auf `notes`, `rows` auf `4` und `cols` auf `50`.

# --hints--

Ihr letztes `fieldset` sollte ein `label`-Element enthalten.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

Das `label`-Element sollte den Text `Any specific concerns or topics you'd like to discuss?` haben.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

Das `label`-Element sollte ein `for`-Attribut haben, das auf `"notes"` gesetzt ist.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

Ihr letztes `fieldset` sollte ein `textarea`-Element enthalten.

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

Ihr `textarea`-Element sollte unter dem `label`-Element stehen.

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

Das `textarea`-Element sollte ein `id`-Attribut haben, das auf `"notes"` gesetzt ist.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

Das `textarea`-Element sollte ein `name`-Attribut haben, das auf `"notes"` gesetzt ist.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

Das `textarea`-Element sollte ein `rows`-Attribut haben, das auf `4` gesetzt ist.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

Das `textarea`-Element sollte ein `cols`-Attribut haben, das auf `50` gesetzt ist.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
