---
id: 68ea8666b8b473455f5b33ef
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Der erste Abschnitt des Formulars konzentriert sich auf die Informationen des Schülers, wie seinen Namen und seine Klasse.

Fügen Sie innerhalb des `form`-Elements ein `fieldset`-Element hinzu. Innerhalb dieses `fieldset`-Elements fügen Sie ein `legend`-Element mit dem Text `Student Information` hinzu.

Unterhalb des `legend`-Elements fügen Sie ein `label`-Element mit dem Text `Full Name: ` und ein `for`-Attribut mit dem Wert `student-name` hinzu.


# --hints--

Sie sollten ein `fieldset`-Element haben.

```js
assert.exists(document.querySelector("fieldset"));
```

Ihr `fieldset`-Element sollte sich innerhalb des `form`-Elements befinden.

```js
assert.exists(document.querySelector("form fieldset"));
```

Sie sollten ein `legend`-Element haben.

```js
assert.exists(document.querySelector("legend"));
```

Ihr `legend`-Element sollte sich innerhalb des `fieldset`-Elements befinden.

```js
assert.exists(document.querySelector("fieldset legend"));
```

Ihr `legend`-Element sollte den Text `Student Information` haben.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

Sie sollten ein `label`-Element haben.

```js
assert.exists(document.querySelector("label"));
```

Ihr `label`-Element sollte den Text `Full Name: ` haben. Vergessen Sie nicht das Leerzeichen nach dem Doppelpunkt.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

Ihr `label`-Element sollte ein `for`-Attribut haben.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

Das `label`-Attribut Ihres `for`-Elements sollte den Wert `student-name` haben.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

Ihr `label`-Element sollte unterhalb des `legend`-Elements sein.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
