---
id: 68eafe1bae175f3b2a87a90d
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Fügen Sie schließlich die Absenden-Schaltfläche für das Formular hinzu.

Fügen Sie unter dem letzten `fieldset`-Element ein `button`-Element mit der Klasse `submit-btn` und dem Attribut `type` mit dem Wert `"submit"` hinzu. Der Textinhalt der Schaltfläche sollte `Submit Form` sein.

Damit ist Ihre HTML-Struktur vollständig!

# --hints--

Sie sollten ein `button`-Element unter dem letzten `fieldset`-Element hinzufügen.

```js
assert.exists(document.querySelector("fieldset:last-of-type + button"));
```

Das `button`-Element sollte die Klasse `"submit-btn"` haben.

```js
assert.exists(document.querySelector("fieldset:last-of-type + button.submit-btn"));
```

Das `button`-Element sollte ein `type`-Attribut mit dem Wert `"submit"` haben.

```js
const button = document.querySelector("fieldset:last-of-type + button");
assert.strictEqual(button?.getAttribute("type"), "submit");
```

Das `button`-Element sollte den Text `Submit Form` haben.

```js
const button = document.querySelector("fieldset:last-of-type + button");
assert.strictEqual(button?.innerText, "Submit Form");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
