---
id: 68ea8a5db73b155708992687
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Der nächste Schritt besteht darin, das zugehörige `input`-Element für die Informationen des Lernenden hinzuzufügen.

Beginnen Sie damit, ein `input`-Element mit einem `type`-Attribut zu erstellen, das auf `"text"` gesetzt ist. Fügen Sie dann ein `name`- und ein `id`-Attribut hinzu, die beide auf `student-name` gesetzt sind.

Als Nächstes fügen Sie ein `placeholder`-Attribut hinzu, das auf `E.g., Jane Doe` gesetzt ist. Und schließlich fügen Sie ein `required`-Attribut hinzu.

# --hints--

Sie sollten ein `input`-Element haben.

```js
assert.exists(document.querySelector("fieldset input"));
```

Ihr `input`-Element sollte ein `type`-Attribut haben, das auf `"text"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Ihr `input`-Element sollte ein `name`-Attribut haben, das auf `"student-name"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

Ihr `input`-Element sollte ein `id`-Attribut haben, das auf `"student-name"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

Ihr `input`-Element sollte ein `placeholder`-Attribut haben, das auf `E.g., Jane Doe` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

Ihr `input`-Element sollte ein `required`-Attribut haben.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
