---
id: 68ea8c330ed07b60879f3d38
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Jetzt ist es Zeit, die Formularelemente hinzuzufügen, um die Noteninformationen der Lernenden zu erfassen.

Beginnen Sie damit, ein weiteres `label`-Element mit dem Text `Student Grade: ` und dem `for`-Attribut, das auf `"grade"` gesetzt ist, hinzuzufügen.

Fügen Sie dann unter Ihrem `label`-Element ein `input`-Element mit dem `type`-Attribut hinzu, das auf `"number"` gesetzt ist. Die `name`- und `id`-Attribute sollten auf `"grade"` gesetzt sein. Das `placeholder`-Attribut sollte auf `"E.g., 4"` gesetzt sein. Schließlich sollte Ihre Nummer `input` erforderlich sein.

# --hints--

Sie sollten insgesamt zwei `label`-Elemente auf der Seite haben.

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

Ihr zweites `label`-Element sollte den Text `Student Grade: ` haben. Vergessen Sie nicht das Leerzeichen nach dem Doppelpunkt.

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

Ihr zweites `label`-Element sollte ein `for`-Attribut haben.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

Das `label`-Attribut Ihres zweiten `for`-Elements sollte den Wert `"grade"` haben.

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

Sie sollten zwei `input`-Elemente auf der Seite haben.

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

Ihr zweites `input`-Element sollte unter Ihrem zweiten `label`-Element sein.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Ihr zweites `input`-Element sollte ein `type`-Attribut haben, das auf `"number"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

Ihr zweites `input`-Element sollte ein `name`-Attribut haben, das auf `"grade"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

Ihr zweites `input`-Element sollte ein `id`-Attribut haben, das auf `"grade"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

Ihr zweites `input`-Element sollte ein `placeholder`-Attribut haben, das auf `"E.g., 4"` gesetzt ist.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

Ihr zweites `input`-Element sollte ein `required`-Attribut haben.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
