---
id: 68eb058081504963447bfa60
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Als Nächstes sollten Sie den Hauptcontainer gestalten. Wählen Sie das Element mit der Klasse `container` aus und setzen Sie seine `background-color` auf den Hex-Code `#ffffff1a`.

Ein Hex-Code ist eine sechsstellige Kombination aus Zahlen und Buchstaben, die in HTML und CSS verwendet wird, um Farben darzustellen. Er beginnt mit einem `#`, gefolgt von drei Paaren:

- Die ersten zwei Ziffern stehen für Rot
- Die nächsten zwei für Grün
- Die letzten zwei für Blau

Zum Beispiel ist `#ffffff` reines Weiß, weil es den Maximalwert für Rot, Grün und Blau hat.

Sie können auch zwei zusätzliche Ziffern am Ende hinzufügen, um die Deckkraft zu steuern (Alpha-Kanal genannt). In `#ffffff1a` macht der `1a` die weiße Farbe halbtransparent. Je niedriger der Alpha-Wert, desto transparenter erscheint die Farbe. Später lernen Sie noch mehr über Hex-Codes!

# --hints--

Sie sollten das Element mit der Klasse `container` ansprechen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

Setzen Sie die Eigenschaft `background-color` auf `#ffffff1a`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.container')
assert.strictEqual(style?.backgroundColor, 'rgba(255, 255, 255, 0.1)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```


```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

--fcc-editable-region--

--fcc-editable-region--
```
