---
id: 65821c1c40ccac44adf38e6b
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`expenses` という名前の空のリストを作成してください。これを使って各支出をストアします。

# --hints--

`expenses` という名前の新しい変数を宣言してください。

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

`expenses` 変数に空のリストを代入してください。

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
