---
id: 6582324b5d37fe4b622bc3c0
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

Lambda関数はPythonの短くて無名の関数で、シンプルで一度きりの処理に最適です。`lambda`キーワードで定義され、次の構文を使います。

```py
lambda x: expr
```

上の例では、`x`が`expr`の式で使われるパラメータを表し、通常の関数のパラメータと同じように機能します。`expr`はlambda関数が呼び出されたときに評価されて返される式です。

`test`という変数を作成し、`x`パラメータを受け取り`x * 2`を返すlambda関数を代入してください。

# --hints--

`test`という変数を宣言してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

`lambda x: x * 2`を`test`変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
