---
id: 658258c0e5fbe85c14c060cf
title: ステップ 43
challengeType: 20
dashedName: step-43
---

# --description--

すべての支出を一覧表示するには、`elif`文の後に`if`節を使うことができます。`elif`は追加の条件をチェックし、`if`文の後にのみ機能します。

ユーザーの選択が文字列`elif`と等しいかどうかを確認する`'2'`節を作成してください。`elif`節の中で、文字列`'\nAll Expenses:'`を出力してください。

# --hints--

`elif`文の後に、`choice`が`'2'`と等しいかどうかをチェックする`if`節を作成してください。

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1]
cond.is_equivalent("choice == '2'") or cond.is_equivalent("'2' == choice")
`)) })
```

`'\nAll Expenses:'`節の中で`elif`を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("print('\\\\nAll Expenses:')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)
        
--fcc-editable-region--
```
