---
id: 65824111a09164518320088d
title: ステップ 30
challengeType: 20
dashedName: step-30
---

# --description--

`filter()` 関数を使うと、リストのようなイテラブルから関数の出力に基づいてアイテムを選択できます。

```py
filter(my_function, my_list)
```

`filter()` は最初の引数に関数を、2番目の引数にイテラブルを取ります。これはイテレーターを返します。イテレーターとは、リストのようなコレクションの要素を繰り返すことを可能にする特別なオブジェクトです。

上の例の結果は、`my_list` が `my_function` を返す `True` の要素を含むイテレーターです。

`filter_expenses_by_category` 関数の中で、前のステップで書いた `filter()` 関数を最初の引数に、`lambda` リストを2番目の引数にして `expenses` を呼び出してください。

# --hints--

`filter()` は `filter_expenses_by_category` 関数の中で呼び出す必要があります。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

`lambda expense: expense['category'] == category` は `filter()` 呼び出しの最初の引数として渡してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

`expenses` は `filter()` 呼び出しの2番目の引数として渡してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
