---
id: 65823634a308cb4d6bcc5fa6
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

ラムダ関数は、リストのようなオブジェクトのコレクション内の各要素に対して指定された関数を実行する`map()`関数と組み合わせると非常に役立ちます。

```py
map(lambda x: x * 2, [1, 2, 3])
```

実行する関数は最初の引数として渡され、イテラブルは2番目の引数として渡されます。

上記の例の結果は`[2, 4, 6]`で、`map()`に渡されたリストの各アイテムがラムダ関数の動作によって2倍になっています。

`print()`の呼び出しを修正して、`map()`を`test`を最初の引数に、`[2, 3, 5, 8]`を2番目の引数にして呼び出した結果を出力してください。まだ読みやすい出力は表示されません。

# --hints--

コード内に`print(map(test, [2, 3, 5, 8]))`があるはずです。

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
