---
id: 6582267ef39f5047411c1743
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

次に、`print_expenses` という名前の関数を定義し、1つのパラメータ `expenses` を受け取るようにしてください。この関数は後でリスト内の各支出を表示するために使います。

新しい関数の本体には `pass` 文を記述してください。

# --hints--

`print_expenses` という名前の新しい関数を定義してください。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

`print_expenses` 関数は1つのパラメータ `expenses` を受け取る必要があります。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
