---
id: 65824561f3478e5371a33ae5
title: ステップ 32
challengeType: 20
dashedName: step-32
---

# --description--

次のステップは、インタラクティブな支出トラッカーのプログラムのエントリーポイントとなるメイン関数を定義することです。

パラメータなしで `main` という名前の関数を定義してください。関数の本体には、このプロジェクトの最初に作成した `expenses` リストを入れてください。このリストを使って支出記録を保存します。

# --hints--

パラメータなしで `main()` という名前の関数を定義してください。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

`expenses` リストを `main()` 関数の本体内に移動してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
