---
id: 65823cfc74aa564ffc460489
title: ステップ 26
challengeType: 20
dashedName: step-26
---

# --description--

今、`map()` を呼び出して、`lambda` 関数を最初の引数として、`expenses` リストを2番目の引数として渡してください。

# --hints--

`map()` 関数は `total_expenses` 関数の中で呼び出す必要があります。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("total_expenses").find_calls("map")) == 1`)) })
```

`lambda` 関数を `map()` 呼び出しの最初の引数として渡してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[0].is_equivalent("lambda expense: expense['amount']")`)) })
```

`expenses` リストを `map()` 呼び出しの2番目の引数として渡してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
