---
id: 6582601b2987045e8e7da994
title: ステップ 48
challengeType: 20
dashedName: step-48
---

# --description--

`print()` を呼び出した後、経費をフィルターしてフィルター済みのリストを出力する必要があります。`expenses_from_category` という変数を宣言し、`filter_expenses_by_category` と `expenses` を引数として渡して `category` を呼び出した結果を代入してください。

# --hints--

`expenses_from_category` 節の中で `elif` という名前の変数を宣言してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[3].has_variable("expenses_from_category")`)) })
```

`filter_expenses_by_category()` を呼び出し、`expenses` と `category` を引数として渡し、その結果を `expenses_from_category` 変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[3].find_variable("expenses_from_category").is_equivalent("expenses_from_category = filter_expenses_by_category(expenses, category)")`)) })
```

`expenses_from_category` の呼び出しの後に `print()` 変数が来るようにしてください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[3].is_equivalent("category = input('Enter category to filter: ')\\nprint(f'\\\\nExpenses for {category}:')\\nexpenses_from_category = filter_expenses_by_category(expenses, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)

        elif choice == '2':
            print('\nAll Expenses:')
            print_expenses(expenses)

        elif choice == '3':
            print('\nTotal Expenses: ', total_expenses(expenses))
--fcc-editable-region--
        elif choice == '4':
            category = input('Enter category to filter: ')
            print(f'\nExpenses for {category}:')
            
--fcc-editable-region--
```
