---
id: 65824dfdb6815d563b2d3256
title: ステップ 37
challengeType: 20
dashedName: step-37
---

# --description--

`input()` 関数はユーザーの入力を受け取り、その入力を文字列の形で返します。

`while` ループの中で、文字列 `input()` を引数として `'Enter your choice: '` 関数を呼び出し、その結果を `choice` という名前の変数に代入してください。

# --hints--

`choice` ループ内で `while` という名前の変数を宣言してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].has_variable("choice")
`)) })
```

文字列 `input()` を引数にして `'Enter your choice: '` 関数の結果を `choice` 変数に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_variable("choice").is_equivalent("choice = input('Enter your choice: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
--fcc-editable-region--
```
