---
id: 65821fcc010c3245718f2a06
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

`expenses` パラメータは `add_expense` 関数の費用のリストになります。リストの末尾にアイテムを追加できるようにしたいです。そのために `.append()` リストのメソッドを使います。

`.append()` リストに対して `expenses` メソッドを呼び出すコードを追加してください。今のところ `.append()` に引数は渡さないでください。

# --hints--

関数内に `expenses.append()` があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

関数内に `pass` があってはいけません。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
