---
id: 65822bd82d708c4895080c35
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

Pythonでは、文字列を定義するのに使ったのと同じ種類の引用符は文字列の中で使えないことが重要です。例えば、文字列 `'I'm a string!'` は有効ではありません。その文字列の中でシングルクォートを使いたい場合は、次のどちらかを行う必要があります。

- バックスラッシュを前に付けてエスケープする：`'I\'m a string!'`
- または文字列を定義するのに二重引用符を使う：`"I'm a string!"`（こちらがお勧めです）。

辞書の値にはキーを使ってアクセスできます。ブラケット記法を使い、大括弧の間にキーを入れる必要があります。

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

現在、f文字列の中で`expense`辞書を補間しています。f文字列の式を修正して、`'amount'`辞書の`'category'`キーと`expense`キーの値にアクセスしてください。

# --hints--

`f'Amount: {expense["amount"]}, Category: {expense["category"]}'`を`print()`呼び出しに渡す必要があります。シングルクォートで囲まれたf文字列の中では二重引用符を使い、その逆も覚えておいてください。

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
