---
id: 65823e87c9741750a22085a7
title: ステップ 28
challengeType: 20
dashedName: step-28
---

# --description--

次に、`filter_expenses_by_category` という名前の関数を定義してください。この関数は `expenses` と `category` の2つのパラメータを受け取ります。関数本体は `pass` を使って記述してください。

# --hints--

`filter_expenses_by_category` という名前の関数を定義する必要があります。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

`filter_expenses_by_category` 関数は `expenses` と `category` をパラメータとして受け取るべきです。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
