---
id: 658238f7604f154ea9a23e1e
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

`sum()` 関数は、引数として渡されたイテラブル内のアイテムの合計を返します。`sum()` を `map()` と `lambda` 関数と一緒に使って、支出の合計金額を求めます。

まずは簡単なテストをして、現在の `print()` 呼び出しを修正し、`list()` 呼び出しを `sum()` 関数の呼び出しに置き換え、その引数に現在の `map()` 呼び出しを渡してください。

# --hints--

`sum(map(test, [2, 3, 5, 8]))` を出力してください。

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
