---
id: 65822308ded1ba4632f66aa3
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

辞書はPythonのもう一つの組み込みデータ型です。辞書は*キー*と*値*のペアの形でデータを集めたものです。辞書は波括弧（`{}`）で定義され、キーと値のペアはカンマで区切られています。各キーの後にはコロン（`:`）が続き、その後に値が続きます。

```py
{'amount': 50.0, 'category': 'Food'}
```

上の例では、`'amount'` と `'category'` がキーで、`50.0` と `'Food'` がそれぞれ対応する値です。

キーが `'amount'` で値が `amount` パラメータの辞書を作成し、新しい辞書を `.append()` の呼び出しに渡してください。

# --hints--

`.append()` メソッドに辞書を引数として渡す必要があります。

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

`{'amount': amount}` メソッドに `.append()` を引数として渡す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
