---
id: 65823040bdb32949d4c1a96f
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

支出の合計金額を計算する関数が必要です。

`total_expenses` という名前の関数を定義し、`expenses` というパラメータを1つ受け取るようにしてください。関数の本体はとりあえず `pass` 文で埋めてください。

# --hints--

`total_expenses` という名前の関数を定義してください。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

`total_expenses` 関数は1つの `expenses` パラメータを受け取る必要があります。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
