---
id: 65825cda2668995d5168e37c
title: ステップ 46
challengeType: 20
dashedName: step-46
---

# --description--

`elif` かどうかをチェックする別の `choice == '4'` 節を作成してください。新しい `elif` の中で、`category` という変数を作成し、経費のカテゴリをフィルターするために `input('Enter category to filter: ')` を代入してください。

# --hints--

`elif` かどうかをチェックする `choice == '4'` 節を作成する必要があります。

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[3]
cond.is_equivalent("choice == '4'") or cond.is_equivalent("'4' == choice")
`)) })
```

新しい `category` 節の中で `elif` という名前の変数を宣言する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[3].has_variable("category")`)) })
```

`input('Enter category to filter: ')` 変数に `category` を代入する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[3].find_variable("category").is_equivalent("category = input('Enter category to filter: ')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
       
        choice = input('Enter your choice: ')

        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)

        elif choice == '2':
            print('\nAll Expenses:')
            print_expenses(expenses)
--fcc-editable-region--           
        elif choice == '3':
            print('\nTotal Expenses: ', total_expenses(expenses))
        
--fcc-editable-region--
```
