---
id: 65822934c8a0454803e3ef12
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

次に、各支出の詳細を表示します。

`for` ループの中で、`pass` を `print()` 呼び出しに置き換え、次の f-string を渡してください：`f'Amount: {expense}, Category: {expense}'`。

# --hints--

`f'Amount: {expense}, Category: {expense}'` ループの中で `for` を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

ループ本体の中に `pass` を含めないでください。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
