---
id: 65823b0167fd894f4a7ea60a
title: ステップ 24
challengeType: 20
dashedName: step-24
---

# --description--

次に、同じロジックを`total_expenses`関数内に実装します。

まずは、`test`関数と`print()`の呼び出しを両方削除してください。

# --hints--

`test`変数と関連する`print()`の呼び出しを削除してください。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_variable("test")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(sum(map(test, [2, 3, 5, 8])))
--fcc-editable-region--

expenses = []
```
