---
id: 65823dde36392f505a39f7c7
title: ステップ 27
challengeType: 20
dashedName: step-27
---

# --description--

最後に、`map()` の呼び出しを `sum()` 関数に渡して合計支出を取得し、その結果を返してください。

# --hints--

現在の `sum()` の呼び出しを引数として `map()` 関数を呼び出し、`total_expenses` 関数から結果を返す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").has_return("sum(map(lambda expense: expense['amount'], expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    map(lambda expense: expense['amount'], expenses)
--fcc-editable-region--

expenses = []
```
