---
id: 65822773f4c60a479ec89f66
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

`print_expenses` 関数の中で、`for` リストの各アイテムを反復する `expenses` ループを作成してください。ループ変数として `expense` を使い、`pass` をループ本体の中に移動させてください。

# --hints--

`for` 関数内で `expenses` を反復する `print_expenses` ループを作成する必要があります。ループ本体の中で `pass` を使うことを忘れないでください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

ループ変数として `expense` を使うべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

`pass` はループ本体の中だけにあるべきです。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
