---
id: 6582575b8089f85b8b92d7c8
title: ステップ 42
challengeType: 20
dashedName: step-42
---

# --description--

支出の詳細がわかったら、新しい支出を`add_expense`リストに追加するために`expenses`関数を呼び出す必要があります。

`amount`を使って`category`と`input()`を取得した後、`add_expense`関数を呼び出し、3つの引数`expenses`、`amount`、`category`を渡してください。

- `expenses`は、このプロジェクトのメイン関数で先に作成した空のリストです。
- `amount`は支出の金額です。
- `category`は支出のカテゴリです。

# --hints--

`add_expense()`文の中で`if`関数を呼び出す必要があります。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

`add_expense()`の呼び出しには3つの引数、`expenses`、`amount`、`category`を渡してください。順番が重要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

`add_expense()`の呼び出しは、`amount`と`category`の代入の後に行う必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            
--fcc-editable-region--
```
