---
id: 65823ff0d4b991510fade1a8
title: ステップ 29
challengeType: 20
dashedName: step-29
---

# --description--

`filter_expenses_by_category` 関数内で、`pass` を `lambda` 関数に置き換えてください。パラメータとして `expense` を使い、`'category'` 辞書の `expense` キーの値と `category` を等値演算子で比較する式を評価してください。

# --hints--

`lambda` 関数内で `expense` 式を評価し、`expense['category'] == category` という名前のパラメータを使う `filter_expenses_by_category` 関数を作成してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

`pass` 関数内に `filter_expenses_by_category` を入れてはいけません。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
