---
id: 658254db6e12485a48130f57
title: ステップ 40
challengeType: 20
dashedName: step-40
---

# --description--

計算を行う前に、支出額を変換する必要があります。`float()` 関数は文字列または整数を引数に取り、浮動小数点数を返します。

`input('Enter amount: ')` を `float()` 関数に渡してください。

# --hints--

`input('Enter amount: ')` の代入で `float()` 関数に `amount` を渡す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = float(input('Enter amount: '))")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
        
        if choice == '1':
            amount = input('Enter amount: ')
--fcc-editable-region--
```
