---
id: 658246c28575d653d1f89b59
title: ステップ 33
challengeType: 20
dashedName: step-33
---

# --description--

`while` ループは、指定された条件が `True` の間、コードの一部を繰り返し実行する別の種類のループです。条件が `False` になるとループは終了します。

```py
while condition:
    <code>
```

`expenses` リストの下に、`while` ループを作成してください。条件には `True` を使い、ループの中で文字列 `'\nExpense Tracker'` を出力してプログラムのタイトルを表示してください。

# --hints--

`while` 関数内で、条件に `True` を使った `main()` ループを作成してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

`'\nExpense Tracker'` ループ内で `while` を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
