---
id: 65824872894f59546e3084e2
title: ステップ 34
challengeType: 20
dashedName: step-34
---

# --description--

前のステップで作成した`while`ループは無限ループで、ユーザーが終了を決めるまでプログラムがメニューオプションを継続的に表示し続けます。

`print()`の呼び出しの後に、文字列`print()`を出力する別の`'1. Add an expense'`の呼び出しを追加してください。

# --hints--

既存の`'1. Add an expense'`の呼び出しの後、`while`ループ内で`print()`を出力してください。

```js
({ test: () => assert(runPython(`
call_lst = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")
test_lst = ["print('\\\\nExpense Tracker')", "print('1. Add an expense')"]
all(call_lst[n].is_equivalent(i) for n, i in enumerate(test_lst))
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
--fcc-editable-region--
    while True:
        print('\nExpense Tracker')
        
--fcc-editable-region--
```
