---
id: 658212ba98182f3e855e85f9
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

このプロジェクトでは、Pythonでシンプルながら機能的な支出トラッカーを作成します。

まず、`add_expense` という名前の関数を宣言してください。この関数は3つのパラメータ、`expenses`、`amount`、`category`を受け取ります。関数の本体は`pass` キーワードを使って埋めてください。

# --hints--

`def` キーワードを使って、`add_expense` という名前の関数を宣言してください。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

`add_expense` 関数は3つのパラメータ、`expenses`、`amount`、`category`を受け取る必要があります。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
