---
id: 65824a1b16631c54fa524154
title: ステップ 35
challengeType: 20
dashedName: step-35
---

# --description--

次に、もう一つ`print()`を呼び出し、文字列`'2. List all expenses'`を渡してください。

# --hints--

既存の`'2. List all expenses'`の呼び出しの後、`while`ループ内で`print()`を出力する必要があります。

```js
({ test: () => assert(runPython(`
call_lst = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")
test_lst = ["print('\\\\nExpense Tracker')", "print('1. Add an expense')", "print('2. List all expenses')"]
all(call_lst[n].is_equivalent(i) for n, i in enumerate(test_lst))
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
--fcc-editable-region--
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        
--fcc-editable-region--
```
