---
id: 658225d560369446ece5518b
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

`expense` リストに追加している辞書に、もう1つのキーと値のペアを追加してください。キーには文字列 `'category'` を使い、値にはパラメータ `category` を使います。

# --hints--

`.append()` メソッドに引数として渡している辞書に、2つ目のキーと値のペアを追加してください。

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

`{'amount': amount, 'category': category}` メソッドに引数として `.append()` を渡す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
