---
id: 6582507654b3ed5712341382
title: ステップ 38
challengeType: 20
dashedName: step-38
---

# --description--

条件文を使ってユーザーの選択を確認します。選択が `'1'` の場合、ユーザーは支出を追加したいことを意味します。

まだ `while` ループの中で、`choice` 変数の下に、`if` 文を書いて `choice` が文字列 `'1'` と等しいかどうかを確認してください。もし真なら、新しい支出を追加するための開始点になります。

`if` 文の本体内で、変数 `amount` を宣言し、空の `input()` 呼び出しを代入してください。

# --hints--

`if` ループ内で `choice` が `'1'` と等しいかどうかを確認する `while` 文を作成してください。

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
cond.is_equivalent("choice == '1'") or cond.is_equivalent("'1' == choice")
`)) })
```

`amount` 文の本体内で `if` という名前の変数を宣言してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("amount")
`)) })
```

`input()` 変数に空の `amount` 呼び出しを代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input()")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

--fcc-editable-region--
```
