---
id: 6582352cbdbcde4d0c4f7b0b
title: ステップ 20
challengeType: 20
dashedName: step-20
---

# --description--

ラムダ関数を呼び出すには、変数名の後に丸括弧を付けて通常の関数構文を使います。

`test` のラムダ関数を呼び出し、引数として `3` を渡してください。その後、結果を出力してください。

# --hints--

コード内に `print(test(3))` があるはずです。

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
