---
id: 66314f2add0a8b57da5d6d91
title: ステップ 22
challengeType: 20
dashedName: step-22
---

# --description--

コンソールに`<map object at 0xd273a8>`のようなものが出力されているはずです。これは`map()`が返すmapオブジェクトの文字列表現です。

読みやすい出力を得るには、mapオブジェクトをリストに変換する必要があります。`map()`の呼び出しを`list()`関数の引数として渡して実行してください。

# --hints--

`map()`の呼び出しを`list()`関数の引数として渡す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
