---
id: 66694fc4bba24f33ca01fa5b
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

Python にはいくつかのリストメソッドがあります。リストの項目を追加したり削除したりするメソッドなどです。

リストの末尾に項目を追加するには `append()` メソッドを使います。例えば：

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` は今 `[4, 5, 6, 7]` です。

`append()` メソッドを使って `3` を `my_list` に追加してみてください。その後リストを出力してください。

# --hints--

`my_list` の最後の要素は `3` であるべきです。

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` の長さは 3 であるべきです。

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

リストの宣言は変更しないでください。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

リストを出力してください。

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

`append()` を使った後にリストを出力してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
