---
id: 666951d2e4b27e3a3c72ca5f
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

リストの単一の要素にアクセスして、その値を取得することもできます。

リストは文字列と同様にゼロから始まるインデックスです。つまり、最初の要素はインデックス`0`にあり、2番目の要素はインデックス`1`にあります。

要素にアクセスするにはブラケット記法を使います。例えば、`example_list[1]`は`1`の2番目の要素、つまりインデックス`example_list`の要素にアクセスしています。

`my_list`の最初の要素を出力してください。

# --hints--

リストの最初の要素は`print(my_list[0])`で出力してください。

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
