---
id: 66694eaaeb7d253049fa7883
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

リストにアイテムを追加するには、大括弧の間に値を書き、カンマで区切ります。

リストの中に数字の `1` と `2` を追加してください。

# --hints--

`my_list` の最初の要素は `1` であるべきです。

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

`my_list` の2番目の要素は `2` であるべきです。

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list` には2つの要素が入っているべきです。

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

`[1, 2]` を `my_list` に直接代入してください。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
