---
id: 6669539c1379793f9cb8917c
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

Python のリストはミュータブルで、リストの要素の値を変更できます。ブラケット記法を使って要素の値を変更できます。

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

これにより `example_list` の値が `[4, 'oh', 6, 7]` になります。

`my_list` の最初の要素を `0` に変更し、その後リストを出力して値を確認してください。

# --hints--

リストの宣言は変更しないでください。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

`my_list` の最初の要素は 0 であるべきです。

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

最初の要素は `my_list[0] = 0` を使って変更してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

リストを出力してください。

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

最初の要素を変更した後にリストを出力してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
