---
id: 66694bfa435cef241f6f9a68
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

プロジェクトを始める前に、リストについて学びます。

数値や文字列とは異なり、<dfn>リスト</dfn>は他の値を格納するコンテナとして機能するデータ型です。

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

リストはすべての値を囲む大括弧と、値の間のカンマで特徴づけられます。例えば `["A", "happy", "list"]` のようになります。リストに値が含まれていない場合は空のリストです：`[]`。

リストは異なるデータ型を含むことができます：`[1, "Up", ["Down", "Twice"]]`。これはすべての可能なデータ型を含みます。また、リストの中に別のリストを含めることもできます！

`my_list` という変数を作成し、空のリストを代入してください。

# --hints--

`my_list` という変数があるはずです。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

変数 `my_list` には空のリストの値が代入されている必要があります。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
