---
id: 666aae6a5d30a71f1fd7749f
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`insert` メソッドはリストの任意の位置に要素を追加できます。最初の引数は要素を追加する位置で、2番目の引数は追加する要素です。たとえば、`example_list` の3番目の位置に新しい要素を追加する方法は次のとおりです。

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

`insert()` を使って、`1` に `my_list` を適切な位置に追加し、昇順になるようにしてからリストを出力してください。

# --hints--

`1` を `my_list` のインデックス `1` に `insert()` を使って追加してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

`my_list` は4つの要素を持つはずです。

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list` は `[0, 1, 2, 3]` と等しくなるはずです。

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

新しい要素を追加した後にリストを出力してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
