---
id: 666ab9013cd5563f2b39947a
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

リストについて学んだので、レビュー用に書いたコードはすべて削除して、プロジェクトを始めましょう。

# --hints--

すべて削除してください。

```js
({
    test: () => runPython(`
    assert not _Node(_code)
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)

my_list.pop()
print(my_list)

--fcc-editable-region--
```
