---
id: 666ab82c874add3c95145213
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`pop()` メソッドはリストから要素を削除するために使えます。デフォルトでは、リストの最後の要素を削除します。メソッドにインデックスを引数として渡すと、そのインデックスの要素を削除します。

```py
fruits_list = ["cherry", "lemon", "tomato", "apple", "orange"]

fruits_list.pop(2)

print(fruits_list) # ["cherry", "lemon", "apple", "orange"]
```

この場合、`fruits_list.pop(2)` はリストのインデックス `2` の要素を削除します。

`pop()` を使って `my_list` から最後の要素を削除し、その後 `my_list` を出力してください。

# --hints--

`pop()` を使うべきです。

```js
({
    test: () => runPython(`
    assert len(_Node(_code).find_calls('pop')) > 0
    `)
})
```

`my_list` は `[0, 1, 2]` であるべきです。

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2]
    `)
})
```

`my_list` を `pop()` を使った後に出力してください。

```js
({
    test: () => runPython(`
    code = _Node(_code)
    assert(
        any(
            code.is_ordered(pop, 'print(my_list)') for pop in ['my_list.pop()', 'my_list.pop(3)']
        )
    )
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)


--fcc-editable-region--
```
