---
id: 5d822fd413a79914d39e98d4
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Dé a las partes de su edificio las propiedades de `ancho` y `altura` con estos valores: `70%` y `10%` para `.bb1a`, `80%` y `10%` para `.bb1b`, `90%` y `10%` para `.bb1c`, y `100%` y `70%` para `.bb1d`. Recuerde que estos porcentajes son relativos al elemento padre y tenga en cuenta que las alturas sumarán un 100% - llenando verticalmente el contenedor.

# --hints--

Debe utilizar un selector de clase para estilizar `.bb1a`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

Debe darle a `.bb1a` un `ancho` de `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

Debe darle a `.bb1a` una `altura` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

Debe utilizar un selector de clase para estilizar `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

Debe darle a `.bb1b` un `ancho` de `80%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

Debe darle a `.bb1b` una `altura` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

Debe utilizar un selector de clase para estilizar `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

Debe darle a `.bb1c` un `ancho` de `90%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

Debe darle a `.bb1c` una `altura` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

Debe utilizar un selector de clase para estilizar `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

Debe darle a `.bb1d` un `ancho` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

Debe darle a `.bb1d` una `altura` de `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

