---
id: 694c84106e9023f289c03b84
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Ваш список починається з одного командира-астронавта. Створіть об’єкт з назвою `firstAstronaut` з такими властивостями:

| Ключ | Значення |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

Ви повинні створити змінну `firstAstronaut`.

```js
assert.isDefined(firstAstronaut);
```

Ваш `firstAstronaut` має бути об’єктом.

```js
assert.isObject(firstAstronaut);
```

Ваш об’єкт `firstAstronaut` має містити правильні значення властивостей.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
