---
id: 694c84330cbe604821b9deef
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Центр управління потребує методу для додавання нових учасників до списку. Створіть порожню функцію з іменем `addCrewMember`, яка приймає два параметри, `crew` і `astronaut`.

# --hints--

Ви повинні створити функцію з іменем `addCrewMember`.

```js
assert.isFunction(addCrewMember);
```

Ваша функція `addCrewMember` має мати параметри `crew` і `astronaut`.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
