---
id: 694c9c186c089f7bdff0f3c6
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Загалом, хорошою практикою є перевірка вхідних даних всередині функції. Використовуйте `for` цикл, щоб ітерувати по масиву `crew` і перевірити, чи має якийсь елемент уже той самий `id`, що й вхідний `astronaut`. Якщо знайдено дубліката, виведіть повідомлення про помилку `console.log("Duplicate ID: " + astronaut.id)` і викличте `return`, щоб достроково завершити функцію.

Ось приклад цієї загальновживаної техніки:

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

Ваша функція має використовувати `for` цикл для перевірки існуючих членів `crew` (будь-яка дійсна форма `for (...)` підходить).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

Всередині `for` циклу слід перевірити наявність дубліката `id`. Якщо він знайдений, виведіть `"Duplicate ID: " + astronaut.id"` у консоль і викличте `return`, щоб достроково завершити функцію.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

При наданні дійсних вхідних даних ваша функція `addCrewMember` не повинна викликати повідомлення про помилку `console.log()` або дострокове завершення `return`.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
