---
id: 694c40d86fe6a78d20b6559c
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви будете використовувати JavaScript, щоб створити список астронавтів для космічної місії.

Для початку створіть порожній масив для свого списку з іменем `squad`.

# --hints--

Ви маєте створити змінну `squad`.

```js
assert.isDefined(squad);
```

Ваша змінна `squad` має бути масивом.

```js
assert.isArray(squad);
```

Ваш масив `squad` має бути порожнім.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
