---
id: 694c84330cbe604821b9deef
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Центр управління потребує методу для додавання нових учасників до списку. Створіть порожню функцію з іменем `addCrewMember`, яка приймає два параметри, `crew` і `astronaut`.

# --hints--

Ви повинні створити функцію з іменем `addCrewMember`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addCrewMember);
```

Ваша функція `addCrewMember` має мати параметри `crew` і `astronaut`.

```js
const explorer = await __helpers.Explorer(code);
const { addCrewMember } = explorer.allFunctions;
assert.lengthOf(addCrewMember?.parameters, 2);
assert.equal(addCrewMember?.parameters?.[0]?.toString(), "crew");
assert.equal(addCrewMember?.parameters?.[1]?.toString(), "astronaut");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
