---
id: 694d8ddd106340237efb0c80
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Метод `splice()` може змінювати масиви, додаючи або видаляючи елементи в будь-якому місці, включно з серединою. Оскільки `splice()` повертає масив, що містить видалені елементи, ви можете використати його, щоб поміняти місцями два елементи в масиві без зміни оригіналу. Це можна зробити в один рядок, використовуючи таку техніку:

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

Ось приклад:

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

Техніка, застосована в третій лінії, працює так:

- `splice(3, 1, copyArray[1])` видаляє елемент за індексом `3` (`55`)

- Вона вставляє елемент з `copyArray[1]` (`97`) на індекс `3`

- `splice()` повертає `[55]`, масив, що містить видалений елемент

- `[0]` витягує `55` з цього масиву

- Це значення (`55`) призначається назад `copyArray[1]`, завершуючи обмін

Використайте цю техніку з масивом `updatedCrew`, щоб поміняти місцями астронавтів на `fromIndex` та `toIndex`.

# --hints--

Ви повинні виконати обмін, викликавши метод `splice()`.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

Ви повинні викликати `splice()` з масивом `updatedCrew`, щоб виконати обмін. Використайте техніку в один рядок.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

Ваша функція `swapCrewMembers` не повинна змінювати оригінальний вхідний масив `crew`.

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
