---
id: 6a3e2b9b7eeca084dd4fb857
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Bem-vindo à cozinha! Assim como um chef depende de receitas para preparar uma refeição, você vai depender de funções para construir software. Nesta oficina, você vai aprender a escrever funções Python reutilizáveis e dominar o escopo, entendendo exatamente quais ingredientes (variáveis) estão disponíveis no seu código a qualquer momento.

Para começar, você deve definir alguns dados básicos para o seu inventário da cozinha. Crie três variáveis: `available_eggs`, `available_flour` e `available_sugar`, e atribua a elas os valores `1`, `2` e `3`, respectivamente.

# --hints--

Você deve definir uma variável chamada `available_eggs` e atribuir a ela o valor `1`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

Você deve definir uma variável chamada `available_flour` e atribuir a ela o valor `2`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

Você deve definir uma variável chamada `available_sugar` e atribuir a ela o valor `3`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
