---
id: 6a3ea8a55317ba7a08e21d5b
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

Agora você finalmente pode responder à pergunta anterior: por que o número global de ovos permaneceu em `1`? Quando uma função retorna um valor, ele não sobrescreve automaticamente suas variáveis existentes. Você deve salvar explicitamente esse novo valor!

Você está usando a variável `available_eggs` dentro da função `make_fried_egg`, mas não é a mesma variável definida no topo do seu código. A que está dentro da função é uma variável local que existe apenas enquanto a função é executada. Quando a função termina, a variável local é descartada e a variável global `available_eggs` permanece inalterada.

Para atualizar seu estoque real na cozinha, sobrescreva a variável global `available_eggs` com o novo número de ovos retornado pela sua função `make_fried_egg`.

No final do seu código, faça a chamada da função `make_fried_egg`, passando `available_eggs` como argumento, e atribua o resultado de volta à variável global `available_eggs`.

# --hints--

Você deve reatribuir a variável `available_eggs` ao resultado de `make_fried_egg(available_eggs)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
