---
id: 6a3eaa2caac22f67811fe40f
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Chegou a hora do momento da verdade! Você atualizou com sucesso a variável global `available_eggs` com o novo valor retornado pela sua função `make_fried_egg`.

Confira seu inventário da cozinha mais uma vez para verificar se o número de ovos realmente diminuiu.

No final do seu código, chame a função `check_kitchen_stock()` novamente. Quando você executar o código, deverá ver que a cozinha agora tem `0` ovos!

Parabéns por completar com sucesso a oficina de rastreamento do inventário da cozinha!

# --hints--

Você deve chamar a função `check_kitchen_stock()` uma segunda vez no final do seu código.

```js
({
  test: () => runPython(`
node = _Node(_code)
assert len(node.find_calls("check_kitchen_stock")) == 2
assert node.is_ordered("available_eggs = make_fried_egg(available_eggs)", "check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

available_eggs = make_fried_egg(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

available_eggs = make_fried_egg(available_eggs)

check_kitchen_stock()
```
