---
id: 6a3e6ad6b9554e31dcc30cd6
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Se você não tiver ovos suficientes, não deve tentar cozinhar! Você precisa de uma forma de parar a função de executar mais adiante.

Você pode usar a palavra-chave `return` para isso. Quando uma função encontra uma declaração `return`, ela para de executar imediatamente e devolve o valor retornado para onde foi chamada.

Por exemplo:

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

Dentro do seu bloco `if`, abaixo da declaração de imprimir, use a palavra-chave `return` para retornar o valor de `available_eggs`.

# --hints--

Você deve retornar `available_eggs` dentro da sua declaração `if`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")
`)
})
```

Você deve retornar `available_eggs` depois da declaração de `print` dentro do seu bloco `if`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs").find_ifs()[0]
assert node.is_ordered("print('The kitchen does not have enough eggs.')", "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
