---
id: 6a3e5ccb99fd6053af9032a1
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Agora que você sabe o que tem na sua cozinha, pode usar esses itens para cozinhar. Para isso, crie uma nova função que aceite parâmetros. Pense nos parâmetros como marcadores de posição para os dados que você passa para a função quando a chama.

Aqui está uma função que aceita dois parâmetros, `a` e `b`, e imprime a soma deles no terminal.

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

Abaixo da sua função `check_kitchen_stock`, defina uma função vazia chamada `use_eggs` que aceite dois parâmetros: `available_eggs` e `eggs_to_use`, nessa ordem.

Note que se você der a um parâmetro o mesmo nome de uma variável global (como `available_eggs`), a função usará seu próprio parâmetro local `available_eggs` em vez da variável global. Isso é conhecido como name shadowing, já que a variável local "sombreia" a global.

Use a palavra-chave `pass` dentro do corpo por enquanto.

# --hints--

Você deve definir uma função chamada `use_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

Sua função `use_eggs` deve ter `available_eggs` e `eggs_to_use` como seus parâmetros.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
