---
id: 6a3e6c0eb5d050ce07781709
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Se a condição do `if` for avaliada como `False` (significando que você *tem* ovos suficientes), a função vai pular o bloco do `if` e continuar executando o restante do código.

Agora você precisa processar seus dados e retornar um novo valor calculado.

Fora do seu bloco `if` (no mesmo nível de indentação da palavra-chave `if`), use uma chamada `print()` e uma f-string para imprimir `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`.

Então, na linha seguinte, `return` o número restante de ovos subtraindo `eggs_to_use` de `available_eggs`.

# --hints--

Você deve ter uma declaração print que imprima `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'` na sua função `use_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

Você deve retornar o número restante de ovos subtraindo `eggs_to_use` de `available_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

Você deve retornar `available_eggs - eggs_to_use` depois da declaração `print`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
