---
id: 6a3e71a0a3a44260e675b157
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Para resolver o mistério de por que a contagem de ovos não mudou e realmente preparar o café da manhã, crie uma nova função. Pense nessa função como sua receita, uma que chama outras funções para realizar o trabalho.

Abaixo da sua declaração `print(available_eggs)`, defina uma função vazia chamada `make_fried_egg` que recebe `available_eggs` como seu único parâmetro.

Use a palavra-chave `pass` dentro do corpo por enquanto.

# --hints--

Você deve definir uma função chamada `make_fried_egg`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

Sua função `make_fried_egg` deve ter `available_eggs` como parâmetro.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
