---
id: 6a3e58891e1ff13d07f4ab48
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Python lê o bloco `def` e memoriza a função, mas não executa o código dentro dela imediatamente.

Para executar o código e ver sua saída, você precisa **fazer a chamada** (ou invocar) da função. Faça isso digitando o nome da função seguido de parênteses.

Por exemplo:

```py
my_function()
```

No final do seu código, fora da função, faça a chamada da sua função `check_kitchen_stock`.

# --hints--

Você deve fazer a chamada da função `check_kitchen_stock` no final do seu código. Certifique-se de que ela não esteja indentada, para que seja executada globalmente.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
