---
id: 6a3e524fc3394ef52ed5b797
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Funções têm seu próprio escopo local. Isso significa que variáveis criadas dentro de uma função são como ingredientes guardados na despensa daquela cozinha: só podem ser acessadas dentro daquela função específica.

No entanto, funções não são completamente isoladas. Elas podem ler variáveis globais definidas fora do seu escopo. Podem usar essas variáveis para fazer cálculos, tomar decisões ou criar novas variáveis locais.

Dentro da sua função `check_kitchen_stock`, substitua a palavra-chave `pass` por uma nova variável chamada `total_items` e atribua a soma de `available_eggs`, `available_flour` e `available_sugar`.

# --hints--

Você deve definir uma variável chamada `total_items` dentro da sua função `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

Sua variável `total_items` deve ser a soma de `available_eggs`, `available_flour` e `available_sugar`.

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
