---
id: 6a3e73eabddd614f91cc235a
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Agora você vai resolver o mistério de antes! Quando sua função `use_eggs` retornou o número restante de ovos, você não salvou esse novo valor em lugar nenhum. Ele simplesmente desapareceu no vazio!

Para realmente atualizar seu inventário, você precisa capturar esse valor retornado atribuindo-o a uma variável.

Dentro da sua função `make_fried_egg`, use uma declaração `if/else` para verificar a variável `has_enough_eggs`.

- Se `has_enough_eggs` for `True`, reatribua sua variável `available_eggs` ao resultado de `use_eggs(available_eggs, 1)` e imprima a mensagem `Made a fried egg. Yummy!`.

- Caso contrário, dentro de um bloco `else`, imprima a mensagem `Could not make a fried egg. Not enough eggs!`.

# --hints--

Você deve ter uma declaração `if` e usar `has_enough_eggs` como condição.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

Dentro do seu bloco `if`, você deve reatribuir `available_eggs` ao resultado de `use_eggs(available_eggs, 1)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

Você deve imprimir a mensagem `Made a fried egg. Yummy!` quando `has_enough_eggs` for `True`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

Você deve ter um bloco `else` que imprime a mensagem `Could not make a fried egg. Not enough eggs!` quando `has_enough_eggs` for `False`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
