---
id: 6a3e6fd08125dfa615cb431b
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Agora que sua função `use_eggs` está completa, você pode testá-la. Ao chamar uma função que requer **parâmetros**, você passa os valores específicos dentro dos parênteses. Esses valores são conhecidos como **argumentos**.

Aqui está um exemplo de como chamar uma função com argumentos:

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

No final do seu código, chame sua função `use_eggs` e passe dois argumentos: a variável global `available_eggs` e o número `1` (representando a quantidade de ovos que você quer usar).

# --hints--

Você deve chamar a função `use_eggs` e passar `available_eggs` e `1` como argumentos.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
