---
id: 6a3e4ae26f2b7e4b5b026cf8
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

É hora de criar uma função para verificar o que você tem na sua cozinha. Em Python, você define funções usando a palavra-chave `def`, seguida pelo nome da função, parênteses `()` e dois pontos `:`.

Aqui está um exemplo de como definir uma função simples:

```py
def my_function():
    pass
```

Defina uma função vazia chamada `check_kitchen_stock`. Use a palavra-chave `pass` dentro do corpo da função por enquanto para que o Python não gere um `SyntaxError`.

Lembre-se de que o Python depende da indentação para definir blocos de código. Então, certifique-se de indentar o corpo da função.

# --hints--

Você deve definir uma função chamada `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
