---
id: 6a3e70b94f67d9ee3ecd8486
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Você acabou de chamar `use_eggs(available_eggs, 1)` para usar um ovo. Isso significa que você deveria ter `0` ovos sobrando na sua cozinha agora, certo?

Você vai descobrir! No final do seu código, use a função `print()` para imprimir o valor da variável `available_eggs`.

Quando você executar o código, vai notar algo surpreendente: a saída ainda é `1`! Por que o número de ovos não mudou? Você vai explorar esse mistério nos próximos passos.

# --hints--

Você deve imprimir a variável `available_eggs` no final do seu código.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
