---
id: 6a3e5475c35dbc138d77c6f4
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora que você calculou seu total, deve exibir o inventário. Em Python, você pode usar a função `print()` junto com **f-strings** (literais de string formatadas) para imprimir facilmente texto e variáveis juntos.

Por exemplo:

```py
name = 'Alice'
print(f'Hello, {name}!') # Output: Hello, Alice!
```

Dentro da sua função `check_kitchen_stock`, abaixo da variável `total_items`, adicione exatamente essas quatro declarações de print:

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

Você deve ter a declaração `print(f'The kitchen has {total_items} total items:')` dentro da sua função.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

A declaração `print(f'The kitchen has {total_items} total items:')` deve ser a primeira chamada `print()` dentro da sua função.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

Você deve ter a declaração `print(f'- {available_eggs} eggs')` dentro da sua função.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

Você deve ter a declaração `print(f'- {available_flour} flour')` dentro da sua função.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

Você deve ter a declaração `print(f'- {available_sugar} sugar')` dentro da sua função.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
