---
id: 66f2836c459cfb16ae76f24f
title: بناء مولد هرم
challengeType: 26
dashedName: lab-pyramid-generator
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك دالة باسم `pyramid` تأخذ ثلاث معلمات.
1. المعلمة الأولى يجب أن تكون سلسلة نصية تمثل حرف النمط الذي ستكرره في الهرم.
1. المعلمة الثانية يجب أن تكون عددًا صحيحًا يمثل عدد الصفوف في الهرم.
1. المعلمة الثالثة يجب أن تكون قيمة منطقية.
1. يجب أن تُرجع دالة `pyramid` سلسلة نصية يُكرر فيها حرف النمط ويُرتب ليشكل هرمًا يكون رأسه متجهًا للأعلى عندما تكون المعلمة الثالثة `false`.
1. عندما تكون المعلمة الثالثة `true` يجب أن يكون رأس الهرم متجهًا للأسفل.
1. صف الرأس يجب أن يحتوي على حرف نمط واحد، وكل صف آخر يجب أن يحتوي على حرفين نمط أكثر من الصف السابق.
1. يجب أن يبدأ كل صف بعدد من الفراغات يكفي لوضع الحرف المركزي لكل صف في نفس العمود، ولا يجب أن يكون هناك فراغات في نهاية كل صف.
1. يجب أن يبدأ الهرم وينتهي بحرف سطر جديد.

على سبيل المثال، استدعاء `pyramid("o", 4, false)` يجب أن يعطي هذا الناتج:

```js

   o
  ooo
 ooooo
ooooooo

```

# --hints--

يجب أن يكون لديك دالة باسم `pyramid`.

```js
assert.isFunction(pyramid);
```

يجب أن تحتوي دالة `pyramid` على ثلاث معلمات.

```js
assert.lengthOf(pyramid, 3);
```

`pyramid("o", 4, false)` يجب أن تُرجع `"\n   o\n  ooo\n ooooo\nooooooo\n"`.

```js
assert.equal(pyramid("o", 4, false), "\n   o\n  ooo\n ooooo\nooooooo\n")
```

`pyramid("p", 5, true)` يجب أن تُرجع `"\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n"`.

```js
assert.equal(pyramid("p", 5, true), "\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function pyramid(char, count, isInverted) {
    const rows = []
    for (let i = 1; i <= count; i++) {
        if (isInverted) {
            rows.unshift(" ".repeat(count - i) + char.repeat(2 * i - 1))
        } else {
            rows.push(" ".repeat(count - i) + char.repeat(2 * i - 1))
        }        
    }
    return "\n" + rows.join("\n") + "\n";
}

console.log(pyramid("#", 10, false))

```
