---
id: 587d824c367417b2b2512c4d
title: 比較兩個元素的屬性
challengeType: 2
forumTopicId: 301588
dashedName: compare-the-properties-of-two-elements
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

# --instructions--

在 `tests/1_unit-tests.js` 中 `Comparisons` 測試套件下標籤為 `#8` 的測試中，將每個 `assert` 更改為 `assert.isAbove` 或 `assert.isAtMost`，以使測試通過（應該評估為 `true`）。請勿更改傳遞給 assert 的引數。

# --hints--

所有測試都應該通過。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該為第一個斷言選擇正確的方法 - `isAbove` 與 `isAtMost`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtMost',
  '5 is at most (<=) 5'
);
```

你應該為第二個斷言選擇正確的方法 - `isAbove` 與 `isAtMost`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'isAbove', '1 is greater than 0');
```

你應該為第三個斷言選擇正確的方法 - `isAbove` 與 `isAtMost`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'isAbove',
  'Math.PI = 3.14159265 is greater than 3'
);
```

你應該為第四個斷言選擇正確的方法 - `isAbove` 與 `isAtMost`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isAtMost',
  '1 - Math.random() is > 0 and <= 1. It is atMost 1 !'
);
```

