---
id: 587d824d367417b2b2512c51
title: 測試陣列是否包含某個項目
challengeType: 2
forumTopicId: 301603
dashedName: test-if-an-array-contains-an-item
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

# --instructions--

在 `Arrays` 測試組中，位於 `tests/1_unit-tests.js` 裡標籤為 `#12` 的測試中，將每個 `assert` 改為 `assert.include` 或 `assert.notInclude`，以使測試通過（應該評估為 `true`）。不要更改傳遞給 assert 的引數。

# --hints--

所有測試都應該通過。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該為第一個斷言選擇正確的方法 - `include` 與 `notInclude`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInclude',
  "It's summer in july..."
);
```

你應該為第二個斷言選擇正確的方法 - `include` 與 `notInclude`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'include',
  'JavaScript is a backend language !!'
);
```

