---
id: 587d824c367417b2b2512c4f
title: 測試一個值是否落在特定範圍內
challengeType: 2
forumTopicId: 301598
dashedName: test-if-a-value-falls-within-a-specific-range
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

```javascript
.approximately(actual, expected, delta, [message])
```

斷言 `actual` 與 `expected` 在 +/- `delta` 範圍內相等。

# --instructions--

在 `Comparisons` 測試套件中，位於 `tests/1_unit-tests.js` 中標籤為 `#10` 的測試內，將每個 `assert` 更改為 `assert.approximately`，以使測試通過（應該計算為 `true`）。

選擇最小的範圍（第 3 個參數）以使測試始終通過。它應該小於 1。

# --hints--

所有測試都應該通過。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該為第一個斷言選擇正確的範圍、區間(用於 STL 時) - `approximately(actual, expected, range)`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'approximately');
assert.equal(
  data.assertions[0].args[2],
  0.5,
  "weirdNumbers(0.5) is in the range (0.5, 1.5]. It's within 1 +/- 0.5"
);
```

你應該為第二個斷言選擇正確的範圍 - `approximately(actual, expected, range)`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'approximately');
assert.equal(
  data.assertions[1].args[2],
  0.8,
  "weirdNumbers(0.2) is in the range (0.2, 1.2]. It's within 1 +/- 0.8"
);
```

