---
id: 587d824e367417b2b2512c58
title: 使用 Chai-HTTP 對 API 端點執行功能測試
challengeType: 2
forumTopicId: 301593
dashedName: run-functional-tests-on-api-endpoints-using-chai-http
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

Mocha 允許你使用名為 `chai-http` 的外掛程式來測試非同步操作，例如對 API 端點的呼叫。

以下是使用 `chai-http` 為名為 `'GET /hello?name=[name] => "hello [name]"'` 的測試套件進行測試的範例：

```js
suite('GET /hello?name=[name] => "hello [name]"', function () {
  test('?name=John', function (done) {
    chai
      .request(server)
      .keepOpen()
      .get('/hello?name=John')
      .end(function (err, res) {
        assert.equal(res.status, 200, 'Response status should be 200');
        assert.equal(res.text, 'hello John', 'Response should be "hello John"');
        done();
      });
  });
});
```

測試會向伺服器發送帶有名稱作為 URL 查詢字串（`?name=John`）的 `GET` 請求。在 `end` 方法的回呼函式中，會接收到回應物件（`res`），並包含 `status` 屬性。

第一個 `assert.equal` 檢查狀態是否相等於 `200`。第二個 `assert.equal` 檢查回應字串（`res.text`）是否相等於 `"hello John"`。

另外，請注意測試回呼函式中的 `done` 參數。在測試結束時呼叫它且不帶任何引數，是用來表示非同步操作已完成。

最後，請注意緊接在 `request` 方法之後的 `keepOpen` 方法。通常你會從命令列執行你的測試，或作為自動化整合處理程序的一部分，並且你可以讓 `chai-http` 自動啟動和停止你的伺服器。

但是，當你提交聯結到你的專案時執行的測試需要你的伺服器保持運作，所以你需要使用 `keepOpen` 方法來防止 `chai-http` 停止你的伺服器。

# --instructions--

在 `tests/2_functional-tests.js` 中，修改 `'Test GET /hello with no name'` 測試（`// #1`）以斷言回應的 `status` 和 `text`，使測試通過。請勿更改傳遞給斷言的引數。

不應該有 URL 查詢。若沒有名稱 URL 查詢，端點會回應 `hello Guest`。

# --hints--

所有測試都應該通過

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該測試 `res.status` 是否等於 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

你應該測試 `res.text` 是否等於 `'hello Guest'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /('|")hello Guest\1/);
```

