---
id: 587d824d367417b2b2512c51
title: 測試數組是否包含項目
challengeType: 2
forumTopicId: 301603
dashedName: test-if-an-array-contains-an-item
---

# --description--

As a reminder, this project is being built upon the following starter project cloned from <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

在 `Arrays` 套件中標記爲 `#12` 的測試下的 `tests/1_unit-tests.js` 中，將每個 `assert` 更改爲 `assert.include` 或 `assert.notInclude` 以使測試通過（應評估爲 `true`）。 不要修改傳給斷言的參數。

# --hints--

應通過所有測試。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

應該爲第一個斷言選擇正確的方法：`include` 或 `notInclude`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInclude',
  "It's summer in july..."
);
```

應該爲第二個斷言選擇正確的方法：`include` 或 `notInclude`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'include',
  'JavaScript is a backend language !!'
);
```

