---
id: 587d824b367417b2b2512c48
title: 使用 Assert.isOK 和 Assert.isNotOK
challengeType: 2
forumTopicId: 301607
dashedName: use-assert-isok-and-assert-isnotok
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

`isOk()` 會測試是否為真值，`isNotOk()` 會測試是否為假值。

若要進一步了解 truthy 和 falsy 值，請嘗試我們的 <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-algorithm-scripting/falsy-bouncer" target="_blank" rel="noopener noreferrer nofollow">Falsy Bouncer</a> 挑戰。

# --instructions--

在 `Basic Assertions` 測試套件中，位於 `tests/1_unit-tests.js` 裡標籤為 `#3` 的測試中，將每個 `assert` 改為 `assert.isOk()` 或 `assert.isNotOk()`，以使測試通過（應該評估為 `true`）。不要更改傳遞給 assert 的引數。

# --hints--

所有測試都應該通過。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

你應該為第一個斷言選擇正確的方法 - `isOk` 與 `isNotOk`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNotOk', 'Null is falsy');
```

你應該為第二個斷言選擇正確的方法 - `isOk` 與 `isNotOk`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isOk', 'A string is truthy');
```

你應該為第三個斷言選擇正確的方法 - `isOk` 與 `isNotOk`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[2].method, 'isOk', 'true is truthy');
```

