---
id: 587d824b367417b2b2512c49
title: 測試真值性
challengeType: 2
forumTopicId: 301596
dashedName: test-for-truthiness
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

`isTrue()` 會測試布林值 `true`，而 `isNotTrue()` 在給定任何非布林值 `true` 時會通過。

```js
assert.isTrue(true, 'This will pass with the boolean value true');
assert.isTrue('true', 'This will NOT pass with the string value "true"');
assert.isTrue(1, 'This will NOT pass with the number value 1');
```

`isFalse()` 和 `isNotFalse()` 也存在，且行為類似於它們的 true 對應函式，除了它們尋找的是布林值 `false`。

# --instructions--

在 `Basic Assertions` 測試套件中，位於 `tests/1_unit-tests.js` 裡標籤為 `#4` 的測試中，將每個 `assert` 改為 `assert.isTrue` 或 `assert.isNotTrue`，以使測試通過（應該計算為 `true`）。不要更改傳遞給 assert 的引數。

# --hints--

所有測試都應該通過。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

你應該為第一個斷言選擇正確的方法 - `isTrue` 與 `isNotTrue`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isTrue', 'True is true');
```

你應該為第二個斷言選擇正確的方法 - `isTrue` 與 `isNotTrue`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isTrue',
    'Double negation of a truthy value is true'
  );
```

你應該為第三個斷言選擇正確的方法 - `isTrue` 與 `isNotTrue`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isNotTrue',
    'A truthy object is not true - neither is a false one'
  );
```

