---
id: 587d824f367417b2b2512c5c
title: 使用無頭瀏覽器模擬操作
challengeType: 2
dashedName: simulate-actions-using-a-headless-browser
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

在接下來的挑戰中，你將使用無頭瀏覽器模擬人類與頁面的互動。

無頭瀏覽器是沒有圖形介面的網頁瀏覽器。它們能夠以與一般瀏覽器相同的方式呈現和解析 HTML、CSS 和 JavaScript，使其在測試網頁時特別有用。

在接下來的挑戰中，你將使用 Zombie.js，這是一個輕量級的無頭瀏覽器，不需要安裝額外的二進位檔。但還有許多其他更強大的無頭瀏覽器選項。

Mocha 允許你在任何實際測試執行之前先執行一些程式碼。這對於像是為資料庫添加條目，並在其餘測試中使用，會很有幫助。

使用無頭瀏覽器時，在執行測試之前，你需要**造訪**你將測試的頁面。

`suiteSetup` 掛鉤只會在測試套件開始時執行一次。

還有其他幾種掛鉤型別，可以在每個測試之前、每個測試之後或測試套件結束時執行程式碼。更多資訊請參閱 Mocha 文件。

# --instructions--

在 `tests/2_functional-tests.js` 中，緊接在 `Browser` 宣告式之後，為變數的 `site` 屬性添加你的專案 URL：

```js
Browser.site = 'http://0.0.0.0:3000'; // Your URL here
```

然後在 `'Functional Tests with Zombie.js'` 測試組的根層級，使用以下程式碼實例化一個新的 `Browser` 物件實體：

```js
const browser = new Browser();
```

並使用 `suiteSetup` 掛鉤以以下程式碼將 `browser` 導向 `/` 路由。**注意**：`done` 作為回呼傳遞給 `browser.visit`，你不應該喚起它。

```js
suiteSetup(function(done) {
  return browser.visit('/', done);
});
```

# --hints--

所有測試都應該通過。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=4');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

