---
id: 587d824f367417b2b2512c5a
title: 使用 Chai-HTTP 對 API 回應執行功能測試 III - PUT 方法
challengeType: 2
forumTopicId: 301590
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iii
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

當你測試 `PUT` 請求時，你經常會同時傳送資料。你隨 `PUT` 請求包含的資料稱為請求的主體。

要傳送 `PUT` 請求和 JSON 物件到 `'/travellers'` 端點，你可以使用 `chai-http` 外掛程式的 `put` 和 `send` 方法：

```js
chai
  .request(server)
  .keepOpen()
  .put('/travellers')
  .send({
    "surname": [last name of a traveller of the past]
  })
  ...
```

路由回應：

```json
{
  "name": [first name],
  "surname": [last name],
  "dates": [birth - death years]
}
```

請參閱伺端程式碼中對 `'/travellers'` 端點的不同回應。

# --instructions--

在 `tests/2_functional-tests.js` 中，修改 `'Send {surname: "Colombo"}'` 測試（`// #3`），並使用 `put` 和 `send` 方法來測試 `'/travellers'` 端點。

使用你的 PUT 請求傳送以下 JSON 物件：

```json
{
  "surname": "Colombo"
}
```

在 `request.end` 回呼中檢查以下內容：

1.  `status` 應該是 `200`
2.  `type` 應該是 `application/json`
3.  `body.name` 應該是 `Cristoforo`
4.  `body.surname` 應該是 `Colombo`

請遵循上述的斷言順序－我們依賴它。此外，完成後務必移除 `assert.fail()`。

# --hints--

所有測試都應該通過。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該測試 `res.status` 是否為 200。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

你應該測試 `res.type` 是否為 `'application/json'`。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

你應該測試 `res.body.name` 是否為 `'Cristoforo'`。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Cristoforo\1/);
```

你應該測試 `res.body.surname` 是否為 `'Colombo'`。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")Colombo\1/);
```

