---
id: 587d824a367417b2b2512c46
title: 瞭解 JavaScript 斷言的工作原理
challengeType: 2
forumTopicId: 301589
dashedName: learn-how-javascript-assertions-work
---

# --description--

Working on these challenges will involve you writing your code using one of the following methods:

- 克隆 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub 倉庫</a> 並完成您的本地項目。
- Use a site builder of your choice to complete the project. Be sure to incorporate all the files from our GitHub repo.

# --instructions--

在 `Basic Assertions` 套件中標記爲 `#1` 的測試下的 `tests/1_unit-tests.js` 中，將每個 `assert` 更改爲 `assert.isNull` 或 `assert.isNotNull` 以使測試通過（應評估爲 `true`）。 不要修改傳給斷言的參數。

# --hints--

應通過所有測試。

```js
  $.get(code + '/_api/get-tests?type=unit&n=0').then(
    (data) => {
      assert.equal(data.state, 'passed');
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

請選擇正確的斷言—isNull 或 isNotNull

```js
  $.get(code + '/_api/get-tests?type=unit&n=0').then(
    (data) => {
      assert.equal(data.assertions[0].method, 'isNull', 'Null is null');
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

You should choose the right assertion - isNull vs. isNotNull.

```js
  $.get(code + '/_api/get-tests?type=unit&n=0').then(
    (data) => {
      assert.equal(data.assertions[1].method, 'isNotNull', '1 is not null');
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

