---
id: 587d824e367417b2b2512c55
title: 測試物件是否具有屬性
challengeType: 2
forumTopicId: 301604
dashedName: test-if-an-object-has-a-property
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

`property` 斷言實際物件具有指定的屬性。

# --instructions--

在 `Objects` 測試套件中，位於 `tests/1_unit-tests.js` 裡標籤為 `#16` 的測試中，將每個 `assert` 改為 `assert.property` 或 `assert.notProperty`，以使測試通過（應該評估為 `true`）。不要更改傳遞給 assert 的引數。

# --hints--

所有測試都應該通過。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該為第一個斷言選擇正確的方法 - `property` 與 `notProperty`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notProperty',
  'A car has not wings'
);
```

你應該為第二個斷言選擇正確的方法——`property` 與 `notProperty`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'property',
  'planes have engines'
);
```

你應該為第三個斷言選擇正確的方法 - `property` 與 `notProperty`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'property', 'Cars have wheels');
```

