---
id: 587d824f367417b2b2512c5b
title: 使用 Chai-HTTP IV 的 PUT 方法對 API 回應執行功能測試
challengeType: 2
forumTopicId: 301591
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iv
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

這個練習與前一個類似。

既然你已經知道如何測試 `PUT` 請求，現在輪到你從頭開始操作了。

# --instructions--

在 `tests/2_functional-tests.js` 中，修改 `'Send {surname: "da Verrazzano"}'` 測試（`// #4`），並使用 `put` 和 `send` 方法來測試 `'/travellers'` 端點。

使用你的 PUT 請求傳送以下 JSON 物件：

```json
{
  "surname": "da Verrazzano"
}
```

在 `request.end` 回呼中檢查以下內容：

1.  `status` 應該是 `200`
2.  `type` 應該是 `application/json`
3.  `body.name` 應該是 `Giovanni`
4.  `body.surname` 應該是 `da Verrazzano`

請遵循上述的斷言順序－我們依賴它。此外，完成後務必移除 `assert.fail()`。

# --hints--

所有測試都應該通過

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該測試 `res.status` 是否為 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

你應該測試 `res.type` 是否為 `'application/json'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

你應該測試 `res.body.name` 是否為 `'Giovanni'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Giovanni\1/);
```

你應該測試 `res.body.surname` 是否為 `'da Verrazzano'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")da Verrazzano\1/);
```

