---
id: 587d824d367417b2b2512c53
title: 測試字串是否包含子字串
challengeType: 2
forumTopicId: 301597
dashedName: test-if-a-string-contains-a-substring
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

`include()` 和 `notInclude()` 也適用於字串！`include()` 斷言實際字串包含預期的子字串。

# --instructions--

在 `Strings` 測試組中，位於 `tests/1_unit-tests.js` 內標籤為 `#14` 的測試中，將每個 `assert` 更改為 `assert.include` 或 `assert.notInclude`，以使測試通過（應該評估為 `true`）。請勿更改傳遞給 assert 的引數。

# --hints--

所有測試都應該通過。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該為第一個斷言選擇正確的方法 - `include` 與 `notInclude`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'include',
  "'Arrow' contains 'row'..."
);
```

你應該為第二個斷言選擇正確的方法 - `include` 與 `notInclude`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notInclude',
  "... a 'dart' doesn't contain a 'queue'"
);
```

