---
id: 587d824b367417b2b2512c4a
title: 用兩個等號斷言相等
challengeType: 2
forumTopicId: 301609
dashedName: use-the-double-equals-to-assert-equality
---

# --description--

As a reminder, this project is being built upon the following starter project cloned from <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`equal()` 使用 `==` 比較對象。

# --instructions--

在 `Equality` 套件中標記爲 `#5` 的測試下的 `tests/1_unit-tests.js` 中，將每個 `assert` 更改爲 `assert.equal` 或 `assert.notEqual` 以使測試通過（應評估爲 `true`）。 不要修改傳給斷言的參數。

# --hints--

不應有未通過的測試

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

應該爲第一個斷言選擇正確的方法：`equal` 或 `notEqual`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'equal',
    'Numbers are coerced into strings with == '
  );
```

應該爲第二個斷言選擇正確的方法：`equal` 或 `notEqual`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notEqual',
    ' == compares object references'
  );
```

應該爲第三個斷言選擇正確的方法：`equal` 或 `notEqual`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'equal',
    "6 * '2' is 12 ! It should be equal to '12'"
  );
```

應該爲第四個斷言選擇正確的方法：`equal` 或 `notEqual`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[3].method, 'notEqual', "6 + '2' is '62'...");
```

