---
id: 587d824a367417b2b2512c46
title: 瞭解 JavaScript 斷言的運作方式
challengeType: 2
forumTopicId: 301589
dashedName: learn-how-javascript-assertions-work
---

# --description--

處理這些挑戰將涉及你使用以下其中一種方法來撰寫你的程式碼：

- 複製<a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">這個 GitHub 倉庫</a>並在區域完成這些挑戰。
- 使用你選擇的網站建置工具來完成專案。務必將我們 GitHub 倉庫中的所有檔案納入其中。

# --instructions--

在 `tests/1_unit-tests.js` 中，位於 `Basic Assertions` 測試套件下標籤為 `#1` 的測試中，將每個 `assert` 更改為 `assert.isNull` 或 `assert.isNotNull`，以使測試通過（應該評估為 `true`）。不要更改傳遞給 assert 的引數。

# --hints--

所有測試都應該通過。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

你應該為第一個斷言選擇正確的方法 - `isNull` 與 `isNotNull`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNull', 'Null is null');
```

你應該為第二個斷言選擇正確的方法 - `isNull` 與 `isNotNull`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isNotNull', '1 is not null');
```

