---
id: 587d824c367417b2b2512c4e
title: 測試一個值是否低於或至少與另一個值一樣大
challengeType: 2
forumTopicId: 301606
dashedName: test-if-one-value-is-below-or-at-least-as-large-as-another
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

# --instructions--

在 `Comparisons` 測試套件中，位於 `tests/1_unit-tests.js` 裡標籤為 `#9` 的測試中，將每個 `assert` 改為 `assert.isBelow` 或 `assert.isAtLeast`，以使測試通過（應該評估為 `true`）。不要更改傳遞給 assert 的引數。

# --hints--

所有測試都應該通過。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該為第一個斷言選擇正確的方法 - `isBelow` 與 `isAtLeast`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtLeast',
  '5 is at least (>=) 5'
);
```

你應該為第二個斷言選擇正確的方法 - `isBelow` 與 `isAtLeast`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isAtLeast',
  '2 * Math.random() is at least 0'
);
```

你應該為第三個斷言選擇正確的方法 - `isBelow` 與 `isAtLeast`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isBelow', '1 is smaller than 2');
```

你應該為第四個斷言選擇正確的方法 - `isBelow` 與 `isAtLeast`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isBelow',
  '2/3 (0.6666) is smaller than 1'
);
```

