---
id: 587d8250367417b2b2512c5d
title: 使用無頭瀏覽器執行功能測試
challengeType: 2
forumTopicId: 301595
dashedName: run-functional-tests-using-a-headless-browser
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

在頁面上有一個輸入表單。它會將資料以 AJAX 請求的方式傳送到 `PUT /travellers` 端點。

當請求成功完成時，客戶端程式碼將包含回應中資訊的 `<div>` 附加到 DOM。

以下是如何使用 Zombie.js 與表單互動的範例：

```js
test('Submit the surname "Polo" in the HTML form', function (done) {
  browser.fill('surname', 'Polo').then(() => {
    browser.pressButton('submit', () => {
      browser.assert.success();
      browser.assert.text('span#name', 'Marco');
      browser.assert.text('span#surname', 'Polo');
      browser.assert.elements('span#dates', 1);
      done();
    });
  });
});
```

首先，`browser` 物件的 `fill` 方法會將表單的 `surname` 欄位填入值 `'Polo'`。`fill` 傳回一個 promise，因此接著會串鏈 `then`。

在 `then` 回呼中，使用 `browser` 物件的 `pressButton` 方法來喚起表單的 `submit` 事件監聽器。`pressButton` 方法是非同步的。

然後，一旦從 AJAX 請求收到回應，就會進行一些斷言以確認：

1.  回應的狀態是 `200`
2.  `<span id='name'></span>` 元素內的文字符合 `'Marco'`
3.  `<span id='surname'></span>` 元素內的文字符合 `'Polo'`
4.  有 `1` 個 `<span id='dates'></span>` 元素。

最後，會喚起 `done` 回呼，這是因為非同步測試所需要的。

# --instructions--

在 `tests/2_functional-tests.js` 中，於 `'Submit the surname "Colombo" in the HTML form'` 測試（`// #5`）中，自動化以下內容：

1.  在表單中填寫姓氏 `Colombo`
2.  按下提交按鈕

並且在 `pressButton` 回呼中：

1.  斷言狀態為 OK `200`
2.  斷言元素 `span#name` 內的文字是 `'Cristoforo'`
3.  斷言元素 `span#surname` 內的文字是 `'Colombo'`
4.  斷言元素 `span#dates` 存在且其數量為 `1`

不要忘記移除 `assert.fail()` 呼叫。

# --hints--

所有測試都應該通過。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該斷言無頭瀏覽器請求成功。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

你應該斷言元素 `span#name` 內的文字是 `'Cristoforo'`。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Cristoforo\1/);
```

你應該斷言元素 `span#surname` 內的文字是 `'Colombo'`。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Colombo\1/);
```

你應該斷言元素 `span#dates` 存在且其數量為 1。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```

