---
id: 587d824c367417b2b2512c4c
title: 使用 .deepEqual 和 .notDeepEqual 斷言深度相等性
challengeType: 2
forumTopicId: 301587
dashedName: assert-deep-equality-with--deepequal-and--notdeepequal
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

`deepEqual()` 斷言兩個物件是深度相等。

# --instructions--

在 `tests/1_unit-tests.js` 中 `Equality` 測試套件下標籤為 `#7` 的測試中，將每個 `assert` 改為 `assert.deepEqual` 或 `assert.notDeepEqual`，以使測試通過（應該計算結果為 `true`）。不要更改傳遞給 assert 的引數。

# --hints--

所有測試都應該通過。

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
})();
```

你應該為第一個斷言選擇正確的 `method` - `deepEqual` 與 `notDeepEqual`。

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'deepEqual',
    'The order of the keys does not matter'
  );
})();
```

你應該為第二個斷言選擇正確的 `method` - `deepEqual` 與 `notDeepEqual`。

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notDeepEqual',
    'The position of elements within an array does matter'
  );
})();
```

