---
id: 587d824b367417b2b2512c4b
title: 使用三重等號來斷言嚴格相等性
challengeType: 2
forumTopicId: 301610
dashedName: use-the-triple-equals-to-assert-strict-equality
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

`strictEqual()` 使用 `===` 比較物件。

# --instructions--

在 `tests/1_unit-tests.js` 中 `Equality` 測試套件下標籤為 `#6` 的測試中，將每個 `assert` 改為 `assert.strictEqual` 或 `assert.notStrictEqual`，以使測試通過（應該計算為 `true`）。不要更改傳遞給 assert 的引數。

# --hints--

所有測試都應該通過。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

你應該為第一個斷言選擇正確的 `method` - `strictEqual` 與 `notStrictEqual`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'notStrictEqual',
    'with strictEqual the type must match'
  );
```

你應該為第二個斷言選擇正確的 `method` - `strictEqual` 與 `notStrictEqual`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'strictEqual', '3*2 = 6...');
```

你應該為第三個斷言選擇正確的 `method` — `strictEqual` 與 `notStrictEqual`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'strictEqual',
    "6 * '2' is 12. Types match !"
  );
```

你應該為第四個斷言選擇正確的 `method` — `strictEqual` 與 `notStrictEqual`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[3].method,
    'notStrictEqual',
    'Even if they have the same elements, the Arrays are notStrictEqual'
  );
```

