---
id: 5f8884f4c46685731aabfc41
title: 使用無頭瀏覽器執行功能測試 II
challengeType: 2
forumTopicId: 301594
dashedName: run-functional-tests-using-a-headless-browser-ii
---

# --description--

提醒你，這個專案是基於從 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 複製的以下起始專案所建立。

# --instructions--

在 `tests/2_functional-tests.js` 中，在 `'Submit the surname "Vespucci" in the HTML form'` 測試（`// #6`）裡，自動化以下內容：

1.  在表單中填寫姓氏 `Vespucci`
2.  按下提交按鈕

並且在 `pressButton` 回呼中：

1.  斷言狀態為 OK `200`
2.  斷言元素 `span#name` 內的文字是 `'Amerigo'`
3.  斷言元素 `span#surname` 內的文字是 `'Vespucci'`
4.  斷言元素 `span#dates` 存在且其數量為 `1`

不要忘記移除 `assert.fail()` 呼叫。

# --hints--

所有測試都應該通過。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你應該斷言無頭瀏覽器請求成功。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

你應該斷言元素 `span#name` 內的文字是 `'Amerigo'`。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Amerigo\1/);
```

你應該斷言元素 `span#surname` 內的文字是 `'Vespucci'`。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Vespucci\1/);
```

你應該斷言元素 `span#dates` 存在且其數量為 1。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```

