---
id: 6733b5632fb854bc6c58dc0d
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

أنشئ الآن كائنًا آخر، بنفس الثلاث خصائص، لتمثيل نوع آخر من الخزامى الإسباني يسمى `Pretty Polly`، وعيّنه في متغير باسم `prettyPolly`.

# --hints--

يجب أن يكون لديك متغير باسم `prettyPolly`.

```js
assert.exists(prettyPolly);
```

يجب أن تعيّن كائنًا إلى متغير `prettyPolly` الخاص بك.

```js
assert.isObject(prettyPolly);
```

يجب أن يحتوي كائن `prettyPolly` على خاصية `commonName` بقيمة `"Spanish lavender"`.

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

يجب أن يحتوي كائن `prettyPolly` على خاصية `scientificName` بقيمة `"Lavandula stoechas"`.

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

يجب أن يحتوي كائن `prettyPolly` على خاصية `cultivar` بقيمة `"Pretty Polly"`.

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

يجب أن يحتوي كائن `prettyPolly` على ثلاث مفاتيح.

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

--fcc-editable-region--

--fcc-editable-region--
```
