---
id: 6733b5e6f210c1bd1afdf8b3
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

تمت إضافة بقية كائنات النباتات لك. حان الوقت الآن لإنشاء الكتالوج الخاص بك.

مثل كائنات JavaScript العادية، تحتوي كائنات `Map` على مجموعة من أزواج المفتاح-القيمة. لإنشاء كائن `Map`، يمكنك استخدام دالة البناء `Map`:

```js
const myMap = new Map();
```

أعلن عن متغير `catalog` وقم بتعيين كائن `Map` له.

# --hints--

يجب أن يكون لديك متغير باسم `catalog`.

```js
assert.exists(catalog);
```

يجب أن تعيّن كائن `Map` إلى متغير `catalog` الخاص بك.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

--fcc-editable-region--

--fcc-editable-region--
```
