---
id: 673339d96c56d36e963e888e
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستبني كتالوجًا لمشتل النباتات أثناء ممارسة استخدام كائنات Map و Set في JavaScript.

لتمثيل كل نبات في كتالوجك، ستستخدم كائنات، كل منها يحتوي على ثلاث خصائص:

- `commonName`، الاسم الشائع للنبات
- `scientificName`، الاسم العلمي (أو الاسم اللاتيني) للنبات
- `cultivar`، نوع النبات

ابدأ بإنشاء كائن يحتوي على الخصائص الثلاثة أعلاه لـ `Spanish lavender` (الاسم العلمي `Lavandula stoechas`) من نوع يسمى `Ballerina` وعيّنه إلى متغير باسم `ballerina`.

# --hints--

يجب أن يكون لديك متغير باسم `ballerina`.

```js
assert.exists(ballerina);
```

يجب أن تعيّن كائنًا إلى متغير `ballerina` الخاص بك.

```js
assert.isObject(ballerina);
```

يجب أن يحتوي كائن `ballerina` على خاصية `commonName` مضبوطة على `"Spanish lavender"`.

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

يجب أن يحتوي كائن `ballerina` على خاصية `scientificName` مضبوطة على `"Lavandula stoechas"`.

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

يجب أن يحتوي كائن `ballerina` على خاصية `cultivar` مضبوطة على `"Ballerina"`.

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

يجب أن يحتوي كائن `ballerina` على ثلاث مفاتيح.

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
