---
id: 6734f0ae9d1cf4dfd43a7805
title: الخطوة 15
challengeType: 1
dashedName: step-15
---

# --description--

لإزالة عنصر من كائن الخريطة، يمكنك استخدام طريقة `delete` التي تأخذ مفتاحًا كمعلمة وتحذف الإدخال المرتبط بهذا المفتاح. تُرجع قيمة منطقية تشير إلى ما إذا كانت عملية الحذف ناجحة.

أنشئ دالة باسم `removePlant` تأخذ كائن نبات كمعلمتها وتحذفه من الكتالوج. اجعلها تُرجع إما `true` أو `false` حسب ما إذا تم حذف النبات المحدد من الكتالوج.

بعد ذلك، سجّل نتيجة استدعاء دالتك الجديدة مع السلسلة النصية `"ballerina"` كمعلمة لها.

# --hints--

يجب أن يكون لديك دالة باسم `removePlant`.

```js
assert.isFunction(removePlant);
```

يجب أن تأخذ دالة `removePlant` معلمة واحدة فقط.

```js
assert.lengthOf(removePlant, 1);
```

يجب أن تستدعي دالة `removePlant` طريقة `delete` على `catalog`.

```js
let flag = false;
const temp = catalog.delete;
try {
    catalog.delete = (p) => {flag = true}
    removePlant("");
    assert.isTrue(flag);
} finally {
    catalog.delete = temp;
}
```

يجب أن تحذف `removePlant(item)` العنصر `item` من الكتالوج.

```js
const testArr = [];
const temp = catalog.delete;
try {
    catalog.delete = (p) => {testArr.push(p)}
    removePlant("test");
    assert.deepEqual(testArr, ["test"]);
} finally {
    catalog.delete = temp;
}
```

يجب أن تُرجع دالة `removePlant` قيمة منطقية تشير إلى نجاح عملية الحذف.

```js
try {
    assert.isFalse(removePlant("test"));
    assert.isTrue(removePlant(royalCrown));
} finally {
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 })
}
```

يجب أن تسجل في وحدة التحكم `removePlant("ballerina")`.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*removePlant\s*\(\s*("|')ballerina\1\s*\)\s*\)/);
```



# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const sellPlants = (plant, size, potsNo) => {
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

--fcc-editable-region--

--fcc-editable-region--

console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```
