---
id: 67349ba230cff237c0f2bdf2
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

تتذكر كائنات الخريطة ترتيب إدخال المفاتيح في الخريطة. لتعيين زوج مفتاح-قيمة في كائن الخريطة الخاص بك، يمكنك استخدام الطريقة `set`. المثال التالي يعين المفتاح `"name"` إلى القيمة `"Jerry"` في كائن الخريطة `myMap`.

```js
myMap.set("name", "Jerry");
```

عيّن المفتاح `"ballerina"` إلى قيمة كائن `ballerina` الخاص بك في خريطتك `catalog`. ثم سجّل الكتالوج في وحدة التحكم.

لاحظ أنه لرؤية كائن الخريطة الخاص بك بشكل صحيح، يجب فتح وحدة تحكم المتصفح.

# --hints--

يجب أن تعيّن المفتاح `"ballerina"` إلى قيمة كائن `ballerina` الخاص بك في خريطتك `catalog`.

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

يجب أن تحتوي خريطتك `catalog` على زوج مفتاح-قيمة واحد فقط.

```js
assert.lengthOf(catalog, 1);
```

يجب أن تسجل خريطتك `catalog` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
--fcc-editable-region--

--fcc-editable-region--
```
