---
id: 6734a606e5968a5a19b77171
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

استمر في إضافة النباتات المتبقية إلى الكتالوج كما فعلت في الخطوات السابقة.

# --hints--

يجب أن تضبط كائن `willowVale` كمفتاح في خريطة `catalog` الخاصة بك.

```js
assert.isTrue(catalog.has(willowVale));
```

يجب أن تكون قيمة المفتاح `willowVale` في خريطة `catalog` كائنًا يحتوي على الخصائص `small` و`medium` و`large`، كل منها بقيمة عدد صحيح موجب.

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

يجب أن تضبط كائن `hidcote` كمفتاح في خريطة `catalog` الخاصة بك.

```js
assert.isTrue(catalog.has(hidcote));
```

يجب أن تكون قيمة المفتاح `hidcote` في خريطة `catalog` كائنًا يحتوي على الخصائص `small` و`medium` و`large`، كل منها بقيمة عدد صحيح موجب.

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

يجب أن تضبط كائن `imperialGem` كمفتاح في خريطة `catalog` الخاصة بك.

```js
assert.isTrue(catalog.has(imperialGem));
```

يجب أن تكون قيمة المفتاح `imperialGem` في خريطة `catalog` كائنًا يحتوي على الخصائص `small` و`medium` و`large`، كل منها بقيمة عدد صحيح موجب.

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

يجب أن تضبط كائن `royalCrown` كمفتاح في خريطة `catalog` الخاصة بك.

```js
assert.isTrue(catalog.has(royalCrown));
```

يجب أن تكون قيمة المفتاح `royalCrown` في خريطة `catalog` كائنًا يحتوي على الخصائص `small` و`medium` و`large`، كل منها بقيمة عدد صحيح موجب.

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

يجب أن تظل مفتاح `ballerina` موجودًا في خريطة `catalog` الخاصة بك.

```js
assert.isTrue(catalog.has(ballerina));
```

يجب أن تظل مفتاح `prettyPolly` موجودًا في خريطة `catalog` الخاصة بك.

```js
assert.isTrue(catalog.has(prettyPolly));
```

يجب أن تحتوي خريطة `catalog` الخاصة بك على ستة عناصر.

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
