---
id: 6734c10a327760a665b7d5b9
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

في الخطوات القليلة القادمة، ستعمل على دالة تسمح لك بتحديث الكتالوج عند بيع النباتات.

أنشئ دالة باسم `sellPlants` تأخذ ثلاث معلمات: كائن النبات، حجم الأصيص، وعدد الأصص للبيع، بهذا الترتيب.

عندما يكون عدد الأصص المحدد أكبر من الأصص المتاحة، اجعل دالتك تُرجع `Not enough <pot-size> size pots for <item-name>. Only <pot-number> left.`. استبدل `<pot-size>` بحجم الأصيص، و`<pot-number>` بعدد الأصص المتبقية لذلك الحجم، و`<item-name>` بالاسم العلمي للنبات متبوعًا بمسافة واسم الصنف بين علامتي اقتباس مفردتين.

لاختبار أن كل شيء يعمل، سجّل نتيجة استدعاء دالتك `sellPlants` مع كائن `ballerina`، و`"small"`، و`25` في وحدة التحكم.

# --hints--

لا يجب تعديل خريطة `catalog` الخاصة بك.

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 20, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

يجب أن يكون لديك دالة باسم `sellPlants`.

```js
assert.isFunction(sellPlants);
```

يجب أن تأخذ دالة `sellPlants` ثلاث معلمات.

```js
assert.lengthOf(sellPlants, 3);
```

`sellPlants(ballerina, "small", 25)` يجب أن تُرجع `Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.`

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")

assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

يجب أن تسجل `sellPlants(ballerina, "small", 25)` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*25\s*\)\s*\)/);
```

`sellPlants(ballerina, "small", 10)` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(sellPlants(ballerina, "small", 10));
// prevent hardcoding
assert.isUndefined(sellPlants(hidcote, "large", 5));
assert.isUndefined(sellPlants(royalCrown, "medium", 5));
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

--fcc-editable-region--

--fcc-editable-region--

console.log(catalog)
```
