---
id: 6734a2385c52c54d916a8b86
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

أضف زوج مفتاح-قيمة آخر إلى الكتالوج. استخدم المتغير `prettyPolly` كمفتاح واجعل القيمة كائنًا يحتوي على الخصائص `small` و`medium` و`large` التي تكون قيمها أعداد صحيحة مرة أخرى.

# --hints--

يجب أن تعيّن كائن `prettyPolly` كمفتاح في خريطة `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

يجب أن تكون قيمة المفتاح `prettyPolly` في خريطة `catalog` كائنًا يحتوي على الخصائص `small` و`medium` و`large`.

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

يجب أن تكون الخاصية `small` للكائن المرتبط بالمفتاح `prettyPolly` في خريطة `catalog` عددًا صحيحًا موجبًا.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

يجب أن تكون الخاصية `medium` للكائن المرتبط بالمفتاح `prettyPolly` في خريطة `catalog` عددًا صحيحًا موجبًا.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

يجب أن تكون الخاصية `large` للكائن المرتبط بالمفتاح `prettyPolly` في خريطة `catalog` عددًا صحيحًا موجبًا.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

يجب أن تظل تملك المفتاح `ballerina` في خريطة `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

يجب أن تحتوي خريطة `catalog` على عنصرين.

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
