---
id: 6734a13a29acdf47f2b8e161
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

على عكس الكائنات العادية التي يجب أن تكون مفاتيحها دائمًا سلسلة نصية، تسمح كائنات الخريطة باستخدام أي نوع بيانات كمفاتيح. في الكتالوج الخاص بك، ستستخدم أزواج المفتاح-القيمة من كائن إلى كائن، حيث تربط كل كائن نبات بكائن يحتوي على عدد الأواني المتاحة لأحجام مختلفة.

عدّل استدعاء `set` لاستخدام كائن `ballerina` كمفتاح. للقيمة، استخدم كائنًا يحتوي على ثلاث خصائص: `small` و`medium` و`large`، كل منها بقيمة عدد صحيح موجب.

# --hints--

يجب أن تعيّن كائن `ballerina` كمفتاح في خريطة `catalog` الخاصة بك.

```js
assert.isTrue(catalog.has(ballerina));
```

يجب أن تكون قيمة المفتاح `ballerina` في خريطة `catalog` كائنًا يحتوي على الخصائص `small` و`medium` و`large`.

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

يجب أن تكون الخاصية `small` للكائن المرتبط بالمفتاح `ballerina` في خريطة `catalog` عددًا صحيحًا موجبًا.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

يجب أن تكون الخاصية `medium` للكائن المرتبط بالمفتاح `ballerina` في خريطة `catalog` عددًا صحيحًا موجبًا.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

يجب أن تكون الخاصية `large` للكائن المرتبط بالمفتاح `ballerina` في خريطة `catalog` عددًا صحيحًا موجبًا.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

يجب أن تحتوي خريطة `catalog` على زوج مفتاح-قيمة واحد فقط.

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
--fcc-editable-region--
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
