---
id: 673521668be7905059b2d555
title: الخطوة 22
challengeType: 1
dashedName: step-22
---

# --description--

أخيرًا، يمكنك استخدام الطريقة `forEach` لتشغيل دالة على كل زوج مفتاح-قيمة في الخريطة الخاصة بك.

```js
myMap.forEach((value, key) => {console.log(key, value)});
```

لاحظ أن الاستدعاء الراجع الممرر إلى `forEach` يصل إلى القيمة عبر المعلمة الأولى والمفتاح عبر المعلمة الثانية لكل مدخل.

عدّل دالتك لتُرجع سلسلة نصية تحتوي على كامل الكتالوج وسجل النتيجة في وحدة التحكم. ابنِ السلسلة النصية باستخدام الهيكل التالي لكل مدخل في `catalog`:

```md
<scientific-name> '<cultivar>': <small-pots> S, <medium-pots> M, <large-pots> L

```

لاحظ أن تمثيل كل مدخل يجب أن ينتهي بحرف سطر جديد.

# --hints--

يجب أن تسجل نتيجة استدعاء `displayCatalog` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*displayCatalog\s*\(\s*\)\s*\)/)
```

يجب أن تستخدم دالة `displayCatalog` الطريقة `forEach` على `catalog`.

```js
let flag = false;
const temp = catalog.forEach;
try {
    catalog.forEach = obj => {flag = true};
    displayCatalog();    
    assert.isTrue(flag);    
} finally {
    catalog.forEach = temp;
}
```

يجب أن تُرجع `displayCatalog` القيمة `--fcc-expected--`، لكن وُجدت القيمة `--fcc-actual--`.

```js
const displayCatalogTest = () => {
    let catalogString = "";
    catalog.forEach((val, key) => {
        catalogString += `${key.scientificName} '${key.cultivar}': ${val.small} S, ${val.medium} M, ${val.large} L
`   })
    return catalogString
}
assert.equal(displayCatalog(), displayCatalogTest())
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

const removePlant = plant => catalog.delete(plant);

const displayCatalog = () => {
--fcc-editable-region--
    for (const entry of catalog.entries()) {
        console.log(entry);
    }
}
displayCatalog()
--fcc-editable-region--
```
