---
id: a202eed8fc186c8434cb6d61
title: Construir um Inversor de String
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

Neste laboratório, você irá construir um inversor de string simples que inverte os caracteres de uma string fornecida.

Por exemplo, `"hello"` deve se tornar `"olleh"`.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `reverseString` que recebe uma string como argumento.
2. A função deve retornar a string invertida.

# --hints--

Você deve ter uma função chamada `reverseString`.

```js
assert.isFunction(reverseString);
```

`reverseString` deve receber uma string como argumento.

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` deve retornar uma string.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` deve retornar a string `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` deve retornar a string `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` deve retornar a string `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
