---
id: 6852db755c9c1274d75aae9a
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將創建一個 PIN 擷取器。PIN 數字隱藏在詩的每一行中。

首先創建一個名為 `pin_extractor` 的函式（程式），並帶有一個參數 `poem`。沒有函式主體會引發錯誤，因此在函式內加入 `pass` 作為暫時的佔位符，使其有效。

# --hints--

你應該有一個名為 `pin_extractor` 的函式（程式）。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

`pin_extractor` 函式（程式）應該有一個名為 `poem` 的參數。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
