---
id: 6852dd8486c7ee0171322950
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

你將需要一個變數來保存秘密程式碼以便解碼，所以請移除 `pass` 並在函式內建立一個名為 `secret_code` 的變數，並將空字串指定給它。

# --hints--

你應該在 `pin_extractor` 函式中有一個 `secret_code` 變數。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

變數 `secret_code` 應該有一個空字串的值。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

你不應該在你的程式碼中再有 `pass`。

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
