---
id: 6852db755c9c1274d75aae9a
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將創建一個密碼提取器，密碼數字隱藏在詩的每一行中。

首先建立一個名為 `pin_extractor` 且帶有參數 `poem` 的函式（程式）。空的函式（程式）會產生錯誤，因此暫時在函式（程式）內加入 `pass` 以使其有效。

# --hints--

你應該有一個名為 `pin_extractor` 的函式（程式）。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

`pin_extractor` 函式（程式）應該有一個名為 `poem` 的參數。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
