---
id: 68530231c05b93324e071104
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

當你在處理該函式時，為了偵錯目的，擁有一個函式呼叫會很有用，所以在函式外建立一個包含這首詩的 `poem` 變數：

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

你可以為此使用多行字串。

然後，使用 `poem` 變數作為引數來呼叫函式。

# --hints--

你應該有一個名為 `poem` 的全域變數。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

`poem` 變數的值應該相等於 `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`。

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

你應該以 `poem` 作為引數呼叫函式（程式）`pin_extractor`。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
