---
id: 68531fbb934d735415053274
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

現在你已經為你的函式（程式）添加了呼叫，你可以在終端機中看到函式內產生的任何錯誤，並且你隨時可以使用 `print` 來查看數值。

每一行中都隱藏著密碼的一個數字，因此在函式中使用 `split` 方法將字串分割成一個行的列表。以換行字元（`\n`）分割行，並將結果列表指定給名為 `lines` 的變數。

# --hints--

你應該在你的 `pin_extractor` 函式中擁有一個名為 `lines` 的變數。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

變數 `lines` 應該有一個值為 `poem.split('\n')`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
