---
id: 685535bf527b2c77236b3a9a
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

第 n 個密碼位元被隱藏為第 n 行第 n 個字的長度。

要找出第 n 個字的長度，下一步是將詩句分割成字的列表：在迴圈內，創建一個變數 `words` 並將 `line` 使用 `split` 方法分割成字後的值指定給它。

然後，為你的程式碼添加一個以 `words` 為引數的 `print` 呼叫。

# --hints--

你應該在 `loop` 裡面有一個 `words` 變數。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

`words` 變數應該有 `line.split()` 的值。

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

你應該使用 `print(words)`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
