---
id: 6855918e87c33c6979d21d3f
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

現在這個函式在從詩中擷取程式碼方面表現良好，但有一個邊緣情況需要考慮：詩句可能比預期的還要短，這會導致函式出錯（你可以嘗試從詩中移除字 `bright`）。

將程式碼行 `secret_code += str(len(words[line_index]))` 放入一個 `if` 述語中，以檢查 `words` 列表中是否有足夠的字。

# --hints--

在迴圈中你應該有一個 `if` 述語。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

if 述語條件應該是 `len(words) > line_index`。


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

在 `if` 述語內你應該有 `secret_code += str(len(words[line_index]))`。


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
