---
id: 68557bd29f7f3c1880f98836
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

在迴圈的最後一行，更新 `print` 呼叫，使其列印 `words` 列表中索引為 `line_index` 的字。

# --hints--

你應該在迴圈結尾有 `print(words[line_index])`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words[line_index])')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(words)
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
