---
id: 68553423d775cc60134ac8bb
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

你需要獨立處理每一行：為 `lines` 創建一個使用 `line` 作為迴圈變數的迴圈。在迴圈內，列印 `line`。

# --hints--

你應該在你的 `pin_extractor` 函式中有一個 `for` 迴圈。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

迴圈應該對 `lines` 進行迭代。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

迴圈應該以 `line` 作為迴圈變數（`for line in lines`）。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

你應該在迴圈內有 `print(line)`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
