---
id: 685578aae65e4106ad2b4300
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

如你在先前的課程中所學，`enumerate` 函式允許在對可迭代物件進行迴圈時追蹤索引：

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

秘密程式碼的第 n 位數字來自第 n 行的第 n 個字，因此你需要知道每一行的編號。

將迴圈改為對 `enumerate(lines)` 進行迭代，並在 `line` 之前添加另一個迴圈變數，名為 `line_index`。

此外，將 `print` 呼叫更新為 `print(line_index, line)`。

# --hints--

迴圈應該對 `enumerate(lines)` 進行迭代。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

該迴圈應該以 `line_index, line` 作為迴圈變數。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

你應該在迴圈內使用 `print(line_index, line)`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
