---
id: 68559f3561d7df287586142e
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

現在針腳拔除器可以運作，但它可以更高效！它可以同時從多首詩中拔除針腳！

在全域的範圍中為你的程式碼添加兩首詩，`poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` 和 `poem3 = 'There\nonce\nwas\na\ndragon'`，並且暫時將帶有對 `pin_extractor` 呼叫的 `print` 註解掉。


# --hints--

`poem2` 和 `poem3` 這兩個變數應該存在並且具有正確的值。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

你應該將 `print(pin_extractor(poem))` 呼叫註解掉。

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

--fcc-editable-region--


print(pin_extractor(poem))
--fcc-editable-region--

```
