---
id: 68592b42b9c7f27cbe123203
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

詩的第三行缺少第三個字，所以針腳比預期的短。

當字數不足時，為你的 `secret_code` 添加一個 `else` 語句，該語句會附加一個 `'0'`，以便使用詩的所有行來尋找數字。

# --hints--

你應該有一個 `else` 語句。


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1] == _Node()
    `)
})
```

`else` 語句應該包含 `secret_code += '0'`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_if('len(words) > line_index').is_equivalent('if len(words) > line_index:\\n  secret_code += str(len(words[line_index]))\\nelse:\\n  secret_code += \\'0\\'')
    `)
})
```

`pin_extractor(poem)` 應該傳回 `5202`。


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5202'
    `)
})
```

`secret_code('There\nonce\nwas\na\ndragon')` 應該傳回 `50000`。

```js
({
    test: () => runPython(`
    assert pin_extractor('There\\nonce\\nwas\\na\\ndragon') == '50000'
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
--fcc-editable-region--
        
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))

```
