---
id: 6859290fc8e07e6b2a8d8523
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

移除兩個 `print` 呼叫，並在函式的最後一行添加一個傳回 `secret_code` 的述語。

# --hints--

你不應該有 `print` 呼叫。

```js
({
    test: () => runPython(`
    assert not _Node(_code).block_has_call('print')
    `)
})
```

`pin_extractor` 函式（程式）應該在結尾傳回 `secret_code`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_return('secret_code')
    `)
})
```

`pin_extractor(poem)` 應該傳回 `5262`。


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5262'
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
--fcc-editable-region--
        print(line_index, line)
        words = line.split()
        print(str(len(words[line_index])))
        secret_code += str(len(words[line_index]))
    
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
