---
id: 68592730aaf34b61ae279168
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

從 `len` 傳回的長度是整數，你需要使用 `str` 函式將它轉換成字串，才能將它添加到 `secret_code` 字串中。

更新 `print` 呼叫以列印長度為字串。

# --hints--

你應該在迴圈結尾有 `print(str(len(words[line_index])))`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(len(words[line_index]))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
