---
id: 685929ee6d500c6f2bce6751
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

將 `pin_extractor(poem)` 傳入 `print` 函式（程式），這樣你就可以看到該函式（程式）傳回的內容。

# --hints--

你應該有函式（程式）呼叫 `print(pin_extractor(poem))`。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('print(pin_extractor(poem))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        secret_code += str(len(words[line_index]))
    return secret_code


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

--fcc-editable-region--
pin_extractor(poem)
--fcc-editable-region--

```
