---
id: 68599f0fc7479410b2269742
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

將函式引數更新為 `poems`，然後在函式的所有現有內容周圍創建一個 `for` 迴圈，該迴圈迭代 `poems` 並使用 `poem` 作為迴圈變數。

# --hints--

`pin_extractor` 函式（程式）應該有一個名為 `poems` 的參數。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poems')
    `)
})
```

你應該有一個 `for` 迴圈來迭代 `poems`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('poems')
    `)
})
```

迴圈應該以 `poem` 作為迴圈變數。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('poem')
    `)
})
```

你應該將函式中現有的主體移動到新的 `for` 迴圈中。

```js
({
    test: () => runPython(`
    loop_body = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0]
    assert loop_body.find_variable('secret_code').is_equivalent('secret_code = ""'), "secret_code definition not found"
    assert loop_body.find_variable('lines').is_equivalent("lines = poem.split('\\\\n')"), "lines definition not found"
    assert loop_body.has_return('secret_code'), "return statement not found"
    inner_loop = loop_body.find_for_loops()[0]
    assert inner_loop.find_for_vars().is_equivalent('line_index, line'), "variable of loop not found"
    assert inner_loop.find_for_iter().is_equivalent('enumerate(lines)'), "iter of loop not found"
    inner_loop_body = inner_loop.find_bodies()[0]
    assert inner_loop_body.find_variable('words').is_equivalent("words = line.split()"), "words variable not found"
    if_clause = inner_loop_body.find_ifs()[0]
    if_conditions = if_clause.find_conditions()
    assert if_conditions[0].is_equivalent('len(words) > line_index'), "if condition not found"
    assert if_conditions[1] == _Node(), "else not found"
    if_bodies = if_clause.find_bodies()
    assert if_bodies[0].find_aug_variable('secret_code').is_equivalent("secret_code += str(len(words[line_index]))"), "secret_code inside if clause not found"
    assert if_bodies[1].find_aug_variable('secret_code').is_equivalent("secret_code += '0'"), "secret_code in else clause not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):

    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code
--fcc-editable-region--
poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))

```
