---
id: 6894c1b715b8d18a18ef3203
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie die Verwendung von Python-Klassen und -Objekten, indem Sie ein Inventar für Musikinstrumente erstellen.

Um eine Klasse in Python zu definieren, verwenden Sie das Schlüsselwort `class` gefolgt vom Klassennamen und einem Doppelpunkt.

```python
class ClassName:
    # class contents go here
```

Denken Sie daran, dass Klassennamen in Python üblicherweise PascalCase verwenden, das heißt, jedes Wort beginnt mit einem Großbuchstaben und es gibt keine Leerzeichen zwischen den Wörtern.

Beginnen Sie damit, eine Klasse namens `MusicalInstrument` zu erstellen. Fügen Sie `pass` vorerst als Klassenrumpf hinzu.

# --hints--

Sie sollten eine Klasse namens `MusicalInstrument` erstellen.

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
