---
id: 6895e4940e39950288e406c2
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Fügen Sie als Nächstes eine weitere Methode zu Ihrer Klasse hinzu, die `get_fact` heißt und eine Tatsache über das Instrument zurückgibt.

Diese Methode sollte nur den `self`-Parameter haben und einen f-String zurückgeben, der `The [instrument name] is part of the [instrument type] family of instruments.` sagt.

Im Gegensatz zur `play`-Methode, die eine Nachricht direkt ausgibt, sollte diese Methode einen String zurückgeben, der an anderer Stelle in Ihrem Code verwendet werden kann.

# --hints--

Sie sollten eine Methode namens `get_fact` mit einem `self`-Parameter definieren.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

Ihre `get_fact`-Methode sollte einen f-String zurückgeben, der den Namen und den Typ des Instruments enthält.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
