---
id: 6895e2becb7890fad818ad4d
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Da Sie nun verstehen, wie Sie auf Objektattribute zugreifen, entfernen Sie die vorhandenen `print`-Aufrufe.

# --hints--

Sie sollten die `print(instrument_1.name)`-Anweisung entfernen.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

Sie sollten die `print(instrument_1.instrument_type)`-Anweisung entfernen.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

Sie sollten die `print(instrument_2.name)`-Anweisung entfernen.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

Sie sollten die `print(instrument_2.instrument_type)`-Anweisung entfernen.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
print(instrument_1.name)
print(instrument_1.instrument_type)

print(instrument_2.name)
print(instrument_2.instrument_type)
--fcc-editable-region--
```
