---
id: 6895e197180c9af6988e1e19
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Erstellen Sie eine zweite Instanz namens `instrument_2` für ein `Trumpet`, das ein `brass`-Instrument ist.

Beachten Sie, wie Sie mehrere Instanzen aus derselben Klassenvorlage erstellen können. Jede Instanz speichert ihre eigenen Attributwerte unabhängig von anderen Instanzen.

# --hints--

Sie sollten eine Instanz namens `instrument_2` der `MusicalInstrument`-Klasse erstellen.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

Sie sollten `Trumpet` und `brass` als Argumente übergeben, wenn Sie `instrument_2` erstellen.

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
--fcc-editable-region--

--fcc-editable-region--
```
