---
id: 6895e235935106f88281b648
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Sie können auf die Attribute Ihrer Instanzen mit der Punktnotation zugreifen:

```py
instance_name.attribute_name
```

So können Sie die Werte abrufen, die während der Initialisierung gespeichert wurden.

Versuchen Sie, die `name` und `instrument_type` Attribute beider Instanzen auszugeben. Sie werden die Werte sehen, die Sie ihnen bei der Erstellung zugewiesen haben.

# --hints--

Sie sollten das `name` Attribut von `instrument_1` ausgeben.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

Sie sollten das `instrument_type` Attribut von `instrument_1` ausgeben.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

Sie sollten das `name` Attribut von `instrument_2` ausgeben.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

Sie sollten das `instrument_type` Attribut von `instrument_2` ausgeben.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
