---
id: 6895dffc64e97df0e62e9ef9
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Methoden sind Funktionen, die innerhalb einer Klasse definiert sind und Aktionen mit den Daten der Klasse ausführen können. So erstellen Sie Methoden:

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

In Python ist die `__init__`-Methode eine spezielle Methode, die als Initialisierungsmethode bekannt ist und aufgerufen wird, wenn Sie eine neue Instanz der Klasse erstellen.

Fügen Sie innerhalb Ihrer `MusicalInstrument`-Klasse eine `__init__`-Methode mit drei Parametern hinzu: `self`, `name` und `instrument_type`.

Der `self`-Parameter ist eine übliche Bezeichnung für die Instanz der Klasse und wird als erster Parameter für alle Instanzmethoden in einer Klasse benötigt.

# --hints--

Sie sollten eine `__init__`-Methode innerhalb Ihrer `MusicalInstrument`-Klasse definieren.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

Ihre `__init__`-Methode sollte drei Parameter haben: `self`, `name` und `instrument_type`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
--fcc-editable-region--
    
    pass
--fcc-editable-region--
```
