---
id: 6895e4e023e07603bcef46f9
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Nun verwenden wir die Methoden, die Sie erstellt haben.

Rufen Sie die `play()`-Methode bei Ihrer `instrument_1`-Instanz auf. Dadurch wird die Nachricht über das Spielen der Oboe ausgegeben.

Um eine Methode bei einer Instanz aufzurufen, verwenden Sie die Punktnotation:

```python
instance_name.method_name()
```

# --hints--

Sie sollten die `play()`-Methode bei der `instrument_1`-Instanz aufrufen.

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_1.play()")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
