---
id: 6895e0452b3370f26f365a46
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

In der `__init__`-Methode müssen Sie die Parameter den Instanzattributen zuweisen. Instanzattribute sind Variablen, die zu einer bestimmten Instanz einer Klasse gehören.

Weisen Sie innerhalb Ihrer `__init__`-Methode den `name`-Parameter `self.name` zu und den `instrument_type`-Parameter `self.instrument_type`.

# --hints--

Sie sollten den `name`-Parameter innerhalb der `self.name`-Methode `__init__` zuweisen.

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

Sie sollten den `instrument_type`-Parameter innerhalb der `self.instrument_type`-Methode `__init__` zuweisen.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        pass
--fcc-editable-region--
```
