---
id: 6895e356fd877efd6e3657ee
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--


Definieren Sie innerhalb der `MusicalInstrument`-Klasse eine weitere Methode namens `play` mit nur dem `self`-Parameter. Der `self`-Parameter ermöglicht es Methoden, auf die Attribute der jeweiligen Instanz zuzugreifen.

# --hints--

Sie sollten eine Methode namens `play` mit einem `self`-Parameter definieren.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
    
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')
```
