---
id: 6895e0bdb0348bf46952ab67
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Jetzt ist es Zeit, Instanzen Ihrer `MusicalInstrument` Klasse zu erstellen.

Um eine Instanz einer Klasse zu erstellen, rufen Sie den Klassennamen wie eine Funktion auf und übergeben die erforderlichen Argumente für die `__init__` Methode (außer `self`, die von Python automatisch behandelt wird).

```python
my_instance = ClassName(arg1, arg2)
```

Wenn Sie diese Instanz erstellen, ruft Python Ihre `__init__` Methode auf und speichert die Werte als Instanzattribute, die nur zu dieser speziellen Instanz gehören.

Außerhalb der Klassendefinition erstellen Sie eine Instanz namens `instrument_1` für ein `Oboe`, das ein `woodwind` Instrument ist.

# --hints--

Sie sollten eine Instanz der `MusicalInstrument` Klasse mit dem Namen `instrument_1` erstellen.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

Sie sollten `Oboe` und `woodwind` als Argumente übergeben, wenn Sie `instrument_1` erstellen.

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

--fcc-editable-region--
```
