---
id: 616965351e74d4689eb6de30
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

En una página web puede haber varios elementos `meta`. Cada elemento `meta` añade información sobre la página que no puede ser expresada por otros elementos HTML.

Añade otro elemento `meta` dentro de la `head`. Dale un atributo `name` con el valor `"viewport"` y un atributo `content` con el valor `"width=device-width, initial-scale=1.0"` para que tu página tenga el mismo aspecto en todos los dispositivos.

# --hints--

Debes tener dos elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

Tu nuevo elemento `meta` debe ser un elemento vacío, no debe tener una etiqueta de cierre `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```


El nuevo `meta` debe tener un atributo `name` definido como `"viewport"` y un atributo `content` definido como `"width=device-width, initial-scale=1.0"`.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

Tus dos elementos `meta` deben estar dentro del elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
