---
id: 61695197ac34f0407e339882
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Como has visto en los proyectos anteriores, las páginas web deben comenzar con la declaración `DOCTYPE html`, seguida de un elemento `html`.

Añade una declaración `DOCTYPE html` en la parte superior del documento y un elemento `html` después de la declaración. Dale al elemento `html` un atributo `lang` con el valor `en`.

# --hints--

Su código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debes incluir un espacio despues de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debe definir el tipo de documento para que sea `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Debe cerrar la declaración `DOCTYPE` con `>` después del tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Tu elemento `html` debe tener una etiqueta de apertura con un atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Su elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/gi));
```

Su declaración `DOCTYPE` debe estar al comienzo de su HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
