---
id: 61696ef7ac756c829f9e4048
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

In this project you'll work with an external CSS file to style the page. We've already created a `styles.css` file for you. But before you can use it, you'll need to link it to the page.

Anida un elemento `link` dentro del elemento `head`. Dale un atributo `rel` con el valor `"stylesheet"` y un atributo `href` con el valor `"styles.css"`.

# --hints--

Debea tener un elemento de `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Tu elemento de `link` debe ser un elemento vacío, no debe tener una etiqueta de cierre `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Tu elemento de `link` debe tener un atributo `rel` con el valor `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Tu elemento de `link` debe tener un atributo `href` con el valor `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Tu elemento de `link` debe estar dentro del elemento `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
