---
id: 67051431a73f1ca25d9a6f25
title: Ein Projekt-Ideen-Board erstellen
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Objektkonstante namens `projectStatus` mit den drei Schlüsseln `PENDING`, `SUCCESS` und `FAILURE` definieren. Jeder Status sollte ein Objekt mit einem `description`-Schlüssel erhalten, dessen Wert jeweils `Pending Execution`, `Executed Successfully` und `Execution Failed` ist.
2. Sie sollten eine Klasse namens `ProjectIdea` mit einem `constructor` definieren, der `title` und `description` als Parameter entgegennimmt. Initialisieren Sie die Eigenschaften `title` und `description` mit den übergebenen Parametern. Die Klasse sollte außerdem eine Eigenschaft namens `status` haben, die standardmäßig auf den Wert `projectStatus.PENDING` gesetzt ist.
3. Sie sollten eine Methode namens `updateProjectStatus` innerhalb der `ProjectIdea`-Klasse definieren. Diese Methode sollte einen `newStatus`-Parameter akzeptieren und die Eigenschaft `status` auf den übergebenen Wert aktualisieren.
4. Sie sollten eine `ProjectIdeaBoard`-Klasse mit einem `constructor` definieren, der einen `title` akzeptiert und ein leeres Array namens `ideas` initialisiert, um Instanzen der `ProjectIdea`-Klasse zu speichern.
5. Sie sollten eine Methode namens `pin` innerhalb der `ProjectIdeaBoard`-Klasse definieren, die eine Instanz der `ProjectIdea`-Klasse akzeptiert und die übergebene Instanz in das `ideas`-Array einfügt.
6. Sie sollten eine Methode namens `unpin` innerhalb der `ProjectIdeaBoard`-Klasse definieren. Diese Methode sollte eine Instanz der `ProjectIdea`-Klasse akzeptieren und diese aus dem `ideas`-Array entfernen.
7. Sie sollten eine Methode namens `count` definieren, die die Anzahl der Projektideen im übergebenen `ProjectIdeaBoard`-Array zurückgibt.
8. Sie sollten eine Methode namens `formatToString` definieren, die den Namen der Projektideen, deren Beschreibung und Status im folgenden Format zurückgibt:

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

Sie sollten eine Objektkonstante namens `projectStatus` mit den korrekten Projektstatus und Beschreibungen definieren.

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

Sie sollten eine `ProjectIdea`-Klasse haben.

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

Ihre `ProjectIdea`-Klasse sollte die Eigenschaften `title` und `description` basierend auf den übergebenen Parametern initialisieren. Außerdem sollte der Status auf `projectStatus.PENDING` gesetzt werden.

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

Ihre `ProjectIdea`-Klasse sollte eine `updateProjectStatus`-Methode haben.

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

Sie sollten eine `ProjectIdeaBoard`-Klasse haben.

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

Ihre `ProjectIdeaBoard` sollte die Eigenschaft `title` basierend auf dem übergebenen Parameter initialisieren und ein leeres Array namens `ideas` anlegen, um Instanzen der `ProjectIdea`-Klasse zu speichern.

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

Ihre `ProjectIdeaBoard`-Klasse sollte eine `pin`-Methode haben.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

Ihre `ProjectIdeaBoard`-Klasse sollte eine `unpin`-Methode haben.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

Ihre `ProjectIdeaBoard`-Klasse sollte eine `count`-Methode haben.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

Ihre `ProjectIdeaBoard`-Klasse sollte eine `formatToString`-Methode haben.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Smart Window Locks", "An automation project allowing users to lock, unlock windows automatically based on weather conditions.")` sollte `{ title: 'Smart Window Locks', description: 'An automation project allowing users to lock, unlock windows automatically based on weather conditions.', status: { description: 'Pending Execution' } }` zurückgeben.

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

Das Aufrufen von `updateProjectStatus(projectStatus.SUCCESS)` auf `new ProjectIdea("Fitness Tracker App", "An app that tracks user workouts, diet, and sleep patterns.")` sollte den Status auf `{ description: 'Executed Successfully' }` aktualisieren.

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

Das Aufrufen von `updateProjectStatus(projectStatus.FAILURE)` auf `new ProjectIdea("Breakfast Chef Robot", "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.")` sollte den Status auf `{ description: 'Execution Failed' }` aktualisieren.

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

Das Aufrufen von `updateProjectStatus(projectStatus.SUCCESS)` auf `new ProjectIdea("Online Used Video Games Store", "An online platform where users can buy second hand physical copies of video games from other users.")` sollte den Status auf `{ description: 'Executed Successfully' }` aktualisieren.

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

Sie sollten in der Lage sein, ein `ProjectIdea`-Objekt mit der `ProjectIdeaBoard`-Methode an Ihr `pin` anzuheften.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

Sie sollten in der Lage sein, ein `ProjectIdea`-Objekt mit der `ProjectIdeaBoard`-Methode von Ihrem `unpin` zu lösen.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

Wenn `new ProjectIdeaBoard("Empty Board")` leer ist, sollte `emptyBoard.formatToString()` `Empty Board has 0 idea(s)\n` zurückgeben.

```js
let emptyBoard = new ProjectIdeaBoard("Empty Board");

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

Wenn Sie `new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.")` an `new ProjectIdeaBoard("Tech Projects Board")` anheften, sollte `techProjects.formatToString()` `Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n` zurückgeben.

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```
