---
id: 698b42492a7f5f737895892e
title: Costruire un convertitore da lista di adiacenza a matrice
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

In questo laboratorio costruirai una funzione che converte una rappresentazione di un grafo tramite lista di adiacenza in una matrice di adiacenza. Una lista di adiacenza è un oggetto in cui ogni chiave rappresenta un nodo e il valore corrispondente è un array di nodi a cui il nodo chiave è connesso. Una matrice di adiacenza è un array 2D in cui la voce in posizione `[i][j]` è `1` se esiste un arco dal nodo `i` al nodo `j`, e `0` altrimenti.

Per esempio, data la lista di adiacenza:

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

La matrice di adiacenza corrispondente sarebbe:

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi definire una funzione chiamata `adjacencyListToMatrix` per convertire una lista di adiacenza in una matrice di adiacenza.
2. La funzione deve prendere come argomento un oggetto che rappresenta la lista di adiacenza di un grafo non pesato (sia non orientato che orientato).
3. La funzione deve:
   - Convertire la lista di adiacenza in una matrice di adiacenza.
   - Stampare ogni riga della matrice di adiacenza.
   - Restituire la matrice di adiacenza.

Per esempio, `adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})` dovrebbe stampare:

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

e restituire `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`.

# --hints--

Devi definire una funzione chiamata `adjacencyListToMatrix`.

```js
assert.isFunction(adjacencyListToMatrix);
```

La funzione `adjacencyListToMatrix` deve avere un parametro.

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

La funzione deve determinare correttamente il numero di nodi dalla lista di adiacenza.

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

La funzione deve impostare correttamente i valori della matrice a `1` per gli archi esistenti.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

La funzione deve stampare ogni riga della matrice.

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

La funzione deve restituire la matrice di adiacenza.

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

Quando riceve la lista di adiacenza `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}`, la funzione deve restituire `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`.

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Quando riceve la lista di adiacenza `{0: [1], 1: [0]}`, la funzione deve restituire `[[0, 1], [1, 0]]`.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Quando riceve la lista di adiacenza `{0: [], 1: [], 2: []}`, la funzione deve restituire `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
