---
id: 6a3e524fc3394ef52ed5b797
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Le funzioni hanno un proprio ambito locale. Questo significa che le variabili create all'interno di una funzione sono come gli ingredienti conservati nella dispensa di quella cucina: possono essere consultate solo all'interno di quella specifica funzione.

Tuttavia, le funzioni non sono completamente isolate. Possono leggere le variabili globali definite esternamente al loro ambito. Possono usare quelle variabili per eseguire calcoli, prendere decisioni o creare nuove variabili locali.

All'interno della tua funzione `check_kitchen_stock`, sostituisci la parola chiave `pass` con una nuova variabile chiamata `total_items` e assegna a essa la somma di `available_eggs`, `available_flour` e `available_sugar`.

# --hints--

Dovresti definire una variabile chiamata `total_items` all'interno della tua funzione `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

La tua variabile `total_items` dovrebbe essere la somma di `available_eggs`, `available_flour` e `available_sugar`.

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
