---
id: 6a3e2b9b7eeca084dd4fb857
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Benvenuto in cucina! Proprio come uno chef si affida alle ricette per preparare un pasto, anche tu ti affiderai alle funzioni per buildare il software. In questo workshop imparerai a scrivere funzioni Python riutilizzabili e a padroneggiare l'ambito, comprendendo esattamente quali ingredienti (variabili) sono disponibili nel tuo codice in ogni momento.

Per iniziare, dovresti definire alcuni dati di base per il tuo inventario di cucina. Crea tre variabili: `available_eggs`, `available_flour` e `available_sugar`, e assegna loro rispettivamente i valori `1`, `2` e `3`.

# --hints--

Dovresti definire una variabile chiamata `available_eggs` e assegnarle il valore `1`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

Dovresti definire una variabile chiamata `available_flour` e assegnarle il valore `2`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

Dovresti definire una variabile chiamata `available_sugar` e assegnarle il valore `3`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
