---
id: 6a3e5475c35dbc138d77c6f4
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora che hai calcolato il totale, dovresti mostrare l'inventario. Ricorda dalle lezioni precedenti che puoi usare la funzione `print()` con le **f-string** per includere variabili nel tuo output.

All'interno della tua funzione `check_kitchen_stock`, sotto la variabile `total_items`, aggiungi queste quattro istruzioni di stampa esatte:

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

Dovresti avere l'istruzione `print(f'The kitchen has {total_items} total items:')` all'interno della tua funzione.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

L'istruzione `print(f'The kitchen has {total_items} total items:')` dovrebbe essere la prima chiamata a `print()` all'interno della tua funzione.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

Dovresti avere l'istruzione `print(f'- {available_eggs} eggs')` all'interno della tua funzione.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

Dovresti avere l'istruzione `print(f'- {available_flour} flour')` all'interno della tua funzione.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

Dovresti avere l'istruzione `print(f'- {available_sugar} sugar')` all'interno della tua funzione.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
