---
id: 6a3e6ad6b9554e31dcc30cd6
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Se non hai abbastanza uova, non dovresti provare a cucinare! Ti serve un modo per fermare l'esecuzione della funzione.

Puoi usare la parola chiave `return` per questo. Quando una funzione incontra un'istruzione `return`, smette di eseguire immediatamente e restituisce il valore al punto in cui è stata chiamata.

Per esempio:

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

All'interno del tuo blocco `if`, sotto l'istruzione di stampa, usa la parola chiave `return` per restituire il valore di `available_eggs`.

# --hints--

Dovresti restituire `available_eggs` all'interno della tua istruzione `if`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")
`)
})
```

Dovresti restituire `available_eggs` dopo l'istruzione `print` all'interno del tuo blocco `if`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs").find_ifs()[0]
assert node.is_ordered("print('The kitchen does not have enough eggs.')", "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
