---
id: 6a3ea8a55317ba7a08e21d5b
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Ora puoi finalmente rispondere alla domanda di prima: perché il numero globale di uova è rimasto a `1`? Quando una funzione restituisce un valore, non sovrascrive automaticamente le tue variabili esistenti. Devi salvare esplicitamente quel nuovo valore!

Stai usando la variabile `available_eggs` all'interno della funzione `make_fried_egg`, ma non è la stessa variabile definita all'inizio del tuo codice. Quella dentro la funzione è una variabile locale che esiste solo mentre la funzione viene eseguita. Quando la funzione termina, la variabile locale viene scartata e la variabile globale `available_eggs` rimane invariata.

Per aggiornare il tuo inventario reale in cucina, sovrascrivi la variabile globale `available_eggs` con il nuovo numero di uova restituito dalla tua funzione `make_fried_egg`.

Alla fine del tuo codice, chiama la funzione `make_fried_egg`, passando `available_eggs` come argomento, e assegna il risultato di nuovo alla variabile globale `available_eggs`.

# --hints--

Dovresti riassegnare la variabile `available_eggs` al risultato di `make_fried_egg(available_eggs)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
