---
id: 6a3e4ae26f2b7e4b5b026cf8
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

È il momento di creare una funzione per controllare cosa hai in cucina. In Python, definisci le funzioni usando la parola chiave `def`, seguita dal nome della funzione, dalle parentesi tonde `()` e da due punti `:`.

Ecco un esempio di come definire una funzione semplice:

```py
def my_function():
    pass
```

Definisci una funzione vuota chiamata `check_kitchen_stock()`. Usa la parola chiave `pass` all'interno del corpo della funzione per ora, così Python non lancerà un errore di sintassi.

Ricorda che Python si basa sull'indentazione per definire i blocchi di codice. Quindi assicurati di indentare il corpo della funzione.

# --hints--

Dovresti definire una funzione chiamata `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
