---
id: 6a3e58891e1ff13d07f4ab48
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Python legge il blocco `def` e memorizza la funzione, ma non esegue immediatamente il codice al suo interno.

Per eseguire il codice e vedere il tuo output, devi **chiamare** (o invocare) la funzione. Fai questo digitando il nome della funzione seguito dalle parentesi tonde.

Per esempio:

```py
my_function()
```

Alla fine del tuo codice, esterno alla funzione, chiama la tua funzione `check_kitchen_stock`.

# --hints--

Dovresti chiamare la funzione `check_kitchen_stock` alla fine del tuo codice. Assicurati che non sia indentata, così viene eseguita globalmente.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
