---
id: 6a3e6fd08125dfa615cb431b
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Ora che la tua funzione `use_eggs` è completa, puoi provarla. Quando chiami una funzione che richiede **parametri**, passi i valori specifici all'interno delle parentesi tonde. Questi valori sono noti come **argomenti**.

Ecco un esempio di come chiamare una funzione con argomenti:

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

In questo esempio, le prime due righe sono la definizione della funzione, dove `a` e `b` sono i parametri. L'ultima riga, `calculate_sum(3, 1)`, è la chiamata della funzione. I valori `3` e `1` sono gli argomenti. Vengono associati ai parametri in base alla posizione, quindi `3` viene assegnato a `a` e `1` viene assegnato a `b`.

Alla fine del tuo codice, chiama la funzione `use_eggs` e passale due argomenti: la variabile globale `available_eggs` e il numero `1` (che rappresenta il numero di uova che vuoi usare).

# --hints--

Dovresti chiamare la funzione `use_eggs` e passarle `available_eggs` e `1` come argomenti.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
