---
id: 6a3e7284e89ce3367e709331
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Prima di preparare un uovo fritto, devi prima controllare se hai gli ingredienti necessari. Puoi farlo confrontando la quantità disponibile di uova con quella richiesta dalla ricetta.

All'interno della tua `make_fried_egg` funzione, sostituisci la parola chiave `pass` creando una variabile chiamata `has_enough_eggs` e assegnandole il risultato del controllo se `available_eggs` è maggiore o uguale a `1`.

# --hints--

Dovresti assegnare a `has_enough_eggs` il risultato del controllo se `available_eggs` è maggiore o uguale a `1`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_variables("has_enough_eggs")[0]
assert node.is_equivalent("has_enough_eggs = available_eggs >= 1") or node.is_equivalent("has_enough_eggs = 1 <= available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
