---
id: 6a3e73eabddd614f91cc235a
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Ora risolverai il mistero di prima! Quando la tua funzione `use_eggs` ha restituito il numero rimanente di uova, non hai salvato quel nuovo valore da nessuna parte. È semplicemente scomparso nel nulla!

Per aggiornare davvero il tuo inventario, devi catturare quel valore restituito assegnandolo a una variabile.

All'interno della tua funzione `make_fried_egg`, usa un'istruzione `if/else` per controllare la variabile `has_enough_eggs`.

- Se `has_enough_eggs` è `True`, riassegna la tua variabile `available_eggs` al risultato di `use_eggs(available_eggs, 1)` e stampa il messaggio `Made a fried egg. Yummy!`.

- Altrimenti, all'interno di un blocco `else`, stampa il messaggio `Could not make a fried egg. Not enough eggs!`.

# --hints--

Dovresti avere un'istruzione `if` e usare `has_enough_eggs` come condizione.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

All'interno del tuo blocco `if`, dovresti riassegnare `available_eggs` al risultato di `use_eggs(available_eggs, 1)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

Dovresti stampare il messaggio `Made a fried egg. Yummy!` quando `has_enough_eggs` è `True`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

Dovresti avere un blocco `else` che stampa il messaggio `Could not make a fried egg. Not enough eggs!` quando `has_enough_eggs` è `False`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
