---
id: 6a3e67fc5ef4045201c28219
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

All'interno della tua funzione `use_eggs`, devi verificare se hai abbastanza uova per cucinare. Puoi farlo usando un'istruzione `if`.

Un'istruzione `if` valuta una condizione. Se la condizione è vera, il codice all'interno del blocco `if` viene eseguito.

Sostituisci la parola chiave `pass` con un'istruzione `if` che verifica se `eggs_to_use` è maggiore (`>`) di `available_eggs`. All'interno del blocco `if`, stampa la stringa `The kitchen does not have enough eggs.`.

# --hints--

Dovresti avere un'istruzione `if` che verifica se `eggs_to_use > available_eggs`.

```js
({
  test: () => runPython(`
  cond = _Node(_code).find_function("use_eggs").find_ifs()[0].find_conditions()[0]
  assert cond.is_equivalent("eggs_to_use > available_eggs") or cond.is_equivalent("available_eggs < eggs_to_use")`)
})
```

Dovresti stampare `The kitchen does not have enough eggs.` all'interno della tua istruzione `if`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_call("print('The kitchen does not have enough eggs.')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
