---
id: 6a3e70b94f67d9ee3ecd8486
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Hai appena chiamato `use_eggs(available_eggs, 1)` per usare un uovo. Questo significa che ora dovresti avere `0` uova rimaste in cucina, giusto?

Lo scoprirai! In fondo al tuo codice, usa la funzione `print()` per stampare il valore della variabile `available_eggs`.

Quando esegui il codice, noterai qualcosa di sorprendente: l'output è ancora `1`! Perché il numero di uova non è cambiato? Esplorerai questo mistero nei prossimi passaggi.

# --hints--

Dovresti stampare la variabile `available_eggs` alla fine del tuo codice.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
