---
id: 6a3e71a0a3a44260e675b157
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Per risolvere il mistero del perché il conteggio delle uova non è cambiato e cucinare davvero la colazione, crea una nuova funzione. Pensa a questa funzione come alla tua ricetta, una che chiama altre funzioni per portare a termine il lavoro.

Sotto la tua istruzione `print(available_eggs)`, definisci una funzione vuota chiamata `make_fried_egg` che prende `available_eggs` come suo unico parametro.

Usa la parola chiave `pass` all'interno del corpo per ora.

# --hints--

Dovresti definire una funzione chiamata `make_fried_egg`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

La tua funzione `make_fried_egg` dovrebbe avere `available_eggs` come parametro.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
