---
id: 6a3e6c0eb5d050ce07781709
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Se la condizione `if` valuta a `False` (dove significa che *hai* abbastanza uova), la funzione salterà il blocco `if` e continuerà a eseguire il resto del codice.

Ora devi elaborare i tuoi dati e restituire un nuovo valore calcolato.

All'esterno del tuo blocco `if` (allo stesso livello di indentazione della parola chiave `if`), usa una chiamata a `print()` e una f-string per stampare `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`.

Poi, nella riga successiva, `return` il numero rimanente di uova sottraendo `eggs_to_use` da `available_eggs`.

# --hints--

Dovresti avere un'istruzione print che stampa `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'` nella tua funzione `use_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

Dovresti restituire il numero rimanente di uova sottraendo `eggs_to_use` da `available_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

Dovresti restituire `available_eggs - eggs_to_use` dopo l'istruzione `print`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
