---
id: 6a3e5ccb99fd6053af9032a1
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Ora che sai cosa c'è nella tua cucina, puoi usare quegli elementi per cucinare. Per farlo, crea una nuova funzione che accetta parametri. Pensa ai parametri come a segnaposto per i dati che passi alla funzione quando la chiami.

Ecco una funzione che accetta due parametri, `a` e `b`, e stampa la loro somma nel terminale.

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

Sotto la tua funzione `check_kitchen_stock`, definisci una funzione vuota chiamata `use_eggs` che accetta due parametri: `available_eggs` e `eggs_to_use`, in questo ordine.

Nota che se dai a un parametro lo stesso nome di una variabile globale (come `available_eggs`), la funzione userà il suo parametro locale `available_eggs` invece di quello globale. Questo è noto come name shadowing perché la variabile locale "oscura" quella globale.

Usa la parola chiave `pass` all'interno del corpo per ora.

# --hints--

Dovresti definire una funzione chiamata `use_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

La tua funzione `use_eggs` dovrebbe avere `available_eggs` e `eggs_to_use` come parametri.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
