---
id: 587d825a367417b2b2512c89
title: 퀵소트 알고리즘 구현하기
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

이 실습에서는 정렬 알고리즘인 퀵소트를 구현할 것입니다.

퀵소트는 배열을 정렬하는 효율적이고 재귀적인 분할 정복 방법입니다. 이 방법에서는 원본 배열에서 피벗 값을 선택합니다. 배열은 피벗 값보다 작은 값과 큰 값으로 이루어진 두 개의 하위 배열로 분할됩니다. 그런 다음 두 하위 배열에 대해 퀵소트 알고리즘을 재귀적으로 호출한 결과를 합칩니다. 이 과정은 빈 배열이나 한 개의 항목만 있는 배열이라는 기본 사례에 도달할 때까지 계속되며, 이 경우 해당 배열을 반환합니다. 재귀 호출이 풀리면서 정렬된 배열이 반환됩니다.

피벗 값 선택이 중요하지만, 여기서는 어떤 피벗이든 상관없습니다. 간단하게 첫 번째나 마지막 요소를 사용할 수 있습니다.

퀵소트는 평균적으로 *O(nlog(n))* 성능을 제공하는 매우 효율적인 정렬 방법입니다. 또한 구현하기도 비교적 쉽습니다. 이러한 특성 덕분에 널리 사용되고 유용한 정렬 방법입니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 정수 배열을 입력으로 받는 `quicksort` 함수가 있어야 합니다.
1. `quicksort` 함수는 입력과 동일한 정수를 포함하지만 가장 낮은 값부터 높은 값 순서로 정렬된 배열을 반환해야 합니다.
1. `quicksort` 함수는 배열을 정렬하기 위해 재귀적으로 자기 자신을 호출해야 합니다.

# --hints--

`quicksort` 함수가 있어야 합니다.

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` 함수는 순서만 바뀌었을 뿐 변경되지 않은 배열을 반환해야 합니다.

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` 함수는 오름차순으로 정렬된 배열을 반환해야 합니다.

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` 함수는 내장 `.sort()` 메서드를 사용하지 않아야 합니다.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
