---
id: 698c5710136539797f296098
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

在 `p` 元素下方，添加一个 `audio` 元素，为它添加一个 `src` 属性，值为 `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`，并添加 `loop` 和 `controls` 属性。

# --hints--

你应该有第二个 `audio` 元素。

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

新建的 `audio` 元素应位于 `p` 元素下方。

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

你的 `audio` 元素应具有 `src` 属性。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

你的 `audio` 元素应具有值为 `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3` 的 `src` 属性。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

你的 `audio` 元素应具有 `loop` 属性。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

你的 `audio` 元素应具有 `controls` 属性。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
