---
id: 698c513508507ec8bb820227
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

要指定音频的媒体资源，你需要为你的 `audio` 元素添加 `src` 属性。

为你的 `src` 属性添加值为 `https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3`。

# --hints--

你的 `audio` 元素应具有 `src` 属性。

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('src'));
```

你的 `audio` 元素应具有值为 `https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3` 的 `src` 属性。

```js
const audio = document.querySelector('audio');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio></audio>
--fcc-editable-region--

</body>
</html>
```
