---
id: 698c4f4e965bd853ad5084e1
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

在 `h1` 下方，添加一个带有文本 `Can't Stay Down` 的 `h2` 元素，这是第一首歌的标题。

在 `h2` 元素下方添加一个带有文本 `Artist: Quincy Larson` 的 `p` 元素。

# --hints--

你应该有一个 `h2` 元素。

```js
assert.exists(document.querySelector('h2'));
```

你的 `h2` 元素应该放在 `h1` 元素下面。

```js
assert.exists(document.querySelector('h1 + h2'));
```

你的 `h2` 元素应包含文本 `Can't Stay Down`。

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

你应该有一个 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

你的 `p` 元素应包含文本 `Artist: Quincy Larson`。

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
