---
id: 698c58745b73289bb6c89ece
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

最后一首歌的标题是 `Scratching the Surface`，艺术家是 Quincy Larson，使用的文件是 `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`。

为你的音乐播放器添加最后一首歌曲，遵循你为前两首歌曲使用的相同模式。

# --hints--

你应该有第三个 `h2` 元素。

```js
assert.exists(document.querySelector('h2:nth-of-type(3)'));
```

你的第三个 `h2` 元素应包含文本 `Scratching the Surface`。

```js
const h2 = document.querySelector('h2:nth-of-type(3)');
assert.strictEqual(h2?.textContent.trim(), 'Scratching the Surface');
```

你应该有一个 `p` 元素。

```js
assert.exists(document.querySelector('p:nth-of-type(3)'));
```

你的 `p` 元素应包含文本 `Artist: Quincy Larson`。

```js
const p = document.querySelector('p:nth-of-type(3)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

你应该有第三个 `audio` 元素。

```js
assert.exists(document.querySelector('audio:nth-of-type(3)'));
```

新建的 `audio` 元素应位于 `p` 元素下方。

```js
assert.exists(document.querySelector('p:nth-of-type(3) + audio:nth-of-type(3)'));
```

你的 `audio` 元素应具有 `src` 属性。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('src'));
```

你的 `audio` 元素应具有值为 `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3` 的 `src` 属性。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

你的 `audio` 元素应具有 `loop` 属性。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('loop'));
```

你的 `audio` 元素应具有 `controls` 属性。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

  <h2>Scratching the Surface</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3" loop controls></audio>
  
</body>
</html>
```
