---
id: 698c4fe24f1f1a17be74a23e
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

接下来，在 `p` 元素下方创建一个 `audio` 元素。在接下来的几个步骤中，你将为音频元素添加必要的属性以播放音乐。 

# --hints--

你应该有一个 `audio` 元素。

```js
assert.exists(document.querySelector('audio'));
```

新建的 `audio` 元素应位于 `p` 元素下方。

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
