---
id: 698c55e7aaacfae48a227459
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

在 `audio` 元素下方为你的页面添加一个带有文本 `Cruising for a Musing` 的新 `h2` 元素。

在 `h2` 下方，添加一个带有文本 `Artist: Quincy Larson` 的 `p` 元素。

# --hints--

你应该有第二个 `h2` 元素。

```js
assert.exists(document.querySelector('h2:nth-of-type(2)'));
```

你的第二个 `h2` 元素应包含文本 `Cruising for a Musing`。

```js
const h2 = document.querySelector('h2:nth-of-type(2)');
assert.strictEqual(h2?.textContent.trim(), 'Cruising for a Musing');
```

你应该有了第二个 `p` 元素。

```js
assert.exists(document.querySelector('p:nth-of-type(2)'));
```

你的第二个 `p` 元素应包含文本 `Artist: Quincy Larson`。

```js
const p = document.querySelector('p:nth-of-type(2)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
