---
id: 698c5229e47d57c6dfaa91bd
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

`controls` 属性为 `audio` 元素提供播放控件，包括暂停/恢复播放、定位和卷控件。 

`controls` 属性是一个布尔属性，不需要值。

为你的 `audio` 元素添加 `controls` 属性。

现在你应该能看到页面上显示了 `audio` 元素。

# --hints--

你的 `audio` 元素应具有 `controls` 属性。

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
