---
id: 698c50535b54b2d75fa985fb
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

`loop` 属性将在播放完成后重新启动音频。想象一个重复播放的互联网表情包。省略 `loop` 属性将使音频只播放一次。

`loop` 属性是一个布尔属性，不需要值。

为你的 `audio` 元素添加 `loop` 属性。

# --hints--

你的 `audio` 元素应具有 `loop` 属性。

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
