---
id: 6ac09f8420b4b67e075d2cd2
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Nesta oficina, você vai criar seu primeiro programa em Python.

Este é seu editor de código, onde você vai escrever código Python. A linha 1 está destacada. Vá até ela e digite:

```py
'Hello'
```

Certifique-se de começar no início da linha e não esqueça as aspas.

Quando terminar, clique no botão "verificar seu código" para ver se está correto.

# --hints--

Você deve digitar `'Hello'`. Certifique-se de começar no início da linha e não esqueça as aspas.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
