---
id: 6ac0b35f3b29a27f7e207295
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

No passo anterior, você criou a variável `sentence`. Agora imprima seu valor no terminal digitando `print(sentence)` na linha destacada.

Com isso, você completou seu primeiro programa em Python.

# --hints--

Você deve ter `print(sentence)` no seu código. Certifique-se de começar no início da linha.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
