---
id: 6ac0a4463662d9738c393e21
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

No passo anterior, você criou uma **string**: uma sequência de caracteres envolta em aspas simples ou aspas duplas. Na programação, strings são usadas para representar texto.

O **terminal** é onde você pode ver a saída do seu código.

Para conseguir ver uma saída, transforme sua linha de código existente em:

```py
print('Hello')
```

Certifique-se de começar no início da linha.

# --hints--

Você deve ter `print('Hello')` no seu código. Certifique-se de começar no início da linha.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
