---
id: 6ac0ae563f6de70930f47ef7
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Ao escrever código, você frequentemente vai querer reutilizar ou referenciar valores existentes no seu código. Variáveis permitem que você faça isso.

Para atribuir um valor a uma variável, escreva o nome da variável seguido do operador de atribuição (`=`) e do valor que você quer atribuir a essa variável:

```py
sentence = 'I am learning Python'
```

Vá para a linha destacada, crie uma variável chamada `sentence` e atribua a ela a string `I am learning Python`.

# --hints--

Você deve ter `sentence = 'I am learning Python'` no seu código. Certifique-se de começar no início da linha.

```js
({ test: () => runPython(`
import re
target_var = _Node(_code).find_variable("sentence")
pattern = r"""sentence\\s*=\\s*(['"])\\s*i\\s+am\\s+learning\\s+python\\s*\\1\\s*"""
assert re.fullmatch(pattern, str(target_var).strip(), flags=re.IGNORECASE)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
