---
id: 6ac0aaba914d81b32567c076
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

`print` é uma função embutida que você pode usar para enviar saída para o terminal. Como você pode ver, o texto `Hello` agora está visível no terminal.

Agora, imprima o texto `freeCodeCamp` no terminal indo até a linha destacada e digitando:

```py
print('freeCodeCamp')
```

Certifique-se de começar no início da linha.

# --hints--

Você deve ter `print('freeCodeCamp')` no seu código. Certifique-se de começar no início da linha.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
