---
id: 6a57c331ae94fa4631b6086d
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

تحقق من أن مفتاح التبديل يعمل. أولًا، نفذ استدعاء `console.log()` مع `collection[102].onDisplay`. بعد ذلك، نفذ استدعاء `toggleDisplayStatus(102)`. أخيرًا، نفذ استدعاء `console.log()` مع `collection[102].onDisplay` مرة أخرى.

# --hints--

يجب أن تنفذ استدعاء `toggleDisplayStatus` مع `102`.

```js
assert.match(__helpers.removeJSComments(code), /toggleDisplayStatus\s*\(\s*102\s*\)/);
```

يجب أن تسجل قيمة `onDisplay` للقطعة الثانية، ثم تنفذ استدعاء `toggleDisplayStatus`، ثم تسجلها مرة أخرى.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*onDisplay|\[\s*["']onDisplay["']\s*\])\s*\)[\s\S]*toggleDisplayStatus\s*\(\s*102\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*onDisplay|\[\s*["']onDisplay["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
