---
id: 6a57c331ae94fa4631b60870
title: الخطوة 17
challengeType: 1
dashedName: step-17
---

# --description--

أخيرًا، أنشئ ملخصًا قابلًا للقراءة لأي قطعة أثرية. ابدأ بالتعامل مع الحالة التي تكون فيها القطعة الأثرية مفقودة.

أنشئ دالة `buildSummary` تأخذ معلمة `id`. استخدم `collection[id]` للبحث عن القطعة الأثرية، وخزنها في متغير باسم `artifact`. إذا لم تُعثر على القطعة الأثرية، أرجع السلسلة النصية `"Artifact not found"`.

# --hints--

يجب أن يكون لديك دالة باسم `buildSummary` مع معلمة باسم `id`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary);
const params = explorer.allFunctions.buildSummary.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

يجب تعيين `collection[id]` إلى متغير باسم `artifact`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary?.variables.artifact);
assert.isTrue(explorer.allFunctions.buildSummary?.variables.artifact?.value.matches("collection[id]"));
```

يجب أن تُرجع `buildSummary` السلسلة النصية `"Artifact not found"` عندما لا تطابق أي قطعة أثرية المعرف `id`.

```js
assert.strictEqual(buildSummary(999), "Artifact not found");
```

يجب ألا تُرجع `buildSummary` السلسلة النصية `"Artifact not found"` عندما توجد القطعة الأثرية.

```js
assert.notStrictEqual(buildSummary(101), "Artifact not found");
assert.notStrictEqual(buildSummary(102), "Artifact not found");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
