---
id: 6a57c331ae94fa4631b60868
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

يقوم القائمون على التنظيم بوضع وسم على القطع الأثرية لتسهيل البحث عنها، ويجب ألا يُضاف الوسم مرتين. قبل إضافة وسم جديد، يمكنك استخدام الطريقة `includes()` للتحقق مما إذا كان موجودًا بالفعل، إذ تُرجع `true` عندما تحتوي المصفوفة على القيمة.

أنشئ دالة باسم `addTag` تأخذ معلمتين هما `id` و `tag`. ابحث عن القطعة الأثرية باستخدام `collection[id]`. إذا كانت موجودة ولم تتضمن مصفوفة `tags` الخاصة بها الوسم `tag`، أضف `tag` إلى `artifact.tags`.

# --hints--

يجب أن يكون لديك دالة باسم `addTag` بمعلمتين هما `id` و `tag`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

استدعاء `addTag` مع وسم جديد يجب أن يضيف ذلك الوسم إلى مصفوفة `tags` الخاصة بالقطعة الأثرية.

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

استدعاء `addTag` مع وسم موجود بالفعل في القطعة الأثرية يجب ألا يضيف نسخة مكررة.

```js
const tagsBefore = [...collection[102].tags];
addTag(102, "bronze");
assert.deepEqual(collection[102].tags, tagsBefore);

const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

يجب أن تضيف `addTag` الوسم إلى أي قطعة أثرية يشير إليها `id`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
addTag(457, "test-tag");
assert.deepEqual(collection[457].tags, ["test-tag"]);
```

يجب ألا ترجع `addTag` خطأ أو تغير المجموعة عندما لا تطابق أي قطعة أثرية الـ `id`.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => addTag(999, "test"));
assert.strictEqual(JSON.stringify(collection), before);
```

يجب أن تضيف `tag` إلى مصفوفة `tags` الموجودة للقطعة الأثرية، بدلًا من استبدال المصفوفة.

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const tagsBefore = collection[458].tags;
const spy = __helpers.spyOn(tagsBefore, "push");
try {
  addTag(458, "pushed");
  assert.deepEqual(spy.calls, [["pushed"]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].tags, tagsBefore);
assert.deepEqual(collection[458].tags, ["pushed"]);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
