---
id: 6a57c331ae94fa4631b60862
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

يتتبع كل قطعة أثرية أيضًا الأماكن التي زارتها، وكيف توصف، وما إذا كانت معروضة.

أضف ثلاث خصائص أخرى إلى القطعة الأثرية المخزنة تحت المفتاح `101`:

- `locations`: مصفوفة تحتوي على كائنين، `{ gallery: "Hall A", year: 2020 }` و `{ gallery: "Hall C", year: 2024 }`
- `tags`: مصفوفة تحتوي على السلاسل النصية `"gold"` و `"egypt"`
- `onDisplay`: القيمة المنطقية `true`

# --hints--

يجب أن تكون خاصية `locations` للقطعة الأثرية مصفوفة من كائنات الموقع المعطاة.

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

يجب أن تكون خاصية `tags` للقطعة الأثرية هي `["gold", "egypt"]`.

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

يجب أن تكون خاصية `onDisplay` للقطعة الأثرية هي `true`.

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
