---
id: 6a57c331ae94fa4631b60866
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

أحيانًا يطلب الموظفون قطعة أثرية لا يملكها المتحف. إذا لم يُخزن أي أثر تحت ذلك الـ `id`، فإن `collection[id]` تكون `undefined`، وقراءة `.title` منها ستؤدي إلى رمي خطأ.

حدّث جملة الإرجاع لتستخدم شرطًا. أرجع `artifact.title` إذا وُجد الأثر، وأرجع السلسلة النصية `"Artifact not found"` إذا لم يوجد.

# --hints--

يجب أن تُعيد الدالة `getArtifactTitle(101)` القيمة `"Golden Mask"` كما هي.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

يجب أن تُرجع الدالة `getArtifactTitle` القيمة `"Artifact not found"` إذا لم يتطابق أي قطعة أثرية مع الـ `id`.

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

يجب أن تُرجع `getArtifactTitle` عنوان أي أثر تجده، وليس فقط `101`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
