---
id: 6a57c331ae94fa4631b60865
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

يبحث الموظفون عن القطع الأثرية عشرات المرات يوميًا، لذا يحتاج المدقق إلى طريقة بحث قابلة لإعادة الاستخدام. يمكن أيضًا أن تأخذ الأقواس المربعة متغيرًا، لذا `collection[id]` تُرجع القطعة الأثرية التي يشير إليها هذا `id`.

أنشئ دالة باسم `getArtifactTitle` تأخذ معلمة `id`. استخدم `collection[id]` للحصول على القطعة الأثرية، خزّنها في متغير باسم `artifact`، وأرجع `artifact.title`.

# --hints--

يجب أن يكون لديك دالة باسم `getArtifactTitle` تحتوي على معلمة باسم `id`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle);
const params = explorer.allFunctions.getArtifactTitle.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

يجب تعيين `collection[id]` إلى متغير باسم `artifact`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle?.variables.artifact);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.variables.artifact?.value.matches("collection[id]"));
```

يجب أن تُرجع الدالة `getArtifactTitle` القيمة `artifact.title`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.hasReturn("artifact.title"));
```

يجب أن تُرجع `getArtifactTitle(101)` القيمة `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

يجب أن تُرجع `getArtifactTitle(102)` القيمة `"Bronze Tablet"`.

```js
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
