---
id: 6a57c331ae94fa4631b60864
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

تبدأ عملية التدقيق بقراءة ما هو مسجل بالفعل. بما أن كل قطعة محفوظة تحت معرفها، يمكنك الوصول إلى واحدة بوضع ذلك المعرف داخل أقواس مربعة، مثل `collection[101]`.

استدعِ `console.log()` مع `collection[101].title` لتسجيل عنوان القطعة الأولى. ثم استدعِ `console.log()` مع `collection[101].curator.name` لتسجيل اسم القائم على التنسيق.

# --hints--

يجب تسجيل `title` القطعة الأولى في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)/);
```

يجب تسجيل اسم القائم على التنسيق `name` للقطعة الأولى في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

يجب تسجيل `title` أولًا، واسم القائم على التنسيق `name` ثانيًا.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
