---
id: 6a57c331ae94fa4631b60869
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

تؤكد أبحاث جديدة أن القناع الذهبي كان ملكيًا. أضف الوسم `"royal"` إلى القطعة الأثرية التي لها `id` يساوي `101`. ثم استدعِ `console.log()` مع `collection[101].tags` لتسجيل الوسوم الخاصة بتلك القطعة الأثرية في وحدة التحكم.

# --hints--

يجب استدعاء `addTag` مع `101` و`"royal"`.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)/);
```

يجب تسجيل `tags` الخاصة بالقطعة الأثرية الأولى في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

يجب تسجيل `tags` بعد إضافة الوسم الجديد.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

يجب أن تتضمن `tags` الخاصة بالقطعة الأثرية الأولى `"royal"`.

```js
assert.include(collection[101].tags, "royal");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
