---
id: 6a57c265abbc95dfebc05098
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة العمل، ستبني مدقق أصول القطع الأثرية. ستخزن قطع المتحف في كائن وتكتب دوالًا للبحث عنها، ووضع وسم عليها، ونقلها بين القاعات، وتلخيص تاريخها.

يُخزن كل أثر تحت معرفه. أنشئ كائنًا باسم `collection` وأضف إليه مدخلاً. اجعل المفتاح `101` والقيمة كائنًا يحتوي على `title` بقيمة `"Golden Mask"` و`category` بقيمة `"Ceremonial"`.

# --hints--

يجب إعلان متغير باسم `collection`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

يجب أن يكون `collection` كائنًا.

```js
assert.isObject(collection);
```

يجب أن يحتوي `collection` على مدخل واحد.

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

يجب أن تكون القطعة الأثرية المخزنة تحت `101` كائنًا.

```js
assert.isObject(collection[101]);
```

يجب أن يحتوي القطعة المخزنة تحت المفتاح `101` على `title` بقيمة `"Golden Mask"`.

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

يجب أن يحتوي القطعة المخزنة تحت المفتاح `101` على `category` بقيمة `"Ceremonial"`.

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
