---
id: 6a57c331ae94fa4631b60863
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

يحتوي المتحف على أكثر من قطعة أثرية مسجلة. أضف قطعة أثرية ثانية إلى `collection` تحت المفتاح `102`، باتباع نفس الشكل كالقطعة الأولى.

استخدم هذه القيم:

```js
{
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
}
```

# --hints--

يجب أن يحتوي `collection` على مدخلين.

```js
assert.strictEqual(Object.keys(collection).length, 2);
```

يجب أن تطابق القطعة الأثرية المخزنة تحت `102` الكائن الموصوف.

```js
assert.deepEqual(collection[102], {
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
});
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
--fcc-editable-region--
  
--fcc-editable-region--
};
```
