---
id: 6a57c331ae94fa4631b60861
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

لكل قطعة أثرية شخص مسؤول عنها.

أضف خاصية `curator` إلى القطعة الأثرية المخزنة تحت المفتاح `101`، واجعل قيمتها كائنًا يحتوي على `id` بقيمة `201` و`name` بقيمة `"Earl Sinclair"`.

# --hints--

يجب أن تحتوي القطعة الأثرية المخزنة تحت `101` على خاصية `curator`.

```js
assert.property(collection[101], "curator");
```

يجب أن يكون `curator` كائنًا.

```js
assert.isObject(collection[101].curator);
```

يجب أن يحتوي `curator` على `id` بقيمة `201`.

```js
assert.strictEqual(collection[101].curator.id, 201);
```

يجب أن يحتوي `curator` على خاصية `name` بقيمة `"Earl Sinclair"`.

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
