---
id: 6735a7370e0ae93a4577c689
title: Projetar um Layout de Jornal
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

Neste laboratório, você irá projetar um layout de jornal usando CSS Grid, incluindo conceitos como linhas da grade e colunas da grade.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório. 

**Histórias de Usuário:**

1. Sua página deve conter um elemento `main` com a classe `newspaper-layout`.
1. O `.newspaper-layout` deve incluir um `header` com a classe `title` contendo um `h1` para exibir o nome do jornal.
1. O `.newspaper-layout` deve incluir um `article` com a classe `feature-article` para a matéria principal.
1. O `.feature-article` deve incluir um elemento `h2` para o título do artigo seguido por um elemento `p` para o conteúdo do artigo.
1. Você deve adicionar mais três elementos `article` para artigos menores, com as classes `small-article1`, `small-article2` e `small-article3`.
1. Cada um dos artigos menores deve incluir um elemento `h3` para o título do artigo seguido por um elemento `p` para o conteúdo do artigo.
1. O `.newspaper-layout` deve incluir um elemento `article` com a classe `secondary-article` para uma seção adicional de notícias.
1. O `.secondary-article` deve incluir um elemento `h2` para o título do artigo seguido por um elemento `p` para o conteúdo do artigo.
1. A `.newspaper-layout` deve incluir um `figure` com a classe `cover-image` para exibir uma imagem que representa o conteúdo do jornal.
1. Os elementos com as classes `title`, `feature-article`, `secondary-article`, `cover-image`, `small-article1`, `small-article2` e `small-article3` devem ter a propriedade `grid-area` definida com o mesmo nome da classe.
1. Seu `.newspaper-layout` deve usar CSS Grid com a propriedade `grid-template-areas` para definir o arranjo das seções:
   - O `.title` deve abranger toda a linha superior.
   - O `.feature-article` deve abranger duas colunas na segunda linha, com a `.cover-image` na terceira coluna.
   - O `.secondary-article` deve abranger duas colunas na terceira linha, com a `.cover-image` na terceira coluna.
   - Os três pequenos artigos devem aparecer na quarta linha.
1. O `.newspaper-layout` deve ter uma propriedade `grid-template-columns` que divide o espaço em três colunas iguais.
1. Você deve definir a propriedade `grid-template-rows` do `.newspaper-layout` como `auto` para a primeira linha e dividir o espaço restante em partes iguais para as outras linhas.
1. Você deve adicionar um espaçamento entre os itens da grade.
1. Garanta que a imagem dentro de `.cover-image` se ajuste ao espaço designado definindo seu `max-width` para `100%`.

**Nota:** Certifique-se de vincular sua folha de estilos no seu HTML e aplicar seu CSS.

# --hints--

Você deve ter apenas um elemento `main`.

```js
assert.exists(document.querySelector('main'));
```

Seu elemento `main` deve ter uma `class` de `newspaper-layout`.

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

Você deve ter um elemento `header` com a `class` `title` dentro do seu elemento `.newspaper-layout`.

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

Seu elemento `header` deve conter um elemento `h1`.

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

Seu elemento `.title` deve conter algum texto.

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

Você deve ter um elemento `article` para cada artigo com as classes `feature-article`, `small-article1`, `small-article2`, `small-article3` e `secondary-article`, respectivamente.

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

O `.feature-article` deve incluir um elemento `h2` seguido por um elemento `p`.

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`, `.small-article2` e `.small-article3` devem incluir um elemento `h3` seguido por um elemento `p`.

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

O `.secondary-article` deve incluir um elemento `h2` seguido por um elemento `p`.

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

Você deve ter um elemento `figure` com a classe `cover-image` dentro do seu `.newspaper-layout`.

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

Seu `.cover-image` deve conter um elemento `img`.

```js
assert.exists(document.querySelector(".cover-image img"));
```

A propriedade `display` do elemento `.newspaper-layout` deve ser definida como `grid`.

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

Seu elemento `.newspaper-layout` deve ter uma propriedade `grid-template-areas` com `title` abrangendo todas as três colunas da primeira linha.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

A segunda linha do modelo de grade deve conter um `feature-article` que abranja duas colunas e um `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

A terceira linha do template da grade deve conter um `secondary-article` que abranja duas colunas e um `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

A quarta linha do modelo de grade deve conter `small-article1`, `small-article2` e `small-article3`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

Seu seletor `.newspaper-layout` deve usar a propriedade `grid-template-columns` para dividir o espaço em três partes iguais. Lembre-se de que você pode usar a função `repeat` ou definir manualmente cada unidade fracionária.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

Seu seletor `.newspaper-layout` deve ter a propriedade `grid-template-rows` definida como `auto 1fr 1fr 1fr`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

Você deve adicionar um espaçamento entre os itens da grade.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title` deve ter uma `grid-area` de `title`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` deve ter uma `grid-area` de `feature-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` deve ter uma `grid-area` de `cover-image`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article` deve ter uma `grid-area` de `secondary-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1` deve ter uma `grid-area` de `small-article1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2` deve ter uma `grid-area` de `small-article2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3` deve ter uma `grid-area` de `small-article3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

O `img` dentro de `.cover-image` deve ter um `max-width` de `100%`.

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
