---
id: 5f332a88dc25a0fd25c7687a
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Назва кафе — `CAMPER CAFE`. Тож додайте `h1` елемент всередину вашого `main` елемента. Напишіть назву кафе великими літерами, щоб вона виділялася.

# --hints--

У вас має бути відкриваючий тег `<h1>`.

```js
assert.match(code, /<h1>/i);
```

У вас має бути закриваючий тег `</h1>`.

```js
assert.match(code, /<\/h1>/i);
```

Ви не повинні змінювати свій існуючий `main` елемент.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Ваш `h1` елемент має бути вкладений у ваш `main` елемент.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Ваш `h1` елемент має містити текст `CAMPER CAFE`.

```js
assert.equal(document.querySelector("h1")?.innerText.trim(), "CAMPER CAFE");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
