---
id: 5f33294a6af5e9188dbdb8f3
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

У цьому практичному занятті ви потренуєтеся з основами CSS (каскадних таблиць стилів), створюючи меню кафе.

Почнемо з додавання вмісту меню. Додайте елемент `main` всередині існуючого елемента `body`. Він згодом міститиме інформацію про ціни на каву та десерти, які пропонує кафе.

# --hints--

Ваш код має містити відкриваючий тег `<main>`.

```js
assert.match(code, /<main>/i);
```

Ваш код має містити кінцевий тег `</main>`.

```js
assert.match(code, /<\/main>/i);
```

Ви не повинні змінювати свій елемент `body`. Переконайтеся, що випадково не видалили кінцевий тег.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Ваш тег `main` має бути всередині тегу `body`.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
