---
id: 5f332b23c2045fb843337579
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Щоб відвідувачі знали, що кафе було засновано в `2020` році, додайте елемент `p` під елементом `h1` з текстом `Est. 2020`.

# --hints--

Ви повинні мати початковий тег `<p>`.

```js
assert.match(code, /<p>/i);
```

Ви повинні мати кінцевий тег `</p>`.

```js
assert.match(code, /<\/p>/i);
```

Не змінюйте наявний елемент `h1`. Переконайтеся, що не видалили кінцевий тег.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

Елемент `p` повинен бути під елементом `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

Елемент `p` повинен містити текст `Est. 2020`.

```js
assert.equal(document.querySelector("p")?.innerText.trim(), "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
</html>
```

