---
id: 5f344fad8bf01691e71a30eb
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

До цього моменту у вас був обмежений контроль над презентацією та зовнішнім виглядом вашого вмісту. Щоб змінити це, додайте елемент `style` всередині елемента `head`.

# --hints--

У вашому коді має бути `style` елемент.

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
```

Ваш код має містити закриваючий тег `</style>`.

```js
assert.match(code, /<\/style\s*>/i);
```

Ваш елемент `style` має бути вкладений у ваш елемент `head`.

```js
assert.exists(document.querySelector('head > style:not(.fcc-hide-header)'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
