---
id: 5f344fbc22624a2976425065
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Створіть елемент `h2` всередині елемента `section` і додайте йому текст `Coffee`.

# --hints--

У вас має бути відкриваючий тег `<h2>`.

```js
assert.match(code, /<h2\s*>/i);
```

У вас має бути кінцевий тег `</h2>`.

```js
assert.match(code, /<\/h2\s*>/i);
```

Не слід змінювати існуючий елемент `section`. Переконайтеся, що ви не видалили кінцевий тег.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

Ваш елемент `h2` має бути всередині елемента `section`.

```js
const h2 = document.querySelector('h2');
assert.equal(h2.parentElement.tagName, 'SECTION');
```

Ваш елемент `h2` має містити текст `Coffee`.

```js
const h2 = document.querySelector('h2');
assert.equal(h2?.innerText.trim(), 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
