---
id: 5f344f9c805cd193c33d829c
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

У попередніх уроках ви навчились додавати властивості `CSS` і значення так:

```css
element {
 property: value;
}
```

Вирівняйте вміст елемента `h1` по центру, встановивши його властивість `text-align` у значення `center`.

# --hints--

У вашому елементі `h1` має бути селектор `style`.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasSelector);
```

Ваша властивість `text-align` має мати значення `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert.isTrue(hasTextAlign);
```

Ваш селектор `h1` має встановлювати властивість `text-align` у значення `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
