---
id: 5f3477cb2e27333b1ab2b955
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Тепер потрібно прив’язати файл `styles.css`, щоб стилі знову були застосовані. Додайте елемент `link` до елемента `head`. Надайте йому атрибут `rel` зі значенням `"stylesheet"` та атрибут `href` зі значенням `"styles.css"`.

**Note:** The `link` element is a void element, which means it doesn't have a closing tag. Void elements should be written as `<link>` rather than `<link></link>`.

# --hints--

Ваш код повинен мати елемент `link`.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

Елемент `link` є порожнім елементом, він не повинен мати кінцевий тег `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Не змінюйте наявний елемент `head`. Переконайтеся, що не видалили кінцевий тег.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

Ви повинні мати один елемент `link`.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

Елемент `link` має бути в межах елемента `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
