---
id: 5f3477aefa51bfc29327200b
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Ви стилізували три елементи, прописавши CSS всередині тегів `style`. Це працює, але оскільки буде ще багато стилів, краще розмістити всі стилі в окремому файлі та посилатися на нього.

A separate `styles.css` file has been created for you. You can switch between files with the tabs at the top of the editor.

Розпочніть з перезапису створених стилів в файл `styles.css`. Make sure you exclude the opening and closing `style` tags.

# --hints--

Файл `styles.css` повинен мати селектор типу `h1, h2, p`.

```js
assert.match(editableContents.replace(/[\s\n]*/g, ""), /(h1|h2|p),(h1|h2|p),(h1|h2|p){/);
```

Селектор повинен встановлювати властивість `text-align` на `center`.

```js
assert.match(editableContents, /text-align:\s*center;?/);
```

Ваш код не повинен містити теґи `<style>`.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
