---
id: 5f3477ae34c1239cafe128be
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Тепер у вас є три селектори типу з однаковим стилем. Ви можете додати одну й ту саму групу стилів до багатьох елементів, створивши список селекторів. Кожен селектор розділяється комами ось так:

```css
selector1, selector2 {
  property: value;
}
```

Видаліть три існуючі селектори типу і замініть їх одним списком селекторів, який вирівнює текст по центру для елементів `h1`, `h2` і `p`.

# --hints--

Ви повинні використати один селектор типу для всіх трьох елементів, `h1, h2, p`. Обов’язково дотримуйтесь цього порядку.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert.exists(hasSelector);
```

У вашому елементі `style` має бути лише один селектор.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert.equal(selectors.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
