---
id: 5f344fc1520b6719f2e35605
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

У меню буде два розділи: один для кави, інший — для десертів. Додайте елемент `section` всередину елемента `main`, щоб мати місце для всіх доступних кав.

# --hints--

Ви повинні мати відкриваючий тег `<section>`.

```js
assert.match(code, /<section\s*>/i);
```

Ви повинні мати кінцевий тег `</section>`.

```js
assert.match(code, /<\/section\s*>/i);
```

Не змінюйте існуючий елемент `main`. Переконайтеся, що ви не видалили кінцевий тег.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Ваш елемент `section` має бути всередині елемента `main`.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert.strictEqual(section.parentElement, main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
