---
id: 5f3477aefa51bfc29327200b
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Ви стилізували три елементи, написавши CSS всередині тегів `style`. Це працює, але оскільки стилів буде набагато більше, найкраще помістити всі стилі в окремий файл і направити на нього посилання.

Для вас створено окремий файл `styles.css`. Ви можете перемикатися між файлами за допомогою вкладок у верхній частині редактора.

Почніть із переписування стилів, які ви створили, у файл `styles.css`. Переконайтеся, що ви виключили відкриваючі та закриваючі теги `style`.

# --hints--

Ваш файл `styles.css` має містити селектор типу `h1, h2, p`.

```js
assert.match(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, ""), __helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/));
```
 
Ваш селектор має встановити властивість `text-align` у значення `center`.

```js
assert.match(code.slice(code.indexOf('</html>')), /text-align:\s*center;?/);
```

Ваш код не повинен містити тегів `<style>`.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
