---
id: 5f3477ae9675db8bb7655b30
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

У попередньому кроці ви використали <dfn>селектор типу</dfn> для стилізації елемента `h1`. Вирівняйте вміст елементів `h2` та `p` по центру, додавши новий селектор типу для кожного з них до існуючого елемента `style`.

# --hints--

Ви не повинні змінювати існуючий селектор `h1`.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasH1);
```

Ви не повинні додавати новий тег `style`. Додайте нові CSS-правила до існуючого тегу `style`.

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

Ви повинні додати новий селектор `h2`.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert.exists(hasH2);
```

Ви повинні додати новий селектор `p`.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert.exists(hasP);
```

Ваш елемент `h1` повинен мати ідентифікатор `text-align` з назвою `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

Ваш елемент `h2` повинен мати ідентифікатор `text-align` з назвою `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert.equal(h2TextAlign, 'center');
```

Ваш елемент `p` повинен мати ідентифікатор `text-align` з назвою `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert.equal(pTextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1 {
        text-align: center;
      }
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
