---
id: 5f35e5c44359872a137bd98f
title: Крок 24
challengeType: 0
dashedName: step-24
---

# --description--

Оскільки основним продуктом кафе є кава, ви можете використати зображення кавових зерен як фон сторінки.

Видаліть коментар та його вміст всередині селектора типу `body`. Після цього додайте властивість `background-image` і встановіть її значення на `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

# --hints--

Вам слід видалити закоментовану властивість `background-color`.

```js
assert.notMatch(code, /\/\*\s*background-color:\s*burlywood;?\s*\*\//i);
```

Ваш селектор `body` не повинен містити коментарів.

```js
assert.notMatch(code, /body\s*{\s*\/\*/i);
```

Вам слід встановити властивість `background-image` у значення `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert.isTrue(hasBackground)
```

Ваш елемент `body` має мати фонове зображення кавових зерен.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');

assert.equal(bodyBackground, `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  /*
  background-color: burlywood;
  */
--fcc-editable-region--
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
