---
id: 5f3c866d5414453fc2d7b480
title: Крок 27
challengeType: 0
dashedName: step-27
---

# --description--

Починаючи з пари coffee/price, що вже існує, додайте наступні каву та ціни, використовуючи елементи `article` з двома вкладеними елементами `p` всередині кожного.

```md
Caramel Macchiato 3.75
Pumpkin Spice 3.50
Hazelnut 4.00
Mocha 4.50
```

Як і раніше, текст першого елемента `p` має містити смак кави, а текст другого елемента `p` — ціну.

# --hints--

У вас має бути п’ять елементів `article`.

```js
assert.lengthOf(document.querySelectorAll('article'),5);
```

Кожен елемент `article` має містити два елементи `p`.

```js
const articles = document.querySelectorAll('article');
assert.lengthOf(articles?.[0]?.children, 2);
assert.lengthOf(articles?.[1]?.children, 2);
assert.lengthOf(articles?.[2]?.children, 2);
assert.lengthOf(articles?.[3]?.children, 2);
assert.lengthOf(articles?.[4]?.children, 2);
```

Ваш перший елемент `article` має містити елементи `p` з текстом `French Vanilla` і `3.00`.

```js
const children = document.querySelector('article')?.children;
assert.match(children?.[0]?.innerText.trim(), /French Vanilla/i);
assert.match(children?.[1]?.innerText.trim(), /3\.00/i);
```

Ваш другий елемент `article` має містити елементи `p` з текстом `Caramel Macchiato` і `3.75`.

```js
const children = document.querySelectorAll('article')?.[1]?.children;
assert.match(children?.[0]?.innerText.trim(), /Caramel Macchiato/i);
assert.match(children?.[1]?.innerText.trim(), /3\.75/i);
```

Ваш третій елемент `article` має містити елементи `p` з текстом `Pumpkin Spice` і `3.50`.

```js
const children = document.querySelectorAll('article')?.[2]?.children;
assert.match(children?.[0]?.innerText.trim(), /Pumpkin Spice/i);
assert.match(children?.[1]?.innerText.trim(), /3\.50/i);
```

Ваш четвертий елемент `article` має містити елементи `p` з текстом `Hazelnut` і `4.00`.

```js
const children = document.querySelectorAll('article')?.[3]?.children;
assert.match(children?.[0]?.innerText.trim(), /Hazelnut/i);
assert.match(children?.[1]?.innerText.trim(), /4\.00/i);
```

Ваш п’ятий елемент `article` має містити елементи `p` з текстом `Mocha` і `4.50`.

```js
const children = document.querySelectorAll('article')?.[4]?.children;
assert.match(children?.[0]?.innerText.trim(), /Mocha/i);
assert.match(children?.[1]?.innerText.trim(), /4\.50/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
