---
id: 5f716bee5838c354c728a7c5
title: Крок 48
challengeType: 0
dashedName: step-48
---

# --description--

Під десертом, який ви щойно додали, додайте решту десертів і ціни, використовуючи ще три `article` елементи, кожен з двома вкладеними `p` елементами. Кожен елемент має містити правильний текст десерту та ціни, і всі вони повинні мати правильні класи.

```md
Cherry Pie 2.75
Cheesecake 3.00
Cinnamon Roll 2.50
```

# --before-each--

```js
const section = [...document.querySelectorAll('section')].at(-1);
const paragraphs = section?.querySelectorAll('p');
const prices = section?.querySelectorAll('p.price');
const desserts = section?.querySelectorAll('p.dessert');
```

# --hints--

У вашому другому `article` має бути чотири `section` елементи.

```js
const articles = section?.querySelectorAll('article');
assert.lengthOf(articles, 4);
```

У вас має бути чотири `.dessert` елементи.

```js
assert.lengthOf(desserts, 4);
```

У вас має бути чотири нові `.price` елементи.

```js
assert.lengthOf(prices, 4);
```

Ваш другий `section` елемент має містити вісім `p` елементів.

```js
assert.lengthOf(paragraphs, 8);
```

Ваші `.dessert` елементи мають містити текст `Donut`, `Cherry Pie`, `Cheesecake` та `Cinnamon Roll`.

```js
assert.match(desserts?.[0]?.innerText.trim(), /donut/i);
assert.match(desserts?.[1]?.innerText.trim(), /cherry pie/i);
assert.match(desserts?.[2]?.innerText.trim(), /cheesecake/i);
assert.match(desserts?.[3]?.innerText.trim(), /cinnamon roll/i);
```

Ваші нові `.price` елементи мають містити текст `1.50`, `2.75`, `3.00` та `2.50`.

```js
assert.match(prices?.[0]?.innerText.trim(), /1\.50/);
assert.match(prices?.[1]?.innerText.trim(), /2\.75/);
assert.match(prices?.[2]?.innerText.trim(), /3\.00/);
assert.match(prices?.[3]?.innerText.trim(), /2\.50/);
```

Між вашими елементами `p` не повинно бути пробілів.

```js
assert.lengthOf(document.querySelectorAll('.item'), 9);
document.querySelectorAll('.item').forEach(item => {
  assert.lengthOf(item.children, 2);
  const [first, second] = item.children;
  assert.strictEqual(first.nextSibling, second);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%;
}
```
