---
id: 5f331e55dfab7a896e53c3a1
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` はウェブページ上には表示されない追加情報を提供する要素の一種です。このような要素は、検索エンジンで利用されたり、ページの表示方法を指定するのに役立ちます。

ブラウザがページの文字をエンコードする方法を指定するため、`head` 要素の中に `meta` 要素をネストし、その `charset` 属性を `utf-8` に設定してください。

# --hints--

`meta` タグが 1 つ必要です。

```js
assert.match(code, /<meta.*?\/?>/is);
```

`meta` 要素は空要素なので、終了タグ `</meta>` は必要ありません。

```js
assert.notMatch(code, /<\/meta>/i);
```

`meta` タグには `charset` 属性が必要です。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

`charset` 属性の値を `utf-8` に設定する必要があります。

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
