---
id: 5f331e55dfab7a896e53c3a1
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

The `title` is one of several elements that provide extra information not visible on the web page, but it is useful for search engines or how the page gets displayed.

ブラウザがページの文字をエンコードする方法を指定するため、`head` 要素の中に `meta` 要素をネストし、その `charset` 属性を `utf-8` に設定してください。

# --hints--

`meta` タグが 1 つ必要です。

```js
assert.match(code, /<meta.*?\/?>/is);
```

`meta` 要素は空要素なので、終了タグ `</meta>` は必要ありません。

```js
assert.notMatch(code, /<\/meta>/i);
```

Your `meta` tag should have a `charset` attribute.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Your `charset` attribute should have a value of `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
