---
id: 5f33071498eb2472b87ddee4
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this project, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. CSS is the language used to style an HTML document. It describes how HTML elements should be displayed on the screen.

As you learned in the last few steps of the Cat Photo App, there is a basic structure needed to start building your web page. HTML 文書には `DOCTYPE` 宣言と `html` 要素が必要です。 `DOCTYPE` は、この文書がどのバージョンの HTML で書かれているかをブラウザに伝えます。 そして `html` 要素は、他すべての要素を内包するルート要素を表します。

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Add the `<!DOCTYPE html>` tag, and an `html` element with a `lang` attribute of `en`.

# --hints--

`<!DOCTYPE html>` 宣言が必要です。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

言語に英語を指定した開始タグ `<html lang="en">` が必要です。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

終了タグ `</html>` が必要です。

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
