---
id: 662a6bc12cde72c32fb526f0
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

Una <dfn>interface</dfn> es como un plano para una clase. Una interface contiene un conjunto de métodos y propiedades que una clase debe implementar.  

Comienza este proyecto declarando una clase vacía llamada `Equation`. Utilizarás esta clase para definir una interfaz, un plano para una ecuación genérica.

# --hints--

Debes definir una clase llamada `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
