---
id: 662bdd364bf2cde1487922a9
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

A diferencia de otros lenguajes de programación, Python no implementa interfaces en su lenguaje principal, pero la biblioteca estándar de Python te permite definir interfaces de manera sencilla.

Para este proyecto, usarás utilidades del módulo `abc`. Por lo tanto, importa este módulo en tu código.

# --hints--

Debes importar el módulo `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
