---
id: 662bd8260da84bdd5feae419
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Quieres que la clase `LinearEquation` implemente y no simplemente herede todos los métodos definidos dentro de la clase `Equation`, que debería actuar como una interfaz.

Actualmente, la clase `Equation` es simplemente la clase padre de `LinearEquation`. En los siguientes pasos aprenderás cómo convertirla en una interfaz formal.

Por ahora, crea una instancia de `Equation` y asígnala a una variable `eq`, y una instancia de `LinearEquation` y asígnala a una variable `lin_eq`.

# --hints--

Debes declarar una variable `eq` y asignarle una instancia de `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

Debes declarar una variable `lin_eq` y asignarle una instancia de `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
