---
id: 662bd552e1c1d2db1b88ba47
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Ahora, defina otra clase llamada `LinearEquation` y haz que hereda de `Equation`. Utilizarás esta clase para representar ecuaciones lineales.

# --hints--

Debes definir una clase llamada `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

Tu clase `LinearEquation` debe heredar de la clase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
