---
id: 662bde88dc84f1e249801b1a
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` significa *Abstract Base Classes*. La clase `ABC` te permite convertir una clase regular en una clase abstracta, que es una clase que actúa como una estructura para clases concretas.

Modifica tu declaración `import` para importar solo la clase `ABC` del módulo `abc`. Puedes importar un objeto específico `x` de un módulo `y` siguiendo la estructura de importación `from y import x`.

Luego, convierte tu clase `Equation` en una clase abstracta haciéndola heredar de `ABC`.

# --hints--

Debes importar `ABC` desde el módulo `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

Tu clase `Equation` debería heredar de `ABC`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
