---
id: 662bd456896f16d9bd03f1a6
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Dentro de la clase `Equation`, define dos nuevos métodos de instancia llamados `solve` y `analyze`.

# --hints--

Debes definir un método llamado `solve` dentro de la clase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

Tu método `solve` debe tomar un parámetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

Debes definir un método llamado `analyze` dentro de la clase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

Tu método `analyze` debe tomar un parámetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
