---
id: 662fa2e2cf27c09f21f4f5d0
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

En Python, los tipos de datos se reconocen durante la ejecución (cuando se ejecuta el código). Por lo tanto, no tienes que especificar el tipo de dato de una variable cuando la declaras. No obstante, puedes anotar una variable para aclarar que contendrá un tipo de dato específico con `variable: <data type> = value` o simplemente `variable: <data type>`. Ten en cuenta que el intérprete de Python no aplica las restricciones de los tipos usados para anotar variables, y normalmente necesitarías herramientas externas para hacerlo.

Dentro de la clase `Equation`, define un atributo de clase `degree`. No le asignes un valor. En cambio, usa una anotación de tipo `int` para mostrar que almacenará un número entero dentro de las clases concretas.

Más adelante, usarás este atributo de clase como parte del proceso de validación de los argumentos pasados para instanciar los objetos de la ecuación.

# --hints--

Debes definir un atributo de clase llamado `degree` y anotarlo con `int` dentro de la clase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
