---
id: 662fc3eba556a6bf800d48c1
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

La función incorporada `isinstance()` toma dos argumentos y devuelve un Booleano que indica si el objeto pasado como primer argumento es una instancia de la clase pasada como segundo argumento.

```py
isinstance(7, int) # True
```

Otra cosa que quieres verificar es que cada argumento sea un número. Después de tu primer `if`, crea un bucle `for` que itere sobre `args` y compruebe si el argumento en la iteración actual no es una instancia de `int` o `float`. Usa la función `isinstance()` y pásale una tupla que contenga `int` y `float` como segundo argumento.

Si el argumento no es un número, lanza un `TypeError` que diga `"Coefficients must be of type 'int' or 'float'"`.

# --hints--

Debes crear un ciclo `for` que itere sobre `args` después de tu sentencia `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

Debes crear una estructura `if` que verifique si el coeficiente actual no es una instancia de `int` ni de `float` dentro del `for` loop.

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

Debes usar la cadena proporcionada para lanzar un `TypeError` dentro de la sentencia `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
