---
id: 662f96576ef178927de87975
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Para ser reconocido como un método abstracto, un método debe estar decorado con el decorador `@abstractmethod`.

Un <dfn>decorator</dfn> se usa en Python para modificar el comportamiento de una función. Aquí tienes un ejemplo de cómo usar un decorator llamado `spam`:

```py
@spam
def foo():
    pass
```

Modifica tu declaración de importación para importar el decorador `abstractmethod` y decora tanto los métodos `solve` como `analyze` de la clase `Equation`. Esto generará dos excepciones.

Una vez que una clase que hereda de `ABC` tiene un método abstracto, la clase ya no puede ser instanciada. Por lo tanto, elimina la instancia `Equation` para deshacerte del error.

El otro error ocurre porque la clase `LinearEquation` debe implementar todos los métodos abstractos definidos en la interfaz. Asegúrate de definirlos también dentro de la clase `LinearEquation`. No debes usar el decorador `abstractmethod` en la clase concreta.

# --hints--

Debes importar `abstractmethod` del módulo `abc`.

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

Debes decorar con `@abstractmethod` el método `solve` dentro de la clase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

Debes decorar con `@abstractmethod` el método `analyze` dentro de la clase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

Debes definir un método llamado `solve` dentro de la clase `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

Tu método `solve` debe tomar un parámetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

Debes definir un método llamado `analyze` dentro de la clase `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

Tu método `solve` debe tomar un parámetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
