---
id: 662f6d7c92381a3049e4c987
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Una interfaz no tiene que definir solo métodos abstractos, sino que también puede implementar métodos para ser heredados por las clases concretas.

Antes de encargarte de la implementación real de `solve` y `analyze`, dentro de la clase `Equation`, define un método `__init__`. No uses ningún decorador en él.

# --hints--

Debes definir un método `__init__` en tu clase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

Tu método `__init__` debe tomar un parámetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
