---
id: 6639f947d3a1818c9322c64a
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

El último paso para validar los coeficientes es verificar que el coeficiente de mayor grado sea diferente de cero. Recuerda que el coeficiente de mayor grado debe pasarse como el primer argumento al instanciar el objeto.

Agrega una estructura `if` para eso y lanza un `ValueError` usando la siguiente cadena para proporcionar un mensaje personalizado: `'Highest degree coefficient must be different from zero'`.

# --hints--

Debes crear una estructura `if` que verifique si el primer coeficiente pasado para instanciar la ecuación es igual a cero.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_conditions()[0]
cond.is_equivalent("args[0] == 0") or cond.is_equivalent("0 == args[0]") or cond.is_equivalent("not args[0]")
`)) })
```

Debes lanzar un `ValueError` dentro de la nueva estructura `if` y usar la cadena proporcionada para devolver un mensaje de error personalizado.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_bodies()[0].has_stmt("raise ValueError('Highest degree coefficient must be different from zero')")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
