---
id: 663a32735b317af9812eb0d7
title: Paso 27
challengeType: 20
dashedName: step-27
---

# --description--

En ecuaciones lineales en la forma \\( ax + b = 0 \\), la pendiente es simplemente el coeficiente \\( a \\), y la intersección con el eje y es el coeficiente \\( b \\).

<img alt="a plot of a linear function" src="https://cdn.freecodecamp.org/curriculum/python/linear-equation.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Vas a usar el método `analyze` para proporcionar información adicional sobre la ecuación. Dentro del método `analyze`, desempaqueta los coeficientes en las variables `slope` y `intercept`.

Luego, devuelve un diccionario con las claves `'slope'` y `'intercept'` y los valores de la pendiente y la intersección en y, respectivamente. Después de eso, llama a `analyze` sobre `lin_eq` e imprime el resultado.


# --hints--

Debes desempaquetar los valores almacenados en el atributo `coefficients` en las variables `slope` e `intercept` dentro del método `analyze`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_stmt("slope, intercept = self.coefficients.values()")`)) })
```

El método `analyze` debería devolver un diccionario con las claves `'slope'` y `'intercept'` y los valores de la pendiente y la ordenada al origen, respectivamente.

```js
({ test: () => runPython(`
eq = LinearEquation(2.2, 1.5)
a = eq.analyze()
assert a['slope'] == 2.2, "Expected different slope"
assert a['intercept'] == 1.5, "Expected different intercept"
`) })
```

Debes llamar al método `analyze` de tu objeto `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(lin_eq.analyze())")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x
--fcc-editable-region--
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)
print(lin_eq.solve())
--fcc-editable-region--
```
