---
id: 663a22ba7420c4d2f7fd2aec
title: Paso 25
challengeType: 20
dashedName: step-25
---

# --description--

Es hora de implementar el método `solve`. Dada una ecuación lineal en la forma \\( ax + b = 0 \\), la solución es \\(x = -\frac{b}{a}\\).

Desempaqueta los coeficientes almacenados en el atributo `coefficients` en las variables `a` y `b`. Ten en cuenta que necesitarás usar el método `.values()`.

Luego, declara una variable `x`, asígnale la solución de la ecuación y devuélvela desde el método `solve`.

# --hints--

Debes desempaquetar los valores almacenados dentro del atributo `coefficients` en las variables `a` y `b`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("a, b = self.coefficients.values()")`)) })
```

Debes declarar una variable llamada `x` y asignarle la solución de la ecuación lineal.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("x = -b/a")`)) })
```

Debes retornar `x` desde tu método `solve`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_return("x")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
--fcc-editable-region--
    def solve(self):
        pass
--fcc-editable-region--
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
