---
id: 663b83a28943e6aa6275a514
title: Paso 19
challengeType: 20
dashedName: step-19
---

# --description--

Aún dentro de la clase `Equation`, define un método `__str__` para darle una representación en cadena adecuada a los objetos de ecuación que vas a crear.

Por ahora, dentro del método `__str__`, declara una variable `terms` y asígnale una lista vacía. Usarás esta variable para almacenar cada término (coeficiente por \( x^n \)) de tu ecuación.

Luego, declara una variable `equation_string`, asígnale el resultado de unir los elementos en la lista `terms` con un espacio. Finalmente, retorna `equation_string`.

# --hints--

Debes definir un método `__str__` dentro de la clase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

Tu método `__str__` debe tomar un parámetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

Debes declarar una variable `terms` y asignarle una lista vacía dentro del método `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

Debes declarar una variable `equation_string` y asignarle el resultado de unir los elementos en `terms` con un espacio dentro del método `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

Debes devolver `equation_string` desde tu método `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
