---
id: 663b93aee129b3c4cc07d0db
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Justo después de la lista `terms`, crea un bucle `for` y usa el método `.items()` para iterar sobre las claves y valores almacenados en el atributo `coefficients`. Usa `n` y `coefficient` como las variables del bucle. 

Dentro del ciclo, crea una estructura `if` que verifique si el coeficiente en la iteración actual tiene un valor falso y omite la iteración en ese caso. Esto es porque no quieres representar coeficientes con el valor cero.

# --hints--

Debes crear un ciclo `for` que itere sobre `coefficients.items()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_for_iter().is_equivalent("self.coefficients.items()")`)) })
```

Tu ciclo `for` debe usar `n` y `coefficient` para iterar sobre `coefficients.items()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_for_vars().is_equivalent("n, coefficient")`)) })
```

Debes crear una estructura `if` para verificar si `coefficient` tiene un valor falso dentro de tu ciclo `for`.

```js
({ test: () => assert(runPython(`
if_cond = _Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
conditions = ["not coefficient", "coefficient == 0", "0 == coefficient"]
any(if_cond.is_equivalent(condition) for condition in conditions)
`)) })
```

Debes usar la palabra clave `continue` dentro de tu nueva estructura `if`.

```js
({ test: () => assert(runPython(`
_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("continue")
`)) })
```

Tu ciclo `for` debe colocarse justo después de la declaración de `terms`.

```js
({ test: () => assert(runPython(`
loop = str(_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0])
_Node(_code).find_class("Equation").find_function("__str__").is_ordered("terms = []", loop, "equation_string = ' '.join(terms)", "return equation_string")
`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    def __str__(self):
        terms = []
        
--fcc-editable-region--
        equation_string = ' '.join(terms)
        return equation_string        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
