---
id: 66759e32b88fb5459b1e0234
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

El método `__init_subclass__` se llama cada vez que la clase que lo define es subclasificada y permite personalizar las clases hijas. El método recibe un parámetro llamado por convención `cls` (que significa "class"), que representa la nueva clase hija.

Define un método `__init_subclass__` en tu clase `Equation` y dale un parámetro `cls`.

# --hints--

Debes definir un método `__init_subclass__` con un parámetro `cls` en tu clase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
