---
id: 6675aaf418b41157f6ccd692
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Es hora de volver al método `__init__`. Dependiendo del tipo de ecuación, necesitarás pasar un número variable de argumentos durante la instanciación.

Agrega un segundo parámetro `args` al método y usa el operador `*` para que acepte un número variable de argumentos.

# --hints--

Tu método `__init__` debe tomar dos parámetros, `self` y `*args`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
