---
id: 6675a38a8b535e4ff3274520
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

La función incorporada `hasattr` toma un objeto como primer argumento y una cadena que representa el nombre de un atributo como segundo argumento. Devuelve un booleano que indica si el objeto tiene el atributo especificado.

Ahora vas a usar el método `__init_subclass__` para verificar si la clase hija tiene el atributo `degree` en el momento de la instanciación.

Crea una estructura `if` para verificar si `cls` no tiene un atributo `degree`. Si es así, lanza un `AttributeError` y usa la cadena `f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"` para proporcionar un mensaje personalizado.

Después de eso, corrige el error que ha aparecido en el terminal declarando un atributo de clase `degree` dentro de la clase `LinearEquation`. Este atributo debe representar el grado de la ecuación, que es el exponente del término \\( x \\) más alto. Por lo tanto, asigna el entero `1` al atributo `degree`.

# --hints--

Debes crear una estructura `if` que verifique si `cls` no tiene el atributo `degree` dentro del método `__init_subclass__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_conditions()[0].is_equivalent("not hasattr(cls, 'degree')")`)) })
```

Debes lanzar un `AttributeError` usando la cadena proporcionada dentro de tu sentencia `if`.

```js
({ test: () => runPython(`
raise_stmt = 'raise AttributeError(f"Cannot create \\'{cls.__name__}\\' class: missing required attribute \\'degree\\'")'
node = _Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_bodies()[0]
assert node.has_stmt(raise_stmt) 
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        pass

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    
--fcc-editable-region--
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
