---
id: 69a5f35669099ed52f8563b1
title: 스마트 식료품 저장실 재고 보충기 만들기
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

이 실습에서는 배열, 객체, 반복문, 조건문 같은 기본 JavaScript 개념을 사용해 작은 식료품 저장실 관리 프로그램을 만듭니다.

식료품 저장실에 도착한 물품을 받고, 각 물품에 대해 처리할 작업을 결정하며, 저장을 위해 결과를 정리하는 과정을 시뮬레이션합니다.

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. 문자열 배열을 받아 `parseShipment(rawData)` 속성을 가진 객체 배열을 반환하는 `{ sku, name, qty, expires, zone }` 함수를 구현해야 합니다.

2. 현재 식료품 저장실과 도착한 물품을 비교해 `planRestock(pantry, shipment)` 형태의 작업 배열을 반환하는 `{ type: "restock" | "discard" | "donate", item }` 함수를 구현해야 합니다.

3. 각 아이템의 `groupByZone(actions)` 속성에 따라 작업을 저장 구역별로 그룹화하는 `zone` 함수를 구현해야 합니다.

4. 원본 목록에 영향을 주지 않도록 식료품 저장실의 깊은 복사본을 반환하는 `clonePantry(pantry)` 함수를 구현해야 합니다. 깊은 복사란 새 배열과 새 객체를 만들어 복사본을 수정해도 원본이 변하지 않도록 하는 것입니다.

5. 모든 함수를 함께 사용해 도착한 물품을 처리하고 최종 그룹화된 결과 객체를 콘솔에 출력해야 합니다.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

하나의 문자열 배열 `parseShipment`을 매개변수로 받는 `rawData` 함수를 정의하세요.

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

`parseShipment` 함수는 배열의 도착 물품 문자열을 `sku`, `name`, `qty`, `expires`, `zone` 속성을 가진 객체로 변환해야 합니다.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A10",
  name: "Tomatoes",
  qty: 5,
  expires: "2027-01-01",
  zone: "fridge"
});
```

도착 물품에 중복된 SKU는 무시해야 합니다.

``` js
const rawData = [
  "C32|Eggs|3|2027-01-01",
  "C32|Eggs|3|2027-01-01"
];

const result = parseShipment(rawData);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C32");
```

도착 물품 문자열에 `zone`가 없으면 `general`으로 할당하세요.

``` js
const rawData = ["B21|Bananas|10|2027-01-01"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].zone, "general");
```

`qty` 값은 숫자로 변환해야 합니다.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

두 개의 매개변수 `planRestock`, `pantry`을 받는 `shipment` 함수를 정의하세요.

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

`planRestock` 함수는 `type`와 `item` 속성을 가진 작업 배열을 반환해야 합니다.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 6, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

도착 물품 중 수량이 `0` 이하인 경우, 식료품 저장실에 있든 없든 작업 유형을 `discard`로 생성해야 합니다.

``` js
const pantry = [{ sku: "D43", name: "Pineapples", qty: 2, expires: "2027-01-01", zone: "general" }];
const shipment = [{ sku: "D43", name: "Pineapples", qty: 0, expires: "2027-01-01", zone: "general" },
                  { sku: "E54", name: "Peppers", qty: -1, expires: "2027-01-01", zone: "fridge" }
                ];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
assert.strictEqual(result[1].type, "discard");
```

도착 물품이 식료품 저장실에 이미 존재하면(같은 `sku`) 작업은 `restock`이어야 합니다.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

도착 물품이 식료품 저장실에 없으면 작업은 `"donate"`이어야 합니다.

``` js
const pantry = [];
const shipment = [{ sku: "C32", name: "Eggs", qty: 3, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

하나의 매개변수 `groupByZone`을 받는 `actions` 함수를 정의하세요.

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

`groupByZone` 함수는 각 객체의 `zone` 속성별로 작업을 그룹화해 반환해야 합니다.

``` js

const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }, 
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

하나의 매개변수 `clonePantry`을 받는 `pantry` 함수를 정의하세요.

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

`clonePantry` 함수는 원본 식료품 저장실 배열 대신 새 배열을 반환해야 합니다.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.isArray(pantryClone);
assert.notStrictEqual(pantryClone, pantry);
```

복제된 식료품 저장실 안의 객체들도 새 객체여야 합니다(깊은 복사).

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.notStrictEqual(pantryClone[0], pantry[0]);
assert.deepEqual(pantryClone[0], pantry[0]);

pantryClone[0].qty = 100;
assert.strictEqual(pantry[0].qty, 4);
```

함수들은 함께 작동해 도착 물품을 처리하고 작업을 그룹화해야 합니다.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];

const shipment = parseShipment(rawData);
const pantryClone = clonePantry(pantry);
const actions = planRestock(pantryClone, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.strictEqual(grouped.fridge[0].item.qty, 5);
```

그룹화된 작업 결과를 구역별로 콘솔에 출력하세요.

``` js
const groupedLog = getLogs().find(log => {
  if (!log || typeof log !== "object" || Array.isArray(log)) return false;

  return Object.values(log).some(zoneActions => {
    if (!Array.isArray(zoneActions)) return false;

    return zoneActions.some(action =>
      action &&
      ["restock", "discard", "donate"].includes(action.type) &&
      action.item &&
      typeof action.item.sku === "string"
    );
  });
});

assert.isOk(groupedLog, "Log the grouped actions object to the console.");
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js

function parseShipment(rawData) {

  // Early return if rawData is not an Array
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let line of rawData) {

    const parts = line.split("|");

    const sku = parts[0].trim();

    if (seenSKUs.includes(sku)) continue; // Avoid duplicates

    const name = parts[1].trim();
    const qty = Number(parts[2].trim()); // Parse qty into a number
    const expires = parts[3].trim();
    const zone = ( parts[4] ) ? parts[4].trim() : "general";

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires,
      zone: zone
    });

    seenSKUs.push(sku);

  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let item of shipment) {

    var exists = false;

    for (let stock of pantry)
    {
      if (stock.sku === item.sku)
      {
        exists = true;
      }
    }

    if (item.qty > 0) {
      if (exists) {
        actions.push({ type: "restock", item: item });
      } else{
        actions.push({ type: "donate", item: item });
      }
    } else {
      actions.push({ type: "discard", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {

  const byZone = {};

  for (const action of actions) {

    const zone = action.item.zone;

    if (!byZone[zone]) {
      byZone[zone] = [];
    }

    byZone[zone].push(action);
  }

  return byZone;

}

function clonePantry(pantry) {
  const pantryClone = [];

  for (let item of pantry) {
    pantryClone.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return pantryClone;
}

const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01", // Restock existing item
  "B21|Bananas|10|2027-01-01", // Donate new item without zone
  "C32|Eggs|3|2027-01-01|fridge", // Donate to a defined zone
  "C32|Eggs|3|2027-01-01", // Duplicated SKU in shipment
  "D43|Pineapples|0|2027-01-01", // Discard with quantity of 0
  "E54|Peppers|-1|2027-01-01|fridge" // Discard even if it's not in pantry
];

const shipment = parseShipment(rawData);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

// Show items expanded in zones
console.log(grouped);

```
