---
id: 5a858944d96184f06fd60d61
title: 最初の CSS グリッドを作成する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
dashedName: create-your-first-css-grid
---

# --description--

任意の HTML 要素を `display` プロパティを `grid` に設定してグリッドコンテナに変えましょう。これにより、CSS Grid に関連する他のすべてのプロパティを使えるようになります。

**注意:** CSS Grid では、親要素を <dfn>コンテナ</dfn> と呼び、その子要素を <dfn>アイテム</dfn> と呼びます。

# --instructions--

`container` クラスを持つ div の display を `grid` に変更してください。

# --hints--

`container` クラスは `display` プロパティを持ち、その値は `grid` である必要があります。

```js
assert.match(code, /.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    display: grid;
  }
</style>
```
