---
id: 5a9036ee38fddaf9a66b5d34
title: CSS Gridの単位を使って列と行のサイズを変更する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cvE8phd'
forumTopicId: 301134
dashedName: use-css-grid-units-to-change-the-size-of-columns-and-rows
---

# --description--

CSS Gridでは、`px`や`em`のような絶対単位や相対単位を使って行や列のサイズを定義できます。これらも使えます：

`fr`：利用可能なスペースの分数で列や行のサイズを設定します。

`auto`：列や行のサイズを自動的にその内容の幅や高さに合わせます。

`%`：列や行のサイズをコンテナのパーセント幅に調整します。

プレビューの出力を生成するコードは以下の通りです：

```css
grid-template-columns: auto 50px 10% 2fr 1fr;
```

このスニペットは5つの列を作成します。最初の列は内容の幅と同じで、2番目の列は50px、3番目の列はコンテナの10％、最後の2つの列は残りのスペースを3つに分割し、4番目の列に2つ、5番目の列に1つ割り当てています。

# --instructions--

幅がそれぞれ1fr、100px、2frの3列のグリッドを作ってください。

# --hints--

`container`クラスには、幅が`grid-template-columns`、`1fr`、`100px`の3列を持つ`2fr`プロパティが必要です。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?1fr\s*?100px\s*?2fr\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: auto 50px 10% 2fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 50px 50px;
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 1fr 100px 2fr;
  }
</style>
```
