---
id: 5a90376038fddaf9a66b5d3c
title: justify-items を使ってすべてのアイテムを水平方向に整列する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpECn'
forumTopicId: 301120
dashedName: align-all-items-horizontally-using-justify-items
---

# --description--

CSSグリッド内のすべてのアイテムを同じ配置にしたい場合があります。これまでに学んだプロパティを使って個別に整列させることもできますが、グリッドコンテナに`justify-items`を設定すれば、すべてのアイテムを一度に水平方向に整列できます。このプロパティは、前の2つのチャレンジで学んだすべての値を受け入れられますが、違いはグリッド内の**すべて**のアイテムを希望の配置に動かすことです。

# --instructions--

このプロパティを使って、すべてのアイテムを水平方向の中央に配置してください。

# --hints--

`container`クラスには、値が`justify-items`の`center`プロパティが必要です。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*justify-items\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    justify-items: center;
  }
</style>
```
