---
id: 5a9036ee38fddaf9a66b5d37
title: grid-gapでギャップをより速く追加する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/ca2qVtv'
forumTopicId: 301118
dashedName: add-gaps-faster-with-grid-gap
---

# --description--

`grid-gap` は、前の2つのチャレンジで使った`grid-row-gap`と`grid-column-gap`の省略形プロパティで、より使いやすくなっています。`grid-gap`に1つの値がある場合は、すべての行と列の間にギャップを作ります。ただし、2つの値がある場合は、最初の値が行間のギャップに、2つ目の値が列間のギャップに使われます。

# --instructions--

`grid-gap`を使って、行間に`10px`のギャップ、列間に`20px`のギャップを設定してください。

# --hints--

`container`クラスには、行間に`grid-gap`のギャップ、列間に`10px`のギャップを作る`20px`プロパティが必要です。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-gap\s*?:\s*?10px\s+?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>
<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-gap: 10px 20px;
  }
</style>
```
