---
id: 5a9036d038fddaf9a66b5d32
title: grid-template-columnsでカラムを追加する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c7NzDHv'
forumTopicId: 301117
dashedName: add-columns-with-grid-template-columns
---

# --description--

単にグリッド要素を作成するだけでは十分ではありません。グリッドの構造も定義する必要があります。グリッドにカラムを追加するには、以下のようにグリッドコンテナに`grid-template-columns`プロパティを使います。

```css
.container {
  display: grid;
  grid-template-columns: 50px 50px;
}
```

これでグリッドに幅50pxのカラムが2つできます。`grid-template-columns`プロパティに与えるパラメータの数がグリッドのカラム数を示し、各パラメータの値が各カラムの幅を示します。

# --instructions--

グリッドコンテナに幅がそれぞれ`100px`のカラムを3つ作ってください。

# --hints--

`container`クラスには`grid-template-columns`プロパティを設定し、`100px`の単位で3つ分指定してください。

```js
const templateColumns = new __helpers.CSSHelp(document).getStyle(
  '.container'
)?.gridTemplateColumns;
assert.strictEqual(templateColumns, '100px 100px 100px');
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 100px 100px 100px;
  }
</style>
```
