---
id: 5a90374338fddaf9a66b5d3a
title: justify-self を使ってアイテムを水平方向に揃える
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpKHq'
forumTopicId: 301122
dashedName: align-an-item-horizontally-using-justify-self
---

# --description--

CSS Gridでは、各アイテムの内容は<dfn>セル</dfn>と呼ばれるボックス内に配置されます。グリッドアイテムの`justify-self`プロパティを使うと、そのセル内で内容の位置を水平方向に揃えられます。デフォルトでは、このプロパティの値は`stretch`で、内容がセルの幅いっぱいに広がります。このCSS Gridのプロパティは他の値も受け入れます：

`start`：内容をセルの左側に揃えます、

`center`：内容をセルの中央に揃えます、

`end`：内容をセルの右側に揃えます。

# --instructions--

`justify-self`プロパティを使って、`item2`クラスのアイテムを中央に揃えてください。

# --hints--

`item2`クラスには、値が`justify-self`の`center`プロパティを設定してください。

```js
assert.match(
  code,
  /.item2\s*?{[\s\S]*justify-self\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }

  .item2 {
    background: LightSalmon;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item2 {
    justify-self: center;
  }
</style>
```
