---
id: 5a94fe4469fb03452672e460
title: minmax関数を使ってアイテムのサイズを制限する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cD97RTv'
forumTopicId: 301131
dashedName: limit-item-size-using-the-minmax-function
---

# --description--

`grid-template-columns`と`grid-template-rows`と一緒に使えるもう一つの組み込み関数に`minmax`があります。これはグリッドコンテナのサイズが変わったときにアイテムのサイズを制限するために使います。これを行うには、アイテムの許容されるサイズ範囲を指定する必要があります。例を示します。

```css
grid-template-columns: 100px minmax(50px, 200px);
```

上のコードでは、`grid-template-columns`が2つの列を作成するように設定されています。1つ目は幅100px、2つ目は最小幅50px、最大幅200pxです。

# --instructions--

`minmax`関数を使って、`1fr`関数内の`repeat`を、最小幅`90px`、最大幅`1fr`の列サイズに置き換え、プレビューパネルのサイズを変更して効果を確認してください。

# --hints--

`container`クラスには、最小幅`grid-template-columns`、最大幅`90px`で3列を繰り返すように設定された`1fr`プロパティが必要です。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?minmax\s*?\(\s*?90px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, 1fr);

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, minmax(90px, 1fr));
  }
</style>
```
