---
id: 5a94fe2669fb03452672e45e
title: "エリアテンプレートを作成せずに grid-area を使う"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c6N7VhK'
forumTopicId: 301135
dashedName: use-grid-area-without-creating-an-areas-template
---

# --description--

前のチャレンジで学んだ `grid-area` プロパティは、別の使い方もできます。もしグリッドに参照するエリアテンプレートがない場合でも、アイテムを配置するためのエリアをその場で作成できます。例えば次のようにします。

```css
item1 {
  grid-area: 1/1/2/4;
}
```

これは、先に学んだライン番号を使って、このアイテムのエリアがどこにあるかを定義しています。上の例の数字は次の値を表しています。

```css
grid-area: horizontal line to start at / vertical line to start at / horizontal
  line to end at / vertical line to end at;
```

つまり、例のアイテムは1行目と2行目の間の行、そして1列目と4列目の間の列を占めます。

# --instructions--

`grid-area` プロパティを使って、`item5` クラスの要素を3番目と4番目の水平ラインの間、そして1番目と4番目の垂直ラインの間に配置してください。

# --hints--

`item5` クラスには、3番目と4番目の水平ラインの間、1番目と4番目の垂直ラインの間の全エリアを埋めるための `grid-area` プロパティが必要です。

```js
assert.match(
  code,
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?3\s*?\/\s*?1\s*?\/\s*?4\s*?\/\s*?4\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: 3/1/4/4;
  }
</style>
```
