---
id: 5a94fe8569fb03452672e464
title: グリッド内にグリッドを作成する
challengeType: 0
forumTopicId: 301128
dashedName: create-grids-within-grids
---

# --description--

要素をグリッドにすると、その直接の子孫要素の動作にのみ影響します。ですから、直接の子孫要素をグリッドにすると、グリッドの中にグリッドができます。

たとえば、`display`クラスの要素に`grid-template-columns`と`item3`のプロパティを設定することで、グリッドの中にグリッドを作成できます。

# --instructions--

`item3`クラスの要素を、`auto`と`1fr`を使って幅が`display`と`grid-template-columns`の2列のグリッドにしてください。

# --hints--

`item3`クラスには、`grid-template-columns`と`auto`を値に持つ`1fr`プロパティが必要です。

```js
assert.match(
  code,
  /.item3\s*?{[\s\S]*grid-template-columns\s*?:\s*?auto\s*?1fr\s*?;[\s\S]*}/gi
);
```

`item3`クラスには、`display`の値を持つ`grid`プロパティが必要です。

```js
assert.match(code, /.item3\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .container {
    font-size: 1.5em;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: auto 1fr;
    grid-template-rows: auto 1fr auto;
    grid-gap: 10px;
    grid-template-areas:
      'advert header'
      'advert content'
      'advert footer';
  }
  .item1 {
    background: LightSkyBlue;
    grid-area: header;
  }

  .item2 {
    background: LightSalmon;
    grid-area: advert;
  }

  .item3 {
    background: PaleTurquoise;
    grid-area: content;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .item4 {
    background: lightpink;
    grid-area: footer;
  }

  .itemOne {
    background: PaleGreen;
  }

  .itemTwo {
    background: BlanchedAlmond;
  }
</style>

<div class="container">
  <div class="item1">header</div>
  <div class="item2">advert</div>
  <div class="item3">
    <div class="itemOne">paragraph1</div>
    <div class="itemTwo">paragraph2</div>
  </div>
  <div class="item4">footer</div>
</div>
```

# --solutions--

```html
<style>
  .item3 {
    grid-template-columns: auto 1fr;
    display: grid;
  }
</style>
```
