---
id: 5a94fe3669fb03452672e45f
title: repeat関数を使って繰り返しを減らす
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cQvqyHR'
forumTopicId: 301133
dashedName: reduce-repetition-using-the-repeat-function
---

# --description--

`grid-template-columns` と `grid-template-rows` を使ってグリッドの構造を定義したとき、作成した各行や列に対して値を入力しました。

同じ高さの行が100行あるグリッドを作りたいとしましょう。100個の値を個別に入力するのはあまり実用的ではありません。幸い、より良い方法があります。`repeat` 関数を使って、繰り返したい列や行の数を指定し、その後にカンマと繰り返したい値を続けて書く方法です。

以下は、各行の高さが50pxの100行のグリッドを作成する例です。

```css
grid-template-rows: repeat(100, 50px);
```

repeat関数を使って複数の値を繰り返すこともでき、グリッド構造を定義するときに他の値の間に関数を挿入することもできます。以下はその例です。

```css
grid-template-columns: repeat(2, 1fr 50px) 20px;
```

これは次のように解釈されます。

```css
grid-template-columns: 1fr 50px 1fr 50px 20px;
```

**注意:** `1fr 50px` は2回繰り返され、その後に20pxが続きます。

# --instructions--

`repeat` を使って `grid-template-columns` プロパティの繰り返しをなくしてください。

# --hints--

`container` クラスには、幅が `grid-template-columns` の列を3回繰り返すように設定された `1fr` プロパティが必要です。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?1fr\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: 1fr 1fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, 1fr);
  }
</style>
```
