---
id: 5a94fe1369fb03452672e45d
title: "grid-area プロパティを使ってグリッドエリアにアイテムを配置する"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cRrqmtV'
forumTopicId: 301132
dashedName: place-items-in-grid-areas-using-the-grid-area-property
---

# --description--

前のチャレンジで示したように、グリッドコンテナのエリアテンプレートを作成した後、付けた名前を参照してカスタムエリアにアイテムを配置できます。これを行うには、アイテムに対して`grid-area`プロパティを次のように使います。

```css
.item1 {
  grid-area: header;
}
```

これにより、グリッドは`item1`クラスを`header`という名前のエリアに配置することがわかります。この場合、アイテムはトップ行全体を使います。なぜなら、その行全体が`header`エリアとして名前付けされているからです。

# --instructions--

`item5`プロパティを使って、`footer`エリアに`grid-area`クラスの要素を配置してください。

# --hints--

`item5`クラスには、値が`grid-area`の`footer`プロパティが必要です。

```js
assert.match(
  __helpers.removeCssComments(code),
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?footer\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: footer;
  }
</style>
```
