---
id: 66ea8239b9eeb76f12299030
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

工作坊的下一部分内容是回顾如何在两个数值之间生成随机数。

首先添加另一个 `console.log()`，将信息 `"Now, generate a random number between two values."` 记录到控制台。

然后，创建一个名为 `min` 的变量，为其赋值 `1`，创建一个名为 `max` 的变量，为其赋值 `100` 。

下一步，你将在这两个值之间生成一个随机数。

# --hints--

你应该有一个 `console.log()`，记录信息 `"Now, generate a random number between two values."`。

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

你应该有一个名为 `min` 的变量。

```js
assert.isNotNull(min);
```

你的 `min` 变量应该是一个数字。

```js
assert.isNumber(min);
```

你应将数字 `1` 赋值给 `min` 变量。

```js
assert.equal(min, 1);
```

你应该有一个名为 `max` 的变量。

```js
assert.isNotNull(max);
```

你的 `max` 变量应该是一个数字。

```js
assert.isNumber(max);
```

你应将数字 `100` 赋值给 `max` 变量。

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
