---
id: 66ea7adae8053065a64f9002
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本工作坊，你将通过构建一个数学机器人来复习使用不同的 `Math` 对象方法。 这个数学机器人会将一些数学运算和信息记录到控制台中，以帮助你了解 JavaScript 的 `Math` 对象。

第一步，首先创建一个名为 `botName` 的变量， 并为其赋值 `"MathBot"` 的字符串值。

然后，创建另一个名为 `greeting` 的变量，并为其赋值 `"Hi there! My name is [botName goes here] and I am here to teach you about the Math object!"`。 使用 `botName` 变量代替 `[botName goes here]` 占位符。 

你可以使用模板字面量或使用 `+` 操作符的字符串连接来实现这一目的。

最后，将 `greeting` 变量记录到控制台。

# --hints--

你应该有一个名为 `botName` 的变量。

```js
assert.isNotNull(botName);
```

你的 `botName` 变量应该是一个字符串。

```js
assert.isString(botName);
```

你应将字符串值 `"MathBot"` 赋值给 `botName` 变量。

```js
assert.strictEqual(botName, "MathBot");
```

你应该有一个名为 `greeting` 的变量。

```js
assert.isNotNull(greeting);
```

你的 `greeting` 变量应该是一个字符串。

```js
assert.isString(greeting);
```

You should assign the following sentence to the `greeting` variable: `"Hi there! My name is [MathBot goes here] and I am here to teach you about the Math object!"`。 确保将 `[MathBot goes here]` 替换为 `botName` 变量。

```js
assert.strictEqual(greeting, "Hi there! My name is MathBot and I am here to teach you about the Math object!");
```

你应将 `greeting` 变量记录到控制台。

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
