---
id: 66ea845f54195c7020cc521b
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

在两个数值之间生成随机数的公式如下：

```js
Math.random() * (maximum - minimum) + minimum;
```

这将产生一个介于两个数值之间的浮点数。

创建名为 `randomNum2` 的变量，将生成介于 `min` 和 `max` 之间的值的赋值给它。

然后，记录 `randomNum2` 变量以查看结果。 尝试刷新页面以查看不同的结果。

# --hints--

你应该有一个名为 `randomNum2` 的变量。

```js
assert.isNotNull(randomNum2);
```

你的 `randomNum2` 变量应该是一个数字。

```js
assert.isNumber(randomNum2);
```

你应赋值生成介于 `min` 和 `max` 之间的值的结果。 确保不要硬编码数值。  You should use the formula provided above.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

你应将 `randomNum2` 变量记录到控制台。

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
