---
id: 66ea7fc048a6e26e33c598fe
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

现在，使用 `Math.random()` 方法生成一个随机数。

创建名为 `randomNum` 的变量，并将调用 `Math.random()` 方法的结果值赋值给它。

然后，将 `randomNum` 变量记录到控制台。

尝试在代码中添加空格，重新运行机器人，查看生成的不同随机数。

# --hints--

你应该有一个名为 `randomNum` 的变量。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.randomNum);
```

你应将调用 `Math.random()` 方法的结果赋值给 `randomNum` 变量。

```js
const explorer = await __helpers.Explorer(code);
const randomNum = explorer.variables.randomNum;
assert.equal(randomNum?.value?.toString(), "Math.random()");
```

你应将 `randomNum` 变量记录到控制台。

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
