---
id: 66ea89c25e1bd871688eba6a
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

之前，你只处理过浮点数和数字 `0` 。

接下来，你将学习如何使用 `Math.floor()` 方法。

首先添加一个 `console.log()`，记录信息 `"The Math.floor() method rounds the value down to the nearest whole integer."`

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该记录信息 `"The Math.floor() method rounds the value down to the nearest whole integer."`

```js
assert.equal(
  getLogs()[5],
  "The Math.floor() method rounds the value down to the nearest whole integer."
);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);
--fcc-editable-region--

--fcc-editable-region--
```
