---
id: 66ea7e7399244c6d4b6052a1
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

工作坊的下一部分将回顾 `Math.random()` 方法的工作原理。 

首先为你的代码添加另一个 `console.log()`，该 `console.log()` 将字串 `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` 记录到控制台。
 
# --hints--

你应该有一个 `console.log()`，它记录字串 `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."`。

```js
assert.match(code, /console\.log\((`|'|")\s*The\s+Math\.random\(\)\s+method\s+returns\s+a\s+pseudo\s+random\s+number\s+greater\s+than\s+or\s+equal\s+to\s+0\s+and\s+less\s+than\s+1\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```
