---
id: 66ea7fc048a6e26e33c598fe
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

现在，使用 `Math.random()` 方法生成一个随机数。

创建名为 `randomNum` 的变量，并将调用 `Math.random()` 方法的结果值赋值给它。

然后，将 `randomNum` 变量记录到控制台。

尝试在代码中添加空格，重新运行机器人，查看生成的不同随机数。

# --hints--

你应该有一个名为 `randomNum` 的变量。

```js
assert.isNotNull(randomNum);
```

你应将调用 `Math.random()` 方法的结果赋值给 `randomNum` 变量。

```js
assert.match(code, /randomNum\s*=\s*Math\.random\(\);?/);
```

你不应为 `randomNum` 变量硬编码一个随机数。 确保使用 `Math.random()` 方法生成随机数。

```js
assert.notMatch(code, /randomNum\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum\s*=\s*\d+;?/);
```

你应将 `randomNum` 变量记录到控制台。

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
