---
id: 66eb3ec3b07dfe8f898646c0
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

As you learned in previous lessons, the `Math.floor()` method rounds the value down to the nearest whole integer.

```js
const price = 10.99;
Math.floor(price); // 10
```

创建一个名为 `numRoundedDown` 的变量，将浮点数 `6.7` 向下取整到最接近的整数后的结果赋值给它。

然后，将 `numRoundedDown` 变量记录到控制台以查看结果。

# --hints--

你应该有一个 `numRoundedDown` 变量。

```js
assert.isNotNull(numRoundedDown);
```

`numRoundedDown` 变量应为数字。

```js
assert.isNumber(numRoundedDown);
```

你的 `numRoundedDown` 变量应等于将浮点数 `6.7` 向下取整到最接近的整数。

```js
assert.equal(numRoundedDown, 6);
```

你不应将 `numRoundedDown` 变量的值硬编码为 `6`。

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

你应将 `numRoundedDown` 变量记录到控制台。

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
