---
id: 66eb488c0129e192542f0865
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

In previous lessons, you learned how to work with the `Math.ceil()` method like this:

```js
const price = 10.01;
Math.ceil(price); // 11
```

`Math.ceil()` 方法将数值向上取整为最接近的整数。

创建一个名为 `numRoundedUp` 的变量，将浮点数 `3.2` 向上取整到最接近的整数赋值给它。

然后，将 `numRoundedUp` 变量记录到控制台以查看结果。

# --hints--

你应该有一个 `numRoundedUp` 变量。

```js
assert.isNotNull(numRoundedUp);
```

`numRoundedUp` 变量应为数字。

```js
assert.isNumber(numRoundedUp);
```

你的 `numRoundedUp` 变量应等于将浮点数 `3.2` 向上取整到最接近的整数。

```js
assert.equal(numRoundedUp, 4);
```

你不应将 `numRoundedUp` 变量的值硬编码为 `4`。

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

你应将 `numRoundedUp` 变量记录到控制台。

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
