---
id: 69218e46f4490b33f2481a1a
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

生成两个值之间随机整数的公式如下：

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

这将产生一个介于两个值之间的整数结果。

创建一个名为 `randomInt` 的变量，并将生成介于 `min` 和 `max` 值之间的结果赋值给它。

然后，记录 `randomInt` 变量以查看结果。尝试刷新页面以查看不同的结果。

# --hints--

你应该有一个名为 `randomInt` 的变量。

```js
assert.isNotNull(randomInt);
```

你的 `randomInt` 变量应该是一个数字。

```js
assert.isNumber(randomInt);
```

你应赋值生成介于 `min` 和 `max` 之间的值的结果。 确保不要硬编码数值。  You should use the formula provided above.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

你应该将 `randomInt` 变量记录到控制台日志。

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
