---
id: 69218a15fe536ee1be38aefb
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

接下来部分的研讨会是回顾如何生成两个值之间的随机整数。 

首先为你的代码添加另一个 `console.log()`，将消息 `"Now, generate a random integer between two values."` 记录到控制台。

在下一步中，你将生成一个介于你已经创建的 `min` 和 `max` 变量值之间的随机整数。

# --hints--

你应该有一个 `console.log()`，它记录消息 `"Now, generate a random integer between two values."`。

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+integer\s+between\s+two\s+values\.\s*\1\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

--fcc-editable-region--

--fcc-editable-region--
```
