---
id: 66ebd4ae2812430bb883c787
title: Crie uma central de eventos
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

Neste laboratório, você utilizará os elementos semânticos HTML para criar a estrutura de uma página web.Você adicionará conteúdo e imagens para que pareça um hub de eventos real.

Atenda às histórias de usuário abaixo e faça todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter um elemento `header`.
1. Dentro do elemento `header`, você deve ter um elemento `h1` que contenha o texto `Event Hub`, e um elemento `nav`.
1. Dentro do elemento `nav`, você deve ter uma lista não ordenada com dois itens contendo links para diferentes seções da página. O primeiro item deve ter o texto `Próximos Eventos`, e o segundo item deve ter o texto `Eventos Passados`.
1. Cada link deve ser representado por um elemento `a` com um atributo `href` que aponte para a seção correspondente da página, `#upcoming-events` e `#past-events`, respectivamente.
1. Você deve ter um elemento `main` que contenha as diferentes seções da página.
1. Dentro do elemento `main`, você deve ter dois elementos `section`.
1. O primeiro elemento `section` deve ter um atributo `id` com o valor `upcoming-events`
1. Dentro da seção `#upcoming-events`, você deve ter:
   
   - Um elemento `h2` com o texto `Upcoming Events`.
   - Dois elementos `article`. Cada artigo deve representar um evento e deve ter :
      - Um elemento `h3` para o título do evento.
      - Um elemento `p` para a descrição do evento. Você pode adicionar uma data na parte inferior se desejar.

1. O segundo elemento `section` deve ter um atributo `id` com o valor `past-events`.
1. Dentro da seção `#past-events`, você deve ter:

   - Um elemento `h2` com o texto `Past Events`.
   - Dois elementos `article`. Cada elemento `article` deve representar um evento passado e deve ter:
     - Um elemento `h3` para o título do evento, 
     - Um elemento `p` para a descrição do evento. Você pode adicionar uma data na parte inferior se desejar.
     - Um elemento de imagem com o atributo `src` apontando para um arquivo de imagem e o atributo `alt` com uma descrição da imagem.

**Nota:** Você pode usar qualquer texto para as descrições dos eventos e datas.Você pode usar as seguintes URLs de imagem para as imagens, se desejar:

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

Você deve ter um elemento `header`.

```js
assert.isNotNull(document.querySelector("header"));
```

Seu elemento `header` deve vir após a tag de abertura `body`.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

Dentro do elemento `header`, você deve ter um elemento `h1` que contenha o texto `Event Hub`.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText, "Event Hub");
```

Dentro do elemento `header`, após o elemento `h1`, você deve ter um elemento `nav`.

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

Seu elemento `nav` deve conter uma lista não ordenada com dois itens.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

O primeiro item na lista não ordenada deve ser um link.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

O segundo item na lista não ordenada deve ser um link.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

O texto do primeiro item na lista não ordenada deve ser `Próximos Eventos`.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText, "Upcoming Events");
```

O primeiro item na lista não ordenada deve ter o `href` definido como `#upcoming-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

O texto do segundo item na lista não ordenada deve ser `Eventos Passados`.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText, "Past Events");
```

O segundo item na lista não ordenada deve ter o `href` definido como `#past-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

Você deve ter um elemento `main` após a tag de fechamento do elemento `header`.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

Dentro do elemento `main`, você deve ter dois elementos `section`.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

Seu primeiro elemento `section` deve ter um atributo `id` com o valor `upcoming-events`.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

Seu segundo elemento `section` deve ter um atributo `id` com o valor `past-events`.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

Dentro da seção `#upcoming-events`, você deve ter um elemento `h2` com o texto `Próximos Eventos`.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText, "Upcoming Events");
```

Dentro da seção `#upcoming-events`, você deve ter dois elementos `article` abaixo do elemento `h2`.

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Ambos os elementos `article` dentro da seção `#upcoming-events` devem conter um elemento `h3` para o título do evento.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Ambos os elementos `article` dentro da seção `#upcoming-events` devem conter um elemento de parágrafo para a descrição do evento.

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Dentro da seção `#past-events`, você deve ter um elemento `h2` com o texto `Eventos Passados`.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText, "Past Events");
```

Dentro da seção `#past-events`, você deve ter dois elementos `article` abaixo do elemento `h2`.

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Ambos os elementos `article` dentro da seção `#past-events` devem ter um elemento `h3` para o título do evento.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Ambos os elementos `article` dentro da seção `#past-events` devem conter um elemento de parágrafo para a descrição do evento.

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Ambos os elementos `article` dentro da seção `#past-events` devem conter um elemento de imagem.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

Ambos os elementos de imagem dentro da seção `#past-events` devem ter o atributo `src` apontando para um arquivo de imagem.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

Ambos os elementos de imagem dentro da seção `#past-events` devem ter o atributo `alt` com uma descrição da imagem.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

Cada elemento `h3` deve conter o título do evento.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText)));
```

Cada elemento `p` deve conter a descrição do evento.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText)));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```

