---
id: a105e963526e7de52b219be9
title: Eine eindeutige sortierte Vereinigung implementieren
challengeType: 26
dashedName: implement-a-unique-sorted-union
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `uniteUnique` haben.
1. Die `uniteUnique`-Funktion sollte zwei oder mehr Arrays als Argumente akzeptieren.
1. Die Funktion sollte ein neues Array zurückgeben, das eindeutige Werte aus den Argument-Arrays enthält, und zwar in der Reihenfolge, in der sie zuerst in den Argumenten gefunden werden. Zum Beispiel würde eine Eingabe wie `[1, 2, 4], [2, 3, 5]` eine Ausgabe von `[1, 2, 4, 3, 5]` haben.

# --hints--

Sie sollten eine `uniteUnique`-Funktion haben.

```js
assert.isFunction(uniteUnique);
```

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` sollte `[1, 3, 2, 5, 4]` zurückgeben.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` sollte `[1, 2, 3, 5]` zurückgeben.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` sollte `[1, 2, 3, 5, 4, 6, 7, 8]` zurückgeben.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` sollte `[1, 3, 2, 5, 4, 6]` zurückgeben.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` sollte `[1, 3, 2, 5, 4]` zurückgeben.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const uniteUnique = (...arrays) => {
  return arrays.reduce((a, b) => {
    return [
      ...a,
      ...b.filter((e, currentIndex) => b.indexOf(e) === currentIndex && a.indexOf(e) === -1)
    ];
  }, []);
};
```
