---
id: 66fe4f33a2cc9b33f4d5cd9b
title: 순열 생성기 만들기
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

이 실습에서는 문자열을 받아 문자열의 문자로 만들 수 있는 모든 가능한 순열을 반환하는 순열 생성기를 만듭니다. 예를 들어, 문자열 `cat`의 가능한 순열은 `cat`, `cta`, `act`, `atc`, `tac`, 그리고 `tca`입니다.

문자열의 순열을 재귀적으로 만드는 방법은 문자열의 고정된 시작 부분(접두사)을 저장하고 나머지 부분의 순열을 만드는 방식으로 작동합니다.

예를 들어, `machine`라는 단어를 생각해 봅시다. 순열을 만드는 첫 번째 단계는 `m`를 문자열의 접두사로 고정하고, 나머지 문자열 `achine`의 순열을 만드는 것입니다.

나머지 문자열에 대해서도 같은 방식으로 순열이 계속됩니다. 한 글자가 접두사에 추가되는데, 예를 들어 `c`가 추가되어 접두사가 `mc`가 됩니다. 그런 다음 `ahine`의 각 순열이 접두사에 연결됩니다.

이 과정은 접두사가 모든 문자를 포함하고 나머지 문자열이 비어 있을 때까지 계속되며, 그때 하나의 순열이 완성됩니다.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `permuteString`라는 이름의 함수를 만들어야 합니다.
2. `permuteString` 함수는 하나의 매개변수인 문자열과 기본값이 있는 두 개의 매개변수, 즉 접두사 값과 결과를 저장하고 반환할 빈 배열을 받아야 합니다. 접두사 값은 순열을 형성하기 위해 문자를 누적하는 데 사용되는 문자열이어야 합니다. 함수는 `permuteString("cat")`처럼 단일 인수로 호출됩니다.
3. 함수 내부에서 전달된 문자열의 길이가 `0`인지 확인해야 합니다. 그렇다면 현재 접두사를 결과에 추가하고 결과를 반환하세요.
4. 입력 문자열의 각 문자에 대해 반복하면서 현재 문자를 문자열에서 제거하고 `permuteString` 함수를 재귀적으로 호출하여 나머지 순열을 만드세요.
5. 최종 결과 배열을 반환해야 합니다.
6. 중복을 제거하여 순열이 고유하도록 해야 합니다.

# --hints--

`permuteString` 함수를 만들어야 합니다.

```js
assert.isFunction(permuteString);
```

`permuteString` 함수에서 재귀를 사용해야 합니다.

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")`는 `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`를 반환해야 합니다.

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")`는 `[ "fcc", "cfc", "ccf" ]`를 반환해야 합니다.

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")`는 `[ "p" ]`를 반환해야 합니다.

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")`는 `[""]`를 반환해야 합니다.

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")`는 `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]`를 반환해야 합니다.

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString`는 올바른 결과를 반환해야 합니다.

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
