---
id: 5a858944d96184f06fd60d61
title: أنشئ أول شبكة CSS خاصة بك
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
dashedName: create-your-first-css-grid
---

# --description--

حوّل أي عنصر HTML إلى حاوية شبكة بتعيين خاصية `display` إلى `grid`. هذا يمنحك القدرة على استخدام جميع الخصائص الأخرى المرتبطة بشبكة CSS.

**ملاحظة:** في شبكة CSS، يُشار إلى العنصر الأب باسم <dfn>الحاوية</dfn> وأبناؤه يُسمون <dfn>عناصر</dfn>.

# --instructions--

غيّر خاصية العرض للعنصر div الذي يحمل صنف `container` إلى `grid`.

# --hints--

يجب أن يحتوي صنف `container` على خاصية `display` بقيمة `grid`.

```js
assert.match(code, /.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    display: grid;
  }
</style>
```
