---
id: 5a9036d038fddaf9a66b5d32
title: إضافة أعمدة باستخدام grid-template-columns
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c7NzDHv'
forumTopicId: 301117
dashedName: add-columns-with-grid-template-columns
---

# --description--

إنشاء عنصر شبكة فقط لا يكفي. يجب تعريف هيكل الشبكة أيضًا. لإضافة أعمدة إلى الشبكة، استخدم الخاصية `grid-template-columns` على حاوية الشبكة كما هو موضح أدناه:

```css
.container {
  display: grid;
  grid-template-columns: 50px 50px;
}
```

سيعطي هذا شبكتك عمودين كل منهما بعرض 50px. عدد المعلمات المعطاة للخاصية `grid-template-columns` يحدد عدد الأعمدة في الشبكة، وقيمة كل معلمة تحدد عرض كل عمود.

# --instructions--

أعطِ حاوية الشبكة ثلاثة أعمدة كل منها بعرض `100px`.

# --hints--

يجب أن يحتوي صنف `container` على خاصية `grid-template-columns` بثلاث وحدات من `100px`.

```js
const templateColumns = new __helpers.CSSHelp(document).getStyle(
  '.container'
)?.gridTemplateColumns;
assert.strictEqual(templateColumns, '100px 100px 100px');
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 100px 100px 100px;
  }
</style>
```
