---
id: 5a90376038fddaf9a66b5d3c
title: محاذاة كل العناصر أفقيًا باستخدام justify-items
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpECn'
forumTopicId: 301120
dashedName: align-all-items-horizontally-using-justify-items
---

# --description--

أحيانًا تريد أن تشترك كل العناصر في شبكة CSS في نفس المحاذاة. يمكنك استخدام الخصائص التي تعلمتها سابقًا ومحاذاتها بشكل فردي، أو يمكنك محاذاتها جميعًا مرة واحدة أفقيًا باستخدام `justify-items` على حاوية الشبكة. هذه الخاصية تقبل كل القيم نفسها التي تعلمتها في التحديين السابقين، والفرق هو أنها تحرك **كل** العناصر في شبكتنا إلى المحاذاة المطلوبة.

# --instructions--

استخدم هذه الخاصية لمركزة كل العناصر أفقيًا.

# --hints--

يجب أن يحتوي صنف `container` على خاصية `justify-items` بقيمة `center`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*justify-items\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    justify-items: center;
  }
</style>
```
