---
id: 5a9036ee38fddaf9a66b5d37
title: إضافة الفواصل بسرعة أكبر باستخدام grid-gap
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/ca2qVtv'
forumTopicId: 301118
dashedName: add-gaps-faster-with-grid-gap
---

# --description--

`grid-gap` هي خاصية مختصرة لـ `grid-row-gap` و `grid-column-gap` من التحديين السابقين، وهي أكثر سهولة في الاستخدام. إذا كان لـ `grid-gap` قيمة واحدة، فستُنشئ فاصلًا بين جميع الصفوف والأعمدة. أما إذا كانت هناك قيمتان، فستستخدم الأولى لتحديد الفاصل بين الصفوف والثانية لتحديد الفاصل بين الأعمدة.

# --instructions--

استخدم `grid-gap` لإضافة فاصل `10px` بين الصفوف وفاصل `20px` بين الأعمدة.

# --hints--

يجب أن تحتوي صنف `container` على خاصية `grid-gap` تُدخل فاصلًا `10px` بين الصفوف وفاصلًا `20px` بين الأعمدة.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-gap\s*?:\s*?10px\s+?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>
<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-gap: 10px 20px;
  }
</style>
```
