---
id: 5a90374338fddaf9a66b5d3a
title: محاذاة عنصر أفقيًا باستخدام justify-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpKHq'
forumTopicId: 301122
dashedName: align-an-item-horizontally-using-justify-self
---

# --description--

في CSS Grid، يقع محتوى كل عنصر داخل صندوق يُسمى <dfn>خلية</dfn>. يمكنك محاذاة موقع المحتوى داخل خلية العنصر أفقيًا باستخدام الخاصية `justify-self` على العنصر الشبكي. افتراضيًا، تكون قيمة هذه الخاصية `stretch`، مما يجعل المحتوى يملأ عرض الخلية بالكامل. تقبل هذه الخاصية في CSS Grid قيمًا أخرى أيضًا:

`start`: يحاذي المحتوى إلى يسار الخلية،

`center`: يحاذي المحتوى إلى وسط الخلية،

`end`: يحاذي المحتوى إلى يمين الخلية.

# --instructions--

استخدم الخاصية `justify-self` لمحاذاة العنصر الذي يحمل الصنف `item2` في الوسط.

# --hints--

يجب أن يحتوي الصنف `item2` على خاصية `justify-self` بقيمة `center`.

```js
assert.match(
  code,
  /.item2\s*?{[\s\S]*justify-self\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }

  .item2 {
    background: LightSalmon;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item2 {
    justify-self: center;
  }
</style>
```
