---
id: 5a9036ee38fddaf9a66b5d34
title: استخدام وحدات CSS Grid لتغيير حجم الأعمدة والصفوف
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cvE8phd'
forumTopicId: 301134
dashedName: use-css-grid-units-to-change-the-size-of-columns-and-rows
---

# --description--

يمكنك استخدام وحدات مطلقة ونسبية مثل `px` و `em` في CSS Grid لتحديد حجم الصفوف والأعمدة. يمكنك أيضًا استخدام هذه الوحدات:

`fr`: يحدد العمود أو الصف ككسر من المساحة المتاحة،

`auto`: يحدد العمود أو الصف بعرض أو ارتفاع محتواه تلقائيًا،

`%`: يضبط العمود أو الصف كنسبة مئوية من عرض الحاوية الخاصة به.

إليك الشفرة التي تولد الناتج في المعاينة:

```css
grid-template-columns: auto 50px 10% 2fr 1fr;
```

ينشئ هذا المقتطف خمسة أعمدة. العمود الأول بعرض محتواه، العمود الثاني بعرض 50px، العمود الثالث بنسبة 10% من الحاوية، وبالنسبة للعمودين الأخيرين؛ تُقسم المساحة المتبقية إلى ثلاثة أجزاء، يُخصص اثنان للعمود الرابع وواحد للعمود الخامس.

# --instructions--

أنشئ شبكة بثلاثة أعمدة بعروض كما يلي: 1fr، 100px، و 2fr.

# --hints--

يجب أن يحتوي صنف `container` على خاصية `grid-template-columns` بثلاثة أعمدة بعروض: `1fr`، `100px`، و `2fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?1fr\s*?100px\s*?2fr\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: auto 50px 10% 2fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 50px 50px;
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 1fr 100px 2fr;
  }
</style>
```
