---
id: 5a94fe0569fb03452672e45c
title: قسّم الشبكة إلى قالب مساحة
challengeType: 0
forumTopicId: 301130
dashedName: divide-the-grid-into-an-area-template
---

# --description--

يمكنك تجميع خلايا الشبكة معًا في <dfn>مساحة</dfn> وإعطاء المساحة اسمًا مخصصًا. افعل ذلك باستخدام `grid-template-areas` على الحاوية كما يلي:

```css
grid-template-areas:
  'header header header'
  'advert content content'
  'advert footer footer';
```

الشيفرة أعلاه تجمع خلايا الشبكة إلى أربع مساحات: `header`، `advert`، `content`، و`footer`. كل كلمة تمثل خلية وكل زوج من علامات الاقتباس يمثل صفًا.

# --instructions--

غيّر القالب بحيث تمتد مساحة `footer` عبر الصف السفلي بأكمله. تعريف المساحات لن يؤثر بصريًا الآن. لاحقًا، ستجعل عنصرًا يستخدم مساحة لترى كيف يعمل.

# --hints--

يجب أن يحتوي صنف `container` على خاصية `grid-template-areas` مشابهة للمثال ولكن مع امتداد مساحة `footer` عبر الصف السفلي بأكمله.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.container\s*?{[\s\S]*grid-template-areas\s*?:\s*?["|']\s*?header\s*?header\s*?header\s*?["|']\s*?["|']\s*?advert\s*?content\s*?content\s*?["|']\s*?["|']\s*?footer\s*?footer\s*?footer\s*?["|']\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
    /* Only change code below this line */
      'header header header'
      'advert content content'
      'advert footer footer';
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;

    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```
