---
id: 5a94fe1369fb03452672e45d
title: "وضع العناصر في المساحات الشبكية باستخدام خاصية grid-area"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cRrqmtV'
forumTopicId: 301132
dashedName: place-items-in-grid-areas-using-the-grid-area-property
---

# --description--

بعد إنشاء قالب مساحة لحاوية الشبكة الخاصة بك، كما في التحدٍّ السابق، يمكنك وضع عنصر في مساحتك المخصصة بالإشارة إلى الاسم الذي أعطيته لها. للقيام بذلك، استخدم خاصية `grid-area` على العنصر بهذا الشكل:

```css
.item1 {
  grid-area: header;
}
```

هذا يُعلم الشبكة بأنك تريد أن يذهب الصنف `item1` إلى المساحة المسماة `header`. في هذه الحالة، سيستخدم العنصر الصف العلوي بأكمله لأن هذا الصف كله مسمى كمساحة `header`.

# --instructions--

ضع عنصرًا من الصنف `item5` في مساحة `footer` باستخدام خاصية `grid-area`.

# --hints--

يجب أن يحتوي الصنف `item5` على خاصية `grid-area` بقيمة `footer`.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?footer\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: footer;
  }
</style>
```
