---
id: 5a94fe4469fb03452672e460
title: تحديد حجم العنصر باستخدام دالة minmax
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cD97RTv'
forumTopicId: 301131
dashedName: limit-item-size-using-the-minmax-function
---

# --description--

هناك دالة مدمجة أخرى تُستخدم مع `grid-template-columns` و`grid-template-rows` تُسمى `minmax`. تُستخدم لتحديد حجم العناصر عندما يتغير حجم حاوية الشبكة. للقيام بذلك، يجب تحديد نطاق الحجم المقبول للعنصر. إليك مثالًا:

```css
grid-template-columns: 100px minmax(50px, 200px);
```

في الشفرة أعلاه، تم تعيين `grid-template-columns` لإنشاء عمودين؛ الأول بعرض 100px، والثاني بعرض أدنى 50px وأقصى 200px.

# --instructions--

باستخدام دالة `minmax`، استبدل `1fr` في دالة `repeat` بحجم عمود له الحد الأدنى للعرض `90px` والحد الأقصى للعرض `1fr`، ثم غيّر حجم لوحة المعاينة لرؤية التأثير.

# --hints--

يجب أن يحتوي صنف `container` على خاصية `grid-template-columns` مضبوطة لتكرار 3 أعمدة بالحد الأدنى للعرض `90px` والحد الأقصى للعرض `1fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?minmax\s*?\(\s*?90px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, 1fr);

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, minmax(90px, 1fr));
  }
</style>
```
