---
id: 5a94fe3669fb03452672e45f
title: تقليل التكرار باستخدام دالة repeat
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cQvqyHR'
forumTopicId: 301133
dashedName: reduce-repetition-using-the-repeat-function
---

# --description--

عندما استخدمت `grid-template-columns` و`grid-template-rows` لتعريف هيكل الشبكة، أدخلت قيمة لكل صف أو عمود أنشأته.

لنفترض أنك تريد شبكة بها 100 صف بنفس الارتفاع. ليس من العملي إدخال 100 قيمة بشكل فردي. لحسن الحظ، هناك طريقة أفضل - باستخدام دالة `repeat` لتحديد عدد المرات التي تريد تكرار العمود أو الصف فيها، متبوعة بفاصلة والقيمة التي تريد تكرارها.

إليك مثالًا ينشئ شبكة من 100 صف، كل صف بارتفاع 50px.

```css
grid-template-rows: repeat(100, 50px);
```

يمكنك أيضًا تكرار قيم متعددة باستخدام دالة repeat وإدخال الدالة بين قيم أخرى عند تعريف هيكل الشبكة. هذا ما يبدو عليه الأمر:

```css
grid-template-columns: repeat(2, 1fr 50px) 20px;
```

هذا يُترجم إلى:

```css
grid-template-columns: 1fr 50px 1fr 50px 20px;
```

**ملاحظة:** يتم تكرار `1fr 50px` مرتين متبوعًا بـ 20px.

# --instructions--

استخدم `repeat` لإزالة التكرار من خاصية `grid-template-columns`.

# --hints--

يجب أن تحتوي صنف `container` على خاصية `grid-template-columns` مضبوطة لتكرار 3 أعمدة بعرض `1fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?1fr\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: 1fr 1fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, 1fr);
  }
</style>
```
