---
id: 5a94fe2669fb03452672e45e
title: استخدم grid-area دون إنشاء قالب للمساحات
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c6N7VhK'
forumTopicId: 301135
dashedName: use-grid-area-without-creating-an-areas-template
---

# --description--

يمكن استخدام الخاصية `grid-area` التي تعلمتها في التحدٍّ السابق بطريقة أخرى. إذا لم يكن لدى شبكة التخطيط قالب مناطق للرجوع إليه، يمكنك إنشاء مساحة مؤقتة لوضع عنصر فيها بهذه الطريقة:

```css
item1 {
  grid-area: 1/1/2/4;
}
```

هذا يستخدم أرقام الخطوط التي تعلمتها سابقًا لتحديد مكان المساحة لهذا العنصر. الأرقام في المثال أعلاه تمثل هذه القيم:

```css
grid-area: horizontal line to start at / vertical line to start at / horizontal
  line to end at / vertical line to end at;
```

لذا سيشغل العنصر في المثال الصفوف بين الخطين 1 و2، والأعمدة بين الخطين 1 و4.

# --instructions--

باستخدام الخاصية `grid-area`، ضع العنصر الذي يحمل الصنف `item5` بين الخطين الأفقيين الثالث والرابع وبين الخطين الرأسيين الأول والرابع.

# --hints--

يجب أن يحتوي الصنف `item5` على خاصية `grid-area` ليملأ كامل المساحة بين الخطين الأفقيين الثالث والرابع، والخطين الرأسيين الأول والرابع.

```js
assert.match(
  code,
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?3\s*?\/\s*?1\s*?\/\s*?4\s*?\/\s*?4\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: 3/1/4/4;
  }
</style>
```
