---
id: 5a94fe8569fb03452672e464
title: إنشاء شبكات داخل شبكات
challengeType: 0
forumTopicId: 301128
dashedName: create-grids-within-grids
---

# --description--

تحويل عنصر إلى شبكة يؤثر فقط على سلوك العناصر التابعة المباشرة له. لذا بتحويل عنصر تابع مباشر إلى شبكة، تحصل على شبكة داخل شبكة.

على سبيل المثال، بتعيين خاصيتي `display` و `grid-template-columns` للعنصر الذي يحمل صنف `item3`، تنشئ شبكة داخل شبكتك.

# --instructions--

حوّل العنصر الذي يحمل صنف `item3` إلى شبكة ذات عمودين بعرض `auto` و `1fr` باستخدام خاصيتي `display` و `grid-template-columns`.

# --hints--

يجب أن يحتوي صنف `item3` على خاصية `grid-template-columns` بقيم `auto` و `1fr`.

```js
assert.match(
  code,
  /.item3\s*?{[\s\S]*grid-template-columns\s*?:\s*?auto\s*?1fr\s*?;[\s\S]*}/gi
);
```

يجب أن يحتوي صنف `item3` على خاصية `display` بقيمة `grid`.

```js
assert.match(code, /.item3\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .container {
    font-size: 1.5em;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: auto 1fr;
    grid-template-rows: auto 1fr auto;
    grid-gap: 10px;
    grid-template-areas:
      'advert header'
      'advert content'
      'advert footer';
  }
  .item1 {
    background: LightSkyBlue;
    grid-area: header;
  }

  .item2 {
    background: LightSalmon;
    grid-area: advert;
  }

  .item3 {
    background: PaleTurquoise;
    grid-area: content;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .item4 {
    background: lightpink;
    grid-area: footer;
  }

  .itemOne {
    background: PaleGreen;
  }

  .itemTwo {
    background: BlanchedAlmond;
  }
</style>

<div class="container">
  <div class="item1">header</div>
  <div class="item2">advert</div>
  <div class="item3">
    <div class="itemOne">paragraph1</div>
    <div class="itemTwo">paragraph2</div>
  </div>
  <div class="item4">footer</div>
</div>
```

# --solutions--

```html
<style>
  .item3 {
    grid-template-columns: auto 1fr;
    display: grid;
  }
</style>
```
