---
id: 5a24c314108439a4d4036144
title: ReduxをReactに接続するためにProviderを使う
challengeType: 6
forumTopicId: 301435
dashedName: use-provider-to-connect-redux-to-react
---

# --description--

前回のチャレンジでは、メッセージの配列を管理するReduxストアを作成し、新しいメッセージを追加するためのアクションを作成しました。次のステップは、ReactがReduxストアと更新をディスパッチするために必要なアクションにアクセスできるようにすることです。React Reduxはこれらのタスクを助ける`react-redux`パッケージを提供しています。

React Reduxは、`Provider`と`connect`という2つの重要な機能を持つ小さなAPIを提供します。`connect`については別のチャレンジで扱います。`Provider`はReact Reduxからのラッパーコンポーネントで、Reactアプリを包みます。このラッパーによって、コンポーネントツリー全体でReduxの`store`と`dispatch`関数にアクセスできるようになります。`Provider`はReduxストアとアプリの子コンポーネントという2つのプロップスを受け取ります。Appコンポーネントの`Provider`を定義すると、次のようになります。

```jsx
<Provider store={store}>
  <App/>
</Provider>
```

# --instructions--

コードエディターには、過去のチャレンジで作成したReduxとReactのコードがすべて表示されています。Reduxストア、アクション、そして`DisplayMessages`コンポーネントが含まれています。新しい部分は一番下の`AppWrapper`コンポーネントだけです。このトップレベルコンポーネントを使って、`Provider`からの`ReactRedux`をレンダリングし、Reduxストアをプロップとして渡してください。その後、`DisplayMessages`コンポーネントを子としてレンダリングします。完了すると、Reactコンポーネントがページにレンダリングされているはずです。

**注意:** ここではReact Reduxがグローバル変数として利用可能なので、ドット記法でProviderにアクセスできます。エディター内のコードはこれを利用しており、`Provider`という定数に設定して、`AppWrapper`のレンダリングメソッドで使えるようにしています。

# --hints--

`AppWrapper`がレンダリングされているはずです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('AppWrapper').length === 1;
  })()
);
```

`Provider`ラッパーコンポーネントには、Reduxストアと等しい`store`プロップが渡されているはずです。

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
      return __helpers
        .removeWhiteSpace(code)
        .includes('<Providerstore={store}>');
    })()
  );
```

`DisplayMessages`は`AppWrapper`の子としてレンダリングされているはずです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('AppWrapper').find('DisplayMessages').length === 1
    );
  })()
);
```

`DisplayMessages`コンポーネントは`h2`、`input`、`button`、`ul`要素をレンダリングしているはずです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h2').length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('ul').length === 1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};



const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
    this.handleChange = this.handleChange.bind(this);
    this.submitMessage = this.submitMessage.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {  
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Render the Provider below this line

  // Change code above this line
};
```

# --solutions--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
 this.handleChange = this.handleChange.bind(this);
 this.submitMessage = this.submitMessage.bind(this);
 }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };  
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Change code below this line
  render() {
    return (
      <Provider store = {store}>
        <DisplayMessages/>
      </Provider>
    );
  }
  // Change code above this line
};
```
