---
id: 5a24c314108439a4d4036147
title: ReduxをReactに接続する
challengeType: 6
forumTopicId: 301426
dashedName: connect-redux-to-react
---

# --description--

`mapStateToProps()` と `mapDispatchToProps()` の両方の関数を書いたので、それらを使って `state` と `dispatch` を Reactコンポーネントの `props` にマッピングできます。この作業は React Redux の `connect` メソッドで行えます。このメソッドは2つのオプションの引数、`mapStateToProps()` と `mapDispatchToProps()` を受け取ります。これは、コンポーネントが `state` にのみアクセスする必要があり、アクションをディスパッチする必要がない場合や、その逆の場合があるためです。

このメソッドを使うには、関数を引数として渡し、その結果をすぐにコンポーネントに対して呼び出します。この構文は少し変わっていて、次のようになります。

```js
connect(mapStateToProps, mapDispatchToProps)(MyComponent)
```

**注意:** `connect` メソッドの引数の1つを省略したい場合は、その代わりに `null` を渡してください。

# --instructions--

コードエディターには `mapStateToProps()` と `mapDispatchToProps()` の関数と、新しい Reactコンポーネント `Presentational` があります。`connect` グローバルオブジェクトの `ReactRedux` メソッドでこのコンポーネントを Redux に接続し、`Presentational` コンポーネントに対してすぐに呼び出してください。結果を接続されたコンポーネントを表す新しい `const` `ConnectedComponent` に代入します。これで Redux に接続できました！`connect` の引数のどちらかを `null` に変更して、テスト結果を確認してみてください。

# --hints--

`Presentational` コンポーネントがレンダリングされるはずです。

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('Presentational').length === 1;
  })()
);
```

`Presentational` コンポーネントは `messages` を通じてプロップ `connect` を受け取るはずです。

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return props.messages === '__INITIAL__STATE__';
  })()
);
```

`Presentational` コンポーネントは `submitNewMessage` を通じてプロップ `connect` を受け取るはずです。

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return typeof props.submitNewMessage === 'function';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line

const ConnectedComponent = connect(mapStateToProps, mapDispatchToProps)(Presentational);
```
