---
id: 669fd6fd12918e3de87854d4
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

يجب أن يكون هناك طريقة للمستخدمين للتنقل بين الأقسام المختلفة في الصفحة على مدونتك.

يُستخدم عنصر `nav` لتوفير روابط تنقل إلى أقسام أخرى في المستند أو أقسام أخرى في الموقع. كثيرًا ما ترى عنصر `nav` يُستخدم للقوائم أو جدول المحتويات.

إليك مثالًا على استخدام عنصر `nav`:

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

أسفل عنصر `figure` الخاص بك، أضف عنصر `nav` بداخله عنصر `ul` مسكن.

داخل عنصر `ul`، أضف ثلاثة عناصر `li`.

# --hints--

يجب أن يحتوي عنصر `nav` على وسم بداية.

```js
assert.match(code, /<nav\>/);
```

يجب أن يحتوي عنصر `nav` على وسم الإغلاق.

```js
assert.match(code, /<\/nav\>/);
```

يجب أن يكون عنصر `nav` داخل عنصر `header` الخاص بك.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

يجب أن يحتوي عنصر `ul` على وسم بداية.

```js
assert.match(code, /<ul\>/);
```

يجب أن يحتوي عنصر `ul` على وسم الإغلاق.

```js
assert.match(code, /<\/ul\>/);
```

يجب أن يكون عنصر `ul` داخل عنصر `nav`.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

يجب أن يحتوي عنصر `ul` على ثلاثة عناصر `li`.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
