---
id: 669fdfc9a5e3da42d2376609
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

الآن بعد أن أضفت قسم about، جرّب النقر على الرابط `About` لترى الصفحة تنتقل إلى ذلك القسم.

القسم التالي في صفحة المدونة سيكون قائمة منشورات تتحدث عن السيد Whiskers.

أضف عنصر `section` آخر مع تعيين `id` إلى `"posts"`.

داخل عنصر `section`، أضف عنصر `h2` يحتوي على نص `Posts`.

# --hints--

يجب أن يحتوي شفرتك على وسم افتتاحي `section`.

```js
assert.match(code, /<section.*>/i);
```

يجب أن يحتوي شفرتك على وسم إغلاق `section`.

```js
assert.match(code, /<\/section>/i);
```

يجب أن يحتوي عنصر `section` على `id` معين إلى `"posts"`.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

يجب أن يكون عنصر `section` داخل عنصر `main` الخاص بك.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

يجب أن يكون عنصر `h2` داخل عنصر `section`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

يجب أن يحتوي عنصر `h2` على نص `Posts`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
