---
id: 669fd003cb89ee3c2402e041
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

في هذا المحتوى التمهيدي، سترغب في عرض صورة للسيد ويسكرز مع تسمية توضيحية.

أسفل عنصر `h1`، ابدأ بإضافة عنصر `figure`.

داخل عنصر `figure`، أضف عنصر `img`.

يجب أن تحتوي الخاصية `src` لعنصر `img` على القيمة `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` ويجب أن يكون نص `alt` بقيمة `"a cat in the garden"`.

أسفل عنصر `img` الخاص بك، أضف `figcaption` مع النص `Mr. Whiskers in the Garden`.

# --hints--

يجب أن يحتوي عنصر `figure` على وسم بداية.

```js
assert.match(code, /<figure\>/);
```

يجب أن يحتوي عنصر `figure` على وسم الإغلاق.

```js
assert.match(code, /<\/figure\>/);
```

يجب أن يكون عنصر `figure` داخل عنصر `header` الخاص بك.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

يجب أن يكون لديك عنصر `img`.

```js
assert.exists(document.querySelector("img"));
```

يجب أن يكون عنصر `img` داخل عنصر `figure` الخاص بك.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

يجب أن يحتوي عنصر `img` على خاصية `src` مضبوطة على `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

يجب أن يحتوي عنصر `img` على خاصية `alt` مضبوطة على `"a cat in the garden"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

يجب أن يحتوي عنصر `figcaption` على وسم بداية.

```js
assert.match(code, /<figcaption\>/);
```

يجب أن يحتوي عنصر `figcaption` على وسم الإغلاق.

```js
assert.match(code, /<\/figcaption\>/);
```

يجب أن يكون عنصر `figcaption` داخل عنصر `figure` الخاص بك.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

يجب أن يحتوي عنصر `figcaption` على النص `Mr. Whiskers in the Garden`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

يجب أن يكون عنصر `figcaption` أسفل عنصر `img`. لقد وضعتهم بالترتيب الخطأ.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
