---
id: 669fde2081f65141ad703fe4
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

القسم الأول في الصفحة سيكون قسم التعريف. سيقدم هذا القسم السيد Whiskers ويعطي المستخدمين فكرة عن موضوع هذا المدونة.

داخل عنصر `main` الخاص بك، أضف عنصر `section` مع خاصية `id` مضبوطة على `"about"`.

داخل عنصر `section`، أضف عنصر `h2` يحتوي على النص `About`.

# --hints--

يجب أن يحتوي عنصر `section` الخاص بك على وسم بداية.

```js
assert.match(code, /<section.*>/);
```

يجب أن يحتوي عنصر `section` الخاص بك على وسم الإغلاق.

```js
assert.match(code, /<\/section\>/);
```

يجب أن يحتوي عنصر `section` الخاص بك على خاصية `id` مضبوطة على `"about"`.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

يجب أن يكون عنصر `section` الخاص بك داخل عنصر `main`.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

يجب أن يحتوي عنصر `h2` الخاص بك على وسم بداية.

```js
assert.match(code, /<h2\>/);
```

يجب أن يحتوي عنصر `h2` الخاص بك على وسم الإغلاق.

```js
assert.match(code, /<\/h2\>/);
```

يجب أن يحتوي عنصر `h2` الخاص بك على النص `About`. تحقق من صحة التهجئة.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

يجب أن يكون عنصر `h2` الخاص بك داخل عنصر `section`.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
