---
id: 669fdc11c9b0263fe0814a7a
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

داخل كل عنصر من عناصر `li`، يجب أن يحتوي على عنصر وسم الرابط.

لعنصر وسم الرابط الأول، يجب أن يكون النص `About` وقيمة الخاصية `href` هي `"#about"`. رمز الهاش أمام `about` يمثل اسم `id`، والذي سيُضاف لاحقًا في المشروع.

لعنصر وسم الرابط الثاني، يجب أن يكون النص `Posts` وقيمة الخاصية `href` هي `"#posts"`.

لعنصر وسم الرابط الثالث، يجب أن يكون النص `Contact` وقيمة الخاصية `href` هي `"#contact"`.

# --hints--

يجب أن يحتوي عنصر وسم الرابط الأول على النص `About` داخل أول عنصر `li` لديك.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

يجب أن يحتوي عنصر وسم الرابط الأول على خاصية `href` مضبوطة على `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

يجب أن يحتوي عنصر وسم الرابط الثاني على النص `Posts` داخل ثاني عنصر `li` لديك.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

يجب أن يحتوي عنصر وسم الرابط الثاني على خاصية `href` مضبوطة على `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

يجب أن يحتوي عنصر وسم الرابط الثالث على النص `Contact` داخل ثالث عنصر `li` لديك.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

يجب أن يحتوي عنصر وسم الرابط الثالث على خاصية `href` مضبوطة على `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
