---
id: 669fcb5106eeb13ab7afcd09
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتدرب على العمل مع HTML الدلالي من خلال بناء صفحة مدونة مخصصة للسيد ويسكرز القط.

أول قسم ستبنيه هو القسم العلوي للصفحة <dfn>header</dfn>.

يُستخدم عنصر `header` لتمثيل المحتوى التمهيدي مثل تنقل الصفحة ومعلومات تمهيدية أخرى.

إليك مثالًا يستخدم عنصر `header`:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

داخل عنصر `body`، أضف عنصر `header`.

# --hints--

يجب أن يكون لديك وسم `header` افتتاحي.

```js
assert.match(code, /<header>/i);
```

يجب أن يكون لديك وسم `header` إغلاق.

```js
assert.match(code, /<\/header>/i);
```

يجب أن يكون عنصر `header` داخل عنصر `body`.

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
