---
id: 66a49f685961e997e337cab1
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

سيتولى `header` عرض العنوان الرئيسي، والصورة، والتنقل في الصفحة للمدونة.

داخل عنصر `header`، أضف عنصر `h1` يحتوي على النص `Welcome to Mr. Whiskers' Blog Page!`.

# --hints--

يجب أن يكون لديك وسم `h1` افتتاحي.

```js
assert.match(code, /<h1>/i);
```

يجب أن يكون لديك وسم `h1` إغلاق.

```js
assert.match(code, /<\/h1>/i);
```

يجب أن يحتوي عنصر `h1` على النص `Welcome to Mr. Whiskers' Blog Page!`. تحقق جيدًا من أي أخطاء إملائية.

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

يجب أن يكون عنصر `h1` مسكنًا داخل عنصر `header`.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
