---
id: 65f058e06f34fd1f0ee6e55d
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Em seguida, dentro da classe `Vector`, crie um método vazio `norm` e dê a ele um parâmetro `self`.

# --hints--

Sua classe `Vector` deve ter um método `norm`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

Seu método `norm` deve ter um parâmetro `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
