---
id: 65f0694fb2296f3caadf8347
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

O comprimento de um vetor $\mathbf{a}$, ou norma, é tipicamente indicado como $\\| \mathbf{a} \\|$. Ele pode ser calculado como a raiz quadrada da soma de seus componentes ao quadrado:
\\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2}
 \\]

Calcule a norma do vetor e retorne o resultado do seu método `norm`.

# --hints--

Seu método `norm` deve usar a fórmula fornecida para calcular e retornar um número que representa a norma da instância atual do vetor. Não aproxime o valor.

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
