---
id: 65f01969115f933073b6be03
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Um vetor é um objeto que possui um comprimento (ou magnitude) e uma direção e não pode ser expresso por um único número.Em física, vetores são comumente usados para representar forças, velocidades, acelerações e outras quantidades.

Neste projeto você vai construir um espaço vetorial, um conjunto no qual uma série de operações é definida entre os elementos desse conjunto.Você aprenderá todos os detalhes muito em breve.

Por enquanto, comece o projeto declarando uma classe vazia chamada `Vector`.

# --hints--

Você deve declarar uma classe vazia chamada `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
