---
id: 65f07c9b1ffb814d856dcffc
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Teste se `norm` funciona como esperado imprimindo o valor retornado pelo método.

# --hints--

Você deve chamar `norm()` em `v1` e imprimir o resultado.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1.norm())")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
--fcc-editable-region--
```
