---
id: 65f03d9f92eac9183a4d3281
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Um vetor pode ser definido por duas coordenadas, `x` e `y`, no plano Euclidiano. A distância entre a origem dos eixos e o ponto `(x, y)` será seu comprimento, ou norma. E a direção do vetor apontará para `(x, y)`.

<img alt="a 2-dimensional vector of coordinates (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Dentro da classe `Vector`, crie um método `__init__` e dê a ele três parâmetros, `self`, `x` e `y`.

# --hints--

Você deve definir um método `__init__` dentro da classe `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

Seu método `__init__` deve receber três parâmetros: `self`, `x` e `y`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
