---
id: 65f06a8e5a57673d700c79c3
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Fora da classe `Vector`, crie uma instância de `Vector` passando os inteiros `2` e `3` como os valores `x` e `y` e atribua-a a uma variável chamada `v1`.

# --hints--

Você deve declarar uma variável `v1` e atribuir a ela `Vector(2, 3)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
