---
id: 65f055b9190fc41ca35549b8
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Python oferece vários métodos que incluem tanto um duplo sublinhado inicial quanto final em seus nomes. Você já pode estar familiarizado com alguns, como `__init__` e `__str__`. Esses métodos, que seguem o padrão de nomenclatura `__<name>__`, são chamados de métodos especiais, métodos mágicos ou métodos *dunder* (que significa double underscore).

Definir métodos especiais em uma classe afeta o comportamento dessa classe. Eles são chamados nos bastidores em situações específicas (por exemplo, `__init__` durante a instanciação, `__str__` quando o objeto é impresso ou passado para `str()`). Neste projeto, você vai aprender alguns dos métodos especiais mais comumente usados.

Por enquanto, atribua `x` ao atributo `x` do objeto `Vector`.

# --hints--

Você deve atribuir `x` a `self.x`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
