---
id: 65f3fa097f26b510db6c710b
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Ótimo, `norm` está funcionando como esperado. Agora, se você tentar imprimir `v1`, você obterá a representação padrão em string de um objeto (algo como `<__main__.Vector object at 0x11eb778>`).

Dentro da classe `Vector`, declare um método vazio `__str__` para implementar uma representação em string legível. Lembre-se de dar a ele um parâmetro `self`.

Preste atenção para não imprimir `v1` até que `__str__` retorne uma string, caso contrário você receberá um `TypeError`, porque `__str__` deve sempre retornar uma string.

# --hints--

Você deve definir um método `__str__` dentro da classe `Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

Seu método `__str__` deve ter um parâmetro `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
