---
id: 65f3fa60a93b84110b3f1708
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Você quer retornar uma representação em string significativa. Por exemplo, no caso de `v1`, você quer retornar uma string contendo a tupla com os valores dos componentes do vetor: `(2, 3)`.

No método `__str__`, retorne uma string que represente o vetor como uma tupla contendo os componentes do vetor na ordem. Então, saia da classe `Vector` e imprima `v1` para verificar o resultado.

# --hints--

Seu método `__str__` deve retornar uma string que representa o vetor como uma tupla contendo os componentes do vetor na ordem.

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

Você deve imprimir `v1`.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
