---
id: 65f3fe07cc763212efe91285
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Um vetor pode ter um número `n` de dimensões (componentes).Aqui está uma representação de um vetor tridimensional:

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Até agora, você criou um vetor 2-dimensional.Você quer ser capaz de representar vetores com um número diferente de dimensões sem reescrever o código necessário para cada caso específico.Para isso, você usará herança.

Comece renomeando a classe `Vector` para `R2Vector` para especificar que esta classe vai lidar com vetores bidimensionais. Lembre-se de modificar também a instanciação de `v1`.

# --hints--

Você deve renomear a classe `Vector` para `R2Vector`. Lembre-se de modificar a instanciação de `v1` também.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
