---
id: 65f4520e363e2642f8112e33
title: Passo 19
challengeType: 20
dashedName: step-19
---

# --description--

Como você pode ver, algo não está muito certo.Os métodos `norm` e `__str__` herdados de `R2Vector` não podem se adaptar a um vetor tridimensional.A implementação deles precisa ser mais flexível.

Todo objeto em Python possui um atributo especial chamado `__dict__`, que é um dicionário que armazena os atributos do objeto.

Remova as chamadas existentes de `print`. Depois, imprima o atributo `__dict__` dos seus vetores `v1` e `v2` para ver como eles são.

# --hints--

Você não deve ter `print(v2)` e `print(v2.norm())` no seu código.

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

Você deve imprimir o atributo `__dict__` de `v1` e `v2`.

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
