---
id: 65f4535bbdb28d436ff3ddc9
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Como você pode ver na saída, `__dict__` contém os valores dos atributos da sua instância.Em vez de adicionar explicitamente os quadrados de `self.x` e `self.y`, você vai iterar sobre os valores armazenados em `__dict__` para calcular o valor da norma.

Dentro do corpo do método `norm`, substitua o conteúdo dos parênteses por uma expressão geradora que eleva cada valor `val` em `self.__dict__.values()` à potência de `2`.Além disso, passe essa expressão geradora como argumento para a função `sum`, para que todos os quadrados sejam somados antes de calcular a raiz quadrada.

# --hints--

Você deve retornar `sum(val**2 for val in self.__dict__.values())**0.5` do método `norm()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in self.__dict__.values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)

```
