---
id: 65f40401e6ef53173c04e27d
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Dentro da sua nova classe, declare um método `__init__`.Dê a `self`, `x`, `y` e `z` como parâmetros.

# --hints--

Você deve definir um método `__init__` dentro da sua classe `R3Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").has_function("__init__")
    `))
})
```

Seu método `__init__` deve ter quatro parâmetros `self`, `x`, `y` e `z`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, x, y, z")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
