---
id: 65f40051d6b09a139f253e8e
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Herança permite que você defina uma classe a partir de uma já existente. A nova classe, chamada child, herda todos os métodos e propriedades da classe existente, chamada parent.

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass
    
Oak().sprout() # Output: Making new leaves!
```

No exemplo acima, a classe filha `Oak` herda de `Tree` e herda o método `sprout` da classe pai `Tree`.

Crie uma nova classe chamada `R3Vector` e siga o exemplo acima para fazê-la herdar da classe `R2Vector`.

# --hints--

Você deve definir uma nova classe `R3Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

Sua nova classe `R3Vector` deve herdar de `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
