---
id: 65f412c208c3791fee305acf
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Em Python, você pode impor o uso de argumentos somente por palavra-chave adicionando um `*` como um argumento adicional na assinatura da função ou método. Modifique ambos os métodos `__init__` adicionando um `*` como o segundo parâmetro (após `self`). Todo parâmetro colocado depois disso exigirá o uso de um argumento nomeado na chamada da função/método.

Isso significa que você precisa modificar também a chamada `super().__init__(x, y)`. Faça isso atribuindo a `x` o valor `x` e a `y` o valor `y`.

Finalmente, modifique a instanciação de `v1` e `v2` usando argumentos nomeados.

# --hints--

Você deve impor argumentos nomeados adicionando um `*` como o segundo parâmetro em ambos os seus métodos `__init__`.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

Você deve modificar sua chamada `super.__init__(x, y)` para usar argumentos nomeados.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

Você deve modificar a atribuição de `v1` para ser `R2Vector(x=2, y=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

Você deve modificar a atribuição de `v2` para ser `R3Vector(x=2, y=2, z=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
