---
id: 65f407ea37ad6e181b90462e
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Você poderia atribuir cada parâmetro `i` a `self.i` como fez antes. Embora neste caso você tenha poucas linhas para repetir, uma forma de evitar essa repetição é usando a função `super()`. `super()` permite que você se refira implicitamente à classe pai: `super().__init__(x, y)` chama o método `__init__` da classe pai.

Adicione uma chamada `super().__init__(x, y)` ao seu método `__init__`.

# --hints--

Você deve ter uma chamada `super().__init__(x, y)` no seu método `__init__`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
