---
id: 65f8720685ec351abef26740
title: Passo 24
challengeType: 20
dashedName: step-24
---

# --description--

Para obter uma representação adequada em string, chame o construtor `tuple()` e passe a expressão geradora que você escreveu na etapa anterior como argumento.

# --hints--

Você deve passar a expressão retornada pelo método `__str__` para o construtor `tuple()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").has_return("tuple(getattr(self, i) for i in vars(self))")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return (getattr(self, i) for i in vars(self))
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
