---
id: 65f8709620b2ce1a62608f5a
title: Passo 23
challengeType: 20
dashedName: step-23
---

# --description--

Quando você precisa acessar dinamicamente alguns atributos a partir de uma entrada do tipo string, a função embutida `getattr()` é o que você precisa. Ele recebe um objeto como seu primeiro argumento e uma string contendo o nome do atributo como seu segundo argumento.

Comece a corrigir o método `__str__` substituindo a string retornada por `__str__()` por uma expressão geradora que itere pelos atributos do objeto e chame `getattr()` para cada atributo `i`.

# --hints--

Você deve retornar uma expressão geradora que itere sobre `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

Você deve retornar uma expressão geradora que use `i` como a variável de iteração.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

Você deve retornar uma expressão geradora que chama `getattr(self, i)` para cada item `i` em `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
