---
id: 65f8749b13774b1d2e4a7fba
title: Passo 27
challengeType: 20
dashedName: step-27
---

# --description--

Enquanto o método `__str__` retorna uma representação em string legível para humanos de um objeto, o método `__repr__` deve retornar a string necessária para instanciar o objeto.

Adicione um método `__repr__` com um parâmetro `self` dentro da classe `R2Vector`.

# --hints--

Sua classe `R2Vector` deve ter um método `__repr__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__repr__")`)) })
```

O seu método `__repr__` deve ter um parâmetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
