---
id: 65f9b17a638f0d0dcce8c354
title: Passo 38
challengeType: 20
dashedName: step-38
---

# --description--

Para criar um espaço vetorial, você precisa definir como os vetores devem se comportar em vários casos.Vetores podem ser somados, formando um novo vetor.

O método especial `__add__` pode ser implementado para substituir o que acontece por padrão quando dois objetos são somados usando o operador `+`.

No momento, tentar adicionar duas instâncias de `R2Vector` ou `R3Vector` geraria uma exceção.Crie um método vazio `__add__` dentro da classe `R2Vector` e dê a ele dois parâmetros: `self` e `other`.

# --hints--

Você deve definir um método `__add__` dentro da classe `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

Seu método `__add__` deve ter dois parâmetros, `self` e `other`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
