---
id: 65f9dd6e5a08af19c196c2df
title: Passo 45
challengeType: 20
dashedName: step-45
---

# --description--

O vetor resultante da subtração de um vetor de outro é obtido calculando a diferença de cada um de seus componentes. Por exemplo, subtrair `(2, 4)` de `(7, 3)` gera um terceiro vetor `(5, -1)`, onde `5` é a diferença entre `7` e `2` e `-1` é a diferença entre `3` e `4`.

Como você fez antes dentro do método `__add__`, declare uma variável `kwargs` após a instrução `if`. Atribua a ele uma compreensão de dicionário que itere pelos atributos do objeto e crie um par chave-valor para cada chave `i`, onde a chave é a mesma chave da iteração atual e seu valor é a diferença entre os valores dos atributos `i` de dois operandos.

# --hints--

Você deve declarar uma variável `kwargs` e atribuir a ela uma compreensão de dicionário que itera sobre os atributos do objeto.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

A compreensão de dicionário atribuída a `kwargs` deve usar a variável `i` para iterar sobre os atributos do objeto.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

A compreensão de dicionário atribuída a `kwargs` deve usar `i` como chave.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

A compreensão de dicionário atribuída a `kwargs` deve atribuir a diferença entre o atributo `i` de `self` e o atributo `i` de `other` à chave `i` para cada `i`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
