---
id: 65f9bc10a9fb1612e066e9e8
title: Passo 42
challengeType: 20
dashedName: step-42
---

# --description--

Ao adicionar dois vetores, cada componente de um vetor é somado ao mesmo componente do outro vetor.Por exemplo, adicionar `(1, 2)` e `(2, 4)` gera um terceiro vetor `(3, 6)`, onde `3` é a soma de `1` e `2` e `6` é a soma de `2` e `4`.

O dicionário `kwargs` conterá os pares chave-valor necessários para instanciar um novo vetor da mesma classe dos dois vetores somados.

Transforme o dicionário vazio em uma compreensão de dicionário que itere por `vars(self)` e para cada chave (`i`) crie um par chave-valor, onde a chave é a mesma chave da iteração atual e seu valor é a soma dos valores dos atributos `i` dos dois operandos.

# --hints--

Você deve transformar o dicionário vazio atribuído a `kwargs` em uma compreensão de dicionário que itere sobre `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

A compreensão de dicionário atribuída a `kwargs` deve usar a variável `i` para iterar sobre `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

A compreensão de dicionário atribuída a `kwargs` deve usar `i` como chave.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

A compreensão de dicionário atribuída a `kwargs` deve atribuir a soma do atributo `i` de `self` e do atributo `i` de `other` à chave `i` para cada `i` em `vars(self)`.Use a função `getattr()` para isso.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
