---
id: 666827a6fd0dbaafe8330ea6
title: Passo 35
challengeType: 20
dashedName: step-35
---

# --description--

`__getattr__` é outro método especial que desempenha um papel no acesso a atributos.

A implementação padrão de `__getattribute__` é lançar um `AttributeError` quando o atributo solicitado não é um atributo da instância ou não está presente na árvore da classe.

Nesse caso, `__getattr__` é chamado se definido pela classe.Você pode considerá-lo como uma espécie de fallback quando o procedimento usual de acesso a atributos falha.

Transforme o método `__getattribute__` em `__getattr__` e modifique a string retornada pelo método para `'calling __getattr__'`.

# --hints--

Você deve definir um método chamado `__getattr__` com dois parâmetros, `self` e `attr`, dentro da classe `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_args("self, attr")`)) })
```

Seu método `__getattr__` deve retornar a string `'calling __getattr__'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_return("'calling __getattr__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __getattribute__(self, attr):                
        return 'calling __getattribute__'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

print(v1.x)
print(getattr(v1, 'x'))

```
