---
id: 66682150151af29efec9727d
title: Passo 33
challengeType: 20
dashedName: step-33
---

# --description--

O método `__getattribute__` é chamado nos bastidores toda vez que você tenta acessar um atributo de instância. Se o atributo não for encontrado no nível da instância, o método irá procurá-lo no nível da classe e eventualmente nas classes pai.

Dentro da sua classe, defina um método `__getattribute__` com dois parâmetros, `self` e `attr`, e faça com que ele retorne a string `'calling __getattribute__'`. Você substituirá momentaneamente a implementação padrão apenas para ver como o acesso ao atributo funciona.

# --hints--

Você deve definir um método chamado `__getattribute__` com dois parâmetros, `self` e `attr`, dentro da classe `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_args("self, attr")`)) })
```

Seu método `__getattribute__` deve retornar a string `'calling __getattribute__'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_return("'calling __getattribute__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

```
