---
id: 68e6aff5325410118e76e7be
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, practicarás trabajando con ARIA construyendo un controlador de audio accesible.

Para comenzar, agrega un elemento `h1` con el texto `Audio Controls`.

# --hints--

Debería tener un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Tu elemento `h1` debe estar dentro de tu elemento `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Tu elemento `h1` debe tener el texto `Audio Controls`.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
