---
id: 68eac57fadcc662ce917e75e
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

A continuación, agrega un elemento `button` con el atributo `type` establecido en el valor `button`.

El texto del elemento `button` debe ser `Play`.

# --hints--

Debería tener un elemento `button`.

```js
assert.exists(document.querySelector('button'));
```

Tu elemento `button` debe tener un atributo `type`.

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

Tu elemento `button` debe tener el atributo `type` configurado en `button`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

Tu elemento `button` debe tener el texto `Play`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
