---
id: 68eb7f2397da251a40e04aa2
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

A continuación, agrega un elemento `div` para contener todos los elementos del control de volumen.

Dentro del elemento `div`, anida un elemento `span` con un atributo `id` establecido en `volume-label`.  
Además, el texto para el elemento `span` debe ser `Volume`.

El `id` será referenciado más adelante por atributos ARIA para que las tecnologías de asistencia puedan identificar este texto como la etiqueta del control de volumen.

# --hints--

Debes tener un elemento `div` debajo de tu elemento `button`.

```js
assert.exists(document.querySelector('button + div'));
```

Debes tener un elemento `span` dentro de tu elemento `div`.

```js
assert.exists(document.querySelector('button + div > span'));
```

Tu elemento `span` debe tener un atributo `id`.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

Tu elemento `span` debe tener un atributo `id` con el valor establecido en `volume-label`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

Tu elemento `span` debe tener el texto `Volume`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
