---
id: 68ec0a48c9d6db3594d768c3
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Ahora es momento de agregar otro elemento `span`.

Este elemento `span` debe tener un atributo `id` establecido en `volume-description`.  
Además, el texto para este elemento `span` debe ser `Adjust the sound level`.

Similar al elemento `span` anterior, el `id` de este `span` será referenciado más adelante por atributos ARIA para que las tecnologías de asistencia puedan identificar este texto como la descripción del control de volumen.

# --hints--

Debes tener dos elementos `span` dentro del elemento `div`.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

Tu segundo elemento `span` debe tener un atributo `id`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

Tu segundo elemento `span` debe tener un atributo `id` con el valor de `volume-description`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

Tu segundo elemento `span` debe tener el texto `Adjust the sound level`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
