---
id: 68ed3a96e0ebcc36565b39e5
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

A continuación, necesitas agregar un elemento `input` con el atributo `type` establecido en `range` para crear el control deslizante de volumen. Configura el atributo `min` en `0`, el atributo `max` en `100` y el atributo `value` en `50` para definir el nivel de volumen predeterminado.

# --hints--

Debes tener un elemento `input` debajo de tu segundo elemento `span`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

Tu elemento `input` debe tener un atributo `type`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

Tu elemento `input` debe tener un atributo `type` con el valor `range`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

Tu elemento `input` debe tener un atributo `min`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

Tu elemento `input` debe tener un atributo `min` con el valor `0`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

Tu elemento `input` debe tener un atributo `max`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

Tu elemento `input` debe tener un atributo `max` con el valor `100`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

Tu elemento `input` debe tener un atributo `value`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

Tu elemento `input` debe tener un atributo `value` con el valor `50`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
