---
id: 68ed488821845990c911a223
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Tu controlador de audio accesible está casi completo, pero hay una última cosa que agregar.

Agrega un elemento `button` con el atributo `type` establecido en `button` y con el texto `Mute` debajo del elemento `div`.

Y con ese cambio, ¡has completado este taller!

# --hints--

Debes tener un elemento `button` debajo del elemento `div`.

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

Tu elemento `button` debe tener un atributo `type`.

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

Tu elemento `button` debe tener un atributo `type` con el valor `button`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

Tu elemento `button` debe tener el texto `Mute`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
