---
id: 68ed4204ab7edc66b6cb2383
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Como recuerdas de lecciones anteriores, el atributo `aria-labelledby` se usa cuando hay texto existente en la página que puede usarse como etiqueta. En este caso, has agregado el atributo `id` a tus elementos `span` que ahora se usarán aquí.

Agrega un atributo `aria-labelledby` al elemento `input` configurado en `volume-label volume-description`.

# --hints--

Tu elemento `input` debe tener un atributo `aria-labelledby`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

Tu elemento `input` debe tener un atributo `aria-labelledby` con el valor `volume-label volume-description`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
