---
id: 612e6afc009b450a437940a1
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

从基本的 HTML 结构开始。 添加 `DOCTYPE` 声明以及 `html`、`head`、`body` 和 `title` 元素。

将此页面的语言设置为英文。 将 `title` 设置为 `Piano`。

# --hints--

你的代码应该包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你应在 `DOCTYPE` 后添加一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你应将文档类型定义为 `html` 。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你应该在 `DOCTYPE` 声明的 `type` 后用 `>` 关闭。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素应该有一个值为 `en` 的 `lang` 属性的开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你应该有一个 `head` 开始标签。

```js
assert(code.match(/<head\s*>/i));
```

你应该有一个 `head` 结束标签。

```js
assert(code.match(/<\/head\s*>/i));
```

你应该有一个 `body` 开始标签。

```js
assert(code.match(/<body\s*>/i));
```

你应该有一个 `body` 结束标签。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素应该是同级元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你的代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

项目应该有一个 `Piano` 标题。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

请记住，标题的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
