---
id: 612e78af05201622d4bab8aa
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

Time to start working on the piano. 在 `body` 元素内创建一个 `div` 元素，并将 `id` 设置为 `piano`。

# --hints--

你应该创建一个新的 `div` 元素。

```js
const div = document.querySelector('div');
assert.exists(div);
```

你的 `div` 元素应该在你的 `body` 元素中。

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

你的 `div` 应该有一个 `id`，设置为 `piano`。

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
