---
id: bad87fee1348bd9aec908856
title: 標籤 Bootstrap 按鈕
challengeType: 0
forumTopicId: 18222
dashedName: label-bootstrap-buttons
---

# --description--

就像我們為水井標記一樣，我們也想為按鈕標記。

為你的每個 `button` 元素提供與其 id 選擇器相對應的文字。

# --hints--

你的 `button` 元素，其 id 為 `target1`，應該有文字 `#target1`。

```js
assert.match(document.querySelector('#target1')?.textContent,/#target1/gi)
```

你的 id 為 `target2` 的 `button` 元素應該有 `#target2` 的文字。

```js
assert.match(document.querySelector('#target2')?.textContent,/#target2/gi)
```

你的 `button` 元素，其 id 為 `target3`，應該有 `#target3` 的文字。

```js
assert.match(document.querySelector('#target3')?.textContent,/#target3/gi)
```

你的 `button` 元素，id 為 `target4`，應該有 `#target4` 的文字。

```js
assert.match(document.querySelector('#target4')?.textContent,/#target4/gi)
```

你的 `button` 元素，id 為 `target5`，應該有 `#target5` 的文字。

```js
assert.match(document.querySelector('#target5')?.textContent,/#target5/gi)
```

你的 `button` 元素，其 id 為 `target6`，應該有 `#target6` 的文字。

```js
assert.match(document.querySelector('#target6')?.textContent,/#target6/gi)
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1"></button>
        <button class="btn btn-default target" id="target2"></button>
        <button class="btn btn-default target" id="target3"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4"></button>
        <button class="btn btn-default target" id="target5"></button>
        <button class="btn btn-default target" id="target6"></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
