---
id: bad87fee1348bd9aec908846
title: 創建 Bootstrap 標題
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

現在讓我們從頭開始建置一些東西，以練習我們的 HTML、CSS 和 Bootstrap 技能。

我們將建置一個 jQuery 遊樂場，並且很快會在我們的 jQuery 挑戰中使用它。

首先，創建一個 `h3` 元素，並帶有文字 `jQuery Playground`。

將 `text-primary` Bootstrap 類別添加到你的 `h3` 元素以設定顏色，並將 `text-center` Bootstrap 類別添加到你的 `h3` 元素以置中。

# --hints--

你應該為你的頁面添加一個 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

你的 `h3` 元素應該有一個結束標籤。

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

你的 `h3` 元素應該透過套用類別 `text-primary` 來著色

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

你的 `h3` 元素應該透過套用類別 `text-center` 來置中。

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

你的 `h3` 元素應該有 `jQuery Playground` 文字。

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```
