---
id: bad87fee1348bd9aec908746
title: 將我們的頁面置於 Bootstrap 的 `container-fluid` 容器中
challengeType: 0
forumTopicId: 18198
dashedName: house-our-page-within-a-bootstrap-container-fluid-div
---

# --description--

現在讓我們確保你頁面上的所有內容都是行動裝置響應式的。

讓你的 `h3` 元素巢狀在具有類別 `container-fluid` 的 `div` 元素內。

# --hints--

你的 `div` 元素應該有 `container-fluid` 類別。

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container-fluid'));
```

你的每個 `div` 元素都應該有結束標籤。

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);

assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

你的 `h3` 元素應該巢狀在 `div` 元素內。

```js
const divElement = document.querySelector('div');
const divChildren = divElement?.querySelectorAll(`:scope ${'h3'}`)
assert.lengthOf(divChildren,1);
```

# --seed--

## --seed-contents--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```

# --solutions--

```html
<div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
</div>
```
