---
id: bad87fee1348bd9aec908848
title: 創建 Bootstrap Wells
challengeType: 0
forumTopicId: 16825
dashedName: create-bootstrap-wells
---

# --description--

Bootstrap 有一個名為 `well` 的類別，可以為你的行創建視覺深度感。

在你的每個 `col-xs-6` `div` 元素中巢狀一個帶有類別 `well` 的 `div` 元素。

# --hints--

你應該在每個具有類別 `col-xs-6` 的 `div` 元素內添加一個具有類別 `well` 的 `div` 元素。

```js
const wells = document.querySelectorAll('div.col-xs-6 > div.well');
assert.lengthOf( wells,2 ); 
```

你的兩個具有類別 `col-xs-6` 的 `div` 元素都應該巢狀在具有類別 `row` 的 `div` 元素內。

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

你的所有 `div` 元素都應該有結束標籤。

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g)?.length , code.match(/<div/g)?.length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">

    </div>
    <div class="col-xs-6">

    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
  </div>
</div>
```
