---
id: bad87fee1348bd9bec908846
title: 創建 Bootstrap 列
challengeType: 0
forumTopicId: 16813
dashedName: create-a-bootstrap-row
---

# --description--

現在我們將為我們的行內元素創建一個 Bootstrap 列。

在 `h3` 頁籤下方創建一個 `div` 元素，並設定類別為 `row`。

# --hints--

你應該在你的 `h3` 元素下方添加一個 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div'),2);
assert.lengthOf(document.querySelectorAll('div.row h3.text-primary'),0);
assert.lengthOf(document.querySelectorAll('div.row + h3.text-primary'),0)
assert.lengthOf(document.querySelectorAll('h3.text-primary + div.row'),1);
```

你的 `div` 元素應該有 `row` 類別

```js
const newDiv = document.querySelectorAll('div')?.[1]; 
assert.isTrue(newDiv?.classList?.contains('row'));
```

你的 `row div` 應該巢狀在 `container-fluid div` 裡面

```js
assert.lengthOf(document.querySelectorAll('div.container-fluid div.row'),1);
```

你的 `div` 元素應該有一個結束標籤。

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>

</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row"></div>
</div>
```
