---
id: bad87fee1348bd9aec908847
title: 拆分你的 Bootstrap 列
challengeType: 0
forumTopicId: 18306
dashedName: split-your-bootstrap-row
---

# --description--

現在我們有了一個 Bootstrap 列，讓我們將它分成兩個行來放置我們的元素。

在你的列中創建兩個 `div` 元素，兩者皆具有類別 `col-xs-6`。

# --hints--

兩個 `div class="col-xs-6"` 元素應該巢狀在你的 `div class="row"` 元素內。

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

你的所有 `div` 元素都應該有結束標籤。

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length , code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">


  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6"></div>
    <div class="col-xs-6"></div>
  </div>
</div>
```
