---
id: 67ed03ac474c48692f41749e
title: بناء جدول تجزئة
challengeType: 27
dashedName: build-a-hash-table
saveSubmissionToDB: true
---

# --description--

في هذا المختبر، ستبني جدول تجزئة من الصفر. جدول التجزئة هو هيكل بيانات يخزن أزواج المفتاح-القيمة. يعمل جدول التجزئة بأخذ المفتاح كمُدخَل ثم تجزئة هذا المفتاح وفقًا لدالة تجزئة محددة.

لغرض هذا المختبر، ستكون دالة التجزئة بسيطة: ستجمع قيم Unicode لكل حرف في المفتاح. ثم يُستخدم قيمة التجزئة كمفتاح فعلي لتخزين القيمة المرتبطة. تُستخدم نفس قيمة التجزئة أيضًا لاسترجاع وحذف القيمة المرتبطة بالمفتاح.

**الهدف:** نفذ قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تعرف صنفًا باسم `HashTable` مع خاصية `collection` مُهيأة إلى قاموس فارغ عند إنشاء نسخة جديدة من `HashTable`. يجب أن يخزن قاموس `collection` أزواج المفتاح-القيمة بناءً على قيمة التجزئة للمفتاح.

2. يجب أن يحتوي صنف `HashTable` على أربع طرق مثيل: `hash`، `add`، `remove`، و `lookup`.

3. يجب أن تقوم طريقة `hash` بما يلي:

- تأخذ سلسلة نصية كمعلمة.
    - تُرجع قيمة مجزأة محسوبة كمجموع قيم Unicode (ASCII) لكل حرف في السلسلة النصية. يمكنك استخدام الدالة `ord` لهذا الحساب.

4. يجب أن تقوم طريقة `add` بما يلي:
    - تأخذ معلمتين تمثلان زوج المفتاح-القيمة، وتحسب تجزئة المفتاح.
    - تستخدم قيمة التجزئة المحسوبة كمفتاح لتخزين قاموس يحتوي على زوج المفتاح-القيمة داخل قاموس `collection`.
    - إذا أنتجت عدة مفاتيح نفس قيمة التجزئة، يجب تخزين أزواج المفتاح-القيمة الخاصة بها في القاموس المسكن الموجود تحت نفس قيمة التجزئة.

5. يجب أن تقوم طريقة `remove` بما يلي:

- تأخذ مفتاحًا كمعامل وتحسب تجزئته.
    - تتأكد من وجود المفتاح في القاموس.
    - تحذف زوج المفتاح-القيمة المقابل من جدول التجزئة.
    - إذا لم يكن المفتاح موجودًا في القاموس، فلا يجب أن ترفع خطأ أو تحذف أي شيء.

6. يجب أن تقوم طريقة `lookup` بما يلي:

- تأخذ مفتاحًا كمعامل.
    - تحسب تجزئة المفتاح، وتُرجع القيمة المقابلة المخزنة داخل جدول التجزئة.
    - إذا لم يكن المفتاح موجودًا في القاموس، يجب أن تُرجع `None`.

# --hints--

يجب أن تعرف صنفًا باسم `HashTable`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_class("HashTable")
    `)) 
})
```

عند إنشاء نسخة جديدة من صنف `HashTable`، يجب تهيئة خاصية `collection` إلى قاموس فارغ.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.collection == {}
    `) 
})
```

يجب أن يحتوي صنف `HashTable` على طريقة `hash`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('hash')`))}})
```

يجب أن تأخذ طريقة `hash` سلسلة نصية كمعلمة.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.hash)
    len(sig.parameters) == 1
  `))
})
```

يجب أن تأخذ طريقة `hash` سلسلة نصية كمعامل وتُرجع مجموع قيم Unicode لكل حرف في السلسلة النصية.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("fcc")
        assert hash_result == 300

        # prevent hardcoding
        assert ht.hash("golf") == 424
        assert ht.hash("read") == 412
`) })
```

يجب أن يحتوي صنف `HashTable` على طريقة `add`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('add')`))}})
```

يجب أن تأخذ طريقة `add` مفتاحًا وقيمة كمعلمتين.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.add)
    len(sig.parameters) == 2
  `))
})
```

يجب أن يحتوي صنف `HashTable` على طريقة `remove`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('remove')`))}})
```

يجب أن تأخذ طريقة `remove` مفتاحًا كمعلمة.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.remove)
    len(sig.parameters) == 1
  `))
})
```

عند محاولة حذف مفتاح غير موجود في القاموس، لا يجب أن ترفع خطأ أو تحذف أي شيء.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("rose", "flower")
        index = ht.hash("rose")
        original = ht.collection.copy()

        ht.remove("tulip")
        ht.remove("sore")

        assert ht.collection == original
        assert "rose" in ht.collection[index]
        assert "tulip" not in ht.collection.get(index, {})
    `) 
})
```

إذا كانت عدة مفاتيح تُجزأ إلى نفس الفهرس، يجب أن تحذف طريقة `remove` زوج المفتاح-القيمة المحدد فقط، وليس القاموس بأكمله عند ذلك الفهرس.

```js
({
  test: () => runPython(`
    ht = HashTable()
    ht.add("rose", "flower")
    ht.add("sore", "pain")  # "rose" and "sore" both hash to the same index

    index = ht.hash("rose")
    ht.remove("rose")

    assert index in ht.collection
    assert "rose" not in ht.collection[index]
    assert "sore" in ht.collection[index]
    assert ht.collection[index]["sore"] == "pain"
  `)
})
```

يجب أن يحتوي صنف `HashTable` على طريقة `lookup`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('lookup')`))}})
```

يجب أن تأخذ طريقة `lookup` مفتاحًا كمعلمة.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.lookup)
    len(sig.parameters) == 1
  `))
})
```

`HashTable().hash('golf')` يجب أن تُرجع `424`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("golf")
        assert hash_result == 424

        # prevent hardcoding
        assert ht.hash("dear") == 412
        assert ht.hash("cat") == 312
    `) 
})
```

`HashTable().add('golf', 'sport')` يجب أن تضيف زوج المفتاح-القيمة إلى القاموس عند المفتاح `424`.

```js

({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")
    expected_value = {424: {'golf': 'sport'}}

    assert ht.collection == expected_value
    `) 
})
```

`HashTable().add('dear', 'friend')` و `HashTable().add('read', 'book')` يجب أن تضيفا كلا زوجي المفتاح-القيمة إلى القاموس عند الفهرس `412` كقاموس مسكن.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("dear", "friend")
    ht.add("read", "book")

    expected_value = {
        "dear": "friend",
        "read": "book"
    }

    assert ht.collection.get(412) == expected_value
    `) 
})
```

عندما يكون المفتاح موجودًا في جدول التجزئة، يجب أن تحذف طريقة `remove()` المفتاح المعطى وقيمته المقابلة من القاموس.

```js
({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")

    expected_value_before_removal = {
        "golf": "sport"
    }

    index = ht.hash("golf")

    assert ht.collection.get(index) == expected_value_before_removal

    ht.remove("golf")

    assert "golf" not in ht.collection.get(index, {})
    `) 
})
```

عندما يكون زوج المفتاح-القيمة `'golf', 'sport'` موجودًا في جدول التجزئة، يجب أن تُرجع `HashTable().lookup('golf')` القيمة `sport`.

```js
({ 
    test: () => runPython(`
    ht = HashTable()

    ht.add("golf", "sport")

    expected_value = "sport"
    assert ht.lookup("golf") == expected_value
    `) 
})

```

عندما لا يكون زوج المفتاح-القيمة `'golf', 'sport'` موجودًا في القاموس، يجب أن تُرجع `HashTable().lookup('golf')` القيمة `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.lookup("golf") is None
`) })

```

عندما يكون المفتاح `'fcc'` موجودًا في القاموس، يجب أن تُرجع `HashTable().lookup('cfc')` القيمة `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("fcc", "coding")
        assert ht.lookup("cfc") is None
`) })
```

عند إضافة `('rose', 'flower')` إلى جدول التجزئة، يجب أن تبدو خاصية `collection` هكذا: `{ 441: { 'rose': 'flower' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("rose", "flower")

    expected_value = {
        "rose": "flower"
    }

    assert ht.collection.get(441) == expected_value

    # prevent hardcoding

    ht2 = HashTable()

    ht2.add("kebab", "food")
    expected_value = {
        "kebab": "food"
    }

    assert ht2.collection.get(501) == expected_value
    `)
})
```

عندما تضيف مفتاحًا ينتج عنه نفس قيمة التجزئة لمفتاح موجود، مثل `fcc` و `cfc`، يجب أن تبدو الـ `collection` هكذا: `{ 300: { 'fcc': 'coding', 'cfc':  'chemical' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("fcc", "coding")
    ht.add("cfc", "chemical")

    expected_value = {
        "fcc": "coding",
        "cfc": "chemical"
    }

    assert ht.collection.get(300) == expected_value

    # prevent hardcoding

    ht2 = HashTable()
    ht2.add("cat", "animal")
    ht2.add("act", "verb")

    expected_value = {
        "cat": "animal",
        "act": "verb"
    }

    assert ht2.collection.get(312) == expected_value
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class HashTable:
    def __init__(self):
        self.collection = {}

    def hash(self, string):
        hashed = 0
        for char in string:
            hashed += ord(char)  
        return hashed

    def add(self, key, val):
        the_hash = self.hash(key)
        if the_hash not in self.collection:
            self.collection[the_hash] = {}
        self.collection[the_hash][key] = val

    def remove(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            del self.collection[the_hash][key]
            if not self.collection[the_hash]: 
                del self.collection[the_hash]

    def lookup(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            return self.collection[the_hash][key]
        return None
```
