---
id: 66ed41f912d0bb1dc62da5dd
title: JavaScriptトリビアボットを作ろう
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. ユーザーへの挨拶メッセージとして`"Hello! I'm your coding fun fact guide!"`をコンソールにログ出力してください。
1. ボットの名前、出身地、お気に入りのコーディング言語をそれぞれストアする`botName`、`botLocation`、`favoriteLanguage`の3つの変数を作成してください。
1. `"My name is (botName) and I live on (botLocation)."`をコンソールにログ出力してください。
1. `"My favorite programming language is (favoriteLanguage)."`をコンソールにログ出力してください。
1. `let`を使って`codingFact`という変数を作成し、ボットのお気に入りのコーディング言語に関する面白い事実を文字列で代入し、`favoriteLanguage`変数を含めてください。
1. `codingFact`をコンソールにログ出力してください。
1. `codingFact`変数にボットのお気に入り言語に関する新しい事実を`favoriteLanguage`変数を使って再代入してください。
1. `codingFact`を再度コンソールにログ出力してください。
1. `codingFact`変数にもう一度、`favoriteLanguage`変数を使ってボットのお気に入り言語に関する別の新しい事実を再代入してください。
1. `codingFact`を3回目のコンソールログとして出力してください。
1. ボットからの別れの言葉として`"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."`をコンソールにログ出力してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`"Hello! I'm your coding fun fact guide!"`をコンソールにログ出力してください。

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

`botName`変数を宣言してください。スペルや大文字・小文字の間違いがないか再確認してください。

```js
assert.exists(botName);
```

`botName`変数は文字列である必要があります。

```js
assert.isString(botName);
```

`botLocation`変数を宣言してください。スペルや大文字・小文字の間違いがないか再確認してください。

```js
assert.exists(botLocation);
```

`botLocation`変数は文字列である必要があります。

```js
assert.isString(botLocation);
```

`favoriteLanguage`変数を宣言してください。スペルや大文字・小文字の間違いがないか再確認してください。

```js
assert.exists(favoriteLanguage);
```

`favoriteLanguage`変数は文字列である必要があります。

```js
assert.isString(favoriteLanguage);
```

`"My name is (botName) and I live on (botLocation)."`をコンソールにログ出力し、変数を文字列に追加してください。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

`"My favorite programming language is (favoriteLanguage)."`をコンソールにログ出力し、変数を文字列に追加してください。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

`let`を使って新しい変数`codingFact`を宣言してください。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

`codingFact`に`favoriteLanguage`を含む値を代入してください。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

`codingFact`をコンソールにログ出力してください。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

`codingFact`に`favoriteLanguage`を含む新しい値を代入し、コンソールにログ出力してください。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

`codingFact`に`favoriteLanguage`を含む値を3回目に代入し、コンソールにログ出力してください。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact); 
```

`"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."`をコンソールにログ出力し、変数の値を追加してください。

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
