---
id: 66faac4139dbbd5dd632c860
title: Formular für eine Bewerbung erstellen
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um die Aufgabe abzuschließen.

**User Stories:**

1. Sie sollten ein `div`-Element mit der Klasse `container` haben.
1. Innerhalb des `div`-Elements sollten Sie ein `form`-Element haben.
1. Das Formular sollte ein `input`-Element mit dem Typ `text` und der ID `name` enthalten, um den vollständigen Namen des Benutzers einzugeben.
1. Sie sollten ein weiteres `input`-Element mit dem Typ `email` und der ID `email` haben, um die E-Mail-Adresse des Benutzers einzugeben.
1. Das Formular sollte ein `select`-Element mit der ID `position` enthalten, das es den Benutzenden ermöglicht, eine Stelle auszuwählen.
1. Sie sollten ein `fieldset`-Element mit der Klasse `radio-group` haben.
1. Innerhalb von `.radio-group` sollten Sie eine Gruppe von `input`-Elementen mit dem Typ `radio` und passenden Bezeichnungen haben, um Verfügbarkeitsoptionen auszuwählen (z. B. Vollzeit, Teilzeit). Die Gruppe `name` sollte `availability` sein.
1. Sie sollten ein `textarea`-Element mit der ID `message` haben, um eine Nachricht einzugeben.
1. Sie sollten jedes `input`-, `select`- und `textarea`-Element mit einem `label`-Element verknüpfen.
1. Sie sollten ein `button`-Element mit dem Typ `submit` zum Absenden des Formulars haben.
1. Fügen Sie eine `:focus`-Pseudoklasse zu den `input`- und `textarea`-Elementen hinzu, um deren Rahmenfarbe zu ändern und den Standardrahmen beim Fokus zu entfernen.
1. Die `input`-, `select`- und `textarea`-Elemente sollten eine `:invalid`-Pseudoklasse haben, die die Rahmenfarbe rot ändert, wenn eine ungültige Eingabe erkannt wird.
1. Die `input`-, `select`- und `textarea`-Elemente sollten eine `:valid`-Pseudoklasse haben, die die Rahmenfarbe grün ändert, wenn eine gültige Eingabe erfolgt.
1. Das `button`-Element sollte eine `:hover`-Pseudoklasse haben, die die Hintergrundfarbe beim Darüberfahren mit der Maus ändert.
1. Verwenden Sie die `:checked`-Pseudoklasse auf `.radio-group input[type="radio"]`, um eine Rahmenfarbe, Hintergrundfarbe und einen Schatten hinzuzufügen, wenn die Optionsschaltfläche ausgewählt ist.
1. Verwenden Sie die `:checked`-Pseudoklasse auf Optionsschaltflächen, um die Textfarbe der zugehörigen `label` zu ändern, wenn die Option ausgewählt ist.
1. Fügen Sie eine `:first-of-type`-Pseudoklasse zum `input`-Element hinzu, um das erste Eingabefeld anders zu gestalten (z. B. abgerundete Ecken).

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet in Ihrem HTML verlinken und Ihr CSS anwenden.

# --hints--

Sie sollten ein `div`-Element mit der Klasse `container` haben.

```js
assert.exists(document.querySelector("div.container"));
```

Innerhalb des `div`-Elements sollten Sie ein `form`-Element haben.

```js
assert.exists(document.querySelector("div.container > form"));
```

Das Formular sollte ein `input`-Element mit dem Typ `text` und der ID `name` enthalten, um den vollständigen Namen des Benutzers einzugeben.

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

Sie sollten ein weiteres `input`-Element mit dem Typ `email` und der ID `email` haben, um die E-Mail-Adresse des Benutzers einzugeben.

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

Das Formular sollte ein `select`-Element mit der ID `position` mit einigen `option`-Elementen enthalten.

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

Sie sollten ein `fieldset`- oder `section`-Element mit der Klasse `.radio-group` haben.

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

Innerhalb von `.radio-group` sollten Sie eine Gruppe von `input`-Elementen mit dem Typ `radio` haben, um Verfügbarkeitsoptionen auszuwählen. Die Gruppe `name` sollte `availability` sein.

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

Sie sollten ein `textarea`-Element mit der ID `message` haben, um eine Nachricht einzugeben.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

Sie sollten jedes `input`-, `select`- und `textarea`-Element mit einem `label`-Element verknüpfen.

```js
const els = document.querySelectorAll('input, select, textarea');
assert.isNotEmpty(els);
els.forEach(el => {
  const label = document.querySelector(`label[for="${el.id}"]`) || el.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

Sie sollten ein `button`-Element mit dem Typ `submit` zum Absenden des Formulars haben.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

Sie sollten eine `:focus`-Pseudoklasse zu den `input`- und `textarea`-Elementen hinzufügen, um deren Rahmenfarbe beim Fokus zu ändern. Verwenden Sie einen Listen-Selektor in der angegebenen Reihenfolge.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'));
```

Sie sollten den Selektor für die `:focus`-Pseudoklasse für `input` und `textarea` verwenden, um auch den Standardrahmen beim Fokus zu entfernen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('outline'), 'none');
```

Die `input`-, `select`- und `textarea`-Elemente sollten eine `:invalid`-Pseudoklasse haben, die die Rahmenfarbe rot ändert, wenn eine ungültige Eingabe erkannt wird. Verwenden Sie einen Listen-Selektor in der angegebenen Reihenfolge.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

Die `input`-, `select`- und `textarea`-Elemente sollten eine `:valid`-Pseudoklasse haben, die die Rahmenfarbe grün ändert, wenn eine gültige Eingabe erfolgt. Verwenden Sie einen Listen-Selektor in der angegebenen Reihenfolge.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

Das `button`-Element sollte eine `:hover`-Pseudoklasse haben, die die Hintergrundfarbe beim Darüberfahren mit der Maus ändert.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

Sie sollten die `:checked`-Pseudoklasse auf `.radio-group input[type="radio"]` verwenden, um eine Rahmenfarbe hinzuzufügen, wenn die Optionsschaltfläche ausgewählt ist.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

Sie sollten die `:checked`-Pseudoklasse auf `.radio-group input[type="radio"]` verwenden, um eine Hintergrundfarbe hinzuzufügen, wenn die Optionsschaltfläche ausgewählt ist.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

Sie sollten die `:checked`-Pseudoklasse auf `.radio-group input[type="radio"]` verwenden, um einen Schatten hinzuzufügen, wenn die Optionsschaltfläche ausgewählt ist.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

Sie sollten die `:checked`-Pseudoklasse auf Optionsschaltflächen verwenden, um die Textfarbe der zugehörigen `label` zu ändern, wenn die Option ausgewählt ist.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

Fügen Sie eine `:first-of-type`-Pseudoklasse zu den `input`-Elementen hinzu, um das erste Eingabefeld anders zu gestalten.

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
