---
id: 698367b9d1c6914a22095aec
title: "implementare l'algoritmo di ordinamento bubble sort"
challengeType: 26
dashedName: implement-the-bubble-sort-algorithm
---

# --description--

In questo laboratorio implementerai l'algoritmo di ordinamento bubble sort. Inizia dall'inizio di un array non ordinato e "fa risalire" i valori non ordinati verso la fine, iterando attraverso l'array finché non è completamente ordinato. Lo fa confrontando elementi adiacenti e scambiandoli se sono fuori ordine. Il metodo continua a iterare sull'array finché non avvengono più scambi, a quel punto l'array è ordinato.

Questo metodo richiede più iterazioni sull'array e nei casi medi e peggiori ha complessità temporale quadratica. Pur essendo semplice, di solito è poco pratico nella maggior parte delle situazioni.

**Obiettivo:** soddisfa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Scrivi una funzione `bubbleSort` che prende come input un array di numeri interi e restituisce un array di questi numeri ordinati dal più piccolo al più grande.

# --hints--

Dovresti avere una funzione chiamata `bubbleSort`.

```js
assert.isFunction(bubbleSort);
```

`bubbleSort` dovrebbe restituire un array ordinato (dal più piccolo al più grande).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    bubbleSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`bubbleSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` dovrebbe restituire un array che non cambia se non nell'ordine.

```js
assert.sameMembers(
  bubbleSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`bubbleSort` non dovrebbe usare il metodo integrato `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    bubbleSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function bubbleSort(array) {
  for (let i = 0; i < array.length; i++) {
    let swapped = false;
    for (let j = 1; j < array.length; j++) {
      if (array[j - 1] > array[j]) {
        let temp = array[j-1];
        array[j-1] =  array[j];
        array[j] = temp;
        swapped = true;
      }
    }
    if (swapped === false) {
      break;
    }
  }
  return array;
}
```
