---
id: 66be24cb4144f955b6bcc550
title: Construye una Página de Baraja de Cartas
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Instrucciones:**

1. Deberías construir una página web que muestre al menos tres cartas de juego.
1. Debes tener un elemento `main` con un ID de `playing-cards`.
1. Dentro de tu elemento `#playing-cards`, deberías tener al menos tres elementos `div`, cada uno con una clase de `card`.
1. Dentro de cada elemento `.card`, deberías tener tres elementos `div`, el primero con una clase de `left`, el segundo con una clase de `middle` y el tercero con una clase de `right`.
1. Tu elemento `#playing-cards` debería usar flexbox para centrar horizontalmente sus elementos secundarios, permitir que se envuelvan, y colocar un `20px` de espacio entre ellos.
1. Cada uno de tus elementos `.card` debería usar flexbox para justificar sus elementos secundarios usando `space-between`.
1. Cada uno de tus elementos `.left` debería usar propiedades de elemento flexbox para alinearse al inicio de su contenedor padre.
1. Cada uno de tus elementos `.middle` debería usar propiedades de elemento flexbox para alinearse en el centro de su contenedor padre.
1. Cada uno de tus elementos `.right` debería usar propiedades de elemento flexbox para alinearse al final de su contenedor padre.
1. Cada uno de tus elementos `.middle` debería usar flexbox para mostrar sus elementos secundarios en una columna.

Aquí hay algunos caracteres que puedes copiar y pegar para usar en tus cartas si quieres: `♠`, `♣`, `♥` y `♦`.

# --hints--

Deberías tener un elemento `main` con un `id` de `playing-cards`.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

Deberías tener al menos tres elementos `div` con una clase de `card` dentro de tu elemento `#playing-cards`.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

Cada uno de tus elementos `.card` debería tener un ancho y una altura.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

Cada uno de tus elementos `.card` debería tener exactamente tres elementos `div` como hijos.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

Deberías tener un `div` con una clase de `left` dentro de cada elemento `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

Deberías tener un `div` con una clase de `middle` dentro de cada elemento `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

Deberías tener un `div` con una clase de `right` dentro de cada elemento `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

Deberías tener un selector `#playing-cards` que establezca los elementos `display` en `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

Deberías tener un selector `#playing-cards` que establezca los elementos `justify-content` en `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

Deberías tener un selector `#playing-cards` que establezca los elementos `flex-wrap` en `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

Deberías tener un selector `#playing-cards` que establezca los elementos `gap` en `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

Deberías tener un selector `.card` que establezca los elementos `display` en `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

Deberías tener un selector `.card` que establezca los elementos `justify-content` en `space-between`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

Deberías tener un selector `.left` que establezca los elementos `align-self` en `flex-start`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

Deberías tener un selector `.middle` que establezca los elementos `align-self` en `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

Deberías tener un selector `.right` que establezca los elementos `align-self` en `flex-end`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

Deberías tener un selector `.middle` que establezca los elementos `display` en `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

Deberías tener un `flex-direction` establecido en `column` en tu clase `.middle`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
