---
id: 641d9a19bff38d34d5a5edb8
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将学习如何使用 Fetch API 从外部来源装载数据并将其显示在充满 freeCodeCamp 作者的页面上。

本次研讨会的所有 `HTML` 和 `CSS` 已为你提供。你可以查看这两个 `file` 来熟悉它们。

准备好后，首先访问 `#author-container` 和 `#load-more-btn` 元素。记住你可以使用 `getElementById` 或 `querySelector` 方法来完成此操作。

然后将这些元素分别赋值给名为 `authorContainer` 和 `loadMoreBtn` 的变量。

# --hints--

你应该有一个名为 `authorContainer` 的变量。

```js
assert.isDefined(authorContainer);
```

你应该将 `#author-container` 元素赋值给 `authorContainer` 变量。

```js
assert.strictEqual(authorContainer, document.querySelector('#author-container'));
```

你应该有一个名为 `loadMoreBtn` 的变量。

```js
assert.isDefined(loadMoreBtn);
```

你应该将 `#load-more-btn` 元素赋值给 `loadMoreBtn` 变量。


```js
assert.strictEqual(loadMoreBtn, document.querySelector('#load-more-btn'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
