---
id: 69b83e35f19ba26ba1fa517a
title: Construir un secuenciador de semáforo
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

En este laboratorio, simularás ciclos configurables de semáforos y registrarás anomalías.

Trabajarás con objetos de configuración que describen las fases de un semáforo. Cada objeto de configuración tiene las siguientes propiedades:

- `fault`: una bandera booleana que provoca la terminación anticipada cuando es `true`.
- `phases`: un arreglo de objetos fase.

Cada objeto fase dentro de `phases` tiene las siguientes propiedades:

- `color`: una cadena que representa el color de la luz (`"green"`, `"yellow"` o `"red"`).
- `duration`: un entero positivo que representa cuánto dura la fase en segundos.

Puedes referirte a los objetos `config1`, `config2`, `config3` y `config4` proporcionados como ejemplos de posibles objetos de configuración.

**Nota**: No agregues declaraciones adicionales de `console.log()`, ya que pueden hacer que las pruebas fallen.

**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

# Historias de usuario

1. Debes tener una función llamada `runSequence` con dos parámetros: `config` y `cycles`, donde `config` representa un objeto de configuración y `cycles` representa el número máximo de veces que la secuencia puede ejecutarse antes de detenerse.

1. Debes implementar `runSequence(config, cycles)` usando un ciclo `for` o `while` para iterar a través de cada fase durante el número dado de ciclos.

1. La función `runSequence` debe:
    - Registrar `No phases found` y retornar inmediatamente si `config.phases` está vacío.
    - Registrar `Faulted phase!` y detener la simulación anticipadamente si `config.fault` está en `true`.
    - Registrar `Invalid phase detected` si `duration <= 0`.
    - Registrar `Switching to [color] for [duration] s` para cada fase válida. Reemplaza `[color]` y `[duration]` con las propiedades correspondientes del objeto fase.
    - Por ejemplo, `runSequence(config1, 1)` debería registrar:

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. Debes tener una función llamada `generateTimeline` con dos parámetros: `config` y `cycles`.

1. La función `generateTimeline` debe:
    - Registrar el tiempo acumulado transcurrido después de cada fase a través de los ciclos en un arreglo, sumando la `duration` de cada fase al total acumulado mientras iteras.
    - Procesar todas las fases con fallas y fases inválidas sin validación, incluso si `config.fault` es `true` o `duration <= 0`.
    - Retornar el arreglo de marcas de tiempo acumuladas.
    - Por ejemplo, `generateTimeline(config1, 1)` debería retornar el arreglo `[5, 7, 11]`.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

No debes hacer cambios en los objetos `config1`, `config2`, `config3` y `config4`.

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

Debes tener una función llamada `runSequence`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.runSequence);
```

Tu función `runSequence` debe aceptar dos parámetros, `config` y `cycles`.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.runSequence?.parameters;
assert.equal(parameters?.length, 2);
assert.equal(parameters?.[0]?.toString(), "config");
assert.equal(parameters?.[1]?.toString(), "cycles");
```

`runSequence(config1, 1)` debería registrar `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` en orden.

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` debería registrar `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`, `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` en orden.

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` debería registrar `Switching to red for 3 s`, `Invalid phase detected`, `Switching to green for 6 s` en orden.

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` debería registrar `Faulted phase!`.

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` debería registrar `No phases found`.

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

Debes tener una función llamada `generateTimeline`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.generateTimeline);
```

Tu función `generateTimeline` debe aceptar dos parámetros, `config` y `cycles`.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.generateTimeline?.parameters;
assert.equal(parameters?.length, 2);
assert.equal(parameters?.[0]?.toString(), "config");
assert.equal(parameters?.[1]?.toString(), "cycles");
```

`generateTimeline(config1, 1)` debería retornar `[5, 7, 11]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` debería retornar `[5, 7, 11, 16, 18, 22]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` debería retornar `[3, 1, 7, 10, 8, 14]`.

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` debería retornar `[5, 7, 13]`.

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` debería retornar `[]`.

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
