---
id: 66da326c02141df538f29ba5
title: Створіть сторінку оформлення замовлення
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати елемент `h1` з текстом `Checkout`.
1. Ви повинні мати два елементи `section` одразу після елемента `h1`.
1. Ви повинні мати елемент `h2` з текстом `Your Cart` у першому розділі.
1. Ви повинні мати зображення елемента у першому розділі з відповідним альтернативним текстом. Ви можете використати це зображення, якщо хочете: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. Ви повинні мати елемент `h2` з текстом `Payment Information` у другому розділі.
1. Ви повинні мати елемент `form` у другому розділі.
1. Ви повинні мати input з `id` і `name` зі значенням `card-name`, а також `type` зі значенням `text` у вашій формі та `label`, пов’язаний з ним.
1. Ви повинні мати input з `id` і `name` зі значенням `card-number`, а також `type` зі значенням `text` у вашій формі та `label`, пов’язаний з ним.
1. Ви повинні мати принаймні два елементи `input` з атрибутом `required`.
1. Ви повинні включити елемент `span` з текстом `*` і `aria-hidden`, встановленим на `true`, всередині елемента `label` для кожного обов’язкового input, щоб обов’язкові поля були візуально позначені.
1. Ви повинні мати елемент `p` з допоміжним текстом, який пояснює потрібний формат номера картки, розміщений одразу після input для номера картки. `p` повинен мати `id` зі значенням `card-number-help` і бути посиланням для input номера картки за допомогою `aria-describedby`.

# --hints--

Ви повинні мати елемент `h1` з текстом `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

На вашій сторінці повинен бути лише один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Після елемента `section` має бути принаймні два елементи `h1`.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

У першому розділі повинен бути елемент `h2` з текстом `Your Cart`.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

У першому розділі має бути принаймні одне зображення з альтернативним текстом.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

У другому розділі повинен бути елемент `h2` з текстом `Payment Information`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

У другому елементі `form` має бути елемент `section`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

У вашій формі повинен бути елемент `input` з `id` і `name`, встановленими на `card-name`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

Ваш input для імені на картці повинен мати `type` зі значенням `text`.

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

У вашій формі повинен бути елемент `input` з `id` і `name`, встановленими на `card-number`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

Ваш input для номера картки повинен мати `type` зі значенням `text`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

Усі ваші елементи `input`, які не є `type` зі значенням `submit`, повинні мати пов’язаний з ними елемент `label`.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

Ви повинні мати принаймні два елементи `input` з атрибутом `required`.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

Ви повинні включити елемент `span` з текстом `*` всередині елемента `label` для кожного обов’язкового input.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

Ваші елементи `span` повинні мати `aria-hidden`, встановлений на `true`.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

Одразу після input для номера картки повинен бути елемент `p` з `id` зі значенням `card-number-help`.

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

Ваш допоміжний текст для номера картки не повинен бути порожнім.

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

Ваш input для номера картки повинен мати `aria-describedby`, встановлений на `card-number-help`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
