---
id: 6823e637c1c0ed56f781b4fc
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Нарешті, JavaScript робить вашу сторінку інтерактивною — він дозволяє вказати сторінці, що робити, коли хтось натискає на кнопку, відправляє форму або виконує багато інших дій.

Для останнього кроку практичного заняття додайте ще один `h3` та `p` елемент на сторінку, що описує JavaScript.

Спочатку додайте `h3` елемент із текстом:

```md
Introduction to JavaScript
```

Потім, нижче цього `h3` елемента, додайте `p` елемент із таким текстом:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

Ви повинні додати третій `h3` елемент на сторінку.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Текст вашого `h3` елемента має бути `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

Ви повинні додати четвертий `p` елемент на сторінку.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Ваш `p` елемент має містити текст `JavaScript adds interactivity to a webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
