---
id: 6823ac607bfdbc46331b2559
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML означає HyperText Markup Language. Це код, який визначає структуру та вміст сторінки. Це ваш редактор коду, де ви писатимете HTML.

Знайдіть лінію 1 у редакторі та введіть цей текст:

```md
Welcome to freeCodeCamp
```

Коли закінчите, натисніть кнопку «перевірити свій код», щоб дізнатися, чи все правильно.

# --hints--

У вашому редакторі має бути текст `Welcome to freeCodeCamp`. Перевірте правильність написання.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
