---
id: 6823d9ac8bdc3853df65a1ff
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

В HTML є шість заголовків: `h1` до `h6`. Вони використовуються, щоб показати важливість розділів на вашій сторінці, причому `h1` є найважливішим, а `h6` — найменш важливим.

Під вашим елементом `p` додайте заголовок `h3` з текстом:

```md
Introduction to HTML
```

# --hints--

Ваш елемент `h3` має мати відкриваючий тег `<h3>`.

```js
assert.exists(document.querySelector("h3"));
```

Ваш елемент `h3` має мати закриваючий тег `</h3>`.

```js
assert.match(code, /<\/h3\s*\>/);
```

Текст вашого елемента `h3` має бути `Introduction to HTML`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
