---
id: 6823d6244511f252c8300eed
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Коли потрібно додати абзац на вебсторінку, ви можете використати елемент `p` ось так:

```html
<p>This is a paragraph element.</p>
```

Перетворіть `Learn the skills to become a full-stack developer` на елемент абзацу.

# --hints--

Ваш елемент `p` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

Ваш елемент `p` має мати кінцевий тег. Кінцеві теги мають такий синтаксис: `</elementName>`.

```js
assert.match(code, /<\/p\s*\>/);
```

Текст вашого елемента `p` має бути `Learn the skills to become a full-stack developer`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

Ваш елемент `p` має бути розташований нижче за ваш елемент `h2`.

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
