---
id: 6823e036ea4b71553558c01b
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Зверніть увагу, що кожен заголовок виглядає трохи інакше у попередньому перегляді. Те саме стосується абзацу.

Далі додайте ще один елемент `p`, який відображатиме наступний текст:

```md
HTML represents the content and structure of a webpage
```

# --hints--

У вас має бути другий елемент `p`.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Ваш елемент `p` повинен містити текст `HTML represents the content and structure of a webpage`. Ви можете скопіювати текст із інструкцій.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
