---
id: 682ba2318000b62f179bdf04
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML складається з елементів. Перший, який ви використаєте, — це елемент `h1`:

```html
<h1>Welcome to freeCodeCamp</h1>
```

Він починається з початкового тегу (`<h1>`), закінчується кінцевим тегом (`</h1>`) і має текст, який відображатиме, між тегами.

Перетворіть свій текст `Welcome to freeCodeCamp` на елемент `h1`, додавши початковий тег перед ним і кінцевий тег після нього.

# --hints--

Ваш елемент `h1` має мати початковий тег `<h1>`.

```js
assert.exists(document.querySelector("h1"));
```

Ваш елемент `h1` має мати кінцевий тег `</h1>`.

```js
assert.match(code, /<\/h1\s*\>/);
```

Ваш елемент `h1` має виглядати так: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
