---
id: 682cd206883fc7b25eb539c5
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Зверніть увагу, що HTML, який ви пишете в редакторі, відображається у попередньому перегляді. У цьому практичному занятті ви напишете HTML для часткової сторінки навчальної програми.

Під вашим елементом `h1` на порожньому рядку введіть наступне:

```md
Full-Stack Curriculum
```

# --hints--

У вашому редакторі має бути текст `Full-Stack Curriculum`. Перевірте правильність написання.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
