---
id: 683921c4769fd23dbadec2fe
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

HTML визначає структуру і вміст вебсторінки, а ось CSS надає їй стилю (а саме — кольори, шрифти, інтервали і макети).

Під `h3`, який ви щойно додали, додайте ще один абзац з таким текстом:

```md
CSS is used to style a webpage
```

# --hints--

Додайте другий елемент `h3` до сторінки.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

Текстом елемента `h3` має бути `Introduction to CSS`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

Додайте третій елемент `p` до сторінки.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Елемент `p` повинен мати текст `CSS is used to style a webpage`. Ви можете скопіювати текст з інструкцій.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
