---
id: 692fb81e768aab5a1a6a0381
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

在 `Spider` 类中，创建一个构造函数，它接受一个类型为 `emojiElement` 的单个参数 `HTMLParagraphElement`。

在构造函数内部，使用 `super` 参数调用 `emojiElement` 函数，以调用父 `Bug` 类的构造函数。然后，将 `emoji` 属性设置为字串值 `"🕷️"`。

# --hints--

你的 `Spider` 类应该有一个构造函数。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.classes.Spider.classConstructor);
```

你的 `Spider` 构造函数应该有一个名为 `emojiElement`、类型为 `HTMLParagraphElement` 的参数。

```js
assert.match(__helpers.removeJSComments(code), /class\s+Spider[\s\S]*?constructor\s*\(\s*emojiElement\s*:\s*HTMLParagraphElement\s*\)/);
```

你的 `Spider` 构造函数应该调用 `super` 并传入 `emojiElement` 参数。

```js
assert.match(__helpers.removeJSComments(code), /class\s+Spider[\s\S]*?super\s*\(\s*emojiElement\s*\)/);
```

你的 `Spider` 构造函数应该将 `this.emoji` 设置为 `"🕷️"`。

```js
assert.match(__helpers.removeJSComments(code), /this\.emoji\s*=\s*["']🕷️["']/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bug Emoji Picker</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <form id="bug-form">
        <label for="species">Species:</Label>
        <select id="species">
          <option value="bee" selected>🐝 Bee</option>
          <option value="spider">🕷️ Spider</option>
        </select>
      </form>
      <div>
        <p id="bug-emoji">🐝</p>
      </div>
    </main>
  </body>
  <script src="index.ts"></script>
</html>
```

```css
:root {
  --gradient-back: #987284;
  --gradient-left: #F9B5AC;
  --gradient-right: #AA7674;
  --card: #FFFFFF;
  --card-shadow: 0 10px 30px rgba(30, 41, 59, 0.12);
  --font: #222;
}

html {
  width: 100%;
  min-height: 100%;
  background: radial-gradient(100% 100% at 10% 10%, var(--gradient-left), transparent 60%),
              radial-gradient(90% 85% at 90% 0%, var(--gradient-right), transparent 55%),
              linear-gradient(180deg, var(--gradient-back), #ffffff);
}

body {
  margin: 5rem auto;
  padding: 2rem 3rem;
  width: fit-content;
  min-width: 15%;
  color: var(--font);
  font: 500 16px/1.4 system-ui;
  background: rgba(255, 255, 255, 0.3);
  border: 1px solid rgba(255, 255, 255, 0.3);
  border-radius: 15%;
  box-shadow: var(--card-shadow);

  backdrop-filter: blue(12px);
  -webkit-backdrop-filter: blur(12px);
}

select, input {
  appearance: none;
  -webkit-appearance: none;
  -moz-appearance: none;

  background: rgba(255, 255, 255, 0.25);
  border-radius: 12px;
  padding: 0.1rem 2rem 0.1rem 0.8rem;
  font-size: 1rem;
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

form { 
  display: flex;
  flex-direction: column;
}

#bug-emoji {
  margin: 0 auto;
  width: fit-content;
  font-size: 128px;
}

#bug-name {
  margin: 0 auto;
  width: fit-content;
  font-size: 24;
  font-weight: 700;
}
```

```ts
abstract class Bug<T> {
  emoji!: T;
  emojiElement!: HTMLParagraphElement;
  constructor(emojiElement: HTMLParagraphElement) {
    this.emojiElement = emojiElement;
  }

  abstract render(): void;
}

class Bee extends Bug<string> {
  constructor(emojiElement: HTMLParagraphElement) {
    super(emojiElement);
    this.emoji = "🐝";
  }

  override render() {
    this.emojiElement.innerText = this.emoji;
  }
}

class Spider extends Bug<string> {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
