---
id: 68f1f993460f06a64c033061
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

`poster` 属性是视频的缩略图。想想你在 YouTube 上观看的视频。它在视频下载时显示。如果省略该属性，下载阶段将显示第一个视频帧。

现在，为你的 `video` 元素添加值为 `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` 的 `poster` 属性。

# --hints--

你的 `video` 元素应具有值为 `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` 的 `poster` 属性。

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
