---
id: 68f1f371fefd448b4f70b898
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在之前的课程中，你了解了 `video` 元素可用的不同属性。`width` 属性决定视频的宽度（以像素为单位）。

为你的 `video` 元素添加值为 `640` 的 `width` 属性。

# --hints--

你的 `video` 元素应具有属性 `width`，其值为 `640`。

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
