---
id: 68f1eb9ad0a7575e4dbc31d5
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

接下来，在 `h1` 下方创建一个 `video` 元素。在接下来的几个步骤中，你将为视频播放器添加必要的属性以使其正常工作。 

# --hints--

你应该有一个 `video` 元素。

```js
assert.exists(document.querySelector('video'));
```

新的 `video` 元素应位于 `h1` 下方。

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
