---
id: 68f1f57723debf9383c75947
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

`loop` 属性将在播放完成后重新启动视频。想象一个重复播放的互联网表情包。省略 `loop` 属性将使视频只播放一次。

`loop` 属性是一个布尔属性，不需要值。

为你的 `video` 元素添加 `loop` 属性。

# --hints--

你的 `video` 元素应具有 `loop` 属性。

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
