---
id: 68f1fac259a4afabb478b960
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

你可能已经注意到你没有链接到实际的 `video`。你将在下一阶段完成这一步。关于视频文件类型，不同浏览器的支持存在差异。为此，你可以在 `video` 元素内使用 `source` 元素，浏览器将选择第一个兼容的 `source`。

这是一个 `source` 元素的示例：

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

`source` 元素是一个空元素，因此它没有闭合标签。 

在你的 `video` 元素内添加一个 `source` 元素。 


# --hints--

你的 `video` 元素内应包含一个 `source` 元素。

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
