---
id: 68f1f8c248306da204655117
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

`muted` 属性将在初始播放时静音音频。如果你启用了 `controls`，用户将能够取消静音音频。省略 `muted` 属性将在初始播放时播放音频。

`muted` 属性是一个布尔属性，不需要值。

为你的 `video` 元素添加 `muted` 属性。

# --hints--

你的 `video` 元素应具有 `muted` 属性。

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
