---
id: 68f1f6f97af9749a42260a43
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

`controls` 属性为 `video` 元素提供播放控件，包括播放、倒带和卷控制。 

`controls` 属性是一个布尔属性，不需要值。

为你的 `video` 元素添加 `controls` 属性。

现在你应该能看到页面上显示的 `video` 元素。

# --hints--

你的 `video` 元素应具有 `controls` 属性。

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
