---
id: 68f26d52b13e0f5b6e2d2e09
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

另一个常见的 MIME 类型是 `video/webm` MIME 类型。

WebM 是由 Google 开发的开源音视频媒体文件格式，主要用于基于网页的媒体内容。它支持 VP8、VP9 和 AV1 等视频编码，以及 Vorbis 和 Opus 等音频编码，使其成为 HTML5 视频和音频元素的热门选择。

在你的第一个 `source` 元素下方，添加另一个 `source` 元素，并为其添加一个 `src` 属性，值为 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`，以及一个 `type` 属性，值为 `video/webm`。

# --hints--

你应该有第二个 `source` 元素。

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

你的第二个 `source` 元素应具有 `src` 属性。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

你的第二个 `source` 元素应具有一个 `src` 属性，值为 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

你的第二个 `source` 元素应具有 `type` 属性。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

你的第二个 `source` 元素应具有值为 `video/webm` 的 `type` 属性。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
