---
id: 68f27139bddf9f6bf2821acf
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

你将添加的最后一个 `source` 元素将用于 `video/quicktime` MIME 类型。

QuickTime 是由 Apple 创建的可扩展多媒体架构，支持播放、流式传输、编码和转码多种数字媒体格式。虽然不如 MP4 格式流行，但你可能需要它来支持遗留应用。

在你的第三个 `source` 元素下方，添加第四个 `source` 元素，并为其添加 `src` 属性，值为 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`，以及 `type` 属性，值为 `video/quicktime`。

恭喜你！你完成了超文本标记语言 `Video` 播放器研讨会。

# --hints--

你应该有第四个 `source` 元素。

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

你的第四个 `source` 元素应具有 `src` 属性。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

你的第四个 `source` 元素应具有一个 `src` 属性，值为 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

你的第四个 `source` 元素应具有 `type` 属性。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

你的第四个 `source` 元素应具有值为 `video/quicktime` 的 `type` 属性。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
