---
id: 68f2701b2160c765f23b60f8
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

另一个常见的 MIME 类型是 `video/ogg` MIME 类型。

Ogg 是一种数字的多媒体容器格式，旨在为数字多媒体的高效流式传输和操作提供支持。它由 Xiph.Org 基金会维护，免费且开放，不受软件专利限制。其名称派生自计算机游戏 Netrek 中的行话“ogging”。

在你的第二个 `source` 元素下方，添加第三个 `source` 元素，并为其添加一个 `src` 属性，值为 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg`，以及一个 `type` 属性，值为 `video/ogg`。

# --hints--

你应该有第三个 `source` 元素。

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

你的第三个 `source` 元素应具有 `src` 属性。

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

你的第三个 `source` 元素应具有一个 `src` 属性，其值为 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg`。

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

你的第三个 `source` 元素应具有 `type` 属性。

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

你的第三个 `source` 元素应具有值为 `video/ogg` 的 `type` 属性。

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
