---
id: 68f26a5f79e97c5032c10642
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

你使用了一个带有 `mp4` 扩展名的视频文件，你需要告诉浏览器，以便它知道如何读取该文件。

你将使用 `type` 属性来指定 `video/mp4` MIME 类型。

MIME（多用途互联网邮件扩展）是一种标准，用于描述除 ASCII 文本之外的其他形式的文档，例如音频、视频和图像。

MP4，正式名称为 MPEG-4 Part 14，是一种数字的多媒体容器格式。它广泛用于存储视频和音频，但也可以包含其他数据类型，如字幕和静态图像。MP4 `file` 设计用于通过互联网流式传输，并且兼容许多设备和平台。

现在，为你的元素添加 `type` 属性和值 `video/mp4`。

# --hints--

你的 `source` 元素应具有 `type` 属性。

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

你的 `source` 元素应具有值为 `video/mp4` 的 `type` 属性。

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--
      
--fcc-editable-region--
    >
  </video>
</body>
</html>
```
