---
id: 68f268880219c2485bd71d2a
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

要指定视频的媒体资源，你需要为 `source` 元素添加 `src` 属性。

为你的元素添加 `src` 属性，值为 `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`。

# --hints--

你的 `source` 元素应具有 `src` 属性。

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('src'));
```

你的 `source` 元素应具有 `src` 属性，值为 `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`。

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('src'),
'https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    <source>
--fcc-editable-region--
  </video>
</body>
</html>
```
