---
id: 5d822fd413a79914d39e98ca
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`html` の下に開始タグと終了タグの `DOCTYPE` タグを追加して、コードを記述し始める場所を作ってください。言語は英語に設定してください。

# --hints--

`DOCTYPE` の宣言はHTMLの最初にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 要素には開始タグがあり、`lang` 属性の値は `en` である必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

`html` タグは正しい順序で配置してください。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

`html` 要素は1つだけにしてください。

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
