---
id: 5d822fd413a79914d39e98c9
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

CSS Variables Skyline プロジェクトへようこそ！まず、ブラウザーがどの種類のドキュメントを読み込んでいるかを認識できるように、ドキュメントの先頭に `!DOCTYPE html` 宣言を追加してください。

# --hints--

コードには `DOCTYPE` 参照が含まれている必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れてください。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプは `html` に設定してください。

```js
assert(code.match(/html/gi));
```

`DOCTYPE` 宣言はタイプの後に `>` で閉じてください。

```js
assert(code.match(/html\s*>/gi));
```

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
