---
id: 5d822fd413a79914d39e98d2
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

背景の建物コンテナ内に、クラス名が`div`の`bb1`をネストしてください。`styles.css`ファイルを開き、`.bb1`に`width`の`10%`と`height`の`70%`を付けてください。"bb"は"background building"の略で、これが最初の建物になります。

# --hints--

新しい`div`要素を作成してください。

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

新しい`div`に`bb1`のクラスを付けてください。

```js
assert.exists(document.querySelector('div.bb1'));
```

`.bb1`のクラスセレクターを使って要素をスタイルしてください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

`.bb1`要素に`width`の`10%`を付けてください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

`.bb1`要素に`height`の`70%`を付けてください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
    
```
