---
id: 5d822fd413a79914d39e98ce
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

すべてに`box-sizing`の`border-box`を追加してください。これにより、追加したボーダーが要素のサイズに影響しなくなります。

# --hints--

`box-sizing`プロパティを使う必要があります。

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

既存の`*`セレクターを活用してください。

```js
// Two selectors create two CSSStyleRule objects
assert.equal(new __helpers.CSSHelp(document).getStyleDeclarations('*').length, 1);
```

すべての要素に`box-sizing`の`border-box`を設定してください。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```
