---
id: 5d822fd413a79914d39e98d4
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

建物の各部分に、これらの値を持つ`width`と`height`のプロパティを設定してください：`70%`と`10%`は`.bb1a`へ、`80%`と`10%`は`.bb1b`へ、`90%`と`10%`は`.bb1c`へ、`100%`と`70%`は`.bb1d`へ。これらのパーセンテージは親要素に対する割合であることを覚えておいてください。また、高さの合計は100%になり、コンテナを縦方向に埋めることになります。

# --hints--

`.bb1a`をスタイル設定するにはクラスセレクターを使ってください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

`.bb1a`に`width`を`70%`で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

`.bb1a`に`height`を`10%`で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

`.bb1b`をスタイル設定するにはクラスセレクターを使ってください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

`.bb1b`に`width`を`80%`で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

`.bb1b`に`height`を`10%`で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

`.bb1c`をスタイル設定するにはクラスセレクターを使ってください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

`.bb1c`に`width`を`90%`で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

`.bb1c`に`height`を`10%`で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

`.bb1d`をスタイル設定するにはクラスセレクターを使ってください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

`.bb1d`に`width`を`100%`で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

`.bb1d`に`height`を`70%`で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```
