---
id: 5d822fd413a79914d39e98cc
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`head`の中に、`meta`の`charset`を持つ`UTF-8`要素、タイトルが`title`の`City Skyline`要素、そして`link`ファイルにリンクする`styles.css`要素をネストしてください。

# --hints--

`meta`要素の中に`head`要素を作成してください。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('meta'));
```

`meta`タグに`charset`の`UTF-8`を付けてください。

```js
const linkElement = document.querySelector('meta');
const charset = linkElement?.getAttribute("charset").toLowerCase();
assert.strictEqual(charset, 'utf-8');
```

コードには`title`要素が必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title`要素は`head`要素の中にあるべきです。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

プロジェクトのタイトルは`City Skyline`にしてください。

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

タイトルの大文字・小文字や綴りに注意してください。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

コードには`link`要素が必要です。

```js
assert.isNotNull(document.querySelector('link'));
```

既存の`head`タグは変更しないでください。終了タグを削除していないか確認してください。

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

空要素の`link`が1つ必要です。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link`要素は`head`要素の中にあるべきです。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link`要素には値が`rel`の`stylesheet`属性が必要です。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link`要素には値が`href`の`styles.css`属性が必要です。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
