---
id: 5d822fd413a79914d39e98d6
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

これで建物のようなものができました。最初の変数を作成する準備ができています。変数の宣言は二つのダッシュ（`-`）で始まり、名前と値を次のように指定します：`--variable-name: value;`。`bb1` クラスのルール内で、`--building-color1` という名前の変数を作成し、値を `#999` にしてください。

# --hints--

`--building-color1` という新しい変数を作成してください。

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('--building-color1'));
```

`--building-color1` 変数は `.bb1` の中で定義してください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
```

`--building-color1` に `#999` の値を与えてください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1').trim(),'#999');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
