---
id: 69386f3bce81c60f2b3feeb8
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Holen Sie das `select`-Element mit der ID `species` aus dem DOM und speichern Sie es in einer Konstanten namens `selectElement`. Verwenden Sie den `!`-Operator zur Zusicherung, dass das Element nicht `null` ist.

# --hints--

Sie sollten eine Konstante namens `selectElement` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.selectElement);
```

Ihre `selectElement` sollte das Element mit der ID `species` mit `document.querySelector` unter Verwendung eines Typ-Arguments von `HTMLSelectElement` abfragen.

```js
assert.match(__helpers.removeJSComments(code), /querySelector\s*<\s*HTMLSelectElement\s*>\s*\(\s*['"]#species['"]\s*\)/);
```

Ihre `selectElement` sollte den Operator zur Zusicherung der Nicht-Null (`!`) verwenden.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.selectElement.value.hasNonNullAssertion());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bug Emoji Picker</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <form id="bug-form">
        <label for="species">Species:</label>
        <select id="species">
          <option value="bee" selected>🐝 Bee</option>
          <option value="spider">🕷️ Spider</option>
        </select>
      </form>
      <div>
        <p id="bug-emoji">🐝</p>
      </div>
    </main>
  </body>
  <script src="index.ts"></script>
</html>
```

```css
:root {
  --gradient-back: #987284;
  --gradient-left: #F9B5AC;
  --gradient-right: #AA7674;
  --card: #FFFFFF;
  --card-shadow: 0 10px 30px rgba(30, 41, 59, 0.12);
  --font: #222;
}

html {
  width: 100%;
  min-height: 100%;
  background: radial-gradient(100% 100% at 10% 10%, var(--gradient-left), transparent 60%),
              radial-gradient(90% 85% at 90% 0%, var(--gradient-right), transparent 55%),
              linear-gradient(180deg, var(--gradient-back), #ffffff);
}

body {
  margin: 5rem auto;
  padding: 2rem 3rem;
  width: fit-content;
  min-width: 15%;
  color: var(--font);
  font: 500 16px/1.4 system-ui;
  background: rgba(255, 255, 255, 0.3);
  border: 1px solid rgba(255, 255, 255, 0.3);
  border-radius: 15%;
  box-shadow: var(--card-shadow);

  backdrop-filter: blur(12px);
  -webkit-backdrop-filter: blur(12px);
}

select, input {
  appearance: none;
  -webkit-appearance: none;
  -moz-appearance: none;

  background: rgba(255, 255, 255, 0.25);
  border-radius: 12px;
  padding: 0.1rem 2rem 0.1rem 0.8rem;
  font-size: 1rem;
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

form { 
  display: flex;
  flex-direction: column;
}

#bug-emoji {
  margin: 0 auto;
  width: fit-content;
  font-size: 128px;
}

```

```ts
abstract class Bug<T> {
  emoji!: T;
  emojiElement!: HTMLParagraphElement;
  constructor(emojiElement: HTMLParagraphElement) {
    this.emojiElement = emojiElement;
  }

  abstract render(): void;
}

class Bee extends Bug<string> {
  constructor(emojiElement: HTMLParagraphElement) {
    super(emojiElement);
    this.emoji = "🐝";
  }

  override render() {
    this.emojiElement.innerText = this.emoji;
  }
}

class Spider extends Bug<string> {
  constructor(emojiElement: HTMLParagraphElement) {
    super(emojiElement);
    this.emoji = "🕷️";
  }

  override render() {
    this.emojiElement.innerText = this.emoji;
  }
}

function isSelect(element: EventTarget | null): element is HTMLSelectElement {
  return element instanceof HTMLSelectElement;
}

const bugEmojiElement = document.querySelector<HTMLParagraphElement>('#bug-emoji')!;
const bugMap: Record<string, Bug<string>> = {
  "bee": new Bee(bugEmojiElement),
  "spider": new Spider(bugEmojiElement),
};

--fcc-editable-region--

--fcc-editable-region--
```
