---
id: 6140c7e645d8e905819f1dd4
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin with the standard boilerplate. Add your `DOCTYPE` declaration, your `html` element with the language set to English, your `head` and `body` elements.

在 `meta` 元素中添加正確的 `charset`、`title` 元素和鏈接 `./styles.css` 文件的 `link` 元素。

將 `title` 設置爲 `Ferris Wheel`。

# --hints--

你的代碼應該包含 `DOCTYPE` 聲明。

```js
assert.match(code, /<!DOCTYPE/gi);
```

你應該在 `DOCTYPE` 聲明後添加一個空格。

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

您應將文檔類型定義爲 `html` 。

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

在類型後，您應用 `>` 來關閉 `DOCTYPE` 聲明。

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

你的 `html` 元素應該有一個帶有值爲 `en` 的 `lang` 屬性的開始標籤。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

您的 `html` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/html\s*>/);
```

您的 `DOCTYPE` 聲明應位於最前面（HTML的開頭）。

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

您應有一個 `head` 開始標籤。

```js
assert.match(code, /<head\s*>/i);
```

您應有一個 `head` 結束標籤。

```js
assert.match(code, /<\/head\s*>/i);
```

您應有一個 `body` 開始標籤。

```js
assert.match(code, /<body\s*>/i);
```

您應有一個 `body` 結束標籤。

```js
assert.match(code, /<\/body\s*>/i);
```

你的 `body` 元素應該是 `head` 元素之後的第一個元素。

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

你的 `head` 和 `body` 元素應該是兄弟元素。

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

你的 `head` 元素應該是 `html` 元素中的第一個元素。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

你的 `head` 元素應該在 `html` 元素中。

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

你的 `body` 元素應該在 `html` 元素中。

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Your code should have a `meta` element.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Your `meta` element should have a `charset` attribute with the value `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

你的 `meta` 元素應該在 `head` 元素中。

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your project should have a title of `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Remember, the casing and spelling matter for the title.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Your code should have a `link` element.

```js
assert.match(code, /<link/);
```

你應該有一個 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Your `link` element should be within your `head` element.

```js
assert.exists(document.querySelector('head > link'));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
