---
id: 6140c7e645d8e905819f1dd4
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin with the standard boilerplate. Add your `DOCTYPE` declaration, your `html` element with the language set to English, your `head` and `body` elements.

Füge dein `meta`-Element dem richtigen `charset` hinzu, dein `title`-Element, und ein `link`-Element für die `./styles.css`-Datei.

Setze den `title` auf `Ferris Wheel`.

# --hints--

Dein Code sollte die `DOCTYPE`-Deklaration enthalten.

```js
assert.match(code, /<!DOCTYPE/gi);
```

Du solltest ein Leerzeichen nach der `DOCTYPE`-Deklaration einfügen.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Du solltest den Dokumententyp als `html` definieren.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Du solltest die `DOCTYPE`-Deklaration mit `>` nach dem Typ schließen.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Dein `html`-Element sollte ein öffnendes Tag mit einem `lang`-Attribut von `en` haben.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Dein `html`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/html\s*>/);
```

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTML stehen.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Du solltest einen einleitenden `head`-Tag haben.

```js
assert.match(code, /<head\s*>/i);
```

Du solltest einen abschließenden `head`-Tag haben.

```js
assert.match(code, /<\/head\s*>/i);
```

Du solltest einen öffnenden `body`-Tag haben.

```js
assert.match(code, /<body\s*>/i);
```

Du solltest einen abschließenden `body`-Tag haben.

```js
assert.match(code, /<\/body\s*>/i);
```

Dein `body`-Element sollte das erste Element nach dem `head`-Element sein.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Deine `head` und `body`-Elemente sollten Geschwister sein.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Dein `head`-Element sollte das erste Element innerhalb des `html`-Elements sein.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Dein `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Dein `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Your code should have a `meta` element.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Your `meta` element should have a `charset` attribute with the value `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Dein `meta`-Element sollte innerhalb des `head`-Elements liegen.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your project should have a title of `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Remember, the casing and spelling matter for the title.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Your code should have a `link` element.

```js
assert.match(code, /<link/);
```

Du solltest ein `link`-Element haben.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Your `link` element should be within your `head` element.

```js
assert.exists(document.querySelector('head > link'));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
