---
id: 689f7996426c1534fa5bea66
title: 골프 점수 변환기 만들기
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

골프 게임에서 각 홀에는 `par`가 있습니다. 이는 골퍼가 공을 홀에 넣어 플레이를 완료하기 위해 예상되는 평균 `strokes` 수를 의미합니다. `par`보다 `strokes`가 얼마나 높거나 낮은지에 따라 다른 별명이 붙습니다.

이 실습에서는 `par`와 `strokes`를 별명으로 변환하는 함수를 작성할 것입니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `golfScore`라는 이름의 함수를 만들어야 합니다.
1. `golfScore`는 코스의 파와 획수 두 개의 숫자 인수를 받아야 합니다.
1. `golfScore`는 문자열을 반환해야 합니다.
1. `golfScore`는 `"Hole-in-one!"`이 `strokes`일 때 `1`를 반환해야 합니다.
1. `golfScore`는 `"Eagle"`가 `strokes`에서 `par`를 뺀 값보다 작거나 같을 때 `2`를 반환해야 합니다.
1. `golfScore`는 `"Birdie"`이 `strokes`에서 `par`를 뺀 값과 같을 때 `1`를 반환해야 합니다.
1. `golfScore`는 `"Par"`가 `strokes`와 같을 때 `par`를 반환해야 합니다.
1. `golfScore`는 `"Bogey"`가 `strokes`에 `par`를 더한 값과 같을 때 `1`를 반환해야 합니다.
1. `golfScore`는 `"Double Bogey"`가 `strokes`에 `par`를 더한 값과 같을 때 `2`를 반환해야 합니다.
1. `golfScore`는 `"Go Home!"`가 `strokes`에 `par`를 더한 값보다 크거나 같을 때 `3`를 반환해야 합니다.

# --hints--

`golfScore`라는 이름의 함수를 만들어야 합니다.

```js
assert.isFunction(golfScore);
```

`golfScore`는 두 개의 매개변수를 받아야 합니다.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore`는 문자열을 반환해야 합니다.

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)`는 문자열 `Hole-in-one!`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)`는 문자열 `Hole-in-one!`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)`는 문자열 `Hole-in-one!`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)`는 문자열 `Hole-in-one!`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)`는 문자열 `Eagle`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)`는 문자열 `Eagle`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)`는 문자열 `Birdie`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)`는 문자열 `Birdie`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)`는 문자열 `Birdie`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)`는 문자열 `Par`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)`는 문자열 `Par`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)`는 문자열 `Par`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)`는 문자열 `Bogey`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)`는 문자열 `Bogey`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)`는 문자열 `Bogey`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)`는 문자열 `Double Bogey`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)`는 문자열 `Double Bogey`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)`는 문자열 `Double Bogey`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)`는 문자열 `Go Home!`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)`는 문자열 `Go Home!`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)`는 문자열 `Go Home!`를 반환해야 합니다.

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
