---
id: 66ba762af611169359d9d369
title: Construye un Artículo de Periódico
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historia de usuario:**

1. Debe establecer el `font-size` raíz de su documento HTML en `24px`.
1. Debe tener un elemento con una clase de `newspaper` que contenga todos sus otros elementos.
1. Su elemento `.newspaper` debe tener un `font-size` de `16px` y una fuente de `Open Sans` con una fuente de reserva de `sans-serif`.
1. Dentro de su elemento `.newspaper`, debe tener al menos siete elementos más: uno para el nombre del periódico que tiene una clase de `name`, uno para la fecha del artículo con una clase de `date`, uno para el titular con una clase de `headline`, uno para el subtítulo con una clase de `sub-headline`, uno para el autor con una clase de `author`, y dos párrafos, cada uno con una clase de `text`. Todos estos elementos deben estar llenos con la información de su artículo.
1. Su elemento `.name` debe tener un `font-size` que sea el doble del `font-size` del elemento raíz y debe usar la fuente `Times New Roman`, con una fuente de reserva de `serif`.
1. Sus elementos `.name` y `.author` deben utilizar CSS para hacer que todos sus caracteres estén en mayúsculas.
1. Su elemento `.headline` debe tener un `font-size` que sea el doble del `font-size` de su elemento padre y debe estar en negrita.
1. Su elemento `.sub-headline` debe tener un `font-weight` de `100`, un `font-size` que sea `1.5` veces el `font-size` de su elemento padre, y debe estar en cursiva.
1. Su elemento `.author` debe utilizar CSS para que esté en negrita.
1. Sus elementos `.text` deben tener un `text-indent` de `20px`.
1. Sus elementos `.text` deben tener un `line-height` que sea el doble del `font-size` de su elemento padre.
1. La primera letra de sus elementos `.text` debe estar en negrita y ser el doble del `font-size` de su elemento padre. Utilice el selector `::first-letter` para esto.

# --hints--

Debe tener un selector `html` que establezca su `font-size` a `24px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

Debe tener un elemento con una clase de `newspaper`.

```js
assert.exists(document.querySelector('.newspaper'));
```

Debe tener un selector `.newspaper` que establezca el `font-size` de sus elementos a `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

Debe tener un selector `.newspaper` que establezca el `font-family` de sus elementos a `'Open Sans', sans-serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

Debe tener un elemento con una clase de `name` dentro de su elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

Su elemento `.name` debe tener el nombre de su periódico.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

Debe tener un elemento con una clase de `date` dentro de su elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

Su elemento `.date` debe tener la fecha.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

Debe tener un elemento con una clase de `headline` dentro de su elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

Su elemento `.headline` debe tener el titular de su artículo.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

Debe tener un elemento con una clase de `sub-headline` dentro de su elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

Su elemento `.sub-headline` debe tener el subtítulo de su artículo.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

Debe tener un elemento con una clase de `author` dentro de su elemento `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

Su elemento `.author` debe tener el autor de su artículo.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

Debe tener al menos dos elementos de párrafo, cada uno con una clase de `text`, dentro de su elemento `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

Sus elementos `.text` deben tener el texto de su artículo.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

Debe tener un selector `.name` que establezca el `font-size` de sus elementos a `2rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

Debe tener un selector `.name` que establezca el `font-family` de sus elementos a `'Times New Roman', serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

El elemento `.name` debe tener la propiedad `text-transform` establecida en `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

El elemento `.author` debe tener la propiedad `text-transform` establecida en `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Debe tener un selector `.headline` que establezca el `font-size` de sus elementos a `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

Debe tener un selector `.headline` que establezca el `font-weight` de sus elementos a `negrita`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

Debe tener un selector `.sub-headline` que establezca el `font-weight` de sus elementos a `100`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

Debe tener un selector `.sub-headline` que establezca el `font-size` de sus elementos a `1.5em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

Debe tener un selector `.sub-headline` que establezca el `font-style` de sus elementos a `cursiva`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

Debe tener un selector `.author` que establezca el `font-weight` de sus elementos a `negrita`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

Debe tener un selector `.text` que establezca el `text-indent` de sus elementos a `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

Debe tener un selector `.text` que establezca el `line-height` de sus elementos a `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

Debe tener un selector `.text::first-letter` que establezca el `font-weight` de sus elementos a `negrita`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

Debe tener un selector `.text::first-letter` que establezca el `font-size` de sus elementos a `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
