---
id: 642db8c409d9991d0b3b2f0d
title: Hatua ya 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Mradi wako unaanza na kontena la msingi la HTML na baadhi ya CSS zinazolingana. Zoezi lako la kwanza litakuwa kutengeneza kwa njia ya programu seli za jedwali lako la hesabu.

Kipengele cha `window` cha ulimwengu kinawakilisha dirisha la kivinjari (au kichupo). Kina sifa ya `onload` inayokuwezesha kufafanua tabia wakati dirisha limepakia ukurasa mzima, ikiwa ni pamoja na majalada ya mitindo na maandishi ya programu.

Anza kwa kuweka sifa ya `onload` ya `window` kuwa kitendakazi cha mshale kisicho na vigezo. Katika kitendakazi hicho, tangaza kigezo cha `container` na weka thamani yake kwa kupata kipengele kwa `id` ya `"container"`.

# --hints--

Unapaswa kupata upatikanaji wa sifa ya `onload` ya kipengele cha `window`.

```js
assert.match(code, /window\.onload/);
```

Unapaswa kuweka sifa ya `onload` kuwa kitendakazi.

```js
assert.isFunction(window.onload);
```

Unapaswa kutumia sintaksia ya mshale.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

Kitendakazi chako cha `onload` hakipaswi kuchukua vigezo vyovyote.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

Unapaswa kutangaza kigezo cha `container` katika kitendakazi chako cha `onload`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

Kigezo chako cha `container` kinapaswa kutangazwa kwa `const`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

Unapaswa kutumia `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

Unapaswa kupata kipengele kwa `id` ya `"container"`.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

Unapaswa kuweka kipengele cha `#container` kwa `container`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
