---
id: 642df9df4b5216350de7b0d2
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Utahitaji kitendakazi kuunda safu ya nambari.

Tangaza kitendakazi tupu cha `range` kinachopokea kigezo cha `start` na `end`. Tumia kijenzi cha `Array()` na rudisha safu tupu kwa njia isiyo ya moja kwa moja.

# --hints--

Unapaswa kutangaza kinachobadilika cha `range`.

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

Unapaswa kutumia `const` kutangaza kinachobadilika chako cha `range`.

```js
assert.match(code, /const\s+range/);
```

Kinachobadilika chako cha `range` kinapaswa kuwa kitendakazi.

```js
assert.isFunction(range);
```

Kitendakazi chako cha `range` kinapaswa kutumia sintaksia ya mshale.

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

Kitendakazi chako cha `range` kinapaswa kupokea kigezo cha `start` kwanza.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

Kitendakazi chako cha `range` kinapaswa kupokea kigezo cha `end` pili.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

Kitendakazi chako cha `range` kinapaswa kutumia kurudisha isiyo ya moja kwa moja. Kumbuka hii inamaanisha hutatumia mabano ya curly.

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

Kitendakazi chako cha `range` kinapaswa kutumia kijenzi cha `Array()`. Kijenzi cha vitu vya msingi hakihitaji neno kuu la `new`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

Hutaki kupeleka chochote kwa kijenzi cha `Array()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
