---
id: 642ddfdea4200e313f80a4b6
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Kumbuka kuwa kitu cha `document` kina njia ya `.createElement()` inayokuwezesha kuunda vipengele vipya vya HTML kwa njia ya mabadiliko.

Katika kitendakazi chako cha `createLabel`, tangaza kigezo cha `label` na weka thamani kipengele kipya cha `div`.

# --hints--

Unapaswa kutangaza kigezo cha `label` katika kitendakazi chako cha `createLabel`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

Kigezo chako cha `label` kinapaswa kutangazwa kwa `const`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

Unapaswa kutumia njia ya `.createElement()` ya kitu cha `document`.

```js
assert.match(code, /document\.createElement\(/);
```

Unapaswa kupitisha mfuatano wa herufi `"div"` kwa njia ya `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

Unapaswa kuweka kipengele chako kipya cha `div` kwa `label`.

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
