---
id: 642e004130958c3975aa3a4a
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Kwa sasa kitendakazi chako cha `range` kinarudisha safu ya data yenye urefu sahihi, lakini thamani zote ni thamani ya `start`. Ili kurekebisha hili, ungana njia ya `.map()` na njia yako ya `.fill()`.

Pitia njia ya `.map()` kitendakazi cha mwitiko kinachopokea `element` na `index` kama vigezo na kurudisha jumla ya vigezo hivyo.

# --hints--

Unapaswa kutumia njia ya `.map()`.

```js
assert.match(code, /\.map\(/);
```

Unapaswa kuunganisha njia ya `.map()` na njia yako ya `.fill()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(\s*start\s*\)\.map\(/);
```

Unapaswa kupitisha kitendakazi cha mwitiko kwa `.map()` ukitumia sintaksia ya mshale.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(\s*start\s*\)\.map\(\s*\(.*\)\s*=>/);
```

Kitendakazi chako cha mwitiko cha `.map()` kinapaswa kupokea `element` kama kigezo cha kwanza.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(\s*start\s*\)\.map\(\s*\(\s*element/);
```

Kitendakazi chako cha mwitiko cha `.map()` kinapaswa kupokea `index` kama kigezo cha pili.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(\s*start\s*\)\.map\(\s*\(\s*element\s*,\s*index\s*\)\s*=>/);
```

Kitendakazi chako cha mwitiko cha `.map()` kinapaswa kutumia kurudisha kwa njia isiyoelezwa wazi.

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(\s*start\s*\)\.map\(\s*\(\s*element\s*,\s*index\s*\)\s*=>\s*\{/);
```

Kitendakazi chako cha mwitiko cha `.map()` kinapaswa kurudisha kwa njia isiyoelezwa wazi jumla ya `element` na `index`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(\s*start\s*\)\.map\(\s*\(\s*element\s*,\s*index\s*\)\s*=>\s*(element\s*\+\s*index|index\s*\+\s*element)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
const range = (start, end) => Array(end - start + 1).fill(start);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
