---
id: 642e0011c45c893845842058
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Kijenzi cha `Array()` kina njia ya `.fill()` ambayo inaweza kutumika kujaza safu ya data na thamani. Unaweza kutumia hii kujaza safu yako ya data na thamani ya `start`.

Unganisha njia ya `.fill()` kwa kijenzi chako cha `Array()`, na upitie thamani ya `start`.

# --hints--

Unapaswa kutumia njia ya `.fill()`.

```js
assert.match(code, /\.fill\(/);
```

Unapaswa kuita njia ya `.fill()` kwenye kijenzi chako cha `Array()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(/);
```

Unapaswa kupitisha `start` kwa njia ya `.fill()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(\s*start\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
const range = (start, end) => Array(end - start + 1);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
