---
id: 642e02be7845f13b014cd2b0
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Sasa baada ya kuwa na kitendakazi cha `range`, unaweza kukitumia kuunda safu ya herufi pia.

Tangaza kitendakazi cha `charRange` ukitumia `const` na sintaksia ya mshale. Kitendakazi hicho kinapaswa kuchukua kigezo cha `start` na `end`. Kitendakazi kinapaswa kurudisha matokeo ya wito wa `range()` kwa `start` na `end` kama hoja.

# --hints--

Unapaswa kutangaza kigezo cha `charRange`.

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

Kigezo chako cha `charRange` kinapaswa kutangazwa kwa `const`.

```js
assert.match(code, /const\s+charRange/);
```

Kigezo chako cha `charRange` kinapaswa kuwa kitendakazi.

```js
assert.isFunction(charRange);
```

Kitendakazi chako cha `charRange` kinapaswa kutumia sintaksia ya mshale.

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

Kitendakazi chako cha `charRange` kinapaswa kuchukua `start` kama kigezo cha kwanza.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

Kitendakazi chako cha `charRange` kinapaswa kuchukua `end` kama kigezo cha pili.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

Kitendakazi chako cha `charRange` kinapaswa kutumia kurudisha kwa njia isiyo wazi.

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

Kitendakazi chako cha `charRange` kinapaswa kuita kitendakazi chako cha `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

Unapaswa kupitisha `start` na `end` kama hoja katika wito lako la `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
