---
id: 64345b810a6e481e5e326849
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

`range()` itarudisha safu ya nambari, ambazo unahitaji kuzibadilisha tena kuwa herufi. Unganisha njia ya `.map()` na wito wako wa `range()`.

Pitia kitendakazi cha mwitiko kinachochukua `code` kama kigezo na kinarejesha kwa njia isiyo ya moja kwa moja thamani ya kupitisha `code` kwa njia ya `String.fromCharCode()`.

# --hints--

Unapaswa kutumia njia ya `.map()`.

```js
assert.lengthOf(code.match(/\.map\(/g), 2);
```

Unapaswa kuunganisha njia ya `.map()` na wito wako wa `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)\.map\(/);
```

Unapaswa kutumia sintaksia ya mshale kwa kitendakazi cha mwitiko cha `.map()`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)\.map\(\s*(\(.*\)|[^\s()]+)\s*=>/);
```

Kitendakazi chako cha mwitiko cha `.map()` kinapaswa kuchukua kigezo cha `code`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)\.map\(\s*(\(\s*code\s*\)|code)\s*=>/);
```

Kitendakazi chako cha mwitiko cha `.map()` kinapaswa kutumia kurudisha thamani isiyo ya moja kwa moja.

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)\.map\(\s*(\(\s*code\s*\)|code)\s*=>\s*\{/);
```

Kitendakazi chako cha mwitiko cha `.map()` kinapaswa kurudisha matokeo ya wito wa `String.fromCharCode()`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)\.map\(\s*(\(\s*code\s*\)|code)\s*=>\s*String\.fromCharCode\(/);
```

Unapaswa kupitisha kigezo `code` kwa `String.fromCharCode()`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)\.map\(\s*(\(\s*code\s*\)|code)\s*=>\s*String\.fromCharCode\(\s*code\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0));
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
