---
id: 643475e13dc727231acd0f72
title: Hatua 17
challengeType: 0
dashedName: step-17
---

# --description--

Sasa katika wito wako uliopangwa ndani wa `.forEach()`, tangaza kigezo cha `input`. Tumia njia ya `.createElement()` ya kipengele cha `document` kuunda kipengele cha `input`. Weka sifa ya kipengele ya `type` kuwa `"text"` na sifa ya kipengele ya `id` kuwa `letter + number`.

# --hints--

Unapaswa kutangaza kigezo cha `input`.

```js
assert.match(code, /(?:var|let|const)\s+input/)
```

Unapaswa kutumia `const` kutangaza kigezo chako cha `input`.

```js
assert.match(code, /const\s+input/)
```

Unapaswa kuita njia ya `.createElement()` ya kipengele cha `document`.

```js
assert.lengthOf(code.match(/document\.createElement\(/g), 2)
```

Unapaswa kupitisha mfuatano wa herufi `"input"` kwa njia ya `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Unapaswa kuweka kipengele chako kipya cha `input` katika kigezo chako cha `input`.

```js
assert.match(code, /const\s+input\s*=\s*document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Unapaswa kupata upatikanaji wa sifa ya kipengele ya `type` ya kipengele chako cha `input`.

```js
assert.match(code, /input\.type/);
```

Unapaswa kuweka sifa ya kipengele ya `type` ya kipengele chako cha `input` kuwa `"text"`.

```js
assert.match(code, /input\.type\s*=\s*('|"|`)text\1/)
```

Unapaswa kupata upatikanaji wa sifa ya kipengele ya `id` ya kipengele chako cha `input`.

```js
assert.match(code, /input\.id/);
```

Unapaswa kuweka sifa ya kipengele ya `id` ya kipengele chako cha `input` kuwa `letter + number`.

```js
assert.match(code, /input\.id\s*=\s*letter\s\+\snumber/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {

    })
--fcc-editable-region--
  })
}
```
