---
id: 6434552bcc0a951a0a99df3b
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Kitendakazi chako cha `range` kinatarajia nambari, lakini thamani zako za `start` na `end` zitakuwa mfuatano wa herufi (hasa, zitakuwa herufi moja kama `A`).

Badilisha thamani zako za `start` na `end` katika wito wako wa `range()` kuwa nambari kwa kutumia njia ya `.charCodeAt()` juu yao, ukipitisha nambari `0` kama hoja kwa njia hiyo.

# --hints--

Unapaswa kutumia njia ya `.charCodeAt()`.

```js
assert.match(code, /\.charCodeAt\(/);
```

Unapaswa kuita njia ya `.charCodeAt()` juu ya `start`.

```js
assert.match(code, /start\.charCodeAt\(/);
```

Unapaswa kupitisha `0` kwa njia ya `.charCodeAt()` ya `start`.

```js
assert.match(code, /start\.charCodeAt\(\s*0\s*\)/);
```

Unapaswa kuita njia ya `.charCodeAt()` juu ya `end`.

```js
assert.match(code, /end\.charCodeAt\(/);
```

Unapaswa kupitisha `0` kwa njia ya `.charCodeAt()` ya `end`.

```js
assert.match(code, /end\.charCodeAt\(\s*0\s*\)/);
```

Unapaswa kutumia njia za `.charCodeAt()` moja kwa moja katika wito wako wa `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start, end);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
