---
id: 643715013330824ecaa70442
title: Hatua 23
challengeType: 0
dashedName: step-23
---

# --description--

Kitendakazi chako kinachofuata kitahesabu thamani ya `median` ya safu ya nambari. Anza kwa kutangaza kitendakazi cha mshale chenye kigezo `nums`.

Ndani ya kitendakazi, tangaza kigezo cha `sorted` na weka thamani yake kuwa matokeo ya kupanga nakala ya safu ya `nums`.

Unapaswa kutumia njia ya `slice()` kwa ajili ya kuunda nakala ya kina ya safu hiyo.

# --hints--

Unapaswa kutangaza kigezo cha `median`.

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

Unapaswa kutumia `const` kutangaza kigezo chako cha `median`.

```js
assert.match(code, /const\s+median/);
```

Kigezo chako cha `median` kinapaswa kuwa kitendakazi.

```js
assert.isFunction(median);
```

Kitendakazi chako cha `median` kinapaswa kutumia sintaksia ya mshale.

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

Kitendakazi chako cha `median` kinapaswa kuwa na kigezo cha `nums`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

Kitendakazi chako cha `median` hakipaswi kutumia kurudisha thamani kwa njia isiyo wazi.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

Kitendakazi chako cha `median` kinapaswa kuwa na kigezo cha `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

Unapaswa kutumia `const` kutangaza kigezo chako cha `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

Unapaswa kutumia `.slice()` kuweka nakala ya safu ya `nums` katika `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

Unapaswa kuunganisha njia ya `.sort()` na njia yako ya `.slice()`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

Unapaswa kupitisha kitendakazi cha mwitiko kwa njia yako ya `sort` ili kupanga nambari kwa usahihi kwa mpangilio wa kuongezeka. Tumia kurudisha thamani kwa njia isiyo wazi kwa uwazi.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
