---
id: 64496df724dd3716a71fe971
title: Hatua 26
challengeType: 0
dashedName: step-26
---

# --description--

Sifa za vitu zinajumuisha jozi za funguo/thamani. Unaweza kutumia majina ya sifa ya ufupisho unapotangaza kitu cha aina ya literal. Unapotumia sintaksia ya sifa ya ufupisho, jina la kigezo linakuwa funguo ya sifa na thamani yake kuwa thamani ya sifa.

Mfano ufuatao unatangaza kitu cha `user` chenye sifa za `userId`, `firstName`, na `loggedIn`.

```js
const userId = 1;
const firstName = "John";
const loggedIn = true;

const user = {
  userId,
  firstName,
  loggedIn,
};

console.log(user); // { userId: 1, firstName: 'John', loggedIn: true }
```

Ili kufuatilia vitendakazi vyote vya jedwali lako la hesabu, tangaza kitu cha `spreadsheetFunctions`. Ukitumia sintaksia ya ufupisho wa sifa, weka `sum`, `average`, na `median` kama sifa kwenye kitu cha `spreadsheetFunctions`.

# --hints--

Unapaswa kutangaza kigezo cha `spreadsheetFunctions`.

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

Unapaswa kutumia `const` kutangaza kigezo chako cha `spreadsheetFunctions`.

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

Kigezo chako cha `spreadsheetFunctions` kinapaswa kuwa kitu.

```js
assert.isObject(spreadsheetFunctions);
```

Kitu chako cha `spreadsheetFunctions` kinapaswa kuwa na sifa ya `sum`.

```js
assert.property(spreadsheetFunctions, "sum");
```

Sifa yako ya `sum` inapaswa kuwa kitendakazi chako cha `sum`.

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

Kitu chako cha `spreadsheetFunctions` kinapaswa kuwa na sifa ya `average`.

```js
assert.property(spreadsheetFunctions, "average");
```

Sifa yako ya `average` inapaswa kuwa kitendakazi chako cha `average`.

```js
assert.equal(spreadsheetFunctions?.average, average);
```

Kitu chako cha `spreadsheetFunctions` kinapaswa kuwa na sifa ya `median`.

```js
assert.property(spreadsheetFunctions, "median");
```

Sifa yako ya `median` inapaswa kuwa kitendakazi chako cha `median`.

```js
assert.equal(spreadsheetFunctions?.median, median);
```

Unapaswa kutumia sintaksia ya ufumbuzi wa muundo kuweka thamani za sifa zako.

```js
const objectText = code.replace(/.*const\s+spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
