---
id: 65ae458e23954c3469e0c209
title: Hatua 18
challengeType: 0
dashedName: step-18
---


# --description--

Katika miradi ya awali ulijifunza kuhusu njia ya `setAttribute`. Njia nyingine ya kusahihisha sifa ya kipengele katika JavaScript ni kutumia sintaksia ifuatayo:

```js
el.attribute = value;
```

Majina ya sifa kwa thamani za sifa za kipengele za HTML zenye kistari, kama `aria-label`, hutegemea camel case, zikawa `ariaLabel`.

```js
el.ariaLabel = "Aria Label Value";
```

Weka sifa ya `aria-label` kwa kipengele cha `input` kuwa na thamani ile ile kama sifa ya `id`.


# --hints--

Unapaswa kuwa na `input.ariaLabel`.

```js
assert.match(code, /input\.ariaLabel\s*=/);
```


Unapaswa kumpa `input.ariaLabel` thamani ile ile kama `input.id`.

```js
assert.match(code, /input\.ariaLabel\s*=\s*((letter\s*\+\s*number)|(input.id))/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
     createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {
     const input = document.createElement("input");
     input.type = "text";
     input.id = letter + number;
 
    })   
--fcc-editable-region--
  })
}
```
