---
id: a6b0bb188d873cb2c8729495
title: 實作一個 HTML 物體轉換器
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

本實驗室是關於將字串中的特殊字元轉換為對應的 HTML 實體。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `convertHTML` 函式（程式），它接受一個字串作為引數。
1. `convertHTML` 函式應該透過將引數字串中的特殊字元轉換為對應的 HTML 實體來傳回一個新的字串。

   - `&` 應該轉換為 `&amp;`。
   - `<` 應該轉換為 `&lt;`。
   - `>` 應該轉換為 `&gt;`。
   - `&quot;` 應該轉換為 `&quot;`。
   - `'` 應該轉換為 `&apos;`。

# --hints--

你應該有一個 `convertHTML` 函式（程式）。

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` 應該傳回字串 `Dolce &amp; Gabbana`。

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` 應該傳回字串 `Hamburgers &lt; Pizza &lt; Tacos`。

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` 應該傳回字串 `Sixty &gt; twelve`。

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` 應該傳回字串 `Stuff in &quot;quotation marks&quot;`。

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` 應該傳回字串 `Schindler&apos;s List`。

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` 應該傳回字串 `&lt;&gt;`。

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` 應該傳回字串 `abc`。

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
