---
id: 65f569725359e10d345bc52a
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Utaanzisha programu inayohesabu na kuchora njia ya mripuko kulingana na pembe, kasi na urefu wa kurusha.

Anza kwa ku-import `math`, utaitumia sana katika mradi huu kwa sababu ina njia muhimu kama `math.radians`, `math.cos`, `math.sin` na nyinginezo.

Pia tengeneza hizi variable ili ziwe na thamani ya kasi ya mvutano wa dunia na baadhi ya alama maalum zitakazosaidia baadaye (tumia nakala na weka kwa hizi).

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

Tumia `import math` ku-import moduli ya `math`.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

Unapaswa kuwa na variable ya `GRAVITATIONAL_ACCELERATION` yenye thamani ya `9.81`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

Unapaswa kuwa na variable ya `PROJECTILE` yenye thamani ya `"∙"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

Unapaswa kuwa na variable ya `x_axis_tick` yenye thamani ya `"T"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

Unapaswa kuwa na variable ya `y_axis_tick` yenye thamani ya `"⊣"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
