---
id: 65fd4b3bedee044b4b957d33
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Tengeneza darasa la `Projectile` lenye njia ya `__init__` inayoweka darasa hilo kwa kutumia hoja tatu: kasi ya mwanzo, urefu wa mwanzo, na pembe ya mwanzo ya kurusha mrengo, kwa mpangilio huu. 

Ndani ya njia ya `__init__`, weka thamani hizi za hoja tatu katika sifa za kipengele binafsi tatu: `__speed`, `__height`, na `__angle`. Pembe iliyotolewa itakuwa kwa nyuzi; hata hivyo, inapaswa kuhifadhiwa ndani kwa radiani. Ili kufanikisha hili, tumia kitendakazi cha `math.radians()` kubadilisha pembe kutoka nyuzi hadi radiani wakati unaiweka katika `__angle`. 

Matumizi ya alama mbili za mstari chini kabla ya jina la sifa ya kipengele husababisha mabadiliko ya jina (name mangling) katika Python. Hii ina maana sifa hizo hazipatikani moja kwa moja kutoka nje ya darasa kwa kutumia majina yaliyotolewa, na lazima zipatikane kwa majina yaliyobadilishwa kama `ball._Projectile__height` ikiwa zinahitajika kutoka nje, ikionyesha kuwa hizi ni kwa matumizi ya ndani tu.

# --hints--

Unapaswa kutumia `class Projectile:` kutengeneza darasa.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

Unapaswa kuwa na njia ya `__init__` ndani ya darasa la `Projectile`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

Kipengele kipya cha darasa la `Projectile` kinapaswa kuwa na thamani za sifa tatu zilizowekwa vizuri.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
