---
id: 65fd4bd84561a14d3e720061
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Kigezo cha darasa `__slots__` kina matumizi maalum katika madarasa ya Python. Kutangaza `__slots__` na kuweka thamani yake kuwa mfululizo wa mfuatano wa herufi kunazuia uundaji wa sifa za kipengele kuwa zile tu zilizomo katika mfululizo huo. Pia, kunazuia uundaji wa sifa ya kipengele maalum ya `__dict__` na kunaruhusu upatikanaji wa sifa za kipengele kwa ufanisi zaidi.

Unapaswa kutumia kigezo cha `__slots__` ndani ya darasa ili kufafanua ni sifa gani za kipengele darasa lina: weka thamani kwa `__slots__` tuple inayojumuisha mfuatano wa herufi 3, kila moja sawa na moja ya majina ya sifa za kipengele zilizotangazwa katika `__init__`.

# --hints--

Unapaswa kuweka thamani kwa `__slots__` tuple inayojumuisha `'__height'`, `'__speed'`, na `'__angle'`.

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
