---
id: 660400f02cf3a76795b82d4e
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Kitu cha kwanza cha kuanzisha ni njia inayohesabu mwelekeo wa mradi, ambao ni nafasi ya usawa iliyosafiriwa kutoka kwa kutupa hadi mradi kugusa ardhi.

Tengeneza njia `__calculate_displacement`, ambayo ina `self` pekee kama kigezo, na rudisha mwelekeo wa mradi.

Tumia fomula ifuatayo kuhesabu mwelekeo wa mradi:
\\[ d = \frac{v \cdot \cos\theta \cdot \left(v\cdot\sin\theta + \sqrt{v^2 \cdot \sin^2\theta + 2 \cdot g \cdot h}\right)}{g} \\]

Ambapo $d$ ni mwelekeo, $v$ ni kasi ya kuanzia, $\theta$ ni pembe na $h$ ni urefu wa kuanzia wa mradi.
Kwa $g$ unaweza kutumia tofauti ya `GRAVITATIONAL_ACCELERATION`.

Unapaswa kutumia njia `math.cos()` na `math.sin()` kwa kazi za trigonometric na `math.sqrt()` kuhesabu mzizi wa mraba. Pia unapaswa kujua kuwa $x^y$ inaandikwa kama `x ** y` katika python. Pia $\sin^2\theta$ inamaanisha kuwa thamani inayotokana na sine kisha inazidishwa yenyewe mara mbili.

Kumbuka kuwa kwa jina la mangling unahitaji kuita njia kama `_Projectile__calculate_displacement` ikiwa unataka kujaribu, au kuitumia kutoka nje ya darasa:

```py
ball = Projectile(10, 3, 45)
displacement_of_ball = ball._Projectile__calculate_displacement() # 12.6173996009878
```

# --hints--

Unapaswa kutangaza njia inayoitwa `__calculate_displacement` na `def __calculate_displacement(self):`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").has_function("__calculate_displacement")`))})
```

Njia ya `__calculate_displacement` inapaswa kuwa na hoja `self` pekee.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_args('self')`))})
```

Njia ya `__calculate_displacement` inapaswa kurudisha thamani sahihi.

```js
({test: () => assert(runPython(`
p = Projectile(20, 21, 22)
disp = p._Projectile__calculate_displacement()
round(disp, 2) == 55.06 and round(disp, 2) != disp
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    
--fcc-editable-region--
```
