---
id: 660fcf3e1b9bb056b2edb567
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Kwa sasa uko tayari kuunda uwakilishi wa mfuatano wa herufi.
Anza kwa kuunda mfano wa darasa la `Projectile`. Chini ya ufafanuzi wa darasa, unda kigezo cha `ball` na weka thamani yake kwa wito wa `Projectile` ukitumia `10, 3, 45` kama hoja.

Fafanua njia ya `__str__` ndani ya darasa, ili kumpa mfano wako uwakilishi sahihi wa mfuatano wa herufi, na ifanye irudishe mfuatano wa herufi wenye muundo ufuatao:

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

Inapaswa kuanza na kumalizika na tabia ya mstari mpya, pembe iandikwe kama nambari kamili katika digrii na uhamisho uchapishwe kwa tarakimu moja ya desimali.

Utapata `math.degrees` kuwa muhimu kubadilisha pembe kutoka radians hadi digrii.

Unapokuwa tayari unaweza `print(ball)` kujaribu kitendakazi chako.

# --hints--

Unapaswa kuwa na kigezo cha `ball`.

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

Kigezo cha `ball` kinapaswa kuwa na thamani ya `Projectile(10, 3, 45)`.

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

Uwakilishi wa mfuatano wa herufi kwa `Projectile(45, 45, 45)` unapaswa kuwa sahihi.

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

Uwakilishi wa mfuatano wa herufi pia unapaswa kuwa sahihi kwa mifano mingine.

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

Unapaswa kuwa na wito wa `print(ball)`.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
