---
id: 6616ec0bf9ee8548ce6b0f08
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Sasa ni wakati wa kufanya kazi ya kuhesabu kuratibu za njia ya mwelekeo! Tengeneza njia ya darasa la `Projectile` iitwayo `__calculate_y_coordinate`, inapaswa kuwa na, mbali na `self`, kigezo cha `x`.

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

Hii ni fomula ya kuhesabu nafasi ya wima $y$ kwa nafasi yoyote ya usawa $x$, ukiwa na pembe ya kuanzia $\theta$, kasi $v_0$ na urefu $y_0$.

Utahitaji kutumia `math.tan()` na `math.cos()` na kumbuka kuwa `x ** y` ni njia ya kuandika $x^y$, na thamani ya $g$ iko katika kigezo `GRAVITATIONAL_ACCELERATION`.

Tekeleza njia hiyo ili irudishe kuratibu ya $y$.

Kama unataka kujaribu njia hiyo, kutoka nje ya darasa, unaweza kuita `ball._Projectile__calculate_y_coordinate()` na nambari kama hoja baada ya `ball` kutangazwa.

# --hints--

Darasa lako la `Projectile` linapaswa kuwa na njia iitwayo `__calculate_y_coordinate`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

Njia ya `__calculate_y_coordinate` inapaswa kuwa na `self, x` kama hoja.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

Njia inapaswa kurudisha thamani sahihi. Kwa `Projectile(12, 12, 12)`, ikitolewa kuratibu 4 za $x$ za 0, 1, 2 na 3, kuratibu ya y inapaswa kuwa takriban 12, 12.18, 12.28, 12.32.

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

Njia inapaswa kurudisha thamani sahihi. Kwa `Projectile(45, 12, 22)`, ikitolewa kuratibu za $x$ za 6, 9, 10 na 12 kuratibu ya y inapaswa kuwa takriban 14.32, 15.41, 15.76, 16.44.

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

Njia ya `__calculate_y_coordinate` haipaswi kuzungusha matokeo.

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
