---
id: 662fa453ec8033322b3ab2e3
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Ni wakati wa kuzungumzia kufunga (encapsulation) na wapokeaji (getters). Umeandika sifa tatu za mfano (instance) kuwa binafsi kwa kutumia alama mbili za mstari chini mwanzoni. Kumbuka kwamba sifa hizi huitwa binafsi kwa mkataba: ingawa bado zinaweza kufikiliwa kutoka nje, imekubalika kutofanya hivyo.

Wapokeaji ni kile kinachoweza kutumika kupata thamani kutoka nje. Ili kufafanua mpokeaji, unafafanua njia inayorejesha thamani ya sifa inayotakiwa na kuipa dekoreta ya `@property`:

```py
class Nest:
    ...
    @property
    def number_of_eggs(self):
        return self.__number_of_eggs
```

 Dekoreta hubadilisha njia kuwa sifa, ikimaanisha kwamba njia haipigiwi kama njia ya kawaida, bali inatumiwa kama sifa:

```py
n = Nest()
print(n.number_of_eggs)
```

Katika mfano hapo juu, sifa binafsi `__number of eggs` inafikiliwa kupitia sifa ya `number_of_eggs` ya `n`.

Tengeneza mpokeaji aliyeitwa `speed` na uifanye irudishe thamani ya sifa binafsi `__speed`.

# --hints--

Unapaswa kufafanua njia ya `speed`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('speed')
`))})
```

Njia yako ya `speed` inapaswa kuwa na kigezo kimoja, `self`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_args('self')
`))})
```

Njia yako ya `speed` inapaswa kuwa na dekoreta ya `@property`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_decorators('property')
`))})
```

Njia yako ya `speed` inapaswa kurudisha thamani ya sifa binafsi `__speed`.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed == 22
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
