---
id: 696920c0c98a1ed58eb86293
title: Crear una vitrina de productos
challengeType: 25
dashedName: lab-product-showcase
demoType: onClick
---

# --description--

Para este laboratorio, se te ha proporcionado todo el HTML y CSS. Completarás una vitrina de productos usando TypeScript.

**Objetivo:** Cumple las historias de usuario a continuación y haz que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una interfaz llamada `Item` con:
   - Una propiedad `type` que solo puede ser `"book"`, `"electronics"` o `"clothing"`.
   - Una propiedad `id` de tipo `string`.
   - Una propiedad `price` de tipo `number`.

1. Debes crear una interfaz llamada `Book` que extienda `Item`. Su propiedad `type` debe estar establecida con el literal `"book"`. También debe tener propiedades `title` y `author` de tipo `string`.

1. Debes crear una interfaz llamada `Electronics` que extienda `Item`. Su propiedad `type` debe estar establecida con el literal `"electronics"`. También debe tener propiedades `item` y `model` de tipo `string`, y una propiedad opcional `warranty` de tipo `number`.

1. Debes crear una interfaz llamada `Clothing` que extienda `Item`. Su propiedad `type` debe estar establecida con el literal `"clothing"`. También debe tener propiedades `item` y `brand` de tipo `string`, y una propiedad opcional `size` que solo puede ser `"S"`, `"M"` o `"L"`.

1. Debes crear un tipo unión llamado `Product` que combine las interfaces `Book`, `Electronics` y `Clothing`.

1. Debes crear una clase genérica llamada `Collection` con un parámetro de tipo `T`. El constructor debe aceptar un arreglo de elementos de tipo `T` y almacenarlos en una propiedad llamada `items`.

1. La clase `Collection` debe tener un método `getAll()` que devuelva todos los elementos de la colección.

1. La clase `Collection` debe tener un método `filter()` que acepte un callback y devuelva solo los elementos del arreglo `items` que cumplan la condición dada. El callback debe recibir un elemento de tipo `T` y devolver un booleano.

1. Debes crear una función llamada `renderProduct` con un parámetro de tipo `Product`. Debe devolver una cadena HTML que represente un producto con:
   - Un elemento con la clase `item` y un ID con el valor del `id` del producto.
   - Un elemento con la clase `price` y el texto del `price` del producto.
   - Un elemento que contenga información adicional del producto.

1. La función `renderProduct` debe usar estrechamiento de tipos para renderizar diferente HTML según el tipo de producto:

   - Para libros, el HTML renderizado debe tener el formato `Book: [title] by [author]`, donde `[title]` y `[author]` deben reemplazarse por el título y el autor, respectivamente.
   - Para electrónicos, el HTML renderizado debe tener el formato `Electronics: [item] - [model]`. Si el producto tiene garantía, añade ` - Warranty: [warranty] year(s)` al formato. Reemplaza `[item]`, `[model]` y `[warranty]` con el artículo, el modelo y la garantía, respectivamente.
   - Para ropa, el HTML renderizado debe tener el formato `Clothing: [item] by [brand]`. Si el producto tiene talla, añade ` - Size [size]` al formato. Reemplaza `[item]`, `[brand]` y `[size]` con el artículo, la marca y la talla, respectivamente.
   - Si se encuentra un tipo de producto inválido, lanza un error con el formato: `Unknown product type: [product]`, donde `[product]` debe reemplazarse por el objeto de producto inválido. Usa `JSON.stringify` para convertir el producto inválido en cadena.

1. Debes crear una instancia `Collection<Product>` que incluya al menos un elemento para cada tipo de producto y asignar la instancia a una variable llamada `products`.

1. Debes crear una función `showProducts` que muestre productos en el elemento `#output`. Debe aceptar un filtro opcional que represente un tipo de producto; cuando se omita, mostrar todos los productos; cuando se proporcione, mostrar solo productos de ese tipo.

1. Debes renderizar productos convirtiendo cada producto en una cadena HTML usando la función `renderProduct` e inyectando el resultado combinado en el elemento `#output`.

1. Al hacer clic en los botones `#all`, `#books`, `#electronics` y `#clothing` se debe mostrar el conjunto correspondiente de productos en el elemento `#output`.

1. Al cargar la página, todos los productos deben mostrarse por defecto.

# --hints--

Debes tener una interfaz llamada `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Item);
```

Tu interfaz `Item` debe tener una propiedad `type` que solo pueda ser `"book"`, `"electronics"` o `"clothing"`.

```js
const explorer = await __helpers.Explorer(code);
const { Item } = explorer.interfaces;
const typeProp = Item.typeProps.type;
assert.isTrue(typeProp.isUnionOf(["'book'", "'electronics'", "'clothing'"]));
```

Tu interfaz `Item` debe tener una propiedad `id` de tipo `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "id", type: "string"}));
```

Tu interfaz `Item` debe tener una propiedad `price` de tipo `number`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "price", type: "number"}));
```

Debes tener una interfaz llamada `Book`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Book);
```

Tu interfaz `Book` debe extender `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.doesExtend("Item"));
```

Tu interfaz `Book` debe tener una propiedad `type` establecida en `"book"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "type", type: "'book'"}));
```

Tu interfaz `Book` debe tener una propiedad `title` de tipo `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "title", type: "string"}));
```

Tu interfaz `Book` debe tener una propiedad `author` de tipo `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "author", type: "string"}));
```

Debes tener una interfaz llamada `Electronics`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Electronics);
```

Tu interfaz `Electronics` debe extender `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.doesExtend("Item"));
```

Tu interfaz `Electronics` debe tener una propiedad `type` establecida en `"electronics"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "type", type: "'electronics'"}));
```

Tu interfaz `Electronics` debe tener una propiedad `item` de tipo `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "item", type: "string"}));
```

Tu interfaz `Electronics` debe tener una propiedad `model` de tipo `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "model", type: "string"}));
```

Tu interfaz `Electronics` debe tener una propiedad opcional `warranty` de tipo `number`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "warranty", type: "number", isOptional: true}));
```

Debes tener una interfaz llamada `Clothing`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Clothing);
```

Tu interfaz `Clothing` debe extender `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.doesExtend("Item"));
```

Tu interfaz `Clothing` debe tener una propiedad `type` establecida en `"clothing"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "type", type: "'clothing'"}));
```

Tu interfaz `Clothing` debe tener una propiedad `item` de tipo `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "item", type: "string"}));
```

Tu interfaz `Clothing` debe tener una propiedad `brand` de tipo `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "brand", type: "string"}));
```

Tu interfaz `Clothing` debe tener una propiedad opcional `size` que solo pueda ser `"S"`, `"M"` o `"L"`.

```js
const explorer = await __helpers.Explorer(code);
const { Clothing } = explorer.interfaces;
assert.isTrue(Clothing.hasTypeProps({name: "size", isOptional: true}));
assert.isTrue(Clothing.typeProps.size.isUnionOf(["'S'", "'M'", "'L'"]));
```

Debes crear un tipo unión llamado `Product` que combine las interfaces `Book`, `Electronics` y `Clothing`.

```js
const explorer = await __helpers.Explorer(code);
const { Product } = explorer.types;
assert.isTrue(Product.isUnionOf(["Book", "Electronics", "Clothing"]));
```

Debes tener una clase `Collection` con un parámetro de tipo `T`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.isTrue(Collection.typeParameters[0].matches("T"));
```

El constructor de tu clase `Collection` debe aceptar un arreglo de elementos de tipo `T` y almacenarlos en una propiedad llamada `items`.

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.items, arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.items, arr);
```

Tu clase `Collection` debe tener un método `getAll`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
console.log(Collection.methods)
console.log(Collection.methods.getAll)
assert.exists(Collection.methods.getAll);
```

Tu método `getAll` debe devolver todos los elementos de la colección.

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.getAll(), arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.getAll(), arr);
```

Tu clase `Collection` debe tener un método `filter`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.exists(Collection.methods.filter);
```

Tu método `filter` debe aceptar una función callback y devolver un arreglo que contenga los elementos de la propiedad `items` que cumplan la condición dada.

```js
let arr = [1, 2, 3, 4, 5, 6];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.filter(n => n % 2 === 0), [2, 4, 6]);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.filter(s => s === "c"), ["c"]);
```

Debes tener una función llamada `renderProduct`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.renderProduct);
```

`renderProduct(product)` debe devolver una cadena que contenga un elemento con la clase `item` y el ID igual a `product.id`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  assert.isNotNull(testEl.querySelector(`[class="item"][id="${p.id}"]`));
})
```

`renderProduct(product)` debe devolver una cadena que contenga un elemento con la clase `price` y el texto igual a `product.price`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const price = testEl.querySelector(`[class="price"]`);
  assert.include(price.innerText, p.price.toString());
})
```

Cuando `renderProduct` se llame con un producto `Book` `p`, debe devolver una cadena que contenga un elemento con el texto siguiendo el formato `Book: [title] by [author]`. Reemplaza `[title]` y `[author]` con las propiedades correspondientes de `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const expectedText = `Book: ${p.title} by ${p.author}`;
  assert.include(testEl.innerText, expectedText);
})
```

Cuando `renderProduct` se llame con un producto `Electronics` `p`, debe devolver una cadena que contenga un elemento con el texto siguiendo el formato `Electronics: [item] - [model]` o `Electronics: [item] - [model] - Warranty: [warranty] year(s)` si `p.warranty` no es indefinido. Reemplaza `[item]`, `[model]` y `[warranty]` con las propiedades correspondientes de `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Electronics: ${p.item} - ${p.model}`;
  const expectedText = !p.warranty ? baseFormat : baseFormat + ` - Warranty: ${p.warranty} year(s)`;
  assert.include(testEl.innerText, expectedText);
  if (!p.warranty) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

Cuando `renderProduct` se llame con un producto `Clothing` `p`, debe devolver una cadena que contenga un elemento con el texto siguiendo el formato `Clothing: [item] by [brand]` o `Clothing: [item] by [brand] - Size [size]` si `p.size` no es indefinido. Reemplaza `[item]`, `[brand]` y `[size]` con las propiedades correspondientes de `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Clothing: ${p.item} by ${p.brand}`;
  const expectedText = !p.size ? baseFormat : baseFormat + ` - Size ${p.size}`;
  assert.include(testEl.innerText, expectedText);
  if (!p.size) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

Cuando `renderProduct` se llame con un producto inválido, debe lanzar un error con el formato `Unknown product type: [p]`. Reemplaza `[p]` con el producto inválido convertido a cadena con `JSON.stringify`.

```js
const invalidProduct = { item: "T-Shirt", price: 19.99 };
const expected = `Unknown product type: ${JSON.stringify(invalidProduct)}`;
try {
  renderProduct(invalidProduct);
  assert.fail();
} catch({name, message}) {
  if (name == "AssertionError") {
    assert.fail("renderProduct should throw with invalid product");
  } else {
    assert.equal(message, expected);
  }
}
```

Debes crear una instancia `Collection<Product>` y asignarla a una variable llamada `products`.

```js
assert.instanceOf(products, Collection);
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.products.value.typeArguments[0].matches("Product"));
```

Tu variable `products` debe contener al menos un elemento para cada tipo de producto.

```js
assert.isTrue(products.items.some(p => p.type === "book"));
assert.isTrue(products.items.some(p => p.type === "electronics"));
assert.isTrue(products.items.some(p => p.type === "clothing"));
```

Debes tener una función llamada `showProducts`.

```js
assert.isFunction(showProducts);
```

Al hacer clic en el botón `#all` se deben mostrar todos los productos en el elemento `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
books.dispatchEvent(new Event("click"));
all.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

Al hacer clic en el botón `#books` se debe mostrar el conjunto correspondiente de productos en el elemento `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "book") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click"));
books.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Al hacer clic en el botón `#electronics` se debe mostrar el conjunto correspondiente de productos en el elemento `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "electronics") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click"));
electronics.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Al hacer clic en el botón `#clothing` se debe mostrar el conjunto correspondiente de productos en el elemento `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "clothing") {
    count++;
  }
});
books.dispatchEvent(new Event("click"));
clothing.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Todos los productos deben ser visibles por defecto cuando la página cargue. Usa el evento `DOMContentLoaded` para asegurarlo.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
document.dispatchEvent(new Event("DOMContentLoaded"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts
interface Item {
  type: "book" | "electronics" | "clothing";
  id: string;
  price: number;
}

interface Book extends Item {
  type: "book";
  title: string;
  author: string;
}

interface Electronics extends Item {
  type: "electronics";
  item: string;
  model: string;
  warranty?: number;
}

interface Clothing extends Item {
  type: "clothing";
  item: string;
  brand: string;
  size?: "S" | "M" | "L";
}

type Product = Book | Electronics | Clothing;

class Collection<T> {
  items: T[];
  constructor(items: T[]) {
    this.items = items;
  }

  getAll(): T[] {
    return this.items;
  }

  filter(callback: (item: T) => boolean): T[] {
    return this.items.filter(callback);
  }
}

function renderProduct(p: Product): string {
  const createProductCard = (id: string, content: string, price: number): string => {
    return `
     <div class="item" id="${id}">
       ${content}
       <div class="price">$${price}</div>
     </div>
   `;
  };

  if (p.type === "book") {
    const content = `<strong>Book:</strong> ${p.title} by ${p.author}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "electronics") {
    const warranty = p.warranty ? ` - Warranty: ${p.warranty} year(s)` : "";
    const content = `<strong>Electronics:</strong> ${p.item} - ${p.model}${warranty}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "clothing") {
    const size = p.size ? ` - Size ${p.size}` : "";
    const content = `<strong>Clothing:</strong> ${p.item} by ${p.brand}${size}`;
    return createProductCard(p.id, content, p.price);
  }

  const _never: never = p;
  throw new Error(`Unknown product type: ${JSON.stringify(p)}`);
}


const products = new Collection<Product>([
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },

  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "b2", type: "book", title: "1984", author: "George Orwell", price: 9.99 },

  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 },

  { id: "b3", type: "book", title: "Brave New World", author: "Aldous Huxley", price: 11.99 },
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "e5", type: "electronics", item: "Smartwatch", model: "Chronex Pulse-Q", warranty: 1, price: 199.99 },

  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 },
]);

const output = document.getElementById("output");

// Track currently active filter category (or undefined for all)
let currentFilter: Product["type"] | undefined = undefined;

function showProducts(filter?: Product["type"]): void {
  if (!output) return;

  let itemsToShow: Product[];

  // Toggle logic
  if (currentFilter === filter || filter === undefined) {
    // Reset to all if same button clicked again or filter is undefined
    itemsToShow = products.getAll();
    currentFilter = undefined;
    setActiveButton(undefined);
  } else {
    itemsToShow = products.filter((p: Product): boolean => p.type === filter);
    currentFilter = filter;
    setActiveButton(filter);
  }

  output.innerHTML = itemsToShow.map(renderProduct).join("");
}

// Active Button Highlighting
function setActiveButton(filter: Product["type"] | undefined): void {
  const buttons: { [key: string]: HTMLElement | null } = {
    all: document.getElementById("all"),
    book: document.getElementById("books"),
    electronics: document.getElementById("electronics"),
    clothing: document.getElementById("clothing")
  };

  for (const key in buttons) {
    buttons[key]?.classList.remove("active");
  }

  if (filter === undefined) {
    buttons["all"]?.classList.add("active");
  } else {
    buttons[filter]?.classList.add("active");
  }
}

// Button Event Handlers
document.getElementById("all")?.addEventListener("click", (): void => showProducts());
document.getElementById("books")?.addEventListener("click", (): void => showProducts("book"));
document.getElementById("electronics")?.addEventListener("click", (): void => showProducts("electronics"));
document.getElementById("clothing")?.addEventListener("click", (): void => showProducts("clothing"));

// Initial Render
document.addEventListener("DOMContentLoaded", (): void => showProducts());
```
