---
id: 614ccc21ea91ef1736b9b578
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the first part of the Accessibility Quiz. As you are becoming a seasoned HTML and CSS developer, we have started you off with the basic boilerplate.

通过为你的 `html` 元素添加一个 `lang` 属性来开始这段无障碍之旅。 这将帮助屏幕阅读器识别页面的语言。

# --hints--

你应该给 `html` 元素一个 `lang` 属性。 _提示：对于英文，你可以使用值 `en` 。_

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
