---
id: 68a23004e96a105df3a21fc4
title: Jenga msaidizi wa kuhesabu kadi
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

Katika mchezo wa kasino Blackjack, mchezaji anaweza kubaini kama ana faida kwenye mkono unaofuata dhidi ya nyumba kwa kufuatilia idadi ya kadi za juu na za chini zilizobaki kwenye karata. Hii huitwa Kuhesabu Kadi.

Kuwa na kadi nyingi za juu zilizobaki kwenye karata kunamfaa mchezaji. Wakati hesabu ni chanya, mchezaji anapaswa kuweka dau kubwa. Wakati hesabu ni sifuri au hasi, mchezaji anapaswa kuweka dau ndogo.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kutumia `let` kutangaza kigezo cha kimataifa kinachoitwa `count` na kukiweka kuwa `0`.
1. Unapaswa kuwa na kitendakazi kinachoitwa `cardCounter`.
1. Kitendakazi cha `cardCounter` kinapaswa kupokea kigezo `card` ambacho kinaweza kuwa nambari au mfuatano wa herufi.
    - Kwa thamani kati ya `2` hadi `10`, kigezo cha `card` kitakuwa nambari.
    - Kwa thamani nyingine zote, kigezo cha `card` kitakuwa mfuatano wa herufi.
1. Kitendakazi cha `cardCounter` kinapaswa kubadilisha kigezo cha kimataifa `count` kulingana na vigezo fulani.
1. Kigezo cha kimataifa `count` kinapaswa kuongezwa kwa `1` kwa kadi `2`, `3`, `4`, `5`, au `6`.
1. Kigezo cha kimataifa `count` kinapaswa kubaki kisibadilike kwa kadi `7`, `8`, `9`.
1. Kigezo cha kimataifa `count` kinapaswa kupunguzwa kwa `1` kwa kadi `10`, `"J"`, `"Q"`, `"K"`, `"A"`.
1. Kitendakazi cha `cardCounter` kinapaswa kurudisha mfuatano wa herufi wenye hesabu ya sasa na mfuatano wa herufi `Bet` ikiwa hesabu ni chanya.
1. Kitendakazi cha `cardCounter` kinapaswa kurudisha mfuatano wa herufi wenye hesabu ya sasa na mfuatano wa herufi `Hold` ikiwa hesabu ni ndogo au sawa na `0`.
1. Katika matokeo ya kitendakazi, hesabu ya sasa na uamuzi wa mchezaji (`Bet` au `Hold`) vinapaswa kutenganishwa na nafasi moja. Kwa mfano, `-3 Hold`.

# --hints--

Unapaswa kutumia `let` kutangaza kigezo cha kimataifa kinachoitwa `count` na kukiweka kuwa `0`.

```js
const explorer = await __helpers.Explorer(code);
const { count } = explorer.variables;
assert.isTrue(count?.matches('let count = 0'));
```

Unapaswa kuwa na kitendakazi kinachoitwa `cardCounter`.

```js
assert.isFunction(cardCounter);
```

Kitendakazi chako kinapaswa kurudisha thamani ya `count` na maandishi (`Bet` au `Hold`) na nafasi moja kati yao.

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

Baada ya kadi `2`, `3`, `4`, `5`, basi kuitwa `cardCounter(6)` kunapaswa kurudisha mfuatano wa herufi `5 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

Baada ya kadi `7`, `8`, basi kuitwa `cardCounter(9)` kunapaswa kurudisha mfuatano wa herufi `0 Hold`.

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

Baada ya kadi `10`, `"J"`, `"Q"`, `"K"`, basi kuitwa `cardCounter("A")` kunapaswa kurudisha mfuatano wa herufi `-5 Hold`.

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

Baada ya kadi `3`, `7`, `"Q"`, `8`, basi kuitwa `cardCounter("A")` kunapaswa kurudisha mfuatano wa herufi `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

Baada ya kadi `2`, `"J"`, `9`, `2`, basi kuitwa `cardCounter(7)` kunapaswa kurudisha mfuatano wa herufi `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

Baada ya kadi `2`, `2`, basi kuitwa `cardCounter(10)` kunapaswa kurudisha mfuatano wa herufi `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

Baada ya kadi `3`, `2`, `"A"`, `10`, basi kuitwa `cardCounter("K")` kunapaswa kurudisha mfuatano wa herufi `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
