---
id: 66a37f37ef5823a313de8c26
title: 달 궤도 만들기
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

이 실습에서는 HTML과 CSS를 사용해 달이 지구 주위를 도는 간단한 애니메이션을 만듭니다. 지구는 시스템의 중심에 위치하고, 달은 그 주위를 공전합니다.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 뷰포트 높이의 `100%`로 높이를 설정하고 `body` 요소에 플렉스박스 레이아웃을 설정하여 모든 요소를 페이지 중앙에 정렬해야 합니다.
   
2. `div` 클래스를 가진 `space`가 있어야 합니다.

3. `.space` 요소 안에는 순서대로 `div`와 `earth` 클래스를 가진 두 개의 `orbit` 요소가 있어야 합니다.

4. `.orbit` 요소 안에는 `div` 클래스를 가진 또 다른 `moon`가 있어야 합니다.
   
5. `div` 요소는 `class`의 `space`를 가지고 페이지 중앙에 위치하며 너비와 높이가 `200px`여야 합니다.

6. `div` 요소는 `class`의 `space`를 가지고 `relative` 위치 지정 방식을 사용해야 합니다.

7. `.earth` 요소는 `absolute` 위치 지정 방식을 사용해야 합니다. `.earth` 요소의 중심을 부모 요소의 수직(top)과 수평(left) 축의 중간 지점에 위치시키세요. 그 후 `.earth` 요소를 자신의 너비와 높이의 절반만큼 뒤로 이동시켜 부모인 `.space` 요소 내에서 중앙에 위치하도록 하세요.

8. `.earth` 요소는 너비와 높이가 `100px`여야 합니다.

9. `.orbit` 요소는 너비와 높이가 `200px`여야 합니다.

10. `.orbit` 요소는 `absolute` 위치 지정 방식을 사용해야 합니다. 그 오른쪽 아래 모서리가 `.space` 요소의 중앙에 위치하도록 `transform` 속성을 사용해 수직과 수평 축 모두 `-50%`만큼 이동시키세요.

11. 달이 지구 주위를 도는 궤도 경로는 원형이어야 합니다.

12. `.moon` 요소는 `absolute` 위치 지정 방식을 사용하고 너비와 높이가 `30px`여야 합니다. `.moon` 요소는 `.orbit` 요소의 상단에 위치하고 수평으로 중앙에 있어야 합니다.

13. `.moon` 요소의 수평 위치를 추가로 조정하여 너비의 절반만큼 왼쪽으로 이동시키세요.

14. `.earth` 선택자는 배경색과 `border-radius`가 `50%`이어야 합니다.

15. `.moon` 선택자는 배경색과 `border-radius`가 `50%`이어야 합니다.

16. `@keyframes orbit` 애니메이션을 정의하여 `.orbit` 요소를 중심을 기준으로 360도 회전시키세요. 이 애니메이션을 `.orbit` 요소에 적용하며, 지속 시간은 `5`초, 타이밍 함수는 선형, 반복은 무한이어야 합니다.

**참고:** HTML에 스타일시트를 연결하고 CSS를 적용하는 것을 잊지 마세요.

# --hints--

`body` 요소의 내용은 뷰포트 높이의 `100%`로 높이를 설정하고 `display`를 `flex`, `justify-content`, `align-items`를 `center`로 설정하여 페이지 중앙에 정렬해야 합니다.

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

`div` 클래스를 가진 메인 `space`가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

`.space` 요소는 너비와 높이가 `200px`여야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

`.space` 요소는 `position` 속성이 `relative`여야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

`div`의 `earth`를 가진 `div` 클래스를 가진 `class`가 `space` 안에 있어야 합니다.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

`.earth` 요소의 첫 번째 자식은 `.space` 요소여야 합니다.

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

`.earth` 요소는 너비와 높이가 `100px`여야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

`.earth` 요소는 `absolute` 위치 지정 방식을 사용해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

`.earth` 요소는 `top`와 `left` 위치가 `50%`이어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

`.earth` 요소는 `50%` 속성을 `transform`로 설정하여 `translate(-50%, -50%)`만큼 뒤로 그리고 위로 이동해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

`div` 클래스를 가진 `orbit`가 `.space` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

`.orbit` 요소는 `.earth` 요소 다음에 와야 합니다.

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

`.orbit` 요소는 너비와 높이가 `200px`여야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

`.orbit` 요소는 `absolute` 위치 지정 방식을 사용해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

`.orbit` 요소는 `.space` 속성을 사용해 수직과 수평 축 모두 `transform`만큼 이동하여 `-50%` 요소의 중앙에 위치해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

`div` 클래스를 가진 `moon`가 `.orbit` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

`.moon` 요소는 너비와 높이가 `30px`여야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

`.moon` 요소는 `absolute` 위치 지정 방식을 사용해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

`.moon` 요소의 상단 가장자리는 포함 블록의 상단 가장자리와 정렬되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

`.moon` 요소의 왼쪽 가장자리를 포함 블록의 중앙(`50%` 왼쪽에서)으로 이동시켜야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

`.moon` 요소는 x축을 따라 `-50%`만큼 변환하여 수평으로 조정해야 합니다.

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

`.earth` 선택자는 배경색을 가져야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

`.earth` 선택자는 `border-radius`가 `50%`이어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

`.moon` 선택자는 배경색을 가져야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

`.moon` 선택자는 `border-radius`가 `50%`이어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

`@keyframes` 규칙이 있어야 합니다.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

새로운 `@keyframes` 규칙의 이름은 `orbit`이어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

`@keyframes orbit` 규칙은 `0%` 선택자를 가져야 합니다.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

`@keyframes orbit` 규칙은 `100%` 선택자를 가져야 합니다.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

`0%` 선택자는 `transform` 속성이 `rotate(0deg) translate(-50%, -50%)`로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

`100%` 선택자는 `0%` 선택자 다음에 와야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

`100%` 선택자는 `transform` 속성이 `rotate(360deg) translate(-50%, -50%)`로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

`.orbit` 선택자는 `animation` 속성이 `orbit 5s linear infinite`로 설정되어야 합니다.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```
