---
id: 69446f18fdcb11aa516f7c5c
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda kitendakazi cha `countup` kinachounda safu ya data yenye nambari kuanzia 1 hadi nambari iliyotolewa kwa kutumia kurudia.

Anza kwa kuunda kitendakazi tupu kinachoitwa `countup` kinachopokea kigezo kimoja.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `countup`.

```js
assert.isFunction(countup);
```

Kitendakazi chako cha `countup` kinapaswa kupokea kigezo kimoja.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
