---
id: 69447054012a22f38c944ff0
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Ongeza tamko la `if` ndani ya kitendakazi chako cha `countup`.

Kagua kama `number` ni ndogo kuliko `1` na rudisha safu ya data tupu.

Hii itatumika kama kesi ya msingi kwa mzunguko wako wa kurudia.

# --hints--

Wakati `number` ni ndogo kuliko `1` kitendakazi kinapaswa kurudisha safu ya data tupu.

```js
assert.deepStrictEqual(countup(0), []);
```

Wakati `number` si ndogo kuliko `1` kitendakazi hakipaswi kurudisha chochote.

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
