---
id: 694470b9a65b278f981ec62e
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Mwishowe, jaribu kitendakazi chako cha `countup` kwa kukiwita na hoja `5` ndani ya tamko la `console.log()`.

Matokeo yanapaswa kuwa `[1, 2, 3, 4, 5]`.

# --hints--

Unapaswa kuitisha kitendakazi `countup` na hoja `5` ndani ya tamko la `console.log()`

```js
assert.deepStrictEqual(countup(5), [1, 2, 3, 4, 5]);
assert.match(code, /console\.log\s*\(\s*countup\s*\(\s*5\s*\)\s*\)/)
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
console.log(countup(5));
```
