---
id: 6948705160b66cf93a916e29
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Tangaza safu ya data tupu iitwayo `countArray` ndani ya kitendakazi cha `countup` juu kidogo ya tamko lako la `if`.

Safu hii ya data itahifadhi thamani ya `number` kila wito wa kitendakazi.

# --hints--

Unapaswa kutangaza safu ya data tupu iitwayo `countArray`.

```js
const explorer = await __helpers.Explorer(code);
const { countup } = explorer.functions;
assert.exists(countup);
const { countArray } = countup.variables;
assert.exists(countArray);
assert.isTrue(countArray.value.matches('[]'));
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
