---
id: 694870bcf3c6874112b6f951
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Kwa nambari zote kubwa kuliko na sawa na `1`, tengeneza kizuizi cha `else` na ndani yake, fanya wito wa kitendakazi cha `countup` kwa `number - 1` na weka thamani ya matokeo kwa `countArray`. 

Kisha tumia `.push()` kuongeza `number` kwa `countArray`, na hatimaye rudisha `countArray`. 

Hii inahakikisha mzunguko wa kurudia (recursion) hatimaye unafikia kesi ya msingi na unamaliza.

# --hints--

Unapaswa kuweka thamani ya matokeo ya wito wa `countup` kwa `number - 1` kwa `countArray`.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

Unapaswa kusukuma `number` ndani ya `countArray`.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

Unapaswa kurudisha `countArray`.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
