---
id: 686b9720ee1d032bd77a480a
title: ISBN 검증기 디버그하기
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

ISBN(국제 표준 도서 번호)는 상업용 도서에 할당된 고유 식별자입니다. 길이는 10자리 또는 13자리입니다. 마지막 문자는 체크 숫자로, 앞의 숫자들인 주요 숫자에서 계산됩니다.

Camperbot이 자체 ISBN 검증기를 만들려고 했지만, 몇 가지 실수를 했습니다.

이 실습에서는 기존 코드를 고쳐 제대로 작동하게 만드세요.

예상 동작:
사용자가 프로그램을 실행하면 `Enter ISBN and length: `라는 프롬프트가 표시됩니다. 사용자는 `ISBN,length` 형식으로 검증하려는 ISBN 코드를 입력할 수 있습니다. 하이픈 없이 ISBN을 입력하고, 쉼표와 길이(`10` 또는 `13`)를 입력하세요.

입력 예시:
ISBN-10 코드의 경우 `1530051126,10`.
ISBN-13 코드의 경우 `9781530051120,13`.

검증 숫자 찾는 방법:

- 주요 숫자는 ISBN-10 코드의 처음 9자리 숫자이거나 ISBN-13 코드의 처음 12자리 숫자입니다. 모두 숫자여야 합니다.
- 마지막 문자는 주어진 체크 숫자입니다. 예상 체크 숫자는 주요 숫자에서 계산되며, 두 값이 일치할 때만 ISBN 코드가 유효합니다.
- 예를 들어, ISBN-10 코드 `1530051126`에서 주요 숫자는 `153005112`이고 주어진 체크 숫자는 `6`입니다. ISBN-13 코드 `9781530051120`에서는 주요 숫자가 `978153005112`이고 주어진 체크 숫자는 `0`입니다.
- 이 실습에서는 상세 계산 로직을 알 필요가 없습니다. `calculate_check_digit_10`과 `calculate_check_digit_13` 함수가 계산을 처리하고 예상 체크 숫자를 문자열로 반환합니다.
- ISBN-10 코드의 체크 숫자는 `0`부터 `9`까지 숫자 또는 대문자 `X`일 수 있습니다.
- ISBN-13 코드의 체크 숫자는 `0`부터 `9`까지 숫자일 수 있습니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 현재 코드에서 `IndentationError`를 고치세요.
1. 사용자가 쉼표로 구분된 값을 입력하지 않아도 프로그램이 `IndexError`를 처리하여 중단되지 않게 하세요.
1. 사용자가 쉼표로 구분된 값을 입력하지 않으면 콘솔에 `Enter comma-separated values.` 메시지를 보여주고 프로그램을 종료하세요.
1. 사용자가 길이에 숫자가 아닌 값을 입력해도 프로그램이 `ValueError`를 처리하여 중단되지 않게 하세요.
1. 사용자가 길이에 숫자가 아닌 값을 입력하면 콘솔에 `Length must be a number.` 메시지를 보여주고 프로그램을 종료하세요.
1. `validate_isbn` 함수에서 경계값 오류를 고치세요.
1. 사용자가 유효한 ISBN 코드를 입력할 때 발생하는 `TypeError`를 고치세요.
1. 사용자가 유효한 ISBN 코드를 입력할 때 발생하는 `IndexError`를 고치세요.
1. 사용자가 주요 숫자에 숫자가 아닌 문자가 포함된 잘못된 ISBN 코드를 입력해도 프로그램이 발생하는 `ValueError`를 처리하여 중단되지 않게 하세요.
1. 사용자가 주요 숫자에 숫자가 아닌 문자가 포함된 잘못된 ISBN 코드를 입력하면 콘솔에 `Invalid character was found.` 메시지를 보여주세요.
1. 사용자가 `1530051126,10`을 입력하면 `Valid ISBN Code.` 메시지를 보여주세요.
1. 사용자가 `9781530051120,13`을 입력하면 `Valid ISBN Code.` 메시지를 보여주세요.

**중요:** 테스트가 제대로 실행되려면 전역 공간에서 `main()` 호출을 주석 처리해야 합니다.

프로젝트를 완료하면 사용자가 입력한 값에 따라 다음 메시지가 표시됩니다.

| ISBN 코드 | 길이 | 메시지 | 예시 입력 |
|-----------|--------|---------| ------------- |
| 유효 | 유효 | `Valid ISBN Code.` | `1530051126,10` |
| 잘못된 숫자 | 유효 | `Invalid ISBN Code.` | `1530051125,10` |
| 지정된 길이와 일치하지 않거나 비어 있음 | 유효 | `ISBN-10 code should be 10 digits long.` 또는 `ISBN-13 code should be 13 digits long.` (입력한 길이에 따라 다름) | `9781530051120,10` 또는 `1530051126,13` |
| 주요 숫자에 숫자가 아닌 문자가 포함됨 | 유효 | `Invalid character was found.` | `15-0051126,10` |
| 어떤 값이든 | 잘못된 숫자 | `Length should be 10 or 13.` | `1530051126,9` |
| 어떤 값이든 | 숫자가 아니거나 비어 있음 | `Length must be a number.` | `1530051125,A` |
| 쉼표로 구분되지 않음 | 쉼표로 구분되지 않음 | `Enter comma-separated values.` | `1530051125` |

원한다면 다음 값을 사용해 코드를 수동으로 테스트할 수 있습니다.

유효한 ISBN-10 코드 입력 예시:
`1530051126,10`
`9971502100,10`
`080442957X,10`

유효한 ISBN-13 코드 입력 예시:
`9781530051120,13`
`9781947172104,13`

# --hints--

나머지 테스트가 제대로 작동하려면 `main` 함수 호출을 주석 처리해야 합니다.

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

`validate_isbn` 함수를 만들어야 합니다.

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

`calculate_check_digit_10` 함수를 만들어야 합니다.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

`calculate_check_digit_13` 함수를 만들어야 합니다.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

`main` 함수를 만들어야 합니다.

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

사용자가 쉼표로 구분되지 않은 값을 입력하면 콘솔에 `Enter comma-separated values.` 메시지가 보여야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

사용자가 길이에 숫자가 아닌 값을 입력하면 콘솔에 `Length must be a number.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

사용자가 주요 숫자에 숫자가 아닌 문자가 포함된 잘못된 ISBN 코드를 입력하면 콘솔에 `Invalid character was found.` 메시지가 보여야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

사용자가 `1530051126,10`을 입력하면 콘솔에 `Valid ISBN Code.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

사용자가 `9781530051120,13`을 입력하면 `Valid ISBN Code.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

사용자가 `1530051125,10`을 입력하면 콘솔에 `Invalid ISBN Code.` 메시지가 보여야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

사용자가 `9781530051120,10`을 입력하면 `ISBN-10 code should be 10 digits long.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

사용자가 `1530051126,13`을 입력하면 `ISBN-13 code should be 13 digits long.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

사용자가 `15-0051126,10`을 입력하면 `Invalid character was found.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

사용자가 `1530051126,9`을 입력하면 `Length should be 10 or 13.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

사용자가 `1530051125,A`을 입력하면 `Length must be a number.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

사용자가 `1530051125`을 입력하면 `Enter comma-separated values.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

사용자가 `9971502100,10`을 입력하면 `Valid ISBN Code.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

사용자가 `080442957X,10`을 입력하면 `Valid ISBN Code.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

사용자가 `9781947172104,13`을 입력하면 `Valid ISBN Code.` 메시지가 표시되어야 합니다.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
