---
id: 66bba6fff611169359d9d36a
title: Construye una Página de Correo Electrónico Confidencial
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

Cumple con las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Instrucciones:**

1. Deberías tener un elemento `main` con un `id` de `email`.
1. Tu elemento `#email` debería tener `padding` de `50px`, un margen superior de `50px`, un `width` de `500px` y un `border` de `2px`.
1. El ancho total de tu elemento `#email`, incluyendo los paddings y bordes, debería ser de `500px`.
1. Deberías tener dos elementos `div`, uno con un id de `confidential` y el otro con un id de `top-secret`, dentro de tu elemento `#email`.
1. Tus elementos `#confidential` y `#top-secret` deberían tener un `display` de `inline-block`.
1. Tus elementos `#confidential` y `#top-secret` deberían tener un `padding`, un margen izquierdo y un `border`.
1. El elemento `#confidential` debería tener el texto `CONFIDENTIAL`.
1. El elemento `#top-secret` debería tener el texto `TOP SECRET`.
1. Tus elementos `#confidential` y `#top-secret` deberían estar rotados utilizando un transform de CSS.
1. Deberías tener al menos tres elementos de párrafo dentro de tu elemento `#email`.
1. Deberías tener al menos tres elementos `span` con una clase de `blurred`, dentro de tus elementos de párrafo.
1. Deberías tener un selector `blurred` que desenfoque el elemento `3px` utilizando un filtro CSS.

# --hints--

Deberías tener un elemento `main` con un `id` de `email`.

```js
assert.exists(document.querySelector('main#email'));
```

Deberías tener un selector `#email` que establezca el `padding` de sus elementos en todos los lados a `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

Deberías tener un selector `#email` que establezca el `margin-top` de sus elementos a `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

Deberías tener un selector `#email` que establezca el `width` de sus elementos a `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

Tu elemento `#email` debería tener un `border` de `2px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

Tu elemento `#email` debería tener un `box-sizing` de `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

Deberías tener un elemento `div` con un `id` de `confidential` dentro de tu elemento `#email`.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

Deberías tener un elemento `div` con un `id` de `top-secret` dentro de tu elemento `#email`.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

Tu elemento `#confidential` debería tener su `display` configurado como `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

Tu elemento `#top-secret` debería tener su `display` configurado como `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

Tu elemento `#confidential` debería tener un `padding` en todos los lados.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Tu elemento `#top-secret` debería tener un `padding` en todos los lados.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Tu elemento `#confidential` debería tener un margen izquierdo.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

Tu elemento `#top-secret` debería tener un margen izquierdo.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

Tu elemento `#confidential` debería tener un `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

Tu elemento `#top-secret` debería tener un `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

El texto del elemento `#confidential` debería ser `CONFIDENTIAL`. Has omitido el texto o tienes un error tipográfico.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

El texto del elemento `#top-secret` debería ser `TOP SECRET`. Has omitido el texto o tienes un error tipográfico.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

Tu elemento `#confidential` debería utilizar una transformación CSS `transform` para rotar el elemento.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

Tu elemento `#top-secret` debería utilizar una transformación CSS `transform` para rotar el elemento.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

Deberías tener al menos tres elementos de párrafo dentro de tu elemento `#email`.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

Deberías tener al menos tres elementos `span` con una clase de `blurred` dentro de tus párrafos.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

Tus elementos `.blurred` no deberían estar vacíos.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

Deberías tener un selector `.blurred` que establezca su elemento `filter` a `blur(3px)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
