---
id: 66b102ee0bdbad34a9f85ba0
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次研討會上，你將通過構建教師聊天機器人繼續學習字符串。

首先，添加一個 `console` 語句，其中包含 `"Hi there!"` 的消息。

# --hints--

你應該有一個 `console` 語句。

```js
assert.match(code, /console\.log(.*)/);
```

你的控制檯語句應該輸出消息 `"Hi there!"`。

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
