---
id: 66b102ee0bdbad34a9f85ba0
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將繼續透過建置 Teacher Chatbot 來學習字串。

首先，為你的程式碼添加一個 `console` 述語，訊息為 `"Hi there!"`。

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

你應該有一個 `console` 述語。

```js
assert.isAbove(logSpy.calls.length, 0);
```

你的主控台述語應該輸出訊息 `"Hi there!"`。

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
