---
id: 66b102ee0bdbad34a9f85ba0
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將繼續透過建置 Teacher Chatbot 來學習字串。 

首先，為你的程式碼添加一個 `console` 述語，訊息為 `"Hi there!"`。

# --hints--

你應該有一個 `console` 述語。

```js
assert.match(code, /console\.log(.*)/);
```

你的主控台述語應該輸出訊息 `"Hi there!"`。

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
