---
id: 66b5928b86e907fdfbf34e56
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

現在是設定機器人名稱的時候了。

創建一個名為 `botName` 的變數，並將字串值 `"teacherBot"` 指定給它。

# --hints--

你應該有一個名為 `botName` 的變數。

```js
assert.isNotNull(botName);
```

你的 `botName` 變數應該是一個字串。

```js
assert.isString(botName);
```

你的 `botName` 變數應該有 `"teacherBot"` 的值。

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
