---
id: 66b59be6ab830800c4df9146
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

現在，是時候使用你在前一步驟中創建的變數。

在之前的課程中，你學會了如何使用像這樣的模板字串：

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

首先建立一個 `sentence` 變數。

使用模板字面量語法，將字串 `Today, you will learn about [topic variable goes here] in [subject variable goes here].` 指定給 `sentence` 變數。

你將使用 `topic` 和 `subject` 變數以及 `${}` 語法來取代 `[topic variable goes here]` 和 `[subject variable goes here]` 佔位符。

最後，將 `sentence` 變數記錄到主控台。

# --hints--

你應該有一個 `sentence` 變數。

```js
assert.isNotNull(sentence);
```

你的 `sentence` 變數應該是一個字串。

```js
assert.isString(sentence);
```

你應該使用模板字面值語法，將字串 `Today, you will learn about [topic variable goes here] in [subject variable goes here].` 指定給 `sentence` 變數。若你需要幫助，請參考範例。

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

你應該將 `sentence` 變數記錄到主控台。

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
