---
id: 66b59829dba144ff1351220f
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

現在是時候使用 `botName` 變數來創建問候語。

在之前的課程中，你學會了如何使用像這樣的模板字面值來串接字串：

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

首先建立一個名為 `greeting` 的變數。

接著，使用模板字面值語法，指定一個字串，內容為 `My name is`，後面接著 `botName` 變數，並以句號（`.`）結尾。

最後，將 `greeting` 變數記錄到主控台。

# --hints--

你應該有一個名為 `greeting` 的變數。

```js
assert.isNotNull(greeting);
```

你應該為你的 `greeting` 變數指定一個字串。

```js
assert.isString(greeting);
```

你應該使用模板字面值將字串 `My name is` 與變數 `botName` 串接，後面接上一個句號（`.`）。

```js
assert.equal(greeting, "My name is teacherBot.");
```

你應該有一個 `console` 述語，將 `greeting` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
