---
id: 66b59b12a745e10011158f7b
title: 步驟4
challengeType: 1
dashedName: step-4
---

# --description--

下一步是創建一些將在未來的機器人消息中使用的變量。

創建一個名爲 `subject` 的變量併爲其分配字符串值 `"JavaScript"`。

然後創建一個名爲 `topic` 的變量併爲其分配字符串值 `"strings"`。

# --hints--

你應該有一個名爲 `subject` 的變量。

```js
assert.isNotNull(subject);
```

你應該爲 `subject` 變量分配一個字符串。

```js
assert.isString(subject);
```

你應該將字符串 `"JavaScript"` 分配給你的 `subject` 變量。

```js
assert.equal(subject, "JavaScript");
```

你應該有一個名爲 `topic` 的變量。

```js
assert.isNotNull(topic);
```

你應該爲 `topic` 變量分配一個字符串。

```js
assert.isString(topic);
```

你應該將字符串 `"strings"` 分配給你的 `topic` 變量。

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
