---
id: 66b59be6ab830800c4df9146
title: 步驟5
challengeType: 1
dashedName: step-5
---

# --description--

現在是時候使用你在上一步中創建的變量了。

In previous lessons, you learned how to work with template literals like this:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

首先創建一個 `sentence` 變量。

使用模板字面量語法，將字符串 `Today, you will learn about [topic variable goes here] in [subject variable goes here].` 分配給 `sentence` 變量。

你將用 `topic` 和 `subject` 變量和 `${}` 語法替換 `[topic variable goes here]` 和 `[subject variable goes here]` 佔位符。

最後，將 `sentence` 變量記錄到控制檯。

# --hints--

你應該有一個 `sentence` 變量。

```js
assert.isNotNull(sentence);
```

你的 `sentence` 變量應該是一個字符串。

```js
assert.isString(sentence);
```

你應該使用模板字面量語法將字符串 `Today, you will learn about [topic variable goes here] in [subject variable goes here].` 分配給 `sentence` 變量。 如果你需要幫助，請參考示例。

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

你應該將 `sentence` 變量記錄到控制檯。

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
