---
id: 66b59b12a745e10011158f7b
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

下一步是創建幾個將用於未來機器人訊息的變數。

創建一個名為 `subject` 的變數，並將字串值 `"JavaScript"` 指定給它。

然後創建一個名為 `topic` 的變數，並將字串值 `"strings"` 指定給它。

# --hints--

你應該有一個名為 `subject` 的變數。

```js
assert.isNotNull(subject);
```

你應該為你的 `subject` 變數指定一個字串。

```js
assert.isString(subject);
```

你應該將字串 `"JavaScript"` 指定到你的 `subject` 變數。

```js
assert.equal(subject, "JavaScript");
```

你應該有一個名為 `topic` 的變數。

```js
assert.isNotNull(topic);
```

你應該將字串指定到你的 `topic` 變數。

```js
assert.isString(topic);
```

你應該將字串 `"strings"` 指定到你的 `topic` 變數。

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
