---
id: 66b6e39031393e30f2c48d0f
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

要取得字串的長度，你可以使用 <dfn>length</dfn> 屬性。此屬性會傳回字串中的字元數量。

這裡有一個範例：

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

首先使用 `length` 屬性取得 `subject` 字串的長度，並將該值記錄到主控台。

# --hints--


你應該使用 `length` 屬性來取得 `subject` 字串的長度。

```js
assert.match(code, /subject\.length/);
```

你應該將 `subject` 字串的長度記錄到主控台。

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
