---
id: 66b6e62423e8a031d6c1f03d
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

現在是取得 `topic` 字串長度的時候了。

你可以在 `console` 述語中這樣使用模板字串：

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

首先將訊息 `Here is an example of using the length property on the word [topic].` 輸出到主控台。 

請記得將 `[topic]` 替換為 `topic` 變數，並使用正確的模板字串語法，就像你在前面的步驟中所做的那樣。

然後，為你的程式碼添加第二個 `console.log` 述語，將 `topic` 字串的長度輸出到主控台。

# --hints--

你應該使用模板字面值語法將訊息 `Here is an example of using the length property on the word [topic].` 輸出到 `console`。

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

你應該有第二個 `console` 述語，將 `topic` 變數的長度輸出到 `console`。

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
