---
id: 66b6f586767a1534f3097353
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

現在是時候存取 `subject` 變數的第二個字母。

首先為你的程式碼添加一個 `console` 述語，輸出訊息 `Here is an example of accessing the second letter in the word [subject].`。請記得將 `[subject]` 替換為 `subject` 變數的實際值，並使用正確的模板字面語法。

然後為你的程式碼添加另一個 `console` 述語，使用括號符號和正確的索引號輸出 `subject` 變數的第二個字母。

# --hints--

你的 `console` 述語應該輸出訊息 `Here is an example of accessing the second letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

你應該有另一個 `console` 述語，使用括號符號索引 `1` 輸出 `subject` 變數的第二個字母。

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
