---
id: 66b6ffb42b88e33943788abf
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

在工作坊的最後一部分，你將複習如何尋找字串中子字串的索引位置。請記得，子字串是字串的一部分。

首先建立一個名為 `learningIsFunSentence` 的變數，並將字串值 `"Learning is fun."` 指定給它。

# --hints--

你應該有一個名為 `learningIsFunSentence` 的變數。

```js
assert.isNotNull(learningIsFunSentence);
```

你的 `learningIsFunSentence` 應該保存一個字串的值。

```js
assert.isString(learningIsFunSentence);
```

你應該將 `"Learning is fun."` 的值指定給變數 `learningIsFunSentence`。

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
