---
id: 66b6f586767a1534f3097353
title: 步驟11
challengeType: 1
dashedName: step-11
---

# --description--

現在是時候訪問 `subject` 變量的第二個字母了。

首先添加一個輸出消息 `Here is an example of accessing the second letter in the word [subject].` 的 `console` 語句，記住用 `subject` 變量的實際值替換 `[subject]`，並使用正確的模板字面量語法。

然後添加另一個 `console` 語句，使用括號符號和正確的索引號輸出 `subject` 變量的第二個字母。

# --hints--

你的 `console` 語句應該輸出消息 `Here is an example of accessing the second letter in the word [subject].`。

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

你應該有另一個 `console` 語句，使用括號表示法索引 `1` 輸出 `subject` 變量的第二個字母。

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
