---
id: 66b6fdb76441c738719039fa
title: 步驟 13
challengeType: 1
dashedName: step-13
---

# --description--

在之前的課程中，你學會了如何像這樣存取字串中的最後一個字元：

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` 永遠會給你字串的最後一個索引號碼。

創建一個名為 `lastCharacter` 的新變數，並將 `subject` 變數中最後一個字元的值指定給它。

然後，將 `lastCharacter` 變數的值記錄到主控台。

# --hints--

你應該有一個名為 `lastCharacter` 的變數。

```js
assert.isNotNull(lastCharacter);
```

你應該將 `subject` 變數中最後一個字元的值指定給 `lastCharacter` 變數。如果你需要幫助，請參考範例。

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

你應該將 `lastCharacter` 變數的值記錄到 `console`。

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
