---
id: 66b6efddeca35833cd6f0b03
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

在之前的課程中，你學會了如何像這樣存取字串中的字元：

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

請記住索引號從 `0` 開始，因此字串中的第一個字母永遠位於索引 `0`。

首先為你的程式碼添加另一個 `console` 述語。 

在 `console` 述語中，使用括號表示法和正確的索引數字輸出 `subject` 變數的第一個字母。

# --hints--

你的 `console` 述語應該使用括號符號和索引 `0` 輸出 `subject` 變數的第一個字母。

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
