---
id: 66b6e80d6c3f0b329c360283
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

本工作坊的下一個零件是檢視如何從字串中存取字元。

首先將訊息 `Here is an example of accessing the first letter in the word [subject].` 輸出到主控台。 

請記得將 `[subject]` 替換為 `subject` 變數，並使用正確的模板字串語法，就像你在前面的步驟中所做的那樣。

# --hints--

你的 `console` 述語應該輸出訊息 `Here is an example of accessing the first letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+first\s+letter\s+in\s+the\s+word\s+\${subject}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

--fcc-editable-region--

--fcc-editable-region--
```
