---
id: 66b6d482bbb9e12f2e5ee1ae
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

在專案的下一部分中，機器人將教你如何使用字串的 `length` 屬性。

首先建立一個名為 `strLengthIntro` 的新變數。

然後使用模板字面量語法，將字串 `Here is an example of using the length property on the word [subject].` 指定給 `strLengthIntro` 變數。 

將 `[subject]` 取代為你先前使用的 `subject` 變數。

最後，將 `strLengthIntro` 變數記錄到主控台。

# --hints--

你應該有一個名為 `strLengthIntro` 的變數。

```js
assert.isNotNull(strLengthIntro);
```

你的 `strLengthIntro` 變數應該是一個字串。

```js
assert.isString(strLengthIntro);
```

你應該使用模板字面值語法，將字串 `Here is an example of using the length property on the word [subject goes here].` 指定給 `strLengthIntro` 變數。

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

你應該將 `strLengthIntro` 變數記錄到主控台。

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
