---
id: 66b6e62423e8a031d6c1f03d
title: 步驟 8
challengeType: 1
dashedName: step-8
---

# --description--

現在是時候獲取 `topic` 字符串的長度了。

你可以在 `console` 語句中使用模板字面量，如下所示：

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

首先向控制檯輸出消息 `Here is an example of using the length property on the word [topic].`。

記住用 `topic` 變量替換 `[topic]`，並使用正確的模板字面量語法，就像在前面的步驟中所做的那樣。

然後，添加第二個 `console.log` 語句，將 `topic` 字符串的長度輸出到控制檯。

# --hints--

你應該使用模板字面量語法將消息 `Here is an example of using the length property on the word [topic].` 輸出到 `console`。

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

你應該有第二個 `console` 語句來將 `topic` 變量的長度輸出到 `console`。

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
