---
id: 66b7137348cfb53fd3ec6c73
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

接著，你將新增一個 `console` 述語，輸出對 `learningIsFunSentence` 變數使用 `indexOf` 方法尋找子字串 `"fun"` 位置的結果。

在該 `console` 述語下方，為你的 `console` 添加一個新的 `console` 述語，輸出使用 `indexOf` 方法尋找子字串 `"learning"` 位置的結果。

注意最後一個 `console` 述語輸出的內容。 

# --hints--

你的 `console` 應該輸出 `learningIsFunSentence` 變數中子字串 `"fun"` 的索引位置。

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(fun)\1\s*\)\s*\)/);
```

你的第二個 `console` 述語應該輸出 `learningIsFunSentence` 變數中子字串 `"learning"` 的索引位置。

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

console.log(learningIsFunSentence.indexOf("Learning"));

--fcc-editable-region--

--fcc-editable-region--
```
