---
id: 66b70b8a611cbf3bcc5c6c5f
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

在之前的課程中，你學會了如何像這樣使用 `indexOf` 方法。

```js
const sentence = "I love to learn.";

// returns index 2
console.log(sentence.indexOf("love"));

// returns -1
console.log(sentence.indexOf("hate"));
```

請記住，`indexOf` 方法會傳回字串中子字串第一次出現的位置索引。如果找不到子字串，則傳回 `-1`。

為你的程式碼添加一個新的 `console` 述語，輸出對 `learningIsFunSentence` 變數使用 `indexOf` 方法以尋找子字串 `"Learning"` 的索引位置的結果。

# --hints--

你的 `console` 應該對 `learningIsFunSentence` 變數使用 `indexOf` 方法來尋找子字串 `"Learning"` 的索引位置。

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(Learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

--fcc-editable-region--

--fcc-editable-region--
```
