---
id: 587d78af367417b2b2512b04
title: 建置產品登陸頁面
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301144
dashedName: build-a-product-landing-page
---

# --description--

**目標：** 建置一個功能上類似於 <a href="https://product-landing-page.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://product-landing-page.freecodecamp.rocks</a> 的應用程式。**請勿複製此示範專案**。

**使用者故事：**

1. 你的產品登陸頁面應該有一個帶有對應 `id="header"` 的 `header` 元素。
1. 你可以在具有對應 `id="header-img"` 的 `header` 元素中看到一個影像（這裡放一個標誌會是個不錯的影像）。
1. 在 `#header` 元素內，你可以看到一個具有對應 `id="nav-bar"` 的 `nav` 元素。
1. 你可以看到在 `nav` 元素內至少有三個可點擊的元素，每個元素都有 `nav-link` 類別。
1. 當你點擊 `nav` 元素中的 `.nav-link` 按鈕時，你會被帶到登陸頁面的對應區段。
1. 你可以觀看內嵌的產品 `video`，其 `id="video"`。
1. 你的登陸頁面有一個帶有對應 `id="form"` 的 `form` 元素。
1. 在表單中，有一個 `input` 欄位，其 `id="email"`，你可以在此輸入電子郵件位址。
1. `#email` 輸入欄位應該有佔位文字，讓使用者知道該欄位的用途。
1. `#email` 輸入欄位使用 HTML5 驗證來確認輸入的文字是電子郵件位址
1. 在表單中，有一個帶有對應 `id="submit"` 的提交 `input`。
1. 當你點擊 `#submit` 元素時，電子郵件會提交到靜態頁面（使用此模擬 URL：`https://www.freecodecamp.com/email-submit`）
1. 導覽列應該始終位於視窗頂部
1. 你的產品登陸頁面應該至少有一個媒體查詢。
1. 你的產品登陸頁面應該至少使用一次 CSS flexbox。

完成使用者故事並通過以下所有測試以完成此專案。展現你個人的風格。祝你程式設計愉快！

**注意：** 請務必在你的 HTML 中添加 `<link rel="stylesheet" href="styles.css">` 以聯結你的樣式表並套用你的 CSS。

# --hints--

你應該有一個 `header` 元素，其 `id` 為 `header`。

```js
const el = document.getElementById('header')
assert(!!el && el.tagName === 'HEADER')
```

你應該有一個 `img` 元素，其 `id` 為 `header-img`。

```js
const el = document.getElementById('header-img')
assert(!!el && el.tagName === 'IMG')
```

你的 `#header-img` 應該是 `#header` 的後代。

```js
const els = document.querySelectorAll('#header #header-img')
assert(els.length > 0)
```

你的 `#header-img` 應該有一個 `src` 屬性。

```js
const el = document.getElementById('header-img')
assert(!!el && !!el.src)
```

你的 `#header-img` 的 `src` 值應該是有效的 URL（以 `http` 開頭）。

```js
const el = document.getElementById('header-img')
assert(!!el && /^http/.test(el.src))
```

你應該有一個 `nav` 元素，其 `id` 為 `nav-bar`。

```js
const el = document.getElementById('nav-bar')
assert(!!el && el.tagName === 'NAV')
```

你的 `#nav-bar` 應該是 `#header` 的子孫元素。

```js
const els = document.querySelectorAll('#header #nav-bar')
assert(els.length > 0)
```

你應該在 `#nav-bar` 中至少有 3 個 `.nav-link` 元素。

```js
const els = document.querySelectorAll('#nav-bar .nav-link')
assert(els.length >= 3)
```

每個 `.nav-link` 元素都應該有一個 `href` 屬性。

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  if (!el.href) assert(false)
})
assert(els.length > 0)
```

每個 `.nav-link` 元素應該聯結到登陸頁面上的對應元素（具有 `href`，其值為另一個元素的 id，例如 `#footer`）。

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  const linkDestination = el.getAttribute('href').slice(1)
  if (!document.getElementById(linkDestination)) assert(false)
})
assert(els.length > 0)
```

你應該有一個 `video` 或 `iframe` 元素，其 `id` 為 `video`。

```js
const el = document.getElementById('video')
assert(!!el && (el.tagName === 'VIDEO' || el.tagName === 'IFRAME'))
```

你的 `#video` 應該有一個 `src` 屬性。

```js
let el = document.getElementById('video')
const sourceNode = el.children;
let sourceElement = null;
if (sourceNode.length) {
  sourceElement = [...video.children].filter(el => el.localName === 'source')[0];
}
if (sourceElement) {
  el = sourceElement;
}
assert(el.hasAttribute('src'));
```

你應該有一個 `form` 元素，其 `id` 為 `form`。

```js
const el = document.getElementById('form')
assert(!!el && el.tagName === 'FORM')
```

你應該有一個 `input` 元素，其 `id` 為 `email`。

```js
const el = document.getElementById('email')
assert(!!el && el.tagName === 'INPUT')
```

你的 `#email` 應該是 `#form` 的子孫。

```js
const els = document.querySelectorAll('#form #email')
assert(els.length > 0)
```

你的 `#email` 應該有帶有佔位文字的 `placeholder` 屬性。

```js
const el = document.getElementById('email')
assert(!!el && !!el.placeholder && el.placeholder.length > 0)
```

你的 `#email` 應該透過將其 `type` 設為 `email` 來使用 HTML5 驗證。

```js
const el = document.getElementById('email')
assert(!!el && el.type === 'email')
```

你應該有一個 `input` 元素，其 `id` 為 `submit`。

```js
const el = document.getElementById('submit')
assert(!!el && el.tagName === 'INPUT')
```

你的 `#submit` 應該是 `#form` 的子孫元素。

```js
const els = document.querySelectorAll('#form #submit')
assert(els.length > 0)
```

你的 `#submit` 應該有 `submit` 的 `type`。

```js
const el = document.getElementById('submit')
assert(!!el && el.type === 'submit')
```

你的 `#form` 應該有一個 `action` 屬性，值為 `https://www.freecodecamp.com/email-submit`。

```js
const el = document.getElementById('form')
assert(!!el && el.action === 'https://www.freecodecamp.com/email-submit')
```

你的 `#email` 應該有一個 `name` 屬性為 `email`。

```js
const el = document.getElementById('email')
assert(!!el && el.name === 'email')
```

你的 `#nav-bar` 應該始終位於視窗頂部。

```js
  const timeout = (milliseconds) => new Promise((resolve) => setTimeout(resolve, milliseconds));

  const header = document.getElementById('header');
  const headerChildren = header.children;
  const navbarCandidates = [header, ...headerChildren];
  
  // Return smallest top position of all navbar candidates
  const getNavbarPosition = (candidates = []) => {
    return candidates.reduce(
      (min, candidate) =>
        Math.min(min, Math.abs(candidate?.getBoundingClientRect().top)),
      Infinity
    );
  };
  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the viewport '
  );

  window.scroll(0, 500);
  await timeout(1);

  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the ' +
      'viewport even after scrolling '
  );
    
  window.scroll(0, 0);
```

你的產品登陸頁面應該至少使用一個媒體查詢。

```js
const htmlSourceAttr = Array.from(document.querySelectorAll('source')).map(el => el.getAttribute('media'))
const cssCheck = new __helpers.CSSHelp(document).getCSSRules('media')
assert(cssCheck.length > 0 || htmlSourceAttr.length > 0);
```

你的產品登陸頁面應該至少使用一次 CSS Flexbox。

```js
const hasFlex = (rule) => ["flex", "inline-flex"].includes(rule.style?.display)
const stylesheet = new __helpers.CSSHelp(document).getStyleSheet()
const cssRules = new __helpers.CSSHelp(document).styleSheetToCssRulesArray(stylesheet)
const mediaRules = new __helpers.CSSHelp(document).getCSSRules('media')
const usesFlex = cssRules.find(rule => hasFlex(rule))
const usesFlexMedia = mediaRules.find(mediaRule => {
  return [...mediaRule.cssRules].find(rule => hasFlex(rule))
})
assert(usesFlex || usesFlexMedia)
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Product Landing Page</title>
  </head>
  <body>
    <header id="header">
      <nav id="nav-bar">
        <img
          id="header-img"
          src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
          max-height="50px"
        />
        <a href="#Features" class="nav-link">Features</a> |
        <a href="#Video" class="nav-link">See our facility!</a> |
        <a href="#Pricing" class="nav-link">Purchase</a>
        <hr />
      </nav>
    </header>
    <main>
      <h1>
        Pokemon Daycare Service
      </h1>
      <section id="Features">
        <h2>What we offer</h2>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">Guaranteed friendly and loving staff!</div>
        </div>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">
            Comfortable environment for Pokemon to explore and play!
          </div>
        </div>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">
            Multiple membership plans to fit your lifestyle!
          </div>
        </div>
      </section>
      <section id="Video">
        <h2>Check us out!</h2>
        A sneak peek into our facility:
        <br />
        <iframe
          id="video"
          width="520"
          height="347"
          src="https://www.youtube.com/embed/Nw-ksH2r6AQ"
          frameborder="0"
          allowfullscreen
          alt="A video tour of our facility"
        >
        </iframe>
      </section>
      <section id="Pricing">
        <h2>Membership Plans</h2>
        <div class="flex-mem">
          <div class="flex-mem-box">
            <font size="+2">Basic Membership</font><br />
            <ul>
              <li>One Pokemon</li>
              <li>Food and berries provided</li>
            </ul>
            <em>$9.99/month</em>
          </div>
          <div class="flex-mem-box">
            <font size="+2">Silver Membership</font><br />
            <ul>
              <li>Up to Three Pokemon</li>
              <li>Food and berries provided</li>
              <li>Grooming and accessories included</li>
            </ul>
            <em>$19.99/month</em>
          </div>
          <div class="flex-mem-box">
            <font size="+2">Gold Membership</font><br />
            <ul>
              <li>Up to six Pokemon!</li>
              <li>Food and berries provided</li>
              <li>Grooming and accessories included</li>
              <li>Personal training for each Pokemon</li>
              <li>Breeding and egg hatching</li>
            </ul>
            <em>$29.99/month</em>
          </div>
        </div>
      </section>
      <form id="form" action="https://www.freecodecamp.com/email-submit">
        <p>Sign up for our newsletter!</p>
        <label for="email"><p>Email:</p><input name="email" id="email" type="email" placeholder="johnsmith@email.com" required></label>
        <input type="submit" id="submit">
      </form>
      <footer>
        <a href="../">Return to Project List</a> |
        <a href="https://www.nhcarrigan.com">Return to HomePage</a>
      </footer>
    </main>
  </body>
</html>
```

```css
body {
  background-color: #3a3240;
  color: white;
}
main {
  max-width: 750px;
  margin: 50px auto;
}
input {
  background-color: #92869c;
}
a:not(.nav-link) {
  color: white;
}
#header-img {
  max-height: 25px;
}
#nav-bar {
  position: fixed;
  width: 100%;
  text-align: center;
  top: 0%;
  background-color: #92869c;
}
h1 {
  text-align: center;
}
body {
  text-align: center;
}
footer {
  text-align: center;
}
#bullet {
  max-height: 25px;
}
.flex-here {
  display: flex;
  justify-content: center;
}
.flex-left {
  height: 25px;
}
.flex-mem {
  display: flex;
  justify-content: center;
}
.flex-mem-box {
  background-color: #92869c;
  border-color: black;
  border-width: 5px;
  border-style: solid;
  margin: 10px;
  padding: 10px;
  color: black;
}
@media (max-width: 350px) {
  #video {
    width: 300;
    height: 200;
  }
}
```
