---
id: 587d825a367417b2b2512c89
title: Implementar o Algoritmo Quicksort
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

Neste laboratório você implementará o algoritmo de ordenação: quicksort.

Quicksort é uma abordagem eficiente, recursiva e de dividir para conquistar para ordenar um array. Neste método, um valor pivô é escolhido no array original. O array é então particionado em dois subarrays de valores menores e maiores que o valor pivô. Em seguida, o resultado da chamada recursiva do algoritmo quicksort em ambos os subarrays é combinado. Isso continua até que o caso base de um array vazio ou com um único item seja alcançado, que é retornado. O desenrolar das chamadas recursivas devolve o array ordenado.

Embora a escolha do valor pivô seja importante, qualquer pivô servirá para os nossos propósitos aqui. Para simplificar, podemos usar o primeiro ou o último elemento.

Quicksort is a very efficient sorting method, providing *O(nlog(n))* performance on average. Ele também é relativamente fácil de implementar. Estes atributos fazem dele um método de ordenação popular e útil.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função `quicksort` que recebe um array de inteiros como entrada.
1. A função `quicksort` deve retornar um array que contenha os mesmos inteiros da entrada, mas em ordem do menor para o maior.
1. A função `quicksort` deve chamar a si mesma recursivamente para ordenar o array.

# --hints--

Você deve ter uma função `quicksort`.

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` deve retornar um array que não foi alterado exceto pela ordem.

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` deve retornar um array ordenado (do menor para o maior).

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` não deve usar o método `.sort()` embutido.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
