---
id: 68e524c94e2e4ab7b4bfc7e2
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите калькулятор знижок, який зможе застосовувати різні стратегії знижок до продуктів. Система визначатиме найкращу ціну для клієнта на основі кількох правил знижок.

Почніть із визначення класу з назвою `Product`. Додайте до свого класу метод `__init__` з такими параметрами: `self`, `name` і `price`. Усередині методу `__init__` призначте параметри `name` і `price` атрибутам екземпляра з такими ж іменами.

# --hints--

Ви повинні створити клас з назвою `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

Ваш клас `Product` має мати метод `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

Ваш метод `__init__` має мати такі параметри: `self`, `name` і `price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

Ви повинні призначити `name` `self.name` всередині методу `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

Ви повинні призначити `price` `self.price` всередині методу `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
