---
id: 68e529226b0defdff26c4a8b
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Тепер ви почнете створювати інтерфейс системи знижок. Python надає спосіб визначати абстрактні класи, які слугують шаблонами для інших класів. Це робиться за допомогою класу `ABC` (Abstract Base Class) з модуля `abc`.

На початку свого коду додайте таку інструкцію імпорту:

```py
from abc import ABC, abstractmethod
```

Це імпортує клас `ABC` та декоратор `abstractmethod`, які ви будете використовувати для визначення методів, які мають бути реалізовані будь-яким класом, що наслідує ваш абстрактний клас.

# --hints--

У вас має бути згадана інструкція імпорту на початку коду.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
