---
id: 68e5293bd00d2fe134f58983
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Тепер додайте другий абстрактний метод з назвою `apply_discount` до вашого класу `DiscountStrategy`. Використовуйте декоратор `@abstractmethod` і надайте йому параметри `self` та `product`.

Додайте підказки типів: `product` має бути типу `Product`, а метод має повертати `float`. Використовуйте `pass` у тілі методу.

# --hints--

Ваш клас `DiscountStrategy` має містити метод з назвою `apply_discount`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

Ваш метод `apply_discount` має мати декоратор `@abstractmethod`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

Ваш метод `apply_discount` має мати параметри `self` та `product: Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

Ваш метод `apply_discount` має мати підказку типу повернення `-> float`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
