---
id: 68e5293bd00d2fe134f58984
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Тепер настав час створити конкретну реалізацію вашого абстрактного класу. Створіть клас з назвою `PercentageDiscount`, який наслідується від `DiscountStrategy`.

Додайте в нього метод `__init__`, який приймає `self` і `percent` як параметри. Вкажіть тип `int` для `percent` і встановіть тип повернення методу `None`. Усередині методу `__init__` збережіть `percent` як атрибут екземпляра.

# --hints--

Ви повинні створити клас з назвою `PercentageDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

Ваш клас `PercentageDiscount` має наслідуватися від `DiscountStrategy`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

Ваш клас `PercentageDiscount` має мати метод `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

Ваш метод `__init__` має мати параметри `self` і `percent: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

Ваш метод `__init__` має мати підказку типу повернення `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

Ви повинні призначити `percent` `self.percent` всередині вашого методу `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
