---
id: 68e52994bc7ea2e3dec530a1
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Тепер вам потрібен спосіб керувати кількома стратегіями знижок і знаходити найкращу ціну.

Створіть клас із назвою `DiscountEngine`. Цей клас керуватиме всіма стратегіями знижок і обчислюватиме найкращу ціну для продукту.

Надайте йому `__init__` метод, який прийматиме `self` і `strategies` як параметри, де `strategies` має типове позначення `list[DiscountStrategy]`. Використання `list[DiscountStrategy]` вказує, що параметр є списком, який містить об’єкти `DiscountStrategy`. Тип повернення методу має бути `None`. Усередині методу `__init__` збережіть `strategies` як атрибут екземпляра.

# --hints--

Ви повинні створити клас із назвою `DiscountEngine`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountEngine")`)
})
```

Ваш клас `DiscountEngine` має мати метод `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("__init__")`)
})
```

Ваш метод `__init__` має мати параметри `self` і `strategies: list[DiscountStrategy]`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_args("self, strategies: list[DiscountStrategy]")`)
})
```

Ваш метод `__init__` має мати тип повернення `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_returns("None")`)
})
```

Ви повинні призначити `strategies` `self.strategies` всередині вашого методу `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_stmt("self.strategies = strategies")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
