---
id: 68e52994bc7ea2e3dec5309e
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

Реалізуйте метод `is_applicable` для `PremiumUserDiscount`. Цей метод має перевіряти, чи є користувач преміум-учасником.

Поверніть `True`, якщо `user_tier.lower()` дорівнює `'premium'`, інакше поверніть `False`. Використання `.lower()` забезпечує нечутливість перевірки до регістру.

# --hints--

Ваш клас `PremiumUserDiscount` має містити метод `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").has_function("is_applicable")`)
})
```

Ваш метод `is_applicable` має мати параметри `self`, `product: Product` та `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").find_function("is_applicable").has_args("self, product:Product, user_tier:str")`)
})
```

Ваш метод `is_applicable` має мати підказку типу повернення `-> bool`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

Ваш метод `is_applicable` має повертати `True`, коли `user_tier.lower()` дорівнює `'premium'`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPremiumUserDiscount(PremiumUserDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPremiumUserDiscount()
    assert discount.is_applicable(test_product, "Premium") == True
    assert discount.is_applicable(test_product, "premium") == True
    assert discount.is_applicable(test_product, "PREMIUM") == True
  `)
})
```

Ваш метод `is_applicable` має повертати `False`, коли `user_tier` не є `'premium'`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPremiumUserDiscount(PremiumUserDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPremiumUserDiscount()
    assert discount.is_applicable(test_product, "regular") == False
    assert discount.is_applicable(test_product, "standard") == False
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

--fcc-editable-region--
class PremiumUserDiscount(DiscountStrategy):
    pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
