---
id: 68e5293bd00d2fe134f58982
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Як обговорювалося в попередньому уроці, абстрактні методи — це методи, які мають бути реалізовані будь-яким класом, що наслідує абстрактний базовий клас. Вони слугують як договір, гарантуючи, що всі стратегії знижок матимуть певні обов’язкові методи.

У вашому класі `DiscountStrategy` визначте метод з ім’ям `is_applicable`. Використайте декоратор `@abstractmethod` над визначенням методу. Метод має мати такі параметри: `self`, `product` і `user_tier`. Використайте інструкцію `pass` у тілі методу.

Додайте підказки типів, щоб зробити сигнатуру методу зрозумілою: `product` має бути типу `Product`, `user_tier` має бути типу `str`.


# --hints--

Ваш клас `DiscountStrategy` має мати метод з ім’ям `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("is_applicable")`)
})
```

Ваш метод `is_applicable` має мати декоратор `@abstractmethod`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_decorators("abstractmethod")`)
})
```

Ваш метод `is_applicable` має мати параметри `self`, `product: Product` і `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

--fcc-editable-region--
class DiscountStrategy(ABC):
    pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
