---
id: 68e529226b0defdff26c4a89
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Зараз, якщо ви створите екземпляр `Product` і виведете його, ви побачите стандартне представлення об’єкта, яке не дуже інформативне. Наприклад, воно може виглядати приблизно як `<__main__.Product object at 0x7f9b8c2d1d30>`.

Щоб перетворити стандартне представлення об’єкта на читабельний рядок, створіть метод `__str__` у вашому класі `Product`, який повертатиме рядок у форматі `name - $price` (де `name` і `price` мають бути замінені на значення відповідних атрибутів).

# --hints--

Ваш клас `Product` має містити метод `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

Ваш метод `__str__` має повертати рядок у форматі `name - $price`.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
--fcc-editable-region--
    
--fcc-editable-region--
```
