---
id: 68e529226b0defdff26c4a8a
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер створіть екземпляр `Product` з іменем `Wireless Mouse` та ціною `50.0`. Призначте його змінній з ім’ям `product`. Потім виведіть `product` у консоль, щоб побачити, як він виглядає.

# --hints--

У вас має бути змінна з ім’ям `product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

Ви повинні призначити екземпляр `Product` змінній `product`.

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

Ваш екземпляр `Product` має мати ім’я `Wireless Mouse` та ціну `50.0`.

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

Ви повинні вивести `product` у консоль.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

Коли ви виводите `Product` з ім’ям `'Wireless Mouse'` та ціною `50.0`, має відобразитися `Wireless Mouse - $50.0`.

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
```
