---
id: 68e5293bd00d2fe134f58981
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Створіть клас із назвою `DiscountStrategy`, який наслідує `ABC`. Це буде абстрактний базовий клас, що визначає інтерфейс для всіх стратегій знижок.

Поки що просто використайте `pass` у тілі класу, щоб створити порожній клас.

# --hints--

Ви маєте створити клас із назвою `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

Ваш клас `DiscountStrategy` має наслідувати `ABC`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
