---
id: 68e52994bc7ea2e3dec530a4
title: Крок 26
challengeType: 20
dashedName: step-26
---

# --description--

Тепер створіть цикл `for`, який ітерує `self.strategies`. Для кожного `strategy` перевірте, чи він застосовний, використовуючи метод `is_applicable`. Якщо так, обчисліть ціну зі знижкою за допомогою `apply_discount` і збережіть її у змінній з ім’ям `discounted`. Потім додайте цю ціну зі знижкою до списку `prices`.

# --hints--

Ви повинні створити цикл `for`, який ітерує `self.strategies`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").find_for_loops()[0].find_for_iter().is_equivalent("self.strategies")
  `)
})
```

Всередині циклу ваша інструкція `if` має бути `if strategy.is_applicable(product, user_tier):`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").find_for_loops()[0].find_ifs()[0].find_conditions()[0].is_equivalent("strategy.is_applicable(product, user_tier)")
  `)
})
```

Якщо стратегія застосовна, ви маєте викликати `apply_discount` і додати результат до `prices`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").find_for_loops()[0].find_ifs()[0].find_body().is_ordered("discounted = strategy.apply_discount(product)","prices.append(discounted)")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8

class DiscountEngine:
    def __init__(self, strategies: list[DiscountStrategy]) -> None:
        self.strategies = strategies

--fcc-editable-region--
    def calculate_best_price(self, product: Product, user_tier: str) -> float:
        prices = [product.price]
        

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
