---
id: 68e5293bd00d2fe134f58985
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Оскільки `PercentageDiscount` наслідує абстрактний клас, ви повинні реалізувати всі абстрактні методи. Почніть із реалізації методу `is_applicable`. Метод `is_applicable` має приймати `self`, `product: Product` і `user_tier: str` як параметри. Він повинен мати підказку типу повернення `bool`.

Цей метод має визначати, чи є відсоткова знижка дійсною. Поверніть `True`, якщо `self.percent` менше або дорівнює `70`, інакше поверніть `False`. Це запобігає нереалістичним знижкам, наприклад 90%.

# --hints--

Ваш клас `PercentageDiscount` має мати метод `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

Ваш метод `is_applicable` має мати параметри `self`, `product: Product` і `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

Ваш метод `is_applicable` має повертати `True`, коли `self.percent` менше або дорівнює `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

Ваш метод `is_applicable` має повертати `False`, коли `self.percent` більше за `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

Ваш метод `is_applicable` має мати підказку типу повернення `bool`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
