---
id: 68e79d91daf8fb4808f93c47
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Після реалізації методу `apply_discount` протестуйте свою стратегію знижки. Створіть екземпляр `PercentageDiscount` з `10` відсотками та призначте його змінній з ім’ям `discount`. Потім застосуйте знижку за допомогою `discount.apply_discount(product)` і виведіть результат, щоб побачити, як працює знижка.

# --hints--

Ви повинні створити змінну з ім’ям `discount`, яка містить екземпляр `PercentageDiscount`.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

Ваш екземпляр `PercentageDiscount` має бути ініціалізований з `10` відсотками.

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

Ви повинні викликати `discount.apply_discount(product)` і вивести результат.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
    
    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--
```
