---
id: 68e9207ddcba053c5cb24413
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

У Python підказки типів — це необов’язкові сигнали, які повідомляють іншим розробникам, яким очікується тип даних змінної або функції. Ось приклад, який додає підказку типу `str` до параметра `name`.

```py
def greet(name: str):
    return 'Hello, ' + name
```

Для параметрів методу `__init__` оновіть параметр `name`, додавши підказку типу `str`, а параметр `price` — підказку типу `float`.

# --hints--

Ви повинні додати підказку типу `str` до параметра `name` і підказку типу `float` до параметра `price` у методі `__init__` класу `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
