---
id: 68ef53ef276acfa40c83cb12
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Додайте підказку типу повернення `bool` до методу `is_applicable`, щоб чітко вказати, що він повертає `bool`.

# --hints--

Ваш метод `is_applicable` має мати підказку типу повернення `-> bool`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_returns("bool")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

--fcc-editable-region--
class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str):
        pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
