---
id: 68f119ad905b29ff59d858c8
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Як і в попередньому кроці, у вашому методі `__str__` додайте підказку типу повернення `str`, оскільки цей метод повертає рядкове представлення об’єкта.

# --hints--

Ваш метод `__str__` має мати підказку типу повернення `str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
