---
id: 587d8251367417b2b2512c64
title: 연결 리스트 내 검색
challengeType: 1
forumTopicId: 301715
dashedName: search-within-a-linked-list
---

# --description--

연결 리스트 클래스에 몇 가지 유용한 메서드를 더 추가해 보겠습니다. `Stack`와 `Queue` 클래스처럼 리스트가 비어 있는지 아닌지 알 수 있다면 유용하지 않을까요?

또한 연결 리스트에서 특정 요소를 찾을 수 있어야 합니다. 자료 구조를 탐색하는 것은 많이 연습할수록 좋습니다! `indexOf`를 인수로 받고, 해당 요소의 `element`를 반환하는 `index` 메서드를 만들어 봅시다. 만약 요소가 연결 리스트에서 발견되지 않으면 `-1`를 반환하세요.

반대로, `elementAt`를 인수로 받고 주어진 `index`에 있는 `element`를 반환하는 `index` 메서드도 구현해 봅시다. 만약 `element`가 발견되지 않으면 `undefined`를 반환하세요.

# --instructions--

연결 리스트가 비어 있는지 확인하는 `isEmpty` 메서드, 주어진 요소의 `indexOf`를 반환하는 `index` 메서드, 그리고 주어진 `elementAt`에 있는 `element`를 반환하는 `index` 메서드를 작성하세요.

# --hints--

`LinkedList` 클래스에 `isEmpty` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.isEmpty === 'function';
  })()
);
```

`isEmpty` 메서드는 연결 리스트에 요소가 하나 이상 있을 때 `false`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.isEmpty() === false;
  })()
);
```

`isEmpty` 메서드는 연결 리스트에 요소가 없을 때 `true`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    return test.isEmpty() === true;
  })()
);
```

`LinkedList` 클래스에 `indexOf` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.indexOf === 'function';
  })()
);
```

`indexOf` 메서드는 연결 리스트에서 주어진 요소의 인덱스를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('cat') === 0;
  })()
);
```

`indexOf` 메서드는 주어진 요소가 연결 리스트에서 발견되지 않으면 `-1`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('pony') === -1;
  })()
);
```

`LinkedList` 클래스에 `elementAt` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.elementAt === 'function';
  })()
);
```

`elementAt` 메서드는 연결 리스트에서 주어진 인덱스에 있는 요소를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(1) === 'dog';
  })()
);
```

`elementAt` 메서드는 주어진 인덱스에 요소가 없으면 `undefined`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(5) === undefined;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
      head = currentNode.next;
    } else {
      while(currentNode.element !== element) {
        previousNode = currentNode;
        currentNode = currentNode.next;
      }

      previousNode.next = currentNode.next;
    }

    length --;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
        head = currentNode.next;
    } else {
        while(currentNode.element !== element) {
            previousNode = currentNode;
            currentNode = currentNode.next;
        }

        previousNode.next = currentNode.next;
    }

    length --;
  };

  this.indexOf = function(element) {
    if (head === null) return -1

    let current = head;
    let index = 0;

    while (current.element !== element && current.next !== null) {
      current = current.next;
      index++
    }

    if (current.element !== element && current.next === null) {
      return -1
    }

    return index;
  }

  this.elementAt = function(index) {
    if (head === null) return undefined;

    let current = head;
    let currentIndex = 0;

    while (currentIndex !== index && current.next !== null) {
      current = current.next;
      currentIndex++
    }

    if (currentIndex !== index && current.next === null) {
      return undefined;
    }

    return current.element;
  }

  this.isEmpty = function() {
    return length === 0;
  }
}
```
