---
id: 587d825a367417b2b2512c87
title: 이중 연결 리스트 만들기
challengeType: 1
forumTopicId: 301626
dashedName: create-a-doubly-linked-list
---

# --description--

지금까지 만든 연결 리스트는 모두 단일 연결 리스트입니다. 여기서는 <dfn>이중 연결 리스트</dfn>를 만들어 보겠습니다. 이름에서 알 수 있듯, 이중 연결 리스트의 노드는 리스트 내 다음 노드와 이전 노드에 대한 참조를 가집니다.

이 덕분에 리스트를 양방향으로 탐색할 수 있지만, 각 노드가 이전 노드에 대한 추가 참조를 포함해야 하므로 더 많은 메모리를 사용합니다.

# --instructions--

`Node` 객체를 제공했고 `DoublyLinkedList`를 시작했습니다. 이중 연결 리스트에 `add`와 `remove`라는 두 개의 메서드를 추가하세요. `add` 메서드는 주어진 요소를 리스트에 추가해야 하며, `remove` 메서드는 리스트 내 주어진 요소의 모든 항목을 제거해야 합니다.

이 메서드들을 작성할 때 첫 번째 또는 마지막 요소 삭제 같은 모든 가능한 예외 상황을 주의 깊게 처리하세요. 또한 빈 리스트에서 항목을 제거하려 하면 `null`를 반환해야 합니다.

# --hints--

`DoublyLinkedList` 자료 구조가 존재해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

`DoublyLinkedList`는 `add`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.add == undefined) {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

`DoublyLinkedList`는 `remove`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.remove == undefined) {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

빈 리스트에서 항목을 제거하면 `null`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.remove(100) == null;
  })()
);
```

`add` 메서드는 리스트에 항목을 추가해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(5);
    test.add(6);
    test.add(723);
    return test.print().join('') == '56723';
  })()
);
```

각 노드는 이전 노드를 추적해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(50);
    test.add(68);
    test.add(73);
    return test.printReverse().join('') == '736850';
  })()
);
```

첫 번째 항목은 리스트에서 제거할 수 있어야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(25);
    return test.print().join('') == '3560';
  })()
);
```

마지막 항목도 리스트에서 제거할 수 있어야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(60);
    return test.print().join('') == '2535';
  })()
);
```

# --before-each--

```js
DoublyLinkedList.prototype = Object.assign(
  DoublyLinkedList.prototype,
  {
  
  print() {
    if (this.head == null) {
      return null;
    } else {
      var result = new Array();
      var node = this.head;
      while (node.next != null) {
        result.push(node.data);
        node = node.next;
      };
      result.push(node.data);
      return result;
    };
  },
  printReverse() {
    if (this.tail == null) {
      return null;
    } else {
      var result = new Array();
      var node = this.tail;
      while (node.prev != null) {
        result.push(node.data);
        node = node.prev;
      };
      result.push(node.data);
      return result;
    };
  }
});
```

# --seed--

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
