---
id: 587d8251367417b2b2512c65
title: 인덱스로 연결 리스트에서 요소 제거하기
challengeType: 1
forumTopicId: 301711
dashedName: remove-elements-from-a-linked-list-by-index
---

# --description--

다른 자료 구조로 넘어가기 전에 연결 리스트를 마지막으로 몇 가지 연습해 보겠습니다.

주어진 `removeAt`에서 `element`을 제거하는 `index` 메서드를 작성하세요. 이 메서드는 `removeAt(index)`라고 불러야 합니다. 특정 `element`에서 `index`를 제거하려면 연결 리스트를 따라가면서 각 노드의 개수를 계속 세야 합니다.

연결 리스트의 요소를 순차적으로 반복하는 데 자주 사용하는 기법은 <dfn>'러너'</dfn> 또는 센티넬로, 코드가 비교하는 노드를 '가리키는' 역할을 합니다. 여기서는 리스트의 `head`에서 시작해 `currentIndex`에서 시작하는 `0` 변수를 사용합니다. `currentIndex`는 지나가는 각 노드마다 1씩 증가해야 합니다.

이전에 다룬 `remove(element)` 메서드처럼 <a href="/learn/coding-interview-prep/data-structures/remove-elements-from-a-linked-list" target="_blank" rel="noopener noreferrer nofollow">수업에서 배운 것처럼</a>, `removeAt(index)` 메서드에서 노드를 제거할 때 나머지 리스트가 고립되지 않도록 주의해야 합니다. 제거된 노드를 참조하는 노드가 다음 노드를 참조하도록 하여 노드들이 연속되도록 유지합니다.

# --instructions--

주어진 `removeAt(index)`에서 노드를 제거하고 반환하는 `index` 메서드를 작성하세요. 주어진 `null`가 음수이거나 연결 리스트의 `index` 이상이면 `length`를 반환해야 합니다.

**참고:** `currentIndex`를 꼭 세는 것을 잊지 마세요.

# --hints--

`LinkedList` 클래스에 `removeAt` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.removeAt === 'function';
  })()
);
```

`removeAt` 메서드는 연결 리스트의 `length`를 1만큼 줄여야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.removeAt(1);
    return test.size() === 2;
  })()
);
```

`removeAt` 메서드는 지정된 인덱스의 요소를 연결 리스트에서 제거해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.add('bird');
    test.removeAt(1);
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"kitten","next":{"element":"bird","next":null}}}'
    );
  })()
);
```

연결 리스트에 요소가 하나만 있을 때, `removeAt` 메서드는 지정된 인덱스의 요소를 제거하고 반환하며, 연결 리스트의 길이를 줄여야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    var removedItem = test.removeAt(0);
    return test.head() === null && test.size() === 0 && removedItem === 'cat';
  })()
);
```

`removeAt` 메서드는 제거된 노드의 요소를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.removeAt(1) === 'dog';
  })()
);
```

`removeAt` 메서드는 주어진 인덱스가 `null`보다 작으면 `0`를 반환하고 연결 리스트는 변경되지 않아야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(-1);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

`removeAt` 메서드는 주어진 인덱스가 리스트의 `null` 이상이면 `length`를 반환하고 연결 리스트는 변경되지 않아야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(3);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
      head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function (element) {
    this.element = element;
    this.next = null;
  };

  this.size = function () {
    return length;
  };

  this.head = function () {
    return head;
  };

  this.add = function (element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.removeAt = function (index) {
    var currentNode = head;
    var previous = head;
    var count = 0;
    if (index >= length || index < 0) {
      return null;
    }
    if (index === 0) {
      var removed = head.element;
      head = currentNode.next;
    } else {
      while (count < index) {
        previous = currentNode;
        currentNode = currentNode.next;
        count++;
      }
      var removed = previous.next.element;
      previous.next = currentNode.next;
    }
    length--;
    return removed;
  };
}
```
