---
id: 587d8251367417b2b2512c63
title: 연결 리스트에서 요소 제거
challengeType: 1
forumTopicId: 301712
dashedName: remove-elements-from-a-linked-list
---

# --description--

연결 리스트 구현에 반드시 필요한 다음 중요한 메서드는 `remove` 메서드입니다. 이 메서드는 제거하려는 요소를 인수로 받아, 리스트를 검색하여 해당 요소를 포함하는 노드를 찾아 제거해야 합니다.

연결 리스트에서 노드를 제거할 때는 나머지 리스트가 고아가 되지 않도록 주의해야 합니다. 모든 노드의 `next` 속성은 리스트에서 다음 노드를 가리킨다는 점을 기억하세요. 예를 들어 중간 요소를 제거할 때는, 그 요소 이전 노드의 `next` 속성이 중간 요소의 `next` 속성(즉, 리스트의 다음 노드)을 가리키도록 연결을 유지해야 합니다.

이 설명이 다소 혼란스러울 수 있으니, 다시 콩가 라인 예시로 돌아가 개념을 명확히 하겠습니다. 콩가 라인에 서 있다고 상상해 보세요. 바로 앞에 있는 사람이 라인을 떠났다면, 그 사람은 더 이상 라인에 있는 누구의 손도 잡고 있지 않고, 당신도 그 사람이 떠난 뒤로 손을 잡고 있지 않습니다. 당신은 앞으로 나아가 다음 사람의 손을 잡습니다.

제거하려는 요소가 `head` 요소라면, `head`를 연결 리스트의 두 번째 노드로 다시 할당합니다.

# --instructions--

요소를 받아 연결 리스트에서 제거하는 `remove` 메서드를 작성하세요.

**참고:** 연결 리스트에서 요소가 제거될 때마다 리스트의 `length`가 1씩 줄어들어야 합니다.

# --hints--

`LinkedList` 클래스에는 `remove` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.remove === 'function';
  })()
);
```

첫 번째 노드가 제거될 때 `remove` 메서드는 `head`를 두 번째 노드로 다시 할당해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.remove('cat');
    return test.head().element === 'dog';
  })()
);
```

`remove` 메서드는 제거된 노드마다 연결 리스트의 `length`를 1씩 줄여야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('hamster');
    test.remove('cat');
    test.remove('fish');
    return test.size() === 2;
  })()
);
```

`remove` 메서드는 제거된 노드의 이전 노드 참조를 제거된 노드의 `next` 참조로 다시 할당해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('snake');
    test.add('kitten');
    test.remove('snake');
    return test.head().next.next.element === 'kitten';
  })()
);
```

`remove` 메서드는 요소가 연결 리스트에 없으면 리스트를 변경하지 않아야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.remove('elephant');
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    if (head === null) {
      return;
    }
    var previous;
    var currentNode = head;

    while (currentNode.next !== null && currentNode.element !== element) {
      previous = currentNode;
      currentNode = currentNode.next;
    }
    
    if (currentNode.next === null && currentNode.element !== element) {
      return;
    }
    else if (previous) {
      previous.next = currentNode.next;
    } else {
      head = currentNode.next;
    }

    length--;
  };
} 
```
