---
id: 587d8254367417b2b2512c71
title: ES6에서 Set 항목 제거하기
challengeType: 1
forumTopicId: 301713
dashedName: remove-items-from-a-set-in-es6
---

# --description--

`delete` 메서드를 사용해 ES6 Set에서 항목을 제거하는 연습을 해보세요.

먼저, ES6 Set을 만드세요:

```js
var set = new Set([1,2,3]);
```

이제 `delete` 메서드로 Set에서 항목을 제거하세요.

```js
set.delete(1);
console.log([...set]) // should return [ 2, 3 ]
```

# --instructions--

이제 정수 1, 2, 3, 4, 5로 구성된 Set을 만드세요.

값 2와 5를 제거한 다음 Set을 반환하세요.

# --hints--

Set에는 값 1, 3, 4가 포함되어야 합니다.

```js
assert(
  (function () {
    var test = checkSet();
    return test.has(1) && test.has(3) && test.has(4) && test.size === 3;
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet(){
  // Only change code below this line
  var set = null;

  // Only change code above this line
  return set;   
}
```

# --solutions--

```js
function checkSet(){
var set = new Set([1,2,3,4,5]);
set.delete(2);
set.delete(5);
return set;}
```
