---
id: 587d8254367417b2b2512c70
title: ES6에서 집합 생성 및 추가하기
challengeType: 1
forumTopicId: 301636
dashedName: create-and-add-to-sets-in-es6
---

# --description--

ES5를 다뤄봤으니 이제 ES6에서 비슷한 작업을 할 차례입니다. 훨씬 더 쉬워질 거예요. ES6에는 내장 자료 구조 `Set`가 있어서 직접 작성했던 많은 작업이 이제 포함되어 있습니다. 한번 살펴보겠습니다:

새로운 빈 집합을 생성하려면:

```js
var set = new Set();
```

값을 가진 집합을 생성할 수 있습니다:

```js
var set = new Set(1);
```

배열로 집합을 생성할 수도 있습니다:

```js
var set = new Set([1, 2, 3]);
```

집합을 생성한 후에는 `add` 메서드를 사용해 원하는 값을 추가할 수 있습니다:

```js
var set = new Set([1, 2, 3]);
set.add([4, 5, 6]);
```

참고로, 집합은 중복 값을 포함할 수 없는 자료 구조입니다:

```js
var set = new Set([1, 2, 3, 1, 2, 3]);
// set contains [1, 2, 3] only
```

# --instructions--

이번 연습에서는 다음 값을 가진 집합을 반환하세요: `1, 2, 3, 'Taco', 'Cat', 'Awesome'`

# --hints--

`Set`에는 `1, 2, 3, Taco, Cat, Awesome` 값만 포함되어야 합니다.

```js
assert(
  (function () {
    var test = checkSet();
    return (
      test.size == 6 &&
      test.has(1) &&
      test.has(2) &&
      test.has(3) &&
      test.has('Taco') &&
      test.has('Cat') &&
      test.has('Awesome')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet() {
  var set = new Set([1, 2, 3, 3, 2, 1, 2, 3, 1]);
  // Only change code below this line

  // Only change code above this line
  console.log(Array.from(set));
  return set;
}

checkSet();
```

# --solutions--

```js
function checkSet(){var set = new Set([1,2,3,'Taco','Cat','Awesome']);
return set;}
```
