---
id: 587d8255367417b2b2512c73
title: ES5 Set() 통합을 위한 전개 연산자와 노트 사용하기
challengeType: 1
forumTopicId: 301720
dashedName: use-spread-and-notes-for-es5-set-integration
---

# --description--

ES6 전개 연산자 `...`를 기억하시나요?

`...`는 ES6에서 반복 가능한 객체를 배열로 변환할 수 있습니다.

Set을 만들어서 전개 함수가 어떻게 작동하는지 살펴보겠습니다.

```js
var set = new Set([1,2,3]);
var setToArr = [...set]
console.log(setToArr) // returns [ 1, 2, 3 ]
```

# --instructions--

이번 연습에서는 Set 객체를 `checkSet` 함수에 전달할 것입니다. 이 함수는 Set의 값을 포함하는 배열을 반환해야 합니다.

이제 ES6 `Set()` 객체를 성공적으로 사용하는 방법을 배웠습니다. 잘했어요!

# --hints--

`checkSet(new Set([1,2,3,4,5,6,7])`는 `[1, 2, 3, 4, 5, 6, 7]`를 반환해야 합니다.

```js
var test = checkSet(new Set([1, 2, 3, 4, 5, 6, 7]));
assert.deepEqual(test, [1, 2, 3, 4, 5, 6, 7]);
```

# --seed--

## --seed-contents--

```js
function checkSet(set){
   // Only change code below this line

   // Only change code above this line
}
```

# --solutions--

```js
function checkSet(set){
return [...set];}
```
