---
id: 587d8250367417b2b2512c5f
title: 스택 클래스 만들기
challengeType: 1
forumTopicId: 301633
dashedName: create-a-stack-class
---

# --description--

지난 섹션에서는 스택이 무엇인지, 그리고 배열을 사용해 스택을 어떻게 나타낼 수 있는지에 대해 이야기했습니다. 이번 섹션에서는 직접 스택 클래스를 만들어 보겠습니다. 배열로 스택을 만들 수 있지만, 때로는 스택에 대한 제어를 제한하는 것이 더 좋을 때가 있습니다. `push`와 `pop` 메서드 외에도 스택에는 유용한 메서드가 더 있습니다. 스택 클래스에 `peek`, `isEmpty`, 그리고 `clear` 메서드를 추가해 보겠습니다.

# --instructions--

스택의 맨 위에 요소를 추가하는 `push` 메서드, 스택 맨 위의 요소를 제거하고 반환하는 `pop` 메서드, 스택 맨 위의 요소를 확인하는 `peek` 메서드, 스택이 비어 있는지 확인하는 `isEmpty` 메서드, 그리고 스택의 모든 요소를 제거하는 `clear` 메서드를 작성하세요. 일반적으로 스택에는 없지만, 컬렉션을 콘솔에 출력하는 `print` 도우미 메서드를 추가했습니다.

# --hints--

`Stack` 클래스에는 `push` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.push === 'function';
  })()
);
```

`Stack` 클래스에는 `pop` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.pop === 'function';
  })()
);
```

`Stack` 클래스에는 `peek` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.peek === 'function';
  })()
);
```

`Stack` 클래스에는 `isEmpty` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.isEmpty === 'function';
  })()
);
```

`Stack` 클래스에는 `clear` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.clear === 'function';
  })()
);
```

`peek` 메서드는 스택의 맨 위 요소를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.peek() === 'CS50' && test.peek() === 'CS50';
  })()
);
```

`pop` 메서드는 스택의 맨 위 요소를 제거하고 반환해야 합니다.

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.pop() === 'CS50' && test.pop() === 'CS61';
  })()
);
```

`isEmpty` 메서드는 스택에 요소가 없으면 true를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new Stack();
    return test.isEmpty();
  })()
);
```

`clear` 메서드는 스택의 모든 요소를 제거해야 합니다.

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    test.clear();
    return test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Stack() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
class Stack {
  constructor() {
    this.collection = [];
  }
  print() {
    console.log(this.collection);
  }
  push(val) {
    this.collection.push(val);
  }
  pop() {
    return this.collection.pop();
  }
  peek() {
    return this.collection[this.collection.length - 1];
  }
  isEmpty() {
    return this.collection.length === 0;
  }
  clear() {
    return (this.collection.length = 0);
  }
}
```
