---
id: 587d825b367417b2b2512c8d
title: ES6 JavaScript Map 만들기
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--

JavaScript의 새 버전은 내장 Map 객체를 제공하여 이전 도전 과제에서 직접 작성했던 많은 기능을 제공합니다. 이 Map 객체는 일반 JavaScript 객체와 비슷하지만, 일반 객체에는 없는 유용한 기능을 제공합니다. 예를 들어, ES6 Map은 추가된 항목의 삽입 순서를 추적합니다. 다음은 메서드에 대한 좀 더 완전한 개요입니다:

- `.has(key)`는 키의 존재 여부에 따라 true 또는 false를 반환합니다
- `.get(key)`는 키와 연관된 값을 반환합니다
- `.set(key, value)`는 새로운 키, 값 쌍을 설정합니다
- `.delete(key)`는 키, 값 쌍을 제거합니다
- `.clear()`는 모든 키, 값 쌍을 제거합니다
- `.keys()`는 삽입 순서대로 모든 키를 포함하는 새 맵 반복자 객체를 반환합니다
- `.values()`는 삽입 순서대로 모든 값을 포함하는 새 맵 반복자 객체를 반환합니다
- `.entries()`는 삽입 순서대로 모든 키, 값 쌍을 포함하는 새 맵 반복자 객체를 반환합니다

# --instructions--

JavaScript Map 객체를 정의하고 myMap이라는 변수에 할당하세요. `freeCodeCamp`, `Awesome!` 키, 값 쌍을 추가하세요.

# --hints--

`myMap` 객체가 존재해야 합니다.

```js
assert(typeof myMap === 'object');
```

`myMap`는 `freeCodeCamp`, `Awesome!` 키 값 쌍을 포함해야 합니다.

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```
