---
id: 587d8257367417b2b2512c7d
title: 이진 탐색 트리의 최솟값과 최댓값 높이 찾기
challengeType: 1
forumTopicId: 301641
dashedName: find-the-minimum-and-maximum-height-of-a-binary-search-tree
---

# --description--

지난 도전 과제에서는 트리가 균형을 잃을 수 있는 상황을 설명했습니다. 균형 개념을 이해하기 위해 다른 트리 속성인 높이를 살펴보겠습니다. 트리에서 높이는 루트 노드부터 특정 리프 노드까지의 거리를 나타냅니다. 가지가 많이 뻗은 트리 구조에서는 경로마다 높이가 다를 수 있지만, 주어진 트리에는 최솟값과 최댓값 높이가 존재합니다. 트리가 균형을 이루면 이 두 값은 최대 1만큼 차이 납니다. 즉, 균형 잡힌 트리에서는 모든 리프 노드가 같은 레벨에 있거나, 그렇지 않더라도 최대 한 레벨 차이만 납니다.

균형 속성은 트리 작업의 효율성을 결정하기 때문에 중요합니다. 지난 도전 과제에서 설명했듯이, 심하게 균형이 깨진 트리에서는 최악의 시간 복잡도가 발생합니다. 동적 데이터 세트를 가진 트리에서는 이 문제를 해결하기 위해 자체 균형 트리가 자주 사용됩니다. 대표적인 예로 AVL 트리, 레드-블랙 트리, B-트리가 있습니다. 이 트리들은 삽입이나 삭제로 인해 균형이 깨질 때 트리를 다시 균형 잡는 추가 내부 로직을 포함합니다.

**참고:** 높이와 비슷한 속성으로 깊이가 있는데, 이는 특정 노드가 루트 노드로부터 얼마나 떨어져 있는지를 나타냅니다.

# --instructions--

이진 트리에 대해 두 개의 메서드 `findMinHeight`와 `findMaxHeight`을 작성하세요. 이 메서드들은 각각 주어진 이진 트리 내에서 최솟값과 최댓값 높이를 정수 값으로 반환해야 합니다. 노드가 비어 있으면 높이를 `-1`로 할당하세요(이것이 기본 사례입니다). 마지막으로 세 번째 메서드 `isBalanced`를 추가하여 트리가 균형 잡혔는지 여부에 따라 `true` 또는 `false`를 반환하세요. 이를 판단하기 위해 방금 작성한 두 메서드를 사용할 수 있습니다.

# --hints--

`BinarySearchTree` 자료 구조가 존재해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

이진 탐색 트리는 `findMinHeight`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMinHeight == 'function';
  })()
);
```

이진 탐색 트리는 `findMaxHeight`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMaxHeight == 'function';
  })()
);
```

이진 탐색 트리는 `isBalanced`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isBalanced == 'function';
  })()
);
```

`findMinHeight` 메서드는 트리의 최솟값 높이를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMinHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMinHeight() == 1;
  })()
);
```

`findMaxHeight` 메서드는 트리의 최댓값 높이를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMaxHeight() == 5;
  })()
);
```

빈 트리는 높이 `-1`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    return test.findMaxHeight() == -1;
  })()
);
```

`isBalanced` 메서드는 트리가 균형이 깨진 이진 탐색 트리라면 `false`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.isBalanced() === false;
  })()
);
```

`isBalanced` 메서드는 트리가 균형 잡힌 이진 탐색 트리라면 `true`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(10);
    test.add(3);
    test.add(22);
    test.add(1);
    test.add(4);
    test.add(17);
    test.add(32);
    return test.isBalanced() === true;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
  this.findMinHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMinHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMinHeight(root.left) + 1;
    }
    const lHeight = this.findMinHeight(root.left);
    const rHeight = this.findMinHeight(root.right);
    return Math.min(lHeight, rHeight) + 1;
  };
  this.findMaxHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMaxHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMaxHeight(root.left) + 1;
    }
    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    return Math.max(lHeight, rHeight) + 1;
  };
  this.isBalanced = function(root = this.root) {
    if (root === null) {
      return true;
    }

    if (root.left === null && root.right === null) {
      return true;
    }

    if (root.left === null) {
      return this.findMaxHeight(root.right) <= 0;
    }

    if (root.right === null) {
      return this.findMaxHeight(root.left) <= 0;
    }

    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    if (Math.abs(lHeight - rHeight) > 1) {
      return false;
    }
    return this.isBalanced(root.left) && this.isBalanced(root.right);
  };
}
```
