---
id: 587d8251367417b2b2512c61
title: 연결 리스트에서 노드 다루기
challengeType: 1
forumTopicId: 301721
dashedName: work-with-nodes-in-a-linked-list
---

# --description--

컴퓨터 과학에서 자주 접하는 또 다른 자료 구조는 <dfn>연결 리스트</dfn>입니다. 연결 리스트는 '노드'라고 하는 데이터 요소들의 선형 모음으로, 각 노드는 다음 노드를 가리킵니다. 연결 리스트의 각 <dfn>노드</dfn>는 두 가지 핵심 정보를 포함합니다: `element` 자체와 다음 `node`에 대한 참조입니다.

당신이 콩가 라인에 있다고 상상해 보세요. 당신은 앞사람의 손을 잡고 있고, 뒤에 있는 사람이 당신의 손을 잡고 있습니다. 당신은 바로 앞사람을 볼 수 있지만, 그 사람이 앞에 있는 다른 사람들의 시야를 가리고 있습니다. 노드는 콩가 라인의 한 사람과 같습니다: 자신이 누구인지 알고 있고 바로 다음 사람만 볼 수 있지만, 앞이나 뒤에 있는 다른 사람들은 알지 못합니다.

# --instructions--

코드 편집기에서 `Kitten`와 `Puppy` 두 노드를 만들었고, `Kitten` 노드를 `Puppy` 노드에 수동으로 연결했습니다.

`Cat`와 `Dog` 노드를 만들고 수동으로 라인에 추가하세요.

# --hints--

`Puppy` 노드는 `Cat` 노드를 참조해야 합니다.

```js
assert(Puppy.next.element === 'Cat');
```

`Cat` 노드는 `Dog` 노드를 참조해야 합니다.

```js
assert(Cat.next.element === 'Dog');
```

# --seed--

## --seed-contents--

```js
var Node = function(element) {
  this.element = element;
  this.next = null;
};
var Kitten = new Node('Kitten');
var Puppy = new Node('Puppy');

Kitten.next = Puppy;
// Only change code below this line
```

# --solutions--

```js
// solution required
```
