---
id: 587d8255367417b2b2512c74
title: 우선순위 큐 클래스 만들기
challengeType: 1
forumTopicId: 301630
dashedName: create-a-priority-queue-class
---

# --description--

이번 도전 과제에서는 우선순위 큐를 만들게 됩니다. 우선순위 큐는 항목들이 우선순위를 나타내는 추가 정보를 가질 수 있는 특별한 유형의 큐입니다. 이 우선순위는 정수로 간단히 나타낼 수 있습니다. 항목의 우선순위는 항목이 큐에서 꺼내지는 순서를 결정할 때 삽입 순서보다 우선합니다. 만약 낮은 우선순위 항목들 뒤에 더 높은 우선순위 항목이 들어오면, 그 높은 우선순위 항목이 다른 모든 항목보다 먼저 꺼내집니다.

예를 들어, 우선순위 큐에 세 개의 항목이 있다고 가정해 봅시다:

```js
[['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

여기서 두 번째 값(정수)은 항목의 우선순위를 나타냅니다. 만약 `['human', 1]`를 우선순위 `1`로 추가한다면(낮은 우선순위가 우선한다고 가정할 때) 이 항목이 가장 먼저 꺼내지는 항목이 됩니다. 컬렉션은 다음과 같이 보일 것입니다:

```js
[['human', 1], ['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

코드 편집기에는 `PriorityQueue`가 이미 작성되어 있습니다. 항목을 우선순위와 함께 추가하는 `enqueue` 메서드, 항목을 제거하고 반환하는 `dequeue` 메서드, 큐에 있는 항목 수를 반환하는 `size` 메서드, 큐 앞에 있는 요소를 반환하는 `front` 메서드, 그리고 큐가 비어 있으면 `isEmpty`을, 그렇지 않으면 `true`를 반환하는 `false` 메서드를 추가해야 합니다.

`enqueue`는 위에 보여진 형식(`['human', 1]`)의 항목을 받아야 하며, 여기서 `1`는 우선순위를 나타냅니다. `dequeue`과 `front`는 항목의 이름만 반환해야 하며, 우선순위는 반환하지 않아야 합니다.

# --hints--

`PriorityQueue` 클래스에는 `enqueue` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.enqueue === 'function';
  })()
);
```

`PriorityQueue` 클래스에는 `dequeue` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.dequeue === 'function';
  })()
);
```

`PriorityQueue` 클래스에는 `size` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.size === 'function';
  })()
);
```

`PriorityQueue` 클래스에는 `front` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.front === 'function';
  })()
);
```

`PriorityQueue` 클래스에는 `isEmpty` 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.isEmpty === 'function';
  })()
);
```

`PriorityQueue` 클래스는 항목이 추가되고 제거될 때 `size` 메서드를 사용해 현재 항목 수를 정확히 추적해야 합니다.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    test.enqueue(['Jon Snow', 1]);
    var size1 = test.size();
    test.dequeue();
    var size2 = test.size();
    test.enqueue(['A', 3]);
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    return size1 === 2 && size2 === 1 && test.size() === 4;
  })()
);
```

`front` 메서드는 항목이 추가되고 제거될 때 큐 앞에 있는 올바른 항목을 반환해야 합니다.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    var front1 = test.front();
    test.enqueue(['Jon Snow', 1]);
    var front2 = test.front();
    test.dequeue();
    test.enqueue(['A', 3]);
    var front3 = test.front();
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    test.dequeue();
    var front4 = test.front();
    return (
      front1 === 'David Brown' &&
      front2 === 'Jon Snow' &&
      front3 === 'David Brown' &&
      front4 === 'A'
    );
  })()
);
```

`isEmpty` 메서드는 큐가 비어 있을 때 `true`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 1]);
    test.enqueue(['B', 1]);
    test.dequeue();
    var first = test.isEmpty();
    test.dequeue();
    return !first && test.isEmpty();
  })()
);
```

우선순위 큐는 낮은 우선순위 항목보다 높은 우선순위 항목을 먼저 반환해야 하며, 그렇지 않은 경우에는 선입선출 순서로 항목을 반환해야 합니다.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 5]);
    test.enqueue(['B', 5]);
    test.enqueue(['C', 5]);
    test.enqueue(['D', 3]);
    test.enqueue(['E', 1]);
    test.enqueue(['F', 7]);
    var result = [];
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    return result.join('') === 'EDABCF';
  })()
);
```

# --seed--

## --seed-contents--

```js
function PriorityQueue () {
  this.collection = [];
  this.printCollection = function() {
    console.log(this.collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function PriorityQueue() {
  this.collection = [];
  this.printCollection = function () {
    console.log(this.collection);
  };
  // Only change code below this line
  this.enqueue = function (newitem) {
    if (this.isEmpty()) {
      return this.collection.push(newitem);
    }

    this.collection = this.collection.reverse();
    var found_index = this.collection.findIndex(function (item) {
      return newitem[1] >= item[1];
    });
    if (found_index === -1) {
      this.collection.push(newitem);
    } else {
      this.collection.splice(found_index, 0, newitem);
    }
    this.collection = this.collection.reverse();
  };
  this.dequeue = function () {
    if (!this.isEmpty()) {
      return this.collection.shift()[0];
    } else {
      return "The queue is empty.";
    }
  };
  this.size = function () {
    return this.collection.length;
  };
  this.front = function () {
    return this.collection[0][0];
  };
  this.isEmpty = function () {
    return this.size() > 0 ? false : true;
  };
  // Only change code above this line
}
```
