---
id: 587d8256367417b2b2512c78
title: 인접 행렬
challengeType: 1
forumTopicId: 301621
dashedName: adjacency-matrix
---

# --description--

그래프를 나타내는 또 다른 방법은 <dfn>인접 행렬</dfn>로 표현하는 것입니다. <dfn>인접 행렬</dfn>은 2차원(2D) 배열로, 각 중첩된 배열은 외부 배열과 같은 수의 요소를 가집니다. 다시 말해, 숫자로 이루어진 행렬 또는 격자이며, 숫자는 간선을 나타냅니다.

**참고**: 행렬의 위쪽과 왼쪽에 있는 숫자는 노드의 레이블일 뿐입니다. 행렬 내부에서 1은 행과 열이 나타내는 정점(노드) 사이에 간선이 존재함을 의미합니다. 0은 간선이나 관계가 없음을 뜻합니다.

<pre>
    1 2 3
  \------
1 | 0 1 1
2 | 1 0 0
3 | 1 0 0
</pre>

위는 매우 간단한 무방향 그래프로, 세 개의 노드가 있고 첫 번째 노드는 두 번째와 세 번째 노드에 연결되어 있습니다. 아래는 같은 내용을 JavaScript로 구현한 예입니다.

```js
var adjMat = [
  [0, 1, 1],
  [1, 0, 0],
  [1, 0, 0]
];
```

인접 리스트와 달리, 행렬의 각 "행"은 그래프의 노드 수와 같은 수의 요소를 가져야 합니다. 여기서는 3x3 행렬이므로 그래프에 노드가 세 개 있다는 뜻입니다. 방향 그래프도 비슷하게 나타냅니다. 아래는 첫 번째 노드가 두 번째 노드를 가리키는 간선을 가지고, 두 번째 노드가 세 번째 노드를 가리키는 간선을 가진 그래프입니다.

```js
var adjMatDirected = [
  [0, 1, 0],
  [0, 0, 1],
  [0, 0, 0]
];
```

그래프의 간선에는 <dfn>가중치</dfn>가 있을 수도 있습니다. 지금까지는 간선의 존재 여부만 이진수(`0` 또는 `1`)로 나타내는 <dfn>가중치 없는</dfn> 간선이었습니다. 애플리케이션에 따라 다양한 가중치를 가질 수 있습니다.

# --instructions--

다섯 개 노드를 가진 무방향 그래프의 인접 행렬을 만드세요. 이 행렬은 다차원 배열이어야 합니다. 이 다섯 노드 사이에는 첫 번째와 네 번째 노드, 첫 번째와 세 번째 노드, 세 번째와 다섯 번째 노드, 네 번째와 다섯 번째 노드 사이에 관계가 있습니다. 모든 간선 가중치는 1입니다.

# --hints--

`undirectedAdjList`에는 다섯 개 노드만 포함되어야 합니다.

```js
assert(
  adjMatUndirected.length === 5 &&
    adjMatUndirected
      .map(function (x) {
        return x.length === 5;
      })
      .reduce(function (a, b) {
        return a && b;
      })
);
```

첫 번째와 네 번째 노드 사이에 간선이 있어야 합니다.

```js
assert(adjMatUndirected[0][3] === 1 && adjMatUndirected[3][0] === 1);
```

첫 번째와 세 번째 노드 사이에 간선이 있어야 합니다.

```js
assert(adjMatUndirected[0][2] === 1 && adjMatUndirected[2][0] === 1);
```

세 번째와 다섯 번째 노드 사이에 간선이 있어야 합니다.

```js
assert(adjMatUndirected[2][4] === 1 && adjMatUndirected[4][2] === 1);
```

네 번째와 다섯 번째 노드 사이에 간선이 있어야 합니다.

```js
assert(adjMatUndirected[3][4] === 1 && adjMatUndirected[4][3] === 1);
```

# --seed--

## --seed-contents--

```js
var adjMatUndirected = [];
```

# --solutions--

```js
var adjMatUndirected = [
  [0, 0, 1, 1, 0],
  [0, 0, 0, 0, 0],
  [1, 0, 0, 0, 1],
  [1, 0, 0, 0, 1],
  [0, 0, 1, 1, 0]
];
```
