---
id: 587d825c367417b2b2512c90
title: 너비 우선 탐색
challengeType: 1
forumTopicId: 301622
dashedName: breadth-first-search
---

# --description--

지금까지 그래프를 나타내는 다양한 방법을 배웠습니다. 다음 질문은 그래프 내 임의의 두 노드 사이의 거리가 얼마인지일 수 있습니다. 여기서 <dfn>그래프 탐색 알고리즘</dfn>이 등장합니다.

<dfn>탐색 알고리즘</dfn>은 그래프 내 노드를 순회하거나 방문하는 알고리즘입니다. 탐색 알고리즘의 한 종류가 너비 우선 탐색 알고리즘입니다.

이 알고리즘은 한 노드에서 시작해 한 간선 떨어진 모든 이웃 노드를 방문합니다. 그다음 각 이웃의 이웃을 방문하는 식으로 모든 노드에 도달할 때까지 진행합니다.

너비 우선 탐색 알고리즘을 구현하는 데 중요한 자료 구조는 큐입니다. 큐는 한쪽 끝에서 요소를 추가하고 반대쪽 끝에서 요소를 제거할 수 있는 배열입니다. 이것은 <dfn>FIFO</dfn> 또는 <dfn>선입선출</dfn> 자료 구조로도 알려져 있습니다.

시각적으로 이 알고리즘은 다음과 같습니다:
<img alt="animation showing the breadth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/breadth-first-search.gif'>

회색 음영은 노드가 큐에 추가되는 것을 나타내고, 검은색 음영은 노드가 큐에서 제거되는 것을 나타냅니다. 노드가 큐에서 제거될 때마다(노드가 검은색으로 변함) 모든 이웃 노드가 큐에 추가되는 것(노드가 회색으로 변함)을 확인하세요.

이 알고리즘을 구현하려면 그래프 구조와 시작할 노드를 입력해야 합니다.

먼저 시작 노드로부터의 거리, 즉 간선 수를 인지해야 합니다. 모든 거리는 `Infinity` 같은 큰 수로 초기화하세요. 이는 시작 노드에서 도달할 수 없는 노드에 대한 계산 문제를 방지합니다. 다음으로 시작 노드에서 이웃 노드로 이동합니다. 이 이웃 노드는 한 간선 떨어져 있으므로, 추적하는 거리 값에 1을 더해야 합니다.

# --instructions--

인접 행렬 그래프(2차원 배열)와 노드 레이블 root를 매개변수로 받는 함수 `bfs()`을 작성하세요. 노드 레이블은 `0`와 `n - 1` 사이의 노드 정수 값이며, `n`는 그래프 내 노드 총 개수입니다.

함수는 노드와 루트로부터의 거리를 키-값 쌍으로 가진 JavaScript 객체를 출력해야 합니다. 노드에 도달할 수 없으면 거리는 `Infinity`여야 합니다.

# --hints--

입력 그래프 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`와 시작 노드 `1`는 `{0: 1, 1: 0, 2: 1, 3: 2}`을 반환해야 합니다.

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 1);
assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: 2 });
```

입력 그래프 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]`와 시작 노드 `1`는 `{0: 1, 1: 0, 2: 1, 3: Infinity}`을 반환해야 합니다.

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 0],
  [0, 0, 0, 0]
];
var results = bfs(graph, 1);
 assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: Infinity });
```

입력 그래프 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`와 시작 노드 `0`는 `{0: 0, 1: 1, 2: 2, 3: 3}`을 반환해야 합니다.

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 0);
 assert.deepEqual(results, { 0: 0, 1: 1, 2: 2, 3: 3 });
```

입력 그래프 `[[0, 1], [1, 0]]`와 시작 노드 `0`는 `{0: 0, 1: 1}`을 반환해야 합니다.

```js
var graph = [
  [0, 1],
  [1, 0]
];
var results = bfs(graph, 0);
assert.deepEqual(results, { 0: 0, 1: 1 });
```

# --seed--

## --seed-contents--

```js
function bfs(graph, root) {
  var nodesLen = {};

  return nodesLen;
};

var exBFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(bfs(exBFSGraph, 3));
```

# --solutions--

```js
function bfs(graph, root) {
  var nodesLen = {};
  // Set all distances to infinity
  for (var i = 0; i < graph.length; i++) {
    nodesLen[i] = Infinity;
  }
  nodesLen[root] = 0; // ...except root node
  var queue = [root]; // Keep track of nodes to visit
  var current; // Current node traversing
  // Keep on going until no more nodes to traverse
  while (queue.length !== 0) {
    current = queue.shift();
    // Get adjacent nodes from current node
    var curConnected = graph[current]; // Get layer of edges from current
    var neighborIdx = []; // List of nodes with edges
    var idx = curConnected.indexOf(1); // Get first edge connection
    while (idx !== -1) {
      neighborIdx.push(idx); // Add to list of neighbors
      idx = curConnected.indexOf(1, idx + 1); // Keep on searching
    }
    // Loop through neighbors and get lengths
    for (var j = 0; j < neighborIdx.length; j++) {
      // Increment distance for nodes traversed
      if (nodesLen[neighborIdx[j]] === Infinity) {
        nodesLen[neighborIdx[j]] = nodesLen[current] + 1;
        queue.push(neighborIdx[j]); // Add new neighbors to queue
      }
    }
  }
  return nodesLen;
}
```
