---
id: 587d8253367417b2b2512c6a
title: 타입이 있는 배열
challengeType: 1
forumTopicId: 301716
dashedName: typed-arrays
---

# --description--

배열은 다양한 요소를 담을 수 있는 JavaScript 객체입니다.

```js
var complexArr = [1, 5, "2", "Word", {"name": "James"}];
```

기본적으로 브라우저가 자동으로 해당 배열에 적절한 메모리 공간을 할당합니다. 데이터가 추가되거나 제거되면 필요에 따라 크기가 변경됩니다.

하지만 고성능 환경이나 다양한 요소 타입을 다룰 때는 배열에 할당할 메모리 크기를 더 구체적으로 지정해야 할 때가 있습니다.

<dfn>타입이 있는 배열</dfn>이 이 문제에 대한 해답입니다. 이제 배열에 할당할 메모리 크기를 직접 지정할 수 있습니다. 아래는 사용 가능한 배열 타입과 각 요소의 바이트 크기에 대한 기본 개요입니다.

<table><tbody><tr><th>타입</th><th>각 요소 크기(바이트)</th></tr><tr><td><code>Int8Array</code></td><td>1</td></tr><tr><td><code>Uint8Array</code></td><td>1</td></tr><tr><td><code>Uint8ClampedArray</code></td><td>1</td></tr><tr><td><code>Int16Array</code></td><td>2</td></tr><tr><td><code>Uint16Array</code></td><td>2</td></tr><tr><td><code>Int32Array</code></td><td>4</td></tr><tr><td><code>Uint32Array</code></td><td>4</td></tr><tr><td><code>Float32Array</code></td><td>4</td></tr><tr><td><code>Float64Array</code></td><td>8</td></tr></tbody></table>

이런 배열을 만드는 방법은 두 가지가 있습니다. 한 가지는 직접 만드는 방법입니다. 아래는 길이가 3인 `Int16Array`를 만드는 방법입니다.

```js
var i8 = new Int16Array(3);
console.log(i8);
// Returns [0, 0, 0]
```

또한 배열이 차지할 데이터 크기(바이트)를 할당하는 <dfn>버퍼</dfn>를 만들어서 생성할 수도 있습니다. **참고**  
버퍼를 사용해 타입이 있는 배열을 만들 때는 위 표에 나온 바이트 수의 배수로 바이트 수를 할당해야 합니다.

```js
// Create same Int16Array array differently
var byteSize = 6; // Needs to be multiple of 2
var buffer = new ArrayBuffer(byteSize);
var i8View = new Int16Array(buffer);
buffer.byteLength; // Returns 6
i8View.byteLength; // Returns 6
console.log(i8View); // Returns [0, 0, 0]
```

<dfn>버퍼</dfn>는 데이터를 담는 일반 목적 객체입니다. 일반적으로 직접 접근할 수 없습니다. 접근하려면 먼저 <dfn>뷰</dfn>를 만들어야 합니다.

```js
i8View[0] = 42;
console.log(i8View); // Returns [42, 0, 0]
```

**참고**  
타입이 있는 배열은 전통적인 배열이 가진 `.pop()`나 `.push()` 같은 메서드를 일부 갖고 있지 않습니다. 또한 타입이 있는 배열은 `Array.isArray()` 검사도 실패합니다. 더 단순하지만, 이는 덜 복잡한 JavaScript 엔진이 구현하기에 유리할 수 있습니다.

# --instructions--

먼저 64바이트 크기의 `buffer`를 만드세요. 그리고 그것을 기반으로 `Int32Array`라는 뷰를 가진 `i32View` 타입이 있는 배열을 만드세요.

# --hints--

`buffer`는 64바이트 크기여야 합니다.

```js
assert(buffer.byteLength === 64);
```

버퍼에 대한 `i32View` 뷰도 64바이트 크기여야 합니다.

```js
assert(i32View.byteLength === 64);
```

버퍼에 대한 `i32View` 뷰는 16개 요소 길이여야 합니다.

```js
assert(i32View.length === 16);
```

# --seed--

## --seed-contents--

```js
var buffer;
var i32View;
```

# --solutions--

```js
var buffer = new ArrayBuffer(64);
var i32View = new Int32Array(buffer);
```
