---
id: 587d8256367417b2b2512c77
title: 인접 리스트
challengeType: 1
forumTopicId: 301620
dashedName: adjacency-list
---

# --description--

그래프는 여러 가지 방법으로 나타낼 수 있습니다. 여기서는 <dfn>인접 리스트</dfn>라고 하는 한 가지 방법을 설명합니다. 인접 리스트는 본질적으로 왼쪽에 노드가 있고 오른쪽에 연결된 다른 모든 노드를 나열하는 글머리 기호 목록입니다. 아래는 인접 리스트의 표현입니다.

<blockquote>Node1: Node2, Node3<br>Node2: Node1<br>Node3: Node1</blockquote>

위는 무방향 그래프입니다. `Node1`가 `Node2`과 `Node3`에 연결되어 있고, 이 정보가 `Node2`과 `Node3`가 보여주는 연결과 일치하기 때문입니다. 방향 그래프의 인접 리스트는 각 행이 방향을 나타냅니다. 위가 방향 그래프라면 `Node2: Node1`는 방향 간선이 `Node2`에서 `Node1`으로 향한다는 뜻입니다. 위 무방향 그래프는 JavaScript 객체 안에 넣어 인접 리스트로 나타낼 수 있습니다.

```js
var undirectedG = {
  Node1: ["Node2", "Node3"],
  Node2: ["Node1"],
  Node3: ["Node1"]
};
```

또한 노드가 문자열 레이블 대신 숫자로만 되어 있는 배열로 더 간단히 나타낼 수도 있습니다.

```js
var undirectedGArr = [
  [1, 2], // Node1
  [0],    // Node2
  [0]     // Node3
];
```

# --instructions--

`James`, `Jill`, `Jenny`, `Jeff`라는 이름의 4개 노드/사람으로 무방향 그래프인 소셜 네트워크를 만드세요. James와 Jeff, Jill과 Jenny, Jeff와 Jenny 사이에 간선/관계가 있습니다.

# --hints--

`undirectedAdjList`에는 노드가 네 개만 있어야 합니다.

```js
assert(Object.keys(undirectedAdjList).length === 4);
```

`Jeff`과 `James` 사이에 간선이 있어야 합니다.

```js
assert(
  undirectedAdjList.James.indexOf('Jeff') !== -1 &&
    undirectedAdjList.Jeff.indexOf('James') !== -1
);
```

`Jill`과 `Jenny` 사이에 간선이 있어야 합니다.

```js
assert(
  undirectedAdjList.Jill.indexOf('Jenny') !== -1 &&
    undirectedAdjList.Jenny.indexOf('Jill') !== -1
);
```

`Jeff`과 `Jenny` 사이에 간선이 있어야 합니다.

```js
assert(
  undirectedAdjList.Jeff.indexOf('Jenny') !== -1 &&
    undirectedAdjList.Jenny.indexOf('Jeff') !== -1
);
```

# --seed--

## --seed-contents--

```js
var undirectedAdjList = {};
```

# --solutions--

```js
var undirectedAdjList = {
  James: ['Jeff'],
  Jill: ['Jenny'],
  Jenny: ['Jill', 'Jeff'],
  Jeff: ['James', 'Jenny']
};
```
