---
id: 587d825a367417b2b2512c8a
title: 최대 힙에 요소 삽입하기
challengeType: 1
forumTopicId: 301703
dashedName: insert-an-element-into-a-max-heap
---

# --description--

이제 또 다른 트리 자료 구조인 이진 힙으로 넘어가겠습니다. 이진 힙은 힙 속성을 만족하는 부분적으로 정렬된 이진 트리입니다. 힙 속성은 부모 노드와 자식 노드 간의 관계를 지정합니다. 모든 부모 노드가 자식 노드보다 크거나 같은 최대 힙이 있을 수 있고, 반대로 자식 노드가 부모 노드보다 크거나 같은 최소 힙도 있을 수 있습니다. 이진 힙은 또한 완전 이진 트리입니다. 이는 트리의 모든 레벨이 완전히 채워져 있고, 마지막 레벨이 부분적으로 채워져 있다면 왼쪽에서 오른쪽으로 채워진다는 뜻입니다.

이진 힙은 왼쪽과 오른쪽 참조를 포함하는 노드로 구성된 트리 구조로 구현할 수 있지만, 힙 속성에 따른 부분 정렬 덕분에 배열로 힙을 나타낼 수 있습니다. 부모-자식 관계가 핵심이며 간단한 산술 연산으로 어떤 부모의 자식과 어떤 자식의 부모도 계산할 수 있습니다.

예를 들어, 이진 최소 힙의 배열 표현을 생각해 보세요:

```js
[ 6, 22, 30, 37, 63, 48, 42, 76 ]
```

루트 노드는 첫 번째 요소인 `6`입니다. 그 자식들은 `22`과 `30`입니다. 이 값들의 배열 인덱스 관계를 보면, 인덱스 `i`의 자식은 `2 * i + 1`와 `2 * i + 2`입니다. 마찬가지로, 인덱스 `0`의 요소는 인덱스 `1`과 `2`에 있는 두 자식의 부모입니다. 더 일반적으로, 다음 식으로 어떤 인덱스의 노드 부모를 찾을 수 있습니다: `Math.floor((i - 1) / 2)`. 이 패턴은 이진 트리가 어떤 크기로 커져도 유지됩니다. 마지막으로, 배열의 첫 번째 요소를 건너뛰면 이 산술을 더 쉽게 만들 수 있습니다. 이렇게 하면 주어진 인덱스 `i`에 있는 요소에 대해 다음 관계가 성립합니다:

예시 배열 표현:

```js
[ null, 6, 22, 30, 37, 63, 48, 42, 76 ]
```

요소의 왼쪽 자식: `i * 2`

요소의 오른쪽 자식: `i * 2 + 1`

요소의 부모: `Math.floor(i / 2)`

이 산술을 이해하면 배열 표현이 매우 유용합니다. 노드 위치를 빠르게 계산할 수 있고, 자식 노드에 대한 참조를 유지할 필요가 없어 메모리 사용량도 줄어듭니다.

# --instructions--

지침: 여기서는 최대 힙을 만들겠습니다. 우선 힙에 요소를 추가하는 `insert` 메서드를 만드세요. 삽입하는 동안 항상 힙 속성을 유지하는 것이 중요합니다. 최대 힙의 경우 루트 요소가 트리에서 가장 큰 값을 가져야 하고 모든 부모 노드는 자식보다 커야 합니다. 배열로 구현된 힙에서는 보통 세 단계로 이 작업을 수행합니다:

<ol>
  <li>새 요소를 배열 끝에 추가합니다.</li>
  <li>요소가 부모보다 크면 두 요소를 교환합니다.</li>
  <li>새 요소가 부모보다 작거나 루트에 도달할 때까지 계속 교환합니다.</li>
</ol>

마지막으로 힙에 추가된 모든 항목을 배열로 반환하는 `print` 메서드를 추가하세요.

# --hints--

`MaxHeap` 자료 구조가 있어야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap`에는 `insert`라는 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap`에는 `print`라는 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

`insert` 메서드는 최대 힙 속성에 따라 요소를 추가해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    test.insert(50);
    test.insert(100);
    test.insert(700);
    test.insert(32);
    test.insert(51);
    test.insert(800);
    const result = test.print();
    const solution = JSON.stringify([null,800,51,700,32,50,100]);
    const solutionWithoutNull = JSON.stringify([800,51,700,32,50,100]);

    return (result.length == 6) ? (JSON.stringify(result) == solutionWithoutNull) : (JSON.stringify(result) == solution);
  })()
);
```

# --seed--

## --seed-contents--

```js
var MaxHeap = function() {
  // Only change code below this line
    
  // Only change code above this line
};
```

# --solutions--

```js
var MaxHeap = function() {
    // Only change code below this line
    this.heap = [];
    this.parent = index => {
      return Math.floor((index - 1) / 2);
    }
    this.insert = element => {
      this.heap.push(element);
      this.heapifyUp(this.heap.length - 1);
    }
    this.heapifyUp = index => {
      let currentIndex = index,
      parentIndex = this.parent(currentIndex);
      while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
        this.swap(currentIndex, parentIndex);
        currentIndex = parentIndex;
        parentIndex = this.parent(parentIndex);
      }
    }
    this.swap = (index1, index2) => {
      [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
    }
    this.print = () => {
      return this.heap;
    }
    // Only change code above this line
};
```
