---
id: 587d8252367417b2b2512c67
title: 연결 리스트의 특정 인덱스에 요소 추가
challengeType: 1
forumTopicId: 301619
dashedName: add-elements-at-a-specific-index-in-a-linked-list
---

# --description--

주어진 인덱스에 요소를 추가하는 addAt(index,element) 메서드를 만들어 보세요. 특정 인덱스에서 요소를 제거하는 것처럼, 연결 리스트를 순회할 때 현재 인덱스를 추적해야 합니다. 현재 인덱스가 주어진 인덱스와 일치하면, 이전 노드의 next 속성을 새로 추가된 노드를 참조하도록 다시 할당해야 합니다. 그리고 새 노드는 현재 인덱스의 다음 노드를 참조해야 합니다. 콩가 라인 예시로 돌아가 보면, 새 사람이 줄 중간에 끼어들고 싶어 합니다. 당신은 줄 중간에 있으니 앞 사람과 손을 떼고, 새 사람이 와서 당신이 손을 잡고 있던 사람과 손을 잡습니다. 이제 당신은 새 사람과 손을 잡고 있습니다.

# --instructions--

주어진 인덱스에 요소를 추가하는 `addAt(index,element)` 메서드를 만드세요. 요소를 추가할 수 없으면 false를 반환하세요. **참고:** 주어진 인덱스가 음수이거나 연결 리스트 길이보다 길면 반드시 확인하세요.

# --hints--

`addAt` 메서드는 주어진 인덱스가 0일 때 `head`를 새 노드로 다시 할당해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'fish');
    return test.head().element === 'fish' && test.head().next.element === 'cat';
  })()
);
```

`addAt` 메서드는 연결 리스트에 새 노드가 추가될 때마다 길이를 1씩 증가시켜야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'cat');
    return test.size() === 3;
  })()
);
```

`addAt` 메서드는 노드를 추가하지 못했을 때 `false`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.addAt(4, 'cat') === false;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element) {
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function() {
    return head;
  };

  this.add = function(element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if (head === null){
        head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };
  this.addAt = function (index, element) {
    if (index > length || index < 0) {
      return false;
    }
    var newNode = new Node(element);
    var currentNode = head;
    if (index === 0) {
      head = newNode;
    } else {
      var previousNode = null;
      var i = 0;
      while (currentNode && i < index) {
        previousNode = currentNode;
        currentNode = currentNode.next;
        i++;
      }
      previousNode.next = newNode;
    }
    newNode.next = currentNode;
    length++;
  }
}
```
