---
id: 587d8257367417b2b2512c7e
title: 이진 탐색 트리에서 깊이 우선 탐색 사용하기
challengeType: 1
forumTopicId: 301719
dashedName: use-depth-first-search-in-a-binary-search-tree
---

# --description--

특정 값을 찾기 위해 이진 탐색 트리를 검색하는 방법은 알고 있습니다. 하지만 전체 트리를 탐색하고 싶다면 어떻게 할까요? 또는 정렬되지 않은 트리가 있고 단순히 값을 검색해야 한다면 어떻게 할까요? 여기서는 트리 자료 구조를 탐색하는 데 사용할 수 있는 몇 가지 트리 순회 방법을 소개합니다. 먼저 깊이 우선 탐색입니다. 깊이 우선 탐색에서는 주어진 부분 트리를 가능한 깊게 탐색한 후에야 다른 부분 트리로 탐색을 계속합니다. 이 작업을 수행하는 세 가지 방법이 있습니다: 중위 순서: 가장 왼쪽 노드에서 시작해 가장 오른쪽 노드에서 끝냅니다. 전위 순서: 잎 노드보다 모든 루트 노드를 먼저 탐색합니다. 후위 순서: 루트 노드보다 모든 잎 노드를 먼저 탐색합니다. 예상할 수 있듯이, 트리가 저장하는 데이터 유형과 찾고자 하는 것에 따라 다른 탐색 방법을 선택할 수 있습니다. 이진 탐색 트리의 경우, 중위 순회는 노드를 정렬된 순서로 반환합니다.

# --instructions--

여기서는 이진 탐색 트리에서 이 세 가지 탐색 방법을 만들어 보겠습니다. 깊이 우선 탐색은 자식 노드가 존재하는 한 계속해서 더 깊은 부분 트리를 탐색하는 본질적으로 재귀적인 작업입니다. 이 기본 개념을 이해하면 노드와 부분 트리를 탐색하는 순서를 단순히 재배열하여 위 세 가지 탐색 중 어느 것이든 만들 수 있습니다. 예를 들어, 후위 순서 탐색에서는 노드 자체를 반환하기 전에 잎 노드까지 재귀적으로 내려가야 하지만, 전위 순서 탐색에서는 먼저 노드를 반환하고 나서 트리를 계속 재귀적으로 탐색합니다. 트리에 `inorder`, `preorder`, `postorder` 메서드를 정의하세요. 이 메서드들은 모두 트리 순회 결과를 나타내는 항목들의 배열을 반환해야 합니다. 배열에는 노드 자체가 아니라 각 노드의 정수 값을 반드시 반환하세요. 마지막으로, 트리가 비어 있으면 `null`을 반환하세요.

# --hints--

`BinarySearchTree` 자료 구조가 존재해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

이진 탐색 트리는 `inorder`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.inorder == 'function';
  })()
);
```

이진 탐색 트리는 `preorder`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.preorder == 'function';
  })()
);
```

이진 탐색 트리는 `postorder`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.postorder == 'function';
  })()
);
```

`inorder` 메서드는 중위 순회 결과로 나온 노드 값들의 배열을 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.inorder().join('') == '012345678910';
  })()
);
```

`preorder` 메서드는 전위 순회 결과로 나온 노드 값들의 배열을 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.preorder().join('') == '710325469810';
  })()
);
```

`postorder` 메서드는 후위 순회 결과로 나온 노드 값들의 배열을 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.postorder().join('') == '024653181097';
  })()
);
```

`inorder` 메서드는 빈 트리일 경우 `null`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    return test.inorder() == null;
  })()
);
```

`preorder` 메서드는 빈 트리일 경우 `null`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    return test.preorder() == null;
  })()
);
```

`postorder` 메서드는 빈 트리일 경우 `null`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    return test.postorder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.result = [];

  this.inorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.inorder(node.left);
    this.result.push(node.value);
    if (node.right) this.inorder(node.right);
    return this.result;
  };
  this.preorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    this.result.push(node.value);
    if (node.left) this.preorder(node.left);
    if (node.right) this.preorder(node.right);
    return this.result;
  };
  this.postorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.postorder(node.left);
    if (node.right) this.postorder(node.right);
    this.result.push(node.value);

    return this.result;
  };
}
```
