---
id: 587d8251367417b2b2512c62
title: 연결 리스트 클래스 만들기
challengeType: 1
forumTopicId: 301628
dashedName: create-a-linked-list-class
---

# --description--

`linked list` 클래스를 만들어 봅시다. 모든 연결 리스트는 몇 가지 기본 속성으로 시작해야 합니다: `head` (리스트의 첫 번째 항목)과 `length` (리스트의 항목 수)입니다. 때때로 연결 리스트 구현에서 리스트의 마지막 요소를 위한 `tail`를 포함하는 경우도 있지만, 지금은 이 두 가지 속성만 사용하겠습니다. 연결 리스트에 요소를 추가할 때마다 `length` 속성은 1씩 증가해야 합니다.

연결 리스트에 항목을 추가하는 방법이 필요하므로, 가장 먼저 만들 메서드는 `add` 메서드입니다.

리스트가 비어 있다면, 연결 리스트에 요소를 추가하는 것은 간단합니다: 그 요소를 `Node` 클래스에 감싸고, 그 노드를 연결 리스트의 `head`에 할당하면 됩니다.

하지만 리스트에 이미 하나 이상의 멤버가 있다면 어떻게 요소를 추가할까요? 연결 리스트의 각 노드가 `next` 속성을 가진다는 것을 기억하세요. 리스트에 노드를 추가하려면, 리스트의 마지막 노드를 찾아서 그 마지막 노드의 `next` 속성이 새 노드를 가리키도록 하면 됩니다. (힌트: 노드의 `next` 속성이 `null`일 때 연결 리스트의 끝에 도달한 것입니다.)

# --instructions--

add 메서드를 작성하세요. 연결 리스트에 처음 푸시하는 노드는 `head`에 할당해야 합니다. 그 이후로는 노드를 추가할 때마다 모든 노드는 이전 노드의 `next` 속성으로 참조되어야 합니다.

참고

연결 리스트에 요소가 추가될 때마다 리스트의 `length`는 1씩 증가해야 합니다.

# --hints--

`LinkedList` 클래스는 `add` 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.add === 'function';
  })()
);
```

`LinkedList` 클래스는 처음 추가된 노드에 `head`를 할당해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    return test.head().element === 'cat';
  })()
);
```

`node` 클래스의 이전 `LinkedList`는 새로 생성된 노드를 참조해야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('fish');
    return test.head().next.element === 'dog' && test.head().next.next.element === 'fish';
  })()
);
```

`size` 클래스의 `LinkedList`는 연결 리스트의 노드 수와 같아야 합니다.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.size() === 2;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() { 
  var length = 0; 
  var head = null; 

  var Node = function(element){
    this.element = element; 
    this.next = null; 
  }; 

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line
    if (head == null) {
      head = new Node(element);
    } 
    else {
      let currentNode = head;
      while (currentNode.next != null) {
        // currentNode.next will be last node of linked list after loop
        currentNode = currentNode.next;
      }
      currentNode.next = new Node(element);
    }
    length++;
    // Only change code above this line
  };
}
```
