---
id: 587d8258367417b2b2512c80
title: 이진 탐색 트리에서 리프 노드 삭제하기
challengeType: 1
forumTopicId: 301637
dashedName: delete-a-leaf-node-in-a-binary-search-tree
---

# --description--

이 도전 과제는 이진 탐색 트리에서 더 어려운 연산인 삭제를 구현하는 세 가지 도전 과제 중 첫 번째입니다. 삭제가 어려운 이유는 노드를 제거하면 트리의 연결이 끊어지기 때문입니다. 이 연결을 신중하게 다시 설정해야 이진 트리 구조가 유지됩니다. 일부 삭제의 경우 트리를 재배열해야 합니다. 일반적으로 노드를 삭제하려 할 때 세 가지 경우 중 하나를 만나게 됩니다: 리프 노드: 삭제 대상 노드가 자식 노드가 없는 경우. 자식 노드 하나: 삭제 대상 노드가 자식 노드가 하나인 경우. 자식 노드 둘: 삭제 대상 노드가 두 개의 자식 노드를 가진 경우. 리프 노드를 제거하는 것은 간단하며, 단순히 노드를 삭제하면 됩니다. 자식 노드가 하나인 노드를 삭제하는 것도 비교적 쉬우며, 삭제한 노드의 부모를 그 자식 노드와 연결하면 됩니다. 그러나 자식 노드가 두 개인 노드를 삭제하는 것은 더 어렵습니다. 이 경우 두 자식 노드를 부모 트리에 다시 연결해야 하기 때문입니다. 이 경우는 세 번째 도전 과제에서 다루겠습니다. 또한 삭제를 처리할 때 몇 가지 예외 상황도 주의해야 합니다. 트리가 비어 있다면? 삭제할 노드가 루트 노드라면? 트리에 요소가 두 개뿐이라면? 지금은 리프 노드를 삭제하는 첫 번째 경우를 처리해 봅시다.

# --instructions--

이진 트리에 `remove`라는 메서드를 만드세요. 여기서 삭제 연산의 로직을 작성할 것입니다. 먼저, remove 내부에 현재 트리에서 삭제하려는 노드를 찾는 함수를 만드세요. 노드가 트리에 없으면 `remove`은 `null`를 반환해야 합니다. 이제 삭제 대상 노드가 자식 노드가 없는 리프 노드라면, 부모의 참조를 `null`로 설정해야 합니다. 이렇게 하면 노드가 트리에서 효과적으로 삭제됩니다. 이를 위해 삭제하려는 노드의 부모도 추적해야 합니다. 또한 삭제 대상 노드가 가진 자식 노드 수를 추적하는 방법을 만드는 것이 유용할 것입니다. 이 값에 따라 삭제가 어떤 경우에 해당하는지 결정할 수 있기 때문입니다. 두 번째와 세 번째 경우는 다음 도전 과제에서 다루겠습니다. 행운을 빕니다!

# --hints--

`BinarySearchTree` 자료 구조가 존재해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

이진 탐색 트리에 `remove`라는 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

빈 트리에서 요소를 삭제하려 하면 `null`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    return test.remove(100) == null;
  })()
);
```

존재하지 않는 요소를 삭제하려 하면 `null`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(15);
    test.add(30);
    return test.remove(100) == null;
  })()
);
```

루트 노드에 자식 노드가 없으면 삭제 시 루트를 `null`로 설정해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(500);
    test.remove(500);
    return test.inorder() == null;
  })()
);
```

`remove` 메서드는 트리에서 리프 노드를 삭제해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(5);
    test.add(3);
    test.add(7);
    test.add(6);
    test.add(10);
    test.add(12);
    test.remove(3);
    test.remove(12);
    test.remove(10);
    return test.inorder().join('') == '567';
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    },
    inorder: function() {
      if (this.root == null) {
        return null;
      } else {
        var result = new Array();
        function traverseInOrder(node) {
          if (node.left != null) {
            traverseInOrder(node.left);
          }
          result.push(node.value);
          if (node.right != null) {
            traverseInOrder(node.right);
          }
        }
        traverseInOrder(this.root);
        return result;
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}

function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
}
```

# --solutions--

```js
// solution required
```
