---
id: 587d8255367417b2b2512c75
title: 원형 큐 만들기
challengeType: 1
forumTopicId: 301625
dashedName: create-a-circular-queue
---

# --description--

이번 도전 과제에서는 원형 큐를 만들게 됩니다. 원형 큐는 컬렉션의 끝에 데이터를 쓰고 나서 컬렉션의 시작 부분부터 덮어쓰기를 시작하는 큐입니다. 이런 유형의 자료 구조는 특정 상황에서 유용합니다. 예를 들어, 원형 큐는 스트리밍 미디어에 사용할 수 있습니다. 큐가 가득 차면 새 미디어 데이터가 오래된 데이터를 덮어씁니다.

이 개념을 설명하는 좋은 방법은 길이가 `5`인 배열을 사용하는 것입니다:

```js
[null, null, null, null, null]
 ^Read @ 0
 ^Write @ 0
```

여기서 읽기와 쓰기 위치가 모두 `0`에 있습니다. 이제 큐에 `a`, `b`, `c` 세 개의 새 레코드가 들어옵니다. 우리의 큐는 이제 다음과 같습니다:

```js
[a, b, c, null, null]
 ^Read @ 0
          ^Write @ 3
```

읽기 헤드가 읽으면서 값을 제거하거나 유지할 수 있습니다:

```js
[null, null, null, null, null]
                   ^Read @ 3
                   ^Write @ 3
```

이제 큐에 `d`, `e`, `f` 값을 씁니다. 쓰기가 배열 끝에 도달하면 다시 시작 부분으로 돌아갑니다:

```js
[f, null, null, d, e]
                ^Read @ 3
    ^Write @ 1
```

이 방법은 일정한 메모리 양만 필요하지만 훨씬 큰 크기의 파일을 처리할 수 있게 해줍니다.

# --instructions--

이번 도전 과제에서는 원형 큐를 구현할 것입니다. 원형 큐는 큐에서 읽고 쓸 수 있도록 `enqueue`와 `dequeue` 메서드를 제공해야 합니다. 클래스 자체는 생성 시 큐 크기를 지정할 수 있는 정수 인수를 받아야 합니다. 코드 편집기에 이 클래스의 시작 버전을 작성해 두었습니다.

큐에 항목을 넣을 때, 쓰기 포인터는 앞으로 이동하며 큐 끝에 도달하면 다시 시작 부분으로 돌아가야 합니다. `enqueue` 메서드는 성공하면 넣은 항목을 반환하고, 그렇지 않으면 `null`을 반환해야 합니다.

마찬가지로, 읽기 포인터는 항목을 꺼낼 때 앞으로 이동해야 합니다. 항목을 꺼내면 그 항목을 반환해야 합니다. 꺼낼 항목이 없으면 `null`를 반환해야 합니다.

쓰기 포인터는 읽기 포인터를 넘어가지 못하게 해야 합니다(아직 읽지 않은 데이터를 덮어쓰지 않도록 클래스가 막아줍니다). 읽기 포인터도 쓴 데이터를 넘어 앞으로 나아갈 수 없습니다.

# --hints--

`enqueue` 메서드는 원형 큐에 항목을 추가해야 합니다.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

읽기 포인터를 넘어 항목을 넣으면 안 됩니다.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    test.enqueue(13);
    test.enqueue(25);
    test.enqueue(59);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

`dequeue` 메서드는 큐에서 항목을 꺼내야 합니다.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 && test.dequeue() === 32 && test.dequeue() === 591
    );
  })()
);
```

항목을 꺼낸 후에는 그 위치를 `null`로 초기화해야 합니다.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(672);
    test.dequeue();
    test.dequeue();
    var print = test.print();
    return print[0] === null && print[1] === null && print[2] === 672;
  })()
);
```

쓰기 포인터를 넘어 꺼내려고 하면 `null`을 반환하고 쓰기 포인터는 이동하지 않습니다.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 &&
      test.dequeue() === 32 &&
      test.dequeue() === 591 &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.enqueue(100) === 100 &&
      test.dequeue() === 100
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
class CircularQueue {
  constructor(size) {

    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line

    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line

    // Only change code above this line
  }
}
```

# --solutions--

```js
class CircularQueue {
  constructor(size) {
    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line
    console.log(this.write, this.max);
    if (this.queue[this.write] === null) {
      this.queue[this.write++] = item;

      if (this.write > this.max) {
        this.write = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line
    if (this.queue[this.read] !== null) {
      let item = this.queue[this.read];
      this.queue[this.read++] = null;
      if (this.read > this.max) {
        this.read = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }
}
```
