---
id: 587d8257367417b2b2512c7b
title: 이진 탐색 트리에 새 요소 추가
challengeType: 1
forumTopicId: 301618
dashedName: add-a-new-element-to-a-binary-search-tree
---

# --description--

이 도전 과제 시리즈에서는 트리 자료 구조를 소개합니다. 트리는 컴퓨터 과학에서 중요하고 다재다능한 자료 구조입니다. 물론 이름은 시각화했을 때 자연 세계에서 익숙한 나무와 매우 비슷하게 생겼기 때문에 붙여졌습니다. 트리 자료 구조는 일반적으로 루트라고 부르는 하나의 노드에서 시작해 추가 노드로 가지가 뻗어나가며, 각 노드는 더 많은 자식 노드를 가질 수 있습니다. 이 자료 구조는 보통 루트 노드를 위쪽에 두고 시각화하는데, 자연 나무를 거꾸로 뒤집은 것처럼 생각할 수 있습니다.

먼저 트리에서 자주 접할 용어를 설명하겠습니다. 루트 노드는 트리의 최상단입니다. 트리 내 데이터 포인트를 노드라고 합니다. 다른 노드로 가지가 뻗어 있는 노드는 그 가지가 연결된 노드(자식)의 부모라고 부릅니다. 예상할 수 있듯이 더 복잡한 가족 관계 용어도 적용됩니다. 부분 트리는 특정 노드의 모든 자손 요소를 의미하며, 가지는 엣지라고도 부를 수 있고, 리프 노드는 자식이 없는 트리의 끝 노드입니다. 마지막으로, 트리는 본질적으로 재귀적인 자료 구조임을 기억하세요. 즉, 노드의 자식은 자신의 부분 트리의 부모가 되고, 이런 식으로 계속됩니다. 트리의 재귀적 특성은 일반적인 트리 연산 알고리즘을 설계할 때 이해하는 것이 중요합니다.

우선, 특정 유형의 트리인 이진 트리에 대해 이야기하겠습니다. 사실, 특정 이진 트리인 이진 탐색 트리에 대해 다룰 것입니다. 이게 무슨 뜻인지 설명하겠습니다. 트리 자료 구조는 한 노드에 여러 가지 가지를 가질 수 있지만, 이진 트리는 각 노드마다 두 개의 가지만 가질 수 있습니다. 더 나아가, 이진 탐색 트리는 자식 부분 트리와 관련해 정렬되어 있어서, 왼쪽 부분 트리의 각 노드 값은 부모 노드 값보다 작거나 같고, 오른쪽 부분 트리의 각 노드 값은 부모 노드 값보다 크거나 같다는 규칙을 가집니다. 이 관계를 시각화하면 이해하기 훨씬 쉽습니다:

<div style='width: 100%; display: flex; justify-content: center; align-items: center;'><img alt="an example of a binary search tree" style='width: 100%; max-width: 350px; background-color: var(--gray-05);' src='https://user-images.githubusercontent.com/18563015/32136009-1e665d98-bbd6-11e7-9133-63184f9f9182.png'></div>

이제 이 정렬 관계가 매우 명확해졌습니다. 루트 노드인 8의 왼쪽에 있는 모든 값은 8보다 작고, 오른쪽에 있는 모든 값은 8보다 큽니다. 또한 이 관계는 각 부분 트리에도 적용됩니다. 예를 들어, 첫 번째 왼쪽 자식은 부분 트리입니다. 3이 부모 노드이고 정확히 두 개의 자식 노드가 있습니다 — 이진 탐색 트리 규칙에 따라 이 노드의 왼쪽 자식(및 그 자식들)은 3보다 작고, 오른쪽 자식(및 그 자식들)은 3보다 크지만 트리의 루트 값보다는 작다는 것을 알 수 있습니다.

이진 탐색 트리는 조회, 삽입, 삭제 같은 여러 일반적인 연산에서 평균적으로 로그 시간 복잡도를 제공하기 때문에 매우 흔하고 유용한 자료 구조입니다.

# --instructions--

간단한 것부터 시작하겠습니다. 여기 이진 탐색 트리 구조의 뼈대와 트리용 노드를 생성하는 함수를 정의했습니다. 각 노드는 왼쪽과 오른쪽 값을 가질 수 있습니다. 존재한다면 이 값들은 자식 부분 트리로 할당됩니다. 이진 탐색 트리에서 새 값을 트리에 추가하는 메서드를 작성하세요. 메서드 이름은 `add`이고, 트리에 추가할 정수 값을 인자로 받아야 합니다. 이진 탐색 트리의 불변 조건을 유지하도록 주의하세요: 왼쪽 자식의 값은 부모 값보다 작거나 같아야 하고, 오른쪽 자식의 값은 부모 값보다 크거나 같아야 합니다. 여기서는 트리에 중복 값을 허용하지 않도록 하겠습니다. 이미 존재하는 값을 추가하려 하면 메서드는 `null`을 반환해야 합니다. 그렇지 않고 추가가 성공하면 `undefined`를 반환하세요.

**힌트:** 트리는 본질적으로 재귀적인 자료 구조입니다!

# --hints--

`BinarySearchTree` 자료 구조가 존재해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

이진 탐색 트리는 `add`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

추가 메서드는 이진 탐색 트리 규칙에 따라 요소를 추가해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    const expectedResult = [1, 4, 7, 8, 34, 45, 73, 87];
    const result = test.inOrder();
    return expectedResult.toString() === result.toString();
  })()
);
```

이미 존재하는 요소를 추가하면 `null`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    return test.add(4) == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    inOrder() {
      if (!this.root) {
        return null;
      }
      var result = new Array();
      function traverseInOrder(node) {
        node.left && traverseInOrder(node.left);
        result.push(node.value);
        node.right && traverseInOrder(node.right);
      }
      traverseInOrder(this.root);
      return result;
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.add = function(element) {
    let current = this.root;
    if (!current) {
      this.root = new Node(element);
      return;
    } else {
      const searchTree = function(current) {
        if (current.value > element) {
          if (current.left) {
            return searchTree(current.left);
          } else {
            current.left = new Node(element);
            return;
          }
        } else if (current.value < element) {
          if (current.right) {
            return searchTree(current.right);
          } else {
            current.right = new Node(element);
            return;
          }
        } else {
          return null;
        }
      };
      return searchTree(current);
    }
  };
}
```
