---
id: 587d825d367417b2b2512c96
title: 깊이 우선 탐색
challengeType: 1
forumTopicId: 301640
dashedName: depth-first-search
---

# --description--

<dfn>너비 우선 탐색</dfn>과 비슷하게, 여기서는 <dfn>깊이 우선 탐색</dfn>이라는 또 다른 그래프 순회 알고리즘에 대해 배웁니다.

너비 우선 탐색이 출발 노드에서 점차적으로 간선 길이를 늘려가며 검색하는 반면, <dfn>깊이 우선 탐색</dfn>은 먼저 한 경로를 가능한 한 깊게 따라 내려갑니다.

경로의 한 끝에 도달하면, 탐색은 방문하지 않은 간선 경로가 있는 마지막 노드로 되돌아가 계속 검색합니다.

아래 애니메이션은 알고리즘이 어떻게 작동하는지 보여줍니다. 알고리즘은 맨 위 노드에서 시작해 번호가 매겨진 순서대로 노드를 방문합니다.

<img alt="animation showing the depth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/depth-first-search.gif'>

너비 우선 탐색과 달리, 노드를 방문할 때마다 모든 이웃을 방문하지 않는 점에 주목하세요. 대신, 먼저 이웃 중 하나를 방문하고 그 경로를 따라 더 이상 방문할 노드가 없을 때까지 계속 내려갑니다.

이 알고리즘을 구현하려면 스택을 사용하세요. 스택은 마지막에 추가된 요소가 가장 먼저 제거되는 배열입니다. 이것은 <dfn>후입선출</dfn> 자료 구조로도 알려져 있습니다. 스택은 깊이 우선 탐색 알고리즘에서 유용한데, 이웃을 스택에 추가할 때 가장 최근에 추가된 이웃을 먼저 방문하고 스택에서 제거하기 때문입니다.

이 알고리즘의 간단한 출력은 주어진 노드에서 도달할 수 있는 노드들의 목록입니다. 따라서 방문한 노드를 추적하는 것도 필요합니다.

# --instructions--

무방향 인접 행렬 `dfs()`과 노드 레이블 `graph`를 매개변수로 받는 함수 `root`를 작성하세요. 노드 레이블은 `0`와 `n - 1` 사이의 노드 숫자 값이며, `n`는 그래프의 총 노드 수입니다.

함수는 `root`에서 도달할 수 있는 모든 노드를 배열로 출력해야 합니다.

# --hints--

입력 그래프 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`와 시작 노드 `1`는 `0`, `1`, `2`, `3`를 포함하는 배열을 반환해야 합니다.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })(),
  [0, 1, 2, 3]
);
```

입력 그래프 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`와 시작 노드 `3`는 `3`, `2`, `1`, `0`를 포함하는 배열을 반환해야 합니다.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3, 2, 1, 0]
);
```

입력 그래프 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]`와 시작 노드 `1`는 네 개의 요소를 포함하는 배열을 반환해야 합니다.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })().length === 4
);
```

입력 그래프 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]`와 시작 노드 `3`는 `3`를 포함하는 배열을 반환해야 합니다.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3]
);
```

입력 그래프 `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]`와 시작 노드 `3`는 한 개의 요소를 포함하는 배열을 반환해야 합니다.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })().length === 1
);
```

입력 그래프 `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]`와 시작 노드 `3`는 `2`와 `3`를 포함하는 배열을 반환해야 합니다.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [2, 3]
);
```

입력 그래프 `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]`와 시작 노드 `3`는 두 개의 요소를 포함하는 배열을 반환해야 합니다.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })().length === 2
);
```

입력 그래프 `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]`와 시작 노드 `0`는 `0`와 `1`를 포함하는 배열을 반환해야 합니다.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })(),
  [0, 1]
);
```

입력 그래프 `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]`와 시작 노드 `0`는 두 개의 요소를 포함하는 배열을 반환해야 합니다.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })().length === 2
);
```

# --seed--

## --seed-contents--

```js
function dfs(graph, root) {

}

var exDFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(dfs(exDFSGraph, 3));
```

# --solutions--

```js
function dfs(graph, root) {
    var stack = [];
    var tempV;
    var visited = [];
    var tempVNeighbors = [];
    stack.push(root);
    while (stack.length > 0) {
        tempV = stack.pop();
        if (visited.indexOf(tempV) == -1) {
            visited.push(tempV);
            tempVNeighbors = graph[tempV];
            for (var i = 0; i < tempVNeighbors.length; i++) {
                if (tempVNeighbors[i] == 1) {
                    stack.push(i);
                }
            }
        }
    }
    return visited;
}
```
