---
id: 587d8250367417b2b2512c60
title: 큐 클래스 만들기
challengeType: 1
forumTopicId: 301631
dashedName: create-a-queue-class
---

# --description--

스택과 마찬가지로, 큐도 요소들의 모음입니다. 하지만 스택과 달리 큐는 FIFO(선입선출) 원칙을 따릅니다. 큐에 추가된 요소는 큐의 꼬리, 즉 끝에 추가되고, 큐의 앞에 있는 요소만 제거할 수 있습니다.

큐를 나타내기 위해 배열을 사용할 수 있지만, 스택과 마찬가지로 큐에 대한 제어를 제한하고 싶습니다.

큐 클래스의 두 가지 주요 메서드는 enqueue와 dequeue 메서드입니다. enqueue 메서드는 요소를 큐의 꼬리에 추가하고, dequeue 메서드는 큐의 앞에 있는 요소를 제거하고 반환합니다. 그 밖에 유용한 메서드로는 front, size, isEmpty 메서드가 있습니다.

# --instructions--

큐의 꼬리에 요소를 추가하는 `enqueue` 메서드, 앞 요소를 제거하고 반환하는 `dequeue` 메서드, 앞 요소를 확인할 수 있는 `front` 메서드, 길이를 보여주는 `size` 메서드, 그리고 큐가 비어 있는지 확인하는 `isEmpty` 메서드를 작성하세요.

# --hints--

`Queue` 클래스는 `enqueue` 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.enqueue === 'function';
  })()
);
```

`Queue` 클래스는 `dequeue` 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.dequeue === 'function';
  })()
);
```

`Queue` 클래스는 `front` 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.front === 'function';
  })()
);
```

`Queue` 클래스는 `size` 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.size === 'function';
  })()
);
```

`Queue` 클래스는 `isEmpty` 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.isEmpty === 'function';
  })()
);
```

`dequeue` 메서드는 큐의 앞 요소를 제거하고 반환해야 합니다.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.dequeue() === 'Smith';
  })()
);
```

`front` 메서드는 큐의 앞 요소 값을 반환해야 합니다.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.front() === 'Smith';
  })()
);
```

`size` 메서드는 큐의 길이를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return test.size() === 1;
  })()
);
```

`isEmpty` 메서드는 큐에 요소가 있으면 `false`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return !test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Queue() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function Queue () { 
    var collection = [];
    this.print = function() {
        console.log(collection);
    };
    // Only change code below this line
    this.enqueue = function(item) {
        collection.push(item);
    }

    this.dequeue = function() {
        return collection.shift();
    }

    this.front = function() {
        return collection[0];
    }

    this.size = function(){
        return collection.length;
    }

    this.isEmpty = function() {
        return collection.length === 0 ? true : false;
    }
    // Only change code above this line
}
```
