---
id: 587d8250367417b2b2512c5e
title: 스택 작동 방식 배우기
challengeType: 1
forumTopicId: 301705
dashedName: learn-how-a-stack-works
---

# --description--

책상 위에 쌓여 있는 책 더미를 아마 알고 있을 거예요. 텍스트 편집기의 실행 취소 기능을 사용해본 적도 있을 겁니다. 그리고 휴대폰에서 이전 화면으로 돌아가기 위해 뒤로 가기 버튼을 누르는 것도 익숙할 거예요.

이 모든 것에 공통점이 뭔지 아세요? 모두 데이터를 뒤로 거슬러 올라갈 수 있는 방식으로 저장한다는 점입니다.

스택의 가장 위에 있는 책은 가장 나중에 올려놓은 책입니다. 만약 스택 맨 위에서 그 책을 제거하면, 그 전에 올려놓은 책이 드러나고 이런 식으로 계속됩니다.

생각해보면, 위 모든 예시에서 <dfn>후입선출(Last-In-First-Out)</dfn> 방식으로 동작하고 있습니다. 이걸 코드로 흉내 내보겠습니다.

이 데이터 저장 방식은 <dfn>스택</dfn>이라고 합니다. 특히, 스택 맨 위에 JavaScript 객체를 추가하는 `push()` 메서드와 현재 스택 맨 위에 있는 JavaScript 객체를 제거하는 `pop()` 메서드를 구현해야 합니다.

# --instructions--

여기 배열로 나타낸 숙제 스택이 있습니다: `"BIO12"`가 바닥에 있고 `"PSY44"`가 스택 맨 위에 있습니다.

주어진 배열을 수정해서 위에서 언급한 JavaScript 메서드를 사용해 `stack`처럼 다루세요. 스택에서 맨 위 요소인 `"PSY44"`를 제거한 다음, `"CS50"`를 새 스택 맨 위 요소로 추가하세요.

# --hints--

`homeworkStack`에는 4개의 요소만 있어야 합니다.

```js
assert(homeworkStack.length === 4);
```

`homeworkStack`의 마지막 요소는 `"CS50"`여야 합니다.

```js
assert(homeworkStack[3] === 'CS50');
```

`homeworkStack`에는 `"PSY44"`가 없어야 합니다.

```js
assert(homeworkStack.indexOf('PSY44') === -1);
```

`homeworkStack`의 초기 선언은 변경하지 마세요.

```js
assert(
  __helpers.removeJSComments(code).match(/=/g).length === 1 &&
    /homeworkStack\s*=\s*\["BIO12"\s*,\s*"HIS80"\s*,\s*"MAT122"\s*,\s*"PSY44"\]/.test(
      __helpers.removeJSComments(code)
    )
);
```

# --seed--

## --seed-contents--

```js
var homeworkStack = ["BIO12","HIS80","MAT122","PSY44"];
// Only change code below this line
```

# --solutions--

```js
// solution required
```
