---
id: 587d8256367417b2b2512c79
title: 접속 행렬
challengeType: 1
forumTopicId: 301644
dashedName: incidence-matrix
---

# --description--

그래프를 나타내는 또 다른 방법은 <dfn>접속 행렬</dfn>로 표현하는 것입니다.

<dfn>접속 행렬</dfn>은 2차원 배열입니다. 일반적으로 접속 행렬은 두 차원 사이에 서로 다른 두 종류의 객체를 연결합니다. 이 행렬은 인접 행렬과 비슷하지만, 여기서 행과 열이 의미하는 바가 다릅니다.

그래프에서는 간선과 노드가 있습니다. 이 둘이 "서로 다른 두 종류의 객체"가 됩니다. 이 행렬은 행이 노드, 열이 간선이 됩니다. 따라서 행과 열의 개수가 같지 않을 수 있습니다.

각 열은 고유한 간선을 나타냅니다. 또한 각 간선은 두 노드를 연결합니다. 두 노드 사이에 간선이 있음을 나타내려면 특정 열의 두 행에 1을 넣으면 됩니다. 아래는 노드 1과 노드 3 사이에 간선이 하나 있는 3개 노드 그래프입니다.

<blockquote>    1<br>   ---<br>1 | 1<br>2 | 0<br>3 | 1</blockquote>

아래는 4개의 간선과 4개의 노드를 가진 접속 행렬 예시입니다. 열은 간선이고 행은 노드입니다.

<blockquote>    1 2 3 4<br>   --------<br>1 | 0 1 1 1<br>2 | 1 1 0 0<br>3 | 1 0 0 1<br>4 | 0 0 1 0</blockquote>

아래는 같은 내용을 JavaScript로 구현한 예입니다.

```js
var incMat = [
  [0, 1, 1, 1],
  [1, 1, 0, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
];
```

방향 그래프를 만들려면 특정 노드에서 나가는 간선에 `-1`를, 노드로 들어오는 간선에 `1`를 사용하세요.

```js
var incMatDirected = [
  [ 0, -1,  1, -1],
  [-1,  1,  0,  0],
  [ 1,  0,  0,  1],
  [ 0,  0, -1,  0]
];
```

그래프의 간선에는 <dfn>가중치</dfn>가 있을 수 있습니다. 지금까지는 간선의 존재 여부만 이진적으로 나타내는 <dfn>가중치 없는</dfn> 간선(`0` 또는 `1`)을 다뤘습니다. 애플리케이션에 따라 다양한 가중치를 가질 수 있으며, 1보다 큰 숫자로 표현합니다.

# --instructions--

5개의 노드와 4개의 간선을 가진 무방향 그래프의 접속 행렬을 만드세요. 이 행렬은 다차원 배열이어야 합니다.

이 5개 노드는 다음과 같은 관계를 가집니다. 첫 번째 간선은 첫 번째 노드와 두 번째 노드 사이입니다. 두 번째 간선은 두 번째 노드와 세 번째 노드 사이입니다. 세 번째 간선은 세 번째 노드와 다섯 번째 노드 사이입니다. 네 번째 간선은 네 번째 노드와 두 번째 노드 사이입니다. 모든 간선 가중치는 1이며 간선 순서가 중요합니다.

# --hints--

`incMatUndirected`에는 노드가 다섯 개만 있어야 합니다.

```js
assert(
  incMatUndirected.length === 5 &&
    incMatUndirected
      .map(function (x) {
        return x.length === 4;
      })
      .reduce(function (a, b) {
        return a && b;
      })
);
```

첫 번째 간선은 첫 번째 노드와 두 번째 노드 사이여야 합니다.

```js
assert(incMatUndirected[0][0] === 1 && incMatUndirected[1][0] === 1);
```

두 번째 간선은 두 번째 노드와 세 번째 노드 사이여야 합니다.

```js
assert(incMatUndirected[1][1] === 1 && incMatUndirected[2][1] === 1);
```

세 번째 간선은 세 번째 노드와 다섯 번째 노드 사이여야 합니다.

```js
assert(incMatUndirected[2][2] === 1 && incMatUndirected[4][2] === 1);
```

네 번째 간선은 두 번째 노드와 네 번째 노드 사이여야 합니다.

```js
assert(incMatUndirected[1][3] === 1 && incMatUndirected[3][3] === 1);
```

# --seed--

## --seed-contents--

```js
var incMatUndirected = [

];
```

# --solutions--

```js
var incMatUndirected = [[1, 0, 0, 0],[1, 1, 0, 1],[0, 1, 1, 0],[0, 0, 0, 1],[0, 0, 1, 0]];
```
