---
id: 587d8259367417b2b2512c84
title: 트라이 검색 트리 만들기
challengeType: 1
forumTopicId: 301634
dashedName: create-a-trie-search-tree
---

# --description--

이번에는 이진 탐색 트리에서 벗어나 트라이라고 하는 또 다른 유형의 트리 자료 구조를 살펴보겠습니다. 트라이는 문자열을 저장하는 데 흔히 사용되는 정렬된 검색 트리로, 더 일반적으로는 키가 문자열인 연관 배열이나 동적 데이터셋을 저장하는 데 쓰입니다. 많은 키가 겹치는 접두사를 가질 때 데이터를 효율적으로 저장하는 데 매우 적합합니다. 예를 들어 사전의 모든 단어가 그렇습니다. 이진 트리와 달리 노드는 실제 값과 연결되지 않습니다. 대신 노드까지의 경로가 특정 키를 나타냅니다. 예를 들어 문자열 code를 트라이에 저장하려면 각 글자마다 하나씩 총 네 개의 노드가 필요합니다: c — o — d — e. 이 노드들을 따라가면 문자열 code가 만들어지며, 이 경로가 우리가 저장한 키입니다. 만약 coding이라는 문자열을 추가한다면 code의 처음 세 노드를 공유하다가 d 이후에서 갈라지게 됩니다. 이렇게 하면 큰 데이터셋도 매우 압축해서 저장할 수 있습니다. 또한 검색도 저장하는 문자열 길이로 제한되기 때문에 매우 빠릅니다. 게다가 이진 트리와 달리 노드는 여러 개의 자식 노드를 가질 수 있습니다. 위 예시에서 짐작할 수 있듯이, 키의 끝을 나타내는 노드에는 메타데이터가 흔히 저장되어 나중에 탐색할 때 해당 키를 다시 찾을 수 있게 합니다. 예를 들어 위 예시에서 codes를 추가한다면 code의 e 노드가 이전에 입력된 키의 끝임을 알 수 있는 방법이 필요합니다. 그렇지 않으면 codes를 추가할 때 이 정보가 사실상 사라지게 됩니다.

# --instructions--

단어를 저장하는 트라이를 만들어 보겠습니다. `add` 메서드를 통해 단어를 받아 트라이 자료 구조에 저장합니다. `isWord` 메서드를 통해 주어진 문자열이 단어인지 쿼리할 수 있고, `print` 메서드를 통해 트라이에 입력된 모든 단어를 가져올 수 있습니다. `isWord`는 불리언 값을 반환해야 하고 `print`는 이 모든 단어를 문자열 값으로 담은 배열을 반환해야 합니다. 이 자료 구조가 올바르게 구현되었는지 확인하기 위해 트리의 각 노드에 `Node` 구조를 제공했습니다. 각 노드는 `keys` 속성을 가진 객체이며, 이 속성은 JavaScript Map 객체입니다. 이 Map은 각 노드의 유효한 키인 개별 글자를 저장합니다. 또한 노드가 단어의 끝을 나타내면 `end` 속성을 `true`로 설정할 수 있도록 만들었습니다.

# --hints--

`Trie`는 `add` 메서드를 가져야 합니다.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

`Trie`는 `print` 메서드를 가져야 합니다.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

`Trie`는 `isWord` 메서드를 가져야 합니다.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.isWord == 'function';
  })()
);
```

`print` 메서드는 트라이에 추가된 모든 항목을 문자열 배열로 반환해야 합니다.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    test.add('jump');
    test.add('jumps');
    test.add('jumped');
    test.add('house');
    test.add('mouse');
    var added = test.print();
    return (
      added.indexOf('jump') != -1 &&
      added.indexOf('jumps') != -1 &&
      added.indexOf('jumped') != -1 &&
      added.indexOf('house') != -1 &&
      added.indexOf('mouse') != -1 &&
      added.length == 5
    );
  })()
);
```

`isWord` 메서드는 트라이에 추가된 단어에 대해서만 `true`를 반환하고, 그 외 단어에 대해서는 `false`를 반환해야 합니다.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    test.add('hop');
    test.add('hops');
    test.add('hopped');
    test.add('hoppy');
    test.add('hope');
    return (
      test.isWord('hop') &&
      !test.isWord('ho') &&
      test.isWord('hopped') &&
      !test.isWord('hopp') &&
      test.isWord('hoppy') &&
      !test.isWord('hoping')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
var Node = function() {
  this.keys = new Map();
  this.end = false;
  this.setEnd = function() {
    this.end = true;
  };
  this.isEnd = function() {
    return this.end;
  };
};
var Trie = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
