---
id: 587d8257367417b2b2512c7c
title: 이진 탐색 트리에 요소가 존재하는지 확인하기
challengeType: 1
forumTopicId: 301623
dashedName: check-if-an-element-is-present-in-a-binary-search-tree
---

# --description--

이제 이진 탐색 트리가 무엇인지 대략 알았으니 좀 더 자세히 살펴보겠습니다. 이진 탐색 트리는 평균적으로 조회, 삽입, 삭제 같은 기본 작업에 대해 로그 시간 복잡도를 제공합니다. 최악의 경우에는 선형 시간이 걸립니다. 왜 그럴까요? 이 기본 작업들은 모두 트리에서 항목을 찾는 작업(삽입의 경우 어디에 넣을지 찾는 작업)을 필요로 하는데, 트리 구조 덕분에 각 부모 노드에서 왼쪽이나 오른쪽으로 분기하면서 남은 트리의 절반을 효과적으로 제외합니다. 이 때문에 탐색 시간은 트리 노드 수의 로그에 비례하게 되어 평균적으로 로그 시간이 걸립니다. 그렇다면 최악의 경우는 어떨까요? 다음 값들을 왼쪽에서 오른쪽으로 차례대로 추가하여 트리를 만든다고 가정해 보세요: `10`, `12`, `17`, `25`. 이진 탐색 트리 규칙에 따라 `12`는 `10`의 오른쪽에, `17`는 그 오른쪽에, `25`는 그 오른쪽에 추가됩니다. 이제 이 트리는 연결 리스트처럼 되어 `25`를 찾으려면 모든 항목을 선형으로 순회해야 합니다. 그래서 최악의 경우 선형 시간이 걸립니다. 문제는 이 트리가 균형이 맞지 않는다는 점입니다. 다음 도전 과제에서 이 점을 좀 더 살펴보겠습니다.

# --instructions--

이번 도전 과제에서는 트리를 위한 유틸리티를 만들겠습니다. 정수 값을 입력으로 받아 이진 탐색 트리에 그 값이 존재하는지 여부를 불리언 값으로 반환하는 `isPresent` 메서드를 작성하세요.

# --hints--

`BinarySearchTree` 자료 구조가 존재해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

이진 탐색 트리에 `isPresent`라는 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isPresent == 'function';
  })()
);
```

`isPresent` 메서드는 트리에 추가된 요소의 존재 여부를 정확히 확인해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    test.add(4);
    test.add(7);
    test.add(411);
    test.add(452);
    return (
      test.isPresent(452) &&
      test.isPresent(411) &&
      test.isPresent(7) &&
      !test.isPresent(100)
    );
  })()
);
```

`isPresent`는 트리가 비어 있는 경우도 처리해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    return test.isPresent(5) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.isPresent = function (value) {
    var current = this.root
    while (current) {
      if (value === current.value) {
        return true;
      }
      current = value < current.value ? current.left : current.right;
    }
    return false;
  }
}
```
