---
id: 587d825b367417b2b2512c8e
title: 해시 테이블 만들기
challengeType: 1
forumTopicId: 301627
dashedName: create-a-hash-table
---

# --description--

이번 도전 과제에서는 해시 테이블에 대해 배웁니다. 해시 테이블은 키-값 쌍의 연관 배열 또는 매핑을 구현하는 데 사용되며, 우리가 방금 공부한 객체와 Map과 비슷합니다. 예를 들어 JavaScript 객체는 해시 테이블로 구현될 수 있습니다(실제 구현은 실행 환경에 따라 다릅니다). 해시 테이블의 작동 방식은 키 입력을 받아 이 키를 결정론적으로 해싱하여 숫자 값으로 변환하는 것입니다. 이 숫자 값이 실제 키로 사용되어 연관된 값이 저장됩니다. 그런 다음 동일한 키에 다시 접근하려 하면 해싱 함수가 키를 처리하여 같은 숫자 결과를 반환하고, 이 숫자를 사용해 연관된 값을 조회합니다. 이로 인해 평균적으로 매우 효율적인 O(1) 조회 시간이 보장됩니다.

해시 테이블은 해시 함수가 지정된 범위 내에서 배열 인덱스를 생성하는 배열로 구현할 수 있습니다. 이 방법에서는 배열 크기 선택과 해싱 함수가 중요합니다. 예를 들어, 해싱 함수가 서로 다른 두 키에 대해 같은 값을 생성하면 어떻게 될까요? 이를 충돌이라고 합니다. 충돌을 처리하는 한 가지 방법은 해당 인덱스에 두 키-값 쌍을 모두 저장하는 것입니다. 그런 다음 조회 시 해당 버킷의 항목들을 순차적으로 반복하여 찾는 키를 찾아야 합니다. 좋은 해싱 함수는 충돌을 최소화하여 효율적인 검색 시간을 유지합니다.

여기서는 해싱이나 해시 테이블 구현의 세부 사항에 신경 쓰지 않고, 해시 테이블이 어떻게 작동하는지 대략적으로 이해하는 데 집중합니다.

# --instructions--

기본적인 해시 테이블 기능을 만들어 보세요. 여러분이 사용할 수 있도록 단순한 해싱 함수를 만들었습니다. 문자열 값을 `hash` 함수에 전달하면 저장소의 키로 사용할 수 있는 해시 값을 반환합니다. 이 해시 값을 기반으로 `this.collection` 객체에 항목을 저장하세요. 세 가지 메서드 `add`, `remove`, `lookup`를 만드세요. 첫 번째는 해시 테이블에 추가할 키-값 쌍을 받아야 합니다. 두 번째는 키를 전달받아 해당 키-값 쌍을 제거해야 합니다. 세 번째는 키를 받아 연관된 값을 반환하거나 키가 없으면 `null`를 반환해야 합니다.

충돌을 처리할 수 있도록 코드를 작성하세요!

**참고:** `remove` 메서드 테스트는 `add`과 `lookup` 메서드가 올바르게 구현되어야 통과합니다.

# --hints--

`hash` 함수는 유효해야 합니다.

```js
let calls = 0;
const ourHash = string => {
  calls++;
  let hashCode = 0;
  for (let i = 0; i < string.length; i++) {
    hashCode += string.charCodeAt(i);
  }
  return hashCode;
};

assert.strictEqual(hash('yek'),ourHash('yek'));

assert.strictEqual(hash('key'),ourHash('key'));
assert.strictEqual(hash('key1'),ourHash('key1'));
assert.strictEqual(hash('key2'),ourHash('key2'));
assert.strictEqual(hash('key3'),ourHash('key3'));

assert.strictEqual(hash('1key'),ourHash('1key'));
assert.strictEqual(hash('ke1y'),ourHash('ke1y'));
assert.strictEqual(hash('altKey'),ourHash('altKey'));

assert.strictEqual(called,calls); 
```

`HashTable` 자료 구조가 존재해야 합니다.

```js
assert.isDefined(HashTable); 
let newHashTable = new HashTable(); 
assert.isObject(newHashTable); 
```

`HashTable`은 `add` 메서드를 가져야 합니다.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.add); 
```

`HashTable`은 `lookup` 메서드를 가져야 합니다.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.lookup); 
```

`HashTable`은 `remove` 메서드를 가져야 합니다.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.remove); 
```

`add` 메서드는 키-값 쌍을 추가해야 하며 `lookup` 메서드는 주어진 키에 연관된 값을 반환해야 합니다.

```js
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
assert.strictEqual(newHashTable.lookup('key'),'value'); 
```

`remove` 메서드는 키를 입력으로 받아 연관된 키-값 쌍을 제거해야 합니다.

```js
let hashValue = hash('key');
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
newHashTable.remove('key'); 
assert.notProperty(newHashTable.collection,hashValue); 
```

`remove` 메서드는 정확한 키-값 쌍만 제거해야 합니다.

```js
let newHashTable = new HashTable(); 
let hashValue = hash('key');
newHashTable.add('key','value');
newHashTable.add('yek','value');
newHashTable.add('altKey','value');
newHashTable.remove('yek');

assert.notExists(newHashTable.lookup('yek'));
assert.exists(newHashTable.lookup('altKey'));
assert.exists(newHashTable.lookup('key'));

newHashTable.remove('key');

assert.notProperty(newHashTable.collection,hashValue); 
assert.exists(newHashTable.lookup('altKey'));

```

항목은 해시 함수를 사용해 추가해야 합니다.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1', 'value1')
newHashTable.add('key2', 'value2');
newHashTable.add('key3', 'value3');
assert.strictEqual(called,3); 
```

해시 테이블은 충돌을 처리해야 합니다.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1','value1');
newHashTable.add('1key', 'value2'); 
newHashTable.add('ke1y', 'value3'); 
assert.strictEqual(newHashTable.lookup('key1'),'value1');
assert.strictEqual(newHashTable.lookup('1key'),'value2');
assert.strictEqual(newHashTable.lookup('ke1y'),'value3');
```

# --seed--

## --seed-contents--

```js
var called = 0;
var hash = string => {
  called++;
  var hashed = 0;
  for (var i = 0; i < string.length; i++) {
    hashed += string.charCodeAt(i);
  }
  return hashed;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
var called = 0;
var hash = (string) => {
  called++;
  var hash = 0;
  for (var i = 0; i < string.length; i++) { hash += string.charCodeAt(i); }
  return hash;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line

  this.add = function(key, val) {
    var theHash = hash(key);
    if (!this.collection.hasOwnProperty(theHash)) {
      this.collection[theHash] = {};
    }
    this.collection[theHash][key] = val;
  }

  this.remove = function(key) {
    var theHash = hash(key);
    var hashedObj = this.collection[theHash];
    if (hashedObj.hasOwnProperty(key)) {
      delete hashedObj[key];
    }
    if (!Object.keys(hashedObj).length) {
      delete this.collection[theHash];
    }
  }

  this.lookup = function(key) {
    var theHash = hash(key);
    if (this.collection.hasOwnProperty(theHash)) {
      return this.collection[theHash][key];
    }
    return null
  }
  // Only change code above this line
};
```
