---
id: 587d825b367417b2b2512c8b
title: 최대 힙에서 요소 제거하기
challengeType: 1
forumTopicId: 301710
dashedName: remove-an-element-from-a-max-heap
---

# --description--

이제 힙에 요소를 추가할 수 있으니, 요소를 어떻게 제거할 수 있는지 살펴보겠습니다. 요소를 제거하고 삽입하는 것은 비슷한 논리를 필요로 합니다. 최대 힙에서는 보통 가장 큰 값을 제거하는데, 이는 트리의 루트에서 값을 추출하는 것을 의미합니다. 이렇게 하면 트리의 힙 속성이 깨지므로, 이를 다시 세워야 합니다. 일반적으로 최대 힙에서는 다음과 같은 방식으로 처리합니다:

<ol>
  <li>힙의 마지막 요소를 루트 위치로 옮깁니다.</li>
  <li>루트의 자식 중 하나가 루트보다 크면, 더 큰 값을 가진 자식과 루트를 교환합니다.</li>
  <li>부모가 두 자식보다 크거나 트리의 마지막 레벨에 도달할 때까지 교환을 계속합니다.</li>
</ol>

# --instructions--

지시사항: 최대 힙에 `remove`라는 메서드를 추가하세요. 이 메서드는 최대 힙에 추가된 가장 큰 값을 반환하고 힙에서 제거해야 합니다. 또한 힙 속성이 유지되도록 힙을 재정렬해야 합니다. 요소를 제거한 후에는 힙에 남아 있는 다음으로 큰 요소가 루트가 되어야 합니다.

# --hints--

`MaxHeap` 자료 구조가 존재해야 합니다.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap`는 `print`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

`MaxHeap`는 `insert`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap`는 `remove`라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

`remove` 메서드는 최대 힙 속성을 유지하면서 최대 힙에서 가장 큰 요소를 제거해야 합니다.

```js
function isHeap(arr, i, n) {
    if( arr[i] < arr[2 * i + 1] || arr[i] < arr[2 * i + 2] ){
        return false;
    }
    if (i > (n - 1) / 2) {
        return true;
    }
    if (isHeap(arr, 2 * i + 1, n) && isHeap(arr, 2 * i + 2, n)) {
        return true;
    }
    return false;
}


assert(
  (function () {
    let test = false;

    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    
    let max = Infinity;
    const [result, vals] = [[], [9, 3, 5, 2, 15, 3, 7, 12, 7, 10, 90]];

    vals.forEach((val) => test.insert(val));

    for (let i = 0; i < vals.length; i++) {
      const curHeap = test.print();
      const arr = curHeap[0] === null ? curHeap.slice(1) : curHeap;
      
      if (!isHeap(arr, 0, arr.length - 1)) {
        return false;
      }

      const removed = test.remove();
      if (!vals.includes(removed)) return false;
      if (removed > max) return false
      max = removed;
      result.push(removed);
    }
    
    for (let i = 0; i < vals.length; i++) {
      if (!result.includes(vals[i])) {
        return false;
      }
    }
    
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
const MaxHeap = function () {
  this.heap = [];
  this.parent = index => {
    return Math.floor((index - 1) / 2);
  }
  this.insert = element => {
    this.heap.push(element);
    this.heapifyUp(this.heap.length - 1);
  }
  this.heapifyUp = index => {
    let currentIndex = index,
    parentIndex = this.parent(currentIndex);
    while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
      this.swap(currentIndex, parentIndex);
      currentIndex = parentIndex;
      parentIndex = this.parent(parentIndex);
    }
  }
  this.swap = (index1, index2) => {
    [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
  }
  this.print = () => {
    return this.heap;
  }
  // Only change code below this line

  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
