---
id: 587d825a367417b2b2512c88
title: 이중 연결 리스트 뒤집기
challengeType: 1
forumTopicId: 301714
dashedName: reverse-a-doubly-linked-list
---

# --description--

이중 연결 리스트에 reverse라는 메서드를 하나 더 만들어 리스트를 제자리에서 뒤집어 보세요. 메서드가 실행되면 head는 이전 tail을 가리키고 tail은 이전 head를 가리켜야 합니다. 이제 head에서 tail로 리스트를 순회하면 원래 리스트와 반대 순서로 노드를 만나야 합니다. 빈 리스트를 뒤집으려 하면 null을 반환해야 합니다.

# --hints--

`DoublyLinkedList` 자료 구조가 존재해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

`DoublyLinkedList`에는 `reverse`라는 메서드가 있어야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.reverse == undefined) {
      return false;
    }
    return typeof test.reverse == 'function';
  })()
);
```

빈 리스트를 뒤집으면 `null`를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.reverse() == null;
  })()
);
```

`reverse` 메서드는 리스트를 뒤집어야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    var n1 = new Node(58, null);
    var n2 = new Node(61, n1);
    var n3 = new Node(32, n2);
    var n4 = new Node(95, n3);
    var n5 = new Node(41, n4);
    n1.next = n2;
    n2.next = n3;
    n3.next = n4;
    n4.next = n5;
    test.head = n1;
    test.tail = n5;
    test.reverse();

    var result = [];
    var current = test.head;
    while (current) {
      result.push(current.data);
      current = current.next;
    }

    return result.join('') == '4195326158';
  })()
);
```

리스트를 뒤집을 때 `next`와 `previous` 참조가 올바르게 유지되어야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    var n1 = new Node(11, null);
    var n2 = new Node(22, n1);
    var n3 = new Node(33, n2);
    var n4 = new Node(44, n3);
    var n5 = new Node(55, n4);
    n1.next = n2;
    n2.next = n3;
    n3.next = n4;
    n4.next = n5;
    test.head = n1;
    test.tail = n5;
    test.reverse();

    var result = [];
    var current = test.tail;
    while (current) {
      result.push(current.data);
      current = current.prev;
    }

    return result.join('') == '1122334455';
  })()
);
```

# --seed--

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
  var Node = function(data, prev) {
    this.data = data;
    this.prev = prev;
    this.next = null;
  };
  var DoublyLinkedList = function() {
    this.head = null;
    this.tail = null;

    this.reverse = function() {
      if (!this.head || !this.head.next) {
        return this.head
      }

      let tail;
      let temp;
      let current = this.head;
      while(current !== null) {
        if(!tail) tail = current;
        temp = current.prev;
        current.prev = current.next;
        current.next = temp;
        current = current.prev;
      }

      this.head = temp.prev;
      this.tail = tail
    }
  };
```
