---
id: 587d825b367417b2b2512c8c
title: 최소 힙으로 힙 정렬 구현하기
challengeType: 1
forumTopicId: 301643
dashedName: implement-heap-sort-with-a-min-heap
---

# --description--

이제 요소를 추가하고 제거할 수 있으니 힙이 사용되는 몇 가지 애플리케이션을 살펴보겠습니다. 힙은 항상 가장 크거나 가장 작은 값을 가진 항목을 첫 번째 위치에 저장하기 때문에 우선순위 큐를 구현하는 데 자주 사용됩니다. 또한 힙 정렬이라는 정렬 알고리즘을 구현하는 데도 사용됩니다. 여기서 그 방법을 알아보겠습니다. 힙 정렬은 최대 힙의 반대인 최소 힙을 사용합니다. 최소 힙은 항상 가장 작은 값을 가진 요소를 루트 위치에 저장합니다.

힙 정렬은 정렬되지 않은 배열을 가져와 배열의 각 항목을 최소 힙에 추가한 다음, 최소 힙에서 모든 항목을 추출하여 새 배열에 넣는 방식으로 작동합니다. 최소 힙 구조 덕분에 새 배열에는 원래 항목들이 가장 작은 것부터 가장 큰 순서로 들어가게 됩니다. 이 알고리즘은 평균 및 최악의 경우 성능이 O(nlog(n))으로 가장 효율적인 정렬 알고리즘 중 하나입니다.

# --instructions--

최소 힙으로 힙 정렬을 구현해 보세요. 최대 힙 코드를 여기서 자유롭게 응용하세요. `MinHeap` 객체를 만들고 `insert`, `remove`, `sort` 메서드를 포함하세요. `sort` 메서드는 최소 힙에 있는 모든 요소를 가장 작은 것부터 가장 큰 것 순서로 정렬한 배열을 반환해야 합니다.

# --hints--

`MinHeap` 자료 구조가 존재해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MinHeap`은 `insert`이라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MinHeap`은 `remove`이라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

`MinHeap`은 `sort`이라는 메서드를 가져야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.sort == 'function';
  })()
);
```

`sort` 메서드는 최소 힙에 추가된 모든 항목을 정렬된 순서로 포함하는 배열을 반환해야 합니다.

```js
assert(
  (() => {
    if (typeof MinHeap === 'undefined') {
      return false;
    }

    const heap = new MinHeap();
    const arr = createRandomArray(25);

    for (let i of arr) {
      heap.insert(i);
    }

    const result = heap.sort();
    arr.sort((a, b) => a - b);

    for (let i = 0; i < arr.length; i++) {
      if (arr[i] !== result[i]) {
        return false;
      }
    }
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
// Generate a randomly filled array
function createRandomArray(size = 5){
  let a = new Array(size);
  for(let i = 0; i < size; i++)
    a[i] = Math.floor(Math.random() * 100);
  
  return a;
}
const array = createRandomArray(25);

var MinHeap = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
