---
id: 587d8255367417b2b2512c72
title: ES6 Set에서 .has와 .size 사용하기
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

ES6 Set 객체에서 사용할 수 있는 .has와 .size 메서드를 살펴보겠습니다.

먼저, ES6 Set을 만드세요

```js
var set = new Set([1,2,3]);
```

.has 메서드는 값이 Set에 포함되어 있는지 확인합니다.

```js
var hasTwo = set.has(2);
```

.size 메서드는 Set의 크기를 나타내는 정수를 반환합니다

```js
var howBig = set.size;
```

# --instructions--

이번 연습에서는 배열과 값을 checkSet() 함수에 전달합니다. 함수는 배열 인수로부터 ES6 Set을 만들어야 합니다. Set에 값 인수가 포함되어 있는지 확인하세요. Set의 크기를 찾으세요. 그리고 이 두 값을 배열로 반환하세요.

# --hints--

`checkSet([4, 5, 6], 3)`는 [ false, 3 ]을 반환해야 합니다

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
