---
id: 5cc0c1b32479e176caf3b422
title: 이진 탐색 트리인지 확인하기
challengeType: 1
forumTopicId: 301624
dashedName: check-if-tree-is-binary-search-tree
---

# --description--

이진 탐색 트리가 무엇인지 이미 알고 있으니, 이 도전 과제에서는 트리가 이진 탐색 트리인지 아닌지 판단하는 방법을 알아봅니다.

이진 탐색 트리의 주요 특징은 노드들이 체계적으로 정렬되어 있다는 점입니다. 노드는 최대 2개의 자식 노드를 가지며, 자식 노드의 값이 부모 노드보다 크거나 같으면 오른쪽에, 작으면 왼쪽에 배치됩니다.

# --instructions--

이번 도전 과제에서는 트리를 위한 유틸리티를 만듭니다. 트리를 입력으로 받아 그 트리가 이진 탐색 트리인지 아닌지를 불리언 값으로 반환하는 JavaScript 메서드 `isBinarySearchTree`를 작성하세요. 가능한 재귀를 사용하세요.

# --hints--

이진 탐색 트리는 `isBinarySearchTree()`로 검사했을 때 true를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    test.push(1);
    test.push(5);
    test.push(3);
    test.push(2);
    test.push(4);
    return isBinarySearchTree(test) == true;
  })()
);
```

이진 탐색 트리가 아닌 트리로 검사했을 때 `isBinarySearchTree()`는 false를 반환해야 합니다.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    test.push(1);
    test.root.left = new Node(1);
    return isBinarySearchTree(test) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype.push = function(val) {
  var root = this.root;

  if (!root) {
    this.root = new Node(val);
    return;
  }

  var currentNode = root;
  var newNode = new Node(val);

  while (currentNode) {
    if (val < currentNode.value) {
      if (!currentNode.left) {
        currentNode.left = newNode;
        break;
      } else {
        currentNode = currentNode.left;
      }
    } else {
      if (!currentNode.right) {
        currentNode.right = newNode;
        break;
      } else {
        currentNode = currentNode.right;
      }
    }
  }
};
```

# --seed--

## --seed-contents--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
}
function isBinarySearchTree(tree) {
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
}
function isBinarySearchTree(tree) {
  if (tree.root == null) {
    return null;
  } else {
    let isBST = true;
    function checkTree(node) {
      if (node.left != null) {
        const left = node.left;
        if (left.value >= node.value) {
          isBST = false;
        } else {
          checkTree(left);
        }
      }
      if (node.right != null) {
        const right = node.right;
        if (right.value < node.value) {
          isBST = false;
        } else {
          checkTree(right);
        }
      }
    }
    checkTree(tree.root);
    return isBST;
  }
};
```
