---
id: 66eaddd04a9e533fba689001
title: Diseña una tarjeta para una publicación de blog
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

En este laboratorio, practicarás cómo estilizar fondos y bordes creando una tarjeta de publicación de blog.

**Objetivo:** Cumple con las historias de usuario a continuación y logra que pasen todas las pruebas para completar el laboratorio.

**Instrucciones:**

1. Deberías tener un elemento `div` con una clase de `blog-post-card` para contener todos tus elementos de tarjeta.
2. Dentro del elemento `.blog-post-card`, deberías tener un elemento de imagen con un atributo `alt` válido y texto, y una clase de `post-img`. Puedes usar `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` para el atributo `src` de la imagen.
3. Deberías tener un `div` con una clase de `post-content` dentro del elemento `.blog-post-card` con lo siguiente:
    - Un elemento `h2` con una clase de `post-title` y texto para el título de tu publicación de blog.
    - Un elemento `p` con una clase de `post-excerpt` y texto para resumir tu publicación de blog.
    - Un elemento `a` con una clase de `read-more` con el texto `Read More`.
4. Deberías aplicar los siguientes estilos al elemento `.blog-post-card`:
    - Un fondo blanco.
    - Esquinas redondeadas.
    - Un ancho de tu elección.
    - La alineación del texto de tu elección.
5. El elemento `.post-img` debería estilizarse para que la imagen llene todo el ancho de la tarjeta y tenga un borde inferior.
6. El elemento `.post-content` debería estilizarse para que haya relleno dentro de la tarjeta.
7. Los elementos `.post-title` y `.post-excerpt` deberían tener un color de texto diferente al predeterminado y márgenes en todos los lados.
8. El elemento `.read-more` debería estilizarse como un botón y tener:
    - Un color de texto diferente al predeterminado.
    - Un color de fondo.
    - Márgenes en todos los lados.
    - Relleno en todos los lados.
    - Esquinas redondeadas.
    - Una propiedad `display` configurada en `inline-block`.
    - Un efecto hover que cambia su color de fondo.

**Nota:** Asegúrate de enlazar tu hoja de estilos en tu HTML y aplicar tu CSS.

# --hints--

Deberías tener un `div` con una clase de `blog-post-card`.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

Deberías tener un elemento `img` con una `class` de `post-img`. Asegúrate de que tu imagen tenga un atributo `alt` con texto y un atributo `src` con un valor.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

Deberías tener un `div` con clase de `post-content`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

Deberías tener un elemento `h2` con una clase de `post-title`. Asegúrate de que tenga algún contenido de texto para el título de tu publicación.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

Deberías tener un `p` con una clase de `post-excerpt`. Asegúrate de que tenga algún contenido de texto para resumir tu publicación.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

Deberías tener un elemento `a` con una `class` de `read-more`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

Tu elemento `.read-more` debería tener el texto `Read More`.

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

Tu elemento `.blog-post-card` debería tener una propiedad `border-radius` con un valor (no debería ser `0` ni un valor negativo).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

Tu elemento `.blog-post-card` debería tener un fondo blanco.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

Deberías apuntar al `.blog-post-card` y establecer su propiedad `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

Deberías apuntar al `.blog-post-card` y establecer su propiedad `text-align`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

Deberías apuntar al `.post-content` y establecer su propiedad `padding`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

Tu elemento `.read-more` debería tener un efecto hover.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

Deberías apuntar al `.read-more` y establecer su propiedad `color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

Deberías apuntar al `.read-more` y establecer su propiedad `background-color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

Deberías apuntar al `.read-more` y establecer su propiedad `margin`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

Deberías apuntar al `.read-more` y establecer su propiedad `display`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

Deberías apuntar al `.read-more` y establecer su propiedad `border-radius`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

Deberías apuntar al `.read-more` y establecer su propiedad `padding`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

Tu elemento `.post-img` debería llenar el `width` de la tarjeta.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

Tu elemento `.post-img` debería tener un valor de `border-bottom`.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

Tus elementos `.post-title` y `.post-excerpt` deberían tener márgenes y colores de texto no predeterminados.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
