---
id: 619665c9abd72906f3ad30f9
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Ви створите щасливого Flappy Penguin і додатково попрактикуєтеся з CSS-трансформаціями та анімаціями.

Почніть із підключення вашого стилю до сторінки.

# --hints--

Ваш код має містити `link` елемент.

```js
assert.exists(document.querySelector('link'));
```

Ваш елемент `link` має бути всередині елемента `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Ваш `link` елемент має мати атрибут `rel` зі значенням `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Ваш `link` елемент має мати атрибут `href` зі значенням `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
--fcc-editable-region--
    
--fcc-editable-region--
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css

```
