---
id: 6565a536ba1f9f25bd30e88b
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste projeto, você implementará o Algoritmo de Luhn. Este algoritmo é uma fórmula para validar uma variedade de números de identificação.

Comece declarando uma função chamada `main`, esta servirá como o ponto de entrada do programa. Use a palavra-chave `pass` para evitar um erro.

# --hints--

Você deve ter `def main():` no seu código.

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

Você deve ter `pass` no seu código.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
