---
id: 656877f71bba2b97acedf9af
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Python vem com classes integradas que podem nos ajudar com a manipulação de strings.Um deles é a classe `str`.Ele tem um método chamado `maketrans` que pode nos ajudar a criar uma tabela de tradução.Esta tabela pode ser usada para substituir caracteres em uma string:

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

O acima, quando chamado em uma string, substituirá todos os caracteres `t` por `c` e todos os caracteres `l` por `b`.

Crie uma variável chamada `card_translation` e atribua a ela uma tabela de tradução para substituir todos os caracteres `-` e ` ` por uma string vazia.

# --hints--

Você deve criar uma variável `card_translation` dentro de `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

Você deve atribuir a `card_translation` um valor de `str.maketrans({'-': '', ' ': ''})`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
