---
id: 65687b48f2201ab32e06c37c
title: Passo 21
challengeType: 20
dashedName: step-21
---

# --description--

Atualmente, seu script lança um `TypeError` porque você está tentando adicionar uma string a um inteiro.Você pode corrigir isso convertendo a variável `digit` para um inteiro antes de adicioná-la a `sum_of_odd_digits`, usando a função embutida `int`:

```python
my_string = '123'
my_int = int(my_string)
```

Converta a variável `digit` para um inteiro antes de adicioná-la a `sum_of_odd_digits`.Então, mova a chamada `print` para o final da função `verify_card_number` para imprimir o valor de `sum_of_odd_digits`.

# --hints--

Você deve ter `sum_of_odd_digits += int(digit)` dentro do loop `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

Você deve ter `print(sum_of_odd_digits)` no final da função `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

Você não deve mais usar `print(digit)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
