---
id: 656877c202f44995ed952608
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Substitua a declaração `pass` por uma variável chamada `card_number` e um valor de `'4111-1111-4555-1142'`.

# --hints--

Você deve ter `card_number = '4111-1111-4555-1142'` dentro da função `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

Você não deve ter um `pass` no seu código.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
