---
id: 65687b182c2a8fb1470e0b97
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Dentro do loop `for`, use o operador `+=` para adicionar o `digit` à variável `sum_of_odd_digits`.

Fazendo isso seu script lança um `TypeError` porque você está tentando adicionar uma string a um inteiro, mas não se preocupe, você vai aprender mais sobre como fazer isso funcionar na próxima etapa.

# --hints--

Você deve ter `sum_of_odd_digits += digit` dentro do loop `for`.

```js
const transformedCode = code.replace(/\r/g, "");
const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
const { function_body } = verify_card_number;

assert.match(function_body, /sum_of_odd_digits\s*\+=\s*digit/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        print(digit)
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
