---
id: 656878585631369a6b2d2191
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Definir a tradução não traduz a string por si só. O método `translate` deve ser chamado na string a ser traduzida com a tabela de tradução como argumento:

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

Crie uma variável chamada `translated_card_number` e atribua a ela o resultado da chamada do método `translate` em `card_number` com `card_translation` como argumento.

# --hints--

Você deve criar uma variável `translated_card_number` dentro de `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

Você deve atribuir a `translated_card_number` um valor de `card_number.translate(card_translation)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
