---
id: 656879c6f35a85a59c06b3a7
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Você já acessou elementos (caracteres) de uma string antes, usando o operador de índice `[]`. Você também pode usar o operador de índice para acessar um intervalo de caracteres em uma string com `string[start:stop:step]`:

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

Onde `start` é o índice inicial (inclusivo), `stop` é o índice final (exclusivo) e `step` é a quantidade de caracteres a serem pulados. Se não especificado, `step` tem valor padrão 1.

Crie uma variável chamada `card_number_reversed` e atribua a ela o valor dos primeiros 4 caracteres de `card_number`.

# --hints--

Você deve ter `card_number_reversed = card_number[0:4]` dentro da função `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
