---
id: 656879a66338b2a461d5d307
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

O algoritmo de Luhn é o seguinte:

1. Da direita para a esquerda, dobre o valor de cada segundo dígito; se o produto for maior que 9, some os dígitos dos produtos.
2. Some a soma de todos os dígitos.
3. Se a soma de todos os dígitos for múltipla de 10, então o número é válido; caso contrário, não é válido.

Assuma um exemplo de um número de conta "7992739871" que terá um dígito verificador adicionado, tornando-o da forma 7992739871x:

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

Substitua a instrução `pass` por uma variável `sum_of_odd_digits` e um valor de `0`.

# --hints--

Você deve ter `sum_of_odd_digits = 0` dentro da função `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

Você não deve ter `pass` dentro da função `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
