---
id: 65687a8253766cac4c99c57f
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Dentro da função `verify_card_number`, crie uma variável `odd_digits` que contenha cada outro dígito da string `card_number_reversed`.

# --hints--

Você deve ter `odd_digits = card_number_reversed[::2]` dentro da função `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /odd_digits\s*=\s*card_number_reversed\s*\[\s*:\s*:\s*2\s*\]/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
