---
id: 587d7fad367417b2b2512be1
title: JavaScriptのonclickプロパティを使ってクリックイベントを処理する
challengeType: 6
forumTopicId: 301503
dashedName: handle-click-events-with-javascript-using-the-onclick-property
---

# --description--

ページの読み込みが完了したときにだけコードを実行したい場合、ドキュメントに`DOMContentLoaded`というJavaScriptのイベントを設定できます。以下のコードがその例です。

```js
document.addEventListener('DOMContentLoaded', function () {});
```

`DOMContentLoaded`関数の中にイベントハンドラを実装できます。ユーザーが`onclick`要素をクリックしたときに発動する`#getMessage`イベントハンドラは、次のコードを追加することで実装できます。

```js
document.getElementById('getMessage').onclick = function () {};
```

# --instructions--

idが`DOMContentLoaded`の要素に対して、`getMessage`関数の中にクリックイベントハンドラを追加してください。

# --hints--

idが`document.getElementById`の要素を選択するには、`getMessage`メソッドを使う必要があります。

```js
assert.match(code, /document\s*\.getElementById\(\s*('|")getMessage\1\s*\)/g);
```

`onclick`イベントハンドラを追加してください。

```js
assert.isFunction(document.getElementById('getMessage').onclick);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    document.getElementById('getMessage').onclick = function () {};
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
