---
id: 587d7fae367417b2b2512be3
title: JavaScript の XMLHttpRequest メソッドで JSON を取得する
challengeType: 6
forumTopicId: 301502
dashedName: get-json-with-the-javascript-xmlhttprequest-method
---

# --description--

外部ソースからデータをリクエストすることもできます。ここで API が登場します。

API（アプリケーションプログラミングインターフェース）は、コンピューター同士が通信するためのツールです。AJAX という技術を使って、API から取得したデータで HTML を更新する方法を学びます。

ほとんどのウェブ API は JSON という形式でデータを転送します。JSON は JavaScript Object Notation の略です。

JSON の構文は JavaScript のオブジェクトリテラル表記に非常によく似ています。JSON は `{` と `}` の間にオブジェクトのプロパティとその現在の値を持ちます。

これらのプロパティと値は「キーと値のペア」と呼ばれることが多いです。

しかし、API から送信される JSON は `bytes` として送られ、アプリケーションでは `string` として受け取ります。これらは JavaScript のオブジェクトに変換できますが、デフォルトでは JavaScript のオブジェクトではありません。`JSON.parse` メソッドは文字列をパースして、それが表す JavaScript オブジェクトを構築します。

freeCodeCamp の Cat Photo API から JSON をリクエストできます。これを行うためにクリックイベントに入れられるコードは以下の通りです。

```js
const req = new XMLHttpRequest();
req.open('GET', '/json/cats.json', true);
req.send();
req.onload = function () {
  const json = JSON.parse(req.responseText);
  document.getElementsByClassName('message')[0].innerHTML =
    JSON.stringify(json);
};
```

各部分が何をしているかのレビューです。JavaScript の `XMLHttpRequest` オブジェクトにはデータ転送に使われる多くのプロパティとメソッドがあります。まず、`XMLHttpRequest` オブジェクトのインスタンスが作成され、`req` 変数に保存されます。次に、`open` メソッドがリクエストを初期化します。この例は API からデータをリクエストしているので、`GET` リクエストです。`open` の第2引数はリクエスト先の API の URL です。第3引数はブール値で、`true` にすると非同期リクエストになります。`send` メソッドがリクエストを送信します。最後に、`onload` イベントハンドラが返されたデータをパースし、`JSON.stringify` メソッドを使って JavaScript オブジェクトを文字列に変換します。この文字列がメッセージテキストとして挿入されます。

# --instructions--

コードを更新して、freeCodeCamp Cat Photo API に `GET` リクエストを作成して送信してください。次に `Get Message` ボタンをクリックしてください。AJAX 関数が `The message will go here` テキストを API からの生の JSON 出力に置き換えます。

# --hints--

コードは新しい `XMLHttpRequest` を作成する必要があります。

```js
assert.match(code, /new\s+?XMLHttpRequest\(\s*?\)/g);
```

コードは `open` メソッドを使って freeCodeCamp Cat Photo API に `GET` リクエストを初期化する必要があります。

```js
assert.match(
  code,
  /\.open\(\s*('|")GET\1\s*,\s*('|")\/json\/cats\.json\2\s*,\s*true\s*\)/g
);
```

コードは `send` メソッドを使ってリクエストを送信する必要があります。

```js
assert.match(code, /\.send\(\s*\)/g);
```

コードには関数に設定された `onload` イベントハンドラが必要です。

```js
assert.match(
  code,
  /\.onload\s*=\s*(function|\(\s*\))\s*(\(\s*\)|\=\>)\s*{/g
);
```

コードは `JSON.parse` メソッドを使って `responseText` をパースする必要があります。

```js
assert.match(code, /JSON\s*\.parse\(\s*.*\.responseText\s*\)/g);
```

コードはクラス `message` の要素を取得し、その inner HTML を JSON データの文字列に変更する必要があります。

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.innerHTML\s*?=\s*?JSON\.stringify\(.+?\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = () => {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
