---
id: 587d7fae367417b2b2512be4
title: APIからJSONデータにアクセスする
challengeType: 6
forumTopicId: 301499
dashedName: access-the-json-data-from-an-api
---

# --description--

前のチャレンジでは、freeCodeCamp Cat Photo APIからJSONデータを取得する方法を学びました。

今回は返されたデータを詳しく見て、JSON形式をよりよく理解しましょう。JavaScriptの記法を思い出してください：

<blockquote>[ ] -> 大括弧は配列を表します。<br>{ } -> 波括弧はオブジェクトを表します。<br>" " -> 二重引用符は文字列を表します。JSONではキー名にも使われます。</blockquote>

APIが返すデータの構造を理解することは重要です。なぜなら、それが必要な値をどのように取得するかに影響するからです。

右側で`Get Message`ボタンをクリックして、freeCodeCamp Cat Photo APIのJSONをHTMLに読み込んでください。

JSONデータの最初と最後の文字は大括弧`[ ]`です。これは返されたデータが配列であることを意味します。JSONデータの2文字目は波括弧`{`で、オブジェクトの開始を示しています。よく見ると、3つの別々のオブジェクトがあることがわかります。JSONデータは3つのオブジェクトの配列で、それぞれのオブジェクトが猫の写真に関する情報を含んでいます。

前に学んだように、オブジェクトはカンマで区切られた「キーと値のペア」を含みます。Cat Photoの例では、最初のオブジェクトに`"id":0`があり、`id`がキーで`0`が対応する値です。同様に、`imageLink`、`altText`、`codeNames`のキーもあります。各猫の写真オブジェクトはこれらと同じキーを持ちますが、値は異なります。

最初のオブジェクトにあるもう一つ興味深い「キーと値のペア」は`"codeNames":["Juggernaut","Mrs. Wallace","ButterCup"]`です。ここで`codeNames`がキーで、その値は3つの文字列の配列です。配列の中にオブジェクトがあったり、キーの値として配列があることもあります。

配列やオブジェクトのデータにアクセスする方法を思い出してください。配列はブラケット記法を使って特定のインデックスのアイテムにアクセスします。オブジェクトはブラケット記法かドット記法で特定のプロパティの値にアクセスします。以下は最初の猫の写真の`altText`プロパティを出力する例です。エディター内の解析済みJSONデータは`json`という変数に保存されています：

```js
console.log(json[0].altText);
```

コンソールには文字列`A white cat wearing a green helmet shaped melon on its head.`が表示されます。

# --instructions--

`id`が2の猫について、`codeNames`配列の2番目の値をコンソールに出力してください。オブジェクト（`json`という変数に保存されています）に対してブラケット記法とドット記法を使って値にアクセスしてください。

# --hints--

コードはブラケット記法とドット記法を使って正しいコード名にアクセスし、`Loki`をコンソールに出力する必要があります。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*json\s*\[2\]\s*(\.\s*codeNames|\[\s*('|`|")codeNames\2\s*\])\s*\[\s*1\s*\]\s*\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line

        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line
        console.log(json[2].codeNames[1]);
        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
