---
id: 587d7faf367417b2b2512be9
title: JavaScript の XMLHttpRequest メソッドでデータを送信する
challengeType: 6
forumTopicId: 301504
dashedName: post-data-with-the-javascript-xmlhttprequest-method
---

# --description--

前の例では、外部リソースからデータを受け取りました。外部リソースがAJAXリクエストに対応していてURLがわかっていれば、データを送信することもできます。

JavaScriptの`XMLHttpRequest`メソッドは、サーバーにデータを送信するためにも使われます。例を見てみましょう。

```js
const xhr = new XMLHttpRequest();
xhr.open('POST', url, true);
xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
xhr.onreadystatechange = function () {
  if (xhr.readyState === 4 && xhr.status === 201) {
    const serverResponse = JSON.parse(xhr.response);
    document.getElementsByClassName('message')[0].textContent =
      serverResponse.userName + serverResponse.suffix;
  }
};
const body = JSON.stringify({ userName: userName, suffix: ' loves cats!' });
xhr.send(body);
```

これらのメソッドは以前にもいくつか見たことがあります。ここでは`open`メソッドが、外部リソースの指定されたURLに対して`POST`としてリクエストを初期化し、3番目のパラメータに`true`を渡して非同期で操作を行うことを示しています。

`setRequestHeader`メソッドは、送信者やリクエストに関する情報を含むHTTPリクエストヘッダーの値を設定します。これは`open`メソッドの後、`send`メソッドの前に呼び出す必要があります。2つのパラメータはヘッダー名と、そのヘッダーの本文として設定する値です。

次に、`onreadystatechange`イベントリスナーがリクエストの状態変化を処理します。`readyState`が`4`の場合は操作が完了しており、`status`が`201`の場合はリクエストが成功したことを意味します。したがって、ドキュメントのHTMLを更新できます。

最後に、`send`メソッドが`body`の値を使ってリクエストを送信します。`body`は`userName`と`suffix`キーで構成されています。

# --instructions--

コードを更新して、APIエンドポイントに`POST`リクエストを送信してください。次に、入力フィールドに名前を入力して`Send Message`をクリックしてください。AJAX関数は`Reply from Server will be here.`をサーバーからのデータで置き換える必要があります。レスポンスは、名前に` loves cats`というテキストを付け加えて表示してください。

# --hints--

コードは新しい`XMLHttpRequest`を作成する必要があります。

```js
assert.match(code, /new\s*XMLHttpRequest\(\s*\)/g);
```

コードは`open`メソッドを使ってサーバーに`POST`リクエストを初期化する必要があります。

```js
assert.match(code, /\.open\(\s*('|")POST\1\s*,\s*url\s*,\s*true\s*\)/g);
```

コードは`setRequestHeader`メソッドを使う必要があります。

```js
assert.match(
  code,
  /\.setRequestHeader\(\s*('|")Content-Type\1\s*,\s*('|")application\/json;\s*charset=UTF-8\2\s*\)/g
);
```

コードには関数に設定された`onreadystatechange`イベントハンドラが必要です。

```js
assert.match(code, /\.onreadystatechange\s*=/g);
```

コードはクラス`message`の要素を取得し、その`textContent`を`userName loves cats`に変更する必要があります。

```js
assert.match(
  code,
  /document\.getElementsByClassName\(\s*('|")message\1\s*\)\[0\]\.textContent\s*=\s*.+?\.userName\s*\+\s*.+?\.suffix/g
);
```

コードは`send`メソッドを使う必要があります。

```js
assert.match(code, /\.send\(\s*body\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line

      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message box">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line
      const xhr = new XMLHttpRequest();
      xhr.open('POST', url, true);
      xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
      xhr.onreadystatechange = function () {
        if (xhr.readyState === 4 && xhr.status === 201) {
          const serverResponse = JSON.parse(xhr.response);
          document.getElementsByClassName('message')[0].textContent =
            serverResponse.userName + serverResponse.suffix;
        }
      };
      const body = JSON.stringify({
        userName: userName,
        suffix: ' loves cats!'
      });
      xhr.send(body);
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```
