---
id: 5ccfad82bb2dc6c965a848e5
title: JavaScript の fetch メソッドで JSON を取得する
challengeType: 6
forumTopicId: 301501
dashedName: get-json-with-the-javascript-fetch-method
---

# --description--

外部データを取得するもう一つの方法は、JavaScript の `fetch()` メソッドを使うことです。これは `XMLHttpRequest` と同等ですが、構文がより理解しやすいとされています。

以下は `/json/cats.json` に GET リクエストを送るコードです。

```js
fetch('/json/cats.json')
  .then(response => response.json())
  .then(data => {
    document.getElementById('message').innerHTML = JSON.stringify(data);
  });
```

注意：`fetch()` メソッドはデフォルトで `GET` を `HTTP` メソッドとして使います。つまり、基本的なデータ取得では明示的に指定する必要はありません。

このコードの各部分を見てみましょう。

最初の行がリクエストを送る部分です。つまり、`fetch(URL)` は指定された URL に `GET` リクエストを送ります。このメソッドは Promise を返します。

Promise が返された後、リクエストが成功すると `then` メソッドが実行され、レスポンスを受け取って JSON 形式に変換します。

`then` メソッドも Promise を返し、次の `then` メソッドで処理されます。2番目の `then` の引数は、探している JSON オブジェクトです！

次に、`document.getElementById()` を使ってデータを受け取る要素を選択します。そして、リクエストから返された JSON オブジェクトを文字列に変換して、その要素の HTML コードを変更します。

# --instructions--

コードを更新して、freeCodeCamp Cat Photo API に `GET` リクエストを作成して送信してください。ただし今回は、`fetch` の代わりに `XMLHttpRequest` メソッドを使います。

# --hints--

取得したデータを使って inner HTML を置き換えるコードにしてください。

```js
const catData = 'dummy data';
const ref = fetch;
fetch = () => Promise.resolve({ json: () => catData });

  try {
    document.getElementById('getMessage').click();
    await new Promise((resolve, reject) => setTimeout(() => resolve(), 250));
  } catch (error) {
    console.log(error);
  } finally {
    fetch = ref;
    assert.equal(
      document.getElementById('message').textContent,
      JSON.stringify(catData)
    );
  }

```

`GET` を使って `fetch` リクエストを送るコードにしてください。

```js
assert.match(code, /fetch\s*\(\s*('|")\/json\/cats\.json\1\s*\)/g);
```

レスポンスを JSON に変換するために `then` を使ってください。

```js
assert.match(
  code,
  /\.then\s*\(\s*\(?(?<var>\w+)\)?\s*=>\s*\k<var>\s*\.json\s*\(\s*\)\s*\)/g
);
```

もう一方の `then` によって JSON に変換されたデータを処理するために `then` を使ってください。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.then\(\(?\w+\)?=>{[^}]*}\)/g);
```

id が `message` の要素を取得し、その inner HTML を JSON データの文字列に変更してください。

```js
assert.match(
  __helpers.removeWhiteSpace(code),
  /document\.getElementById\(('|")message\1\)\.innerHTML=JSON\.stringify\(?\w+\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      fetch('/json/cats.json')
        .then(response => response.json())
        .then(data => {
          document.getElementById('message').innerHTML = JSON.stringify(data);
        });
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
