---
id: 5a858944d96184f06fd60d61
title: Crea la tua prima griglia CSS
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
dashedName: create-your-first-css-grid
---

# --description--

Trasforma qualsiasi elemento HTML in un contenitore griglia impostando la sua proprietà `display` su `grid`. Questo ti permette di usare tutte le altre proprietà associate alla griglia CSS.

**Nota:** Nella griglia CSS, l'elemento genitore è chiamato <dfn>contenitore</dfn> e i suoi figli sono chiamati <dfn>elementi</dfn>.

# --instructions--

Cambia la proprietà display del div con la classe `container` in `grid`.

# --hints--

La classe `container` dovrebbe avere una proprietà `display` con valore `grid`.

```js
assert.match(code, /.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    display: grid;
  }
</style>
```
