---
id: 5a9036ee38fddaf9a66b5d37
title: aggiungere spazi più velocemente con grid-gap
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/ca2qVtv'
forumTopicId: 301118
dashedName: add-gaps-faster-with-grid-gap
---

# --description--

`grid-gap` è una proprietà abbreviata per `grid-row-gap` e `grid-column-gap` delle due sfide precedenti, più comoda da usare. Se `grid-gap` ha un solo valore, creerà uno spazio tra tutte le righe e le colonne. Tuttavia, se ci sono due valori, il primo imposterà lo spazio tra le righe e il secondo tra le colonne.

# --instructions--

Usa `grid-gap` per introdurre uno spazio `10px` tra le righe e uno spazio `20px` tra le colonne.

# --hints--

La classe `container` dovrebbe avere una proprietà `grid-gap` che introduce uno spazio `10px` tra le righe e uno spazio `20px` tra le colonne.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-gap\s*?:\s*?10px\s+?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>
<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-gap: 10px 20px;
  }
</style>
```
