---
id: 5a90376038fddaf9a66b5d3c
title: Allinea tutti gli elementi orizzontalmente usando justify-items
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpECn'
forumTopicId: 301120
dashedName: align-all-items-horizontally-using-justify-items
---

# --description--

A volte vuoi che tutti gli elementi nella tua griglia CSS condividano lo stesso allineamento. Puoi usare le proprietà imparate in precedenza e allinearli singolarmente, oppure puoi allinearli tutti insieme orizzontalmente usando `justify-items` sul contenitore griglia. Questa proprietà può accettare tutti i valori che hai imparato nelle due sfide precedenti, con la differenza che sposterà **tutti** gli elementi della nostra griglia verso l’allineamento desiderato.

# --instructions--

Usa questa proprietà per centrare tutti i nostri elementi orizzontalmente.

# --hints--

La classe `container` dovrebbe avere una proprietà `justify-items` con valore `center`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*justify-items\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    justify-items: center;
  }
</style>
```
