---
id: 5a94fe3669fb03452672e45f
title: Riduci la ripetizione usando la funzione repeat
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cQvqyHR'
forumTopicId: 301133
dashedName: reduce-repetition-using-the-repeat-function
---

# --description--

Quando hai usato `grid-template-columns` e `grid-template-rows` per definire la struttura di una griglia, hai inserito un valore per ogni riga o colonna creata.

Supponiamo che tu voglia una griglia con 100 righe della stessa altezza. Non è molto pratico inserire 100 valori singolarmente. Fortunatamente, c’è un modo migliore: usare la funzione `repeat` per specificare quante volte vuoi che la tua colonna o riga venga ripetuta, seguita da una virgola e dal valore che vuoi ripetere.

Ecco un esempio che creerebbe una griglia di 100 righe, ciascuna alta 50px.

```css
grid-template-rows: repeat(100, 50px);
```

Puoi anche ripetere più valori con la funzione repeat e inserire la funzione tra altri valori quando definisci la struttura di una griglia. Ecco come appare:

```css
grid-template-columns: repeat(2, 1fr 50px) 20px;
```

Questo si traduce in:

```css
grid-template-columns: 1fr 50px 1fr 50px 20px;
```

**Nota:** `1fr 50px` è ripetuto due volte seguito da 20px.

# --instructions--

Usa `repeat` per eliminare la ripetizione dalla proprietà `grid-template-columns`.

# --hints--

La classe `container` dovrebbe avere una proprietà `grid-template-columns` impostata per ripetere 3 colonne con larghezza `1fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?1fr\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: 1fr 1fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, 1fr);
  }
</style>
```
