---
id: 642df9df4b5216350de7b0d2
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Necesitarás una función para generar un rango de números.

Declara una función `range` vacía que tome un parámetro `start` y `end`. Usa el constructor `Array()` y devuelve implícitamente un array vacío.

# --hints--

Debes declarar una variable `range`.

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

Debes usar `const` para declarar tu variable `range`.

```js
assert.match(code, /const\s+range/);
```

Tu variable `range` debe ser una función.

```js
assert.isFunction(range);
```

Tu función `range` debe usar la sintaxis de flecha.

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

Tu función `range` debe tomar primero un parámetro `start`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

Tu función `range` debe tomar un parámetro `end` en segundo lugar.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

Tu función `range` debe usar un retorno implícito. Recuerda que esto significa que no usarás llaves.

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

Tu función `range` debe usar el constructor `Array()`. Los constructores primitivos no necesitan la palabra clave `new`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

No debes pasar nada al constructor `Array()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
