---
id: 642dfb07e7fa6736251541c8
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Tu arreglo debe tener el tama침o del rango. Puedes calcularlo encontrando la diferencia entre `end` y `start`, y sumando `1` al resultado.

Pasa este c치lculo como el argumento para tu constructor `Array()`.

# --hints--

Debes restar `start` de `end`.

```js
assert.match(code, /end\s*-\s*start/);
```

Debes sumar `1` a tu c치lculo de `end - start`.

```js
assert.match(code, /end\s*-\s*start\s*\+\s*1/);
```

Debes pasar tu c치lculo al constructor `Array()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
const range = (start, end) => Array();
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
