---
id: 642ddfdea4200e313f80a4b6
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Recuerda que el objeto `document` tiene un método `.createElement()` que te permite crear dinámicamente nuevos elementos HTML.

En tu función `createLabel`, declara una variable `label` y asígnale un nuevo elemento `div`.

# --hints--

Debes declarar una variable `label` en tu función `createLabel`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

Tu variable `label` debe declararse con `const`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

Debes usar el método `.createElement()` del objeto `document`.

```js
assert.match(code, /document\.createElement\(/);
```

Debes pasar la cadena `"div"` al método `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

Debes asignar tu nuevo elemento `div` a `label`.

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
