---
id: 6434552bcc0a951a0a99df3b
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Tu función `range` espera números, pero tus valores de `start` y `end` serán cadenas (específicamente, serán caracteres individuales como `A`).

Convierte tus valores `start` y `end` en tu llamada a `range()` a números utilizando el método `.charCodeAt()` en ellos, pasando el número `0` como argumento a ese método.

# --hints--

Debes usar el método `.charCodeAt()`.

```js
assert.match(code, /\.charCodeAt\(/);
```

Debes llamar al método `.charCodeAt()` en `start`.

```js
assert.match(code, /start\.charCodeAt\(/);
```

Debes pasar `0` al método `.charCodeAt()` de `start`.

```js
assert.match(code, /start\.charCodeAt\(\s*0\s*\)/);
```

Debes llamar al método `.charCodeAt()` en `end`.

```js
assert.match(code, /end\.charCodeAt\(/);
```

Debes pasar `0` al método `.charCodeAt()` de `end`.

```js
assert.match(code, /end\.charCodeAt\(\s*0\s*\)/);
```

Debes usar los métodos `.charCodeAt()` directamente en tu llamada a `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start, end);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
