---
id: a26cbbe9ad8655a977e1ceb5
title: 构建一个最长单词查找器应用
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

在本实验中，你将构建一个函数，返回所提供句子中最长单词的长度。

例如，在句子 `"The quick brown fox jumped over the lazy dog"` 中，最长的单词是 `"jumped"`，长度为 6。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 `findLongestWordLength` 的函数，该函数以一个字串作为参数。
2. 该函数应返回字串中最长单词的长度。

# --hints--

你应该创建一个名为 `findLongestWordLength` 的函数。

```js
assert.isFunction(findLongestWordLength);
```

`findLongestWordLength` 应该只有一个参数。

```js
assert.lengthOf(findLongestWordLength, 1);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` 应返回一个数字。

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` 应返回 `6`。

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

`findLongestWordLength("May the force be with you")` 应返回 `5`。

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

`findLongestWordLength("Google do a barrel roll")` 应返回 `6`。

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

`findLongestWordLength("Googling do a barrel roll")` 应该返回 `8`。

```js
assert.strictEqual(findLongestWordLength('Googling do a barrel roll'), 8);
```

`findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` 应返回 `8`。

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

`findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` 应返回 `19`。

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
