---
id: 66fe4f33a2cc9b33f4d5cd9b
title: Einen Permutationsgenerator erstellen
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

In diesem Labor erstellen Sie einen Permutationsgenerator, der einen String entgegennimmt und alle möglichen Permutationen der Zeichen im String zurückgibt. Zum Beispiel sind die möglichen Permutationen des Strings `cat` `cat`, `cta`, `act`, `atc`, `tac` und `tca`.

Die rekursive Methode zur Erstellung von Permutationen eines Strings funktioniert, indem der feste Anfangsteil des Strings (Präfix) gespeichert wird und Permutationen des restlichen Teils erstellt werden.

Zum Beispiel betrachten wir das Wort `machine`. Die erste Runde der Permutationserstellung würde darin bestehen, `m` als Präfix des Strings festzulegen und dann Permutationen des restlichen Strings, `achine`, zu erstellen.

Für den restlichen String werden die Permutationen auf die gleiche Weise fortgesetzt. Ein Buchstabe wird zum Präfix hinzugefügt, vielleicht `c`, sodass das Präfix `mc` wird. Dann wird jede der Permutationen von `ahine` an das Präfix angehängt.

Dies wird fortgesetzt, bis das Präfix alle Buchstaben enthält und der restliche String leer ist, was bedeutet, dass eine Permutation erstellt wurde.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `permuteString` erstellen.
2. Die `permuteString`-Funktion sollte einen Parameter, einen String, und dann zwei Parameter mit Standardwerten annehmen: einen Präfixwert und ein leeres Array zum Speichern und Zurückgeben der Ergebnisse. Der Präfixwert sollte ein String sein, der verwendet wird, um Zeichen zu sammeln und eine Permutation zu bilden. Die Funktion wird mit nur einem Argument aufgerufen, wie `permuteString("cat")`.
3. Innerhalb der Funktion sollten Sie prüfen, ob die Länge des übergebenen Strings `0` ist. Falls ja, fügen Sie das aktuelle Präfix zu den Ergebnissen hinzu und geben die Ergebnisse zurück.
4. Durchlaufen Sie jedes Zeichen im Eingabestring und entfernen Sie bei jeder Iteration das aktuelle Zeichen aus dem String. Rufen Sie dann die `permuteString`-Funktion rekursiv mit aktualisierten Argumenten auf, um die verbleibenden Permutationen zu erstellen.
5. Sie sollten das finale Ergebnis-Array zurückgeben.
6. Sie sollten sicherstellen, dass die Permutationen eindeutig sind, indem Sie Duplikate entfernen.

# --hints--

Sie sollten eine Funktion `permuteString` haben.

```js
assert.isFunction(permuteString);
```

Sie sollten Rekursion in Ihrer `permuteString`-Funktion verwenden.

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` sollte `[ "far", "fra", "afr", "arf", "rfa", "raf" ]` zurückgeben.

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` sollte `[ "fcc", "cfc", "ccf" ]` zurückgeben.

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` sollte `[ "p" ]` zurückgeben.

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` sollte `[""]` zurückgeben.

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")` sollte `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]` zurückgeben.
`.

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` sollte die korrekten Ergebnisse zurückgeben.

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
