---
id: 66c06fad3475cd92421b9ac2
title: بناء حاسبة للسنة الكبيسة
challengeType: 26
dashedName: build-a-leap-year-calculator
---

# --description--

السنة الكبيسة هي السنة التي تقبل القسمة على `4`، باستثناء السنوات التي تقبل القسمة على `100` ولا تقبل القسمة على `400`. على سبيل المثال، `2000` سنة كبيسة، لكن `1900` ليست كذلك. أيضًا، تحتوي السنة الكبيسة على يوم إضافي في فبراير، وهو اليوم التاسع والعشرون من الشهر.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. أعلن دالة باسم `isLeapYear` تأخذ رقمًا كمعلمة.
2. خارج الدالة، أعلن متغيرًا باسم `year` يخزن قيمة السنة التي تريد التحقق منها.
3. داخل الدالة، استخدم عبارة `if/ else` أو معامل ثلاثي للتحقق مما إذا كانت السنة كبيسة.
4. للتحقق من كبيسة السنة، استوفِ الشروط التالية:

- إذا كانت السنة تقبل القسمة على `4`، فهي سنة كبيسة.
   - إلا إذا كانت السنة تقبل القسمة أيضًا على `100`، فهي ليست سنة كبيسة.
   - إلا إذا كانت السنة تقبل القسمة أيضًا على `400`، فهي سنة كبيسة.

5. إذا كانت السنة كبيسة، أرجع `[year] is a leap year.`. وإلا، أرجع `[year] is not a leap year.`. ستستبدل `[year]` بالمعلمة المعرفة في دالة `isLeapYear`.
6. يجب استدعاء دالة `isLeapYear` مع `year` كمعلمة وتعيين النتيجة إلى متغير باسم `result`.
7. يجب إخراج متغير `result` إلى وحدة التحكم باستخدام `console.log()`.

# --hints--

يجب إعلان دالة باسم `isLeapYear`.

```js
assert.isFunction(isLeapYear);
```

يجب أن تحتوي دالة `isLeapYear` على معلمة واحدة.

```js
assert.lengthOf(isLeapYear, 1)
```

يجب إعلان متغير `year` وتعيين قيمة له للتحقق مما إذا كانت سنة كبيسة.

```js
assert.isDefined(year);
assert.isNumber(year);
assert.isAtLeast(year, 0);
assert.isTrue(Number.isInteger(year));
```

يجب ألا يكون متغير `year` فارغًا.

```js
assert.isNotNull(year);
```

مع تعيين قيمة `2024` لمتغير `year`، يجب أن تكون قيمة `result` هي `2024 is a leap year.`

```js
assert.strictEqual(isLeapYear(2024), '2024 is a leap year.');
```

مع تعيين قيمة `2000` لمتغير `year`، يجب أن تكون قيمة `result` هي `2000 is a leap year.`

```js
assert.strictEqual(isLeapYear(2000), '2000 is a leap year.');
```

مع تعيين قيمة `1900` لمتغير `year`، يجب أن تكون قيمة `result` هي `1900 is not a leap year.`

```js
assert.strictEqual(isLeapYear(1900), '1900 is not a leap year.');

```

يجب استدعاء دالة `isLeapYear` وتمرير `year` كمعلمة.

```js
assert.match(__helpers.removeJSComments(code), /isLeapYear\(\s*year\s*\)/);
```

يجب إعلان متغير `result`.

```js
assert.isDefined(result);
```

يجب تخزين نتيجة استدعاء دالة `isLeapYear` في متغير باسم `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*isLeapYear\(\s*year\s*\)/);
```

يجب إخراج `result` إلى وحدة التحكم باستخدام `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function isLeapYear(year) {
  if ((year % 4 === 0 && year % 100 !== 0) || year % 400 === 0) {
    return(year + " is a leap year.");
  } else {
    return(year + " is not a leap year.");
  }
}

const year = 2024;
const result = isLeapYear(year);
console.log(result);
```

