---
id: 5dc2385ff86c76b9248c6eb7
title: 步驟5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 has some elements that identify different content areas. These elements make your HTML easier to read and help with Search Engine Optimization (SEO) and accessibility.

`main` 元素用於表示 HTML 文檔正文的主要內容。 `main` 元素裏的內容應該是文檔中唯一的，不應該在文檔的其他部分重複。

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Identify the main section of this page by adding a `<main>` opening tag before the `h1` element, and a `</main>` closing tag after the `p` element.

# --hints--

Your `main` element should have an opening tag. 開始標籤的語法是：`<elementName>`。

```js
assert(document.querySelector('main'));
```

Your `main` element should have a closing tag. 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert(code.match(/<\/main\>/));
```

Your `main` element's opening tag should be below the `body` element's opening tag. 你把順序寫錯了。

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Your `main` element's opening tag should be above the `h1` element. 你把順序寫錯了。

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H1'));
```

Your `main` element's closing tag should be below the `p` element. 你的順序是錯誤的。

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode.contains(pNode));
assert.match(pNode.textContent.toLowerCase(), /everyone loves cute cats online!/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```

