---
id: 5ddbd81294d8ddc1510a8e56
title: 步驟11
challengeType: 0
dashedName: step-11
---

# --description--

鏈接的文本必須放置在錨元素（`a`）的開始和結束標籤之間。

這是一個鏈接文本是 `click here to go to freeCodeCamp.org` 的示例：

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

將錨點文本 `cat photos` 添加到錨點元素中。 這將成爲鏈接的文本。

# --hints--

你的錨點元素（`a`）應該有一個開始標籤。 開始標籤的語法是：`<elementName>`。

```js
assert(document.querySelector('a'));
```

你的錨元素（`a`）應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert(code.match(/<\/a\>/));
```

你的錨點元素（`a`）中的文本應該是 `cat photos`。 確保將鏈接的文本放置在錨元素（`a`）的開始和結束標籤之間。

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

