---
id: 5ef9b03c81a63668521804e3
title: Step 60
challengeType: 0
dashedName: step-60
---

# --description--

Add another checkbox after the one you just added. The `id` attribute value should be `lazy` and the `name` attribute value should be the same as the last checkbox.

還要在新複選框的右側添加一個 `label` 元素，其中包含文本 `Lazy`。 確保使用 `for` 屬性將 `label` 元素與新複選框相關聯。

# --hints--

你需要添加一個新複選框。

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'), 2);
```

你的新複選框應該有一個值爲 `lazy` 的 `id` 屬性和一個值爲 `personality` 的 `name` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且/或所有的屬性名稱前面也要有一個空格。

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'lazy' && checkbox.getAttribute('name') === 'personality'
  )
);
```

你的新複選框應該在第一個複選框之後。 你把順序寫錯了。

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('loving') < checkboxes.indexOf('lazy'));
```

在新複選框的右側，應該有帶有文本 `Lazy` 的 `label` 元素。

```js
const nextElementSibling = document.querySelectorAll('input[type="checkbox"]')?.[1]?.nextElementSibling;
assert.isOk(nextElementSibling?.nodeName === 'LABEL' && nextElementSibling?.innerText?.replace(/\s+/g, '')?.match(/^Lazy$/i)
);
```

新的 `label` 應該有一個 `for` 屬性，其值與新複選框的 `id` 屬性值相同。 你可能落下了這個值或者有拼寫錯誤。

```js
assert.strictEqual(
  document.querySelectorAll('input[type="checkbox"]')?.[1]?.nextElementSibling?.getAttribute('for'),
    'lazy'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

