---
id: 5ef9b03c81a63668521804dc
title: 步驟 46
challengeType: 0
dashedName: step-46
---

# --description--

對於需要從多個選項中獲得一個答案的問題，你可以使用單選按鈕。

Here is an example of a radio button with the text set as `cat`:

```html
<input type="radio"> cat
```

請記住，`input` 元素是一個空元素。

Before the text input, add a radio button with the text set as:

`Indoor`

# --hints--

你應該爲單選按鈕創建一個 `input` 元素。 請檢查語法。

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

你的 `input` 元素應該有一個開始標籤，而不是一個結束標籤。

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

你應該只爲單選按鈕添加一個 `input` 元素。 刪除多餘的。

```js
assert(document.querySelectorAll('form > input'), 2);
```

你的新 `input` 元素應該在現有的 `input` 上方，`type` 屬性應該設置爲 `text` 。 你把順序寫錯了。

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

你的新 `input` 元素沒有 `type` 屬性。 檢查確認開始標籤名稱後應該有一個空格。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

你的新 `input` 元素應該只有一個 `type` 屬性。 刪除多餘的。

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

你的新 `input` 元素應該具有一個 `type` 屬性，其值爲 `radio`。 你可能忽略了這個值或者有拼寫錯誤。 請記住，應該把屬性值放在引號中。

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

儘管你已將新的 `input` 元素的 `type` 屬性設置爲 `radio`，但建議始終將屬性的值放在引號中。

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

`radio` 按鈕的 `Indoor` 文本應該位於它之後而不是之前。

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

`Indoor` 文本應直接位於 `radio` 按鈕的右側。 你可能忽略了文本或有拼寫錯誤。

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

