---
id: 671141f948cbab359e74cc93
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

添加 `p` 標籤，將 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` 變成一個段落。

# --hints--

你應該添加一個 `p` 元素包裹 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.`。

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P);
assert.strictEqual(P.innerText.trim(), "See more cat photos in our gallery.");
```

在新段落中仍然應該有錨元素。

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P.querySelector('a'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
