---
id: 671141d8e32fe934c26fa1be
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

在錨元素之前添加文字 `See more`，在錨元素之後添加文字 `in our gallery`。

# --hints--

你應該使用現有的錨元素，而不是創建一個新的元素。

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

The text inside the anchor element should not change, it must be `cat photos`.

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

You should have the words `See more` before the anchor element. Make sure you include the space after the last word.

```js
assert.match(code, /See more <a href=/)
```

你應該在錨元素後添加文字 `in our gallery`。 Make sure you include the space before the first word.

```js
assert.match(code, /<\/a> in our gallery/)
```

在你的代碼中應該有一個 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` 語句。

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
