---
id: 5cc0bd7a49b71cb96132e54c
title: 숫자 범위 생성기 만들기
challengeType: 26
dashedName: build-a-range-of-numbers-generator
---

# --description--

이 실습에서는 지정한 범위 내의 숫자 배열을 생성하는 `rangeOfNumbers` 함수를 만드세요.

예시:

- `rangeOfNumbers(3, 9)`는 `[3, 4, 5, 6, 7, 8, 9]`를 반환합니다.
- `rangeOfNumbers(5, 5)`는 `[5]`를 반환합니다.

요구사항:

- 재귀를 사용하세요 (함수가 자기 자신을 호출해야 합니다)
- 반복문이나 배열 메서드를 사용하지 마세요 (`for`, `while`, `Array.from`, `forEach`, `map`, `filter`, `reduce`)

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `rangeOfNumbers`라는 이름의 함수를 만들고 두 개의 매개변수 `startNum`과 `endNum`를 받으세요.
1. 함수는 `startNum` 매개변수가 나타내는 숫자로 시작하고 `endNum` 매개변수가 나타내는 숫자로 끝나는 정수 배열을 반환해야 합니다 (포함).
1. `startNum`은 항상 `endNum`보다 작거나 같습니다.
1. 함수는 자기 자신을 호출하는 재귀를 사용해야 합니다. 반복문 구문(`for`, `while`, `Array.from()`)이나 고차 함수(`forEach`, `map`, `filter`, `reduce`)를 사용하지 마세요.
1. 기본 사례로 `startNum`이 `endNum`과 같을 때, 그 숫자 하나만 포함하는 배열을 반환해야 합니다.
1. 재귀 사례에서는 매개변수를 수정해 자기 자신을 호출하여 배열을 만들고, 현재 숫자를 결과에 추가해야 합니다.
1. 배열을 저장하거나 만들기 위해 전역 변수를 사용하지 마세요.

# --hints--

`rangeOfNumbers`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(rangeOfNumbers);
```

`rangeOfNumbers` 함수는 두 개의 매개변수를 받아야 합니다.

```js
assert.match(rangeOfNumbers.toString(), /rangeOfNumbers\s*\(\s*\w+\s*,\s*\w+\s*\)/);
```

함수는 배열을 반환해야 합니다.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

코드에 반복문 구문(`for`, `while`, `Array.from()`)이나 고차 함수(`forEach`, `map`, `filter`, `reduce`)를 사용하지 마세요.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

`rangeOfNumbers`은 이 도전을 해결하기 위해 재귀(자기 자신 호출)를 사용해야 합니다.

```js
assert.isAtLeast(
  __helpers.removeJSComments(rangeOfNumbers.toString()).match(/rangeOfNumbers\s*\(.+\)/g)?.length,
  2
);
```

`rangeOfNumbers(1, 5)`은 `[1, 2, 3, 4, 5]`를 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)`은 `[6, 7, 8, 9]`를 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)`는 `[4]`를 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

`rangeOfNumbers(10, 15)`는 `[10, 11, 12, 13, 14, 15]`를 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(10, 15), [10, 11, 12, 13, 14, 15]);
```

`rangeOfNumbers(2, 8)`는 `[2, 3, 4, 5, 6, 7, 8]`를 반환해야 합니다.

```js
assert.deepStrictEqual(rangeOfNumbers(2, 8), [2, 3, 4, 5, 6, 7, 8]);
```

배열을 저장하기 위해 전역 변수를 사용하지 마세요.

```js
rangeOfNumbers(1, 3);
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
