---
id: 64861c02ff1ef4fa62a9e132
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

`position` 객체 안에 `x`이라는 새 키를 `proportionalSize(10)` 값으로 추가하세요. 그다음 `y`이라는 또 다른 키를 `proportionalSize(400)` 값으로 추가하세요.

여기서 `proportionalSize` 함수를 사용해야 플레이어 위치가 항상 화면 크기에 비례하도록 할 수 있습니다. 이는 화면 크기와 상관없이 플레이어가 화면을 자유롭게 움직일 수 있게 하려는 중요한 이유입니다.

# --hints--

`x` 객체에 `position`라는 새 키를 추가해야 합니다.

```js
assert.match(code, /this\.position/);
const player = new Player();

assert.property(player.position, 'x');
```

`x`의 값을 `proportionalSize(10)`로 설정해야 합니다.

```js
assert.match(code, /this\.position/);
const player = new Player();

assert.propertyVal(player.position, 'x', proportionalSize(10));
```


`y` 객체에 `position`라는 키를 추가해야 합니다.

```js
assert.match(code, /this\.position/);
const player = new Player();

assert.property(player.position, 'y');
```

`y`의 값을 `proportionalSize(400)`로 설정해야 합니다.

```js
assert.match(code, /this\.position/);
const player = new Player();

assert.propertyVal(player.position, 'y', proportionalSize(400));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Intermediate OOP by Building a Platformer Game</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="start-screen">
      <h1 class="main-title">freeCodeCamp Code Warrior</h1>
      <p class="instructions">
        Help the main player navigate to the yellow checkpoints.
      </p>
      <p class="instructions">
        Use the keyboard arrows to move the player around.
      </p>
      <p class="instructions">You can also use the spacebar to jump.</p>

      <div class="btn-container">
        <button class="btn" id="start-btn">Start Game</button>
      </div>
    </div>

    <div class="checkpoint-screen">
      <h2>Congrats!</h2>
      <p>You reached the last checkpoint.</p>
    </div>

    <canvas id="canvas"></canvas>

    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #0a0a23;
  --section-bg-color: #ffffff;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
}

.start-screen {
  background-color: var(--section-bg-color);
  width: 100%;
  position: absolute;
  top: 50%;
  left: 50%;
  margin-right: -50%;
  transform: translate(-50%, -50%);
  border-radius: 30px;
  padding: 20px;
  padding-bottom: 5px;
}

.main-title {
  text-align: center;
}

.instructions {
  text-align: center;
  font-size: 1.2rem;
  margin: 15px;
  line-height: 2rem;
}

.btn {
  cursor: pointer;
  width: 100px;
  margin: 10px;
  color: #0a0a23;
  font-size: 18px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

.btn-container {
  display: flex;
  align-items: center;
  justify-content: center;
}

.checkpoint-screen {
  position: absolute;
  left: 0;
  right: 0;
  margin-left: auto;
  margin-right: auto;
  width: 100%;
  text-align: center;
  background-color: var(--section-bg-color);
  border-radius: 20px;
  padding: 10px;
  display: none;
}

#canvas {
  display: none;
}

@media (min-width: 768px) {
  .start-screen {
    width: 60%;
    max-width: 700px;
  }

  .checkpoint-screen {
    max-width: 300px;
  }
}

```

```js
const startBtn = document.getElementById("start-btn");
const canvas = document.getElementById("canvas");
const startScreen = document.querySelector(".start-screen");
const checkpointScreen = document.querySelector(".checkpoint-screen");
const checkpointMessage = document.querySelector(".checkpoint-screen > p");
const ctx = canvas.getContext("2d");
canvas.width = innerWidth;
canvas.height = innerHeight;
const gravity = 0.5;
let isCheckpointCollisionDetectionActive = true;

const proportionalSize = (size) => {
  return innerHeight < 500 ? Math.ceil((size / 500) * innerHeight) : size;
}

--fcc-editable-region--

class Player {
  constructor() {
    this.position = {

    }
  }
}

--fcc-editable-region--

```
