---
id: 694a9ff65d4472dd451d8255
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop, eserciterai come usare i booleani e le istruzioni condizionali in Python costruendo un calcolatore per la prenotazione di biglietti del cinema.

Inizia dichiarando variabili che archiviano informazioni sull'utente e sulla proiezione del film.

Come hai imparato nelle lezioni precedenti, le variabili si assegnano usando l'operatore di assegnazione (`=`) in questo modo:

```py
# Assigns the value 10 to variable x
x = 10
```

Crea una variabile chiamata `base_price` per archiviare il prezzo base del biglietto del cinema e imposta il suo valore a `15`. Crea un'altra variabile chiamata `age` per archiviare l'età dell'utente e imposta il suo valore a `21`.

# --hints--

Dovresti dichiarare una variabile chiamata `base_price`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

Dovresti assegnare il valore intero `15` a `base_price`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

Dovresti dichiarare una variabile chiamata `age`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

Dovresti assegnare il valore intero `21` a `age`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
