---
id: 694cc1b515976889afd4cded
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Ora devi archiviare alcuni valori stringa riguardanti il biglietto del film. Puoi archiviare stringhe in Python in questo modo:

```py
name = 'Alice'
```

Crea una variabile chiamata `seat_type` per archiviare il tipo di sedile che l’utente ha selezionato e imposta il suo valore sulla stringa `Gold`. Crea un’altra variabile chiamata `show_time` per archiviare l’orario dello spettacolo del film e imposta il suo valore sulla stringa `Evening`.

Ricorda di racchiudere entrambi i valori tra virgolette singole o doppie.

# --hints--

Dovresti dichiarare una variabile chiamata `seat_type`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

Dovresti assegnare la stringa `Gold` a `seat_type`. Ricorda di racchiudere il valore tra virgolette singole o doppie.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

Dovresti dichiarare una variabile chiamata `show_time`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

Dovresti assegnare la stringa `Evening` a `show_time`. Ricorda di racchiudere il valore tra virgolette singole o doppie.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
