---
id: 6958e1875e592b58e8852757
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Ora crea una variabile chiamata `extra_charges` e impostala su `0`. Questa rappresenterà i costi extra da applicare al biglietto del cinema nei fine settimana.

Poi, crea un'istruzione `if` per verificare se `is_weekend` è truthy. All'interno del corpo dell'istruzione `if`, aggiorna il valore `extra_charges` a `2` e stampa `Extra charges will be applied` nel terminale.

# --hints--

Dovresti dichiarare una variabile chiamata `extra_charges`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

Dovresti assegnare il valore intero `0` a `extra_charges`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

Dovresti avere una quarta istruzione `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

La condizione della tua quarta istruzione `if` dovrebbe essere `is_weekend`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

Dovresti assegnare il valore intero `2` alla variabile `extra_charges` all'interno della tua nuova istruzione `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

Dovresti stampare `Extra charges will be applied` all'interno della tua nuova istruzione `if`. Ricorda di racchiudere il messaggio tra virgolette singole o doppie.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
