---
id: 6958de0d773289a89515a6b8
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora verificherai se l'utente può prenotare uno spettacolo serale in base alla sua età.

Crea un'istruzione `if` per controllare se `age` è maggiore o uguale a `21`. All'interno del corpo dell'istruzione `if`, stampa `User is eligible for Evening shows`.

# --hints--

Dovresti avere due istruzioni `if` nel tuo codice.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

La condizione della tua seconda istruzione `if` dovrebbe essere `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Dovresti stampare `User is eligible for Evening shows` all'interno della tua nuova istruzione `if`. Ricorda di racchiudere il messaggio tra virgolette singole o doppie.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
