---
id: 6958dfa0c74b237e1b56ce04
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

L'utente ha diritto a uno sconto per iscritti se è un membro.

Crea una variabile chiamata `discount` e imposta il suo valore su `0`. Questa variabile archivierà lo sconto che l'utente ottiene sul biglietto del film.

# --hints--

Dovresti dichiarare una variabile chiamata `discount`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

Dovresti assegnare il valore intero `0` a `discount`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
