---
id: 6958e2377cacbc46a7f7383a
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Ora, aggiungi una clausola `else` alla tua istruzione `if` e stampa `No extra charges will be applied` all'interno del corpo `else`. Questo stamperà il messaggio solo quando la condizione dei costi aggiuntivi è falsa.

Poi, sotto la clausola `else`, usa la chiamata `print()` per mostrare un messaggio che visualizza `Extra charges:` seguito dal valore aggiornato di `extra_charges` e controlla l'output nel terminale.

# --hints--

Dovresti avere una clausola `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

Dovresti stampare `No extra charges will be applied` all'interno della tua clausola `else`. Ricorda di racchiudere il messaggio tra virgolette singole o doppie.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

Dovresti avere `print('Extra charges:', extra_charges)` sotto la tua istruzione `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
