---
id: 6958e59845777391b86d7331
title: Passaggio 18
challengeType: 20
dashedName: step-18
---

# --description--

In Python, un'istruzione `if` può anche essere inserita all'interno del corpo di un'altra istruzione `if`. Questo si chiama un'istruzione `if` annidata.

Un'istruzione `if` annidata ti permette di verificare una condizione aggiuntiva solo dopo che la prima condizione è già stata soddisfatta. L'istruzione `if` interna verrà eseguita solo se la condizione `if` esterna è vera.

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

Ora calcolerai le spese di servizio in base al tipo di sedile che l'utente ha selezionato.

All'interno del corpo dell'ultima istruzione `if`, sotto la riga `print('Ticket booking condition satisfied')`, crea una variabile chiamata `service_charges` e assegnale il valore `0`. Assicurati di indentare il codice di quattro spazi per mantenerlo all'interno del corpo dell'istruzione `if` esterna.

Poi, crea un'istruzione `if` annidata per verificare se `seat_type` è uguale a `Premium`. All'interno del corpo dell'istruzione `if` annidata, aggiorna il valore `service_charges` a `5`.

# --hints--

Dovresti dichiarare una variabile chiamata `service_charges` all'interno del corpo della tua ultima istruzione `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

Dovresti assegnare il valore intero `0` a `service_charges`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

Dovresti avere un'istruzione `if` all'interno del corpo della tua ultima istruzione `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

La condizione della nuova istruzione `if` dovrebbe essere `seat_type == 'Premium'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

Dovresti assegnare il valore intero `5` alla variabile `service_charges` all'interno della tua nuova istruzione `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
