---
id: 6958dd85587e10c99162dd81
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Come hai imparato nelle lezioni precedenti, in Python un'istruzione `if` può essere usata per eseguire codice a seconda che una condizione sia vera.

Un'istruzione `if` consiste nella parola chiave `if`, seguita da una condizione e due punti. Il codice da eseguire quando la condizione è vera, che deve essere indentato, si chiama corpo dell'istruzione `if`.

```py
if condition:
    # Code to execute if condition is True
```

In questo passaggio, verificherai se l'utente è idoneo a prenotare un biglietto del cinema in base alla sua età.

Crea un'istruzione `if` per verificare se `age` è maggiore di `17`. All'interno del corpo dell'istruzione `if`, stampa `User is eligible to book a ticket`. Questo stamperà il messaggio solo quando l'età dell'utente è maggiore di `17`.

Ricorda di indentare il corpo dell'istruzione `if` e di racchiudere il messaggio tra virgolette singole o virgolette doppie all'interno della chiamata `print()`.

# --hints--

Dovresti avere un'istruzione `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

La condizione dell'istruzione `if` dovrebbe essere `age > 17`.

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

Dovresti avere una chiamata `print()` all'interno della tua istruzione `if`. Ricorda di indentare il corpo dell'istruzione `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

Dovresti stampare `User is eligible to book a ticket` all'interno della tua istruzione `if`. Ricorda di racchiudere il messaggio tra virgolette singole o virgolette doppie.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
