---
id: 6958e0ecc121b0ae3dd9c5e0
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Dovresti anche gestire il caso in cui l’utente non abbia diritto a uno sconto per l’iscrizione.

Aggiungi una clausola `else` alla tua istruzione `if` e stampa `User does not qualify for membership discount` all’interno del corpo `else`.

Puoi stampare un messaggio seguito da una variabile in questo modo:

```py
print('Total:', total)
```

Vuoi anche mostrare il valore aggiornato di `discount`. Sotto l’istruzione `if...else`, usa la chiamata `print()` per mostrare un messaggio che visualizza `Discount:` seguito dal valore aggiornato di `discount`. Poi, controlla l’output nel terminale.

# --hints--

Dovresti avere una clausola `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

Dovresti stampare `User does not qualify for membership discount` all’interno della tua clausola `else`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

Dovresti avere `print('Discount:', discount)` sotto la tua istruzione `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
