---
id: 696099c1ac72a4e585de1d5b
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Alcune informazioni possono essere solo vere o false. Come hai imparato nelle lezioni precedenti, questo può essere rappresentato usando valori booleani.

Crea una variabile chiamata `is_member` per indicare se l'utente è un membro e assegna il suo valore a `True`.

Sotto la variabile `is_member` crea un'altra variabile chiamata `is_weekend` per indicare se la proiezione del film è nel fine settimana e assegna il suo valore a `False`. Non racchiudere il valore tra virgolette.

# --hints--

Dovresti dichiarare una variabile chiamata `is_member`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

Dovresti assegnare il valore booleano `True` a `is_member`. Non racchiudere il valore tra virgolette.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

Dovresti dichiarare una variabile chiamata `is_weekend`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

Dovresti assegnare il valore booleano `False` a `is_weekend`. Non racchiudere il valore tra virgolette.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
