---
id: 696099abb6acc539b27055a1
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Nel passaggio precedente, hai verificato una condizione usando un'istruzione `if`. La clausola `else` serve a gestire il caso in cui la condizione sia falsa. Ecco la sintassi di un'istruzione `if...else`:

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

Ora, aggiungi una clausola `else` alla tua istruzione `if` e stampa `User is not eligible for Evening shows` all'interno del corpo `else`. Questo verrà stampato solo quando la condizione nell'istruzione `if` è falsa.

# --hints--

Dovresti avere una clausola `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

Dovresti stampare `User is not eligible for Evening shows` all'interno della tua clausola `else`. Ricorda di racchiudere il messaggio tra virgolette singole o doppie.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
