---
id: 6960b6a6ea06e27dcc7eef37
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

Ora verificherai se l'utente soddisfa le condizioni per prenotare un biglietto del cinema. Gli utenti con età `21` o superiore possono sempre prenotare biglietti senza alcuna restrizione.

Crea un'istruzione `if` per verificare se `age` è maggiore o uguale a `21`. All'interno del corpo dell'istruzione `if`, stampa `Ticket booking condition satisfied` nel terminale.

Poi, aggiungi una clausola `else` alla tua istruzione `if` e stampa `Ticket booking failed due to restrictions` all'interno del corpo `else`.

# --hints--

Dovresti avere una quinta istruzione `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

La condizione della tua quinta istruzione `if` dovrebbe essere `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Dovresti stampare `Ticket booking condition satisfied` all'interno della tua nuova istruzione `if`. Ricorda di racchiudere il messaggio tra virgolette singole o doppie.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

La tua nuova istruzione `if` dovrebbe avere una clausola `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

Dovresti stampare `Ticket booking failed due to restrictions` all'interno della clausola `else`. Ricorda di racchiudere il messaggio tra virgolette singole o doppie.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
