---
id: 6960ba6e6c95ed4ad896c6f2
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

C’è un’eccezione in più alle regole di prenotazione. Gli utenti tra `18` e `21` possono prenotare spettacoli serali se sono membri.

Quando si usano più operatori logici in un’istruzione `if`, le condizioni collegate con `and` vengono valutate prima delle condizioni collegate con `or`. Le parentesi tonde `()` si usano in Python per raggruppare le condizioni e controllare l’ordine in cui vengono valutate.

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

Aggiorna la condizione della riga `if age >= 21 or age >= 18 and show_time != 'Evening':` per aggiungere un’altra condizione usando l’operatore `or` per verificare se `is_member` è vera. Usa le parentesi tonde `()` per raggruppare insieme le condizioni `show_time != 'Evening'` e `is_member` come mostrato nell’esempio sopra.

# --hints--

L’ultima condizione dell’istruzione `if` dovrebbe essere `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
