---
id: 696cb889fc4fdeb337e52864
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

In Python, le condizioni `if` non devono confrontare esplicitamente i valori con `True` o `False`. Invece, si basano sulla veridicità dei valori. I valori truthy valutano a `True` in un contesto booleano, come la condizione di un'istruzione `if`. Questi includono numeri diversi da zero, stringhe non vuote e il valore booleano `True`.

```py
if value:
   print('value is truthy')
```

Crea un'istruzione `if` per verificare se `is_member` è truthy. All'interno del corpo dell'istruzione `if`, aggiorna il valore `discount` a `3` e stampa `User qualifies for membership discount` nel terminale.

# --hints--

Dovresti avere tre istruzioni `if` nel tuo codice.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

La condizione della tua terza istruzione `if` dovrebbe essere `is_member`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

Dovresti assegnare il valore intero `3` alla variabile `discount` all'interno della tua nuova istruzione `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

Dovresti stampare `User qualifies for membership discount` all'interno della tua nuova istruzione `if`. Ricorda di racchiudere il messaggio tra virgolette singole o doppie.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
