---
id: 68b94f23f7b3ee14078c89c5
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將使用 `iframe` 元素來顯示一個影片。基本的 HTML 範本已為你準備好。

首先建立一個 `h1` 元素，內容為 `iframe Video Display`。

# --hints--

你應該有一個 `h1` 元素。

```js
assert.exists(document.querySelector('h1'))
```

你的 `h1` 元素應該有文字 `iframe Video Display`。

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
