---
id: 68ba9c4f1688914d72876098
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

下一個你將添加的屬性是 `referrerpolicy`。它是決定當你的頁面連接到另一個頁面時，你分享多少細節的規則。

為你的元素添加 `referrerpolicy` 屬性並將其設為 `strict-origin-when-cross-origin`。這會在相同網站上分享完整位址，在其他網站上只分享網站名稱，而在不安全網站上則不分享任何資訊。

# --hints--

你的 `iframe` 元素應該設定 `referrerpolicy` 屬性為 `strict-origin-when-cross-origin`。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
