---
id: 68ba9c4f1688914d72876096
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

其中一個屬性是 `allow`。它就像一個權限列表，告訴瀏覽器 `iframe` 被允許使用哪些特性。

這裡有一個帶有 `allow` 屬性的 `iframe` 元素：

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

為你的 `allow` 屬性添加值 `accelerometer`、`autoplay` 和 `clipboard-write`。

`accelerometer` 讓 `iframe` 使用動作感測器，以便偵測裝置傾斜和旋轉等動作。`autoplay` 讓影片自動開始播放，`clipboard-write` 讓 iframe 將資料寫入使用者的剪貼簿。

# --hints--

你的 `iframe` 元素應該有一個設定為 `accelerometer` 的 `allow` 屬性。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

你的 `iframe` 元素的 `allow` 屬性應該包含 `autoplay` 作為其中一個值。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

你的 `iframe` 元素的 `allow` 屬性應該包含 `clipboard-write` 作為其中一個值。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
