---
id: 68ba9c4f1688914d72876093
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

現在，創建一個 `iframe` 元素。暫時不要在裡面放任何東西。

# --hints--

你應該有一個 `iframe` 元素。

```js
assert.exists(document.querySelector('iframe'))
```

你的 `iframe` 元素應該是空的。

```js
assert.equal(document.querySelector('iframe')?.innerHTML.trim(), '')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
