---
id: 68ba9c4f1688914d72876094
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在第一課關於 `iframe` 元素中，你學到它是一個與 `img` 相同的替換元素。這表示它也可以使用 `width` 和 `height` 屬性來決定它應該有多高和多寬。

為你的 `iframe` 元素設定 `width` 為 `560` 和 `height` 為 `315`。

# --hints--

你的 `iframe` 元素應該有一個設定為 `560` 的 `width` 屬性。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

你的 `iframe` 元素應該有一個設定為 `315` 的 `height` 屬性。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
