---
id: 68ba9c4f1688914d72876097
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

為你的 `allow` 屬性添加 `encrypted-media`、`gyroscope` 和 `web-share` 到現有值中。

這三者將允許使用加密媒體擴充套件來保護 `video`，授予存取裝置的動作和方向感應器，並允許透過裝置的原生分享對話窗分享 `iframe` 內容。 

# --hints--

你的 `iframe` 元素的 `allow` 屬性應該包含 `encrypted-media` 作為其中一個值。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

你的 `iframe` 元素的 `allow` 屬性應該包含 `gyroscope` 作為其中一個值。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

你的 `iframe` 元素的 `allow` 屬性應該包含 `web-share` 作為其中一個值。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
