---
id: 68ba9c4f1688914d72876095
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

`iframe` 元素也帶有一個 `src` 屬性，其值顯示要顯示的資源的 URL 或路徑。

為你的 `iframe` 元素添加 `src` 屬性，值為 `https://www.youtube.com/embed/I0_951_MPE0`。

此時，你應該會看到影片顯示在頁面上，但你還需要添加一些屬性。

# --hints--

你的 `iframe` 元素應該有一個 `src` 屬性，設定為 `https://www.youtube.com/embed/I0_951_MPE0`。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
