---
id: 616968c2c94c8071b349c146
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

`body` の中に、テキストが `h1` の `CSS Color Markers` 要素をネストしてください。

# --hints--

コードには `h1` 要素が必要です。

```js
const title = document.querySelector('h1');
assert.exists(title);
```

開く `h1` タグがあるはずです。

```js
assert.match(code, /<h1\s*>/i);
```

閉じる `h1` タグがあるはずです。

```js
assert.match(code, /<\/h1\s*>/i);
```

`h1` 要素は `body` 要素の中にある必要があります。

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

`h1` 要素にはテキスト `CSS Color Markers` が必要です。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

`h1` のテキストは大文字・小文字やスペルが正確であることを忘れないでください。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
