---
id: 61696ef7ac756c829f9e4048
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

このプロジェクトでは、外部CSSファイルを使ってページのスタイルを設定します。`styles.css`ファイルはすでに作成されています。しかし、それを使う前に、ページにリンクする必要があります。

`link`要素の中に`head`要素をネストしてください。`rel`属性に`"stylesheet"`を設定し、`href`属性に`"styles.css"`を設定してください。

# --hints--

`link`要素は1つだけにしてください。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link`要素は空要素で、終了タグを持ってはいけません `</link>`。

```js
assert.notMatch(code, /<\/link>/i);
```

`link`要素には、値が`rel`の`"stylesheet"`属性が必要です。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

`link`要素には、値が`href`の`"styles.css"`属性が必要です。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

`link`要素は`head`要素の中に入っている必要があります。

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
