---
id: 616d3a67ccf800ad94ec89ae
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

これから、最終的に色のマーカーとしてスタイルを適用する要素を追加します。

まず、`body` 要素の中に、`div` 属性が `class` に設定された `container` 要素を追加してください。`div` 要素が `h1` 要素の下に配置されていることを確認してください。

# --hints--

開く `div` タグがあるはずです。

```js
assert.match(code, /<div.*>/i);
```

閉じる `div` タグがあるはずです。

```js
assert.match(code, /<\/div\s*>/i);
```

`div` 要素は `body` 要素の中にあるはずです。

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

`div` 要素は `h1` 要素の後にあるはずです。

```js
assert.exists(document.querySelector('h1 + div'));
```

`div` 要素は `class` 属性が `container` に設定されているはずです。

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}
```
