---
id: 616d525007b8c5d8b3308b1c
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

マーカーはページの中央に配置すると見栄えが良くなります。これを簡単に行う方法は、`margin` <dfn>省略形プロパティ</dfn>を使うことです。

要素のマージンエリアは、`margin-top` や `margin-left` のようなプロパティで個別に設定できますが、`margin` という省略形プロパティを使うと複数のマージンエリアを同時に設定しやすくなります。

マーカーをページの中央に配置するには、`margin` プロパティを `auto` に設定してください。これにより、`margin-top`、`margin-right`、`margin-bottom`、および `margin-left` がすべて `auto` に設定されます。

# --hints--

`.marker` の CSS ルールには、`margin` プロパティを `auto` に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
