---
id: 616d55bd160a95e22453a081
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

色付きで中央に配置されたマーカーが1つできたので、他のマーカーも追加しましょう。

`.container` 要素の中に、さらに2つの `div` 要素を追加し、それぞれに `marker` のクラスを付けてください。

# --hints--

最初の新しい `div` 要素は開始タグを持つ必要があります。

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][2]);
```

最初の新しい `div` 要素は終了タグを持つ必要があります。

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][2]);
```

2番目の新しい `div` 要素は開始タグを持つ必要があります。

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][3]);
```

2番目の新しい `div` 要素は終了タグを持つ必要があります。

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][3]);
```

新しい `div` 要素は `div` の中で、クラス `container` を持つ要素の中にある必要があります。

```js
const containerChildren = [...document.querySelector('.container')?.children];
containerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(containerChildren, 3);
```

新しい `div` 要素は両方とも `marker` のクラスを持つ必要があります。

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert.isNotEmpty(containerChildren);
containerChildren.forEach(child => assert.isTrue(child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
