---
id: 61764fdda535587e1fefb3aa
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

2番目のマーカー`two`にクラス`div`を、3番目のマーカー`three`にクラス`div`を追加してください。

# --hints--

`two`要素内の2番目のマーカー`div`要素にクラス`.container`を追加する必要があります。

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('two'));
```

2番目のマーカー`div`にはクラス`marker`と`two`が付いているはずです。

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('marker'));
assert.isTrue(containerSecondChild.classList?.contains('two'));
```

`three`要素内の3番目のマーカー`div`要素にクラス`.container`を追加してください。

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('three'));
```

3番目のマーカー`div`にはクラス`marker`と`three`が付いているはずです。

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('marker'));
assert.isTrue(containerThirdChild?.classList?.contains('three'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

```
