---
id: 6176519636a76b810ab1219a
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

クラス `two` を対象にする CSS ルールを作成し、その `background-color` プロパティを `green` に設定してください。

また、クラス `three` を対象にする別の CSS ルールを作成し、その `background-color` を `blue` に設定してください。

# --hints--

クラス `two` を対象にするにはクラスセレクターを使う必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.two'));
```

`.two` の CSS ルールには `background-color` プロパティを `green` に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

クラス `three` を対象にするにはクラスセレクターを使う必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.three'));
```

`.three` の CSS ルールには `background-color` プロパティを `blue` に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
