---
id: 617bc3386dc7d07d6469bf20
title: ステップ 27
challengeType: 0
dashedName: step-27
---

# --description--

オレンジ色を作るために、`rgb` の値で赤の強さを上げ、緑の強さを下げたことに気づきましたか。これは、オレンジ色が赤と黄色を混ぜて作られるからです。

三次色のスプリンググリーンを作るには、シアンと緑を組み合わせます。`rgb` の CSS ルール内の `.two` 関数を更新して、緑を最大値に設定し、青を `127` にしてください。

# --hints--

`.two` の CSS ルールには、`background-color` プロパティが `rgb(0, 255, 127)` に設定されている必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 255);
--fcc-editable-region--
}

.three {
  background-color: rgb(255, 0, 255);
}

```
