---
id: 617bd6ec666b1da2587e4e37
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

三つの三次色があります：シャルトリューズグリーン（緑＋黄）、アジュール（青＋シアン）、ローズ（赤＋マゼンタ）です。

シャルトリューズグリーンを作るには、`rgb` の `.one` 関数を更新して、赤を `127` に設定し、緑を最大値にしてください。

アジュールの場合は、`rgb` の `.two` 関数を更新して、緑を `127` に設定し、青を最大値にしてください。

そして、時にブライトピンクとも呼ばれるローズの場合は、`rgb` の `.three` 関数を更新して、青を `127` に設定し、赤を最大値にしてください。

# --hints--

`.one` の CSS ルールには、`background-color` プロパティが `rgb(127, 255, 0)` に設定されている必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(127, 255, 0)');
```

`.two` の CSS ルールには、`background-color` プロパティが `rgb(0, 127, 255)` に設定されている必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 255)');
```

`.three` の CSS ルールには、`background-color` プロパティが `rgb(255, 0, 127)` に設定されている必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
