---
id: 617bbb6b97a83f6d1f7d6e38
title: ステップ 26
challengeType: 0
dashedName: step-26
---

# --description--

セカンダリーカラーの練習ができたので、<dfn>ターシャリーカラー</dfn>の作り方をレビューしましょう。ターシャリーカラーは、プライマリーカラーと近くのセカンダリーカラーを組み合わせて作られます。

ターシャリーカラーのオレンジを作るには、`rgb` の `.one` 関数を更新して、赤を最大値（`255`）に設定し、緑を `127` に設定してください。

# --hints--

`.one` の CSS ルールには、`background-color` プロパティが `rgb(255, 127, 0)` に設定されている必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: rgb(255, 255, 0);
--fcc-editable-region--
}

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
