---
id: 618a132676346ac9f7fd59dd
title: Step 34
challengeType: 0
dashedName: step-34
---

# --description--

中央の赤い色に自然と目が引き寄せられるのに気づきましたか？サイトをデザインするとき、この効果を使って重要な見出しやボタン、リンクに注目を集めることができます。

補色以外にも重要な色の組み合わせはいくつかありますが、それらは後ほど学びます。

今は、`rgb` の CSS ルール内で `.two` 関数を使って `background-color` を黒に設定してください。

# --hints--

`.two` の CSS ルールには `background-color` プロパティが `rgb(0, 0, 0)` に設定されている必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 0);
}

```
