---
id: 619b7fd56aa2253778bcf5f7
title: ステップ 43
challengeType: 0
dashedName: step-43
---

# --description--

<dfn>HSL</dfn> カラーモデル、つまり色相、彩度、明度は、色を表す別の方法です。

CSS の `hsl` 関数は3つの値を受け取ります：色相は0から360の数値、彩度は0から100のパーセンテージ、明度も0から100のパーセンテージです。

カラーホイールを想像すると、色相の赤は0度、緑は120度、青は240度にあります。

彩度は色の強さを表し、0%はグレー、100%は純粋な色です。彩度と明度の値にはパーセント記号 `%` を付ける必要があります。

明度は色の明るさを表し、0%は完全な黒、100%は完全な白で、50%が中間の明るさです。

`.blue` の CSS ルールで、`hsl` 関数を使って `background-color` プロパティを純粋な青に変更してください。色相は `240`、彩度は `100%`、明度は `50%` に設定します。

# --hints--

あなたの `.blue` CSS ルールには `background-color` プロパティが `hsl(240, 100%, 50%)` に設定されているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

.blue {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

```
