---
id: 61a4ae5f29eb5584187201c3
title: Step 56
challengeType: 0
dashedName: step-56
---

# --description--

このマーカーでは、グラデーションに16進数カラーコードを使います。

`linear-gradient` 関数を使い、`gradientDirection` を `180deg` に設定してください。そして最初の色の引数には、赤が `55`、緑が `68`、青が `0D` の値の16進数カラーコードを使ってください。

# --hints--

`.green` の CSS ルールには、値が `background` の `linear-gradient(180deg, #55680D)` プロパティがあるはずです。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.green\{.*?background:linear-gradient\(180deg,#55680D\)/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}

.green {
--fcc-editable-region--
  background: #007F00;
--fcc-editable-region--
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
