---
id: 61b0936d9e7fc020b43b1bb8
title: ステップ 73
challengeType: 0
dashedName: step-73
---

# --description--

袖が消えたように見えますが、心配しないでください — まだそこにあります。起きているのは、新しいキャップ `div` がマーカーの幅全体を占めていて、袖を次の行に押し下げているということです。

これは、`display` 要素のデフォルトの `div` プロパティが `block` だからです。つまり、2つの `block` 要素が隣り合うと、実際のブロックのように積み重なります。たとえば、マーカー要素はすべて重なって積み重なっています。

2つの `div` 要素を同じ行に配置するには、それらの `display` プロパティを `inline-block` に設定してください。

`cap` と `sleeve` の両方のクラスを対象にする新しいルールを作成し、`display` を `inline-block` に設定してください。

# --hints--

`cap` と `sleeve` の両方のクラスを対象にするには、クラスセレクターを使うべきです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

CSSルールには `display` プロパティを `inline-block` に設定してください。

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
