---
id: 61b093329e7fc020b43b1bb5
title: ステップ 70
challengeType: 0
dashedName: step-70
---

# --description--

色を設定するために`rgb`関数を使っていました。`rgb`カラーにアルファチャンネルを追加するには、代わりに`rgba`関数を使ってください。

`rgba`関数は`rgb`関数と同じように動作しますが、アルファチャンネル用に`0`から`1.0`までのもう1つの数値を受け取ります。

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

アルファチャンネルは`hsl`や`hex`カラーでも使えます。すぐにその方法を学びます。

`.sleeve`ルールでは、`rgba`関数を使って`background-color`プロパティを純粋な白で不透明度50％に設定してください。

# --hints--

あなたの`.sleeve` CSSルールには、`background-color`プロパティが`rgba(255, 255, 255, 0.5)`に設定されている必要があります。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor, 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
--fcc-editable-region--
  background-color: white;
--fcc-editable-region--
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
