---
id: 61b30995968123ceb6b76167
title: Step 83
challengeType: 0
dashedName: step-83
---

# --description--

影の端が鋭くなっていることに気づきましたか。これは `blurRadius` プロパティにオプションの `box-shadow` 値があるためです。

```css
box-shadow: offsetX offsetY blurRadius color;
```

もし `blurRadius` の値が含まれていなければ、デフォルトで `0` となり、鋭い端が作られます。`blurRadius` の値が大きいほど、ぼかし効果が強くなります。

`.green` の CSS ルールに、`box-shadow` プロパティを追加し、`5px` に `offsetX`、`5px` に `offsetY`、`5px` に `blurRadius`、`green` に `color` の値を設定してください。

# --hints--

`.green` の CSS ルールには、`box-shadow` 省略形プロパティがあり、値は `5px 5px 5px green` であるべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.boxShadow, 'green 5px 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
  box-shadow: -5px -5px red;
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
--fcc-editable-region--
  
--fcc-editable-region--
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
