---
id: 691f7bbf76229bb97c60827d
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 `for...of` 반복문을 사용하여 문자열 배열에서 특정 문자열이 몇 번 나타나는지 세는 함수를 만드는 연습을 합니다.

그 함수를 작성하기 전에, 문자열의 각 문자를 콘솔에 출력하는 더 간단한 함수를 먼저 만듭니다.

`printCharacters`라는 이름의 빈 함수를 정의하고 매개변수 `str`를 지정하는 것부터 시작하세요.

# --hints--

`printCharacters` 함수를 만들어야 합니다.

```js
assert.isFunction(printCharacters);
```

`printCharacters` 함수는 `str`라는 하나의 매개변수를 가져야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
