---
id: 6920d189e36b6b60a6330fee
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이제 루프 안에서 `char` 변수를 콘솔에 출력하세요.

# --hints--

루프 안에 `console.log` 문장이 있고 `char`가 인수로 있어야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)[1]

assert.match(loopCodeBlock, /console\.log\(char\)/);
```

`printCharacters` 함수가 인수 `'test'`와 함께 호출되면, 함수 안의 `console.log()` 문장은 다음 인수들로 네 번 호출되어야 합니다: `'t'`, `'e'`, `'s'`, `'t'`. 대신 다음 인수들로 호출되었습니다: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
