---
id: 6920c9cda1c7f056f75e4b1a
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이제 첫 번째 `for...of` 반복문을 작성할 시간입니다. 이전 수업에서 `for...of` 반복문을 사용해 반복 가능한 값(예: 문자열이나 배열)을 순회할 수 있다는 것을 배웠습니다.

다음은 `for...of` 반복문의 예시입니다:

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

함수 안에 빈 코드 블록이 있는 `for...of` 반복문을 추가하세요. 이 반복문은 `str` 인수의 모든 문자를 순회해야 합니다.

# --hints--

`printCharacters` 함수 안에는 `for...of` 반복문이 있어야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

`for...of` 반복문은 함수의 `str` 인수를 순회해야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
