---
id: 6921e10ae2681205b516fbe6
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

이제 `count`를 적절히 증가시키는 로직을 루프에 추가할 차례입니다.

루프 안에서 조건문을 사용해 변수 `count`가 변수 `1`와 같으면 `word`에 `match`를 더하세요. 그렇지 않으면 `count`의 값은 변경하지 마세요.

# --hints--

`getMatchedWordCount([], "bar")`, `getMatchedWordCount(["foo"], "bar")`, `getMatchedWordCount(["foo", "bar", "foo"], "bar")`, `getMatchedWordCount(["foo", "bar", "foo"], "foo")`를 순서대로 호출하면 각각 `"number"` 타입의 값을 반환해야 하는데, 대신 `--fcc-actual--` 타입의 값을 반환했습니다.

```js
const obj = { method: (arg1, arg2) => typeof getMatchedWordCount(arg1, arg2) };
const spy = __helpers.spyOn(obj, 'method');
obj.method([], "bar");
obj.method(["foo"], "bar");
obj.method(["foo", "bar", "foo"], "bar");
obj.method(["foo", "bar", "foo"], "foo");
assert.deepEqual(spy.returns, ["number", "number", "number", "number"]);
spy.restore();
```

`getMatchedWordCount([], "bar")`는 `0`를 반환해야 하는데, `--fcc-actual--`를 반환했습니다.

```js
const resultString = String(getMatchedWordCount([], "bar"));
assert.strictEqual(resultString, "0");
```

`getMatchedWordCount(["foo"], "bar")`는 `0`를 반환해야 하는데, `--fcc-actual--`를 반환했습니다.

```js
const resultString = String(getMatchedWordCount(["foo"], "bar"));
assert.equal(resultString, "0");
```

`getMatchedWordCount(["foo", "bar", "foo"], "bar")`는 `1`를 반환해야 하는데, `--fcc-actual--`를 반환했습니다.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "bar"));
assert.equal(resultString, "1");
```

`getMatchedWordCount(["foo", "bar", "foo"], "foo")`는 `2`를 반환해야 하는데, `--fcc-actual--`를 반환했습니다.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "foo"));
assert.equal(resultString, "2");
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
