---
id: 6921d688b8d58581d34459b9
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

다음 몇 단계에서는 문자열 배열에서 특정 문자열이 몇 번 나타나는지 세는 함수를 만들 것입니다.

먼저, `getMatchedWordCount`라는 이름의 빈 함수를 `sentence`과 `match` 매개변수를 순서대로 포함하여 정의하세요.

# --hints--

`getMatchedWordCount` 함수를 만들어야 합니다.

```js
assert.isFunction(getMatchedWordCount);
```

`getMatchedWordCount` 함수는 두 개의 매개변수 `sentence`와 `match`를 순서대로 가져야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
