---
id: 6921de8d2a111d0341cff702
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

좋아요, 함수가 `0`를 반환하지만 아직 아무것도 세지 않습니다.

이를 고치려면 `getMatchedWordCount` 안에 `for...of` 반복문을 만들어 빈 코드 블록으로 `sentence`의 각 단어를 순회하세요.

# --hints--

`getMatchedWordCount` 함수는 코드 블록 안에 `for...of` 반복문이 있어야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

`for...of` 반복문은 함수의 `sentence` 인수를 순회해야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
--fcc-editable-region--
  
--fcc-editable-region--
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
