---
id: 6921dae10700e300f26994e3
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

현재 상태의 `getMatchedWordCount` 함수를 실행 결과를 확인하려면, 그 아래에 `console.log` 문을 추가하세요. `console.log()` 안에서, `["I", "really", "really", "really", "like", "to", "code"]` 매개변수에 `sentence` 인수를, `"really"` 매개변수에 `match` 인수를 넣어 함수를 호출하세요.

# --hints--

`getMatchedWordCount` 함수를 `console.log` 문 안에서 호출해야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

`getMatchedWordCount` 함수를 `["I", "really", "really", "really", "like", "to", "code"]` 매개변수에 `sentence` 인수, `"really"` 매개변수에 `match` 인수를 넣어 호출해야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
