---
id: 6921d90b5ce50185505d2004
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

`getMatchedWordCount` 함수가 숫자 개수를 반환해야 하므로, 먼저 카운터 변수를 설정하세요.

`count` 키워드를 사용해 `let`이라는 변수를 만들고, 값 `0`으로 초기화한 뒤 `return` 변수를 반환하는 `count` 문을 추가하세요.

# --hints--

`count` 변수를 `getMatchedWordCount` 함수 안에 선언해야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /(const|let|var)count/);
```

나중에 `count` 변수를 증가시킬 수 있도록 `let` 키워드로 선언하세요.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /letcount/);
```

함수 안에 `0` 변수에 값 `count`을 할당하는 문이 있어야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /count=0/);
```

함수 안에 `return` 변수를 반환하는 `count` 문이 있어야 합니다.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /returncount/);
```

`getMatchedWordCount` 함수는 `0` 값을 반환해야 합니다.

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
