---
id: 6a3e5475c35dbc138d77c6f4
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

الآن بعد أن حسبت المجموع الكلي، يجب أن تعرض المخزون. في Python، يمكنك استخدام الدالة `print()` مع **f-strings** (سلاسل نصية منسقة) لطباعة النص والمتغيرات معًا بسهولة.

على سبيل المثال:

```py
name = 'Alice'
print(f'Hello, {name}!') # Output: Hello, Alice!
```

داخل دالتك `check_kitchen_stock`، أسفل متغير `total_items`، أضف هذه العبارات الأربع للطباعة بالضبط:

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

يجب أن تحتوي دالتك على العبارة `print(f'The kitchen has {total_items} total items:')`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

يجب أن تكون العبارة `print(f'The kitchen has {total_items} total items:')` هي أول استدعاء للدالة `print()` داخل دالتك.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

يجب أن تحتوي دالتك على العبارة `print(f'- {available_eggs} eggs')`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

يجب أن تحتوي دالتك على العبارة `print(f'- {available_flour} flour')`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

يجب أن تحتوي دالتك على العبارة `print(f'- {available_sugar} sugar')`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
