---
id: 6a3e71a0a3a44260e675b157
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

لحل لغز عدم تغير عدد البيض وطهي الإفطار فعليًا، أنشئ دالة جديدة. فكر في هذه الدالة كأنها وصفتك، التي تستدعي دوالًا أخرى لإنجاز المهمة.

أسفل عبارة `print(available_eggs)`، عرّف دالة فارغة باسم `make_fried_egg` تأخذ `available_eggs` كمعلمتها الوحيدة.

استخدم الكلمة المفتاحية `pass` داخل جسم الدالة في الوقت الحالي.

# --hints--

يجب أن تعرف دالة باسم `make_fried_egg`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

يجب أن تحتوي دالة `make_fried_egg` على `available_eggs` كمعلمة.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
