---
id: 6a3e5ccb99fd6053af9032a1
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

الآن بعد أن عرفت ما يوجد في مطبخك، يمكنك استخدام تلك العناصر للطهي. للقيام بذلك، أنشئ دالة جديدة تقبل معلمات. فكر في المعلمات كعناصر نائب للبيانات التي تمررها إلى الدالة عند استدعائها.

إليك دالة تقبل معلمتين، `a` و `b`، وتطبع مجموعهما على الطرفية.

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

أسفل دالتك `check_kitchen_stock`، عرّف دالة فارغة باسم `use_eggs` تقبل معلمتين: `available_eggs` و `eggs_to_use`، بهذا الترتيب.

لاحظ أنه إذا أعطيت معلمة نفس اسم متغير عام (مثل `available_eggs`)، ستستخدم الدالة معلمتها المحلية `available_eggs` بدلاً من المتغير العام. يُعرف هذا بتظليل الاسم إذ "يظلل" المتغير المحلي المتغير العام.

استخدم الكلمة المفتاحية `pass` داخل جسم الدالة في الوقت الحالي.

# --hints--

يجب أن تعرف دالة باسم `use_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

يجب أن تحتوي دالة `use_eggs` على `available_eggs` و `eggs_to_use` كمعلمات لها.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
