---
id: 6a3e524fc3394ef52ed5b797
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

تحتوي الدوال على نطاق محلي خاص بها. هذا يعني أن المتغيرات التي تُنشأ داخل دالة تشبه المكونات المحفوظة في مخزن مطبخ معين: لا يمكن الوصول إليها إلا داخل تلك الدالة المحددة.

مع ذلك، الدوال ليست معزولة تمامًا. يمكنها قراءة المتغيرات العامة المعرفة خارج نطاقها. يمكنها استخدام تلك المتغيرات لإجراء حسابات، اتخاذ قرارات، أو إنشاء متغيرات محلية جديدة.

داخل دالتك `check_kitchen_stock`، استبدل الكلمة المفتاحية `pass` بمتغير جديد اسمه `total_items` وقم بتعيين مجموع `available_eggs` و`available_flour` و`available_sugar` له.

# --hints--

يجب تعريف متغير اسمه `total_items` داخل دالتك `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

يجب أن يكون متغير `total_items` هو مجموع `available_eggs` و`available_flour` و`available_sugar`.

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
