---
id: 6a3e70b94f67d9ee3ecd8486
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

لقد استدعيت للتو `use_eggs(available_eggs, 1)` لاستخدام بيضة واحدة. هذا يعني أنه يجب أن يتبقى لديك `0` بيضات في مطبخك الآن، أليس كذلك؟

ستعرف ذلك! في أسفل كودك، استخدم الدالة `print()` لطباعة قيمة المتغير `available_eggs`.

عندما تشغل الكود، ستلاحظ شيئًا مفاجئًا: الناتج لا يزال `1`! لماذا لم يتغير عدد البيض؟ ستستكشف هذا اللغز في الخطوات التالية.

# --hints--

يجب طباعة المتغير `available_eggs` في نهاية كودك.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
