---
id: 6a3e73eabddd614f91cc235a
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

الآن ستُحل اللغز الذي ذكرناه سابقًا! عندما أعادت دالة `use_eggs` العدد المتبقي من البيض، لم تحفظ هذه القيمة الجديدة في أي مكان. اختفت ببساطة في الفراغ!

لتحديث المخزون فعليًا، يجب التقاط القيمة المعادة عن طريق تعيينها إلى متغير.

داخل دالة `make_fried_egg`، استخدم عبارة `if/else` للتحقق من المتغير `has_enough_eggs`.

- إذا كان `has_enough_eggs` يساوي `True`، أعد تعيين المتغير `available_eggs` إلى نتيجة `use_eggs(available_eggs, 1)` واطبع الرسالة `Made a fried egg. Yummy!`.

- وإلا، ضمن كتلة `else`، اطبع الرسالة `Could not make a fried egg. Not enough eggs!`.

# --hints--

يجب أن تحتوي على عبارة `if` وتستخدم `has_enough_eggs` كشرط لها.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

داخل كتلة `if`، يجب أن تعيد تعيين `available_eggs` إلى نتيجة `use_eggs(available_eggs, 1)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

يجب أن تطبع الرسالة `Made a fried egg. Yummy!` عندما يكون `has_enough_eggs` يساوي `True`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

يجب أن تحتوي على كتلة `else` تطبع الرسالة `Could not make a fried egg. Not enough eggs!` عندما يكون `has_enough_eggs` يساوي `False`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
