---
id: 6a3e58891e1ff13d07f4ab48
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

يقرأ Python كتلة `def` ويتذكر الدالة، لكنه لا يُنفذ الكود داخلها فورًا.

لتنفيذ الكود ورؤية الناتج، يجب عليك **استدعاء** الدالة. افعل ذلك بكتابة اسم الدالة متبوعًا بالأقواس العادية.

على سبيل المثال:

```py
my_function()
```

في أسفل الكود الخاص بك، خارج الدالة، استدعِ دالة `check_kitchen_stock` الخاصة بك.

# --hints--

يجب استدعاء دالة `check_kitchen_stock` في نهاية الكود. تأكد من أنها ليست متداخلة (غير مُزاحة)، لكي تُنفذ على المستوى العام.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
