---
id: 6a3e6fd08125dfa615cb431b
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

الآن بعد أن أتممت دالة `use_eggs`، يمكنك اختبارها. عند استدعاء دالة تتطلب **معلمات**، تمرر القيم المحددة داخل الأقواس العادية. تُعرف هذه القيم باسم **المعلمات**.

إليك مثالًا على كيفية استدعاء دالة مع معلمات:

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

في هذا المثال، السطران الأولان هما تعريف الدالة، حيث `a` و `b` هما المعلمتان. السطر الأخير، `calculate_sum(3, 1)`، هو استدعاء الدالة. القيم `3` و `1` هما المعلمتان. يتم مطابقتها مع المعلمات حسب الموضع، لذا تُعيَّن القيمة `3` إلى `a` والقيمة `1` تُعيَّن إلى `b`.

في أسفل الكود الخاص بك، استدعِ دالة `use_eggs` ومرر لها معلمتين: المتغير العام `available_eggs` والرقم `1` (الذي يمثل عدد البيض الذي تريد استخدامه).

# --hints--

يجب أن تستدعي دالة `use_eggs` وتمرر `available_eggs` و `1` كمعلمات.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
