---
id: 6a3e6c0eb5d050ce07781709
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

إذا كانت نتيجة شرط `if` تساوي `False` (أي أن لديك ما يكفي من البيض)، ستتخطى الدالة كتلة `if` وتستمر في تنفيذ بقية الكود.

الآن عليك معالجة بياناتك وإرجاع قيمة جديدة محسوبة.

خارج كتلة `if` (على نفس مستوى المسافة البادئة لكلمة `if`)، استخدم استدعاء `print()` وسلسلة نصية بصيغة f-string لطباعة `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`.

ثم، في السطر التالي، استخدم `return` لإرجاع العدد المتبقي من البيض بطرح `eggs_to_use` من `available_eggs`.

# --hints--

يجب أن تحتوي دالة `use_eggs` على جملة طباعة تطبع `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

يجب أن ترجع العدد المتبقي من البيض بطرح `eggs_to_use` من `available_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

يجب أن ترجع `available_eggs - eggs_to_use` بعد جملة `print`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
