---
id: 6a3e7284e89ce3367e709331
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

قبل صنع بيضة مقلية، يجب أولاً التحقق من وجود المكونات المطلوبة. يمكنك فعل ذلك بمقارنة كمية البيض المتوفرة مع الكمية اللازمة للوصفة.

داخل دالة `make_fried_egg`، استبدل الكلمة المفتاحية `pass` بإنشاء متغير اسمه `has_enough_eggs` وتعيين نتيجة التحقق مما إذا كان `available_eggs` أكبر من أو يساوي `1`.

# --hints--

يجب تعيين `has_enough_eggs` إلى نتيجة التحقق مما إذا كانت قيمة `available_eggs` أكبر من أو تساوي `1`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_variables("has_enough_eggs")[0]
assert node.is_equivalent("has_enough_eggs = available_eggs >= 1") or node.is_equivalent("has_enough_eggs = 1 <= available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
