---
id: 6a3e6ad6b9554e31dcc30cd6
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

إذا لم يكن لديك بيض كافٍ، فلا تحاول الطهي! تحتاج إلى طريقة لإيقاف الدالة من الاستمرار في التنفيذ.

يمكنك استخدام الكلمة المفتاحية `return` لهذا الغرض. عندما تصل الدالة إلى جملة `return`، تتوقف عن التنفيذ فورًا وترجع القيمة إلى المكان الذي تم استدعاؤها منه.

على سبيل المثال:

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

داخل كتلة `if` الخاصة بك، أسفل جملة الطباعة، استخدم الكلمة المفتاحية `return` لإرجاع قيمة `available_eggs`.

# --hints--

يجب أن ترجع `available_eggs` داخل جملة `if` الخاصة بك.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")
`)
})
```

يجب أن ترجع `available_eggs` بعد جملة `print` داخل كتلة `if` الخاصة بك.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs").find_ifs()[0]
assert node.is_ordered("print('The kitchen does not have enough eggs.')", "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
