---
id: 6a3e67fc5ef4045201c28219
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

داخل دالة `use_eggs`، عليك التحقق مما إذا كان لديك بيض كافٍ للطهي. يمكنك فعل ذلك باستخدام جملة `if`.

جملة `if` تُقيّم شرطًا. إذا كان الشرط صحيحًا، يُنفذ الكود داخل كتلة الـ `if`.

استبدل الكلمة المفتاحية `pass` بجملة `if` تتحقق مما إذا كانت قيمة `eggs_to_use` أكبر من (`>`) قيمة `available_eggs`. داخل كتلة الـ `if`، اطبع السلسلة النصية `The kitchen does not have enough eggs.`.

# --hints--

يجب أن تحتوي على جملة `if` تتحقق مما إذا كان `eggs_to_use > available_eggs`.

```js
({
  test: () => runPython(`
  cond = _Node(_code).find_function("use_eggs").find_ifs()[0].find_conditions()[0]
  assert cond.is_equivalent("eggs_to_use > available_eggs") or cond.is_equivalent("available_eggs < eggs_to_use")`)
})
```

يجب أن تطبع `The kitchen does not have enough eggs.` داخل جملة الـ `if`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_call("print('The kitchen does not have enough eggs.')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
