---
id: 6a3e4ae26f2b7e4b5b026cf8
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

حان الوقت لإنشاء دالة للتحقق مما لديك في المطبخ. في Python، تُعرّف الدوال باستخدام الكلمة المفتاحية `def`، تليها اسم الدالة، ثم الأقواس العادية `()`، ونقطتان رأسيتان `:`.

إليك مثالًا على كيفية تعريف دالة بسيطة:

```py
def my_function():
    pass
```

عرّف دالة فارغة باسم `check_kitchen_stock()`. استخدم الكلمة المفتاحية `pass` داخل جسم الدالة الآن حتى لا ترمي Python خطأ نحويًا.

تذكر أن Python تعتمد على المسافة البادئة لتعريف كتل الكود. لذا تأكد من إضافة المسافة البادئة لجسم الدالة.

# --hints--

يجب أن تعرف دالة باسم `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
