---
id: 6a3e2b9b7eeca084dd4fb857
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

مرحبًا بك في المطبخ! تمامًا كما يعتمد الطاهي على الوصفات لتحضير وجبة، ستعتمد أنت على الدوال لبناء البرمجيات. في هذه الورشة، ستتعلم كيفية كتابة دوال Python قابلة لإعادة الاستخدام وإتقان النطاق، مع فهم دقيق لأي مكونات (المتغيرات) متاحة في كودك في أي وقت.

للبدء، عرّف بعض البيانات الأساسية لمخزون المطبخ. أنشئ ثلاثة متغيرات: `available_eggs` و`available_flour` و`available_sugar`، وقم بتعيين القيم `1` و`2` و`3` لها على التوالي.

# --hints--

عرّف متغيرًا باسم `available_eggs` وقم بتعيين القيمة `1` له.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

عرّف متغيرًا باسم `available_flour` وقم بتعيين القيمة `2` له.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

عرّف متغيرًا باسم `available_sugar` وقم بتعيين القيمة `3` له.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
