---
id: 6a3ea8a55317ba7a08e21d5b
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

يمكنك الآن أخيرًا الإجابة على السؤال السابق: لماذا بقي العدد العالمي للبيض عند `1`؟ عندما ترجع دالة قيمة، لا تُعيد تلقائيًا كتابة المتغيرات الموجودة لديك. يجب عليك حفظ تلك القيمة الجديدة صراحةً!

أنت تستخدم المتغير `available_eggs` داخل الدالة `make_fried_egg`، لكنه ليس نفس المتغير المُعرف في أعلى كودك. المتغير داخل الدالة هو متغير محلي لا يوجد إلا أثناء تشغيل الدالة. عند انتهاء الدالة، يُتخلص من المتغير المحلي ويبقى المتغير العالمي `available_eggs` دون تغيير.

لتحديث مخزون مطبخك الفعلي، اكتب فوق المتغير العالمي `available_eggs` بعدد البيض الجديد الذي ترجعه دالتك `make_fried_egg`.

في أسفل كودك، استدعِ الدالة `make_fried_egg`، مرر المتغير `available_eggs` كمعلمة، ثم عيّن النتيجة مرة أخرى إلى المتغير العالمي `available_eggs`.

# --hints--

يجب إعادة تعيين المتغير `available_eggs` إلى نتيجة `make_fried_egg(available_eggs)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
