---
id: 61fd7a160ed17960e971f28b
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

في عنصر `tr` الرابع، أضف عنصر `th` بالنص `Total Assets`. غلف النص `Assets` داخل عنصر `span` مع تعيين الخاصية `class` إلى `sr-only`.

بعد ذلك، أضف ثلاثة عناصر `td` بالنصوص التالية (بالترتيب): `$579`، `$736`، `$809`. اجعل الخاصية `class` في عنصر `td` الثالث تساوي `current`.

# --hints--

يجب أن يحتوي عنصر `tr` الرابع على عنصر `th`.

```js
assert.exists(
  document
    .querySelector('tbody')
    ?.querySelectorAll('tr')?.[3]
    ?.querySelector('th')
);
```

يجب أن يحتوي عنصر `th` على النص `Total Assets`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableHeader = tableRow.querySelector('th');
assert.strictEqual(tableHeader.innerText.trim(), 'Total Assets');
```

يجب أن تغلف النص `Assets` داخل عنصر `span`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const span = tableRow.querySelector('th > span');
assert.strictEqual(span.textContent.trim(), 'Assets');
```

يجب أن يحتوي عنصر `span` على الخاصية `class` مع القيمة `sr-only`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const span = tableRow.querySelector('th > span');
assert.isTrue(span?.classList?.contains('sr-only'));
```

يجب أن يكون لديك ثلاثة عناصر `td`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableDatas = tableRow.querySelectorAll('td');
assert.lengthOf(tableDatas, 3);
```

يجب أن يحتوي عنصر `td` الأول على النص `$579`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '$579');
```

يجب أن يحتوي عنصر `td` الثاني على النص `$736`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[1];
assert.strictEqual(tableData?.textContent.trim(), '$736');
```

يجب أن يحتوي عنصر `td` الثالث على النص `$809`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[2];
assert.strictEqual(tableData.textContent.trim(), '$809');
```

يجب أن تحتوي الخاصية `class` في عنصر `td` الثالث على القيمة `current`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[2];
assert.isTrue(tableData?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
--fcc-editable-region--
                
--fcc-editable-region--
              </tr>
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
