---
id: 61fd66c687e610436494c6f1
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

داخل عنصر `section`، أضف عنصر `h1` يحتوي على عنصر `span` مسكن.

# --hints--

يجب أن يحتوي عنصر `section` على عنصر `h1`.

```js
assert.strictEqual(
  document.querySelector('section')?.children?.[0]?.localName,
  'h1'
);
```

يجب أن يحتوي عنصر `h1` على عنصر `span`.

```js
assert.strictEqual(
  document.querySelector('h1')?.children?.[0]?.localName,
  'span'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
