---
id: 61fd71d596e8f253b9408b39
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

يُستخدم العنصر `tr` للدلالة على صف في الجدول. أضف عنصر `tr` داخل عنصر `thead` الخاص بك. في عنصر `tr` الجديد، أضف عنصر `td`، يليه ثلاثة عناصر `th`.

يشير العنصر `td` إلى خلية بيانات، بينما يشير العنصر `th` إلى خلية قسم علوي.

# --hints--

يجب أن يحتوي عنصر `thead` على عنصر `tr`.

```js
assert.strictEqual(
  document.querySelector('thead')?.children?.[0]?.localName,
  'tr'
);
```

يجب أن يحتوي عنصر `tr` على عنصر `td` كأول عنصر فرعي.

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[0]?.localName,
  'td'
);
```

يجب أن يحتوي عنصر `tr` على ثلاثة عناصر `th` بعد عنصر `td`.

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[1]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[2]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[3]?.localName,
  'th'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
--fcc-editable-region--
              
--fcc-editable-region--
            </thead>
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
