---
id: 61fd719788899952e67692b9
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

تُستخدم عناصر `thead` و `tbody` لتحديد أي جزء من جدولك هو القسم العلوي، وأي جزء يحتوي على البيانات أو المحتوى الأساسي.

أضف عنصرَي `thead` و `tbody` إلى أول `table` لديك، أسفل عنصر `caption`.

# --hints--

يجب أن يحتوي أول عنصر `table` على عنصر `thead`.

```js
const table = document.querySelectorAll('table')?.[0];
assert.isNotNull(table?.querySelector('thead'));
```

يجب أن يحتوي أول عنصر `table` على عنصر `tbody`.

```js
const table = document.querySelectorAll('table')?.[0];
const tbody = table?.querySelector('tbody');
assert.isNotNull(tbody);
```

يجب أن يكون عنصر `thead` مباشرة أسفل عنصر `caption`.

```js
assert.strictEqual(
  document.querySelector('caption')?.nextElementSibling?.localName,
  'thead'
);
```

يجب أن يكون عنصر `tbody` مباشرة أسفل عنصر `thead`.

```js
assert.strictEqual(
  document.querySelector('thead')?.nextElementSibling?.localName,
  'tbody'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            
--fcc-editable-region--
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
