---
id: 61fd5a93fd62bb35968adeab
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ابدأ مشروعك بإضافة عنصر `link` داخل عنصر `head` لورقة الأنماط الخاصة بك.

# --hints--

يجب أن تحتوي شفرتك على عنصر `link`.

```js
assert.match(code, /<link/);
```

يجب أن يكون عنصر `link` داخل عنصر `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بقيمة `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
