---
id: 61fd67a656743144844941cb
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

يُعلن قارئو الشاشة عن عناصر HTML بناءً على تدفق المستند. في النهاية، نريد أن يحتوي كشف الميزانية على عنوان "Balance Sheet" وعنوان فرعي "AcmeWidgetCorp". مع ذلك، هذا الترتيب لا يكون منطقيًا إذا أعلنه قارئ الشاشة.

أضف إلى عنصر `span` الموجود خاصية `class` بقيمة `flex`، وأضف داخله عنصرين `span` جديدين. اجعل النص في الأول `AcmeWidgetCorp`. اجعل النص في الثاني `Balance Sheet`. ستستخدم CSS لعكس ترتيب النص على الصفحة، لكن ترتيب HTML سيكون أكثر منطقية لقارئ الشاشة.

# --hints--

يجب أن يحتوي عنصر `span` الموجود على خاصية `class` بقيمة `flex`.

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

يجب أن يحتوي عنصر `span` الموجود على عنصرين جديدين `span` داخله.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

يجب ألا تحتوي عناصر `span` الجديدة على خاصية `class`.

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

يجب أن يحتوي أول عنصر `span` جديد على النص `AcmeWidgetCorp`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

يجب أن يحتوي ثاني عنصر `span` جديد على النص `Balance Sheet`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
            
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
