---
id: 61fd75ea7f663457612dba02
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

داخل كل عنصر جديد من عناصر `th`، قم بتسكين عنصر `span` مع تعيين الخاصية `class` إلى `sr-only year`. أعطها النص التالي (بالترتيب): `2019`، `2020`، و`2021`.

عيّن الخاصية `class` للعنصر الثالث من نوع `th` إلى `current`.

اترك عنصر `td` فارغًا. هذا العنصر موجود فقط لضمان أن يحتوي الجدول على تخطيط بأربعة أعمدة وربط الأقسام العلوية بالأعمدة الصحيحة.

# --hints--

يجب أن يحتوي كل عنصر من عناصر `th` على عنصر `span`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.lengthOf(th?.children, 1);
  assert.strictEqual(th?.children?.[0]?.localName, 'span');
});
```

يجب أن تحتوي كل عناصر `span` الجديدة على الخاصية `class` المعينة إلى `sr-only year`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.isTrue(th?.children?.[0]?.classList?.contains('sr-only'));
  assert.isTrue(th?.children?.[0]?.classList?.contains('year'));
});
```

يجب أن يحتوي أول عنصر `span` على النص `2019`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent.trim(),
  '2019'
);
```

يجب أن يحتوي ثاني عنصر `span` على النص `2020`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent.trim(),
  '2020'
);
```

يجب أن يحتوي ثالث عنصر `span` على النص `2021`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent.trim(),
  '2021'
);
```

يجب أن يحتوي العنصر الثالث من نوع `th` على الخاصية `class` المعينة إلى `current`.

```js
assert.isTrue(
  document
    .querySelector('table')
    ?.querySelectorAll('th')?.[2]
    ?.classList?.contains('current')
);
```

يجب أن يكون عنصر `td` فارغًا.

```js
const table = document.querySelector('table');
const tableData = table.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '');
assert.isEmpty(tableData?.children);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
--fcc-editable-region--
                <th></th>
                <th></th>
                <th></th>
--fcc-editable-region--
              </tr>
            </thead>
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
