---
id: 61fd778081276b59d59abad6
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

في أول `tr` لديك، أضف عنصر `th` بالنص `Cash This is the cash we currently have on hand.`. غطِ كل هذا النص ما عدا `Cash ` داخل عنصر `span` مع تعيين `class` إلى `description`.

بعد ذلك، أضف ثلاثة عناصر `td` بالنصوص التالية (بالترتيب): `$25`، `$30`، `$28`. أعطِ العنصر الثالث `td` خاصية `class` بقيمة `current`.

# --hints--

يجب أن يحتوي أول `tr` على عنصر `th`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.isNotNull(tableRow?.querySelector('th'));
```

يجب أن يحتوي عنصر `th` على النص `Cash This is the cash we currently have on hand.`

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Cash This is the cash we currently have on hand.'
);
```

يجب أن تغلف النص `This is the cash we currently have on hand.` داخل عنصر `span`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'This is the cash we currently have on hand.'
);
```

يجب أن يحتوي عنصر `span` على الخاصية `class` بقيمة `description`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

يجب أن يكون لديك ثلاثة عناصر `td`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

يجب أن يحتوي أول عنصر `td` على النص `$25`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$25');
```

يجب أن يحتوي ثاني عنصر `td` على النص `$30`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$30');
```

يجب أن يحتوي ثالث عنصر `td` على النص `$28`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$28');
```

يجب أن يحتوي ثالث عنصر `td` على الخاصية `class` بقيمة `current`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
--fcc-editable-region--
                
--fcc-editable-region--
              </tr>
              <tr class="data"></tr>
              <tr class="data"></tr>
              <tr class="total"></tr>
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
