---
id: bad87fee1348bd9aec908857
title: استخدام التعليقات لتوضيح الشفرة
challengeType: 0
forumTopicId: 18347
dashedName: use-comments-to-clarify-code
---

# --description--

عندما نبدأ باستخدام jQuery، سنُعدّل عناصر HTML دون الحاجة إلى تغييرها فعليًا في HTML.

تأكد من أن الجميع يعرفون أنه لا يجب تعديل أي جزء من هذه الشفرة مباشرة.

تذكر أنه يمكنك بدء تعليق بـ `<!--` وإنهاؤه بـ `-->`

أضف تعليقًا في أعلى ملف HTML يقول `Code below this line should not be changed`

# --hints--

يجب أن تبدأ تعليقًا بـ `<!--` في أعلى ملف HTML.

```js
assert.match(code,(/^\s*<!--/));
```

يجب أن يحتوي تعليقك على النص `Code below this line should not be changed`.

```js
assert.match(code,(/<!--(?!(>|->|.*-->.*this line))\s*.*this line.*\s*-->/gi));
```

يجب أن تُغلق تعليقك بـ `-->`.

```js
assert.match(code,(/-->.*\n+.+/g));
```

يجب أن يكون لديك نفس عدد بدايات التعليقات ونهاياتها.

```js
assert.equal(code.match(/<!--/g).length,code.match(/-->/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<!-- Code below this line should not be changed -->
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
