---
id: bad87fee1348bd9aec908849
title: أضف عناصر داخل صناديق Bootstrap الخاصة بك
challengeType: 0
forumTopicId: 16636
dashedName: add-elements-within-your-bootstrap-wells
---

# --description--

الآن نحن داخل عدة عناصر `div` في كل عمود من صفنا. هذا هو العمق الذي نحتاجه فقط. الآن يمكننا إضافة عناصر `button` الخاصة بنا.

قم بتسكين ثلاثة عناصر `button` داخل كل عنصر `div` لديك يحمل اسم الصنف `well`.

# --hints--

يجب تسكين ثلاثة عناصر `button` داخل كل عنصر `div` لديك بالصنف `well`.

```js
const buttonOne = document.querySelectorAll('div.well')?.[0];
const buttonOneChildren = buttonOne?.querySelectorAll(`:scope ${'button'}`);
assert.lengthOf(buttonOneChildren,3);

const buttonTwo = document.querySelectorAll('div.well')?.[1]
const buttonTwoChildren = buttonTwo?.querySelectorAll(`:scope ${'button'}`);
assert.lengthOf(buttonTwoChildren,3);
```

يجب أن يكون لديك ما مجموعه 6 عناصر `button`.

```js
assert.lengthOf(document.querySelectorAll('button'), 6);
```

يجب أن تحتوي كل عناصر `button` على وسم الإغلاق.

```js
assert.match(code,/<\/button>/g);
assert.match(code,/<button/g);
assert.equal(code.match(/<\/button>/g).length,code.match(/<button/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">



      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">



      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
  </div>
</div>
```
