---
id: bad87fee1348bd9aec908846
title: إنشاء عنوان Bootstrap
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

لننشئ شيئًا من الصفر لممارسة مهاراتك في HTML وCSS وBootstrap.

سننشئ مساحة لعب jQuery، والتي سنستخدمها قريبًا في تحديات jQuery الخاصة بنا.

ابدأ بإنشاء عنصر `h3`، مع النص `jQuery Playground`.

لون عنصر `h3` باستخدام صنف Bootstrap `text-primary`، ووسّطه باستخدام صنف Bootstrap `text-center`.

# --hints--

يجب أن تضيف عنصر `h3` إلى صفحتك.

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

يجب أن يحتوي عنصر `h3` على وسم الإغلاق.

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

يجب أن يكون عنصر `h3` ملونًا بتطبيق الصنف `text-primary`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

يجب أن يكون عنصر `h3` في الوسط بتطبيق الصنف `text-center`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

يجب أن يحتوي عنصر `h3` على النص `jQuery Playground`.

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```
