---
id: bad87fee1348bd9aec908853
title: إضافة خاصيات id إلى عناصر Bootstrap
challengeType: 0
forumTopicId: 16639
dashedName: add-id-attributes-to-bootstrap-elements
---

# --description--

تذكّر أنه بالإضافة إلى خاصيات الصنف، يمكنك إعطاء كل عنصر من عناصر الصفحة خاصية `id`.

يجب أن يكون كل id فريدًا لعنصر محدد ويُستخدم مرة واحدة فقط في الصفحة.

لنُعطِ كل عنصر `div` من صنف `well` خاصية id فريدة.

تذكّر أنه يمكنك إعطاء عنصر خاصية id بهذه الطريقة:

```html
<div class="well" id="center-well">
```

أعطِ الـ well على اليسار الخاصية id بقيمة `left-well`. وأعطِ الـ well على اليمين الخاصية id بقيمة `right-well`.

# --hints--

يجب أن يكون للـ `well` الأيسر الخاصية id بقيمة `left-well`.

```js
const column = document.querySelectorAll('.col-xs-6')[0];
const leftWall = column?.querySelectorAll(`:scope ${'#left-well'}`);
assert.lengthOf(leftWall,1)
```

يجب أن يكون للـ `well` الأيمن الخاصية id بقيمة `right-well`.

```js
const column = document.querySelectorAll('.col-xs-6')[1];
const rightWall = column?.querySelectorAll(`:scope ${'#right-well'}`);
assert.lengthOf(rightWall,1)
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well" id="left-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well" id="right-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
  </div>
</div>
```
