---
id: bad87fee1348bd9aec908847
title: قسم صف Bootstrap الخاص بك
challengeType: 0
forumTopicId: 18306
dashedName: split-your-bootstrap-row
---

# --description--

الآن بعد أن أصبح لدينا صف Bootstrap، قسّمه إلى عمودين لاستيعاب عناصرنا.

أنشئ عنصرين `div` داخل صفك، كلاهما يحمل الصنف `col-xs-6`.

# --hints--

يجب أن تكون عنصران `div class="col-xs-6"` مسكنين داخل عنصر `div class="row"` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

يجب أن تحتوي كل عناصر `div` على وسم الإغلاق.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length , code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">


  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6"></div>
    <div class="col-xs-6"></div>
  </div>
</div>
```
