---
id: bad87fee1348bd9aec908848
title: إنشاء صناديق Bootstrap
challengeType: 0
forumTopicId: 16825
dashedName: create-bootstrap-wells
---

# --description--

يحتوي Bootstrap على صنف يسمى `well` يمكنه خلق إحساس بصري بالعمق لأعمدتك.

قم بتسكين عنصر `div` واحد يحمل الصنف `well` داخل كل عنصر `div` من صنف `col-xs-6` لديك.

# --hints--

يجب إضافة عنصر `div` يحمل الصنف `well` داخل كل عنصر `div` من صنف `col-xs-6` لديك.

```js
const wells = document.querySelectorAll('div.col-xs-6 > div.well');
assert.lengthOf( wells,2 ); 
```

يجب أن يكون كلا عنصري `div` من صنف `col-xs-6` مسكنين داخل عنصر `div` من صنف `row`.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

يجب أن تحتوي كل عناصر `div` على وسم الإغلاق.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g)?.length , code.match(/<div/g)?.length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">

    </div>
    <div class="col-xs-6">

    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
  </div>
</div>
```
