---
id: bad87fee1348bd9aec908850
title: تطبيق نمط زر Bootstrap الافتراضي
challengeType: 0
forumTopicId: 16657
dashedName: apply-the-default-bootstrap-button-style
---

# --description--

لدى Bootstrap صنف زر آخر يسمى `btn-default`.

طبّق كلا الصنفين `btn` و `btn-default` على كل عنصر من عناصر `button` الخاصة بك.

# --hints--

يجب تطبيق الصنف `btn` على كل عنصر من عناصر `button` الخاصة بك.

```js
assert.lengthOf(document.querySelectorAll('.btn'),6);
```

يجب تطبيق الصنف `btn-default` على كل عنصر من عناصر `button` الخاصة بك.

```js
assert.lengthOf(document.querySelectorAll('.btn-default'), 6);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
      </div>
    </div>
  </div>
</div>
```
