---
id: bad87fee1348bd9aec908746
title: ضع صفحتنا داخل عنصر div من نوع container-fluid في Bootstrap
challengeType: 0
forumTopicId: 18198
dashedName: house-our-page-within-a-bootstrap-container-fluid-div
---

# --description--

لنضمن الآن أن كل المحتوى على صفحتك متجاوب مع الأجهزة المحمولة.

لنقم بتسكين عنصر `h3` داخل عنصر `div` يحمل الصنف `container-fluid`.

# --hints--

يجب أن يحمل عنصر `div` الصنف `container-fluid`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container-fluid'));
```

يجب أن يحتوي كل عنصر `div` على وسم إغلاق.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);

assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

يجب أن يكون عنصر `h3` مسكنًا داخل عنصر `div`.

```js
const divElement = document.querySelector('div');
const divChildren = divElement?.querySelectorAll(`:scope ${'h3'}`)
assert.lengthOf(divChildren,1);
```

# --seed--

## --seed-contents--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```

# --solutions--

```html
<div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
</div>
```
