---
id: 69babdf3f896711c11516373
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

خلال كل دورة محاكاة:

- يصل عدد معين من الحضور إلى البوابة (من `queue` الخاصة بها).
- تعالج البوابة الحضور بناءً على `capacity` الخاصة بها.
- إذا وصل عدد الحضور أكثر مما تستطيع البوابة التعامل معه، سيبقى بعضهم (تجاوز السعة).

الآن ستبني دالة تتعامل مع هذه المنطق لبوابة واحدة في دورة واحدة.

أنشئ دالة فارغة باسم `processGateFlow` تستقبل معلمتين:

- `gate`: كائن البوابة التي تتم معالجتها.
- `tickIndex`: موقع دورة المحاكاة الحالية في مصفوفة `queue`.

# --hints--

يجب أن تنشئ دالة باسم `processGateFlow`.

```js
assert.isFunction(processGateFlow);
```

يجب أن تحتوي دالتك `processGateFlow` على المعلمتين `gate` و `tickIndex`.

```js
const regex = __helpers.functionRegex('processGateFlow', ['gate', 'tickIndex']);
assert.match(__helpers.removeJSComments(processGateFlow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
