---
id: 69baafd6d683dd9921b75db3
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستستخدم JavaScript لمحاكاة تدفق الحضور في مهرجان موسيقي.

ستعمل مع مجموعتي بيانات تصف بوابات المهرجان خلال نوبات الصباح والليل:

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

يحتوي كل كائن بوابة على الخصائص التالية:

- `id`: سلسلة نصية تحدد البوابة.
- `capacity`: عدد الحضور الذين يمكن للبوابة معالجتهم في كل فترة زمنية.
- `queue`: مصفوفة من الأرقام تمثل عدد الحضور الذين يصلون إلى البوابة خلال فترة زمنية محددة.

لبدء العمل، أنشئ دالة فارغة باسم `initializeThroughput` تستقبل معلمة باسم `gates`.

# --hints--

يجب أن تنشئ دالة باسم `initializeThroughput`.

```js
assert.isFunction(initializeThroughput);
```

يجب أن تحتوي دالة `initializeThroughput` على معلمة `gates`.

```js
const regex = __helpers.functionRegex('initializeThroughput', ['gates']);
assert.match(__helpers.removeJSComments(initializeThroughput.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
