---
id: 69bac067cd15c4685d07906d
title: الخطوة 15
challengeType: 1
dashedName: step-15
---

# --description--

الآن بعد أن لديك دوال لمعالجة العلامات الدقيقة الفردية والتعامل مع تجاوز السعة، حان الوقت لإنشاء دالة تحكم لباب واحد. ستتعامل هذه الدالة مع كل الإجراءات الخاصة بباب واحد خلال علامة دقيقة واحدة من المحاكاة.

أنشئ دالة فارغة باسم `handleGateAtTick` مع المعلمات التالية:

- `gates`: المصفوفة الكاملة لكائنات الأبواب.
- `gate`: الباب الحالي الذي تتم معالجته.
- `tickIndex`: فهرس العلامة الدقيقة الحالي للمحاكاة.
- `throughputSummary`: كائن يتتبع إجمالي الحضور المعالج لكل باب.

# --hints--

يجب أن تنشئ دالة باسم `handleGateAtTick`.

```js
assert.isFunction(handleGateAtTick);
```

يجب أن تحتوي دالة `handleGateAtTick` على المعلمات `gates` و `gate` و `tickIndex` و `throughputSummary`.

```js
const regex = __helpers.functionRegex('handleGateAtTick', ['gates', 'gate', 'tickIndex', 'throughputSummary']);
assert.match(__helpers.removeJSComments(handleGateAtTick.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
