---
id: 69bac027df64d4433c708d44
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

عند إعادة توجيه تجاوز السعة، يجب إرسال الحضور إلى البوابة التالية في مصفوفة `gates`. للقيام بذلك، تحتاج إلى إيجاد فهرس البوابة التالية. عادةً، يمكنك إضافة 1 إلى الفهرس الحالي. لكن إذا كانت البوابة الحالية هي الأخيرة في المصفوفة، يجب العودة إلى البوابة الأولى.

يمكنك استخدام عامل باقي القسمة (`%`) للتعامل مع هذا:

```js
const nextIndex = (currentIndex + 1) % array.length;
```

يعمل هذا لأن عندما يكون `currentIndex + 1` مساويًا لطول المصفوفة `array.length`، تصبح النتيجة 0.

باستخدام هذا الأسلوب، أنشئ متغيرًا باسم `nextGateIndex` يخزن فهرس البوابة التالية في مصفوفة `gates`.

# --hints--

يجب أن تستخدم تقنية باقي القسمة للوصول إلى الفهرس التالي من `currentIndex` في مصفوفة `gates` وتعيينه إلى متغير باسم `nextGateIndex`. اتبع نفس نمط السطر الواحد في المثال.

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+nextGateIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'nextGateIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)nextGateIndex=\((currentIndex\+1|1\+currentIndex)\)%gates\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use the modulo technique to create a variable named `nextGateIndex` that stores the index of the next gate."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
--fcc-editable-region--

--fcc-editable-region--
}
```
