---
id: 69cfdc80443b715a8aa452a7
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

لإكمال دالة `initializeThroughput` الخاصة بك:

- استخدم حلقة `for` للتكرار خلال كل بوابة في مصفوفة `gates` المُدخَلة. لكل بوابة، أضف خاصية إلى `summary` حيث:
  - المفتاح هو `id` الخاص بالبوابة.
  - القيمة هي `0`.
- بعد الحلقة، أرجع كائن `summary`.

يجب ألا تُغيّر دالة `initializeThroughput` مصفوفة `gates` المصدرية المُدخلة.

# --hints--

يجب أن تُرجع دالة `initializeThroughput` `{}` عند استدعائها مع `[]`، دون تعديل المصفوفة المصدرية.

```js
const input = [];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  {},
  "Calling `initializeThroughput([])` should return an empty object."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given an empty array."
);
```

يجب أن تُرجع دالة `initializeThroughput` `{ North: 0 }` عند استدعائها مع `[{ id: "North", capacity: 5, queue: [1, 2, 3] }]`، دون تعديل المصفوفة المصدرية.

```js
const input = [{ id: "North", capacity: 5, queue: [1, 2, 3] }];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0 },
  "Calling `initializeThroughput([{ id: 'North', ... }])` should return `{ North: 0 }`."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given one gate."
);
```

يجب أن تُرجع دالة `initializeThroughput` `{ North: 0, East: 0, South: 0 }` عند استدعائها مع `[{ id: "North", capacity: 5, queue: [] }, { id: "East", capacity: 3, queue: [] }, { id: "South", capacity: 4, queue: [] }]`، دون تعديل المصفوفة المصدرية.

```js
const input = [
  { id: "North", capacity: 5, queue: [] },
  { id: "East", capacity: 3, queue: [] },
  { id: "South", capacity: 4, queue: [] }
];

const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0, East: 0, South: 0 },
  "Calling `initializeThroughput` with multiple gates should return an object with an entry for each gate id set to 0."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given multiple gates."
);
```

يجب أن تستخدم حلقة `for` للتكرار خلال كل كائن بوابة داخل مصفوفة `gates` المُدخلة.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));
assert.match(cleaned, /for\s*\(/);
```

يجب أن تُرجع دالة `initializeThroughput` كائن `summary` بعد الحلقة، كسطر نهائي في الدالة.

```js
const funcStr = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));

assert.match(
  funcStr,
  /returnsummary;?}$/,
  "Your `initializeThroughput` function should end with `return summary`."
);
```

يجب ألا تُغيّر دالة `initializeThroughput` مصفوفة `gates` المصدرية المُدخلة.

```js
const morningCopy = JSON.parse(JSON.stringify(morningGates));
const nightCopy = JSON.parse(JSON.stringify(nightGates));

initializeThroughput(morningGates);
initializeThroughput(nightGates);

assert.deepEqual(
  morningGates,
  morningCopy,
  "`initializeThroughput` should not mutate the `morningGates` seed array."
);

assert.deepEqual(
  nightGates,
  nightCopy,
  "`initializeThroughput` should not mutate the `nightGates` seed array."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
--fcc-editable-region--
  
--fcc-editable-region--
}
```
