---
id: 6140d0069049f5139d78da40
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

La propiedad `transform-origin` se usa para establecer el punto alrededor del cual se aplica la transformación CSS. Por ejemplo, cuando aplicas una transformación `rotate` (como harás más adelante en este proyecto), el `transform-origin` determinará alrededor de qué punto se rota el elemento.

Asigna al selector `.line` una propiedad `transform-origin` de `0% 0%`. Esto desplazará el punto de origen al `0%` desde la izquierda y `0%` desde arriba, estableciéndolo en la esquina superior izquierda del elemento.

# --hints--

Tu selector `.line` debe tener una propiedad `transform-origin` establecida en `0% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert.oneOf(transformOrigin, ['0% 0%', '0% 0% 0px']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
