---
id: 6140d10d50636e14695013b2
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Crea un selector para apuntar a el segundo elemento `.line`. Establece la propiedad `transform` a `rotate(60deg)`.

Recuerda que la propiedad `transform` te permite manipular la forma de un elemento. En este caso, usar el valor `rotate(60deg)` rotará el elemento alrededor de su punto de `transform-origin` en 60 grados en sentido de las agujas del reloj.

# --hints--

Debe crear un selector `.line:nth-of-type(2)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

Tu selector `.line:nth-of-type(2)` debe tener una propiedad `transform` establecida a `rotate(60deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform, 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
