---
id: 6140c7e645d8e905819f1dd4
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Para comenzar, agrega un elemento `link` para el archivo `./styles.css`.

# --hints--

Tu código debe tener un elemento `link`.

```js
assert.match(code, /<link/);
```

Debes tener un elemento `link`.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Su elemento `link` debe estar dentro de su elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Su elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>

  </body>
</html>

```

```css

```
