---
id: 587d7fb5367417b2b2512c03
title: استخدم مؤشر الإدراج لاستخدام أحدث إصدار فرعي من تبعية
challengeType: 2
forumTopicId: 301531
dashedName: use-the-caret-character-to-use-the-latest-minor-version-of-a-dependency
---

# --description--

مماثل لكيفية السماح للمد الذي تعلمناه في التحدٍّ السابق لـ npm بتثبيت أحدث تصحيح (PATCH) لتبعية، يسمح مؤشر الإدراج (`^`) لـ npm بتثبيت التحديثات المستقبلية أيضًا. الفرق هو أن مؤشر الإدراج يسمح بالتحديثات الفرعية (MINOR) والتصحيحات (PATCH) معًا.

الإصدار الحالي لـ `@freecodecamp/example` يجب أن يكون `~1.2.13` مما يسمح لـ npm بتثبيت أحدث إصدار من `1.2.x`. إذا استخدمت مؤشر الإدراج (^) كبادئة للإصدار بدلاً من ذلك، سيسمح npm بالتحديث إلى أي إصدار من `1.x.x`.

```json
"package": "^1.3.8"
```

هذا سيسمح بالتحديثات إلى أي إصدار من `1.x.x` للحزمة.

# --instructions--

استخدم مؤشر الإدراج (`^`) كبادئة لإصدار `@freecodecamp/example` في تبعياتك، واسمح لـ npm بتحديثه إلى أي إصدار فرعي جديد.

**ملاحظة:** لا يجب تغيير أرقام الإصدارات نفسها.

# --hints--

`"dependencies"` يجب أن تتضمن `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

يجب أن يتطابق إصدار `"@freecodecamp/example"` مع `"^1.x.x"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\^1\./,
        'Wrong version of "@freecodecamp/example". It should be ^1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
