---
id: 587d7fb4367417b2b2512bff
title: إضافة إصدار إلى ملف package.json الخاص بك
challengeType: 2
forumTopicId: 301525
dashedName: add-a-version-to-your-package-json
---

# --description--

الـ `version` هو أحد الحقول المطلوبة في ملف package.json الخاص بك. يصف هذا الحقل الإصدار الحالي لمشروعك. إليك مثالاً:

```json
"version": "1.2.0",
```

# --instructions--

أضف `version` إلى ملف package.json الخاص بمشروعك.

# --hints--

يجب أن يحتوي ملف package.json على مفتاح "version" صالح

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.version, '"version" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
