---
id: 587d7fb5367417b2b2512c04
title: إزالة حزمة من تبعياتك
challengeType: 2
forumTopicId: 301530
dashedName: remove-a-package-from-your-dependencies
---

# --description--

لقد اختبرت الآن عدة طرق لإدارة تبعيات مشروعك باستخدام قسم التبعيات في ملف package.json. كما أضفت حزمًا خارجية بإضافتها إلى الملف وأخبرت npm أيضًا بأنواع الإصدارات التي تريدها باستخدام أحرف خاصة مثل المد أو مؤشر الإدراج.

لكن ماذا لو أردت إزالة حزمة خارجية لم تعد بحاجة إليها؟ ربما خمّنت ذلك بالفعل، فقط احذف زوج المفتاح والقيمة المقابل لتلك الحزمة من تبعياتك.

تنطبق نفس الطريقة على إزالة حقول أخرى في ملف package.json أيضًا.

# --instructions--

احذف الحزمة `@freecodecamp/example` من تبعياتك.

**ملاحظة:** تأكد من وجود العدد الصحيح من الفواصل بعد إزالتها.

# --hints--

يجب ألا تتضمن `"dependencies"` الحزمة `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.notProperty(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" still includes "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
