---
id: 587d7fb4367417b2b2512c00
title: وسّع مشروعك باستخدام حزم خارجية من npm
challengeType: 2
forumTopicId: 301527
dashedName: expand-your-project-with-external-packages-from-npm
---

# --description--

أحد أكبر الأسباب لاستخدام مدير الحزم هو إدارته القوية للتبعيات. بدلًا من التأكد يدويًا من حصولك على كل التبعيات عند إعداد مشروع على جهاز جديد، يقوم npm بتثبيت كل شيء تلقائيًا. لكن كيف يعرف npm بالضبط ما يحتاجه مشروعك؟ تعرّف على قسم `dependencies` في ملف package.json الخاص بك.

في هذا القسم، تُخزن الحزم التي يحتاجها مشروعك باستخدام الصيغة التالية:

```json
"dependencies": {
  "package-name": "version",
  "express": "4.14.0"
}

```

# --instructions--

أضف الإصدار `1.1.0` من الحزمة `@freecodecamp/example` إلى حقل `dependencies` في ملف `package.json` الخاص بك.

**ملاحظة:** الحزمة `@freecodecamp/example` هي حزمة وهمية تُستخدم كأداة تعليمية.

# --hints--

يجب أن يتضمن `"dependencies"` الحزمة `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

يجب أن يكون إصدار `"@freecodecamp/example"` هو `"1.1.0"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^[\^\~]?1\.1\.0/,
        'Wrong version of "@freecodecamp/example" installed. It should be 1.1.0'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
