---
id: 587d7fb4367417b2b2512bfe
title: أضف ترخيصًا إلى package.json الخاص بك
challengeType: 2
forumTopicId: 301523
dashedName: add-a-license-to-your-package-json
---

# --description--

حقل `license` هو المكان الذي تخبر فيه المستخدمين بما يُسمح لهم بفعله في مشروعك.

بعض التراخيص الشائعة للمشاريع مفتوحة المصدر تشمل MIT و BSD. معلومات الترخيص ليست مطلوبة، وقوانين حقوق الملكية الفكرية في معظم البلدان تمنحك ملكية ما تنشئه افتراضيًا. مع ذلك، من الجيد دائمًا أن توضح صراحة ما يمكن للمستخدمين فعله وما لا يمكنهم فعله. إليك مثالًا على حقل الترخيص:

```json
"license": "MIT",
```

# --instructions--

املأ حقل `license` في ملف package.json الخاص بمشروعك كما تراه مناسبًا.

# --hints--

يجب أن يحتوي package.json على مفتاح "license" صالح

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.license, '"license" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
