---
id: 587d7fb5367417b2b2512c02
title: استخدم رمز المد دائمًا لاستخدام أحدث إصدار تصحيحي من تبعية
challengeType: 2
forumTopicId: 301532
dashedName: use-the-tilde-character-to-always-use-the-latest-patch-version-of-a-dependency
---

# --description--

في التحدٍّ السابق، أخبرت npm أن تدرج إصدارًا محددًا فقط من الحزمة. هذه طريقة مفيدة لتجميد التبعيات إذا كنت تريد التأكد من أن أجزاء مشروعك المختلفة تبقى متوافقة مع بعضها. لكن في معظم الحالات، لا تريد أن تفوت إصلاحات الأخطاء البرمجية لأنها غالبًا ما تتضمن تصحيحات أمنية مهمة ولا تكسر الأمور أثناء ذلك، على أمل.

للسماح لتبعية npm بالتحديث إلى أحدث إصدار تصحيحي (PATCH)، يمكنك وضع رمز المد (`~`) كبادئة لإصدار التبعية. إليك مثالًا على كيفية السماح بالتحديثات لأي إصدار `1.3.x`.

```json
"package": "~1.3.8"
```

# --instructions--

في ملف package.json، القاعدة الحالية لكيفية ترقية npm لـ `@freecodecamp/example` هي استخدام إصدار محدد (`1.2.13`). لكن الآن، تريد السماح بأحدث إصدار `1.2.x`.

استخدم رمز المد (`~`) كبادئة لإصدار `@freecodecamp/example` في تبعياتك، واسمح لـ npm بتحديثه إلى أي إصدار تصحيحي جديد.

**ملاحظة:** لا يجب تغيير أرقام الإصدارات نفسها.

# --hints--

يجب أن تتضمن `"dependencies"` العنصر `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

يجب أن يتطابق إصدار `"@freecodecamp/example"` مع `"~1.2.13"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\~1\.2\.13/,
        'Wrong version of "@freecodecamp/example". It should be ~1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
