---
id: 587d7fb4367417b2b2512bfd
title: أضف الكلمات المفتاحية إلى package.json الخاص بك
challengeType: 2
forumTopicId: 301526
dashedName: add-keywords-to-your-package-json
---

# --description--

حقل `keywords` هو المكان الذي يمكنك فيه وصف مشروعك باستخدام الكلمات المفتاحية ذات الصلة. إليك مثالًا:

```json
"keywords": [ "descriptive", "related", "words" ],
```

كما ترى، هذا الحقل مُنظّم كمصفوفة من السلاسل النصية المحاطة بعلامات اقتباس مزدوجة.

# --instructions--

أضف مصفوفة من السلاسل النصية المناسبة إلى حقل `keywords` في ملف package.json الخاص بمشروعك.

يجب أن تكون إحدى الكلمات المفتاحية هي "freecodecamp".

# --hints--

يجب أن يحتوي package.json على مفتاح "keywords" صالح

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.keywords, '"keywords" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

يجب أن يكون حقل "keywords" مصفوفة

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.isArray(data.keywords, '"keywords" is not an array');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

يجب أن تتضمن "keywords" كلمة "freecodecamp"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.include(
        data.keywords,
        'freecodecamp',
        '"keywords" does not include "freecodecamp"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
