---
id: 587d7fb3367417b2b2512bfc
title: أضف وصفًا إلى package.json الخاص بك
challengeType: 2
forumTopicId: 301522
dashedName: add-a-description-to-your-package-json
---

# --description--

الجزء التالي من ملف package.json الجيد هو حقل `description`؛ حيث يكتب وصف قصير لكنه مفيد عن مشروعك.

إذا كنت تخطط يومًا لنشر حزمة على npm، فهذه هي السلسلة النصية التي يجب أن تبيع فكرتك للمستخدم عندما يقرر ما إذا كان سيثبت الحزمة أم لا. مع ذلك، هذا ليس الاستخدام الوحيد للوصف، فهو طريقة رائعة لتلخيص ما يفعله المشروع. وهو مهم بنفس القدر في أي مشروع Node.js لمساعدة المطوّرين الآخرين، أو القائمين على الصيانة في المستقبل، أو حتى نفسك في المستقبل على فهم المشروع بسرعة.

بغض النظر عما تخطط له لمشروعك، يُنصح بشدة بوجود وصف. إليك مثالًا:

```json
"description": "A project that does something awesome",
```

# --instructions--

أضف `description` إلى ملف package.json الخاص بمشروعك.

**ملاحظة:** تذكر استخدام علامات الاقتباس المزدوجة لأسماء الحقول (") والفواصل (,) لفصل الحقول.

# --hints--

يجب أن يحتوي package.json على مفتاح "description" صالح

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.description, '"description" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
