---
id: 66ba762af611169359d9d369
title: Створіть газетну статтю
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні встановити кореневий `font-size` вашого HTML-документа на `24px`.
1. Ви повинні мати елемент із класом `newspaper`, який містить усі ваші інші елементи.
1. Ваш елемент `.newspaper` повинен мати властивість `font-size` зі значенням `16px` і шрифт `Open Sans` із запасним варіантом шрифту `sans-serif`.
1. Усередині вашого елемента `.newspaper` має бути щонайменше сім інших елементів: один для назви газети з класом `name`, один для дати статті з класом `date`, один для заголовка з класом `headline`, один для підзаголовка з класом `sub-headline`, один для автора з класом `author` і два абзаци, кожен із класом `text`. Усі ці елементи мають бути заповнені інформацією вашої статті.
1. Ваш елемент `.name` повинен мати властивість `font-size`, яка вдвічі більша за властивість `font-size` кореневого елемента, і використовувати шрифт `Times New Roman` із запасним варіантом шрифту `serif`.
1. Елементи `.name` і `.author` повинні використовувати CSS, щоб усі їхні символи були у верхньому регістрі.
1. Елемент `.headline` повинен мати властивість `font-size`, яка вдвічі більша за властивість `font-size` батьківського елемента, і бути жирним.
1. Елемент `.sub-headline` повинен мати властивість `font-weight` зі значенням `100`, властивість `font-size`, яка в `1.5` разів більша за властивість `font-size` батьківського елемента, і бути курсивним.
1. Ваш елемент `.author` повинен бути жирним за допомогою CSS.
1. Елементи `.text` повинні мати властивість `text-indent` зі значенням `20px`.
1. Елементи `.text` повинні мати властивість `line-height`, яка вдвічі більша за властивість `font-size` батьківського елемента.
1. Перша літера ваших елементів `.text` повинна бути жирною і вдвічі більшою за властивість `font-size` батьківського елемента. Використовуйте селектор `::first-letter` для цього.

# --hints--

Ви повинні мати селектор `html`, який встановлює його властивість `font-size` у значення `24px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

Ви повинні мати елемент із класом `newspaper`.

```js
assert.exists(document.querySelector('.newspaper'));
```

Ви повинні мати селектор `.newspaper`, який встановлює властивість `font-size` своїх елементів у значення `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

Ви повинні мати селектор `.newspaper`, який встановлює властивість `font-family` своїх елементів у значення `'Open Sans', sans-serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

Ви повинні мати елемент із класом `name` всередині вашого елемента `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

Ваш елемент `.name` повинен містити назву вашої газети.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

Ви повинні мати елемент із класом `date` всередині вашого елемента `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

Ваш елемент `.date` повинен містити дату.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

Ви повинні мати елемент із класом `headline` всередині вашого елемента `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

Ваш елемент `.headline` повинен містити заголовок вашої статті.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

Ви повинні мати елемент із класом `sub-headline` всередині вашого елемента `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

Ваш елемент `.sub-headline` повинен містити підзаголовок вашої статті.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

Ви повинні мати елемент із класом `author` всередині вашого елемента `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

Ваш елемент `.author` повинен містити автора вашої статті.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

Ви повинні мати щонайменше два абзаци, кожен із класом `text`, всередині вашого елемента `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

Ваші елементи `.text` повинні містити текст вашої статті.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

Ви повинні мати селектор `.name`, який встановлює властивість `font-size` своїх елементів у значення `2rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

Ви повинні мати селектор `.name`, який встановлює властивість `font-family` своїх елементів у значення `'Times New Roman', serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

Елемент `.name` повинен мати властивість `text-transform`, встановлену у значення `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Елемент `.author` повинен мати властивість `text-transform`, встановлену у значення `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Ви повинні мати селектор `.headline`, який встановлює властивість `font-size` своїх елементів у значення `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

Ви повинні мати селектор `.headline`, який встановлює властивість `font-weight` своїх елементів у значення `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

Ви повинні мати селектор `.sub-headline`, який встановлює властивість `font-weight` своїх елементів у значення `100`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

Ви повинні мати селектор `.sub-headline`, який встановлює властивість `font-size` своїх елементів у значення `1.5em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

Ви повинні мати селектор `.sub-headline`, який встановлює властивість `font-style` своїх елементів у значення `italic`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

Ви повинні мати селектор `.author`, який встановлює властивість `font-weight` своїх елементів у значення `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

Ви повинні мати селектор `.text`, який встановлює властивість `text-indent` своїх елементів у значення `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

Ви повинні мати селектор `.text`, який встановлює властивість `line-height` своїх елементів у значення `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

Ви повинні мати селектор `.text::first-letter`, який встановлює властивість `font-weight` своїх елементів у значення `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

Ви повинні мати селектор `.text::first-letter`, який встановлює властивість `font-size` своїх елементів у значення `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
