---
id: 619692ff79f5770fc6d8c0b4
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Fügen Sie oberhalb des `.ground`-Elements ein `div` mit einem `class` von `penguin` hinzu. Dieses `div` wird Flappy Penguin enthalten.

# --hints--

Sie sollten ein neues `div` innerhalb des `body` hinzufügen.

```js
assert.lengthOf(document.querySelectorAll('body > div'), 2);
```

Sie sollten dem `div` eine `class` von `penguin` geben.

```js
assert.include(document.querySelector('body > div:not(.ground)')?.className, 'penguin');
```

Sie sollten `div.penguin` vor `div.ground` platzieren.

```js
assert.strictEqual(document.querySelector('.ground')?.previousElementSibling, document.querySelector('.penguin'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
