---
id: 619665c9abd72906f3ad30f9
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Sie werden einen fröhlichen Flappy Penguin erstellen und dabei weiter CSS-Transformationen und Animationen üben.

Beginnen Sie damit, Ihr Stylesheet mit der Seite zu verknüpfen.

# --hints--

Ihr Code sollte ein `link`-Element enthalten.

```js
assert.exists(document.querySelector('link'));
```

Ihr `link`-Element sollte sich innerhalb Ihres `head`-Elements befinden.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
--fcc-editable-region--
    
--fcc-editable-region--
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css

```
