---
id: 61969aa6acef5b12200f672e
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Um etwas Landschaft im Hintergrund zu erstellen, fügen Sie zwei Berge hinzu.

Fügen Sie oberhalb des `.penguin`-Elements ein `div` mit einem `class` von `left-mountain` hinzu.

# --hints--

Sie sollten ein neues `div` innerhalb von `body` hinzufügen. Erwartet wurden --fcc-expected-- `div`-Elemente, aber es wurden --fcc-actual-- gefunden.

```js
assert.lengthOf(document.querySelectorAll('body > div'), 3);
```

Sie sollten dem `div` eine `class` von `left-mountain` geben.

```js
assert.include(document.querySelector('body > div:not(.ground, .penguin)')?.className, 'left-mountain');
```

Sie sollten `.left-mountain` vor `.penguin` platzieren.

```js
assert.strictEqual(document.querySelector('.penguin')?.previousElementSibling, document.querySelector('.left-mountain'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
