---
id: 6197cff995d03905b0cca8ad
title: Schritt 28
challengeType: 0
dashedName: step-28
---

# --description--

Ihr Pinguin wird aus zwei Hauptabschnitten bestehen: dem Kopf und dem Körper.

Fügen Sie innerhalb von `.penguin` zwei neue `div`-Elemente hinzu. Das erste mit einem `class` von `penguin-head`, und das zweite mit einem `class` von `penguin-body`.

# --hints--

Sie sollten zwei neue `div`-Elemente zu `.penguin` hinzufügen. Erwartet wurden --fcc-expected-- `.penguin > div`-Elemente, aber gefunden wurden --fcc-actual--.

```js
assert.lengthOf(document.querySelectorAll('.penguin > div'), 2);
```

Dem ersten `div` sollten Sie ein `class` von `penguin-head` geben.

```js
assert.include(document.querySelector('.penguin > div:nth-of-type(1)')?.className, 'penguin-head');
```

Dem zweiten `div` sollten Sie ein `class` von `penguin-body` geben.

```js
assert.include(document.querySelector('.penguin > div:nth-of-type(2)')?.className, 'penguin-body');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="sun"></div>
    <div class="penguin">
--fcc-editable-region--
      
--fcc-editable-region--
    </div>

    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;
  transform: rotate(45deg);
  left: 110px;
  top: 225px;
}

.sun {
  width: 200px;
  height: 200px;
  background-color: yellow;
  position: absolute;
  border-radius: 50%;
  top: -75px;
  right: -75px;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}

```
