---
id: 587d7daa367417b2b2512b6c
title: 使用 join 方法將陣列合併成字串
challengeType: 1
forumTopicId: 18221
dashedName: combine-an-array-into-a-string-using-the-join-method
---

# --description--

`join` 方法用來將陣列的元素串接在一起以生成字串。它接受一個用於分隔字串中陣列元素的引數。

這裡有一個範例：

```js
const arr = ["Hello", "World"];
const str = arr.join(" ");
```

`str` 的值會是字串 `Hello World`。
# --instructions--

在 `sentensify` 函式（程式）內使用 `join` 方法（行為）等，將字串 `str` 中的字組合成一句話。該函式應該傳回一個字串。 例如，`I-like-Star-Wars` 會被轉換成 `I like Star Wars`。在此挑戰中，請勿使用 `replace` 方法（行為）。

# --hints--

你的程式碼應該使用 `join` 方法。

```js
assert(__helpers.removeJSComments(code).match(/\.join/g));
```

你的程式碼不應該使用 `replace` 方法。

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?replace/g));
```

`sentensify("May-the-force-be-with-you")` 應該傳回一個字串。

```js
assert(typeof sentensify('May-the-force-be-with-you') === 'string');
```

`sentensify("May-the-force-be-with-you")` 應該傳回字串 `May the force be with you`。

```js
assert(sentensify('May-the-force-be-with-you') === 'May the force be with you');
```

`sentensify("The.force.is.strong.with.this.one")` 應該傳回字串 `The force is strong with this one`。

```js
assert(
  sentensify('The.force.is.strong.with.this.one') ===
    'The force is strong with this one'
);
```

`sentensify("There,has,been,an,awakening")` 應該傳回字串 `There has been an awakening`。

```js
assert(
  sentensify('There,has,been,an,awakening') === 'There has been an awakening'
);
```

# --seed--

## --seed-contents--

```js
function sentensify(str) {
  // Only change code below this line


  // Only change code above this line
}

sentensify("May-the-force-be-with-you");
```

# --solutions--

```js
function sentensify(str) {
  return str.split(/\W/).join(' ');
}
```
