---
id: 587d7b8f367417b2b2512b62
title: 在 Prototype 上實作 map
challengeType: 1
forumTopicId: 301230
dashedName: implement-map-on-a-prototype
---

# --description--

如你先前使用 `Array.prototype.map()` 或簡稱 `map()` 所見，`map` 方法會傳回一個與其被呼叫的陣列長度相同的陣列。只要其回呼函式沒有改變原始陣列，它也不會改變原始陣列。

換句話說，`map` 是純函式，其輸出完全取決於其輸入。此外，它將另一個函式作為其引數。

如果你實作自己的 `map` 方法，可能會學到很多關於它的知識。建議你使用 `for` 迴圈或 `Array.prototype.forEach()`。

# --instructions--

撰寫你自己的 `Array.prototype.myMap()`，其行為應該與 `Array.prototype.map()` 完全相同。你不應該使用內建的 `map` 方法。可以在 `myMap` 方法中使用 `this` 存取 `Array` 實體。

# --hints--

`[23, 65, 98, 5, 13].myMap(item => item * 2)` 應該相等於 `[46, 130, 196, 10, 26]`。

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item * 2;
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`["naomi", "quincy", "camperbot"].myMap(element => element.toUpperCase())` 應該傳回 `["NAOMI", "QUINCY", "CAMPERBOT"]`。

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element.toUpperCase();
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`[1, 1, 2, 5, 2].myMap((element, index, array) => array[index + 1] || array[0])` 應該傳回 `[1, 2, 5, 2, 1]`。

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array[index + 1] || array[0];
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

你的程式碼不應該使用 `map` 方法。

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?map/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    newArray.push(callback(this[i], i, this));
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const doubled_s = s.myMap(item => item * 2);
```
