---
id: 587d7da9367417b2b2512b67
title: 使用 concat 而非 push 為陣列的末端添加元素
challengeType: 1
forumTopicId: 301226
dashedName: add-elements-to-the-end-of-an-array-using-concat-instead-of-push
---

# --description--

函式式程式設計完全是關於建立和使用不會變異的函式（程式）。

上一個挑戰介紹了 `concat` 方法，作為將陣列合併成新陣列而不改變原始陣列的方式。將 `concat` 與 `push` 方法比較。`push` 將項目添加到它所呼叫的同一個陣列的末端，這會改變該陣列。以下是一個範例：

```js
const arr = [1, 2, 3];
arr.push(4, 5, 6);
```

`arr` 的修改後值會是 `[1, 2, 3, 4, 5, 6]`，這不是函式式程式設計的方式。

`concat` 提供一種將新條款合併到陣列末端而不產生任何變異副作用的方法。

# --instructions--

修改 `nonMutatingPush` 函式，使其使用 `concat` 將 `newItem` 合併到 `original` 的末端，且不會變更 `original` 或 `newItem` 陣列。該函式應傳回一個陣列。

# --hints--

你的程式碼應該使用 `concat` 方法。

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

你的程式碼不應該使用 `push` 方法。

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?push/g));
```

`first` 陣列不應該改變。

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

第二個 `array` 不應該改變。

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingPush([1, 2, 3], [4, 5])` 應該傳回 `[1, 2, 3, 4, 5]`。

```js
assert(
  JSON.stringify(nonMutatingPush([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingPush(original, newItem) {
  // Only change code below this line
  return original.push(newItem);

  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingPush(first, second);
```

# --solutions--

```js
function nonMutatingPush(original, newItem) {
  return original.concat(newItem);
}
const first = [1, 2, 3];
const second = [4, 5];
```
