---
id: 587d7dab367417b2b2512b6f
title: 使用 some 方法檢查陣列中是否有任何元素符合條件
challengeType: 1
forumTopicId: 301314
dashedName: use-the-some-method-to-check-that-any-elements-in-an-array-meet-a-criteria
---

# --description--

`some` 方法用於陣列，以檢查是否有*任何*元素通過特定測試。它會傳回布林值－如果有任何值符合條件，則為 `true`，否則為 `false`。

例如，以下程式碼會檢查 `numbers` 陣列中是否有任何元素小於 10：

```js
const numbers = [10, 50, 8, 220, 110, 11];

numbers.some(function(currentValue) {
  return currentValue < 10;
});
```

`some` 方法會傳回 `true`。

# --instructions--

在 `checkPositive` 函式中使用 `some` 方法來檢查 `arr` 中是否有任何元素為正數。該函式應該傳回布林值。

# --hints--

你的程式碼應該使用 `some` 方法。

```js
assert(__helpers.removeJSComments(code).match(/\.some/g));
```

`checkPositive([1, 2, 3, -4, 5])` 應該傳回 `true`。

```js
assert(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` 應該傳回 `true`。

```js
assert(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([-1, -2, -3, -4, -5])` 應該傳回 `false`。

```js
assert(!checkPositive([-1, -2, -3, -4, -5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.some(elem => elem > 0);
}
```
