---
id: 587d7da9367417b2b2512b66
title: 使用 concat 方法合併兩個陣列
challengeType: 1
forumTopicId: 301229
dashedName: combine-two-arrays-using-the-concat-method
---

# --description--

<dfn>串接</dfn>是指將條款一個接一個地連接起來。JavaScript 提供了用於字串和陣列的 `concat` 方法，兩者的運作方式相同。對於陣列，該方法會在一個陣列上呼叫，然後將另一個陣列作為 `concat` 的引數提供，該陣列會被添加到第一個陣列的末端。它會傳回一個新的陣列，且不會改變任一原始陣列。以下是一個範例：

```js
[1, 2, 3].concat([4, 5, 6]);
```

傳回的陣列會是 `[1, 2, 3, 4, 5, 6]`。

# --instructions--

在 `nonMutatingConcat` 函式中使用 `concat` 方法，將 `attach` 串接到 `original` 的末端。該函式應該傳回串接後的陣列。

# --hints--

你的程式碼應該使用 `concat` 方法。

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

`first` 陣列不應該改變。

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

第二個 `array` 不應該改變。

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingConcat([1, 2, 3], [4, 5])` 應該傳回 `[1, 2, 3, 4, 5]`。

```js
assert(
  JSON.stringify(nonMutatingConcat([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingConcat(original, attach) {
  // Only change code below this line


  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingConcat(first, second);
```

# --solutions--

```js
function nonMutatingConcat(original, attach) {
  return original.concat(attach);
}
const first = [1, 2, 3];
const second = [4, 5];
```
