---
id: 587d7b8f367417b2b2512b64
title: 在 Prototype 上實作 filter 方法
challengeType: 1
forumTopicId: 301231
dashedName: implement-the-filter-method-on-a-prototype
---

# --description--

如果你實作自己的 `filter` 方法，可能會學到很多。建議你使用 `for` 迴圈或 `Array.prototype.forEach()`。

# --instructions--

撰寫你自己的 `Array.prototype.myFilter()`，其行為應該與 `Array.prototype.filter()` 完全相同。你不應該使用內建的 `filter` 方法。可以在 `myFilter` 方法中使用 `this` 存取 `Array` 實體。

# --hints--

`[23, 65, 98, 5, 13].myFilter(item => item % 2)` 應該相等於 `[23, 65, 5, 13]`。

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item % 2;
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

`["naomi", "quincy", "camperbot"].myFilter(element => element === "naomi")` 應該傳回 `["naomi"]`。

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element === "naomi";
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

`[1, 1, 2, 5, 2].myFilter((element, index, array) => array.indexOf(element) === index)` 應該傳回 `[1, 2, 5]`。

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array.indexOf(element) === index;
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

你的程式碼不應該使用 `filter` 方法。

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?filter/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myFilter = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myFilter = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    if (callback(this[i], i, this)) newArray.push(this[i]);
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const odd_s = s.myFilter(item => item % 2 === 1);
```
