---
id: 587d7b8e367417b2b2512b5c
title: 了解函式式程式設計術語
challengeType: 1
forumTopicId: 301240
dashedName: understand-functional-programming-terminology
---

# --description--

FCC 團隊情緒波動，現在想要兩種茶：綠茶和紅茶。一般事實：客戶端情緒波動相當常見。

有了這些資訊，我們需要重新檢視上一個挑戰中的 `getTea` 函式（程式），以處理各種茶的請求。我們可以修改 `getTea`，使其接受一個函式（程式）作為參數，從而能夠改變它準備的茶的型別。這讓 `getTea` 更加靈活，並且當客戶端請求改變時，給程式設計師更多的控制。

但首先，讓我們涵蓋一些函式（程式）術語：

<dfn>回呼</dfn> 是被傳入或傳遞到另一個函式中，以決定該函式呼叫時機的函式。你可能看過它們被傳遞到其他方法中，例如在 `filter` 中，回呼函式告訴 JavaScript 如何篩選陣列的條件。

可以指定給變數、傳入另一個函式，或從另一個函式傳回，就像任何其他一般值一樣的函式，稱為 <dfn>一級</dfn> 函式。在 JavaScript 中，所有函式都是一級函式。

接受函式作為引數，或傳回函式作為返回值的函式，稱為<dfn>高階</dfn>函式。

當函式被傳入或從另一個函式傳回時，則那些被傳入或傳回的函式可以稱為 <dfn>lambda</dfn>。

# --instructions--

準備 27 杯綠茶和 13 杯紅茶，並分別將它們儲存在 `tea4GreenTeamFCC` 和 `tea4BlackTeamFCC` 變數中。請注意，`getTea` 函式（程式）已被修改，現在它將函式（程式）作為第一個引數。

注意：資料（茶杯數量）作為最後一個引數提供。我們會在後續課程中進一步討論這個內容。

# --hints--

`tea4GreenTeamFCC` 變數應該保存 27 杯綠茶給團隊。

```js
assert(tea4GreenTeamFCC.length === 27);
```

`tea4GreenTeamFCC` 變數應該保存綠茶杯數。

```js
assert(tea4GreenTeamFCC[0] === 'greenTea');
```

`tea4BlackTeamFCC` 變數應該保存 13 杯黑茶。

```js
assert(tea4BlackTeamFCC.length === 13);
```

`tea4BlackTeamFCC` 變數應該保存黑茶杯數。

```js
assert(tea4BlackTeamFCC[0] === 'blackTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareGreenTea = () => 'greenTea';

// Function that returns a string representing a cup of black tea
const prepareBlackTea = () => 'blackTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (prepareTea, numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4GreenTeamFCC = null;
const tea4BlackTeamFCC = null;
// Only change code above this line

console.log(
  tea4GreenTeamFCC,
  tea4BlackTeamFCC
);
```

# --solutions--

```js
const prepareGreenTea = () => 'greenTea';
const prepareBlackTea = () => 'blackTea';

const getTea = (prepareTea, numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

const tea4BlackTeamFCC = getTea(prepareBlackTea, 13);
const tea4GreenTeamFCC = getTea(prepareGreenTea, 27);
```
