---
id: 587d7da9367417b2b2512b69
title: 使用 `sort` 方法按字母順序排序陣列
challengeType: 1
forumTopicId: 18303
dashedName: sort-an-array-alphabetically-using-the-sort-method
---

# --description--

`sort` 方法根據回呼函式對陣列的元素進行排序。

例如：

```js
function ascendingOrder(arr) {
  return arr.sort(function(a, b) {
    return a - b;
  });
}

ascendingOrder([1, 5, 2, 3, 4]);
```

這將傳回值 `[1, 2, 3, 4, 5]`。

```js
function reverseAlpha(arr) {
  return arr.sort(function(a, b) {
    return a === b ? 0 : a < b ? 1 : -1;
  });
}

reverseAlpha(['l', 'h', 'z', 'b', 's']);
```

這會傳回值 `['z', 's', 'l', 'h', 'b']`。

JavaScript 預設的排序方法是依照字串的 Unicode 編碼點值，這可能會回傳意料之外的結果。因此，建議提供一個回呼函式來指定如何排序陣列的項目。當提供這樣一個通常稱為 `compareFunction` 的回呼函式時，陣列元素會依照 `compareFunction` 的返回值進行排序：如果 `compareFunction(a,b)` 對兩個元素 `a` 和 `b` 傳回小於 0 的值，則 `a` 會排在 `b` 前面。如果 `compareFunction(a,b)` 對兩個元素 `a` 和 `b` 傳回大於 0 的值，則 `b` 會排在 `a` 前面。如果 `compareFunction(a,b)` 對兩個元素 `a` 和 `b` 傳回相等 0 的值，則 `a` 和 `b` 的順序保持不變。

# --instructions--

在 `alphabeticalOrder` 函式中使用 `sort` 方法，將 `arr` 的元素依字母順序排序。該函式應該傳回排序後的陣列。

# --hints--

你的程式碼應該使用 `sort` 方法。

```js
assert(__helpers.removeJSComments(code).match(/\.sort/g));
```

`alphabeticalOrder(["a", "d", "c", "a", "z", "g"])` 應該傳回 `["a", "a", "c", "d", "g", "z"]`。

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'd', 'c', 'a', 'z', 'g'])) ===
    JSON.stringify(['a', 'a', 'c', 'd', 'g', 'z'])
);
```

`alphabeticalOrder(["x", "h", "a", "m", "n", "m"])` 應該傳回 `["a", "h", "m", "m", "n", "x"]`。

```js
assert(
  JSON.stringify(alphabeticalOrder(['x', 'h', 'a', 'm', 'n', 'm'])) ===
    JSON.stringify(['a', 'h', 'm', 'm', 'n', 'x'])
);
```

`alphabeticalOrder(["a", "a", "a", "a", "x", "t"])` 應該傳回 `["a", "a", "a", "a", "t", "x"]`。

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'a', 'a', 'a', 'x', 't'])) ===
    JSON.stringify(['a', 'a', 'a', 'a', 't', 'x'])
);
```

# --seed--

## --seed-contents--

```js
function alphabeticalOrder(arr) {
  // Only change code below this line

  return arr
  // Only change code above this line
}

alphabeticalOrder(["a", "d", "c", "a", "z", "g"]);
```

# --solutions--

```js
function alphabeticalOrder(arr) {
  return arr.sort();
}
```
