---
id: 587d7dab367417b2b2512b6e
title: 使用 every 方法檢查陣列中每個元素是否符合條件
challengeType: 1
forumTopicId: 301312
dashedName: use-the-every-method-to-check-that-every-element-in-an-array-meets-a-criteria
---

# --description--

`every` 方法用於陣列，以檢查 *每個* 元素是否通過特定測試。它傳回布林值－如果所有值都符合條件，則為 `true`，否則為 `false`。

例如，以下程式碼會檢查 `numbers` 陣列中的每個元素是否都小於 10：

```js
const numbers = [1, 5, 8, 0, 10, 11];

numbers.every(function(currentValue) {
  return currentValue < 10;
});
```

`every` 方法會在這裡傳回 `false`。

# --instructions--

在 `checkPositive` 函式中使用 `every` 方法來檢查 `arr` 中的每個元素是否為正數。該函式應該傳回一個布林值。

# --hints--

你的程式碼應該使用 `every` 方法。

```js
assert(__helpers.removeJSComments(code).match(/\.every/g));
```

`checkPositive([1, 2, 3, -4, 5])` 應該傳回 `false`。

```js
assert.isFalse(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` 應該傳回 `true`。

```js
assert.isTrue(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([1, -2, 3, -4, 5])` 應該傳回 `false`。

```js
assert.isFalse(checkPositive([1, -2, 3, -4, 5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.every(num => num > 0);
}
```
