---
id: 587d7daa367417b2b2512b6b
title: 使用 split 方法將字串拆分成陣列
challengeType: 1
forumTopicId: 18305
dashedName: split-a-string-into-an-array-using-the-split-method
---

# --description--

`split` 方法將字串拆分成字串陣列。它接受一個用於分隔符的引數，該分隔符可以是用來拆分字串的字元或正規表達式。例如，如果分隔符是空格，你會得到一個字陣列；如果分隔符是空字串，你會得到字串中每個字元的陣列。

以下是兩個範例，先以空白分割一個字串，再以正規表達式用數字分割另一個字串：

```js
const str = "Hello World";
const bySpace = str.split(" ");

const otherString = "How9are7you2today";
const byDigits = otherString.split(/\d/);
```

`bySpace` 的值會是 `["Hello", "World"]`，而 `byDigits` 的值會是 `["How", "are", "you", "today"]`。

由於字串是不可變的，`split` 方法使得操作它們更加容易。

# --instructions--

在 `splitify` 函式內使用 `split` 方法將 `str` 拆分成字的陣列。該函式應傳回該陣列。請注意，字不一定總是以空格分隔，且陣列中不應包含標點符號。

# --hints--

你的程式碼應該使用 `split` 方法。

```js
assert(__helpers.removeJSComments(code).match(/\.split/g));
```

`splitify("Hello World,I-am code")` 應該傳回 `["Hello", "World", "I", "am", "code"]`。

```js
assert(
  JSON.stringify(splitify('Hello World,I-am code')) ===
    JSON.stringify(['Hello', 'World', 'I', 'am', 'code'])
);
```

`splitify("Earth-is-our home")` 應該傳回 `["Earth", "is", "our", "home"]`。

```js
assert(
  JSON.stringify(splitify('Earth-is-our home')) ===
    JSON.stringify(['Earth', 'is', 'our', 'home'])
);
```

`splitify("This.is.a-sentence")` 應該傳回 `["This", "is", "a", "sentence"]`。

```js
assert(
  JSON.stringify(splitify('This.is.a-sentence')) ===
    JSON.stringify(['This', 'is', 'a', 'sentence'])
);
```

# --seed--

## --seed-contents--

```js
function splitify(str) {
  // Only change code below this line


  // Only change code above this line
}

splitify("Hello World,I-am code");
```

# --solutions--

```js
function splitify(str) {
  return str.split(/\W/);
}
```
