---
id: 587d7b8d367417b2b2512b5b
title: 了解函式式程式設計
challengeType: 1
forumTopicId: 301233
dashedName: learn-about-functional-programming
---

# --description--

函式式程式設計是一種程式設計風格，其中解決方案是簡單、獨立的函式，且在函式範圍外沒有任何副作用：`INPUT -> PROCESS -> OUTPUT`

函式式程式設計是關於：

獨立函式（程式）－不依賴程式狀態，包括可能變動的全域變數

純函式（程式）－相同的輸入總是產生相同的輸出

具有有限副作用的函式（程式）－對函式（程式）外部程式狀態的任何變更或突變都受到嚴格控制

# --instructions--

freeCodeCamp 的成員恰巧喜歡喝茶。

在程式碼編輯器中，`prepareTea` 和 `getTea` 函式已經為你定義好。呼叫 `getTea` 函式以取得 40 杯茶給團隊，並將它們儲存在 `tea4TeamFCC` 變數中。

# --hints--

`tea4TeamFCC` 變數應該保存 40 杯茶給團隊。

```js
assert(tea4TeamFCC.length === 40);
```

`tea4TeamFCC` 變數應該保存綠茶杯數。

```js
assert(tea4TeamFCC[0] === 'greenTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareTea = () => 'greenTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4TeamFCC = null;
// Only change code above this line
```

# --solutions--

```js
const prepareTea = () => 'greenTea';

const getTea = (numOfCups) => {
  const teaCups = [];
  
  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }

  return teaCups;
};

const tea4TeamFCC = getTea(40); 
```
