---
id: 587d7b88367417b2b2512b45
title: '使用高階函式（函式（程式）、函數（數學））`map`、`filter` 或 `reduce` 來解決複雜問題'
challengeType: 1
forumTopicId: 301311
dashedName: use-higher-order-functions-map-filter-or-reduce-to-solve-a-complex-problem
---

# --description--

現在你已經透過使用像是 `map()`、`filter()` 和 `reduce()` 這些高階函式（函式、函數）完成了幾個挑戰，你現在可以將它們應用於解決更複雜的挑戰。

# --instructions--

使用任意組合的 `map()`、`filter()` 和 `reduce()` 完成 `squareList` 函式的程式碼。當傳入一個實數陣列時，該函式應傳回一個新的陣列，僅包含正整數（小數不是整數）的平方。實數陣列的範例為 `[-3, 4.8, 5, 3, -3.2]`。

**注意：**你的函式（程式）不應該使用任何類型的 `for` 或 `while` 迴圈或 `forEach()` 函式（程式）。

# --hints--

`squareList` 應該是一個 `function`。

```js
assert.typeOf(squareList, 'function'),
  '<code>squareList</code> should be a <code>function</code>';
```

不應使用 `for`、`while` 和 `forEach`。

```js
assert(!__helpers.removeJSComments(code).match(/for|while|forEach/g));
```

應該使用 `map`、`filter` 或 `reduce`。

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/\.(map|filter|reduce)\(/g)
);
```

該函式（程式）應該傳回一個 `array`。

```js
assert(Array.isArray(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])));
```

`squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])` 應該傳回 `[16, 1764, 36]`。

```js
assert.deepStrictEqual(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2]), [
  16,
  1764,
  36
]);
```

`squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3])` 應該傳回 `[9, 100, 49]`。

```js
assert.deepStrictEqual(squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3]), [
  9,
  100,
  49
]);
```

# --seed--

## --seed-contents--

```js
const squareList = arr => {
  // Only change code below this line
  return arr;
  // Only change code above this line
};

const squaredIntegers = squareList([-3, 4.8, 5, 3, -3.2]);
console.log(squaredIntegers);
```

# --solutions--

```js
const squareList = arr => {
  const positiveIntegers = arr.filter(num => {
    return num >= 0 && Number.isInteger(num);
  });
  const squaredIntegers = positiveIntegers.map(num => {
    return num ** 2;
  });
  return squaredIntegers;
};
```
