---
id: 587d7b8e367417b2b2512b5f
title: 傳遞引數以避免函式依賴外部資源
challengeType: 1
forumTopicId: 301234
dashedName: pass-arguments-to-avoid-external-dependence-in-a-function
---

# --description--

最後一個挑戰更接近函式程式設計原則，但仍然缺少一些東西。

我們沒有更改全域變數的值，但函式（程式）`incrementer` 若沒有全域變數 `fixedValue` 就無法運作。

函式（程式）設計的另一個原則是始終明確地宣告你的依賴關係。這表示如果一個函式依賴於某個變數或物件的存在，則將該變數或物件直接作為引數傳入該函式。

這個原則帶來幾個良好的結果。函式（程式）更容易測試，你確切知道它接受什麼輸入，而且它不會依賴你程式中的其他任何東西。

這可以讓你在修改、移除或添加新程式碼時更有信心。你會知道哪些可以或不可以更改，並且能看到潛在的陷阱所在。

最後，該函式（程式）無論程式碼的哪一部分執行它，對相同的一組輸入總是會產生相同的輸出。

# --instructions--

讓我們更新 `incrementer` 函式（程式），以明確宣告它的依賴項。

撰寫 `incrementer` 函式，使其接受一個引數，然後在將值增加一後傳回結果。

# --hints--

你的函式（程式）`incrementer` 不應該改變 `fixedValue`（其值為 `4`）。

```js
assert(fixedValue === 4);
```

你的 `incrementer` 函式應該接受一個引數。

```js
assert(incrementer.length === 1);
```

你的 `incrementer` 函式應該傳回比 `fixedValue` 值大 1 的值。

```js
const __newValue = incrementer(fixedValue);
assert(__newValue === 5);
```

# --seed--

## --seed-contents--

```js
// The global variable
let fixedValue = 4;

// Only change code below this line
function incrementer() {


  // Only change code above this line
}
```

# --solutions--

```js
let fixedValue = 4;

function incrementer(fixedValue) {
  return fixedValue + 1;
}
```
