---
id: 660ef19b95d3308e7dd31bb6
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

<dfn>コンソール</dfn> を使用すると、JavaScript の出力を表示して確認できます。 コンソールに情報を送信するには `console.log()` を使います。 例えば、次のコードはコンソールに `"Naomi"` と表示します。

```js
let developer = "Naomi";
console.log(developer);
```

上記のコードでは、`console.log()` の中で `developer` という名前の変数にアクセスしています。 丸括弧の間の値がコンソールに表示される値となることに注目してください。

では、`character` という変数の値をコンソールに表示してください。 JavaScript コンソールの表示・非表示を切り替えるには「コンソール」ボタンをクリックします。

# --hints--

The `character` variable initialization should not be changed.

```js
assert.strictEqual(character, 'Hello');
```

コード内で `console` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

`console` の `log` プロパティを使用してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

`.log()` メソッドを呼び出すには、丸括弧を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

変数 `character` をコンソールに表示してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
