---
id: 660ee6e3a242da6bd579de69
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

プログラミングにおける最も重要な概念の 1 つが変数です。 <dfn>変数</dfn>は、ある値を格納している特定のメモリアドレスを指すものです。 変数には名前が付けられ、その名前を使ってコード内の至る所からその値にアクセスすることができます。

変数の宣言とは、変数に名前を付けることです。 JavaScript では、変数の宣言によく `let` キーワードが使われます。 例えば、`hello` という変数を宣言する方法は次のとおりです:

```js
let hello;
```

変数名は特有のルールに従います。変数名には、アルファベットの文字、数字、ドル記号、アンダースコアが使用できますが、スペースを含めることはできず、また、数字で始めることもできません。

では、`let` キーワードを使用して、`character` という名前の変数を宣言してください。

_注_: JavaScript では、文をセミコロン `;` で終えることが一般的です 。

# --hints--

`let` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

`character` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

`let` を使って、変数 `character` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

宣言はセミコロンで終わる必要があります。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
