---
id: 660ef5105b8ba095307a0e50
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

ではその後にもう一度、`character` という変数の値をコンソールに表示してください。 `"Hello"` という文字列の後に、`"World"` という文字列が表示されるはずです。

# --hints--

2 回目の `console.log` を使用してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

コンソールに `character` を 2 回出力する必要があります。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


新しい `console.log()` は再代入の後に置く必要があります。

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
