---
id: 660ef5c1904955978a986a5c
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

`let` を使用して、変数 `count` を宣言してください。 その変数に、<dfn>数値</dfn>の `8` を代入してください。 数値を扱う際は、引用符は使いません。 例えば、次のように記述します。

```js
let money = 100;
```

# --hints--

コードに 2 つ目の `let` キーワードを追加してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

`let` を使って `count` という変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

変数 `count` に数値 `8` を割り当ててください。

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
