---
id: 660ef55dd468079679ee0092
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

では、次の変数を宣言する準備をしましょう。 すべての `console.log` 文と、変数 `character` への再代入を行っている箇所を削除してください。

変数 `secondCharacter` も削除してください。変数 `character` の初期化を行っているコードだけ、変更せずに残してください。

# --hints--

コード内に `console.log()` ステートメントが無いようにしてください。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

`character` の値を再代入しないでください。

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

変数 `character` の値を `"Hello"` に設定する必要があります。

```js
assert.equal(character, "Hello");
```

変数 `secondCharacter` は必要ありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
