---
id: 660eebd83100d37862268781
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

現在、変数 `character` には値がありません。 変数に値を代入するには、<dfn>代入</dfn>演算子 `=` を使用します。 例えば、次のように記述します。

```js
let hello = "Hello";
```

変数の宣言と同時に値を代入することは、<dfn>初期化 (initialization)</dfn> と呼ばれます。

変数 `character` の宣言と同時に、値 `"Hello"` を代入して初期化してください。

# --hints--

代入演算子 `=` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

文字列 `"Hello"` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

`"Hello"` という文字列にはダブルクォートを使ってください。

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

変数 `character` に `"Hello"` を代入してください。

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
