---
id: 660ef31a5be625914a0102cd
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

変数が `let` キーワードで宣言されている場合、その変数は<dfn>再代入</dfn> (値の変更) が可能です。 次の例では、変数 `programmer` の値が `"Naomi"` から `"CamperChan"` に変更されます。

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

再代入の際は、`let` キーワードは使わ**ない**ことに注意してください。

では、`console.log` の後で、変数 `character` に値 `"World"` を再代入してください。

# --hints--

コード内で `character` を 2 回使用してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

`character` を再代入するために、代入演算子を使用してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

コード内で `"World"` という文字列を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

Your `character` variable should have the value `"World"`.

```js
assert.equal(character, "World");
```

Your reassignment should not use `let`.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
