---
id: 660ef6355e8f5a9e67fe5f46
title: ステップ 16
challengeType: 1
dashedName: step-16
---

# --description--

数値 (`number`) のデータ型に対しては、足し算などの計算を実行できます。 コンソールに `count + 1` を出力してみましょう。

# --hints--

コード内で `console` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

`console` の `.log()` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

コンソールに `count + 1` を出力する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
