---
id: 660f48e1d3682f39e81843c4
title: ステップ 94
challengeType: 1
dashedName: step-94
---

# --description--

ピラミッド生成機はこの状態でも動作しています。 しかし、再び無限ループに陥る可能性があります。

あなたは `done` が `count` と等しくないことしかチェックしていないので、もし `done` が `count` よりも**大きい**場合、無限ループとなってしまいます。

では、`done` が `count` 以下かどうかをチェックするように、ループの条件を変更してください。

# --hints--

`while` ループで `done` が `count` 以下かどうかをチェックするようにしてください。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
