---
id: 660f4a83373de83ea101685f
title: ステップ 100
challengeType: 1
dashedName: step-100
---

# --description--

逆方向にループしたいので、`i` が `0` より大きい間ループを実行する必要があります。 そのためには<dfn>大なり</dfn>演算子 `>` が使えます。

ループの条件を、`i` が `0` より大きい場合に設定してください。

# --hints--

`for` ループは、`i` が `0` より大きい間実行される必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; false; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
