---
id: 660f39b444fd6f16d1e49c1f
title: ステップ 73
challengeType: 1
dashedName: step-73
---

# --description--

次は、`padRow` の呼び出しに `i + 1` を 渡す代わりに `1` からループを開始することもできます。 つまり 1 始まりのインデックスを使用するループを作成します。

では、イテレーターを `0` ではなく `1` から開始するように変更しましょう。

# --hints--

`for` ループのイテレーター変数 `i` を `1` で初期化してください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
