---
id: 660f0f980e98e8cf77f1ce31
title: ステップ 30
challengeType: 1
dashedName: step-30
---

# --description--

`rows` の宣言を変更して、また空の配列を割り当ててください。

また、`'Hello'` という文字列をダブルクォートを使う形に再度変更してください。 一般的には、どちらを使っても構いませんが、プロジェクト全体を通して一貫性を持たせるようにしましょう。

# --hints--

配列 `rows` は空にしてください。

```js
assert.empty(rows);
```

文字列 `"Hello"` にはダブルクォートを使用してください。

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
