---
id: 660f033cf051ebb50ea3bf48
title: ステップ 18
challengeType: 1
dashedName: step-18
---

# --description--

プログラミングでは、大量のデータを扱うことがよくあります。 データの整理や管理に役立つさまざまなデータ型があります。 最も基本的なデータ構造の 1 つが<dfn>配列</dfn> (array) です。

<dfn>配列</dfn> は、一連の値を保持することができる、プリミティブでないデータ型です。 プリミティブデータ型と異なり、プリミティブでないデータ型は、より複雑なデータを保持することが可能です。 文字列や数値などのプリミティブデータ型は、一度に 1 つの値しか保持できません。

配列は、角括弧 `[]` を用いて表されます。 こちらが、変数の値として空の配列を持たせる例です。

```js
let array = [];
```

では、変数 `rows` を宣言して、その変数に空の配列を代入してください。

# --hints--

変数 `rows` が必要です。

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

`let` を使って変数 `rows` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

変数 `rows` には空の配列を割り当ててください。

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
