---
id: 660f165270622fd4ec0da3f7
title: ステップ 31
challengeType: 1
dashedName: step-31
---

# --description--

変数を `let` キーワードで宣言すると、再代入が可能になります。 つまり、このコードの `character` は後で全く異なる値に変更できます。

ですがこのプロジェクトでは、これらの変数の値は変更したくありません。 そのため、代わりに `const` を使って宣言してください。 `const` 変数は特別です。

まず、`const` 変数は、`let` 変数のように再代入することはできません。 次のようなコードを書くとエラーが発生します:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

また、`const` 変数は初期化せずに使うことができません。 次のコードではエラーが発生します:

```js
const firstName;
```

では、`let` キーワードをすべて `const` に置き換えてください。

# --hints--

`const` を使って変数 `character` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

`const` を使って変数 `count` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

`const` を使って変数 `rows` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

コード内で `let` を使わないようにしてください。

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
