---
id: 660f280dda5040f707c76b4a
title: ステップ 45
challengeType: 1
dashedName: step-45
---

# --description--

これで一連の `#` 記号が表示されるようになりましたが、まだピラミッドの形にはなっていません。 ありがたいことに、変数 `i` はループ内で現在の「行」番号を表しているため、これを使えばピラミッドのような構造を作成できます。

そのためには、文字列に利用可能な `.repeat()` メソッドを使用します。 このメソッドは引数として数値を受け取り、指定された回数だけターゲット文字列を繰り返します。 例えば、`.repeat()` を使用して文字列 `"Code! Code! Code!"` を生成する例がこちらです。

```js
const activity = "Code! ";
activity.repeat(3);
```

では、`character` に対して `.repeat()` メソッドを使用して、繰り返しの回数として `i` を渡してください。

# --hints--

`.repeat()` メソッドを使用する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

変数 `character` に対して `.repeat()` メソッドを使用する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

`.repeat()` メソッドに `i` を渡す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
