---
id: 660f540c2176ea5dec01306d
title: ステップ 118
challengeType: 1
dashedName: step-118
---

# --description--

すばらしい！ 変数 `character`、`count`、`inverted` にさまざまな値を入れて動作を確認しましょう。

次のプロジェクトに進む準備ができたら、`character` を `"!"`、`count` を `10`、`inverted` を `false` に設定して提出してください。

初めての JavaScript プロジェクトの完成、おめでとうございます！

# --hints--

`character` を `"!"` に設定してください。

```js
assert.equal(character, "!");
```

`count` を `10` に設定してください。

```js
assert.equal(count, 10);
```

`inverted` を `false` に設定してください。

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
