---
id: 660f3ba3cceef11b6ba08b59
title: ステップ 75
challengeType: 1
dashedName: step-75
---

# --description--

残念ながら、今度はピラミッドの一番下の段が消えてしまいました。 これは、別の <dfn>Off-by-one エラー</dfn>が起きてしまったためです。

`count` が `8` で、条件は `i` が `count` より小さいこととしていたので、元のループは `i` の値が `0` から `7` まで繰り返していました。 現在のループは `i` の値が `1` から `7` までで実行されています。

ループを修正して、`i` が `8` のときも実行されるようにしましょう。 つまり、`i` が `count` <dfn>以下</dfn>の間、ループを実行します。 これには小なりイコール演算子 `<=` が使用できます。

では、ループ条件を変更して `i` が `count` 以下の間実行されるようにしてください。

# --hints--

`for` ループは `i` が `count` 未満かどうかをチェックするべきではありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

`for` ループが `i` が `count` 以下であるかどうかをチェックするようにしてください。

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
