---
id: 660f039ff313dbb696b007ca
title: ステップ 19
challengeType: 1
dashedName: step-19
---

# --description--

配列内に値 (<dfn>要素</dfn> とも呼びます) を入れる場合は、値をコンマで区切ります。 こちらが、2 つの文字列を持つ配列の例です。

```js
let array = ["first", "second"];
```

`rows` の宣言を変更して、`"Naomi"`、`"Quincy"`、`"CamperChan"` という文字列を入れた配列にしてください。 配列内の値の順番は重要なため、この順番通りにしてください。 文字列では大文字・小文字が区別されます。

# --hints--

配列の最初の要素は、文字列 `"Naomi"` である必要があります。

```js
assert.equal(rows[0], "Naomi");
```

配列の 2 番目の要素は、文字列 `"Quincy"` である必要があります。

```js
assert.equal(rows[1], "Quincy");
```

配列の 3 番目の要素は、文字列 `"CamperChan"` である必要があります。

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
