---
id: 660f447efc0e722f016c1be0
title: ステップ 81
challengeType: 1
dashedName: step-81
---

# --description--

テキストが再び表示されるようになりました！ これは、`"false"` が文字列であり、文字列はブール値として評価されると `true` になるためです。 これは、`"false"` が真値であることを意味します。

<dfn>真値 (truthy value)</dfn> とは、ブール値として評価されると「真」とみなされる値のことです。 JavaScript ではほとんどの値は真値です。

<dfn>偽値 (falsy value)</dfn> はその逆で、ブール値として評価されると「偽」とみなされる値のことです。 JavaScriptには決められた偽値のリストがあります。 偽値には `false`、`0`、`""`、`null`、`undefined`、`NaN` などが含まれます。

では、`if` 条件を偽値である空の文字列 `""` に変更してみてください。

# --hints--

`if` 文には条件として `""` が必要です。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
