---
id: 660f45ccf4ca5c31f253005a
title: ステップ 86
challengeType: 1
dashedName: step-86
---

# --description--

この状態で、もし `continueLoop`を true に変更してしまうと、`while` ループが永久に実行され続けます。 これは<dfn>無限ループ</dfn>と呼ばれます。このような状態にならないよう注意してください。 無限ループになるとシステムがロックされ、その状態を脱するにはそのシステムの再起動が必要になる場合があります。

ではこの問題を避けるために、まずインクリメント演算子を使ってループ内で変数 `done` の値を増やしましょう。

# --hints--

`while` ループ内で変数 `done` をインクリメントしてください。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
