---
id: 660f4455f457ef2e3ec6920f
title: ステップ 80
challengeType: 1
dashedName: step-80
---

# --description--

文字列が出力されなくなりました。なぜかというと、`false` (偽) は `true` (真) ではないからです。 では、他の値についてはどうでしょうか？

条件を文字列の `"false"` に変更してみてください。

# --hints--

`if` 文の条件として、文字列 `"false"` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*('|")false\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (false) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
