---
id: 660f255022991ef34ed0ee88
title: ステップ 44
challengeType: 1
dashedName: step-44
---

# --description--

表示されるものが数字では、ピラミッドらしい見た目にはなりませんね。 配列 `rows` の内容を成形して出力できるようになったので、今度は最初のループを変更しましょう。

配列に `i` をプッシュする代わりに、変数 `character` の値をプッシュしてください。

# --hints--

変数 `i` をプッシュする必要はありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

変数 `character` をプッシュしてください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
