---
id: 660f09a2694b59c3a10ee304
title: ステップ 23
challengeType: 1
dashedName: step-23
---

# --description--

では、最初の console.log と `rows[rows.length - 1]` への代入を削除してください。 `rows` を出力する 2 つ目の console.log 文は、後ほど使うため残しておきます。

# --hints--

`console.log(rows[0])` の文を削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

`rows[rows.length - 1]` の再代入を削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

`console.log(rows)` の文は削除しないでください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
