---
id: 660f3dd626be3a1ffe27e5d1
title: ステップ 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript では複数行のコメントもサポートされています。 複数行のコメントは `/*` で始まり、`*/` で終わります。

一行コメントと異なり、複数行コメントは複数の行をまとめて囲むことができます。

では、`/*` と `*/` を使用して、現在の `for` ループ (本体を含む) を複数行コメントにしてください。

# --hints--

`/*` で複数行コメントを開始してください。

```js
assert.match(code, /\/\*/);
```

`*/` で複数行コメントを終了してください。

```js
assert.match(code, /\*\//);
```

`for` ループ全体をコメントアウトしてください。

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
