---
id: 660f44f10ea40f300b896a5e
title: ステップ 84
challengeType: 1
dashedName: step-84
---

# --description--

`if...else if...else` 文の練習が終わったので、使用したコードを削除してください。

削除できたら、`let` を使って `continueLoop` という変数を宣言し、それにブール値 `false` を割り当ててください。 もう一つ、`let` を使って `done` という変数を宣言し、それに値 `0` を割り当ててください。

# --hints--

`if` 文を削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

`else if` 文を削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

`else` 文を削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

`let` を使って `continueLoop` という変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

変数 `continueLoop` の値は `false` にしてください。

```js
assert.isFalse(continueLoop);
```

`let` を使って `done` という変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

変数 `done` の値は `0` にしてください。

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
