---
id: 660f4934fb48f63abd5ae371
title: ステップ 95
challengeType: 1
dashedName: step-95
---

# --description--

`done` を使って生成済みの行数を追跡する方式も機能しますが、実際にはロジックをもう少し整理することができます。

配列には特別な `length` プロパティがあります。これを使うと配列に含まれる<dfn>要素</dfn>の数を確認することができます。 このプロパティには `myArray.length` という構文でアクセスできます。

ただし、`padRow` 呼び出しで `rows.length` を使う際には Off-by-one エラーの可能性に注意してください。現状は `done` が関数呼び出しの*前*にインクリメントされているためです。

ではまず、条件を変更して `rows.length` が `count` より小さいかどうかをチェックするようにしてください。

# --hints--

`while` ループは `rows.length` が `count` より小さいかどうかをチェックする必要があります。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
