---
id: 660f20473aef47e9b8c9afc6
title: ステップ 40
challengeType: 1
dashedName: step-40
---

# --description--

`result` の値を出力するログ文を追加してください。 使用するコンソールによっては、何も表示されない場合があります。

# --hints--

`console.log` 呼び出しを追加してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

変数 `result` をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
