---
id: 660f07d231941bc11719f664
title: ステップ 21
challengeType: 1
dashedName: step-21
---

# --description--

配列は<dfn>変更可能 (mutable)</dfn> と見なされる点が特殊です。 つまり、あるインデックスに位置する値を直接変更できます。

例えば、次のコードは配列の 2 番目の要素に数値 `25` を代入します。

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

では、配列 `rows` の **3 番目**の要素を、数値 `10` に変更しましょう。 その後、配列 `rows` をコンソールに表示してください。

# --hints--

もう一度 `rows` 配列にブラケット記法を使用してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

`rows` 配列の 3 番目の要素にアクセスしてください。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

`rows` 配列の 3 番目の要素に代入演算子を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

`rows` 配列の 3 番目の要素に、値 `10` を代入してください。

```js
assert.equal(rows[2], 10);
```

コード内に 2 つ目の `console.log` ステートメントを追加してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

`rows` 配列をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
