---
id: 660f2eccfe3f820304af1b39
title: ステップ 64
challengeType: 1
dashedName: step-64
---

# --description--

行をどのように成形したら良いか知るために、`padRow` 関数に現在の行番号と全体の行数を伝える必要があります。

それらの情報を渡す関数パラメーター (仮引数) を作成するのが良いでしょう。 `padRow` 関数に仮引数 `rowNumber` と `rowCount` を持たせてください。 複数のパラメータはコンマで区切ります。

```js
function name(first, second) {

}
```

# --hints--

関数 `padRow` には仮引数 `rowNumber` が必要です。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

仮引数 `rowNumber` の後にコンマを追加してください。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

関数 `padRow` には仮引数 `rowCount` が必要です。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
