---
id: 660f4de78f775e480ba2e451
title: ステップ 107
challengeType: 1
dashedName: step-107
---

# --description--

配列の `.unshift()` メソッドを使うと、配列の**先頭**に値を追加することができます。これは配列の末尾に値を追加する `.push()` と対照的です。 そして `.unshift()` は配列の新しい長さを返します。

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

では、`const` を使って `unshifted` という変数を宣言し、配列 `numbers` に対して `.unshift()` を呼び出した結果を代入してください。 引数として `5` を渡してください。 その後、変数 `unshifted` を出力してください。

# --hints--

変数 `unshifted` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

配列 `numbers` に対して `.unshift()` メソッドを呼び出してください。

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

`.unshift()` 呼び出しの引数として `5` を渡してください。

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

`.unshift()` 呼び出しの結果を変数 `unshifted` に代入してください。

```js
assert.equal(unshifted, 4);
```

変数 `unshifted` をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
