---
id: 660f0c34aad72dc712b97624
title: ステップ 27
challengeType: 1
dashedName: step-27
---

# --description--

このプロジェクトに不可欠なもう一つのメソッドが `.pop()` メソッドです。 このメソッドは配列から最後の要素を削除し、その要素を <dfn>return</dfn> (返却) します。

メソッドが値を「返す」とは、その値があなたの手元に戻ってきて、コードの他の部分で使える状態になることと考えるといいでしょう。

では、`popped` という名前の新しい変数を作成し、その変数に `rows.pop()` の結果を代入してください。 その後、`popped` をコンソールにログ出力してください。

# --hints--

`popped` という名前の変数を宣言する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

`popped` という変数の宣言には `let` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

`.pop()` メソッドを呼び出す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

配列 `rows` に対して `.pop()` メソッドを呼び出す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

変数 `popped` をログに出力する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
