---
id: 660f4ae5b3924c3fc3373973
title: ステップ 101
challengeType: 1
dashedName: step-101
---

# --description--

イテレーション文にも変更が必要です。 各ループで `i` に `1` を足す代わりに、`1` を引く必要があります。

以前に `i = i + 1` の処理を行ったように、イテレーション文を変更して、`i` に自身から `1` を引いた値を代入してください。

# --hints--

`for` ループのイテレーション文として `i = i - 1` を使う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*=\s*i\s*-\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
