---
id: 660f47afe4c98536715d5fa4
title: ステップ 91
challengeType: 1
dashedName: step-91
---

# --description--

<dfn>厳密不等価演算子</dfn> `!==` を使うと、2 つの値が等しくない、または型が異なるかどうかをチェックできます。 構文は等価演算子に似ており、`value !== 4` のように記述します。

では、`while` ループの条件を変更して、`done` が `count` と等しくないことをチェックするようにしてください。

# --hints--

`while` ループの条件で、`done` と `count` が等しくないことをチェックしてください。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
