---
id: 660f2a70ad6225fa503e71c3
title: ステップ 46
challengeType: 1
dashedName: step-46
---

# --description--

近づいてきました！ この時点で、<dfn>Off-by-one エラー</dfn>と呼ばれる問題に直面しています。これは、JavaScript のようなゼロ始まりのインデックスを使う言語でよく見られる問題です。

配列 `rows` の最初のインデックスは `0` なので、`for` ループは `i = 0` で始まります。 しかし、文字列をゼロ回繰り返しても、出力するものが何もありません。

これを修正するには、`.repeat()` の呼び出し時に `i` の値に `1` を加えてください。 ただし、ループ条件のような `i` への再代入はしないでください。

# --hints--

`.repeat()` メソッドで `i` に `1` を加えてください。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
