---
id: 660f4e74f7fd3f4a99ac2e50
title: ステップ 108
challengeType: 1
dashedName: step-108
---

# --description--

配列には `.shift()` メソッドもあります。 これは、配列の最後の要素を削除する `.pop()` とは対照的に、配列の**最初**の要素を削除します。 こちらが `.shift()` メソッドの使用例です。

```js
const numbers = [1, 2, 3];
numbers.shift();
```

配列 `numbers` は `[2, 3]` になります。

では、配列 `numbers` のすぐ下に `shifted` という変数を宣言し、配列 `numbers` に対して `.shift()` を呼び出した結果を代入してください。 次の行で、変数 `shifted` をコンソールにログ出力してください。

# --hints--

`const` を使って変数 `shifted` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

配列 `numbers` に対して `.shift()` を呼び出してください。

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

`.shift()` の呼び出し結果を変数 `shifted` に代入してください。

```js
assert.equal(shifted, 1);
```

変数 `shifted` をログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
