---
id: 660f17d4e9f227d86e834abd
title: ステップ 33
challengeType: 1
dashedName: step-33
---

# --description--

ピラミッドを生成するには、複数の行を作る必要があります。 条件が満たされるまでタスクを繰り返し実行する必要がある場合、<dfn>ループ</dfn>を使用します。 ループを書く方法はいろいろあります。

基本的な `for` ループから始めましょう。 `for` ループは次の構文を使用します:

```js
for (iterator; condition; iteration) {
  logic;
}
```

ループの各コンポーネントについては、この先のステップで詳しく説明します。 今のところは、3 つのコンポーネントに文字列 `"iterator"`、`"condition"`、`"iteration"` をそのまま使用して `for` ループを書いてください。 ループの<dfn>本文</dfn>、つまり波括弧 `{}` の中のセクションは、空にしてください。


# --hints--

`for` ループが必要です。

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

`for` ループの最初のコンポーネントは文字列 `"iterator"` にしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

`for` ループの 2 番目のコンポーネントは文字列 `"condition"` にしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

`for` ループの 3 番目のコンポーネントは文字列 `"iteration"` にしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

`for` ループの本文は空にする必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
