---
id: 660f2b6fd54ac1fc142804dd
title: ステップ 47
challengeType: 1
dashedName: step-47
---

# --description--

このピラミッドを成形するロジックは複雑になりそうです。つまり、このコードを関数に抽出する良いタイミングです。

<dfn>関数</dfn>とは、アプリケーションの他の場所で再利用可能なコードのブロックです。 関数は次の構文で宣言します。

```js
function name(parameter) {

}
```

`function` キーワードは、JavaScript に対して変数 `name` が関数になることを伝えます。 `parameter` は、関数の使用時に渡される値を表す変数です。 関数は、必要に応じて任意の数の<dfn>パラメーター (仮引数) </dfn>を持つことができます。 `for` ループと同様に、波括弧の間は<dfn>関数本体</dfn>です。

では、関数 `padRow` を宣言してください。 パラメーター変数はまだ作成しないでください。 関数本体は空にしてください。 命名規則として、キャメルケースを使うことを覚えておいてください。

# --hints--

`function` キーワードを使用してください。

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

関数 `padRow` を宣言してください。

```js
assert.isFunction(padRow);
```

関数 `padRow()` にはパラメーターを持たせてはいけません。

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

関数 `padRow()` の本体は空にしてください。

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
