---
id: 660f3ce51f70571e1c5227c8
title: ステップ 76
challengeType: 1
dashedName: step-76
---

# --description--

次に、コメントは、なぜその手法を採用しているのか説明したり、未来の自分に向けた TODO メモ (やることメモ) を残すのに役立ちます。

JavaScript では、`//` を使ってコードに一行コメントを残すことができます。

Add a single-line comment above your `for` loop to remind yourself to change the code to a different kind of loop.

# --hints--

一行コメントは `//` で始めてください。

```js
assert.match(code, /\/\//);
```

少なくとも 5 文字以上の一行コメントを書いてください。

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
