---
id: 660f1bf673487ae0bb25b900
title: ステップ 37
challengeType: 1
dashedName: step-37
---

# --description--

現在、ループは 8 回実行されるようになっているはずです。 ループの本文内で、イテレーター変数 `i` の値を出力し、何が起きるか確認しましょう。

# --hints--

`console.log()` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

`i` の値をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

`for` ループ内で `i` の値をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
