---
id: 660f5179b3b0ca558f6b4d4f
title: ステップ 113
challengeType: 1
dashedName: step-113
---

# --description--

`if` 文を使用して、`inverted` が true かどうかをチェックしてください。 この場合、等価演算子は必要ないことを思い出してください。

# --hints--

`if` 文を使用する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(/);
```

`if` 文では `inverted` が `true` かどうかをチェックしてください。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted/);
```

`if` 条件で比較演算子を使用しないでください。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {

  rows.unshift(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
