---
id: 660f3b664421471aa595170f
title: ステップ 74
challengeType: 1
dashedName: step-74
---

# --description--

ピラミッドが少しおかしくなってしまいました。 これは、`0` の代わりに `1` からループを始めるようにしたので、`i` を `padRow` に渡すときに 1 を足す必要がなくなったためです。

`padRow` 呼び出しの第一引数を、`i` に変更してください。

# --hints--

`i + 1` を `padRow` に渡してはいけません。

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

`i` を `padRow` に渡してください。

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
