---
id: 660f38c34a4de6141c0c369f
title: ステップ 71
challengeType: 1
dashedName: step-71
---

# --description--

これで、あなたのピラミッドジェネレーターは期待通りに動作するようになりました。 せっかくなので、ここまで書いたコードをさらに深く見ていきましょう。

加算演算子だけが変数に値を加える唯一の方法ではありません。 <dfn>加算代入演算子</dfn>は、「変数の元の値にこの値を加え、その結果を変数に代入する」という意味の省略形として使用できます。 例えば、次の 2 つのステートメントは同じ結果になります。

```js
test = test + 1;
test += 1;
```

Update your iteration statement in the `for` loop to use addition assignment.

# --hints--

`for` ループで `i = i + 1` を使用しないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

`for` ループは `i` に対して加算代入演算子を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
