---
id: 660f374d532dc41189cc9cc2
title: ステップ 69
challengeType: 1
dashedName: step-69
---

# --description--

さて、少し数学の時間です。 三段のピラミッドを考えてみましょう。 中央に揃えて配置する場合、次のようになります。

```md
··#··
·###·
#####
```

スペースは見やすくするために中黒で置き換えています。 このパターンを推測すると、行の始めと終わりのスペースがある法則に従っていることがわかります。

では、スペースの文字列を `rowCount - rowNumber` 回繰り返すように、コードを変更してください。

コンソールを開いて結果を確認してください。

# --hints--

文字列 `" "` に対して `.repeat()` を呼び出して、スペースを `rowCount - rowNumber` 回繰り返してください。

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
