---
id: 660f53ad3d39175c5d4335ac
title: ステップ 117
challengeType: 1
dashedName: step-117
---

# --description--

あなたのピラミッド表示プログラムは、当初の計画以上の機能を備えて完成に近づいています！ 次のステップは、コードをきれいにすることです。

コードから、一行と複数行のすべてのコメントを削除してください。

# --hints--

コード内に一行コメントが無いようにしてください。

```js
assert.notMatch(code, /\/\//);
```

コード内に複数行コメントが無いようにしてください。

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

コード内にコメントが無いようにしてください。

```js
assert.equal(code, __helpers.removeJSComments(code));
```

複数行コメントでコメントアウトされていたコードも削除してください。

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
