---
id: 660f1a00ac619ddc1e259a66
title: ステップ 35
challengeType: 1
dashedName: step-35
---

# --description--

`for` ループの<dfn>条件 (condition) </dfn>は、ループが何回繰り返すべきかを指示します。 `condition` が false になると、ループは停止します。

JavaScript では、ブール値は `true` または `false` のいずれかです。 これらは文字列ではありません。違いについては後ほど詳しく学びます。

ここでは<dfn>小なり</dfn>演算子 (`<`) を使用します。 これにより、左側の値が右側の値より小さいかどうかをチェックできます。 例えば、`count < 3` は、`count` が `2` の場合は `true` になり、`count` が `4` の場合は `false` になります。

では、`"condition"` という文字列を、`i` が `count` より小さいかどうかをチェックする条件に置き換えてください。

# --hints--

小なり演算子を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /</);
```

小なり演算子を使用して、`i` が `count` より小さいかをチェックしてください。

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

`for` ループの条件として`i < count` を使用する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
