---
id: 660f1b6e60bd9edf902c81fd
title: ステップ 36
challengeType: 1
dashedName: step-36
---

# --description--

<dfn>イテレーション</dfn>ステートメントは、ループが実行されるたびにイテレーター変数をどのように処理するかを指示します。

変数を再代入するとき、再代入の前に現在の値を参照することもできます。 これにより、現在の数に3 を足すというようなことができます。 例えば、`bees = bees + 3;` は `bees` の値を 3 増やします。

この構文を使用して、文字列 `"iteration"` を、`i` を 1 増やす再代入ステートメントに置き換えてください。

# --hints--

変数 `i` に 1 を足してください。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

`i + 1` を、変数 `i` に再代入してください。

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

`for` ループが各反復処理で `i` を `1` 増やすようにしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
