---
id: 660f0da9bf1035c9097af20a
title: ステップ 28
challengeType: 1
dashedName: step-28
---

# --description--

`"freeCodeCamp"` がコンソールに出力されたと思います。 これは `.pop()` が配列から削除された値を返すためです。そしてその前のステップで、あなたは配列の末尾に `"freeCodeCamp"` を push で追加していました。

では、`.push()` は何を返すのでしょうか？ 既存の `rows.push()` を新しい変数 `pushed` に代入し、出力してください。

# --hints--

変数 `pushed` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

変数 `pushed` の宣言には `let` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

`rows.push("freeCodeCamp")` を変数 `pushed` に代入してください。

```js
assert.equal(pushed, 4);
```

You should log your `pushed` variable.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
