---
id: 660f23b53db70af0f2620e78
title: ステップ 43
challengeType: 1
dashedName: step-43
---

# --description--

さて、すべての数字が同じ行に表示されるようになりました。 ですがこのままではピラミッドの形にならなさそうです。

行 (row) ごとに、毎回改行を入れましょう。 しかし、JavaScript のコード内でリターンキーを押して引用符の間に改行を挿入すると、解析エラーが発生します。 代わりに、<dfn>エスケープシーケンス</dfn> `\n` を使う必要があります。これは、文字列がログに出力されるときに改行として解釈されます。 例えば、次のように記述できます。

```js
lineOne = lineOne + "\n" + lineTwo;
```

2 つ目の加算演算子を使って、既存の `result` と `row` の間に改行を連結してください。

# --hints--

エスケープシーケンス `\n` を使用する必要があります。 文字列とするため、引用符で囲まなければならないことを覚えておいてください。

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

エスケープシーケンス `\n` を変数 `result` に連結する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*('|")\\n\1/);
```

変数 `row` をエスケープシーケンス `\n` に連結する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*('|")\\n\1\s*\+\s*row/);
```

連結した結果全体を、変数 `result` に再割り当てする必要があります。 セミコロンを忘れないようにしましょう。

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*result\s*\+\s*('|")\\n\1\s*\+\s*row;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
