---
id: 660f18f059fe0fda192ce394
title: ステップ 34
challengeType: 1
dashedName: step-34
---

# --description--

ループには適切なイテレーターが必要です。 <dfn>イテレーター変数</dfn>は、ループの繰り返しやロジックの流れを制御するために `for` ループ内で宣言する変数です。

ループのイテレータ変数として `i` を使用することは一般的な習慣です。 `for` ループでは、この変数を丸括弧 `()` の中で宣言できます。 次の例は、`index` 変数を宣言し、それに `100` を割り当てている `for` ループです。

```js
for (let index = 100; "second"; "third") {

}
```

では、文字列 `"iterator"` を変数 `i` を宣言する `let` 宣言に置き換えてください。 初期値として `0` を割り当ててください。 そうすると、ループが **1 回目**に実行される時、変数 `i` の値が `0` になります。

# --hints--

`let` を使って変数 `i` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

変数 `i` に `0` を割り当ててください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

`for` ループがイテレーター変数 `i` を `0` から開始するようにしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
