---
id: 660f2fbd45b520046cac68e8
title: ステップ 65
challengeType: 1
dashedName: step-65
---

# --description--

以前のステップで戻り値について学んだことを思い出してください。 関数は、アプリケーションが別の場所で使用するための値を<dfn>返す</dfn>ことができます。

関数では、`return` キーワードを使用して戻り値を指定します。 例えば、次の関数は第一引数に与えられた値を返します。

```js
function name(parameter) {
  return parameter;
}
```

では、`return` キーワードを使用して、変数 `character` の値を `rowNumber` 回繰り返したものを返してください。

# --hints--

`.repeat()` メソッドを使用する必要があります。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

変数 `character` に対して `.repeat()` メソッドを使用する必要があります。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

`.repeat()` 呼び出しに `rowNumber` を渡す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

`return` キーワードを使用する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

`.repeat()` 呼び出しの結果を返す必要があります。

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
