---
id: 660f46460f9c36330ebc07d8
title: ステップ 88
challengeType: 1
dashedName: step-88
---

# --description--

JavaScript の等価演算子は奇妙な挙動を引き起こすことがあります。 例えば、`"0" == 0` は、片方は文字列でもう片方は数値であるにもかかわらず、true となります。

<dfn>厳密等価</dfn>演算子 `===` は、2 つの値が等しくかつ同じ型であるかどうかをチェックするために使用されます。 一般的なルールとして、等価演算子としては常にこちらを使用すべきです。 厳密等価演算子を使えば、`"0" === 0` は false となります。ゼロという同じ値を持っているように見えるかもしれませんが、同じ型ではないためです。

では、`done == count` の条件を変更して、厳密等価演算子を使うようにしてください。

# --hints--

`if` 条件では厳密等価演算子を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
