---
id: 660f0a55847d6cc485f29ba5
title: ステップ 26
challengeType: 1
dashedName: step-26
---

# --description--

JavaScript における<dfn>メソッド</dfn>とは、特定の値またはオブジェクトに関連付けられた関数です。 既にお馴染みの例として、`console` オブジェクトに属する `.log()` メソッドがあります。

配列には固有のメソッドがあります。まず最初に `.push()` メソッドを見ていきましょう。 このメソッドを使うと、配列の末尾に値を「押し込む」ことができます。 下記は、配列の末尾に数値 `12` を追加する例です。

```js
array.push(12);
```

では `.push()` を使用して、文字列 `"freeCodeCamp"` を配列 `rows` の末尾に追加しましょう。 このコードは `console.log` の前に追加して、配列に加えた変更を確認できるようにしてください。

# --hints--

コード内で `.push()` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

配列 `rows` の `.push()` メソッドを使用してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

`.push()` メソッドに文字列 `"freeCodeCamp"` を渡す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
