---
id: 660f4a1472f8e63d76162ce5
title: ステップ 99
challengeType: 1
dashedName: step-99
---

# --description--

ピラミッドを逆さまに<dfn>反転</dfn>させるにはどうしたらいいでしょうか？ 試してみましょう！

まず、新しい `for` ループを作成するところから始めましょう。 イテレーター変数 `i` を宣言して値に `count` を代入し、条件文とイテレーション文には、今のところブール値 `false` を使用してください。

# --hints--

コードには `for` ループが必要です。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

`for` ループは `i` を `count` の値で初期化する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count/);
```

`for` ループは条件として `false` を使用する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*false/);
```


`for` ループは反復として `false` を使用する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
