---
id: 660f34626216270c682e2f7b
title: ステップ 66
challengeType: 1
dashedName: step-66
---

# --description--

<dfn>関数呼び出し</dfn>により、関数を実際に使用できます。 気付いていなかったかもしれませんが、これまで使用してきた `.push()` のようなメソッドも、関数呼び出しでした。

関数の名前と`()`を記述することで、関数が呼び出されます。 例えば `test` 関数を呼び出す方法は次のとおりです。

```js
test();
```

Replace the `character.repeat(i + 1)` in your `.push()` call with a function call for your `padRow` function. Do not add any arguments to it yet.

# --hints--

`push` の呼び出しで `i + 1` を使用しないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

`.push()` の呼び出しで `character.repeat` を使用しないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

`.push()` の呼び出しで `padRow` を呼び出してください。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

You should not have arguments in your `padRow` call.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
