---
id: 660f4cffb1459d45e34902d1
title: ステップ 106
challengeType: 1
dashedName: step-106
---

# --description--

逆さピラミッドを作成するには、逆方向にループする以外の方法もあります。

そのためには、新しい配列メソッドをいくつか学ぶ必要があります。 まず、`const` を使用して `numbers` という変数を宣言するところから始めましょう。 その変数に、`1`, `2`, `3` の要素を持つ配列を代入してください。 そして配列 `numbers` をログに出力してください。

# --hints--

`const` を使用して変数 `numbers` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

変数 `numbers` は配列にしてください。

```js
assert.isArray(numbers);
```

配列 `numbers` の要素は順番に `1`, `2`, `3` としてください。

```js
assert.deepEqual(numbers, [1,2,3]);
```

配列 `numbers` をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
