---
id: 660f50a21fe7645252804f2b
title: ステップ 112
challengeType: 1
dashedName: step-112
---

# --description--

逆さピラミッドと通常のピラミッドを切り替える方法があったら良さそうですね。

まず `inverted` という変数を宣言し、値 `true` を代入してください。 コード内ではこの変数を変更しないのですが、後ほどテストでこの変数を変更するため、`let` を使用してください。

# --hints--

`inverted` という変数を `let` で宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

`inverted` を `true` という値で初期化してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
