---
id: 660f229d2dbe09ef2954a4a1
title: ステップ 42
challengeType: 1
dashedName: step-42
---

# --description--

以前のループでは、加算演算子 `+` を使って `i` の値を `1` 増やしました。

文字列の値でも同様に、新しい文字列を既存の文字列に追加することができます。 例えば、`hello = hello + " World";` は、変数 `hello` に格納されている既存の文字列に、文字列 `" World"` を付け加えます。 これを<dfn>連結</dfn>といいます。

では、`for...of` ループで加算演算子を使って、`result` の値に `row` の値を連結してください。

# --hints--

変数 `result` に加算演算子を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

変数 `result` に `row` を連結してください。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

連結の結果を変数 `result` に再代入してください。

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
