---
id: 660f3915b41a441537ec9f5e
title: ステップ 72
challengeType: 1
dashedName: step-72
---

# --description--

`i` を `1` ずつ増やしているだけなので、<dfn>インクリメント演算子</dfn> `++` を使うこともできます。 この演算子は変数の値を 1 増やし、その変数に代入されている値を更新します。 例えば、次の例では `test` が `8` になります。

```js
let test = 7;
test++;
```

では、ループの繰り返し処理の加算代入をインクリメント演算子に置き換えてください。

# --hints--

`for` ループで `i` に加算代入を使用しないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

`for` ループで `i` にインクリメント演算子を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
