---
id: 660f207334fabaeac3269c38
title: ステップ 41
challengeType: 1
dashedName: step-41
---

# --description--

文字列 `result` の操作には、別の種類のループを使用します。 具体的には `for...of` ループを使用します。このループは、イテラブル (反復可能) なオブジェクトの各アイテムに対して反復処理を行い、一時的にそのアイテムを変数に代入します。

`for...of` ループの構文は次のようになります。

```js
for (const value of iterable) {

}
```

この場合、変数は 1 回の反復のためだけに存在し、ループ全体に渡って存在するものではないため、`const` を使用できます。

では、配列 `rows` を走査するために `for...of` ループを作成し、各値を変数 `row` に代入してください。

# --hints--

`for` キーワードをもう 1 つ使用してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

変数 `row` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

`const` を使って変数 `row` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

`for...of` ループで変数 `row` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

変数 `row` は、`of` キーワードを使用して `rows` から取り出す必要があります。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

`for...of` ループの本文は空にしてください。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
