---
id: 660f505d02b2bd513a1c3468
title: ステップ 111
challengeType: 1
dashedName: step-111
---

# --description--

ピラミッドが逆さまではなくなりました。 これは、新しい行を配列の**末尾**に追加しているためです。

ループの本体を変更して、新しい行を配列の先頭に追加してください。

# --hints--

`rows` の `unshift` メソッドを使ってください。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

`.unshift()` メソッドの引数として `padRow()` 呼び出しを渡してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

`padRow()` 呼び出しの第一引数として `i` を渡してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

`padRow()` 呼び出しの第二引数として `count` を渡してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
