---
id: 660f17294346b7d69e79db3d
title: ステップ 32
challengeType: 1
dashedName: step-32
---

# --description--

これでピラミッド表示プログラムの開発を始める準備が整いました。 変数 `character` は、ピラミッドを構成するブロックの役割をします。

その目的だと、`"Hello"` ではあまりうまく機能しないかもしれませんね。 変数 `character` の値をハッシュ記号 `"#"` に変更しましょう。

# --hints--

変数 `character` の値をハッシュ記号にしてください。

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
