---
id: 660f383d4c772c12ff59904b
title: ステップ 70
challengeType: 1
dashedName: step-70
---

# --description--

完全な式を引数として渡すこともできます。 関数はその式を評価した結果を受け取ります。 For example, these two function calls below would yield the same result. (Note that the order of operations rule PEMDAS—Parenthesis, Exponents, Multiplication, Division, Addition, Subtraction—applies):

```js
test(2 * 3 + 1);
test(7);
```

先ほどのパターンをもう一度見てみましょう。

```md
··#··
·###·
#####
```

では、`character` の値が `2 * rowNumber - 1` 回繰り返されるようにコードを変更してください。

再度コンソールを開いて結果を確認してください。

# --hints--

関数内で `.repeat()` 呼び出しに `2 * rowNumber - 1` を渡す必要があります。

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
