---
id: 660f4774e3e0df35a68bb5f2
title: ステップ 90
challengeType: 1
dashedName: step-90
---

# --description--

ピラミッドを再生成するために、引数として `done` と `count` を用いて `padRow` を呼び出し、その結果を `rows` 配列に追加してください。これは、最初のループで行ったことと似ています。

# --hints--

ループ内で `rows` に対して `.push()` メソッドを呼び出してください。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(/);
```

`.push()` メソッドの括弧内で `padRow` 関数を呼び出してください。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(/);
```

`padRow` の呼び出しでは、第一引数として `done` を渡してください。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done/);
```

`padRow` の呼び出しでは、第二引数として `count` を渡してください。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--

--fcc-editable-region--
  if (done === count) {
    continueLoop = false;
  } 
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
