---
id: 660f061d259bbebc37461080
title: ステップ 20
challengeType: 1
dashedName: step-20
---

# --description--

配列内の値にアクセスするには、その値の<dfn>インデックス</dfn>を使います。 インデックスとは、その値の配列内での位置を表す番号です。インデックスは `0` から始まります。

値にアクセスするには、`array[0]` というような<dfn>ブラケット記法</dfn>を使用します。

では、`console.log` とブラケット記法を使用して、配列 `rows` の最初の値を表示してください。

# --hints--

コード内に `console.log()` ステートメントが無いようにしてください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

`rows` 配列にアクセスしてください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

`rows` 配列にブラケット記法を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

`rows` 配列の最初の要素に、ブラケット記法を使用してアクセスしてください。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

`rows` 配列の最初の要素をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
