---
id: 660f359af3e32e0f1a6880b7
title: ステップ 68
challengeType: 1
dashedName: step-68
---

# --description--

同じ文字列がコンソールに表示されたはずです。 `padRow` 関数がしていることは以前のコードと全く同じですが、これで独立した再利用可能な部品になりました。

では、加算演算子を使用して、文字列 `character` の繰り返しの前後に１つずつスペース `" "` を連結してください。

`+` という演算子を使用して、次のように文字列を結合できることを思い出してください。

```js
" " + "string"
```

# --hints--

戻り値の先頭にスペースを 1 つ結合してください。

```js
assert.match(padRow(1, 1), /^\s/);
```

戻り値の末尾にスペースを 1 つ結合してください。

```js
assert.match(padRow(1, 1), /\s$/);
```

`padRow()` 関数は、`character` の繰り返しの前後に 1 つずつスペースを加えたものを返すようにしてください。

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
