---
id: 6610c105bbdacc114d6cdc44
title: ステップ 48
challengeType: 1
dashedName: step-48
---

# --description--

関数を使うためには、関数を呼び出す必要があります。 <dfn>関数呼び出し</dfn>は、その場所で関数のコードを実行するという指示です。 関数呼び出しの構文は、関数名の後に丸括弧を付けます。 例えば次のコードは、関数 `test` を定義した後、呼び出しています。

```js
function test() {

}
test();
```

では、関数 `padRow` を呼び出しましょう。

# --hints--

関数 `padRow` を呼び出してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

関数 `padRow` の本体は空にしてください。

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
