---
id: 6610c538372aa61cc0f5b122
title: ステップ 57
challengeType: 1
dashedName: step-57
---

# --description--

JavaScript の変数は特定の<dfn>スコープ</dfn>で使用可能です。 言い換えると、変数が宣言される場所によって、その変数がコードのどこで使えるかが変わります。

最初に説明するスコープは「グローバルスコープ」です。 関数や `for` ループのような「ブロック」の外で宣言された変数は、<dfn>グローバルスコープ</dfn>に存在します。 あなたの変数 `character`、`count`、および `rows` はすべてグローバルスコープにあります。

変数がグローバルスコープにある場合、関数定義内でその変数にアクセスできます。 こちらはグローバル変数 `title` を使った関数の例です。

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

この例は `"Professor Naomi"` を返します。 では `padRow` 関数を変更して、変数 `character` を `name` パラメータの先頭に連結した値を返すようにしましょう。

# --hints--

`padRow` 関数は `character` を `name` の先頭に連結する必要があります。

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

`padRow` 関数は `character + name` の結果を返す必要があります。

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
