---
id: 6610c71600966a2191d3a64a
title: ステップ 59
challengeType: 1
dashedName: step-59
---

# --description--

関数を呼び出すことによって、関数から返される値を利用できます。 関数呼び出しをそのまま戻り値として扱うこともできますし、戻り値を変数に格納してから扱うこともできます。 この方法を使うと、ローカルスコープの変数に割り当てられた値を、関数の外で使用することができます。

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

`"Testing"` の値を関数の外で使用できるように、`padRow` 関数から返しましょう。`return` 文を変更して変数 `test` のみを返すようにしてください。


# --hints--

`padRow` 関数は変数 `test` を返す必要があります。

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
