---
id: 6610c6541c82551f95e765ab
title: ステップ 58
challengeType: 1
dashedName: step-58
---

# --description--

変数は、関数内でも宣言することができます。 そのような変数は<dfn>ローカルスコープ</dfn>や<dfn>ブロックスコープ</dfn>にあると言います。 関数内で宣言された変数は、その関数内でのみ使用できます。 関数の外でその変数にアクセスしようとすると、参照エラー (reference error) が発生します。

実際に確認するために、`padRow` 関数内で `const` を使用して変数 `test` を宣言してください。 変数は `"Testing"` という値で初期化してください。

次に、その関数の下で `test` をコンソールにログ出力してみましょう。 `test` は関数のローカルスコープの外では未定義のため、エラーが発生します。 エラーを確認したら、テストをパスさせて次に進むために `console.log` を削除してください。

# --hints--

関数内で `test` という変数を宣言してください。

```js
assert.match(padRow.toString(), /var\s+test/);
```

`test` を `"Testing"` という値で初期化してください。 セミコロンを忘れないようにしましょう。

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

変数 `test` は `return` キーワードより前に置く必要があります。

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
