---
id: 6610c424b7119919b62932f4
title: ステップ 53
challengeType: 1
dashedName: step-53
---

# --description--

関数が返すのはハードコードされた値でなくて構いません。 例えば変数に格納された値を返すことができます。 仮引数は関数用の特別な変数なので、これを返すこともできます。

では、`padRow` 関数を変更して、仮引数 `name` をそのまま `return` してください。

# --hints--

`padRow` 関数では、仮引数 `name` の値を返してください。

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
