---
id: 6610b8017d1671f2814e8c77
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

初期化せずに変数を宣言した場合、その変数は<dfn>初期化されていない</dfn>状態となります。 現在、変数 `secondCharacter` は初期化されていません。

では、`console.log()` を追加して、変数 `secondCharacter` の値がどうなっているか見てみましょう。

# --hints--

2 つ目の `console.log` を追加する必要があります。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

変数 `secondCharacter` をログに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
