---
id: 6610c1d97b1671140f95cfbb
title: ステップ 50
challengeType: 1
dashedName: step-50
---

# --description--

次に、変数 `call` の値を出力するログ文を追加してください。

# --hints--

`console.log` の呼び出しを追加してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2)
```

変数 `call` をログに出力してください。 セミコロンを忘れないようにしましょう。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*call\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
