---
id: 6610c83b52583e245a079217
title: ステップ 61
challengeType: 1
dashedName: step-61
---

# --description--

これで変数 `call` に `"Testing"` という値が入りました。 そしてこの関数ではもう引数 `name` を使わなくなりました。

なので仮引数 `name` を関数宣言から削除して、`padRow` の呼び出しからも `"CamperChan"` の文字列を削除しましょう。

また、`padRow` 関数内にある `console.log` を両方とも削除してください。
# --hints--

`padRow` 関数に仮引数 `name` を含めないでください。

```js
assert.notMatch(padRow.toString(), /name/);
```

`"CamperChan"` を `padRow` の呼び出しに渡さないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

`padRow` 関数の呼び出し自体は引き続き必要です。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

`return` キーワードの前に `console.log` を含めないでください。

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

`return` キーワードの後に `console.log` を含めないでください。

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
