---
id: 6610bf6fa14d700beed1b109
title: ステップ 87
challengeType: 1
dashedName: step-87
---

# --description--

<dfn>等価</dfn>演算子 `==` は、2 つの値が等しいかどうかをチェックするために使用されます。 2 つの値を比較するには、`value == 8` のようなステートメントを使用します。

では、ループ内の `done++` の下に `if` 文を追加してください。 その if 文では、等価演算子を使用して `done` が `count` と等しいかどうかチェックしてください。

# --hints--

ループ内で `if` 文を使用してください。 これは `done++` の後に追加する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

`if` 文の条件では、等価演算子を使って `done` と `count` を比較してください。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
