---
id: 6610c2d8d67563174fcf96dc
title: ステップ 52
challengeType: 1
dashedName: step-52
---

# --description--

コードに直接値を記述することを、<dfn>ハードコード</dfn>と言います。例えば、いま関数内にある `"Hello!"` という文字列がそうです。 関数内に値をハードコーディングすると、関数を思ったように再利用できない可能性があります。

代わりに、関数に<dfn>仮引数 (parameter)</dfn> を定義することができます。 仮引数は、関数を呼び出すときに値を与えられる特別な変数で、関数内でその関数のコードの結果を動的に変更するために使用できます。

関数に仮引数を追加するには、丸括弧内に変数名を追加します。 例えば、次の `demo` 関数には仮引数 `name` があります。

```js
function demo(name) {

}
```

`name` という仮引数が役に立ちそうなので、`padRow` 関数に同じものを追加してください。

# --hints--

`padRow` 関数に仮引数 `name` を持たせてください。

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
