---
id: 6610c48c4ea0891afa7c4696
title: ステップ 54
challengeType: 1
dashedName: step-54
---

# --description--

コンソールを再確認すると、戻り値を定義したにもかかわらず `padRow` 関数が `undefined` を返していることに気付くでしょう。 これは、関数を**呼び出す**ときに仮引数に値を与える必要があるためです。

関数呼び出しに値を渡すとき、その値は <dfn>実引数 (argument)</dfn> と呼ばれます。 次の例は、仮引数 `name` に実引数として `"Naomi"` を渡して `demo` 関数を呼び出すコードです。

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

では、仮引数 `name` に実引数としてあなたの名前を渡して `padRow` 関数を呼び出してください。 あなたの名前は文字列であるため、引用符を使用する必要があることに注意してください。

# --hints--

`padRow()` の呼び出しに文字列を渡してください。

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
