---
id: 6610c21b3ef82015573ffbbe
title: ステップ 51
challengeType: 1
dashedName: step-51
---

# --description--

変数 `call` に代入するものを定義したはずなのに、値が `undefined` となっているようですね。 これは、関数 `padRow` が値を返していないためです。 デフォルトでは、関数は `undefined` を値として返します。

何か他のものを返すには、`return` キーワードを使用する必要があります。 以下は、文字列`"Functions are cool!"`を返す関数の例です。

```js
function demo() {
  return "Functions are cool!";
}
```

では、関数で `return` キーワードを使用して、文字列 `"Hello!"` を返してください。

# --hints--

関数 `padRow` で `return` キーワードを使用する必要があります。

```js
assert.match(padRow.toString(), /return/);
```

関数 `padRow` が文字列 `"Hello!"` を返すようにしてください。

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
