---
id: 6610c8cfe4cf4d278e35c156
title: ステップ 63
challengeType: 1
dashedName: step-63
---

# --description--

期待通り、あなたの関数は再び `undefined` を返すようになりました。 変数 `call` ももう必要ありませんので、`call` の宣言と、変数 `call` を出力する `console.log` を削除してください。

# --hints--

`call` の宣言は必要ありません。

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

変数 `call` のログ出力は必要はありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
