---
id: 6610b8f6a98d25f4d485a94d
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

初期化されていない変数のデフォルト値は `undefined` です。 これは、未定義の値を表す特別なデータ型です。

この状態から、初期化されていない変数に値を割り当てることができます。 次に例を示します。

```js
let uninitialized;
uninitialized = "assigned";
```

では、変数の宣言の次の行で、変数 `secondCharacter` に文字列 `"Test"` を代入してください。 コンソールを開いて、ログがどのように変化したか確認しましょう。

# --hints--

変数 `secondCharacter` は初期化しないでください。 「初期化」とは、変数を宣言すると同時に値を代入することを意味します。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

`secondCharacter` に代入演算子を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

変数 `secondCharacter` に文字列 `"Test"` を代入してください。

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
