---
id: 6610c16c4fa0df12c0e30675
title: ステップ 49
challengeType: 1
dashedName: step-49
---

# --description--

関数 `padRow` を呼び出していますが、まだその関数呼び出しで何もしていません。 JavaScript のすべての関数は値を<dfn>返します (return)</dfn>。つまり、その関数を呼び出して得られるあらかじめ定義された結果を、他の場所で使用できるように提供します。

では、関数 `padRow` を呼び出した結果を確認するため、変数 `call` を宣言し、既存の `padRow` 呼び出しをその変数に代入してください。

# --hints--

変数 `call` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

変数 `call` の宣言には `const` を使ってください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

変数 `call` に、`padRow` 呼び出しの結果を代入してください。

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
