---
id: 6610b741b54b90f0c0fb3d58
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

変数名が複数の単語で構成される場合、大文字の使い方を定めた命名規則があります。 JavaScript では、<dfn>キャメルケース</dfn>という命名規則が使われます。

キャメルケースでは、変数名の最初の単語はすべて小文字で表記し、それ以降の単語は一文字目を大文字にします。 キャメルケースの例を以下に示します。

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

では、キャメルケースを使った新しい変数 `secondCharacter` を宣言してみましょう。

# --hints--

変数 `secondCharacter` を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

変数 `secondCharacter` を宣言するために `let` を使用してください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

変数 `secondCharacter` に値を代入してはいけません。 セミコロンを忘れないでください。

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
