---
id: 664599653fcd6e97104f9261
title: ステップ 83
challengeType: 1
dashedName: step-83
---

# --description--

`if...else if` 条件がすべて `false` だった場合に、異なるコードを実行したいことがあります。 そのような場合には `else` ブロックを追加します。

`else` ブロックは、`if` ブロックと `else if` ブロックの条件がいずれも満たされなかった場合にのみ評価されます。

`else if` ブロックの後に `else` ブロックを追加した例がこちらです。

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

では、`else if` ブロックに続けて `else` ブロックを追加してください。 `else` ブロックの中では、文字列 `"This is the else block"` をコンソールに出力してください。

コンソールで結果を確認するには、`else if` 文の `<` を `>` に書き換えてください。 これにより条件が `false` になり、`else` ブロックが実行されます。

# --hints--

`else` ブロックが必要です。

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

`else` ブロックでは、文字列 `"This is the else block"` をコンソールに出力する必要があります。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```

