---
id: 66458f0a05df478aa627629e
title: ステップ 82
challengeType: 1
dashedName: step-82
---

# --description--

テキストがまた表示されなくなりましたね！ 空の文字列は `false` と評価される<dfn>偽値</dfn>だとわかりました。 真偽値については後のプロジェクトで詳しく学習します。

JavaScript には `if` 文に加えて <dfn>else if</dfn> 文もあります。 `else if` 文を使用すると、1 つのコードブロック内で複数の条件をチェックできます。

`else if` 文の構文は次のとおりです。

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

最初の条件が `false` の場合、JavaScript は次の条件をチェックします。 そして 2 番目の条件が `false` の場合、3 番目の条件をチェックします。その後も同様です。

では、`if` 文に続けて、`5` が `10` より小さいかどうかをチェックする `else if` 文を追加してください。 そして `else if` 文のコードブロック内で、コンソールに文字列 `"5 is less than 10"` を出力してください。

コンソールに表示される結果を確認してください。

# --hints--

`else if` 文が必要です。

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

`else if` 文では `5` が `10` より小さいかどうかをチェックしてください。

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

コンソールに文字列 `"5 is less than 10"` を出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```

