---
id: 664ed97c55a99e5ffab759dc
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

先に進む前に、一旦ここまでに学習したことを復習しましょう。

`let` キーワードを使用して、`profession` と `age` という名前の変数を宣言してください。 変数 `profession` は文字列 `"teacher"` で初期化してください。一方、変数 `age` は初期化しないでください。

両方の変数をコンソールに出力して、結果を確認してください。

# --hints--

変数 `profession` を宣言してください。

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

変数 `profession` は `let` を使用して宣言してください。

```js
assert.match(code, /let\s+profession/);
```

文字列 `"teacher"` を変数 `profession` に代入してください。

```js
assert.strictEqual(profession, "teacher");
```

変数 `age` を宣言してください。

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

変数 `age` は `let` を使用して宣言してください。

```js
assert.match(code, /let\s+age/);
```

変数 `age` に値を代入してはいけません。

```js
assert.isUndefined(age);
```

`age` をコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

`profession` をコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
