---
id: 664ed97c55a99e5ffab759dc
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

先に進む前に、一旦ここまでに学習したことを復習しましょう。

`let` キーワードを使用して、`profession` と `age` という名前の変数を宣言してください。 変数 `profession` は文字列 `"teacher"` で初期化してください。一方、変数 `age` は初期化しないでください。

両方の変数をコンソールに出力して、結果を確認してください。

# --hints--

You should declare a `profession` variable.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

You should use `let` to declare the `profession` variable.

```js
assert.match(code, /let\s+profession/);
```

You should assign the string `"teacher"` to the `profession` variable.

```js
assert.strictEqual(profession, "teacher");
```

You should declare an `age` variable.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

You should use `let` to declare the `age` variable.

```js
assert.match(code, /let\s+age/);
```

You should not assign a value to the `age` variable.

```js
assert.isUndefined(age);
```

You should log `age` to the console.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

You should log `profession` to the console.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
