---
id: 664ed9344ba3495e1f67f8ea
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

これで変数の宣言と初期化について思い出せましたね。では、復習に使ったコードは削除してください。

変数 `character` は削除しないよう気を付けてください。

# --hints--

変数 `profession` は必要ありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

変数 `age` は必要ありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

`age` を出力する `console.log` ステートメントは必要ありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

`profession` を出力する `console.log` ステートメントは必要ありません。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

変数 `character` は削除しないでください。

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
