---
id: 665732e62eb8b30c249b5c2c
title: ステップ 25
challengeType: 1
dashedName: step-25
---

# --description--

これで、配列について次のレッスンに進む準備ができました。

一つ前のステップで書いたコードをすべて削除してください。

# --hints--

変数 `cities` を参照しているコードをすべて削除する必要があります。

```js
assert.notMatch(__helpers.removeJSComments(code), /cities/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let cities = ["London", "New York", "Mumbai"];
console.log(cities);
cities[cities.length - 1] = "Mexico City";
console.log(cities);
--fcc-editable-region--
console.log(rows);
```
