---
id: 665732da4815b70bb083915e
title: ステップ 24
challengeType: 1
dashedName: step-24
---


# --description--

ここまでの数ステップで、配列の操作について学びました。 ここで少し復習の時間を取りましょう。

まず、変数 `cities` を宣言し、それを文字列 `"London"`、`"New York"`、`"Mumbai"` の配列として初期化してください。 次に、その変数をコンソールにログ出力してください。

ログ出力した後、`cities` の最後の要素を文字列 `"Mexico City"` に変更し、もう一度変数 `cities` をログ出力してください。

コードが正しければ、コンソールに次の出力が表示されます。

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

`let` を使って変数 `cities` を宣言してください。

```js
assert.match(code, /let\s+cities/);
```

変数 `cities` に、文字列 `"London"`、`"New York"`、`"Mumbai"` の配列を割り当ててください。

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

`console.log()` を使って、配列 `cities` 全体をコンソールにログ出力してください。

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

配列 `cities` の最後の要素を、文字列 `"Mexico City"` に更新してください。 配列の最後の要素にアクセスするには `array[array.length - 1]` という書き方が使えることを思い出してください。

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

コードには `console.log(cities)` ステートメントが 2 つあるはずです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
