---
id: 66643c9a9e0ad1c787f5fe2b
title: ステップ 56
challengeType: 1
dashedName: step-56
---

# --description--

復習が完了したので、関数 `addTwoNumbers`、変数 `sum`、およびログステートメントを削除してください。

# --hints--

コード内に関数 `addTwoNumbers` が無いようにしてください。

```js
assert.notMatch(code, /addTwoNumbers/);
```

コード内に変数 `sum` が無いようにしてください。

```js
assert.notMatch(code, /sum/);
```

変数 `sum` をログに出力しないようにしてください。

```js
assert.notMatch(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
function addTwoNumbers(num1, num2) {
  return num1 + num2;
}

const sum = addTwoNumbers(5, 10);
console.log(sum)

--fcc-editable-region--

const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
