---
id: 68de40f53fcaf106f51cc20f
title: Costruisci un lettore audio e video HTML
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `h1` per il titolo principale della pagina.
2. Dovresti avere due elementi `section`.
3. All’interno del primo elemento `section`, dovresti avere un elemento `h2` per il titolo del video in riproduzione.
4. Sotto l’elemento `h2`, dovresti avere un elemento `video` con gli attributi `controls` e `width`. L’attributo `width` dovrebbe essere impostato su `640`.
5. All’interno dell’elemento `video`, dovresti avere un elemento `source` con un attributo `src` che punta a un file video e un attributo `type`.
   - Puoi usare `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.
6. All’interno del secondo elemento `section`, dovresti avere un elemento `h2` per il titolo della canzone in riproduzione.
7. Sotto l’elemento `h2`, dovresti avere un elemento `audio` con gli attributi `controls` e `loop`, e un attributo `src` che punta a un file audio.
   - Puoi usare `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`.
   - Oppure `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`.

# --hints--

Dovresti avere un elemento `h1` per il titolo principale della pagina.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

Il tuo `h1` non dovrebbe essere vuoto.

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

Dovresti avere solo un elemento `h1`

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Dovresti avere due elementi `section`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

All’interno del primo elemento `section`, dovresti avere un elemento `h2` per il titolo del video in riproduzione.

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Sotto l’elemento `h2`, dovresti avere un elemento `video`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

L’elemento `video` dovrebbe avere un attributo `controls`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

L’elemento `video` dovrebbe avere un attributo `width` impostato su `640`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

All’interno dell’elemento `video`, dovresti avere un elemento `source`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

L’elemento `source` dovrebbe avere un attributo `src` che punta a un file video.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

L’elemento `source` dovrebbe avere un attributo `type`

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

L’attributo `type` dovrebbe corrispondere al tipo di media dell’estensione del file dell’attributo `src`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

All’interno del secondo elemento `section`, dovresti avere un elemento `h2` per il titolo della canzone in riproduzione.

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Tutti i tuoi elementi `h2` dovrebbero contenere del testo.

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

Sotto l’elemento `h2`, dovresti avere un elemento `audio`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

L’elemento `audio` dovrebbe avere un attributo `controls`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

L’elemento `audio` dovrebbe avere un attributo `loop`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

L’elemento `audio` dovrebbe avere un attributo `src` che punta a un file audio.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
