---
id: 5f344fc1520b6719f2e35605
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--

There will be two sections on the menu, one for coffees and one for desserts. Add a `section` element within the `main` element so you have a place to put all the coffees available.

# --hints--

Unapaswa kuwa na tagi ya kufungua ya `<section>`.

```js
assert(code.match(/<section\s*>/i));
```

Unapaswa kuwa na tagi ya kufunga ya `</section>`.

```js
assert(code.match(/<\/section\s*>/i));
```

Hupaswi kubadilisha kipengele chako cha `main` kilichopo. Hakikisha kuwa hukufuta tagi ya kufunga.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Kipengele chako cha `section` kinapaswa kuwa ndani ya kipengele chako cha `main`.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
