---
id: 5f3c866d5414453fc2d7b480
title: Hatua ya 31
challengeType: 0
dashedName: step-31
---

# --description--

Starting below the existing coffee/price pair, add the following coffee and prices using `article` elements with two nested `p` elements inside each. As before, the first `p` element's text should contain the coffee flavor and the second `p` element's text should contain the price.

```bash
Caramel Macchiato 3.75
Pumpkin Spice 3.50
Hazelnut 4.00
Mocha 4.50
```

# --hints--

Unapaswa kuwa na vipengele vitano vya `article`.

```js
assert.lengthOf(document.querySelectorAll('article'),5);
```

Kila kipengele cha `article` kinafaa kuwa na vipengele viwili vya `p`.

```js
const articles = document.querySelectorAll('article');
assert.lengthOf(articles?.[0]?.children, 2);
assert.lengthOf(articles?.[1]?.children, 2);
assert.lengthOf(articles?.[2]?.children, 2);
assert.lengthOf(articles?.[3]?.children, 2);
assert.lengthOf(articles?.[4]?.children, 2);
```

Kipengele chako cha kwanza cha `article` kinapaswa kuwa na vipengele vya `p` vyenye maandishi `French Vanilla` na `3.00`.

```js
const children = document.querySelector('article')?.children;
assert.match(children?.[0]?.innerText, /French Vanilla/i);
assert.match(children?.[1]?.innerText, /3\.00/i);
```

Kipengele chako cha pili cha `article` kinapaswa kuwa na vipengele vya `p` vyenye maandishi `Caramel Macchiato` na `3.75`.

```js
const children = document.querySelectorAll('article')?.[1]?.children;
assert.match(children?.[0]?.innerText, /Caramel Macchiato/i);
assert.match(children?.[1]?.innerText, /3\.75/i);
```

Kipengele chako cha tatu cha `article` kinapaswa kuwa na vipengele vya `p` vyenye maandishi `Pumpkin Spice` na `3.50`.

```js
const children = document.querySelectorAll('article')?.[2]?.children;
assert.match(children?.[0]?.innerText, /Pumpkin Spice/i);
assert.match(children?.[1]?.innerText, /3\.50/i);
```

Kipengele chako cha nne cha `article` kinapaswa kuwa na vipengele vya `p` vyenye maandishi `Hazelnut` na `4.00`.

```js
const children = document.querySelectorAll('article')?.[3]?.children;
assert.match(children?.[0]?.innerText, /Hazelnut/i);
assert.match(children?.[1]?.innerText, /4\.00/i);
```

Kipengele chako cha tano cha `article` kinapaswa kuwa na vipengele vya `p` vyenye maandishi `Mocha` na `4.50`.

```js
const children = document.querySelectorAll('article')?.[4]?.children;
assert.match(children?.[0]?.innerText, /Mocha/i);
assert.match(children?.[1]?.innerText, /4\.50/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
