---
id: 587d7faf367417b2b2512be9
title: Надсилання даних методом JavaScript XMLHttpRequest
challengeType: 6
forumTopicId: 301504
dashedName: post-data-with-the-javascript-xmlhttprequest-method
---

# --description--

У попередніх прикладах ви отримували дані з зовнішнього ресурсу. Ви також можете надсилати дані на зовнішній ресурс, якщо цей ресурс підтримує AJAX-запити і ви знаєте URL.

Метод `XMLHttpRequest` у JavaScript також використовується для надсилання даних на сервер. Ось приклад:

```js
const xhr = new XMLHttpRequest();
xhr.open('POST', url, true);
xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
xhr.onreadystatechange = function () {
  if (xhr.readyState === 4 && xhr.status === 201) {
    const serverResponse = JSON.parse(xhr.response);
    document.getElementsByClassName('message')[0].textContent =
      serverResponse.userName + serverResponse.suffix;
  }
};
const body = JSON.stringify({ userName: userName, suffix: ' loves cats!' });
xhr.send(body);
```

Ви вже бачили кілька таких методів раніше. Тут метод `open` ініціалізує запит як `POST` до вказаного URL зовнішнього ресурсу і передає `true` як третій параметр — що означає виконання операції асинхронно.

Метод `setRequestHeader` встановлює значення заголовка HTTP-запиту, який містить інформацію про відправника та запит. Його потрібно викликати після методу `open`, але до методу `send`. Два параметри — це назва заголовка і значення, яке потрібно встановити як тіло цього заголовка.

Далі, обробник події `onreadystatechange` реагує на зміну стану запиту. `readyState` зі значенням `4` означає, що операція завершена, а `status` зі значенням `201` означає, що запит був успішним. Тому можна оновити HTML документа.

Нарешті, метод `send` надсилає запит із значенням `body`. `body` складається з `userName` і `suffix` ключів.

# --instructions--

Оновіть код так, щоб він робив `POST` запит до кінцевої точки API. Потім введіть своє ім’я у поле введення і натисніть `Send Message`. Ваша AJAX-функція має замінити `Reply from Server will be here.` на дані з сервера. Відформатуйте відповідь так, щоб відображалося ваше ім’я, додане до тексту ` loves cats`.

# --hints--

Ваш код має створити новий `XMLHttpRequest`.

```js
assert.match(code, /new\s*XMLHttpRequest\(\s*\)/g);
```

Ваш код має використовувати метод `open` для ініціалізації `POST` запиту до сервера.

```js
assert.match(code, /\.open\(\s*('|")POST\1\s*,\s*url\s*,\s*true\s*\)/g);
```

Ваш код має використовувати метод `setRequestHeader`.

```js
assert.match(
  code,
  /\.setRequestHeader\(\s*('|")Content-Type\1\s*,\s*('|")application\/json;\s*charset=UTF-8\2\s*\)/g
);
```

Ваш код має мати обробник події `onreadystatechange`, встановлений на функцію.

```js
assert.match(code, /\.onreadystatechange\s*=/g);
```

Ваш код має отримати елемент з класом `message` і змінити його `textContent` на `userName loves cats`.

```js
assert.match(
  code,
  /document\.getElementsByClassName\(\s*('|")message\1\s*\)\[0\]\.textContent\s*=\s*.+?\.userName\s*\+\s*.+?\.suffix/g
);
```

Ваш код має використовувати метод `send`.

```js
assert.match(code, /\.send\(\s*body\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line

      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message box">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line
      const xhr = new XMLHttpRequest();
      xhr.open('POST', url, true);
      xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
      xhr.onreadystatechange = function () {
        if (xhr.readyState === 4 && xhr.status === 201) {
          const serverResponse = JSON.parse(xhr.response);
          document.getElementsByClassName('message')[0].textContent =
            serverResponse.userName + serverResponse.suffix;
        }
      };
      const body = JSON.stringify({
        userName: userName,
        suffix: ' loves cats!'
      });
      xhr.send(body);
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```
