---
id: 587d7faf367417b2b2512be8
title: "Отримати дані геолокації, щоб знайти GPS-координати користувача"
challengeType: 6
forumTopicId: 18188
dashedName: get-geolocation-data-to-find-a-users-gps-coordinates
---

# --description--

Ще одна крута річ, яку ви можете зробити, — це отримати доступ до поточного місцезнаходження користувача. Кожен браузер має вбудований навігатор, який може надати цю інформацію.

Навігатор отримає поточну довготу та широту користувача.

Ви побачите запит на дозвіл або блокування цього сайту для визначення вашого поточного місцезнаходження. Завдання можна виконати в будь-якому випадку, якщо код правильний.

Якщо вибрати дозвіл, текст на телефоні з вихідними даними зміниться на вашу широту та довготу.

Ось код, який це робить:

```js
if (navigator.geolocation) {
  navigator.geolocation.getCurrentPosition(function (position) {
    document.getElementById('data').innerHTML =
      'latitude: ' +
      position.coords.latitude +
      '<br>longitude: ' +
      position.coords.longitude;
  });
}
```

Спочатку перевіряється, чи існує об’єкт `navigator.geolocation`. Якщо так, викликається метод `getCurrentPosition` цього об’єкта, який ініціює асинхронний запит на позицію користувача. Якщо запит успішний, виконується функція зворотного виклику в методі. Ця функція отримує значення широти та довготи об’єкта `position` за допомогою точкової нотації та оновлює HTML.

# --instructions--

Додайте приклад коду всередину тегів `script`, щоб перевірити поточне місцезнаходження користувача та вставити його в HTML.

# --hints--

Ваш код має використовувати `navigator.geolocation` для доступу до поточного місцезнаходження користувача.

```js
assert.match(code, /navigator\.geolocation\.getCurrentPosition/g);
```

Ваш код має використовувати `position.coords.latitude` для відображення широти користувача.

```js
assert.match(code, /position\.coords\.latitude/g);
```

Ваш код має використовувати `position.coords.longitude` для відображення довготи користувача.

```js
assert.match(code, /position\.coords\.longitude/g);
```

Ви маєте відобразити позицію користувача всередині елемента `div` за допомогою `id="data"`.

```js
assert.match(
  code,
  /document\.getElementById\(\s*('|")data\1\s*\)\.innerHTML/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  // Add your code below this line

  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data"></div>
```

# --solutions--

```html
<script>
  // Add your code below this line
  if (navigator.geolocation) {
    navigator.geolocation.getCurrentPosition(function(position) {
      document.getElementById('data').innerHTML = "latitude: " + position.coords.latitude + "<br>longitude: " + position.coords.longitude;
    });
  }
  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data">

</div>

</section>
```
