---
id: 587d7fae367417b2b2512be3
title: Отримати JSON за допомогою методу JavaScript XMLHttpRequest
challengeType: 6
forumTopicId: 301502
dashedName: get-json-with-the-javascript-xmlhttprequest-method
---

# --description--

Ви також можете запитувати дані з зовнішнього джерела. Саме тут у гру вступають API.

Пам’ятайте, що API — або інтерфейси програмування застосунків — це інструменти, які комп’ютери використовують для спілкування один з одним. Ви навчитеся оновлювати HTML за допомогою даних, які ми отримуємо з API, використовуючи технологію під назвою AJAX.

Більшість веб-API передають дані у форматі, який називається JSON. JSON означає JavaScript Object Notation.

Синтаксис JSON дуже схожий на нотацію літерального об’єкта JavaScript. JSON має властивості об’єкта та їхні поточні значення, розташовані між `{` і `}`.

Ці властивості та їхні значення часто називають «парами ключ-значення».

Однак JSON, який передається через API, надсилається як `bytes`, а ваш застосунок отримує його як `string`. Їх можна перетворити на об’єкти JavaScript, але за замовчуванням вони не є об’єктами JavaScript. Метод `JSON.parse` аналізує рядок і створює описаний ним об’єкт JavaScript.

Ви можете запросити JSON з Cat Photo API від freeCodeCamp. Ось код, який ви можете вставити у свій обробник події click, щоб це зробити:

```js
const req = new XMLHttpRequest();
req.open('GET', '/json/cats.json', true);
req.send();
req.onload = function () {
  const json = JSON.parse(req.responseText);
  document.getElementsByClassName('message')[0].innerHTML =
    JSON.stringify(json);
};
```

Ось повторення того, що робить кожна частина. Об’єкт JavaScript `XMLHttpRequest` має низку властивостей і методів, які використовуються для передачі даних. Спочатку створюється екземпляр об’єкта `XMLHttpRequest` і зберігається у змінній `req`. Далі метод `open` ініціалізує запит — у цьому прикладі запитуються дані з API, тому це `GET` запит. Другий аргумент для `open` — це URL API, з якого ви запитуєте дані. Третій аргумент — булеве значення, де `true` робить запит асинхронним. Метод `send` надсилає запит. Нарешті, обробник події `onload` аналізує отримані дані і застосовує метод `JSON.stringify`, щоб перетворити об’єкт JavaScript у рядок. Цей рядок потім вставляється як текст повідомлення.

# --instructions--

Оновіть код, щоб створити і надіслати `GET` запит до Cat Photo API від freeCodeCamp. Потім натисніть кнопку `Get Message`. Ваша AJAX-функція замінить текст `The message will go here` на сирі вихідні дані JSON з API.

# --hints--

Ваш код має створити новий `XMLHttpRequest`.

```js
assert.match(code, /new\s+?XMLHttpRequest\(\s*?\)/g);
```

Ваш код має використовувати метод `open`, щоб ініціалізувати `GET` запит до Cat Photo API від freeCodeCamp.

```js
assert.match(
  code,
  /\.open\(\s*('|")GET\1\s*,\s*('|")\/json\/cats\.json\2\s*,\s*true\s*\)/g
);
```

Ваш код має використовувати метод `send`, щоб надіслати запит.

```js
assert.match(code, /\.send\(\s*\)/g);
```

Ваш код має мати обробник події `onload`, встановлений на функцію.

```js
assert.match(
  code,
  /\.onload\s*=\s*(function|\(\s*\))\s*(\(\s*\)|\=\>)\s*{/g
);
```

Ваш код має використовувати метод `JSON.parse`, щоб проаналізувати `responseText`.

```js
assert.match(code, /JSON\s*\.parse\(\s*.*\.responseText\s*\)/g);
```

Ваш код має отримати елемент з класом `message` і змінити його внутрішній HTML на рядок даних JSON.

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.innerHTML\s*?=\s*?JSON\.stringify\(.+?\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = () => {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
