---
id: 587d7fad367417b2b2512be1
title: Обробка подій click у JavaScript за допомогою властивості onclick
challengeType: 6
forumTopicId: 301503
dashedName: handle-click-events-with-javascript-using-the-onclick-property
---

# --description--

Ви хочете, щоб ваш код виконувався лише після того, як сторінка повністю завантажиться. Для цього можна приєднати JavaScript-подію до документа під назвою `DOMContentLoaded`. Ось код, який це робить:

```js
document.addEventListener('DOMContentLoaded', function () {});
```

Ви можете реалізувати обробники подій, які розміщуються всередині функції `DOMContentLoaded`. Ви можете реалізувати обробник події `onclick`, який спрацьовує, коли користувач натискає на елемент `#getMessage`, додавши такий код:

```js
document.getElementById('getMessage').onclick = function () {};
```

# --instructions--

Додайте обробник події click всередині функції `DOMContentLoaded` для елемента з ідентифікатором `getMessage`.

# --hints--

Ваш код має використовувати метод `document.getElementById` для вибору елемента з ідентифікатором `getMessage`.

```js
assert.match(code, /document\s*\.getElementById\(\s*('|")getMessage\1\s*\)/g);
```

Ваш код має додати обробник події `onclick`.

```js
assert.isFunction(document.getElementById('getMessage').onclick);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    document.getElementById('getMessage').onclick = function () {};
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
