---
id: 587d7fae367417b2b2512be4
title: Отримати доступ до JSON-даних з API
challengeType: 6
forumTopicId: 301499
dashedName: access-the-json-data-from-an-api
---

# --description--

У попередньому завданні ви побачили, як отримати JSON-дані з API freeCodeCamp Cat Photo.

Тепер ви детальніше розглянете отримані дані, щоб краще зрозуміти формат JSON. Нагадаємо деякі позначення в JavaScript:

<blockquote>[ ] -> Квадратні дужки позначають масив.<br>{ } -> Фігурні дужки позначають об’єкт.<br>" " -> Подвійні лапки позначають рядок. Вони також використовуються для імен ключів у JSON.</blockquote>

Розуміння структури даних, які повертає API, важливе, бо це впливає на те, як ви отримуєте потрібні значення.

Праворуч натисніть на кнопку `Get Message`, щоб завантажити JSON API freeCodeCamp Cat Photo у HTML.

Перший і останній символ, який ви бачите в JSON-даних, — це квадратні дужки `[ ]`. Це означає, що повернені дані — це масив. Другий символ у JSON-даних — фігурна `{` дужка, яка починає об’єкт. Ретельно придивившись, ви побачите, що є три окремі об’єкти. JSON-дані — це масив із трьох об’єктів, де кожен об’єкт містить інформацію про фото кота.

Раніше ви дізналися, що об’єкти містять «пари ключ-значення», які розділені комами. У прикладі Cat Photo перший об’єкт має `"id":0`, де `id` — це ключ, а `0` — відповідне значення. Аналогічно, є ключі для `imageLink`, `altText` і `codeNames`. Кожен об’єкт фото кота має ці самі ключі, але з різними значеннями.

Ще одна цікава «пара ключ-значення» в першому об’єкті — `"codeNames":["Juggernaut","Mrs. Wallace","ButterCup"]`. Тут `codeNames` — ключ, а його значення — масив із трьох рядків. Можна мати масиви об’єктів, а також ключ із масивом як значенням.

Пам’ятайте, як отримувати доступ до даних у масивах і об’єктах. Масиви використовують дужкову нотацію, щоб отримати елемент за індексом. Об’єкти використовують або дужкову, або точкову нотацію, щоб отримати значення певної властивості. Ось приклад, який виводить властивість `altText` першого фото кота — зверніть увагу, що розпарсені JSON-дані в редакторі збережені у змінній `json`:

```js
console.log(json[0].altText);
```

У консолі буде виведено рядок `A white cat wearing a green helmet shaped melon on its head.`.

# --instructions--

Для кота з `id` рівним 2 виведіть у консоль друге значення в масиві `codeNames`. Ви повинні використати дужкову та точкову нотації для об’єкта (який збережений у змінній `json`), щоб отримати це значення.

# --hints--

Ваш код має використовувати дужкову та точкову нотації, щоб отримати правильне ім’я коду, і вивести `Loki` у консоль.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*json\s*\[2\]\s*(\.\s*codeNames|\[\s*('|`|")codeNames\2\s*\])\s*\[\s*1\s*\]\s*\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line

        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line
        console.log(json[2].codeNames[1]);
        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
