---
id: 5ccfad82bb2dc6c965a848e5
title: Отримання JSON за допомогою методу fetch у JavaScript
challengeType: 6
forumTopicId: 301501
dashedName: get-json-with-the-javascript-fetch-method
---

# --description--

Ще один спосіб отримати зовнішні дані — це використати метод JavaScript `fetch()`. Він є еквівалентним до `XMLHttpRequest`, але синтаксис вважається простішим для розуміння.

Ось код для виконання GET-запиту до `/json/cats.json`

```js
fetch('/json/cats.json')
  .then(response => response.json())
  .then(data => {
    document.getElementById('message').innerHTML = JSON.stringify(data);
  });
```

Примітка: Метод `fetch()` використовує `GET` як метод `HTTP` за замовчуванням. Це означає, що вам не потрібно явно його вказувати для базового отримання даних.

Розгляньте кожну частину цього коду.

Перший рядок — це той, що робить запит. Отже, `fetch(URL)` робить `GET` запит до вказаного URL. Метод повертає проміс.

Після повернення проміса, якщо запит був успішним, виконується метод `then`, який приймає відповідь і перетворює її у формат JSON.

Метод `then` також повертає проміс, який обробляється наступним методом `then`. Аргумент у другому `then` — це JSON-об’єкт, який ви шукаєте!

Тепер він вибирає елемент, який отримає дані, використовуючи `document.getElementById()`. Потім змінює HTML-код елемента, вставляючи рядок, створений із JSON-об’єкта, отриманого з запиту.

# --instructions--

Оновіть код, щоб створити і надіслати `GET` запит до API freeCodeCamp Cat Photo. Але цього разу використайте метод `fetch` замість `XMLHttpRequest`.

# --hints--

Ваш код має використовувати отримані дані, щоб замінити внутрішній HTML

```js
const catData = 'dummy data';
const ref = fetch;
fetch = () => Promise.resolve({ json: () => catData });

  try {
    document.getElementById('getMessage').click();
    await new Promise((resolve, reject) => setTimeout(() => resolve(), 250));
  } catch (error) {
    console.log(error);
  } finally {
    fetch = ref;
    assert.equal(
      document.getElementById('message').textContent,
      JSON.stringify(catData)
    );
  }

```

Ваш код має зробити `GET` запит з `fetch`.

```js
assert.match(code, /fetch\s*\(\s*('|")\/json\/cats\.json\1\s*\)/g);
```

Ваш код має використовувати `then` для перетворення відповіді у JSON.

```js
assert.match(
  code,
  /\.then\s*\(\s*\(?(?<var>\w+)\)?\s*=>\s*\k<var>\s*\.json\s*\(\s*\)\s*\)/g
);
```

Ваш код має використовувати `then` для обробки даних, перетворених у JSON іншим `then`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.then\(\(?\w+\)?=>{[^}]*}\)/g);
```

Ваш код має отримати елемент з ідентифікатором `message` і змінити його внутрішній HTML на рядок JSON-даних.

```js
assert.match(
  __helpers.removeWhiteSpace(code),
  /document\.getElementById\(('|")message\1\)\.innerHTML=JSON\.stringify\(?\w+\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      fetch('/json/cats.json')
        .then(response => response.json())
        .then(data => {
          document.getElementById('message').innerHTML = JSON.stringify(data);
        });
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
