---
id: 673b567e3ba535dda140d278
title: 构建投票系统
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

在本实验中，你将构建一个投票系统，该系统使用 `Map` 来创建轮询，使用 `Set` 来防止重复投票。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该初始化一个 `poll` 变量为一个新建的 `Map` 对象。

2. 你应该有一个函数 `addOption`，它接受一个参数 `option`。

3. 在 `addOption` 函数中：

   - 如果 `option` 尚不存在于轮询中，应将其添加到轮询中，值为一个空的 `Set` 以跟踪投票者。你还应返回消息 `Option "<option>" added to the poll.`

   - 如果 `option` 已经存在，它应返回消息 `Option "<option>" already exists.`

   - 如果你尝试添加一个空选项，函数应返回消息 `Option cannot be empty.`

4. 你应该有一个函数 `vote`，它接受两个参数，`option`（投票选项）和 `voterId`（投票者的唯一 ID）。

5. 在 `vote` 函数中：

   - 如果 `option` 不存在于轮询中，函数应返回消息 `Option "<option>" does not exist.`

   - 如果存在该 `option`，函数应查看 `voterId` 是否已对该 `option` 投票。

   - 如果投票者已经投票，该函数应返回消息 `Voter <voterId> has already voted for "<option>"`。

   - 如果投票者尚未投票，应将 `voterId` 添加到此选项的 `voters` 的 `Set` 中。该函数应返回消息 `Voter <voterId> voted for "<option>"`。

6. 你的 `poll` 中应至少有三个选项。

7. 你的 `poll` 应该至少有三票。

8. 你应该有一个函数 `displayResults`，它以以下形式返回轮询结果：

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

你应该有一个 `poll` 变量初始化为一个新建的 `Map` 对象。

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

你应该定义一个函数 `addOption`，它接受一个参数 `option`。

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

你应该定义一个函数 `vote`，它接受两个参数，`option` 和 `voterId`。

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

你应该定义一个函数 `displayResults` 来显示轮询结果。

```js
assert.isFunction(displayResults);
```

你的 `poll` 中应至少有三个选项。

```js
assert.isAtLeast(poll.size, 3);
```

你的 `poll` 应该至少有三票。

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

你应该确保每个投票选项映射到一个 `Set` 对象。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` 应该返回 `Option "Egypt" added to the poll.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

添加一个空选项应返回 `"Option cannot be empty."`

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

当已经添加了 `Turkey` 时，`addOption("Turkey")` 应该返回 `Option "Turkey" already exists.`。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

当 `Malaysia` 存在于投票选项中时，`vote("Malaysia", "traveler1")` 应该返回 `Voter traveler1 voted for "Malaysia".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` 应该更新该选项的投票者 `Set`。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

当 `traveler1` 尝试再次为 `Algeria` 投票时，`vote("Algeria", "traveler1")` 应该返回 `Voter traveler1 has already voted for "Algeria".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

重复投票不应增加 `Set` 的大小。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

当 `Nigeria` 不在投票选项中时，`vote("Nigeria", "traveler2")` 应该返回 `Option "Nigeria" does not exist.`。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

一个唯一的选项应能够接收多个投票。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

当选项 `Turkey` 获得 2 票，选项 `Morocco` 获得 1 票，且选项 `Spain` 没有获得任何票时，`displayResults()` 应该返回 `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"`。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` 应该以正确的形式返回结果。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
