---
id: 66cc095dfe1682753d2ab030
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Um dos conceitos que você aprendeu nos vídeos da aula anterior foi como retornar valores de uma função.

Aqui está um lembrete de como retornar um valor de uma função:

```js
function myFunction() {
  return "Hello World";
}
```

Dentro da função `addTwoAndSeven`, retorne a soma de `2` e `7`.

# --hints--

A função `addTwoAndSeven` deve retornar a soma de `2` e `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Your `addTwoAndSeven` function should return the sum using the `+` operator.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {

}
--fcc-editable-region--
```
