---
id: 66cc0a9e06e00b75d6782be9
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

In the previous lecture videos, you learned how to call (invoke) a function. Chamar uma função significa executar o código dentro da função.

Aqui está um lembrete de como chamar uma função:

```js
function myFunction() {
  return 'Hello World';
}

// function call
myFunction();
```

Adicione uma instrução `console.log` e, dentro dela, chame a função `addTwoAndSeven`.

Agora você deve ver a soma de `2` e `7` registrada no console.

# --hints--

Você deve ter um `console.log` no código.

```js
assert.match(code, /console\.log(.*)/);
```

Você deve chamar a função `addTwoAndSeven` dentro do `console.log`.

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}
--fcc-editable-region--

--fcc-editable-region--
```
