---
id: 66cc06012e46aa6bc9b8c001
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você vai revisar como trabalhar com funções ao criar uma calculadora.

In the previous lectures, you learned how to declare functions like this:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

Comece criando uma função chamada `addTwoAndSeven`.

Você pode optar por usar a sintaxe de função regular ou a sintaxe das arrow functions.

# --hints--

Você deve declarar uma função chamada `addTwoAndSeven`.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
