---
id: 66cc0f1ae40802781b2ea972
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Nos vídeos da aula anterior, você aprendeu como trabalhar com parâmetros de função.

Um parâmetro de função é uma variável definida na declaração da função e que atua como um placeholder.

Aqui está um exemplo de uma função que tem um parâmetro:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Declare uma função chamada `calculateSum`, que receba dois parâmetros, `num1` e `num2`.

# --hints--

Você deve declarar uma função chamada `calculateSum`.

```js
assert.isFunction(calculateSum);
```

A função `calculateSum` deve receber dois parâmetros, `num1` e `num2`.

```js
assert.match(code, /calculateSum\s*(?:=\s*)?\(\s*num1\s*,\s*num2\s*\)\s*/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
