---
id: 66cc1a3a39aef47d6473cb2f
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

O próximo passo é adicionar a funcionalidade de divisão à calculadora.

Declare uma função chamada `calculateQuotient` que receba dois parâmetros, `num1` e `num2`.

Dentro da função, retorne o quociente dos dois parâmetros.

Em seguida, adicione um `console.log` que chame a função `calculateQuotient` com os argumentos `7` e `11`.

# --hints--

Você deve ter uma função chamada `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

A função `calculateQuotient` deve ter dois parâmetros chamados `num1` e `num2`.

```js
const regex = __helpers.functionRegex('calculateQuotient', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

A função `calculateQuotient` deve retornar o quociente dos dois parâmetros.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

Você deve ter um `console.log` que chame a função `calculateQuotient` com os argumentos `7` e `11`.

```js
assert.match(code, /console\.log\s*\(\s*calculateQuotient\s*\(\s*7\s*,\s*11\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
