---
id: 66cc14ca14e65e7a9ebecb80
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Agora que a calculadora pode somar dois números, é hora de criar uma função que subtraia dois números.

Declare uma função chamada `calculateDifference` que receba dois parâmetros, `num1` e `num2`.

Dentro da função, retorne a diferença entre os dois parâmetros.

# --hints--

Você deve ter uma função chamada `calculateDifference`.

```js
assert.isFunction(calculateDifference);
```

A função `calculateDifference` deve ter dois parâmetros chamados `num1` e `num2`.

```js
assert.match(code, /calculateDifference\s*(?:=\s*)?\(\s*num1\s*,\s*num2\s*\)\s*/);
```

A função `calculateDifference` deve retornar a diferença entre os dois parâmetros.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
