---
id: 66cc0de2920bec775f610424
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Mesmo que as funções funcionem como esperado, há muita repetição no código.

Por estar criando uma calculadora, você não quer ter que criar uma função para cada possível combinação de números que deseja somar.

Seria melhor criar uma única função reutilizável que possa adicionar dois números quaisquer.

Remova todo o código que você escreveu até agora.

Na próxima etapa, você revisará como trabalhar com parâmetros e argumentos em funções.

# --hints--

Você não deve ter uma função chamada `addTwoAndSeven`.

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

Você não deve ter uma chamada de função para `addTwoAndSeven`.

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

Você não deve ter uma função chamada `addThreeAndFour`.

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

Você não deve ter uma chamada de função para `addThreeAndFour`.

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
