---
id: 5dc174fcf86c76b9248c6eb2
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

超文本标记语言元素有像 `<h1>` 这样的起始标签和像 `</h1>` 这样的结束标签。元素显示的文本位于其起始标签和结束标签之间。

将下面 `h1` 元素的文本从 `Hello World` 更改为 `CatPhotoApp`，并观察浏览器预览中的变化。

完成后，按“检查你的代码”按钮查看代码是否正确。

# --hints--

代码中应包含文本 `CatPhotoApp`。 你可能想要检查你的拼写。

```js
assert(code.match(/catphotoapp/i));
```

你的 `h1` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('h1'));
```

你的 `h1` 标签应为小写。按照惯例，所有超文本标记语言标签均以小写书写。

```js
assert.notMatch(code, /<\/?H1>/);
```

你的 `h1` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert(code.match(/<\/h1\>/));
```

你有不止一个 `h1` 元素。 删除多余的 `h1` 元素。

```js
assert(document.querySelectorAll('h1').length === 1);
```

你的 `h1` 元素的文本应该是 `CatPhotoApp`。 你要么遗漏了文本或拼写有误，要么没有把文本放在 `h1` 元素的开始和结束标签之间。

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

