---
id: 5dc24073f86c76b9248c6ebb
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>属性</dfn>是位于元素的开始标签中的特殊词汇，用于控制元素的行为。 在 `img` 元素中的 `src` 属性指定图像的 URL（图像所在的位置）。

下面是一个 `img` 元素的 `src` 属性指向 freeCodeCamp logo 的例子：

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

在现有的 `img` 元素中，添加一个 `src` 属性，URL 如下：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

你的代码应该有一个 `img` 元素。 你可能已经删除了 `img` 元素，或者你没有用引号包围 `src` 属性的值。

```js
assert(document.querySelector('img'));
```

你的 `img` 元素应该有一个 `src` 属性。 你要么遗漏了属性，要么拼写有误。 确保元素名称和属性名称之间有一个空格。

```js
assert(document.querySelector('img').src);
```

你的 `img` 元素的 `src` 属性应设置为 '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'。 你要么遗漏了 URL，要么拼写有误。 URL 的大小写很重要。

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

虽然你已经将 `img` 元素的 `src` 设置为正确的 URL，但建议始终将属性的值用引号括起来。

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

你的 `img` 元素应该以 `>` 或 `/>` 结束。

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

