---
id: 5dc24614f86c76b9248c6ebd
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

你可以使用锚元素（`a`）链接到另一个页面。

这是一个链接到 `https://www.freecodecamp.org` 的示例：

```html
<a href="https://www.freecodecamp.org"></a>
```

在段落之后添加一个链接到 `https://freecatphotoapp.com` 的锚元素。 At this point, the link won’t show up in the preview.

# --hints--

你的锚元素（`a`）应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert(document.querySelector('a'));
```

你的锚元素（`a`）应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert(code.match(/<\/a\>/));
```

你的锚元素（`a`）应该在 `p` 元素下方。 你把它们的顺序写错了。

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

你的锚元素（`a`）缺少 `href` 属性。 检查开始标签的名称后面是否有空格和/或所有属性名称之前是否有空格。

```js
assert(document.querySelector('a').hasAttribute('href'));
```

你的锚元素（`a`）应该链接到 `https://freecatphotoapp.com`。 你要么遗漏了 URL，要么拼写有误。

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

虽然你将锚元素（`a`）的 `href` 属性指向了正确的链接，但建议始终以引号来包裹属性的值。

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

