---
id: 5dfb655eeacea3f48c6300b3
title: 步骤 25
challengeType: 0
dashedName: step-25
---

# --description--

`figure` 元素代表自包含的内容，允许你将图像与标题相关联。

将刚刚添加的图像嵌套在一个 `figure` 元素中。

# --hints--

你的 `figure` 元素应有一个起始标签。起始标签的语法如下：`<elementName>`。

```js
assert(document.querySelector('figure'));
```

你的 `figure` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert(code.match(/<\/figure\>/));
```

在第二个 `section` 元素的结束标签上方应该有一个 `figure` 元素。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.lastElementChild?.nodeName, 'FIGURE');
```

千层面 `img` 元素应该嵌套在 `figure` 元素中。

```js
assert.strictEqual(document.querySelector('figure > img')?.src.toLowerCase(), 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
      
        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
      
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

