---
id: 5ef9b03c81a63668521804db
title: 第 43 步
challengeType: 0
dashedName: step-43
---

# --description--

为了防止用户在缺少必须信息时提交表单，你需要为 `input` 元素添加`required` 属性。

这是一个带有`required`属性的input元素的示例：

```html
<input type="text" name="firstName" required>
```

无需为 `required` 属性设置值。 只需将单词 `required` 添加到 `input` 元素，确保它和其他属性之间有空格。

# --hints--

你要么删除了 `input` 元素，要么它的语法无效。 所有属性的值都应该放在引号中。

```js
assert(document.querySelectorAll('input').length);
```

你的 `form` 应该只包含 `input` 元素。 删除 `form` 元素中的任何其他 HTML 元素或文本。

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

你的 `input` 元素应该具有 `required` 属性。 请记住，你只需在 `input` 元素的标签中添加单词 `required`。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('required'));
```

不应为 `required` 属性赋予值。

```js
assert.strictEqual(document.querySelector('input')?.getAttribute('required'), '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

