---
id: 671141d8e32fe934c26fa1be
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

在锚点元素前添加单词 `See more `，在锚点元素后添加单词 ` in our gallery`。

# --hints--

你应该使用现有的锚元素，而不是创建一个新。

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

锚元素内的文字不应更改，必须是 `cat photos`。

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

你的代码中应包含 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery`。

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
