---
id: 66e2d680e129e1423116a541
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utajenga mchambuzi wa sentensi ambao utachukua sentensi na kupata hesabu ya idadi ya maneno, vokali, konsonanti, na alama za uandishi.

Kuanza, tengeneza kitendakazi cha `getVowelCount` chenye kigezo kinachoitwa `sentence`. Kitendakazi chako kinapaswa kurudisha jumla ya idadi ya vokali katika sentensi.

# --hints--

Unapaswa kutengeneza kitendakazi cha `getVowelCount`.

```js
assert.isFunction(getVowelCount);
```

Kitendakazi chako cha `getVowelCount` kinapaswa kuwa na kigezo cha `sentence`.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

Kitendakazi chako cha `getVowelCount` kinapaswa kurudisha nambari.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

Wakati sentensi ni `"Apples are tasty fruits"`, kitendakazi cha `getVowelCount` kinapaswa kurudisha `7`.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

Wakati sentensi ni `"Hello, World!"`, kitendakazi kinapaswa kurudisha `3`.

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

Hesabu yako ya vokali inapaswa kuwa isiyoegemea herufi kubwa au ndogo.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

Kitendakazi chako cha `getVowelCount` kinapaswa kurudisha hesabu sahihi ya vokali kwa sentensi yoyote.

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
