---
id: 66e2edb23013ca596f8d9568
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Sasa ni wakati wa kujaribu kitendakazi chako cha `getPunctuationCount`.

Tengeneza kigezo cha `punctuationCount` na weka thamani yake kuwa matokeo ya wito wa kitendakazi cha `getPunctuationCount` ukiwa na hoja ya `"WHAT?!?!?!?!?"`

Baada ya hapo, andika yafuatayo kwenye konsoli: `"Punctuation Count: [Punctuation count goes here]"`. Badilisha `[Punctuation count goes here]` na jina halisi la kigezo. Unaweza kuchagua kutumia mfuatano wa herufi wa template au kuunganisha mfuatano wa herufi kwa kutumia opereta wa `+` hapa.

# --hints--

Unapaswa kutengeneza kigezo cha `punctuationCount`.

```js
assert.isNotNull(punctuationCount)
```

Kigezo chako cha `punctuationCount` kinapaswa kuwekwa kuwa matokeo ya `getPunctuationCount("WHAT?!?!?!?!?")`.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

Unapaswa kuandika yafuatayo kwenye konsoli: `"Punctuation Count: [Punctuation count goes here]"`. Badilisha `[Punctuation count goes here]` na jina halisi la kigezo. Hakikisha unatumia sintaksia sahihi ya kuunganisha mfuatano wa herufi hapa.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
