---
id: 66e2eab8a5638f57b637b7cc
title: Hatua ya 5
challengeType: 1
dashedName: step-5
---

# --description--

Sasa unapaswa kuhesabu idadi ya alama za uandishi.

Unda kitendakazi cha `getPunctuationCount` chenye kigezo cha `sentence`.

Ndani ya kitendakazi hicho, tengeneza mzunguko wa kuhesabu idadi ya herufi za alama za uandishi katika `sentence` zitakazopitishwa kwenye kitendakazi wakati kitendakazi kitapigiwa wito. Kwa madhumuni yetu, herufi ya alama za uandishi ni herufi yoyote isiyo nafasi (`" "`) au herufi ya kawaida.

Kitendakazi chako cha `getPunctuationCount` lazima kirudishe nambari.

# --hints--

Unapaswa kuunda kitendakazi cha `getPunctuationCount`.

```js
assert.isFunction(getPunctuationCount);
```

Kitendakazi chako cha `getPunctuationCount` kinapaswa kuwa na kigezo cha `sentence`.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

Kitendakazi chako cha `getPunctuationCount` kinapaswa kurudisha nambari.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

Wakati sentensi ni `"What's going on here?"`, kitendakazi cha `getPunctuationCount` kinapaswa kurudisha `2`.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

Wakati sentensi ni `"What????!"`, kitendakazi cha `getPunctuationCount` kinapaswa kurudisha `5`.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

Kitendakazi chako cha `getPunctuationCount` kinapaswa kurudisha hesabu sahihi ya alama za uandishi kwa sentensi yoyote.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
