---
id: 66e2e99caa0cb3570f4822b4
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Sasa ni wakati wa kujaribu kitendakazi chako cha `getConsonantCount`.

Tengeneza kigezo cha `consonantCount` na weka thamani yake kuwa matokeo ya kuita kitendakazi cha `getConsonantCount` kwa hoja ya `"Coding is fun"`

Baada ya hapo, andika yafuatayo kwenye konsoli: `"Consonant Count: [Consonant count goes here]"`. Badilisha `[Consonant count goes here]` na jina halisi la kigezo. Unaweza kuchagua kutumia mfuatano wa herufi wa template au kuunganisha mfuatano wa herufi kwa kutumia kigezo cha `+` hapa.

# --hints--

Unapaswa kutengeneza kigezo cha `consonantCount`.

```js
assert.isNotNull(consonantCount)
```

Kigezo chako cha `consonantCount` kinapaswa kuwekwa kuwa matokeo ya `getConsonantCount("Coding is fun")`.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

Unapaswa kuandika kigezo cha `consonantCount` kwenye konsoli kwa kigezo cha `"Consonant Count: ${consonantCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
