---
id: 66e2f06b191e305c00574e4d
title: Hatua ya 7
challengeType: 1
dashedName: step-7
---

# --description--

Mwishowe, hesabu idadi ya maneno kwa kuunda kitendakazi cha `getWordCount` chenye kigezo cha `sentence`. Kitendakazi hicho kinapaswa kurudisha jumla ya maneno katika sentensi inayotumwa wakati kinapoitwa.

# --hints--

Unapaswa kuunda kitendakazi cha `getWordCount`.

```js
assert.isFunction(getWordCount);
```

Kitendakazi chako cha `getWordCount` kinapaswa kuwa na kigezo cha `sentence`.

```js
assert.match(getWordCount.toString(), /sentence/);
```

Kitendakazi chako cha `getWordCount` kinapaswa kurudisha nambari.

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

Wakati sentensi ni `"When are you gonna start learning to code?"`, kitendakazi cha `getWordCount` kinapaswa kurudisha `8`.

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

Wakati sentensi ni `"What's going on?"`, kitendakazi cha `getWordCount` kinapaswa kurudisha `3`.

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

Hesabu yako ya maneno inapaswa kuwa isiyoegemea herufi kubwa au ndogo.

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

Kitendakazi chako cha `getWordCount` kinapaswa kurudisha hesabu sahihi ya maneno kwa sentensi yoyote.

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both front-end and back-end development"), 7);
```

Kitendakazi chako cha `getWordCount` kinapaswa kurudisha hesabu sahihi ya maneno kwa mfuatano wa herufi tupu, au mfuatano wa herufi wenye nafasi tu.

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
