---
id: 66e2e5d890ac8a54c9d703bd
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Ni wakati wa kuhesabu konsonanti. Unda kitendakazi cha `getConsonantCount` chenye kigezo cha `sentence`.

Ndani ya kitendakazi, tumia mzunguko kuhesabu idadi ya konsonanti katika `sentence` itakayopitishwa kwenye kitendakazi wakati kitakapoitwa. Konsonanti ni herufi yoyote isiyo mojawapo ya herufi zifuatazo: `"aeiou"`.

Kitendakazi chako cha `getConsonantCount` lazima kirudishe nambari.

# --hints--

Unapaswa kuunda kitendakazi cha `getConsonantCount`.

```js
assert.isFunction(getConsonantCount);
```

Kitendakazi chako cha `getConsonantCount` kinapaswa kuwa na kigezo cha `sentence`.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

Kitendakazi chako cha `getConsonantCount` kinapaswa kurudisha nambari.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

Wakati sentensi ni `"Coding is fun"`, kitendakazi cha `getConsonantCount` kinapaswa kurudisha `7`.

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

Wakati sentensi ni `"Hello, World!"`, kitendakazi cha `getConsonantCount` kinapaswa kurudisha `7`.

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

Hesabu yako ya konsonanti inapaswa kuwa isiyoegemea herufi kubwa au ndogo.

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

Kitendakazi chako cha `getConsonantCount` kinapaswa kurudisha hesabu sahihi ya konsonanti kwa sentensi yoyote.

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
