---
id: 66e2df8c5676fe51cead8ec8
title: Hatua ya 2
challengeType: 1
dashedName: step-2
---

# --description--

Sasa ni wakati wa kujaribu kitendakazi chako cha `getVowelCount`.

Tengeneza kigezo cha `vowelCount` na weka thamani yake kuwa matokeo ya kuita kitendakazi cha `getVowelCount` ukiwa na hoja ya `"Apples are tasty fruits"`

Baada ya hapo, andika yafuatayo kwenye konsoli: `"Vowel Count: [vowel count goes here]"`. Badilisha `[vowel count goes here]` na jina halisi la kigezo. Unaweza kuchagua kutumia mfuatano wa herufi wa template au kuunganisha mfuatano wa herufi kwa kutumia kigezo cha `+` hapa.

# --hints--

Unapaswa kutengeneza kigezo cha `vowelCount`.

```js
assert.isNotNull(vowelCount)
```

Kigezo chako cha `vowelCount` kinapaswa kuwekwa kuwa matokeo ya `getVowelCount("Apples are tasty fruits")`.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

Unapaswa kuandika kigezo cha `vowelCount` kwenye konsoli ukiwa na kigezo cha `"Vowel Count: ${vowelCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
