---
id: 63db7f4677d06d7500a13321
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本項目中，你將通過編寫一個簡易的十進制到二進制的轉換器來學習兩種進制間的差異。 此外，你還會初步學習遞歸的用法。

我們已經爲你提供了此項目所需的 HTML 和 CSS 代碼。

當你準備開始時，使用 `.getElementById()` 方法獲取 id 爲 `"number-input"` 的 `input` 元素，並將其存儲在名爲 `numberInput` 的變量中。 使用同樣的方法獲取 id 爲 `"convert-btn"` 的 `button` 元素並將其存儲在名爲 `convertBtn` 的變量中，以及獲取 id 爲 `"result"` 的 `output` 元素並將其存儲在名爲 `result` 的變量中。


**注意**：該項目只會將正數轉換爲二進制。

# --hints--

你應該使用 `document.getElementById()` 方法獲取`#number-input` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)number-input\1\s*\)/);
```

你應該使用 `const` 聲明一個名爲 `numberInput`的變量，並將 `#number-input` 元素賦值給它。

```js
assert.match(code, /const\s+numberInput\s*=\s*document\.getElementById\(\s*('|"|`)number-input\1\s*\)/);
```

你應該使用 `document.getElementById()` 方法獲取 `#convert-btn` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)convert-btn\1\s*\)/);
```

你應該使用 `const` 聲明一個名爲 `convertBtn` 的變量，將 `#convert-btn` 元素賦值給它。

```js
assert.match(code, /const\s+convertBtn\s*=\s*document\.getElementById\(\s*('|"|`)convert-btn\1\s*\)/);
```

你應該使用 `document.getElementById()` 方法獲取 `#result` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)result\1\s*\)/);
```

你應該使用 `const` 聲明一個名爲 `result` 的變量，將 `#result` 元素賦值給它。

```js
assert.match(code, /const\s+result\s*=\s*document\.getElementById\(\s*('|"|`)result\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Decimal to Binary Converter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>
      Decimal to Binary <br />
      Converter
    </h1>
    <section class="input-container">
      <label for="number-input">Enter a decimal number:</label>
      <input
        value=""
        type="number"
        name="decimal number input"
        id="number-input"
        class="number-input"
      />
      <button class="convert-btn" id="convert-btn">Convert</button>
    </section>
    <section class="output-container">
      <output id="result" for="number-input"></output>
      <h2>Call stack</h2>
      <div id="animation-container"></div>
    </section>
    <script src="script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #f5f6f7;
  --dark-blue: #1b1b32;
  --orange: #f1be32;
}

body {
  display: flex;
  flex-direction: column;
  align-items: center;
  font-family: Menlo, Consolas, Monaco, Liberation Mono, Lucida Console,
    monospace;
  font-size: 1.125rem;
  color: var(--light-grey);
  background-color: var(--dark-blue);
  padding: 0 4px;
}

h1 {
  font-size: 2.125rem;
  text-align: center;
  margin: 20px 0;
}

h2 {
  font-size: 1.5rem;
  text-align: center;
  margin: 20px 0;
}

.input-container {
  display: flex;
  flex-direction: column;
  gap: 10px;
  justify-content: center;
  align-items: center;
  width: clamp(320px, 50vw, 460px);
  margin: 10px auto;
}

.input-container label {
  white-space: nowrap;
  word-spacing: -6px;
}

.convert-btn {
  font-size: inherit;
  font-family: inherit;
  background-color: var(--orange);
  width: 100%;
  height: 2rem;
  padding: 0 6px;
  border: none;
  border-radius: 2px;
  cursor: pointer;
}

.number-input {
  font-size: inherit;
  padding: 0.3rem;
  width: 100%;
}

.output-container {
  margin-inline: auto;
  width: clamp(320px, 50vw, 460px);
}

#result {
  display: flex;
  align-items: center;
  justify-content: center;
  font-size: 2rem;
  text-align: center;
  min-height: 80px;
  margin-block-start: 20px;
  padding: 15px;
  border: 2px solid var(--orange);
  border-radius: 2px;
}

#animation-container {
  display: flex;
  flex-direction: column-reverse;
  justify-content: end;
  gap: 1rem;
  margin-block-end: 1rem;
  min-height: 40vh;
  border: 2px dashed var(--orange);
  padding: 1rem;
}

.animation-frame {
  font-family: -apple-system, BlinkMacSystemFont, avenir next, avenir, segoe ui,
    helvetica neue, helvetica, Cantarell, Ubuntu, roboto, noto, arial,
    sans-serif;
  padding: 15px 10px;
  border: 5px solid var(--orange);
  font-size: 1.2rem;
  text-align: center;
}

@media screen and (min-width: 36em) {
  body {
    font-size: 1rem;
  }

  .input-container {
    flex-direction: row;
    width: unset;
  }

  .number-input {
    width: unset;
  }
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
