---
id: 66ed41f912d0bb1dc62da5dd
title: Jenga bot ya maswali ya JavaScript
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuandika `"Hello! I'm your coding fun fact guide!"` kwenye konsoli kama ujumbe wa salamu kwa mtumizi.
1. Unapaswa kutangaza vigezo vitatu: `botName`, `botLocation`, na `favoriteLanguage`, vinavyohifadhi jina la bot, mahali anapotoka, na lugha yake ya kupenda ya kuandika programu, mtawalia.
1. Unapaswa kuandika `"My name is (botName) and I live on (botLocation)."` kwenye konsoli.
1. Unapaswa kuandika `"My favorite programming language is (favoriteLanguage)."` kwenye konsoli.
1. Unapaswa kutumia `let` kutangaza kigezo kipya `codingFact` na kuweka thamani ya mfuatano wa herufi unaohusu ukweli wa kufurahisha kuhusu lugha ya kupenda ya bot na ujumuishe matumizi ya kigezo `favoriteLanguage`.
1. Unapaswa kuandika `codingFact` kwenye konsoli.
1. Unapaswa kuweka thamani mpya kwa kigezo `codingFact` kuhusu ukweli mwingine wa lugha ya kupenda ya bot ukitumia tena kigezo `favoriteLanguage`.
1. Unapaswa kuandika tena `codingFact` kwenye konsoli.
1. Unapaswa kuweka tena thamani mpya kwa kigezo `codingFact` kuhusu ukweli mwingine wa lugha ya kupenda ya bot ukitumia kigezo `favoriteLanguage`.
1. Unapaswa kuandika `codingFact` kwenye konsoli mara ya tatu.
1. Unapaswa kuandika `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` kwenye konsoli kama taarifa ya kuaga kutoka kwa bot.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuandika `"Hello! I'm your coding fun fact guide!"` kwenye konsoli.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

Unapaswa kutangaza kigezo `botName`. Hakikisha hakuna makosa ya tahajia au herufi kubwa/dogo.

```js
assert.exists(botName);
```

Kigezo chako cha `botName` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(botName);
```

Unapaswa kutangaza kigezo `botLocation`. Hakikisha hakuna makosa ya tahajia au herufi kubwa/dogo.

```js
assert.exists(botLocation);
```

Kigezo chako `botLocation` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(botLocation);
```

Unapaswa kutangaza kigezo `favoriteLanguage`. Hakikisha hakuna makosa ya tahajia au herufi kubwa/dogo.

```js
assert.exists(favoriteLanguage);
```

Kigezo chako `favoriteLanguage` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(favoriteLanguage);
```

Unapaswa kuandika kwenye konsoli `"My name is (botName) and I live on (botLocation)."` na kuongeza vigezo kwenye mfuatano wa herufi.

```js
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

Unapaswa kuandika kwenye konsoli `"My favorite programming language is (favoriteLanguage)."` na kuongeza kigezo kwenye mfuatano wa herufi.

```js
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

Unapaswa kutumia `let` kutangaza kigezo kipya `codingFact`.

```js
const explorer = await __helpers.Explorer(code);
const { codingFact } = explorer.variables;
assert.isTrue(codingFact?.matches(`let codingFact = ${codingFact?.value?.toString()}`));
```

Unapaswa kumpa `codingFact` thamani inayojumuisha `favoriteLanguage`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.codingFact);
assert.include(String(codingFact), favoriteLanguage);
```

Unapaswa kuandika `codingFact` kwenye konsoli.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

Unapaswa kuweka thamani mpya kwa `codingFact` inayojumuisha `favoriteLanguage`, na kuandika kwenye konsoli.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.exists(loggingCodingFacts);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
```

Unapaswa kuweka thamani kwa `codingFact` kwa mara ya tatu inayojumuisha `favoriteLanguage`, na kuandika kwenye konsoli.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.lengthOf(loggingCodingFacts, 3);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.isNotEmpty(codingFact);
```

Unapaswa kuandika kwenye konsoli `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` na kuongeza thamani za vigezo.

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
