---
id: 66cc0f1ae40802781b2ea972
title: 步骤6
challengeType: 1
dashedName: step-6
---

# --description--

In previous lessons, you learned how to work with function parameters.

函数参数是在函数声明中定义的变量，用作占位符。

下面是一个有参数的函数示例：

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

声明一个名为 `calculateSum` 的函数，它有 `num1` 和 `num2` 两个参数。

# --hints--

你应声明一个名为 `calculateSum` 的函数。

```js
assert.isFunction(calculateSum);
```

你的 `calculateSum` 函数应包含 `num1` 和 `num2` 两个参数,。

```js
const regex = __helpers.functionRegex('calculateSum', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
