---
id: 66cc095dfe1682753d2ab030
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

你在之前的课程中学到的一个概念是如何从函数中返回值。

下面是如何从函数中返回值的提示：

```js
function myFunction() {
  return "Hello World";
}
```

在你的 `addTwoAndSeven` 函数中，返回 `2` 与 `7` 的和。

# --hints--

你的 `addTwoAndSeven` 函数应返回 `2` 与 `7` 的和。

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

你的 `addTwoAndSeven` 函数应使用 `+` 运算符返回和。

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
const validExpressions = ["2 + 7", "7 + 2", "(2 + 7)", "(7 + 2)"];
assert.isTrue(validExpressions.some(expr => addTwoAndSeven?.hasReturn(expr)));
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
