---
id: 66cc0a9e06e00b75d6782be9
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

In previous lessons, you learned how to call (invoke) a function. 调用函数意味着执行函数内部的代码。

下面是提醒大家如何调用函数：

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

添加一个 `console.log` 语句，并在其中调用 `addTwoAndSeven` 函数。

现在你应该可以看到控制台记录了 `2` 和 `7` 的和。

# --hints--

你的代码中应该有一个 `console.log`。

```js
assert.match(code, /console\.log(.*)/);
```

你应在 `console.log` 中调用 `addTwoAndSeven` 函数。

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}
--fcc-editable-region--

--fcc-editable-region--
```
