---
id: 66cc0a9e06e00b75d6782be9
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

在之前的课程中，你学习了如何调用（invoke）一个函数。调用函数意味着执行函数内部的代码。

下面是提醒大家如何调用函数：

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

添加一个 `console.log` 语句，并在其中调用 `addTwoAndSeven` 函数。

现在你应该可以看到控制台记录了 `2` 和 `7` 的和。

# --hints--

你应该在 `console.log` 内调用 `addTwoAndSeven` 函数。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addTwoAndSeven())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
