---
id: 66cc0de2920bec775f610424
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

尽管这些函数的运行符合预期，但你的代码中存在大量重复。

由于你正在构建一个计算器，因此不必为每种可能的数字加法组合创建一个函数。

最好是创建一个可重复使用的函数，将任意两个数字相加。

删除迄今为止编写的所有代码。

下一步，你将学习如何在函数中使用参数。

# --hints--

你不应有名为 `addTwoAndSeven` 的函数。

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
assert.notExists(addTwoAndSeven);
```

你不应有名为 `addTwoAndSeven` 的函数。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addTwoAndSeven())";
assert.notInclude(strippedCode, unexpectedStatement);
```

你不应有名为 `addThreeAndFour` 的函数。

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
assert.notExists(addThreeAndFour);
```

你不应有名为 `addThreeAndFour` 的函数。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addThreeAndFour())";
assert.notInclude(strippedCode, unexpectedStatement);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
