---
id: 66cc0de2920bec775f610424
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

尽管这些函数的运行符合预期，但你的代码中存在大量重复。

由于你正在构建一个计算器，因此不必为每种可能的数字加法组合创建一个函数。

最好是创建一个可重复使用的函数，将任意两个数字相加。

删除迄今为止编写的所有代码。

下一步，你将学习如何在函数中使用参数。

# --hints--

你不应有名为 `addTwoAndSeven` 的函数。

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

你不应有名为 `addTwoAndSeven` 的函数。

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

你不应有名为 `addThreeAndFour` 的函数。

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

你不应有名为 `addThreeAndFour` 的函数。

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
