---
id: 66cc0ba5881acb7692cfc4de
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

现在，是时候添加另一个函数了。

声明一个名为 `addThreeAndFour` 的函数，返回 `3` 与 `4` 的和。

然后在一个 `console.log` 里面调用 `addThreeAndFour` 函数，查看结果。

# --hints--

你应声明一个名为 `addThreeAndFour` 的函数。

```js
assert.isFunction(addThreeAndFour);
```

你的 `addThreeAndFour` 函数应返回 `3` 与 `4` 的和。

```js
assert.equal(addThreeAndFour(), 7);
```

你的 `addThreeAndFour` 函数应使用 `+` 运算符返回和。

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

你应在一个 `console.log` 中调用 `addThreeAndFour` 函数。

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());
--fcc-editable-region--

--fcc-editable-region--
```
