---
id: 66cc095dfe1682753d2ab030
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

你在之前的课程中学到的概念之一是如何从函数返回值。

下面是如何从函数中返回值的提示：

```js
function myFunction() {
  return "Hello World";
}
```

在你的 `addTwoAndSeven` 函数中，返回 `2` 与 `7` 的和。

# --hints--

你的 `addTwoAndSeven` 函数应返回 `2` 与 `7` 的和。

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

你的 `addTwoAndSeven` 函数应使用 `+` 运算符返回和。

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--

--fcc-editable-region--
}
```
