---
id: 66cc2739f687e484d50bb6f1
title: 步骤18
challengeType: 1
dashedName: step-18
---

# --description--

计算器的最后一项功能是平方根功能。

平方是一个数字乘以它自己。 因此，平方根就是与自己相乘后得到原数的那个数。

例如，`9` 的平方根是 `3`，因为 `3 * 3 = 9`。

要在 JavaScript 中计算一个数字的平方根，可以使用 <dfn>Math.sqrt()</dfn> 方法。

这是一个示例：

```js
// result: 3
Math.sqrt(9);
```

声明一个名为 `calculateSquareRoot` 的函数，接受参数 `num` 并返回 `num` 的平方根。

# --hints--

你应有一个名为 `calculateSquareRoot` 的函数。

```js
assert.isFunction(calculateSquareRoot);
```

你的 `calculateSquareRoot` 函数应接受参数 `num`。

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的 `calculateSquareRoot` 函数应返回 `num` 的平方根。

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
