---
id: 66cc12fa504b0479dac479a0
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

函数的优点是可以用不同的参数来调用，这样就可以用不同的值来重复使用相同的代码。

下面是一些使用不同字符串参数的函数调用示例：

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

添加一个 `console.log`，在其中调用 `calculateSum` 函数，并传入参数 `2` 和 `5`。

# --hints--

你应有一个 `console.log`，在其中调用 `calculateSum` 函数，并传入参数 `2` 和 `5`。

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
