---
id: 66cc16abc30b367ba1d1f9ef
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

现在是测试 `calculateDifference` 函数的时候了。

首先，添加一个 `console.log`，在其中调用 `calculateDifference` 函数，并传入参数 `22` 和 `5`。

然后，再添加一个 `console.log`，在其中调用 `calculateDifference` 函数，并传入参数 `12` 和 `1`。

最后，添加第三个 `console.log`，调用 `calculateDifference` 函数，并传入参数 `17` 和 `9`。

# --hints--

你应有一个 `console.log`，在其中调用 `calculateDifference` 函数，并传入参数 `22` 和 `5`。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(22,5))";
assert.include(strippedCode, expectedStatement);
```

你应有一个 `console.log`，在其中调用 `calculateDifference` 函数，并传入参数 `12` 和 `1`。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(12,1))";
assert.include(strippedCode, expectedStatement);
```

你应有一个 `console.log`，在其中调用 `calculateDifference` 函数，并传入参数 `17` 和 `9`。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(17,9))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
