---
id: 66cc10958d72e57907f898b2
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

在 `calculateSum` 函数中，你需要返回两个参数之和。

# --hints--

你的 `calculateSum` 函数应返回两个参数之和。

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
