---
id: 66cc06012e46aa6bc9b8c001
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在这个工作坊中，你将通过构建一个计算器应用来复习如何使用函数。 

在之前的课程中，你学习了如何像这样声明函数：

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

首先创建一个名为 `addTwoAndSeven` 的函数。

你可以选择使用常规函数语法或箭头函数语法。

# --hints--

你应声明一个名为 `addTwoAndSeven` 的函数。

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
