---
id: 66cc14ca14e65e7a9ebecb80
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

现在你的计算器已经可以将两个数字相加，是时候创建一个可以将两个数字相减的函数了。

声明一个名为 `calculateDifference` 的函数，它有 `num1` 和 `num2` 两个参数。

在函数内部，返回两个参数之间的差值。

# --hints--

你应有一个名为 `calculateDifference` 的函数。

```js
assert.isFunction(calculateDifference);
```

你的 `calculateDifference` 函数应该有两个参数，分别称为 `num1` 和 `num2`。

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的 `calculateDifference` 函数应返回两个参数的差值。

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
