---
id: 66cc16abc30b367ba1d1f9ef
title: 步骤11
challengeType: 1
dashedName: step-11
---

# --description--

现在是测试 `calculateDifference` 函数的时候了。

首先，添加一个 `console.log`，在其中调用 `calculateDifference` 函数，并传入参数 `22` 和 `5`。

然后，再添加一个 `console.log`，在其中调用 `calculateDifference` 函数，并传入参数 `12` 和 `1`。

最后，添加第三个 `console.log`，调用 `calculateDifference` 函数，并传入参数 `17` 和 `9`。

# --hints--

你应有一个 `console.log`，在其中调用 `calculateDifference` 函数，并传入参数 `22` 和 `5`。

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*22\s*,\s*5\s*\)\s*\)\s*;?/);
```

你应有一个 `console.log`，在其中调用 `calculateDifference` 函数，并传入参数 `12` 和 `1`。

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*12\s*,\s*1\s*\)\s*\)\s*;?/);
```

你应有一个 `console.log`，在其中调用 `calculateDifference` 函数，并传入参数 `17` 和 `9`。

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*17\s*,\s*9\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
