---
id: 66cc21d23238dc8240a8a182
title: 步骤 16
challengeType: 1
dashedName: step-16
---

# --description--

如果能让计算器计算一个数字的平方就更好了。 一个数字的平方就是这个数字乘以它自己。

要在 JavaScript 中计算一个数字的平方，可以使用 <dfn>Math.pow()</dfn> 方法。 `Math.pow()` 方法需要两个参数：底数和指数。

你还可以使用指数运算符（`**`）来计算一个数字的平方。

这是一个例子：

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

声明一个名为 `calculateSquare` 的函数，接受参数 `num` 并返回 `num` 的平方。

# --hints--

你应有一个名为 `calculateSquare` 的函数。

```js
assert.isFunction(calculateSquare);
```

你的 `calculateSquare` 函数应接受参数 `num`。

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的 `calculateSquare` 函数应返回 `num` 的平方。

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
