---
id: 66cc1a3a39aef47d6473cb2f
title: 第 13 步
challengeType: 1
dashedName: step-13
---

# --description--

下一步是在计算器中添加除法功能。

声明一个名为 `calculateQuotient` 的函数，它有 `num1` 和 `num2` 两个参数。

在函数内部，返回两个参数的商。

然后，添加一个 `console.log`，在其中调用 `calculateQuotient` 函数，并传入参数 `7` 和 `11`。

# --hints--

你应该声明一个名为 `calculateQuotient` 的函数。

```js
assert.isFunction(calculateQuotient);
```

你的 `calculateQuotient` 函数应该有两个参数，分别称为 `num1` 和 `num2`。

```js
const explorer = await __helpers.Explorer(code);
const { calculateQuotient } = explorer.allFunctions;
assert.lengthOf(calculateQuotient?.parameters, 2);
assert.equal(calculateQuotient?.parameters[0].toString(), "num1");
assert.equal(calculateQuotient?.parameters[1].toString(), "num2");
```

你的 `calculateQuotient` 函数应返回两个参数的商。

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

你应有一个 `console.log`，在其中调用 `calculateQuotient` 函数，并传入参数 `7` 和 `11`。

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateQuotient(7,11))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
