---
id: 66cdf76685e4cb5a8726e27b
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

要进一步测试 `calculateSum` 函数，你需要使用不同的参数来调用它。

再添加一个 `console.log`，在其中调用 `calculateSum` 函数，并传入参数 `10` 和 `10`。

在这个 `console.log` 的下面，再添加一个 `console.log`，调用 `calculateSum` 函数，并传入参数 `5` 和 `5`。

# --hints--

你应有一个 `console.log`，在其中调用 `calculateSum` 函数，并传入参数 `10` 和 `10`。

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

你应有一个 `console.log`，在其中调用 `calculateSum` 函数，并传入参数 `5` 和 `5`。

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));

--fcc-editable-region--

--fcc-editable-region--
```
