---
id: 66ebd4ae2812430bb883c787
title: イベントハブを作成する
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

このラボでは、セマンティックHTML要素を使ってウェブページの構造を作成します。コンテンツや画像を追加して、本物のイベントハブのように見せます。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `header` 要素が必要です。
1. `header` 要素の中に、テキスト `h1` を含む `Event Hub` 要素と `nav` 要素が必要です。
1. `nav` 要素の中に、ページの異なるセクションへのリンクを含む2つの項目からなる番号なしリストが必要です。最初の項目はテキスト `Upcoming Events`、2番目の項目はテキスト `Past Events` を持つ必要があります。
1. 各リンクは `a` 要素で表され、対応するページのセクション `href` と `#upcoming-events` にリンクする `#past-events` 属性を持つ必要があります。
1. ページの異なるセクションを含む `main` 要素が必要です。
1. `main` 要素の中に、2つの `section` 要素が必要です。
1. 最初の `section` 要素は、値 `id` の `upcoming-events` 属性を持つ必要があります。
1. `#upcoming-events` セクションの中には以下が必要です：
   
   - テキスト `h2` を持つ `Upcoming Events` 要素。
   - 2つの `article` 要素。各記事はイベントを表し、以下を持つ必要があります：
      - イベントタイトル用の `h3` 要素。
      - イベント説明用の `p` 要素。必要に応じて日付を下部に追加できます。

1. 2番目の `section` 要素は、値 `id` の `past-events` 属性を持つ必要があります。
1. `#past-events` セクションの中には以下が必要です：

   - テキスト `h2` を持つ `Past Events` 要素。
   - 2つの `article` 要素。各記事要素は過去のイベントを表し、以下を持つ必要があります：
     - イベントタイトル用の `h3` 要素、
     - イベント説明用の `p` 要素。必要に応じて日付を下部に追加できます。
     - 画像ファイルを指す `src` 属性と画像の説明を持つ `alt` 属性を持つ画像要素。

**注意:** イベント説明や日付は任意のテキストを使用できます。画像には以下のURLを使ってもかまいません：

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`。
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`。

# --hints--

`header` 要素が必要です。

```js
assert.isNotNull(document.querySelector("header"));
```

`header` 要素は開く `body` タグの後に配置してください。

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

`header` 要素の中に、テキスト `h1` を含む `Event Hub` 要素が必要です。

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

`header` 要素の中で、`h1` 要素の後に `nav` 要素が必要です。

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

`nav` 要素には2つの項目からなる番号なしリストが必要です。

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

番号なしリストの最初の項目はリンクである必要があります。

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

番号なしリストの2番目の項目もリンクである必要があります。

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

番号なしリストの最初の項目のテキストは `Upcoming Events` である必要があります。

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

番号なしリストの最初の項目は `href` を `#upcoming-events` に設定してください。

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

番号なしリストの2番目の項目のテキストは `Past Events` である必要があります。

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

番号なしリストの2番目の項目は `href` を `#past-events` に設定してください。

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

`main` 要素は `header` 要素の終了タグの後に配置してください。

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

`main` 要素の中に2つの `section` 要素が必要です。

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

最初の `section` 要素は、値 `id` の `upcoming-events` 属性を持つ必要があります。

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

2番目の `section` 要素は、値 `id` の `past-events` 属性を持つ必要があります。

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

`#upcoming-events` セクションの中に、テキスト `h2` を持つ `Upcoming Events` 要素が必要です。

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

`#upcoming-events` セクションの中で、`article` 要素の下に2つの `h2` 要素が必要です。

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`article` セクション内の両方の `#upcoming-events` 要素はイベントタイトル用の `h3` 要素を持つ必要があります。

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`article` セクション内の両方の `#upcoming-events` 要素はイベント説明用の段落要素を持つ必要があります。

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

`#past-events` セクションの中に、テキスト `h2` を持つ `Past Events` 要素が必要です。

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

`#past-events` セクションの中で、`article` 要素の下に2つの `h2` 要素が必要です。

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`article` セクション内の両方の `#past-events` 要素はイベントタイトル用の `h3` 要素を持つ必要があります。

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`article` セクション内の両方の `#past-events` 要素はイベント説明用の段落要素を持つ必要があります。

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

`article` セクション内の両方の `#past-events` 要素は画像要素を持つ必要があります。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

`#past-events` セクション内の両方の画像要素は `src` 属性で画像ファイルを指す必要があります。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

`#past-events` セクション内の両方の画像要素は `alt` 属性で画像の説明を持つ必要があります。

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

各 `h3` 要素はイベントタイトルを持つ必要があります。

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

各 `p` 要素はイベント説明を持つ必要があります。

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```
