---
id: 69babfd65c4c4632a0c76f63
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Wenn ein Tor während eines Ticks nicht alle Teilnehmenden verarbeiten kann, bleiben einige Teilnehmende in der Warteschlange. Um diesen Überlauf zu bewältigen, erstellen Sie eine Funktion, die die überlaufenden Teilnehmenden zu einem anderen Tor umleitet.

Erstellen Sie eine leere Funktion namens `rerouteOverflow` mit den folgenden Parametern:

- `gates`: Das vollständige Array der Tor-Objekte.
- `currentGate`: Das Tor, an dem der Überlauf aufgetreten ist.
- `tickIndex`: Die aktuelle Tick-Position in der Warteschlange.
- `overflowAmount`: Die Anzahl der Teilnehmenden, die nicht verarbeitet werden konnten.

# --hints--

Sie sollten eine Funktion namens `rerouteOverflow` erstellen.

```js
assert.isFunction(rerouteOverflow);
```

Ihre `rerouteOverflow`-Funktion sollte die Parameter `gates`, `currentGate`, `tickIndex` und `overflowAmount` haben.

```js
const regex = __helpers.functionRegex('rerouteOverflow', ['gates', 'currentGate', 'tickIndex', 'overflowAmount']);
assert.match(__helpers.removeJSComments(rerouteOverflow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

--fcc-editable-region--

--fcc-editable-region--
```
