---
id: 69babdf3f896711c11516373
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Während jedes Simulationsschritts:

- Kommt eine bestimmte Anzahl von Teilnehmern am Tor an (von seinem `queue`).
- Verarbeitet das Tor die Teilnehmer basierend auf seinem `capacity`.
- Wenn mehr Teilnehmer ankommen, als das Tor verarbeiten kann, bleiben einige übrig (Überlauf).

Sie werden nun eine Funktion erstellen, die diese Logik für ein Tor in einem einzelnen Schritt handhabt.

Erstellen Sie eine leere Funktion namens `processGateFlow` mit den folgenden Parametern:

- `gate`: Das zu verarbeitende Gate-Objekt.
- `tickIndex`: Die aktuelle Tick-Position im `queue` Array.

# --hints--

Sie sollten eine Funktion namens `processGateFlow` erstellen.

```js
assert.isFunction(processGateFlow);
```

Ihre `processGateFlow` Funktion sollte die Parameter `gate` und `tickIndex` haben.

```js
const explorer = await __helpers.Explorer(code);
const { processGateFlow } = explorer.allFunctions;
const parameters = processGateFlow?.parameters;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), "gate");
assert.equal(parameters[1].toString(), "tickIndex");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
