---
id: 69baafd6d683dd9921b75db3
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop verwenden Sie JavaScript, um den Ablauf der Besucher bei einem Musikfestival zu simulieren.

Sie arbeiten mit zwei Datensätzen, die die Festival-Eingänge während der Morgen- und Nachtschichten beschreiben:

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

Jedes Tor-Objekt enthält die folgenden Eigenschaften:

- `id`: Ein String, der das Tor identifiziert.
- `capacity`: Die Anzahl der Besucher, die das Tor pro Tick verarbeiten kann.
- `queue`: Ein Array von Zahlen, das darstellt, wie viele Besucher während eines bestimmten Ticks am Tor ankommen.

Erstellen Sie zunächst eine leere Funktion namens `initializeThroughput` mit einem Parameter `gates`.

# --hints--

Sie sollten eine Funktion mit dem Namen `initializeThroughput` erstellen.

```js
assert.isFunction(initializeThroughput);
```

Ihre Funktion `initializeThroughput` sollte einen Parameter `gates` haben.

```js
const explorer = await __helpers.Explorer(code);
const { initializeThroughput } = explorer.allFunctions;
const parameters = initializeThroughput?.parameters;
assert.lengthOf(parameters, 1);
assert.equal(parameters[0].toString(), "gates");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
