---
id: 69bac067cd15c4685d07906d
title: Schritt 15
challengeType: 1
dashedName: step-15
---

# --description--

Da Sie nun Funktionen haben, um einzelne Ticks zu verarbeiten und Überläufe zu handhaben, ist es Zeit, eine Steuerfunktion für ein einzelnes Gate zu erstellen. Diese Funktion übernimmt alle Aktionen für ein Gate während eines einzelnen Ticks der Simulation.

Erstellen Sie eine leere Funktion namens `handleGateAtTick` mit den Parametern:

- `gates`: Das vollständige Array der Gate-Objekte.
- `gate`: Das aktuell verarbeitete Gate.
- `tickIndex`: Der aktuelle Tick-Index der Simulation.
- `throughputSummary`: Ein Objekt, das die insgesamt verarbeiteten Teilnehmer pro Gate verfolgt.

# --hints--

Sie sollten eine Funktion namens `handleGateAtTick` erstellen.

```js
assert.isFunction(handleGateAtTick);
```

Ihre `handleGateAtTick`-Funktion sollte die Parameter `gates`, `gate`, `tickIndex` und `throughputSummary` haben.

```js
const regex = __helpers.functionRegex('handleGateAtTick', ['gates', 'gate', 'tickIndex', 'throughputSummary']);
assert.match(__helpers.removeJSComments(handleGateAtTick.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
