---
id: 69bac18fe3729f6761344365
title: Schritt 27
challengeType: 1
dashedName: step-27
---

# --description--

Sie müssen wissen, wie viele Ticks die Simulation laufen soll. Erinnern Sie sich daran:

- Jedes Tor hat ein `queue`-Array, das zeigt, wie viele Besucher bei jedem Tick eintreffen.  
- Da alle Tore dieselbe Anzahl an Ticks haben, können Sie einfach die Länge der Warteschlange des ersten Tors verwenden.

Erstellen Sie eine Variable mit dem Namen `maxTicks` und weisen Sie ihr den Wert von `gates[0].queue.length` zu.

Dieser Wert wird verwendet, um Ihre Simulationsschleife zu steuern und sicherzustellen, dass jeder Tick für alle Tore verarbeitet wird.

# --hints--

Sie sollten eine Variable mit dem Namen `maxTicks` erstellen und ihr den Wert von `gates[0].queue.length` zuweisen.

```js
const rawFuncStr = __helpers.removeJSComments(simulateFestival.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+maxTicks/,
  "You must have at least one space between the declaration keyword (let/const) and 'maxTicks'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)maxTicks=gates\[0\]\.queue\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `simulateFestival` function, create a variable named `maxTicks` and assign it the value of `gates[0].queue.length`."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
  console.log("\nProcessing " + gate.id + "...");
  console.log(
    gate.queue[tickIndex] + " attendees arriving."
  );
  const result = processGateFlow(gate, tickIndex);
  throughputSummary[gate.id] += result.processed;
  if (result.overflow > 0) {
    console.log(
      "Overflow of " + result.overflow +
      " attendees. Rerouting..."
    );
    rerouteOverflow(gates, gate, tickIndex, result.overflow);
  }
}

function printSummary(summary) {
  console.log("\nThroughput Summary");
  for (const gateId in summary) {
    console.log(
      gateId + ": " + summary[gateId] +
      " attendees processed"
    );
  }
}

function simulateFestival(gates, timeBlock) {
  console.log("\n" + timeBlock + " Simulation");
  const throughputSummary = initializeThroughput(gates);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
