---
id: 69bac027df64d4433c708d44
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Wenn Sie Überlauf umleiten, sollten Sie die Teilnehmer zum nächsten Tor im `gates`-Array schicken. Dafür müssen Sie den Index des nächsten Tors finden. Normalerweise könnten Sie zum aktuellen Index 1 addieren. Wenn das aktuelle Tor jedoch das letzte im Array ist, müssen Sie zum ersten Tor zurückspringen.

Sie können den Modulo-Operator (`%`) verwenden, um dies zu handhaben:

```js
const nextIndex = (currentIndex + 1) % array.length;
```

Das funktioniert, weil das Ergebnis 0 wird, wenn `currentIndex + 1` gleich `array.length` ist.

Verwenden Sie diesen Ansatz, um eine Variable namens `nextGateIndex` zu erstellen, die den Index des nächsten Tors im `gates`-Array speichert.

# --hints--

Sie sollten die Modulo-Technik verwenden, um den nächsten Index aus `currentIndex` im `gates`-Array zuzugreifen und zuzuweisen, und zwar in einer Variablen namens `nextGateIndex`. Folgen Sie dem gleichen einzeiligen Muster wie im Beispiel.

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+nextGateIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'nextGateIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)nextGateIndex=\((currentIndex\+1|1\+currentIndex)\)%gates\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use the modulo technique to create a variable named `nextGateIndex` that stores the index of the next gate."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
