---
id: 69cfdc80443b715a8aa452a7
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Um Ihre `initializeThroughput`-Funktion zu vervollständigen:

- Verwenden Sie eine `for`-Schleife, um jedes Tor im Eingabe-`gates`-Array durchzulaufen. Fügen Sie für jedes Tor eine Eigenschaft zu `summary` hinzu, wobei:
  - Der Schlüssel die `id` des Tors ist.
  - Der Wert `0` ist.
- Geben Sie nach der Schleife das `summary`-Objekt zurück.

Ihre `initializeThroughput`-Funktion darf das Quell-Eingabe-Array `gates` nicht verändern.

# --hints--

Ihre `initializeThroughput`-Funktion sollte `{}` zurückgeben, wenn sie mit `[]` aufgerufen wird, ohne das ursprüngliche Eingabe-Array zu verändern.

```js
const input = [];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  {},
  "Calling `initializeThroughput([])` should return an empty object."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given an empty array."
);
```

Ihre `initializeThroughput`-Funktion sollte `{ North: 0 }` zurückgeben, wenn sie mit `[{ id: "North", capacity: 5, queue: [1, 2, 3] }]` aufgerufen wird, ohne das ursprüngliche Eingabe-Array zu verändern.

```js
const input = [{ id: "North", capacity: 5, queue: [1, 2, 3] }];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0 },
  "Calling `initializeThroughput([{ id: 'North', ... }])` should return `{ North: 0 }`."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given one gate."
);
```

Ihre `initializeThroughput`-Funktion sollte `{ North: 0, East: 0, South: 0 }` zurückgeben, wenn sie mit `[{ id: "North", capacity: 5, queue: [] }, { id: "East", capacity: 3, queue: [] }, { id: "South", capacity: 4, queue: [] }]` aufgerufen wird, ohne das ursprüngliche Eingabe-Array zu verändern.

```js
const input = [
  { id: "North", capacity: 5, queue: [] },
  { id: "East", capacity: 3, queue: [] },
  { id: "South", capacity: 4, queue: [] }
];

const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0, East: 0, South: 0 },
  "Calling `initializeThroughput` with multiple gates should return an object with an entry for each gate id set to 0."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given multiple gates."
);
```

Sie sollten eine `for`-Schleife verwenden, um jedes Gate-Objekt im Eingabe-`gates`-Array zu durchlaufen.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));
assert.match(cleaned, /for\s*\(/);
```

Ihre `initializeThroughput`-Funktion sollte `summary` nach der Schleife als letzte Zeile in der Funktion zurückgeben.

```js
const funcStr = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));

assert.match(
  funcStr,
  /returnsummary;?}$/,
  "Your `initializeThroughput` function should end with `return summary`."
);
```

Ihre `initializeThroughput`-Funktion darf das Quell-Eingabe-Array `gates` nicht verändern.

```js
const morningCopy = JSON.parse(JSON.stringify(morningGates));
const nightCopy = JSON.parse(JSON.stringify(nightGates));

initializeThroughput(morningGates);
initializeThroughput(nightGates);

assert.deepEqual(
  morningGates,
  morningCopy,
  "`initializeThroughput` should not mutate the `morningGates` seed array."
);

assert.deepEqual(
  nightGates,
  nightCopy,
  "`initializeThroughput` should not mutate the `nightGates` seed array."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
--fcc-editable-region--
  
--fcc-editable-region--
}
```
