---
id: 6ac0a4463662d9738c393e21
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Nel passaggio precedente, hai creato una **stringa**: una sequenza di caratteri racchiusa tra virgolette singole o virgolette doppie. In programmazione, le stringhe sono usate per rappresentare testo.

Il **terminale** è dove puoi vedere l'output del tuo codice.

Per poter vedere un output, trasforma la tua riga di codice esistente in:

```py
print('Hello')
```

Assicurati di iniziare dall'inizio della riga.

# --hints--

Dovresti avere `print('Hello')` nel tuo codice. Assicurati di iniziare dall'inizio della riga.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
