---
id: 6ac09f8420b4b67e075d2cd2
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop creerai il tuo primo programma Python.

Questo è il tuo editor di codice, dove scriverai codice Python. La riga 1 è evidenziata. Vai lì e digita:

```py
'Hello'
```

Assicurati di iniziare dall'inizio della riga e di non dimenticare le virgolette.

Quando hai finito, clicca sul pulsante "controlla il tuo codice" per vedere se è corretto.

# --hints--

Dovresti digitare `'Hello'`. Assicurati di iniziare dall'inizio della riga e di non dimenticare le virgolette.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
