---
id: 6ac0b35f3b29a27f7e207295
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Nel passaggio precedente, hai creato la variabile `sentence`. Ora stampa il suo valore nel terminale digitando `print(sentence)` sulla riga evidenziata.

Con questo, hai completato il tuo primo programma Python.

# --hints--

Dovresti avere `print(sentence)` nel tuo codice. Assicurati di iniziare dall'inizio della riga.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
