---
id: 6ac0aaba914d81b32567c076
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

`print` è una funzione integrata che puoi usare per inviare output al terminale. Come puoi vedere, il testo `Hello` è ora visibile sul terminale.

Ora, stampa il testo `freeCodeCamp` sul terminale andando alla riga evidenziata e digitando:

```py
print('freeCodeCamp')
```

Assicurati di iniziare dall'inizio della riga.

# --hints--

Dovresti avere `print('freeCodeCamp')` nel tuo codice. Assicurati di iniziare dall'inizio della riga.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
