---
id: 6564b8c9349bd76dc037967b
title: 步骤 29
challengeType: 20
dashedName: step-29
---

# --description--

`re` 模块中的 `search()` 函数分析作为参数传递的字符串，寻找正则表达式模式与字符串匹配的第一个位置。

声明一个名为 `quote` 的变量，并将字符串 `'Not all those who wander are lost.'` 分配给该变量。 然后打印 `pattern.search(quote)` 的结果。

# --hints--

你应该有一个 `quote` 变量。

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

你应该将提供的字符串分配给新的 `quote` 变量。

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

你应该打印 `pattern.search(quote)`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
