---
id: 656477845006313fbfea0ad1
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

尽管效果可能看起来与 `random.choice()` 相同，但 `secrets` 可确保你的操作系统能够提供的最安全的随机源。

现在，删除两个 `print()` 调用。

# --hints--

你应该删除两个 `print()` 调用。

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
