---
id: 656494269ccce754411a2112
title: 步骤 27
challengeType: 20
dashedName: step-27
---

# --description--

`re` 模块允许你在代码中使用*正则表达式*。 你很快就会学到更多有关正则表达式的知识。

现在，在代码顶部添加 `import` 语句以导入 `re` 模块。

# --hints--

你应该导入 `re` 模块。

```js
({ test: () => assert.match(code, /^import\s+re/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

import secrets
import string
--fcc-editable-region--

def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
