---
id: 6564642ce403fe29cc23fee7
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

Python 模块是一个包含用于执行特定任务的代码的文件。 Python 标准库包含许多可以在代码中导入和使用的模块。 你可以通过使用 `import` 语句后跟模块名称来实现此目的。

通过导入 `string` 模块启动此项目。

# --hints--

你应该导入 `string` 模块。

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
