---
id: 65649a631440b757f9f7b428
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

正则表达式是用于匹配字符串中特定字符组合的模式。 当你需要在字符串中匹配或查找模式以进行验证、字符替换等操作时，它是 `if`/`else` 条件语句的有效替代方案。

`re` 模块中的 `compile()` 函数将作为参数传递的字符串编译成正则表达式对象，该对象可被其他 `re` 方法使用。

声明一个新的 `pattern` 变量，并将 `re.compile('i')` 的值赋给该变量。

# --hints--

你应该声明一个 `pattern` 变量。

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

你应该将 `re.compile('i')` 赋值给你的 `pattern` 变量。

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
