---
id: 6564d68c34027a8072a704f4
title: 步骤 42
challengeType: 20
dashedName: step-42
---

# --description--

点字符是通配符，可以匹配字符串中的任何字符 - 默认情况下换行符除外。 修改 `pattern` 以匹配整个字符串，方法是将当前模式替换为 `.` 后跟 `+` 量词。

# --hints--

你应该修改 `pattern` 变量以匹配整个 `quote` 字符串。

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), ".+") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '[^a-z]t'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
