---
id: 656489e0a5d0954c1a64fc93
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

现在打印 `all_characters` 变量来查看它是什么样子。

# --hints--

你应该打印你的 `all_characters` 变量。

```js
({ test: () => assert.match(code, /^print\s*\(\s*all_characters\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
# Combine all characters
all_characters = letters + digits + symbols

--fcc-editable-region--
```
