---
id: 6564d97a5ef99783216a6229
title: 步骤 44
challengeType: 20
dashedName: step-44
---

# --description--

Python 提供了一种特殊类型的字符串，称为*原始*字符串。 原始字符串以 `r` 为前缀。 与常规字符串的主要区别在于它们如何处理反斜杠字符：在原始字符串中，反斜杠被视为文字字符而不是转义字符。 编写正则表达式时，使用原始字符串是一种很好的做法，因为它们通常可以包含很多 `\` 字符。

通过在 `pattern` 字符串前添加 `r` 作为前缀，将其转换为原始字符串。

# --hints--

你应该将 `pattern` 变量修改为 `r'\.'`。

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
