---
id: 656475bbf1c2573de1d2c69c
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

每次代码运行时，你应该会看到 `all_characters` 字符串中的一个随机字符。 这正是你要创建随机密码的目的。

不过，`random` 所依赖的算法使得生成的伪随机数是可预测的。 因此，尽管 `random` 模块适合于大多数常见的应用，因其决定性规律，它不能用于加密目的。 

不要导入 `random`，而是导入 `secrets` 模块。 然后将 `print()` 调用改为使用 `secrets.choice(all_characters)`。

# --hints--

你应该导入 `secrets` 模块而不是 `random` 模块。

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

你应该修改你现有的 `print()` 调用，以打印 `secrets.choice(all_characters)`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
