---
id: 65649122c7f77f519aaf0975
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

最后，声明一个变量 `new_password`，并将调用 `generate_password` 的结果赋给它。 将 `8` 作为参数传递给 `generate_password` 调用。

# --hints--

你应该调用 `generate_password` 并将 `8` 作为参数传递。

```js
({ test: () => assert.match(code, /generate_password\s*\(\s*8\s*\)/) })
```

你应该将 `generate_password(8)` 分配给变量 `new_password`。

```js
({
  test: () => {
    assert(__userGlobals.has("new_password"));
    assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*\)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--

--fcc-editable-region--
```
