---
id: 6564dd65c3c2fa873a83d213
title: 步骤 47
challengeType: 20
dashedName: step-47
---

# --description--

在字符类中，你可以通过在方括号内逐个写入范围（不带任何其他字符）来组合多个范围。 例如：`[a-d3-6]` 是 `[a-d]` 和 `[3-6]` 的组合。

现在，修改最后一个正则表达式模式以匹配任何非字母数字字符。 将 `a-z`、`A-Z` 和 `0-9` 范围组合成一个字符类，并添加 `^` 作为第一个字符来否定该模式。

# --hints--

你应该完成第四个模式以匹配任何非字母数字字符。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\[\^a-zA-Z0-9\]\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
