---
id: 6564683821b2ee3174e7250e
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

你可以通过点符号访问导入模块内定义的实用程序。 点符号的工作原理是将一个点加上实用程序名称附加到模块名称。 例如，下面是如何访问 `ascii_lowercase` 常量：

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

在这个项目中，你将使用来自 `string` 模块的不同常量。 

声明一个名为 `letters` 的新变量，并将 `string.ascii_letters` 分配给该变量。

# --hints--

你应该声明一个名为 `letters` 的变量。

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

你应该将 `string.ascii_letters` 分配给你的 `letters` 变量。

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
