---
id: 65646ffeaed2d238c562a014
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

将 `import` 语句放在代码顶部是一种常见的惯例。 此外，如果有多个 `import` 语句，请按字母顺序对其进行排序，以提高可读性。

在代码顶部，导入 `random` 模块。

# --hints--

你应该导入 `random` 模块。

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
