---
id: 6564a767a7241362f7d8d664
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

独立的单下划线用于表示你不关心的值或不会在代码中使用的值。 你的迭代变量实际上并未被使用。

将你的 `i` 变量修改为单个下划线。

# --hints--

你应该将 `i` 变量修改为 `_`。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/for\s+_\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--
    password = ''
    # Generate password
    for i in range(length):
        password += secrets.choice(all_characters)
    
--fcc-editable-region--
```
