---
id: 65648f4b2281ba50051ae39c
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

接下来，编写一个 `for` 循环，并以 `i` 作为循环变量。 使用 `range()` 函数迭代到 `length` 的值。

在循环内部，使用加法赋值运算符将 `all_characters` 中的随机字符添加到 `password` 的当前值。 为此，请使用 `secrets` 模块中的 `choice()` 函数。

# --hints--

你应该编写一个 `for` 循环来迭代 `range(length)`。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

你应该使用 `+=` 运算符将 `all_characters` 中的随机字符添加到 `password` 的当前值。

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password
    
--fcc-editable-region--
```
