---
id: 656474314bf37d3c83a85143
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

`random` 模块中的 `choice()` 函数接受一个序列并返回该序列的一个随机项。

修改你的 `print()` 调用以使用 `choice()` 函数并传递 `all_characters` 作为参数。

# --hints--

你应该修改你现有的 `print()` 调用以打印 `random.choice(all_characters)`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
