---
id: 656470d517833a39bb8b5608
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

`random` 模块包含一个伪随机数生成器。 它的大多数功能依赖于 `random()` 函数，该函数返回范围在 `0.0`（包含）和 `1.0` （不包含）之间的浮点数。

调用 `random` 模块中的 `random()` 函数并打印结果。

# --hints--

你应该打印调用 `random.random()` 的结果。

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
