---
id: 656479aa5f298441c190bf8f
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

 声明一个 `generate_password` 函数，并在函数体内编写除 `import` 行之外的所有代码。

# --hints--

你应该声明一个名为 `generate_password` 的函数。

```js
({
  test: () => assert(runPython(`
  _Node(_code).has_function('generate_password')
  `))
})
```

导入语句仍应位于函数之外。

```js
({
  test: () => assert(runPython(`
  imports_list = _Node(_code).find_imports()
  any([imp.is_equivalent('import secrets') for imp in imports_list]) and any([imp.is_equivalent('import string') for imp in imports_list])
  `))
})
```

四个变量声明应该移到函数内部。

```js
({
  test: () => assert(runPython(`
  func = _Node(_code).find_function('generate_password')
  vars = ['letters', 'digits', 'symbols', 'all_characters']
  func.find_variable('letters').is_equivalent('letters = string.ascii_letters') and \\
  func.find_variable('digits').is_equivalent('digits = string.digits') and \\
  func.find_variable('symbols').is_equivalent('symbols = string.punctuation') and \\
  func.find_variable('all_characters').is_equivalent('all_characters = letters + digits + symbols') and \\
  all(not _Node(_code).has_variable(var) for var in vars)
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

--fcc-editable-region--
```
