---
id: 6564d75a923d21815caaa445
title: 第 43 步
challengeType: 20
dashedName: step-43
---

# --description--

如果你需要匹配模式中具有特殊含义的字符，例如 `.` 或 `+`，则可以通过在前面添加反斜杠字符 `\` 来转义它。 例如，这与文字加号匹配：`\+`。

修改 `pattern` 以便它匹配单个文字点。

# --hints--

你应该将 `pattern` 变量修改为 `\.`。

```js
({ test: () => assert.match(code, /^pattern\s*=\s*("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '.+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
