---
id: 65647c71a5d2bd431596f629
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

你的 `generate_password` 函数需要一些参数。 首先添加 `length` 参数。

# --hints--

你的 `generate_password` 函数应该采用 `length` 参数。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == "(length)"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import secrets
import string


def generate_password():    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

--fcc-editable-region--
```
