---
id: 6564cbb6311a40783b3f5de6
title: 步骤 36
challengeType: 20
dashedName: step-36
---

# --description--

字符类还允许你指示要匹配的字符范围。 你需要指定以连字符分隔的起始字符和结束字符，`-`。 字符遵循 Unicode 顺序。

修改你的 `pattern` 变量以匹配 `quote` 变量中以小写字母开头的任何字母 `t`。 为此，请使用从 `a` 到 `z` 的字符范围。

# --hints--

你应该使用 `[a-z]` 类修改 `pattern` 变量，使其匹配 `quote` 中以小写字母开头的任何字母 `t`。

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
