---
id: 65646ab6bced52360f44670f
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

现在声明一个名为 `all_characters` 的变量，并将连接现有变量的结果赋给它。

# --hints--

你应该声明一个名为 `all_characters` 的变量。

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

你应该连接 `letters`、`digits` 和 `symbols` 并将结果分配给你的 `all_characters` 变量。

```js
({ test: () => assert(runPython(`
    import string
    from itertools import permutations
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == all_characters for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
