---
id: 656469e62377ab34a5bcb8d1
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

声明两个名为 `digits` 和 `symbols` 的新变量，并分别分配给 `string.digits` 和 `string.punctuation`。

# --hints--

你应该声明一个名为 `digits` 的变量。

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

你应该将 `string.digits` 分配给你的 `digits` 变量。

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

你应该声明一个名为 `symbols` 的变量。

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

你应该将 `string.punctuation` 分配给你的 `symbols` 变量。

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
