---
id: 65648e4c5b316c4ec5e4fddc
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

在函数的底部，声明一个 `password` 变量并为该变量分配一个空字符串。

# --hints--

你的函数内部应该有一个 `password` 变量，其值为空字符串。 注意缩进.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/password\s*=\s*("|')\1/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--

--fcc-editable-region--
```
