---
id: 656ef5bc5c0cd464be1df675
title: 步骤-23
challengeType: 20
dashedName: 步骤-23
---

# --description--

接下来，你将为你的函数提供更多参数，这些参数将作为生成密码的约束。

通过在现有的 `length` 参数后按顺序添加 `nums`、`special_chars`、`uppercase` 和 `lowercase` 来修改函数声明。

# --hints--

你的函数应该以 `length`、`nums`、`special_chars`、`uppercase` 和 `lowercase` 作为参数。 The order matters.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)

    return password

# new_password = generate_password(8)
# print(new_password)
```
