---
id: 656ef54070c72161e6feb90f
title: 步骤 22
challengeType: 20
dashedName: step-22
---

# --description--

看起来一切都很好，但如果有一种方法可以检查生成的密码是否符合特定功能就更好了。 例如，特殊字符、数字或大写/小写字母的最小数量。 你很快就会处理好这件事的。

现在，注释掉代码的最后两行。

# --hints--

你应该将代码的最后两行变成注释。

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
