---
id: 656ef89b4b486b7a16077864
title: 步骤 26
challengeType: 20
dashedName: step-26
---

# --description--

元组是 Python 中的另一种内置数据结构。 元组非常类似于列表，但它们是用括号 `()` 定义的，而不是用方括号。 此外，与列表不同，元组是不可变的。

```py
my_tuple = ('larch', 1, True)
```

你的 `constraints` 列表将存储元组。 每个元组的第一项将是一个约束参数。

通过向列表添加元组来修改 `constraints` 列表分配。 使用 `nums` 作为第一项，使用空字符串作为第二项。

# --hints--

你应该向 `constraints` 列表中添加一个元组。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

你应该将 `(nums, '')` 添加到 `constraints` 列表中。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
