---
id: 656eec2f40d18056cc58b229
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

由于在解析的字符串中未找到 `'i'`，因此返回值 `None`。

现在，将传递给 `re.compile()` 的字符串修改为 `'l'` 并查看结果。

# --hints--

你应该将 `pattern` 变量修改为 `re.compile('l')`。

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)    
--fcc-editable-region--
pattern = re.compile('i')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
