---
id: 656ef7f792734072dedd8319
title: 步骤 25
challengeType: 20
dashedName: step-25
---

# --description--

在 `for` 循环之后，创建一个 `constraints` 变量并为该变量分配一个空列表。

# --hints--

你应该在 `for` 循环后声明一个 `constraints` 变量并为该变量分配一个空列表。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/^(\s*)for.*:.*^\1constraints\s*=\s*\[\s*\]/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--
    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
