---
id: 656f0bc259c1f6b57486ed68
title: 步骤 33
challengeType: 20
dashedName: step-33
---

# --description--

要检查生成的密码是否满足所需的特性，你将使用 `re` 模块中的 `findall()` 函数。 它类似于 `search`，但它返回一个包含所有匹配模式出现的列表。

将 `search()` 调用替换为 `findall()` 并检查输出。

# --hints--

你应该修改现有的 `print()` 调用，将 `search()` 替换为 `findall()`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.search(pattern, quote))
--fcc-editable-region--
```
