---
id: 656f0d6bc60b58b9777a36c9
title: 步骤 34
challengeType: 20
dashedName: step-34
---

# --description--

字符类由方括号表示，并与括号中指定的字符中的一个字符匹配。 例如，`ma[dnt]` 与 `mad`、`man` 或 `mat` 匹配。

修改你的模式以匹配 `w` 后跟 `h` 或 `a`。

# --hints--

你的模式应该使用字符类来匹配 `w` 后跟 `h` 或 `a`。

```js
({ test: () => assert(runPython(`
    pattern == "w[ha]" or pattern == "w[ah]"    
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
