---
id: 6577319039f4f7de9251b822
title: 步骤 61
challengeType: 20
dashedName: step-61
---

# --description--

除了使用循环和计数器变量之外，你还可以使用不同的方法实现相同的结果，你将在接下来的几个步骤中实现该方法。

`all()` 是一个内置 Python 函数，如果给定可迭代对象内的所有元素的计算结果都为 `True`，则返回 `True`。 否则，返回 `False`。

用单个 `if` 语句替换现有的 `for` 循环和两个 `if` 语句。 对于 `if` 条件，使用对 `all()` 函数的调用并将空列表作为函数调用的参数传递。

# --hints--

你应该用单个 `if` 语句替换现有的 `for` 循环和两个 `if` 语句。

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

你的新 `if` 条件应该是 `all([])`。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

你的新 `if` 主体内应该有 `break`。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1

        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
