---
id: 6577333feab1e8e927014f03
title: 步骤 63
challengeType: 20
dashedName: step-63
---

# --description--

拥有 `all([expression for i in iterable])` 意味着通过对 `iterable` 中的每个 `i` 评估 `expression` 来创建一个新列表。 在 `all()` 函数迭代新创建的列表后，该列表会被自动删除，因为不再需要它。

可以使用生成器表达式来节省内存。 生成器表达式遵循列表推导的语法，但它们使用括号而不是方括号。

通过删除方括号，将列表推导更改为生成器表达式。

# --hints--

你应该通过删除方括号将列表理解转换为生成器表达式。

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        if all(
            [
                constraint <= len(re.findall(pattern, password))
                for constraint, pattern in constraints
            ]
        ):
--fcc-editable-region--
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
