---
id: 657732654845d8e2fb1217e6
title: 步骤62
challengeType: 20
dashedName: step-62
---

# --description--

现在，`all()` 正在将一个空列表作为参数。使用推导语法填充该空列表，使列表保存对 `constraints` 列表中每个 `constraint`-`pattern` 元组计算表达式 `constraint <= len(re.findall(pattern, password))` 的结果。

这样，只有在满足所有要求后，你才会跳出 `while` 循环。

# --hints--

你应该将 `[constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints]` 传递给 `all()` 函数。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("all([constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints])")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        if all([]):
--fcc-editable-region--        
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
