---
id: 657dadf4d8b93c1704f3a57c
title: 第 66 步
challengeType: 20
dashedName: step-66
---

# --description--

它确实有效，但仍有几点可以改进。 首先，调用具有 5 个参数的函数可能会造成混淆，不知道将哪个值分配给哪个参数。

你可以使用关键字参数调用函数，即明确写出参数名称，后跟赋值运算符和值。 例如：

```py
def add(x, y):
    return x + y

add(x=1, y=7) # 8
```

修改你的函数调用以使用关键字参数。

# --hints--

你的 `new_password` 变量应具有 `generate_password(length=8, nums=1, special_chars=1, uppercase=1, lowercase=1)` 的值。

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*length\s*=\s*8\s*,\s*nums\s*=\s*1\s*,\s*special_chars\s*=\s*1\s*,\s*uppercase\s*=\s*1\s*,\s*lowercase\s*=\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
   
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    
--fcc-editable-region-- 
new_password = generate_password(8, 1, 1, 1, 1)
print(new_password)
--fcc-editable-region--
```
