---
id: 61fd5a93fd62bb35968adeab
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Comience su proyecto proporcionando a su elemento `head` un elemento `link` para su hoja de estilo.

# --hints--

Su código debe tener un elemento `link`.

```js
assert.match(code, /<link/);
```

Tu elemento `link` debe estar dentro de tu elemento `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Su elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

Su elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
