---
id: 61fd6cc9475a784b7776233e
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Debajo de su elemento `div` existente, agregue un nuevo elemento `div` con una `class` establecida en `table-wrap`. Este será el contenedor para sus tablas.

Dentro de eso, agregue tres elementos `table`. Usará CSS para dar estilo a estas en una sola tabla, pero usar tablas separadas ayudará a los lectores de pantalla a entender el flujo del documento.

# --hints--

Debe crear un nuevo elemento `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Su nuevo elemento `div` debe tener la `class` establecida en `table-wrap`.

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

Su elemento `.table-wrap` debe venir después de su `div` existente.

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

Su elemento `.table-wrap` debe tener tres elementos `table`.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
