---
id: 68e97fe79367ad7b5dd6c9cd
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Nesta etapa, adicione um elemento `h1` com o texto `XYZ Bookstore`.

# --hints--

Você deve ter um elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

O texto do seu elemento `h1` deve ser `XYZ Bookstore`. Verifique novamente erros de ortografia ou capitalização.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

Seu elemento `h1` deve estar dentro do elemento `body`.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>
--fcc-editable-region--
<body>

</body>
--fcc-editable-region--
</html>
```
