---
id: 68e97fe79367ad7b5dd6c9cd
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você vai construir uma página de livraria criando cartões de livros que exibem informações sobre diferentes livros. Você vai praticar a organização de conteúdo usando elementos `div`, classes e IDs.

Comece adicionando um elemento `h1` com o texto `XYZ Bookstore`.

# --hints--

Você deve ter um elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

O texto do seu elemento `h1` deve ser `XYZ Bookstore`. Verifique novamente erros de ortografia ou capitalização.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
