---
id: 68ebdbacdd3fa474132cc975
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Abaixo do elemento `h1`, adicione um elemento `p` com este texto: `Browse our collection of amazing books!`.

# --hints--

Você deve ter um elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

O texto do seu elemento `p` deve ser: `Browse our collection of amazing books!`. 

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

O elemento `p` deve estar abaixo do elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>

--fcc-editable-region--
```
