---
id: 68eca2a795b333ca5fee30a8
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Agora crie um segundo cartão de livro. Adicione outro elemento `div` com o atributo `class` definido como `card`. Observe como você pode reutilizar o mesmo nome de classe para múltiplos elementos para aplicar uma estilização consistente.

# --hints--

Você deve ter dois elementos `div` com um atributo `class` com valor `card`.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

O segundo elemento com uma classe `card` deve estar dentro do elemento que possui uma classe `card-container`.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

Os dois elementos com uma classe `card` devem ser irmãos.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>

  --fcc-editable-region--
  </div>
</body>

</html>
```
