---
id: 68eca6364c5616d393389a30
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

Abaixo do elemento com a classe `card-container`, adicione um novo elemento `p` com este texto:

```md
Review your selections and continue to checkout.
```

Abaixo do elemento `p`, crie um elemento `div` com o atributo `class` definido como `btn-container`. Este contêiner irá agrupar seus elementos de botão de navegação.

# --hints--

O novo elemento `p` deve estar abaixo do elemento com uma classe `card-container`.

```js
assert.exists(document.querySelector('.card-container + p'));
```

O texto do seu novo elemento `p` deve ser `Review your selections and continue to checkout.`.

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

Você deve ter um elemento `div` com uma `class` de `btn-container`.

```js
assert.exists(document.querySelector('div.btn-container'));
```

Seu elemento `div` com uma `class` de `btn-container` deve vir depois do elemento `p` recém-adicionado.

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  <div class="card" id="dave-cooking-book">
    <h2>Dave's Cooking Adventure</h2>
    <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    <button class="btn">Buy Now</button>
  </div>
</div>

--fcc-editable-region--
```
