---
id: 68eca5c0065c82d256d29ca3
title: Passo 14
challengeType: 0
dashedName: step-14
---

# --description--

Dentro do segundo cartão, adicione um elemento `button` com o atributo `class` definido como `btn` e o texto `Buy Now`.

Ambos os elementos `button` agora compartilham a mesma `class`, o que significa que eles podem ser estilizados de forma consistente juntos.

# --hints--

Você deve ter um elemento `button` aninhado dentro do segundo elemento que possui uma classe `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

O elemento `button` do seu segundo cartão deve ter um atributo `class`.

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

O elemento `button` do seu segundo cartão deve ter uma `class` com o valor `btn`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

O texto do elemento `button` do seu segundo cartão deve ser `Buy Now`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  --fcc-editable-region--
  <div class="card" id="dave-cooking-book">
    <h2>Dave's Cooking Adventure</h2>
    <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
  </div>
  --fcc-editable-region--
</div>
```
