---
id: 68ec99e478211dc578699944
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Abaixo do elemento `h2` no primeiro elemento que possui a classe `card`, adicione um elemento `p` com o seguinte texto:

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

Você deve ter um elemento `p` aninhado dentro do elemento que possui uma classe `card`.

```js
assert.exists(document.querySelector('.card p'));
```

O texto do seu elemento `p` deve ser `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`.

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

Seu elemento `p` deve estar abaixo do seu elemento `h2`.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  --fcc-editable-region--
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>

  </div>
  --fcc-editable-region--
</div>
```
