---
id: 68eca11e99e3c5c894ca9d69
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

O elemento `button` é usado para criar botões clicáveis em uma página da web. Botões são elementos interativos que os usuários podem clicar para executar ações.

Adicione um elemento `button` dentro do elemento que tem uma `class` de `card`, dê ao botão um atributo `class` definido como `btn` e o texto `Buy Now`.

# --hints--

Você deve ter um elemento `button` aninhado dentro do elemento que possui uma classe `card`.

```js
assert.exists(document.querySelector('.card button'));
```

Seu elemento `button` deve ter um atributo `class`.

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

Seu elemento `button` deve ter uma `class` que tenha um valor de `btn`.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

O texto do seu elemento `button` deve ser `Buy Now`. 

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
