---
id: 68ec6d9a315221aa31e54816
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

O atributo `class` é usado para identificar um ou mais elementos para estilização. Diferente do atributo `id`, os nomes das classes não precisam ser únicos: múltiplos elementos podem compartilhar a mesma classe.

Aqui está um exemplo:

```html
<p class="example">example paragraph</p>
```

Adicione um atributo `class` ao seu elemento `div` e defina seu valor como `card-container`.

# --hints--

Seu elemento `div` deve ter um atributo `class`.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

Seu elemento `div` deve ter um atributo `class` com o valor `card-container`.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
--fcc-editable-region-- 
  <div>

  </div>
--fcc-editable-region--
</body>

</html>
```
