---
id: 68ec6cd5b7e8f5a8f7319e32
title: passo 6
challengeType: 0
dashedName: step-6
---

# --description--

O elemento `div` é usado como um container para agrupar outros elementos HTML. Você usará principalmente o elemento `div` quando quiser agrupar elementos HTML que compartilharão um conjunto de estilos CSS.

Abaixo do elemento `p`, adicione um elemento `div`. Este `div` será um container para seus cartões de livro.

**Nota**: Esta oficina não aplica CSS. Classes e elementos agrupados são úteis para estilização CSS, mas nesta oficina eles são usados apenas para estruturar e agrupar conteúdo. Você aprenderá como a estilização funciona em um módulo posterior.

# --hints--

Você deve ter um elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

Seu elemento `div` deve estar abaixo do seu elemento `p`.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
