---
id: 68ec9332a9b5b2b32487bd00
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

O atributo `id` adiciona um identificador único a um elemento HTML. Cada `id` deve ser único dentro de uma página e deve ser usado apenas uma vez.

Os valores de `id` não podem conter espaços e devem conter apenas letras, dígitos, underscores e hífens.

Aqui está um exemplo:

```html
<p id="para">example paragraph</p>
```

Adicione um atributo `id` ao seu elemento que tenha uma classe `card` e defina seu valor como `sally-adventure-book`.

# --hints--

Seu elemento com uma classe `card` deve ter um atributo `id`.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

Seu elemento com uma classe `card` deve ter um `id` com o valor `sally-adventure-book`.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
