---
id: 68eca4c4b0f952cf09fabe09
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Dentro do segundo elemento que possui a classe `card`, adicione um elemento `h2` com o texto `Dave's Cooking Adventure`.

# --hints--

Você deve ter um elemento `h2` aninhado dentro do segundo elemento que tem uma classe `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

O elemento `h2` dentro do segundo elemento que possui uma classe `card` deve ter um texto `Dave's Cooking Adventure`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText.trim(), "Dave's Cooking Adventure");
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  --fcc-editable-region--
  <div class="card" id="dave-cooking-book">

  </div>
  --fcc-editable-region--
</div>
```
