---
id: 68eca6e5a6759ed4ea0034dc
title: Passo 19
challengeType: 0
dashedName: step-19
---

# --description--

Dentro do elemento com uma classe `btn-container`, adicione dois elementos `button`:

Primeiro botão:

- Id: `view-cart-btn`
- Classe: `btn`
- Texto: `View Cart`

Segundo botão:

- Id: `checkout-btn`
- Classe: `btn`
- Texto: `Checkout`


Parabéns! Você construiu com sucesso a estrutura de uma página de livraria usando divs, classes e ids para organizar seu conteúdo.

# --hints--

Você deve ter dois elementos `button` aninhados dentro do elemento `.btn-container`.

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

Seu primeiro botão dentro do elemento com a classe `btn-container` deve ter um atributo `id` com o valor `view-cart-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

Seu primeiro botão dentro do elemento que possui uma classe `btn-container` deve ter um atributo `class` com o valor `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

O texto do seu primeiro botão deve ser `View Cart`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

Seu segundo botão dentro do elemento com uma classe `btn-container` deve ter um atributo `id` com o valor `checkout-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

Seu segundo botão dentro do elemento com `class="btn-container"` deve ter um atributo `class` com o valor `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

O texto do seu segundo botão deve ser `Checkout`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
--fcc-editable-region--
  <p>Review your selections and continue to checkout.</p>
  <div class="btn-container">

  </div>
--fcc-editable-region--
</body>

</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
