---
id: 68ec98b38d83a3c28dd30efe
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Dentro do primeiro elemento que possui a classe `card`, adicione um elemento `h2` com o texto `Sally's SciFi Adventure`.

# --hints--

Você deve ter um elemento `h2` aninhado dentro do elemento que possui uma classe `card`.

```js
assert.exists(document.querySelector('.card h2'));
```

O texto do seu elemento `h2` deve ser `Sally's SciFi Adventure`. 

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
