---
id: 68eca5412b5dd9d06b3d6404
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Abaixo do elemento `h2` no segundo card, adicione um elemento `p` com este texto:

```md
This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.
```

# --hints--

Você deve ter um elemento `p` aninhado dentro do segundo elemento que tem uma classe `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('p'));
```

O texto do elemento `p` do seu segundo cartão deve ser `This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('p')?.innerText.trim(), "This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.");
```

O elemento `p` do seu segundo cartão deve estar abaixo do seu elemento `h2`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
  --fcc-editable-region--
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
