---
id: 69134b6cb43aa5f254950a10
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você vai construir uma página de livraria criando cartões de livros que exibem informações sobre diferentes livros. Você vai praticar a organização de conteúdo usando elementos `div`, classes e IDs.

Comece sua página da livraria criando o boilerplate HTML.

Adicione a declaração `<!DOCTYPE html>` e os elementos `html` e `head`.

# --hints--

Seu código deve incluir uma declaração `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Seu código deve incluir as tags de abertura e fechamento `<html>`.

```js
assert.match(code,/<html>[\s\S]*<\/html>/i);
```

Seu código deve incluir as tags `<head>` de abertura e fechamento.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

Seu elemento `head` deve estar aninhado dentro do elemento `html`.

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
