---
id: 69134d14150d51f8066a3246
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Agora, melhore a estrutura do seu documento HTML para garantir que sua página esteja codificada corretamente.

Dentro do elemento `head`, adicione o elemento `<meta charset="UTF-8">`.

Por fim, adicione um elemento `body` abaixo da seção `head`. É aqui que todo o conteúdo visível da sua página será colocado.

# --hints--

Seu código deve incluir um elemento `<meta charset="UTF-8">`.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Seu código deve incluir as tags de abertura e fechamento `<body>`.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

Seu elemento `body` deve estar abaixo do elemento `head`.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
<head>
  <title>XYZ Bookstore Page</title>
</head>

</html>
--fcc-editable-region--
```
