---
id: 691f85085787a8c4ca589fbb
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Adicione o elemento `title` dentro do elemento `head`. 

Defina o título da página para `XYZ Bookstore Page`.

# --hints--

Seu código deve incluir as tags de abertura e fechamento `<title>`. 

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

Seu título deve ser `XYZ Bookstore Page`.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

Seu elemento `title` deve estar aninhado dentro do elemento `head`.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
  --fcc-editable-region--
<!DOCTYPE html>
<html>
<head>

</head>
  --fcc-editable-region-- 
</html>
```
