---
id: 66c07238b01053abaf812065
title: 建置一個階乘計算機
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

階乘是整數與所有比它小的整數的乘積。例如，`5` 的階乘是 `5 * 4 * 3 * 2 * 1 = 120`。在本實驗中，你將創建一個階乘計算器，該計算器從使用者取得一個數字並計算該數字的階乘。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該宣告一個變數 `num` 並指定一個你選擇的數字。指定的數字應該在 `1` 到 `20`（含）之間。
1. 創建一個名為 `factorialCalculator` 的函式（程式），該函式（程式）接受一個數字作為引數並傳回該數字的階乘。
1. 在函式（程式）內，宣告一個 `result` 變數並指定其值為 `1`。使用迴圈，從 `1` 到輸入的數字（包含）逐一迭代每個數字，並將 `result` 變數乘以當前數字後，再將結果指定給 `result` 變數。你可以選擇使用 `for` 迴圈、`while` 迴圈或 `do...while` 迴圈。
1. 你應該以 `num` 作為引數呼叫 `factorialCalculator` 函式，並將結果指定到變數 `factorial`。
1. 你應該將最終輸出以格式 `Factorial of [num] is [factorial]` 儲存，並指定給變數 `resultMsg`。
1. 你應該將 `resultMsg` 的值輸出到主控台。

# --hints--

你應該有一個 `num` 變數。

```js
assert.isDefined(num);
```

你應該為 `num` 變數指定一個值

```js
assert.isNotNull(num);
```

`num` 的值應該介於 `1` 和 `20` 之間。

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

`num` 值應該是一個數字。

```js
assert.isNumber(num);
```

你應該有一個名為 `factorialCalculator` 的函式（程式）。

```js
assert.isFunction(factorialCalculator);
```

`factorialCalculator` 函式（程式）應該接受一個數字作為引數。

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

`5` 的階乘應該傳回 `120`。

```js
assert.strictEqual(factorialCalculator(5), 120);
```

`7` 的階乘應該傳回 `5040`。

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

你應該使用 `num` 變數作為引數來呼叫你的 `factorialCalculator` 函式（程式）。

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s*$/m);
```

你應該定義一個 `factorial` 變數並將 `factorialCalculator` 函式（程式）的結果指定給它。

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

你的 `factorialCalculator` 應該產生正確的結果。

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

你的 `resultMsg` 應該包含一個格式為 `Factorial of [num] is [factorial]` 的字串。

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

你應該將 `resultMsg` 的值輸出到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```
