---
id: 69a030fbe487f185f55b8627
title: 第 13 步
challengeType: 1
dashedName: step-13
---

# --description--

現在，你將著手編寫一個函式，用來從列表中移除節點。

創建一個帶有 `remove` 和 `list` 參數的 `element` 函式。

# --hints--

你應該創建一個 `remove` 函式。

```js
assert.isFunction(remove)
```

你的 `remove` 函式應該有 `list` 和 `element` 參數。

```js
const removeParams = __helpers.getFunctionParams(remove.toString());
const removeParamsNames = removeParams.map(param => param.name)
assert.include(removeParamsNames, "list");
assert.include(removeParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

--fcc-editable-region--

--fcc-editable-region--

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
