---
id: 69a012c6224fac85832247dd
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

注意主控台中 `true`、`{ head: null, length: 0 }` 和 `true` 的結果嗎？這是因為你還沒有為列表 `myList` 指定任何值，也沒有將 `length` 屬性遞增。

繼續為 `node` 指定 `list.head`，然後在其下遞增 `length` 的 `list` 屬性。

# --hints--

你應該將 `node` 指定給 `list.head`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.head\s*=\s*node;?/);
```

你應該遞增 `length` 的 `list` 屬性。

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.length\+\+;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
