---
id: 69a00eaa43458e618331873c
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

現在你將著手編寫一個函式，用來將節點添加到列表中。

創建一個帶有 `add` 和 `list` 參數的 `element` 函式。

# --hints--

你應該創建一個 `add` 函式。

```js
assert.isFunction(add)
```

你的 `add` 函式應該有 `list` 和 `element` 參數。

```js
const addParams = __helpers.getFunctionParams(add.toString());
const addParamsNames = addParams.map(param => param.name)
assert.include(addParamsNames, "list");
assert.include(addParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
