---
id: 69a00a051b9a18694d00c16b
title: 第 1 步驟
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將實作一個鏈結串列。

首先，創建一個名為 `initList` 的函式。在裡面，傳回一個物件，該物件有一個 `head` 屬性設定為 `null`，以及一個 `length` 屬性設定為 `0`。

這個函式是你用來初始化鏈結串列的呼叫。`head` 代表列表中的第一個元素，最初它是 `null`，因為列表中沒有元素。`length` 用來追蹤列表中節點（元素）的數量，當你新增或移除節點時，會對它進行遞增或遞減。

# --hints--

你應該創建一個 `initList` 函式。

```js
assert.isFunction(initList)
```

你的 `initList` 函式應該傳回一個物件。

```js
assert.isObject(initList());
```

你的 `initList` 函式應該傳回一個物件，該物件有一個 `head` 屬性設定為 `null`。

```js
assert.propertyVal(initList(), "head", null);
```

你的 `initList` 函式應該傳回一個物件，該物件有一個 `length` 屬性設定為 `0`。

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
