---
id: 69a03ddce82b0626f329b001
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

現在，創建一個 `while` 迴圈，檢查 `current` 是否不嚴格 `null` 且 `current.element` 是否不嚴格相等於 element。

這表示迴圈應該持續遍歷列表直到到達列表末端或找到要移除的元素。

# --hints--

你應該創建一個 `while` 迴圈，條件為 `current !== null && current.element !== element`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /while\s+\(current\s*!==\s*null\s*&&\s*current\.element\s*!==\s*element\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    --fcc-editable-region--
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
