---
id: 69a03f70b50b4c4d9a0f0113
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

在迴圈內，將 `previous` 重新指定為 `current`，並將 `current` 重新指定為 `current.next`。這表示 `current` 會檢查下一個節點，而 `previous` 則永遠保持在前一個節點。

# --hints--

你應該將 `previous` 變數重新指定為 `current`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /previous\s*=\s*current;?/);
```

你應該將 `current` 變數重新指定為 `current.next`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) { 
        --fcc-editable-region--
                
        --fcc-editable-region--
    }
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
