---
id: 69a013a86c97dac3784b2636
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

現在，主控台應該會顯示類似 `{ head: { element: 42, next: null }, length: 1 }` 的內容。

不過，這並不是處理新增節點為列表中第一個元素（也就是頭節點）時的正確方式。

要正確處理，請使用 if 述語來檢查列表是否為空，方法是呼叫 `isEmpty` 函式並傳入 `list`。如果條件為真，請將你的 `list.head = node` 行移動到 `if` 述語的區塊內。

# --hints--

你應該有一個 `if` 述語，條件是 `isEmpty(list)`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

你應該將你的 `list.head = node` 行移動到 `if` 述語內。

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
