---
id: 69a00cd4ec513345d40e95ec
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

接下來你要處理的是一個可以讓你檢查列表是否為空的函式（程式）。

為此，創建一個名為 `isEmpty` 的函式，並帶有一個 `list` 參數。在函式內部，傳回一個檢查，判斷 `length` 的 `list` 屬性是否嚴格相等於 `0`。

# --hints--

你應該創建一個 `isEmpty` 函式。

```js
assert.isFunction(isEmpty)
```

你的 `isEmpty` 函式應該有一個 `list` 參數。

```js
const isEmptyParams = __helpers.getFunctionParams(isEmpty.toString());
assert.deepEqual(isEmptyParams[0].name, "list");
```

你的 `isEmpty` 函式應該傳回 `list.length === 0`。

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
