---
id: 69a016ac6edea7d8c8020c98
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

現在，創建一個 `while` 迴圈，條件是 `current.next !== null`。這表示只要目前節點後面還有下一個節點，迴圈就應該持續執行。

# --hints--

你應該創建一個 `while` 迴圈，條件是 `current.next !== null`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /while\s+\(current\.next\s*!==\s*null\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
