---
id: 69a02ee82bbeaee00e44c67d
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

這就是你需要做的全部，來將元素添加到列表。為了看看結果，請再為 `myList` 列表添加兩個你選擇的數字。

# --hints--

你應該為 `myList` 列表添加兩個你選擇的數字。

```js
const matches = __helpers.removeJSComments(code).match(/add\s*\(\s*myList\s*,\s*\d+\s*\);?/g) || [];

assert.isAtLeast(matches.length, 3);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
--fcc-editable-region--

--fcc-editable-region--
console.log(myList)
console.log(isEmpty(myList))
```
