---
id: 69a04466917c66b132323a1b
title: 第 18 步
challengeType: 1
dashedName: step-18
---

# --description--

如果 `previous` 不是 `null`，這表示要移除的元素不在列表的開頭，而是在列表中的某個位置，因此應該透過將前一個節點直接連結到下一個節點來跳過目前的節點。

為此，請創建另一個 `if` 述語，檢查 previous 是否嚴格不等於 `null`。如果是這種情況，將 `previous.next` 設為 `current.next`。

# --hints--

你應該創建另一個條件為 `if` 的 `previous !== null` 述語。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*/);
```

你應該將 `previous.next` 設為 `current.next`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*previous\.next\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--
    
--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
