---
id: 69a0102e761a245bb1cbd8e7
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

此時，你應該測試鏈結串列以加深理解。

創建一個 `myList` 變數並將其設為 `initList()`。在其下方，使用 `isEmpty()` 作為參數，將 `myList` 的呼叫結果記錄到主控台。接著，呼叫 `add`，傳入 `myList` 和你選擇的數字，然後將 `myList` 記錄到主控台。在那之下，再次使用 `isEmpty()` 作為參數，將 `myList` 的呼叫結果記錄到主控台。

# --hints--

你應該創建一個 `myList` 變數並設為 `initList()`。

```js
assert.deepEqual(myList, initList());
```

你應該將 `isEmpty(myList)` 的呼叫結果記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\);?/)
```

你應該呼叫 `add`，傳入 `myList` 和你選擇的數字作為參數。

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

你應該將你的 `myList` 變數記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?/)
```

你應該再次將 `isEmpty(myList)` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?\s*console\.log\(isEmpty\(myList\)\);?/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
