---
id: 69a01864d343dc2f2484d131
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

找到最後一個節點後，你需要將新節點連接到它。為此，請在 while 迴圈下方但仍在 `current.next` 語句內，將 `node` 設為 `else`。

# --hints--

你應該在 `current.next` 迴圈下方將 `node` 設為 `while`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\.next\s*=\s*node;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
