---
id: 69a0179a85b00b6127e0af45
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

在 `while` 迴圈內，將 `current` 更新為 `current.next`，使迴圈能夠每次往前移動一個節點，並在到達具有 `next` 屬性且值為 `null` 的節點時停止，也就是列表的結尾。

# --hints--

你應該在 `current` 迴圈內將 `current.next` 變數的值更新為 `while`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
