---
id: 69a014bd9b1730e94136279f
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

如果列表不為空，你需要找到它的尾端，將新元素添加到那裡，並讓它指向 `null`。這表示你必須從 `head` 開始迴圈遍歷列表。

為了開始，打開一個 `else` 語句。在裡面，使用 `let` 創建一個 `current` 變數並將它設為 `list.head`。

# --hints--

你應該有一個 `else` 語句。

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

你應該在 `current` 語句內創建一個 `list.head` 變數並設為 `else`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
