---
id: 69a03aaa1e6fe8152788ce6a
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

在 `remove` 函式（程式）內，使用 `let` 創建一個 `previous` 變數並設定為 `null`，再使用 `let` 創建另一個 `current` 變數並設定為 `list.head`。這些就是你將在巡訪列表時使用的變數。

`current` 用來追蹤目前正在檢查的節點，從列表的頭節點開始，每次往前移動一個節點。另一方面，`previous` 則用來追蹤目前節點的前一個節點。

這點在移除節點時非常重要，因為你必須更新前一個節點的 next 值，使其指向被移除節點的下一個節點，才能保持列表的連結。

# --hints--

你應該在 `let` 函式（程式）內使用 `previous` 創建一個 `null` 變數並設定為 `remove`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

你應該在 `let` 函式（程式）內使用 `current` 創建一個 `list.head` 變數並設定為 `remove`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
