---
id: 69a00f673f171e29159e3ad8
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

這個節點需要兩個屬性，一個是要添加的元素，另一個是指向下一個節點的指標。

在 add 函式（程式）內，創建一個名為 `node` 的物件，它應該有一個名為 `element` 的屬性，以及一個 `next` 屬性，該屬性設為 `null`，這是指向列表下一個元素的指標。

# --hints--

你應該在你的 `node` 函式（程式）內創建一個名為 `add` 的物件。

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

你的 `node` 物件應該有一個名為 `element` 的屬性，並設為 `element`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

你的 `node` 物件應該有一個名為 `next` 的屬性，並設為 `null`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
