---
id: 694a9ff65d4472dd451d8255
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви потренуєтеся використовувати булеві значення та умовні інструкції в Python, створюючи калькулятор для бронювання квитків на фільм.

Почніть з оголошення змінних, які зберігатимуть інформацію про користувача та показ фільму.

Як ви дізналися в попередніх уроках, змінним призначають значення за допомогою оператора присвоєння (`=`) таким чином:

```py
# Assigns the value 10 to variable x
x = 10
```

Створіть змінну з ім’ям `base_price`, щоб зберігати базову ціну квитка на фільм, і призначте їй значення `15`. Створіть іншу змінну з ім’ям `age`, щоб зберігати вік користувача, і призначте їй значення `21`.

# --hints--

Ви повинні оголосити змінну з ім’ям `base_price`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

Ви повинні призначити цілочисельне значення `15` змінній `base_price`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

Ви повинні оголосити змінну з ім’ям `age`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

Ви повинні призначити цілочисельне значення `21` змінній `age`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
