---
id: 694cc1b515976889afd4cded
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Тепер потрібно зберегти деякі рядкові значення про квиток на фільм. Ви можете зберігати рядки в Python таким чином:

```py
name = 'Alice'
```

Створіть змінну з ім’ям `seat_type`, щоб зберегти тип місця, яке вибрав користувач, і призначте їй значення рядка `Gold`. Створіть іншу змінну з ім’ям `show_time`, щоб зберегти час показу фільму, і призначте їй значення рядка `Evening`.

Пам’ятайте, що обидва значення потрібно взяти в одинарні або подвійні лапки.

# --hints--

Ви повинні оголосити змінну з ім’ям `seat_type`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

Ви повинні призначити рядок `Gold` змінній `seat_type`. Пам’ятайте, що значення потрібно взяти в одинарні або подвійні лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

Ви повинні оголосити змінну з ім’ям `show_time`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

Ви повинні призначити рядок `Evening` змінній `show_time`. Пам’ятайте, що значення потрібно взяти в одинарні або подвійні лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
