---
id: 6958e2377cacbc46a7f7383a
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Тепер додайте умову `else` до вашої інструкції `if` і виведіть `No extra charges will be applied` всередині тіла `else`. Це виведе повідомлення лише тоді, коли умова додаткових зборів буде хибною.

Потім, під умовою `else`, використайте виклик `print()` для відображення повідомлення, яке показує `Extra charges:`, а потім оновлене значення `extra_charges`, і перевірте вихідні дані в терміналі.

# --hints--

Ви повинні мати умову `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

Ви повинні вивести `No extra charges will be applied` всередині вашої умови `else`. Не забудьте взяти повідомлення в одинарні або подвійні лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

Ви повинні мати `print('Extra charges:', extra_charges)` під вашою інструкцією `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
