---
id: 6958de0d773289a89515a6b8
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер ви перевірите, чи дозволено користувачу забронювати вечірнє шоу залежно від його віку.

Створіть інструкцію `if`, щоб перевірити, чи `age` більший або дорівнює `21`. Усередині тіла інструкції `if` виведіть `User is eligible for Evening shows`.

# --hints--

У вашому коді має бути дві інструкції `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

Умова другої інструкції `if` має бути `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Ви маєте вивести `User is eligible for Evening shows` всередині нової інструкції `if`. Не забудьте взяти повідомлення в одинарні або подвійні лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
