---
id: 6958e59845777391b86d7331
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

У Python інструкцію `if` також можна розмістити всередині тіла іншої інструкції `if`. Це називається вкладеною інструкцією `if`.

Вкладена інструкція `if` дозволяє перевірити додаткову умову лише після того, як перша умова вже виконана. Внутрішня інструкція `if` виконається лише якщо зовнішня умова `if` є істинною.

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

Тепер ви розрахуєте сервісні збори залежно від типу місця, яке вибрав користувач.

Всередині тіла останньої інструкції `if`, нижче лінії `print('Ticket booking condition satisfied')`, створіть змінну з ім’ям `service_charges` і призначте їй значення `0`. Переконайтеся, що ваш код відступлений на чотири пробіли, щоб залишатися всередині тіла зовнішньої інструкції `if`.

Потім створіть вкладену інструкцію `if`, щоб перевірити, чи дорівнює `seat_type` значенню `Premium`. Всередині тіла вкладеної інструкції `if` оновіть значення `service_charges` на `5`.

# --hints--

Ви повинні оголосити змінну з ім’ям `service_charges` всередині тіла останньої інструкції `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

Ви повинні призначити цілочисельне значення `0` змінній `service_charges`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

У вас має бути інструкція `if` всередині тіла останньої інструкції `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

Умова нової інструкції `if` має бути `seat_type == 'Premium'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

Ви повинні призначити цілочисельне значення `5` змінній `service_charges` всередині нової інструкції `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
