---
id: 6958e1875e592b58e8852757
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Тепер створіть змінну з іменем `extra_charges` і призначте їй значення `0`. Це буде представляти додаткові збори, які застосовуються до квитка на фільм у вихідні.

Потім створіть інструкцію `if`, щоб перевірити, чи є `is_weekend` істинним. Усередині тіла інструкції `if` оновіть значення `extra_charges` на `2` і виведіть `Extra charges will be applied` у термінал.

# --hints--

Ви повинні оголосити змінну з іменем `extra_charges`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

Ви повинні призначити цілочисельне значення `0` змінній `extra_charges`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

У вас має бути четверта інструкція `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

Умова вашої четвертої інструкції `if` має бути `is_weekend`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

Ви повинні призначити цілочисельне значення `2` змінній `extra_charges` всередині нової інструкції `if`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

Ви повинні вивести `Extra charges will be applied` всередині нової інструкції `if`. Не забудьте взяти повідомлення в одинарні або подвійні лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
