---
id: 6958e0ecc121b0ae3dd9c5e0
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Ви також повинні опрацювати випадок, коли користувач не підпадає під знижку за членством.

Додайте умову `else` до вашої інструкції `if` і виведіть `User does not qualify for membership discount` всередині тіла `else`.

Ви можете вивести повідомлення, а потім змінну ось так:

```py
print('Total:', total)
```

Також потрібно показати оновлене значення `discount`. Під інструкцією `if...else` використайте виклик `print()`, щоб показати повідомлення, яке містить `Discount:` і оновлене значення `discount`. Потім перевірте вихідні дані в терміналі.

# --hints--

Ви повинні мати умову `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

Ви повинні вивести `User does not qualify for membership discount` всередині вашої умови `else`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

Ви повинні мати `print('Discount:', discount)` під вашою інструкцією `if...else`.

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
