---
id: 6958dd85587e10c99162dd81
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Як ви дізналися в попередніх уроках, у Python інструкцію `if` можна використовувати для запуску коду залежно від того, чи є умова істинною.

Інструкція `if` складається з ключового слова `if`, за яким слідує умова та двокрапка. Код, який виконується, коли умова істинна, і який має бути відступлений, називається тілом інструкції `if`.

```py
if condition:
    # Code to execute if condition is True
```

У цьому кроці ви перевірите, чи користувач має право забронювати квиток на фільм залежно від його віку.

Створіть інструкцію `if`, щоб перевірити, чи `age` більше за `17`. Усередині тіла інструкції `if` виведіть `User is eligible to book a ticket`. Це повідомлення буде виводитися лише тоді, коли вік користувача більший за `17`.

Не забудьте зробити відступ для тіла інструкції `if` і взяти повідомлення в одинарні або подвійні лапки всередині виклику `print()`.

# --hints--

У вас має бути інструкція `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

Умова інструкції `if` має бути `age > 17`.

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

Усередині інструкції `print()` має бути виклик `if`. Не забудьте зробити відступ для тіла інструкції `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

Ви маєте вивести `User is eligible to book a ticket` усередині інструкції `if`. Не забудьте взяти повідомлення в одинарні або подвійні лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
