---
id: 6960b6a6ea06e27dcc7eef37
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Тепер ви перевірите, чи користувач відповідає умовам для бронювання квитка на фільм. Користувачі віком `21` або старші завжди можуть бронювати квитки без обмежень.

Створіть інструкцію `if`, щоб перевірити, чи `age` більший або дорівнює `21`. Усередині тіла інструкції `if` виведіть `Ticket booking condition satisfied` у термінал.

Потім додайте умову `else` до вашої інструкції `if` і виведіть `Ticket booking failed due to restrictions` усередині тіла `else`.

# --hints--

Ви повинні мати п’яту інструкцію `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

Умова вашої п’ятої інструкції `if` має бути `age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

Ви повинні вивести `Ticket booking condition satisfied` усередині нової інструкції `if`. Не забудьте взяти повідомлення в одинарні або подвійні лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

Нова інструкція `if` має містити умову `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

Ви повинні вивести `Ticket booking failed due to restrictions` усередині умови `else`. Не забудьте взяти повідомлення в одинарні або подвійні лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
