---
id: 696099c1ac72a4e585de1d5b
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Деяка інформація може бути лише істинною або хибною. Як ви вже дізналися в попередніх уроках, це можна представити за допомогою булевих значень.

Створіть змінну з іменем `is_member`, щоб вказати, чи є користувач учасником, і призначте їй значення `True`.

Під змінною `is_member` створіть ще одну змінну з іменем `is_weekend`, щоб вказати, чи показ фільму відбувається у вихідні, і призначте їй значення `False`. Не беріть значення в лапки.

# --hints--

Ви повинні оголосити змінну з іменем `is_member`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

Ви повинні призначити булеве значення `True` змінній `is_member`. Не беріть значення в лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

Ви повинні оголосити змінну з іменем `is_weekend`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

Ви повинні призначити булеве значення `False` змінній `is_weekend`. Не беріть значення в лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
