---
id: 6960ba6e6c95ed4ad896c6f2
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Існує ще одне виключення з правил бронювання. Користувачі віком між `18` та `21` можуть бронювати вечірні шоу, якщо вони є членами.

Коли в `if` інструкції використовується кілька логічних операторів, умови, з’єднані `and`, оцінюються раніше за умови, з’єднані `or`. У Python для групування умов і контролю порядку їх оцінки використовують круглі дужки `()`.

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

Оновіть умову в `if age >= 21 or age >= 18 and show_time != 'Evening':` лінії, додавши ще одну умову з оператором `or`, щоб перевірити, чи `is_member` є істинним. Використайте круглі дужки `()`, щоб згрупувати умови `show_time != 'Evening'` та `is_member` разом, як показано в наведеному вище прикладі.

# --hints--

Остання умова вашої `if` інструкції має бути `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
