---
id: 696099abb6acc539b27055a1
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

У попередньому кроці ви перевірили умову за допомогою інструкції `if`. Умова `else` використовується для обробки випадку, коли умова є хибною. Ось синтаксис інструкції `if...else`:

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

Тепер додайте умову `else` до вашої інструкції `if` і виведіть `User is not eligible for Evening shows` всередині тіла `else`. Це виведе повідомлення лише тоді, коли умова в інструкції `if` є хибною.

# --hints--

Ви повинні мати умову `else`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

Ви повинні вивести `User is not eligible for Evening shows` всередині умови `else`. Не забудьте взяти повідомлення в одинарні або подвійні лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
