---
id: 6960a5165f7efe52a0c0aa39
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

У Python оператор `and` використовується для перевірки, чи виконуються кілька умов одночасно. Ось як ви можете використати його, щоб поєднати дві умови в `if` інструкції:

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

Знижка для учасників має застосовуватися лише до тих, у кого `age` більший або дорівнює `21`.

Оновіть умову в `if is_member:` лінії, використовуючи оператор `and`, щоб поєднати існуючу умову з іншою, яка перевіряє, чи `age` більший або дорівнює `21`.

# --hints--

Умова третьої `if` інструкції має бути `is_member and age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
