---
id: 6960b308bf41cc39c9b4b438
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

У Python оператор `or` використовується, щоб перевірити, чи принаймні одна з двох умов є істинною. Ось як ви можете його використати в інструкції `if`:

```py
if condition1 or condition2:
    # Code to execute if any condition is True
```

Додаткові збори також мають застосовуватися, якщо шоу відбувається ввечері. Оновіть умову в лінії `if is_weekend:`, використовуючи оператор `or`, щоб поєднати існуючу умову з другою, яка перевіряє, чи `show_time` дорівнює рядку `Evening`.

# --hints--

Умова вашої останньої інструкції `if` має бути `is_weekend or show_time == 'Evening'`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend or show_time == 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
--fcc-editable-region--
if is_weekend:
--fcc-editable-region--
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)
```
