---
id: 696cb889fc4fdeb337e52864
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

У Python умови `if` не обов’язково явно порівнюють значення з `True` або `False`. Замість цього вони покладаються на правдивість значень. Правдиві значення оцінюються як `True` у булевому контексті, наприклад, у умові інструкції `if`. До них належать ненульові числа, непорожні рядки та булеве значення `True`.

```py
if value:
   print('value is truthy')
```

Створіть інструкцію `if`, щоб перевірити, чи `is_member` є правдивим. Усередині тіла інструкції `if` оновіть значення `discount` на `3` і виведіть `User qualifies for membership discount` у термінал.

# --hints--

У вашому коді має бути три інструкції `if`.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

Умова вашої третьої інструкції `if` має бути `is_member`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

Ви маєте призначити ціле число `3` змінній `discount` усередині нової інструкції `if`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

Ви маєте вивести `User qualifies for membership discount` усередині нової інструкції `if`. Не забудьте взяти повідомлення в одинарні або подвійні лапки.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
