---
id: 66ebd4ae2812430bb883c787
title: Створіть центр подій
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

У цій лабораторній роботі ви використаєте семантичні HTML-елементи, щоб створити структуру вебсторінки. Ви додасте вміст і зображення, щоб вона виглядала як справжній центр подій.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути елемент `header`.
1. Всередині елемента `header` має бути елемент `h1`, який містить текст `Event Hub`, і елемент `nav`.
1. Всередині елемента `nav` має бути невпорядкований список із двох елементів, що містять посилання на різні розділи сторінки. Перший елемент має містити текст `Upcoming Events`, а другий — текст `Past Events`.
1. Кожне посилання має бути представлене елементом `a` з атрибутом `href`, який посилається на відповідний розділ сторінки, `#upcoming-events` та `#past-events` відповідно.
1. У вас має бути елемент `main`, який містить різні розділи сторінки.
1. Всередині елемента `main` має бути два елементи `section`.
1. Перший елемент `section` має мати атрибут `id` зі значенням `upcoming-events`.
1. Усередині розділу `#upcoming-events` має бути:
   
   - Елемент `h2` з текстом `Upcoming Events`.
   - Два елементи `article`. Кожен article має представляти подію і містити:
      - Елемент `h3` для назви події.
      - Елемент `p` для опису події. Ви можете додати дату внизу, якщо хочете.

1. Другий елемент `section` має мати атрибут `id` зі значенням `past-events`.
1. Усередині розділу `#past-events` має бути:

   - Елемент `h2` з текстом `Past Events`.
   - Два елементи `article`. Кожен article має представляти минулу подію і містити:
     - Елемент `h3` для назви події,
     - Елемент `p` для опису події. Ви можете додати дату внизу, якщо хочете.
     - Елемент зображення з атрибутом `src`, що вказує на файл зображення, та атрибутом `alt` з описом зображення.

**Примітка:** Ви можете використовувати будь-який текст для описів подій і дат. Якщо хочете, можете використати такі URL-адреси зображень:

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

У вас має бути елемент `header`.

```js
assert.isNotNull(document.querySelector("header"));
```

Ваш елемент `header` має йти після відкриваючого тегу `body`.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

Всередині елемента `header` має бути елемент `h1`, який містить текст `Event Hub`.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

Всередині елемента `header`, після елемента `h1`, має бути елемент `nav`.

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

Ваш елемент `nav` має містити невпорядкований список із двох елементів.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

Перший елемент у невпорядкованому списку має бути посиланням.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

Другий елемент у невпорядкованому списку має бути посиланням.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

Текст першого елемента у невпорядкованому списку має бути `Upcoming Events`.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

Перший елемент у невпорядкованому списку має мати атрибут `href`, встановлений у `#upcoming-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

Текст другого елемента у невпорядкованому списку має бути `Past Events`.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

Другий елемент у невпорядкованому списку має мати атрибут `href`, встановлений у `#past-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

У вас має бути елемент `main` після закриваючого тегу елемента `header`.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

Всередині елемента `main` має бути два елементи `section`.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

Ваш перший елемент `section` має мати атрибут `id` зі значенням `upcoming-events`.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

Ваш другий елемент `section` має мати атрибут `id` зі значенням `past-events`.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

Всередині розділу `#upcoming-events` має бути елемент `h2` з текстом `Upcoming Events`.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

Всередині розділу `#upcoming-events` має бути два елементи `article` під елементом `h2`.

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Обидва елементи `article` всередині розділу `#upcoming-events` мають мати елемент `h3` для назви події.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Обидва елементи `article` всередині розділу `#upcoming-events` мають мати елемент абзацу для опису події.

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Всередині розділу `#past-events` має бути елемент `h2` з текстом `Past Events`.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

Всередині розділу `#past-events` має бути два елементи `article` під елементом `h2`.

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Обидва елементи `article` всередині розділу `#past-events` мають мати елемент `h3` для назви події.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Обидва елементи `article` всередині розділу `#past-events` мають мати елемент абзацу для опису події.

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Обидва елементи `article` всередині розділу `#past-events` мають мати елемент зображення.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

Обидва елементи зображення всередині розділу `#past-events` мають мати атрибут `src`, що вказує на файл зображення.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

Обидва елементи зображення всередині розділу `#past-events` мають мати атрибут `alt` з описом зображення.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

Кожен елемент `h3` має містити назву події.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

Кожен елемент `p` має містити опис події.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```
