---
id: 6998b99a8e195cd0c1f211df
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, crearás un rastreador de vida silvestre simple usando objetos de JavaScript.

Aquí tienes un ejemplo de un objeto:

```js
const animal = {}; // empty object
```

Esto crea un objeto vacío.

Ahora crea una variable llamada `tiger` y asígnale un objeto vacío.

# --hints--

Debes crear una variable llamada `tiger`.

```js
assert.exists(tiger);
```

La variable `tiger` debe asignarse a un objeto vacío.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
