---
id: 6999adc72c02812a28c01e31
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

En este paso, crearás una función que actualice una propiedad de un objeto.

Aquí tienes un ejemplo de cómo actualizar una propiedad dentro de una función:

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
};

console.log(changeAge(dog, 6)); // { age: 6 }
```

En este ejemplo, la propiedad `age` se actualiza a un nuevo valor.

Ahora crea una función llamada `updateAge`. La función debe recibir dos parámetros: `animal` y `newAge`.

Dentro de la función, actualiza la propiedad `age` del objeto `animal` a `newAge`. Devuelve el objeto `animal` actualizado.

Después de crear la función, usa `console.log` para llamar a `updateAge(elephant, 12)` y así puedas ver el objeto `elephant` actualizado en la consola.

# --hints--

Debes crear una función llamada `updateAge`.

```js
assert.isFunction(updateAge);
```

La función `updateAge` debe tener dos parámetros: `animal` y `newAge`.

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

La función `updateAge` debe devolver el objeto `animal` actualizado.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

Debes registrar `updateAge(elephant, 12)` en la consola.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

Llamar a `updateAge(elephant, 12)` debe actualizar la propiedad age.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` debe usar los parámetros de la función y funcionar con cualquier objeto.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
