---
id: 6999b0198d18d402c36cc5fe
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

En este paso, crearás una función que elimina una propiedad de un objeto.

Aquí tienes un ejemplo de cómo eliminar una propiedad usando `delete`:

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

En este ejemplo, la propiedad `canFly` se elimina del objeto `bird`.

Ahora crea una función llamada `removeEndangeredStatus`. La función debe recibir un parámetro llamado `animal`.

Dentro de la función, elimina la propiedad `isEndangered` del objeto `animal` usando la palabra clave `delete`. Devuelve el objeto `animal` actualizado.

Después de crear la función, usa `console.log` para llamar a `removeEndangeredStatus(tiger)` y así puedas ver el objeto actualizado en la consola.

# --hints--

Debes crear una función llamada `removeEndangeredStatus`.

```js
assert.isFunction(removeEndangeredStatus);
```

La función `removeEndangeredStatus` debe tener un solo parámetro, `animal`.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

La función `removeEndangeredStatus` debe devolver el objeto `animal` actualizado.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` debe usar la palabra clave `delete` para eliminar la propiedad `isEndangered`.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

Debes mostrar `removeEndangeredStatus(tiger)` en la consola.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

Llamar a `removeEndangeredStatus(tiger)` debe eliminar la propiedad `isEndangered`.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` debe usar los parámetros de la función y funcionar con cualquier objeto.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
