---
id: 6999b07c34c2cde21ee6e1d8
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

En este paso, crearás una función que verifica si un objeto tiene una propiedad específica.

Aquí tienes un ejemplo de cómo usar `hasOwnProperty`:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

En este ejemplo, la función `hasColor` usa `hasOwnProperty` para comprobar si la propiedad `color` existe en un objeto.

Ahora crea una función llamada `hasHabitat`. La función debe recibir un parámetro llamado `animal`.

Usa el método `hasOwnProperty` para devolver `true` si el objeto `animal` tiene una propiedad llamada `"habitat"` y `false` en caso contrario.

Después de crear la función, usa `console.log` para llamar a `hasHabitat(tiger)` y `hasHabitat(elephant)` para que puedas ver ambos resultados en la consola.

# --hints--

Debes crear una función llamada `hasHabitat`.

```js
assert.isFunction(hasHabitat);
```

La función `hasHabitat` debe tener un solo parámetro, `animal`.

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` debe usar el método `hasOwnProperty` para verificar la propiedad `habitat`.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

Debes mostrar `hasHabitat(tiger)` en la consola.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

Debes mostrar `hasHabitat(elephant)` en la consola.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

Llamar a `hasHabitat(tiger)` debe devolver `true` después de agregar un hábitat.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Llamar a `hasHabitat(elephant)` debe devolver `false` si no existe un hábitat.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` debe usar los parámetros de la función y funcionar con cualquier objeto.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
