---
id: 69999ff7f7862fdfb885ea5d
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Ahora actualiza tu objeto `tiger` agregando una nueva propiedad llamada `age`.

Asigna a la propiedad `age` el valor `5`.

# --hints--

`tiger` debe tener una propiedad `age`.

```js
assert.property(tiger, "age");
```

`tiger.age` debe ser igual a `5`.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
