---
id: 6999b0d79ea938edcb032237
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

En este paso final del taller Wildlife Tracker, crearás una función que accede a una propiedad de un objeto usando la notación de corchetes.

Aquí tienes un ejemplo de cómo acceder a una propiedad usando la notación de corchetes:

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

En este ejemplo, la notación de corchetes te permite acceder a una propiedad usando una variable.

Ahora crea una función llamada `getProperty`.

La función debe recibir dos parámetros: `animal` y `propertyName`.

Devuelve el valor de la propiedad usando la notación de corchetes.

Después de crear la función, usa `console.log` para llamar a `getProperty(tiger, "species")` y `getProperty(elephant, "age")` y así puedas ver los valores devueltos en la consola.

Con eso, ¡el taller Wildlife Tracker está completo!

# --hints--

Debes crear una función llamada `getProperty`.

```js
assert.isFunction(getProperty);
```

La función `getProperty` debe tener dos parámetros: `animal` y `propertyName`.

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty` debe usar la notación de corchetes para acceder a la propiedad.

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

Debes mostrar `getProperty(tiger, "species")` en la consola.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

Debes mostrar `getProperty(elephant, "age")` en la consola.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

Llamar a `getProperty(tiger, "species")` debe devolver `"Tiger"`.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

Llamar a `getProperty(elephant, "age")` debe devolver `12`.

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty` debe usar los parámetros de la función y funcionar con cualquier objeto.

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
