---
id: 6999a0c1698b93d5d41156ed
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Ahora agrega otra propiedad al objeto `tiger` llamada `isEndangered`.

Asigna a la propiedad `isEndangered` el valor `true`.

# --hints--

`tiger` debe tener una propiedad `isEndangered`.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` debe ser igual a `true`.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
