---
id: 6999faf886634b72be7a819b
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Para este rastreador de vida silvestre, actualizarás un objeto con nuevas propiedades.

Aquí tienes un ejemplo de cómo crear un objeto con una propiedad:

```js
const animal = {
  name: "Lion" // property: value
};
```

Ahora actualiza el objeto `tiger` para que incluya una propiedad llamada `species`.

Asigna a esa propiedad el valor `"Tiger"`.

# --hints--

`tiger` debe ser un objeto.

```js
assert.isObject(tiger);
```

`tiger` debe tener una propiedad `species`.

```js
assert.property(tiger, "species");
```

`tiger.species` debe ser igual a `"Tiger"`.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
