---
id: 6999ac171928136168654774
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

En este paso, crearás una función para acceder a la propiedad `species` de un objeto.

Aquí tienes un pequeño ejemplo de cómo acceder a una propiedad de un objeto:

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

En este ejemplo, la función recibe un objeto como parámetro y devuelve su propiedad `species`.

Ahora crea una función llamada `getSpecies`.

La función debe recibir un parámetro llamado `animal`.

Devuelve la propiedad `species` del `animal` usando la notación de punto.

Después de crear la función, usa `console.log` para llamar a `getSpecies(tiger)` y así puedas ver el resultado en la consola.

# --hints--

Debes crear una función llamada `getSpecies`.

```js
assert.isFunction(getSpecies);
```

La función `getSpecies` debe tener un solo parámetro, `animal`.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Debes mostrar `getSpecies(tiger)` en la consola.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` debe acceder a la propiedad species usando la notación de punto.

```js
assert.match(code, /animal\.species/);
```

Llamar a `getSpecies(tiger)` debe devolver `"Tiger"`.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` debe usar el parámetro de la función y funcionar con cualquier objeto.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
