---
id: 699a00de9f564bb0effaa14d
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

En este paso, crearás una función para acceder a la propiedad `age` de un objeto.

Aquí tienes un ejemplo:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

En este ejemplo, la función recibe un objeto como parámetro y devuelve su propiedad `age`.

Ahora crea una función llamada `getAge`.

La función debe recibir un parámetro llamado `animal`.

Devuelve la propiedad `age` del `animal` usando la notación de punto.

Después de crear la función, usa `console.log` para llamar a `getAge(tiger)` y así puedas ver el resultado en la consola.

# --hints--

Debes crear una función llamada `getAge`.

```js
assert.isFunction(getAge);
```

La función `getAge` debe tener un solo parámetro, `animal`.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` debe acceder a la propiedad `age` usando la notación de punto.

```js
assert.match(code, /animal\.age/);
```

Debes mostrar `getAge(tiger)` en la consola.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

Llamar a `getAge(tiger)` debe devolver `5`.

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` debe usar el parámetro de la función y funcionar con cualquier objeto.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
