---
id: 69ba8da6b296391d913096a0
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Ahora crea un segundo objeto llamado `elephant`.

Agrega las siguientes propiedades:

- `species` con el valor `"Elephant"`
- `age` con el valor `10`
- `isEndangered` con el valor `true`

# --hints--

Debes crear una variable llamada `elephant`.

```js
assert.exists(elephant);
```

`elephant` debe tener una propiedad `species` igual a `"Elephant"`.

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` debe tener una propiedad `age` igual a `10`.

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` debe tener una propiedad `isEndangered` igual a `true`.

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
