---
id: af7588ade1100bde429baf20
title: بناء كاشف الحرف المفقود
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة باسم `fearNotLetter`.
1. يجب أن تقبل دالة `fearNotLetter` معلمة واحدة: سلسلة نصية تمثل نطاقًا من الحروف مرتبة أبجديًا وقد يكون فيها حرف مفقود.
1. يجب أن تجد الدالة الحرف المفقود في نطاق الحروف الممرر وتُرجعه.
1. إذا كانت كل الحروف موجودة في النطاق، يجب أن تُرجع الدالة `undefined`.

# --hints--

يجب أن تملك دالة `fearNotLetter`.

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` يجب أن تُرجع السلسلة النصية `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` يجب أن تُرجع السلسلة النصية `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` يجب أن تُرجع السلسلة النصية `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` يجب أن تُرجع السلسلة النصية `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
