---
id: 66b205e6eacba4c4e54ea434
title: Construye un Enmascarador de Correo Electrónico
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

En este laboratorio, enmascararás la parte del nombre de usuario de una dirección de correo electrónico con asteriscos. Enmascaramiento es un término utilizado para ocultar o reemplazar información sensible con asteriscos u otros caracteres.

Por ejemplo, si la dirección de correo electrónico fuera `myEmail@email.com`, entonces la dirección de correo enmascarada será `m*****l@email.com`.

**Objetivo:** Cumple las historias de usuario a continuación y pasa todas las pruebas para completar el laboratorio.

**Historias de Usuario:**

1. Crea una función llamada `maskEmail` que tome `email` como argumento.
2. Dentro de la función, debes enmascarar el `email` y añadir el nombre de dominio a él.  Recuerda que puedes usar métodos como `slice`, `repeat`, `indexOf` o incluso `replace` para ayudarte.
3. Fuera de la función, declara una variable llamada `email` para almacenar la dirección de correo electrónico que deseas enmascarar.
4. Llama la función `maskEmail` con la variable `email` y muestra el resultado en la consola.
5. `maskEmail("apple.pie@example.com")` debería devolver `"a*******e@example.com"`.
6. `maskEmail("freecodecamp@example.com")` debería devolver `"f**********p@example.com"`.
7. `maskEmail("info@test.dev")` debería devolver `"i**o@test.dev"`.
8. `maskEmail("user@domain.org")` debería devolver `"u**r@domain.org"`.

# --hints--

Debes definir una función llamada `maskEmail`.

```js
assert.isFunction(maskEmail);
```

La función `maskEmail` debe tomar una cadena, `email` como argumento.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

Fuera de la función, deberías tener una variable `email`.

```js
assert.isDefined(email);
```

Debes asignar una dirección de correo electrónico válida a tu variable `email`.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` debería devolver `"a*******e@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` debería devolver `"f**********p@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` debería devolver `"i**o@test.dev"`.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` debería devolver `"u**r@domain.org"`.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

Tu `maskEmail` debería producir el resultado correcto.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

Debes registrar el resultado de llamar a `maskEmail` con `email` como argumento.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
