---
id: af7588ade1100bde429baf20
title: 构建一个缺失字母检测器
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个名为 `fearNotLetter` 的函数。
1. `fearNotLetter` 函数应接受一个参数：一个表现按字母顺序排列且可能缺少一个字母的字母范围的字串。
1. 该函数应查找传入字母范围中缺失的字母并返回它。
1. 如果所有字母都在该范围内，函数应返回 `undefined`。

# --hints--

你应该有一个 `fearNotLetter` 函数。

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` 应该返回字符串 `d`。

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` 应返回 `i`。

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` 应该返回字符串 `u`。

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` 应该返回字符串 `e`。

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` 应返回 `undefined`。

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
