---
id: 68c15a2a1f0e70ca7154aba5
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop bauen Sie ein Gehaltsverfolgungssystem fĂ¼r Mitarbeitende.

Beginnen Sie damit, eine Klasse mit dem Namen `Employee` zu erstellen. Erstellen Sie darin eine `__init__`-Methode mit den Parametern `self`, `name` und `level`. Innerhalb der `__init__`-Methode weisen Sie `name` und `level` den Instanzattributen mit demselben Namen zu.

# --hints--

Sie sollten eine Klasse mit dem Namen `Employee` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

Ihre `Employee`-Klasse sollte eine `__init__`-Methode mit drei Parametern `self`, `name` und `level` haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

Ihre `__init__`-Methode sollte `self.name` auf `name` setzen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Ihre `__init__`-Methode sollte `self.level` auf `level` setzen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

```
