---
id: 68c16924498728d259e4781c
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Erstellen Sie nun eine Instanz der `Employee`-Klasse, indem Sie die Strings `Charlie Brown` und `trainee` übergeben. Weisen Sie die Instanz einer Variablen mit dem Namen `charlie_brown` zu.

# --hints--

Sie sollten eine Variable mit dem Namen `charlie_brown` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

Sie sollten `Employee('Charlie Brown', 'trainee')` `charlie_brown` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

--fcc-editable-region--

--fcc-editable-region--

```
