---
id: 68c18b1d0fce081f85312d97
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

In einer vorherigen Lektion haben Sie gelernt, dass ein Attribut mit einem einzelnen Unterstrich als Präfix konventionell für die interne Verwendung gedacht ist.

Ändern Sie beide `name`- und `level`-Attribute in `_name` und `_level`, da diese nicht von außerhalb ihrer Klasse verändert werden sollen.

Beachten Sie, dass dies nicht verhindert, dass auf das Attribut von außerhalb der Klasse zugegriffen oder es verändert wird. Außerdem gibt es in Python immer eine Möglichkeit, auch auf private Attribute (mit doppeltem Unterstrich als Präfix) zuzugreifen.

# --hints--

Ihre `__init__`-Methode sollte `self._name` auf `name` setzen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

Ihre `__init__`-Methode sollte `self._level` auf `level` setzen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
--fcc-editable-region--
        self.name = name
        self.level = level
--fcc-editable-region--

charlie_brown = Employee('Charlie Brown', 'trainee')

```
