---
id: 68c27fa1697ceb7c578b2c9f
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Der `@property`-Decorator wird in Python verwendet, um eine Methode in eine Eigenschaft zu verwandeln. Er wird typischerweise genutzt, um Getter-Methoden zu definieren, also Methoden, die den Wert eines Attributs abrufen:

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

Erstellen Sie eine Methode namens `name` mit einem `self`-Parameter und dekorieren Sie sie mit `@property`. Geben Sie innerhalb der Methode `self._name` zurück.
    
# --hints--

Ihre `Employee`-Klasse sollte eine `name`-Methode mit einem `self`-Parameter haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

Ihre `name`-Methode sollte mit `@property` dekoriert sein.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

Ihre `name`-Methode sollte `self._name` zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

--fcc-editable-region--
    
--fcc-editable-region--

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

```
