---
id: 68c28d659d345270d4a61084
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Die `__repr__`-Methode ist eine spezielle Methode, die eine String-Darstellung des Objekts zurückgeben soll, mit der es instanziiert werden kann.

Zum Beispiel sollte die `__repr__`-Methode von `Employee('Charlie', 'developer')` den String `Employee('Charlie', 'developer')` zurückgeben, der derselbe String ist, mit dem das Objekt erstellt wurde.

Geben Sie Ihrer `Employee`-Klasse eine `__repr__`-Methode mit einem `self`-Parameter und lassen Sie sie einen String zurückgeben, mit dem das Objekt instanziiert werden kann.

# --hints--

Ihre `Employee`-Klasse sollte eine `__repr__`-Methode mit einem `self`-Parameter haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

Ihre `__repr__`-Methode sollte den String zurückgeben, der zur Instanziierung des Objekts verwendet wird.

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'
        
--fcc-editable-region--
    
--fcc-editable-region--
    
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

```
