---
id: 68c281cecda94ba2c6838e17
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Jetzt, da Sie einen Getter für `name` definiert haben, können Sie auf das `_name`-Attribut über die `name`-Eigenschaft zugreifen. Geben Sie also `charlie_brown.name` in der Konsole aus.

# --hints--

Sie sollten `charlie_brown.name` in der Konsole ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--

--fcc-editable-region--

```
