---
id: 68c2d2414f5e111c02f4db5e
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Die `__repr__`-Methode wird im Hintergrund aufgerufen, wenn Sie die `repr`-Funktion aufrufen. Um das in Aktion zu sehen, geben Sie das Ergebnis des Aufrufs von `repr(charlie_brown)` am Ende Ihres Codes aus.

# --hints--

Sie sollten `repr(charlie_brown)` ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(repr(charlie_brown))")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
        
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--

--fcc-editable-region--

```
