---
id: 68c28b7cc3ab8f4daf144b5e
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Folgen Sie dem, was Sie in den vorherigen Schritten gemacht haben, und erstellen Sie einen Getter für das `_level`-Attribut.

# --hints--

Ihre `Employee`-Klasse sollte eine `level`-Methode mit einem `self`-Parameter haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_args("self")`)) })
```

Ihre `level`-Methode sollte mit `@property` dekoriert sein.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_decorators("property")`)) })
```

Ihre `level`-Methode sollte `self._level` zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_return("self._level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

```
