---
id: 68c27eb292bfec6b9402d58f
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Fügen Sie der `Employee`-Klasse eine `__str__`-Methode hinzu. Lassen Sie sie einen f-String im Format `name: level` zurückgeben, wobei `name` und `level` durch die entsprechenden Attribute ersetzt werden.

Geben Sie danach `charlie_brown` in der Konsole aus.

# --hints--

Ihre `Employee`-Klasse sollte eine `__str__`-Methode mit einem `self`-Parameter haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

Ihre `__str__`-Methode sollte einen f-String im Format `name: level` zurückgeben, wobei `name` und `level` durch die Werte von `self._name` bzw. `self._level` ersetzt werden.

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

Sie sollten `charlie_brown` in der Konsole ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

--fcc-editable-region--
    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--

```
