---
id: 68c7f7012a700243eff1cbc0
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Das Level, mit dem ein Mitarbeiter instanziiert wird, sollte aus bestimmten Levels ausgewählt werden. Sie verwenden `_base_salaries`, um das Level zu validieren und das richtige Gehalt für den Mitarbeiter festzulegen.

Erstellen Sie nach Ihrer bestehenden `if`-Anweisung eine weitere `if`, die überprüft, ob `level` nicht in `Employee._base_salaries` ist.

Innerhalb der neuen `if`-Anweisung lösen Sie eine `ValueError` mit der Nachricht `Invalid value '{level}' for 'level' attribute.` aus, wobei `{level}` durch den Wert des `level`-Arguments ersetzt werden sollte.

# --hints--

Wenn `level` nicht in `Employee._base_salaries` ist, sollten Sie eine `ValueError` mit der Nachricht `Invalid value '{level}' for 'level' attribute.` auslösen, wobei `{level}` durch den Wert des `level`-Arguments ersetzt werden sollte.

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

Sie sollten keine Ausnahme auslösen, wenn `level` in `Employee._base_salaries` ist.

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
