---
id: 68c9cab4b1118da59eecfc56
title: Schritt 31
challengeType: 20
dashedName: step-31
---

# --description--

Fügen Sie nach den vorhandenen `if`-Anweisungen eine weitere hinzu, um eine `ValueError` auszulösen, wenn `new_level` bereits das ausgewählte Level ist.

Beachten Sie, dass `_level` während der Initialisierung noch nicht existiert. Verwenden Sie daher `hasattr(self, '_level')`, um zu prüfen, ob es existiert, bevor Sie vergleichen. So vermeiden Sie eine `AttributeError`, wenn das Objekt zum ersten Mal erstellt wird.

Für die Nachricht verwenden Sie `'{level}' is already the selected level.`, wobei `{level}` durch das aktuelle Level ersetzt werden sollte.

# --hints--

Sie sollten eine dritte `if`-Anweisung innerhalb Ihres `level`-Setters haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

Wenn `new_level` gleich `self.level` ist, sollten Sie eine `ValueError` mit der Nachricht `'{level}' is already the selected level.` auslösen, wobei `{level}` durch das aktuelle Level ersetzt werden sollte.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "'trainee' is already the selected level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

Ihre `if`-Anweisung sollte `hasattr(self, '_level')` verwenden, um zu prüfen, ob `_level` existiert, bevor Sie vergleichen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2].find_conditions()[0].is_equivalent("hasattr(self, '_level') and new_level == self.level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
