---
id: 68c9a624fd54288f694ebdbf
title: Schritt 25
challengeType: 20
dashedName: step-25
---

# --description--

Sie können die Eigenschaft `name` gerne auf etwas anderes als einen String setzen, um den Validierungsprozess in Aktion zu sehen. Stellen Sie danach `charlie_brown`s `name` wieder her, indem Sie jede Codezeile entfernen, die seinen Namen ändert.

# --hints--

Sie sollten `charlie_brown`s `name` wiederherstellen.

```js
({ test: () => runPython(`
  assert charlie_brown.name == 'Charlie Brown'
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not isinstance(level, str):
            raise TypeError("'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
--fcc-editable-region--
charlie_brown.name = 'Snoopy'
--fcc-editable-region--
```
