---
id: 68ca758f8160b11757f877ae
title: Schritt 32
challengeType: 20
dashedName: step-32
---

# --description--

Erstellen Sie schließlich eine vierte `if`-Anweisung, die eine `ValueError` mit der Nachricht `Cannot change to lower level.` auslöst, wenn das Grundgehalt der neuen Stufe niedriger ist als das Grundgehalt der aktuellen Stufe.

Verwenden Sie `hasattr(self, '_level')`, um eine `AttributeError` während der Initialisierung zu vermeiden.

# --hints--

Sie sollten eine vierte `if`-Anweisung innerhalb Ihres `level`-Setters haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[3]`)) })
```

Ihre `if`-Anweisung sollte `hasattr(self, '_level')` verwenden, um zu prüfen, ob `_level` existiert, bevor Sie vergleichen.

```js
({ test: () => assert(runPython(`"hasattr(self, '_level')" in str(_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[3].find_conditions()[0])`)) })
```

Ihre `if`-Anweisung sollte `Employee._base_salaries[new_level]` mit `Employee._base_salaries[self.level]` vergleichen, um zu prüfen, ob die neue Stufe ein niedrigeres Grundgehalt hat.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[3].find_conditions()[0]
cond.is_equivalent("hasattr(self, '_level') and Employee._base_salaries[new_level] < Employee._base_salaries[self.level]") or cond.is_equivalent("hasattr(self, '_level') and Employee._base_salaries[self.level] > Employee._base_salaries[new_level]")
`)) })
```

Wenn `new_level` niedriger als `self.level` ist, sollten Sie eine `ValueError` mit der Nachricht `Cannot change to lower level.` auslösen.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'junior')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "Cannot change to lower level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

Sie sollten keine Ausnahme auslösen, wenn `new_level` höher als `self.level` ist.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'junior')
  try:
    emp.level = "senior"
  except Exception:
    assert False, "Expected not to raise exception with valid new_level"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
        if hasattr(self, '_level') and new_level == self.level:
            raise ValueError(f"'{self.level}' is already the selected level.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

```
