---
id: 69b014adbc624861e6ff3dd1
title: Schritt 21
challengeType: 20
dashedName: step-21
---

# --description--

Da der `name` Setter jetzt erstellt ist, aktualisieren Sie `__init__`, damit `self.name = name` anstelle von `self._name = name` verwendet wird. Dadurch wird die `name` Zuweisung über den Setter delegiert, und im nächsten Schritt wird eine Validierung hinzugefügt.

Passen Sie außerdem die `if` Bedingung so an, dass nur `level` geprüft wird, und aktualisieren Sie die `TypeError` Meldung entsprechend:

# --hints--

Ihre `__init__` Methode sollte `self.name = name` anstelle von `self._name = name` verwenden.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Ihre `__init__` Methode sollte nur `isinstance(level, str)` prüfen und die `TypeError` Meldung entsprechend aktualisieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_if("not isinstance(level, str)").find_body().has_stmt("raise TypeError(\\"'level' attribute must be of type 'str'.\\")")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
    
    def __init__(self, name, level):
--fcc-editable-region--
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
--fcc-editable-region--
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        self._name = new_name
        
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

```
