---
id: 69899123889dcf5b19be1847
title: Tekeleza algoriti ya utafutaji wa kina kwanza
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

Katika maabara hii, utatekeleza algoriti ya kutembea kwenye grafu inayoitwa <dfn>utafutaji wa kina kwanza</dfn> (DFS).

Wakati <dfn>utafutaji wa upana kwanza</dfn> (BFS) unachunguza majirani hatua kwa hatua, <dfn>utafutaji wa kina kwanza</dfn> unaingia kwa kina zaidi iwezekanavyo katika njia moja ya miunganisho kabla ya kurudi nyuma.

Mara algoriti inapofikia mwisho wa njia (nodi isiyo na majirani wasiochunguzwa), inarudi nyuma hadi nodi ya karibuni ambayo bado ina miunganisho isiyochunguzwa na kuendelea na utafutaji kutoka hapo.

Ili kutekeleza algoriti hii kwa njia ya mzunguko, utahitaji kutumia **stack**. Katika JavaScript, `Array` ya kawaida hutumika vizuri kama stack kwa kutumia njia za `.push()` na `.pop()`, ikifuata kanuni ya **Last-In-First-Out (LIFO)**. Hii inahakikisha kuwa jirani aliyoongezwa hivi karibuni kwenye stack ndiye anayechunguzwa kwanza.

Matokeo rahisi ya algoriti hii ni orodha ya nodi zinazoweza kufikiwa kutoka nodi fulani. Kwa hiyo, pia utahitaji kufuatilia nodi unazotembelea.

**Lengo**: Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kufafanua kitendakazi kinachoitwa `dfs`.
1. Kitendakazi cha `dfs` kinapaswa kuchukua hoja mbili:

    - `graph`: Matriki ya ujirani (safu ya safu za data).
    - `root`: Lebo ya nodi ya nambari (kielezo cha kuanzia) ambayo ni thamani ya nambari ya nodi kati ya `0` na `n - 1`, ambapo `n` ni jumla ya nodi katika grafu.

1. Kitendakazi kinapaswa kurudisha safu ya lebo zote za nodi zinazoweza kufikiwa kutoka nodi ya kuanzia.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `dfs` kinachochukua hoja mbili.

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` inapaswa kurudisha orodha yenye `1`, `2`, `3`, na `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` inapaswa kurudisha orodha yenye `1`, `2`, `3`, na `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` inapaswa kurudisha `[3]`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` inapaswa kurudisha orodha yenye `3` na `2`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` inapaswa kurudisha orodha yenye `0` na `1`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
