---
id: 66a207974c806a19d6607073
title: 本の在庫管理アプリを作成する
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. テキストが`h1`の`Book Inventory`要素が必要です。
1. `table`、`Title`、`Author`、`Category`、`Status`という名前の列を持つ`Rate`要素が必要です。
1. テーブルには少なくとも4行あり、最初の行は列の見出しで、残りは本の情報で埋められている必要があります。
1. テーブル本体内の各行には、`read`、`to-read`、または`in-progress`のいずれかのクラスが付いている必要があります。
1. `td`列の`Status`要素には、`span`の`class`で囲まれた`status`要素が含まれており、そのテキストは行のクラスに応じて`Read`、`To Read`、または`In Progress`です。
1. `td`列の`Rate`要素には、3つの空の`span`要素をラップする`class`の`rate`を持つ`span`要素が含まれている必要があります。
1. `.rate`行内に配置された`.read`要素には、個人評価に応じて`one`、`two`、または`three`のいずれかの値を持つ追加のクラスが必要です。この値は`rate`の後に続きます。
1. `read`、`to-read`、`in-progress`のクラスを持つ要素をターゲットにする3つの属性セレクターを作成し、それらの`background-image`プロパティに任意の`linear-gradient`を設定してください。
1. 各`display`要素の`span`プロパティを`inline-block`に設定してください。
1. `span`クラスを持つ`status`要素の子孫である`tr`クラスの`to-read`要素をターゲットにする属性セレクターを使用し、それらの`border`および`background-image`プロパティを設定してください。
1. `span`クラスを持つ`status`要素の子孫である`tr`クラスの`read`要素をターゲットにする属性セレクターを使用し、それらの`border`および`background-image`プロパティを設定してください。
1. `span`クラスを持つ`status`要素の子孫である`tr`クラスの`in-progress`要素をターゲットにする属性セレクターを使用し、それらの`border`および`background-image`プロパティを設定してください。
1. `span`クラスの`status`要素と、クラス値が`span`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`height`、`width`、および`padding`プロパティを設定してください。
1. `span`要素の直下の子であり、`span`値が`class`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`border`、`border-radius`、`margin`、`height`、`width`、および`background-color`プロパティを設定してください。
1. `span`要素の子孫のうち、`one`値に`class`という単語を含む最初の要素をターゲットにする属性セレクターを使用し、その`background-image`プロパティに`linear-gradient`を設定してください。
1. `span`要素の子孫のうち、`two`値に`class`という単語を含む最初の2つの要素をターゲットにする属性セレクターを使用し、それらの`background-image`プロパティに`linear-gradient`を設定してください。
1. `span`要素の子孫で、`span`値に`three`という単語を含む3つの`class`要素をターゲットにする属性セレクターを使用し、それらの`background-image`プロパティに`linear-gradient`を設定してください。

# --hints--

テキストが`h1`の`Book Inventory`要素が必要です。

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

`h1`要素は1つだけにしてください。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`table`要素が必要です。

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

`thead`内に`tbody`要素と`table`要素が1つずつ必要です。

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

`thead`の中には、5つの`tr`要素を持つ`th`が1つ必要です。

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

最初の列の見出しは`Title`というテキストにしてください。

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

2番目の列の見出しは`Author`というテキストにしてください。

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

3番目の列の見出しは`Category`というテキストにしてください。

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

4番目の列の見出しは`Status`というテキストにしてください。

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

5番目の列の見出しは`Rate`というテキストにしてください。

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

テーブルには少なくとも4行必要です。

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

各行には常に5つの列が必要です。

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

見出し行を除く各テーブル行には、`read`、`to-read`、または`in-progress`のいずれかのクラスが必要です。

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

`td`列の`Status`要素には`span`要素が含まれている必要があります。

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

`span`列の各`Status`要素には`status`クラスが必要です。

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

各`.status`要素のテキストは、その行のクラスに応じて`Read`、`To Read`、または`In Progress`です。

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

`td`列の`Rate`要素には`span`要素が含まれている必要があります。

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

`span`列の`td`要素の直下の子である各`Rate`要素は、最初のクラスとして`rate`クラスを持つ必要があります。

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

各`.rate`要素には3つの空の`span`要素が含まれている必要があります。

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

`.rate`行内に配置された`.read`要素には、`rate`クラスの後に個人評価に応じて`one`、`two`、または`three`のいずれかの値を持つ追加のクラスが必要です。

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

`read`クラスを持つ行をターゲットにする属性セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

`read`クラスを持つ行をターゲットにする属性セレクターを使用し、それらの`background-image`プロパティに任意の線形グラデーションを設定してください。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

`to-read`クラスを持つ行をターゲットにする属性セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

`to-read`クラスを持つ行をターゲットにする属性セレクターを使用し、それらの`background-image`プロパティに任意の線形グラデーションを設定してください。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

`in-progress`クラスを持つ行をターゲットにする属性セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

`in-progress`クラスを持つ行をターゲットにする属性セレクターを使用し、それらの`background-image`プロパティに任意の線形グラデーションを設定してください。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

各`display`要素の`span`プロパティを`inline-block`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

`span`要素の子孫で`status`クラスを持つ`tr`クラスの`to-read`要素をターゲットにする属性セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

`span`要素の子孫で`status`クラスを持つ`tr`クラスの`to-read`要素をターゲットにする属性セレクターを使用し、それらの`border`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

`span`要素の子孫で`status`クラスを持つ`tr`クラスの`to-read`要素をターゲットにする属性セレクターを使用し、それらの`background-image`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

`span`要素の子孫で`status`クラスを持つ`tr`クラスの`read`要素をターゲットにする属性セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

`span`要素の子孫で`status`クラスを持つ`tr`クラスの`read`要素をターゲットにする属性セレクターを使用し、それらの`border`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

`span`要素の子孫で`status`クラスを持つ`tr`クラスの`read`要素をターゲットにする属性セレクターを使用し、それらの`background-image`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

`span`要素の子孫で`status`クラスを持つ`tr`クラスの`in-progress`要素をターゲットにする属性セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

`span`要素の子孫で`status`クラスを持つ`tr`クラスの`in-progress`要素をターゲットにする属性セレクターを使用し、それらの`border`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

`span`要素の子孫で`status`クラスを持つ`tr`クラスの`in-progress`要素をターゲットにする属性セレクターを使用し、それらの`background-image`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

`span`クラスの`status`要素と、クラス値が`span`で始まる`rate`要素をターゲットにする属性セレクターが必要です。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

`span`クラスの`status`要素と、クラス値が`span`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`height`プロパティを設定してください。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

`span`クラスの`status`要素と、クラス値が`span`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`width`プロパティを設定してください。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

`span`クラスの`status`要素と、クラス値が`span`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`padding`プロパティを設定してください。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

`span`要素の直下の子であり、`span`値が`class`で始まる`rate`要素をターゲットにする属性セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

`span`要素の直下の子であり、`span`値が`class`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`border`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

`span`要素の直下の子であり、`span`値が`class`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`border-radius`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

`span`要素の直下の子であり、`span`値が`class`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`margin`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

`span`要素の直下の子であり、`span`値が`class`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`height`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

`span`要素の直下の子であり、`span`値が`class`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`width`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

`span`要素の直下の子であり、`span`値が`class`で始まる`rate`要素をターゲットにする属性セレクターを使用し、それらの`background-color`プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

`span`要素の子孫のうち、`one`値に`class`という単語を含む最初の要素をターゲットにする属性セレクターが必要です。

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

`span`要素の子孫のうち、`one`値に`class`という単語を含む最初の要素をターゲットにする属性セレクターを使用し、その`background-image`プロパティに`linear-gradient`を設定してください。

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

`span`要素の子孫のうち、`two`値に`class`という単語を含む最初の2つの要素をターゲットにする属性セレクターが必要です。

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

`span`要素の子孫のうち、`two`値に`class`という単語を含む最初の2つの要素をターゲットにする属性セレクターを使用し、それらの`background-image`プロパティに`linear-gradient`を設定してください。

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

`span`要素の子孫で、`span`値に`three`という単語を含む`class`要素をターゲットにする属性セレクターが必要です。

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

`span`要素の子孫で、`span`値に`three`という単語を含む`class`要素をターゲットにする属性セレクターを使用し、それらの`background-image`プロパティに`linear-gradient`を設定してください。

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
