---
id: 662a6bc12cde72c32fb526f0
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

<dfn>Інтерфейс</dfn> це як план для класу. Інтерфейс містить набір методів та властивостей, які має реалізувати клас.

Розпочніть цей проєкт з оголошення порожнього класу під назвою `Equation`. Ви будете використовувати цей клас для визначення інтерфейсу, плану загального рівняння.

# --hints--

Визначте новий клас під назвою `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
