---
id: 662bdd364bf2cde1487922a9
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

На відміну від інших мов програмування, Python не реалізує інтерфейси у своїй основній мові, але стандартна бібліотека Python дозволяє вам визначати інтерфейси простим способом.

Для цього проєкту ви будете використовувати утиліти з модуля `abc`. Тому імпортуйте цей модуль у код.

# --hints--

Імпортуйте модуль `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass

    def analyze(self):
        pass


class LinearEquation(Equation):
    pass


eq = Equation()
lin_eq = LinearEquation()
```
