---
id: 662bd552e1c1d2db1b88ba47
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Тепер визначте інший клас під назвою `LinearEquation` та зробіть його спадковим від `Equation`. Ви використовуватимете цей клас для представлення лінійних рівнянь.

# --hints--

Визначте клас під назвою `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

Клас `LinearEquation` має успадкувати клас `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
