---
id: 662bde88dc84f1e249801b1a
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` означає *Abstract Base Classes(абстрактні базові класи)*. Клас `ABC` дозволяє вам перетворити звичайний клас на абстрактний клас, який є класом, який діє як схема для конкретних класів.

Змініть інструкцію `import`, щоб імпортувати лише клас `ABC` із модуля `abc`. Ви можете імпортувати певний об’єкт `x` із модуля `y`, дотримуючись конструкції імпорту `from y import x`.

Потім перетворіть клас `Equation` на абстрактний клас, успадкувавши його від `ABC`.

# --hints--

Імпортуйте `ABC` з модуля `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

Клас `Equation` має успадкувати від `ABC`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass

    def analyze(self):
        pass

class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
