---
id: 662bd456896f16d9bd03f1a6
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

У класі `Equation` визначте два нові методи екземпляру під назвою `solve` та `analyze`.

# --hints--

Визначте метод під назвою `solve` в межах класу `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

Метод `solve` має приймати параметр `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

Визначте метод під назвою `analyze` в межах класу `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

Метод `analyze` має приймати параметр `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
