---
id: 662bd8260da84bdd5feae419
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Ви хочете, щоб клас `LinearEquation` реалізовував, а не просто успадковував усі методи, визначені всередині класу `Equation`, який має діяти як інтерфейс.

Наразі клас `Equation` є просто батьківським класом `LinearEquation`. У наступних кроках ви дізнаєтесь, як перетворити його на формальний інтерфейс.

Наразі створіть екземпляр `Equation` та призначте його до змінної `eq`, а також екземпляр `LinearEquation` та призначте його до змінної `lin_eq`.

# --hints--

Оголосіть змінну `eq` та призначте до неї екземпляр `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

Оголосіть змінну `lin_eq` та призначте до неї екземпляр `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass

    def analyze(self):
        pass


class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
