---
id: 662f96576ef178927de87975
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Щоб розпізнавати метод як абстрактний, його слід прикрашати декоратором `@abstractmethod`.

A <dfn>decorator</dfn> is used in Python to modify the behavior of a function. Here's an example of how to use a decorator named `spam`:

```py
@spam
def foo():
    pass
```

Змініть інструкцію імпорту, щоб імпортувати декоратор `abstractmethod` і прикрасити методи `solve` та `analyze` класу `Equation`. Це підніме два винятки.

Коли клас, що успадковує `ABC`, має абстрактний метод, екземпляр класу більше не можна створити. Тому видаліть екземпляр `Equation`, щоб позбутися помилки.

Інша помилка виникає через те, що клас `LinearEquation` має реалізувати усі абстрактні методи, визначені в інтерфейсі. Не забудьте також визначити їх у класі `LinearEquation`. Не можна використовувати декоратор `abstractmethod` у конкретному класі.

# --hints--

Імпортуйте `abstractmethod` з модуля `abc`.

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

Прикрасьте з `@abstractmethod` метод `solve` у класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

Прикрасьте з `@abstractmethod` метод `analyze` у класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

Визначте метод під назвою `solve` в межах класу `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

Метод `solve` має приймати параметр `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

Визначте метод під назвою `analyze` в межах класу `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

Метод `solve` має приймати параметр `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass

    def analyze(self):
        pass

class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
