---
id: 662fc3eba556a6bf800d48c1
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Вбудована функція `isinstance()` приймає два аргументи та повертає булеве значення, яке вказує, чи об’єкт, переданий як перший аргумент, є екземпляром класу, переданого як другий аргумент.

```py
isinstance(7, int) # True
```

Інша річ, яку ви хочете перевірити, це те, що кожен аргумент є числом. Після першого `if` створіть цикл `for`, який ітерує через `args` та перевіряє, чи аргумент на поточній ітерації не є екземпляром `int` або `float`. Використовуйте функцію `isinstance()` та передайте до неї кортеж, що містить `int` та `float` як другий аргумент.

Якщо аргумент не є числом, підійміть `TypeError` яка каже `"Coefficients must be of type 'int' or 'float'"`.

# --hints--

Створіть цикл `for`, який ітерує через `args` після інструкції `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

Створіть інструкцію `if`, яка перевіряє, чи поточний коефіцієнт не є екземпляром `int` або `float` у циклі `for`.

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

Використайте наданий рядок, щоб підняти `TypeError` в інструкції `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int

    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass

class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
```
