---
id: 662fa2e2cf27c09f21f4f5d0
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

У Python типи даних розпізнаються під час виконання (під час виконання коду). Тому вам не потрібно вказувати тип даних змінної, коли ви її оголошуєте. Тим не менш, ви можете анотувати змінну, щоб уточнити, що вона буде містити певний тип даних за допомогою `variable: <data type> = value` або просто `variable: <data type>`. Зауважте, що інтерпретатор Python не підтримує типи, які використовуються для анотування змінних, та зазвичай для цього вам потрібні зовнішні інструменти.

Усередині класу `Equation` визначте атрибут класу `degree`. Поки що не присвоюйте йому значення. Натомість використовуйте анотацію типу `int`, щоб показати, що вона зберігатиме ціле число всередині конкретних класів.

Пізніше ви використовуватимете цей атрибут класу як частину процесу перевірки аргументів, переданих для створення екземпляра об’єктів рівняння.

# --hints--

Визначте атрибут класу з іменем `degree` та позначте його з `int` у класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
