---
id: 662fbcef5f05e1b84f541a0c
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Кожен об’єкт рівняння створюватиметься шляхом передачі стількох аргументів, скільки коефіцієнтів у рівнянні, починаючи з n-го ступеня \\( x \\) до нульового ступеня, включаючи можливий коефіцієнт зі значенням `0`.

Наприклад, `LinearEquation(4, 5)` представлятиме рівняння \\( 4x + 5 = 0 \\), де `4` є коефіцієнтом першого (найвищого тут) ступеня та `5` коефіцієнт нульового ступеня.

Вам потрібно перевірити, чи передана правильна кількість аргументів для створення екземпляра об’єкта рівняння.

Усередині методу `__init__` створіть інструкцію `if`, щоб перевірити, чи довжина `args` відрізняється від кількості коефіцієнтів, які повинно мати рівняння (`degree + 1`). Якщо так, підійміть `TypeError` та використовуйте такий рядок, щоб надати спеціальне повідомлення: `f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`.

Потім виправте помилку, передавши `2` та `3` для створення екземпляра `lin_eq`.

# --hints--

Створіть інструкцію `if`, яка перевіряє, чи кількість коефіцієнтів, що використовуються для створення екземпляра рівняння, відрізняється від `degree + 1`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

Підійміть `TypeError` у нової інструкції `if` та використайте наданий рядок для повернення спеціального повідомлення про помилку.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

Передайте `2` та `3`, щоб створити екземпляр `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass

class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass

lin_eq = LinearEquation()
--fcc-editable-region--
```
