---
id: 6639f947d3a1818c9322c64a
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Останнім кроком валідації коефіцієнтів є перевірка того, що найвищий коефіцієнт степеня відрізняється від нуля. Пам’ятайте, що при створенні екземпляра об’єкта першим аргументом який має бути переданим це найвищий коефіцієнт ступеня.

Додайте інструкцію `if` для цього та підійміть `ValueError`, використовуючи такий рядок, щоб надати спеціальне повідомлення: `'Highest degree coefficient must be different from zero'`.

# --hints--

Створіть інструкцію `if`, яка перевіряє, чи перший коефіцієнт переданий для створення екземпляра рівняння є рівним нулю.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_conditions()[0]
cond.is_equivalent("args[0] == 0") or cond.is_equivalent("0 == args[0]") or cond.is_equivalent("not args[0]")
`)) })
```

Підійміть `ValueError` у нової інструкції `if` та використайте наданий рядок для повернення спеціального повідомлення про помилку.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_bodies()[0].has_stmt("raise ValueError('Highest degree coefficient must be different from zero')")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )        

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass

class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
```
