---
id: 663b83a28943e6aa6275a514
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Все ще у класі `Equation` визначте метод `__str__`, щоб надати належне рядкове представлення об’єктам рівняння, які ви збираєтеся створити.

Наразі в методі `__str__` оголосіть змінну `terms` та призначте до неї порожній список. Ви використовуватимете цю змінну, щоб зберігати кожен член (коефіцієнт помножений на \\( x^n \\)) рівняння.

Потім оголосіть змінну `equation_string`, призначте до неї результат об’єднання елементів у списку `terms` з допомогою пробілу. Нарешті, поверніть `equation_string`.

# --hints--

Визначте метод `__str__` в межах класу `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

Метод `__str__` має приймати параметр `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

Оголосіть змінну `terms` та призначте до неї порожній список в межах методу `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

Оголосіть змінну `equation_string` та призначте до неї результат об’єднання елементів у `terms` із пробілом в межах методу `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

Поверніть `equation_string` з методу `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int

    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--

--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass

class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
