---
id: 664e4a590b52ba8d2adff19f
title: Крок 30
challengeType: 20
dashedName: step-30
---

# --description--

Дискримінант квадратного рівняння у вигляді \\( ax^2 + bx + c = 0 \\), зазвичай позначається великою грецькою буквою дельта, дорівнює \\( Δ = b^2 - 4ac \\).

В межах класу `QuadraticEquation` визначте метод `__init__`. Використовуйте `super()`, щоб викликати метод `__init__` з батьківського класу. Потім визначте новий атрибут під назвою `delta`, який зберігає значення дискримінанта рівняння.

# --hints--

Визначте метод `__init__` в межах класу `QuadraticEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("__init__")`)) })
```

Метод `__init__` має приймати два параметри: `self` та `*args`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("__init__").has_args("self, *args")`)) })
```

Викличте `super().__init__(*args)` в межах методу `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("__init__").has_call("super().__init__(*args)")`)) })
```

Оголосіть атрибут `delta` у межах методу `__init__` та призначити до нього значення дискримінанта рівняння.

```js
({ test: () => runPython(`
eq = QuadraticEquation(2, -3, -4)
assert eq.delta == 41
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int

    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass

class LinearEquation(Equation):
    degree = 1

    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2
--fcc-editable-region--

--fcc-editable-region--    
    def solve(self):
        pass

    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
