---
id: 6675aaf418b41157f6ccd692
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Настав час повернутися до методу `__init__`. Залежно від типу рівняння вам потрібно буде передати змінну кількість аргументів під час створення екземпляра.

Додайте до методу другий параметр `args` та використовуйте оператор `*`, щоб змусити його приймати змінну кількість аргументів.

# --hints--

Метод `__init__` має приймати два параметри: `self` та `*args`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
