---
id: 6675a38a8b535e4ff3274520
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Вбудована функція `hasattr` приймає об’єкт як перший аргумент, та рядок, що представляє назву атрибута, як другий аргумент. Вона повертає булеве значення, яке вказує, чи має об’єкт вказаний атрибут.

Тепер ви збираєтеся використати метод `__init_subclass__`, щоб перевірити, чи дочірній клас має атрибут `degree` на момент створення екземпляра.

Створіть інструкцію `if`, щоб перевірити, чи `cls` не має атрибута `degree`. Якщо так, викликайте `AttributeError` та використовуйте рядок `f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"`, щоб надати спеціальне повідомлення.

Після цього виправте помилку, що з’явилася в терміналі, оголосивши атрибут класу `degree` всередині класу `LinearEquation`. Цей атрибут має представляти ступінь рівняння, який є показником старшого члена \\( x \\). Тому призначте ціле число `1` до атрибута `degree`.

# --hints--

Створіть інструкцію `if`, яка перевіряє, чи `cls` не має атрибута `degree` у методі `__init_subclass__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_conditions()[0].is_equivalent("not hasattr(cls, 'degree')")`)) })
```

Підійміть `AttributeError`, використовуючи наданий рядок в інструкції `if`.

```js
({ test: () => runPython(`
raise_stmt = 'raise AttributeError(f"Cannot create \\'{cls.__name__}\\' class: missing required attribute \\'degree\\'")'
node = _Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_bodies()[0]
assert node.has_stmt(raise_stmt) 
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int

    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        pass

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):

--fcc-editable-region--
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
