---
id: 66759e32b88fb5459b1e0234
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Метод `__init_subclass__` викликається кожного разу, коли клас, який його визначає, підкласується, та це дозволяє налаштовувати дочірні класи. Метод приймає параметр, названий відповідно до конвенції `cls` (що означає «клас»), який представляє новий дочірній клас.

Визначте метод `__init_subclass__` у класі `Equation` та надайте йому параметр `cls`.

# --hints--

Визначте метод `__init_subclass__` з параметром `cls` у класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int

    def __init__(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
