---
id: 68e76b1ce8216249d04df670
title: Crea una classe planet
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una classe chiamata `Planet`.
1. La classe `Planet` dovrebbe avere un metodo `__init__` che:
   - Ha quattro parametri: `self`, `name`, `planet_type` e `star`.
   - Genera un `TypeError` con il messaggio `name, planet type, and star must be strings` se uno qualsiasi degli argomenti passati non è di tipo stringa.
   - Genera un `ValueError` con il messaggio `name, planet_type, and star must be non-empty strings` se uno qualsiasi degli argomenti passati è una stringa vuota.
   - Assegna i valori passati agli attributi dell'istanza `name`, `planet_type` e `star`.
1. La classe `Planet` dovrebbe avere un metodo `orbit` che restituisce una stringa nel formato `{name} is orbiting around {star}...`.
1. La classe `Planet` dovrebbe avere un metodo `__str__` che restituisce una stringa nel formato `Planet: {name} | Type: {planet_type} | Star: {star}`.
1. Dovresti creare tre istanze della classe `Planet` chiamate `planet_1`, `planet_2` e `planet_3`.
1. Dovresti stampare ogni oggetto pianeta per vedere l'output del metodo `__str__`.
1. Dovresti chiamare il metodo `orbit` su ogni oggetto pianeta e stampare il risultato.

# --hints--

Dovresti creare una classe chiamata `Planet`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

La classe `Planet` dovrebbe avere un metodo `__init__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

Il metodo `__init__` dovrebbe avere quattro parametri: `self`, `name`, `planet_type` e `star`, in questo ordine.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

Il metodo `__init__` dovrebbe generare un `TypeError` con il messaggio `name, planet type, and star must be strings` se il primo argomento passato non è di tipo stringa.

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Il metodo `__init__` dovrebbe generare un `TypeError` con il messaggio `name, planet type, and star must be strings` se il secondo argomento passato non è di tipo stringa.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Il metodo `__init__` dovrebbe generare un `TypeError` con il messaggio `name, planet type, and star must be strings` se il terzo argomento passato non è di tipo stringa.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Il metodo `__init__` dovrebbe generare un `ValueError` con il messaggio `name, planet_type, and star must be non-empty strings` se il primo argomento passato è una stringa vuota.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Il metodo `__init__` dovrebbe generare un `ValueError` con il messaggio `name, planet_type, and star must be non-empty strings` se il secondo argomento passato è una stringa vuota.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Il metodo `__init__` dovrebbe generare un `ValueError` con il messaggio `name, planet_type, and star must be non-empty strings` se il terzo argomento passato è una stringa vuota.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Dovresti assegnare `name` a `self.name` nel metodo `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

Dovresti assegnare `planet_type` a `self.planet_type` nel metodo `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

Dovresti assegnare `star` a `self.star` nel metodo `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

La classe `Planet` dovrebbe avere un metodo `orbit`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

Il metodo `orbit` dovrebbe restituire una stringa nel formato `{name} is orbiting around {star}...`.

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

La classe `Planet` dovrebbe avere un metodo `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

Il metodo `__str__` dovrebbe restituire una stringa nel formato `Planet: {name} | Type: {planet_type} | Star: {star}`.

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

Dovresti creare tre istanze della classe `Planet` chiamate `planet_1`, `planet_2` e `planet_3`.

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

Dovresti stampare ogni oggetto pianeta per vedere l'output del metodo `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

Dovresti chiamare il metodo `orbit` su ogni oggetto pianeta e stampare il risultato.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
