---
id: 5d822fd413a79914d39e98cf
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

You can see the `body` (it's the inner-most box on your page); the box around it is the `html` element. Make your `body` fill the whole viewport by giving it a `height` of `100vh`. Remove the default `margin` from the `body` by setting the `margin` to `0`. Finally, set the `overflow` property to `hidden` to hide any scroll bars that appear when something extends past the viewport.

# --hints--

Unapaswa kutumia kichaguzi cha `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` yako inapaswa kuwa na `height` ya `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

`body` yako inapaswa kuwa na `margin` ya `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

`body` yako inapaswa kuwa na sifa ya `overflow` iliyowekwa kuwa `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

