---
id: 68c16924498728d259e4781c
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

現在創建一個 `Employee` 類別的實體，並傳入字串 `Charlie Brown` 和 `trainee`。將該實體指定給名為 `charlie_brown` 的變數。

# --hints--

你應該有一個名為 `charlie_brown` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

你應該將 `Employee('Charlie Brown', 'trainee')` 指定給 `charlie_brown`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level


--fcc-editable-region--
```
