---
id: 68c18b1d0fce081f85312d97
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

在先前的課程中，你學到以單一底線作為前置式的 `attribute` 是慣例上用於內部使用。

將 `name` 和 `level` 屬性修改為 `_name` 和 `_level`，因為這些不應該從它們的類別外部被修改。

請注意，這並不會阻止在類別外部存取或修改該屬性。此外，在 Python 中總是有方法可以存取私有屬性（以雙底線為前置式）。

# --hints--

你的 `__init__` 方法應該將 `self._name` 設定為 `name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

你的 `__init__` 方法應該將 `self._level` 設定為 `level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
