---
id: 68c15a2a1f0e70ca7154aba5
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將建置一個員工薪資追蹤系統。

首先創建一個名為 `Employee` 的類別。在其中創建一個帶有 `self`、`name` 和 `level` 參數的 `__init__` 方法。在 `__init__` 方法內，將 `name` 和 `level` 指定給同名的實體屬性。

# --hints--

你應該有一個名為 `Employee` 的類別。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

你的 `Employee` 類別應該有一個帶有三個參數 `self`、`name` 和 `level` 的 `__init__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

你的 `__init__` 方法應該將 `self.name` 設定為 `name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

你的 `__init__` 方法應該將 `self.level` 設定為 `level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
