---
id: 68c28c6517e7de5e85f8ebe5
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

既然你已經有 `_name` 和 `_level` 屬性的取值函式，請更新 `__str__` 傳回的字串，改為使用 `self.name` 和 `self.level`。這將呼叫取值函式，而非直接存取屬性。

# --hints--

你的 `__str__` 方法應該傳回 `f'{self.name}: {self.level}'`

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_return("f'{self.name}: {self.level}'")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)
print(charlie_brown.level)
```
