---
id: 68c28cfab52fef6913d18166
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

現在移除最後兩個 `print` 呼叫。

# --hints--

你的程式碼中不應該有 `print(charlie_brown.name)`。

```js
({ test: () => assert(runPython(`not _Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

你的程式碼中不應該有 `print(charlie_brown.level)`。

```js
({ test: () => assert(runPython(`not _Node(_code).has_stmt("print(charlie_brown.level)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
print(charlie_brown.name)
print(charlie_brown.level)
--fcc-editable-region--
```
