---
id: 68c27fa1697ceb7c578b2c9f
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

`@property` 裝飾器在 Python 中用於將方法轉換成屬性。它通常用來定義取值函式，這些取值函式是用來取得屬性值的方法。

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

創建一個名為 `name` 的方法，帶有 `self` 參數，並使用 `@property` 裝飾。在方法內部，傳回 `self._name`。
    
# --hints--

你的 `Employee` 類別應該有一個帶有 `self` 參數的 `name` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

你的 `name` 方法應該使用 `@property` 裝飾。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

你的 `name` 方法應該傳回 `self._name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
