---
id: 68c28c29026ad559f8ab4329
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

現在將 `charlie_brown.level` 列印到主控台。

# --hints--

你應該列印 `charlie_brown.level`。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.level)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

--fcc-editable-region--
```
