---
id: 68c28b7cc3ab8f4daf144b5e
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

依照你在前面步驟所做的，為 `_level` 屬性創建取值函式。

# --hints--

你的 `Employee` 類別應該有帶有 `self` 參數的 `level` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_args("self")`)) })
```

你的 `level` 方法應該使用 `@property` 裝飾。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_decorators("property")`)) })
```

你的 `level` 方法應該傳回 `self._level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_return("self._level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

```
