---
id: 68c28d659d345270d4a61084
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

`__repr__` 方法是一個特殊的方法，應該傳回一個可以用來實例化該物件的字串表述。

例如，`Employee('Charlie', 'developer')` 的 `__repr__` 方法應該傳回字串 `Employee('Charlie', 'developer')`，這是用來創建該物件的相同字串。

為你的 `Employee` 類別添加一個帶有 `self` 參數的 `__repr__` 方法，並使其傳回一個可用來實例化該物件的字串。

# --hints--

你的 `Employee` 類別應該有一個帶有 `self` 參數的 `__repr__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

你的 `__repr__` 方法應該傳回用來實例化該物件的字串。

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
