---
id: 68c27eb292bfec6b9402d58f
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

為你的 `Employee` 類別添加一個 `__str__` 方法。讓它傳回一個格式為 `name: level` 的 f-string，並將 `name` 和 `level` 替換為對應的屬性。

之後，將 `charlie_brown` 列印到主控台。

# --hints--

你的 `Employee` 類別應該有一個帶有 `self` 參數的 `__str__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

你的 `__str__` 方法應該傳回一個格式為 `name: level` 的 f-string，分別將 `name` 和 `level` 替換為 `self._name` 和 `self._level` 的值。

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

你應該將 `charlie_brown` 列印到主控台。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--
```
