---
id: 68c2d31a86d5672ca49d4521
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

現在是時候為 `__init__` 方法添加一些驗證。在該方法的開始，創建一個 `if` 述語來檢查 `name` 或 `level` 是否不是 `str` 的實體。

在 `if` 述語中，引發一個帶有訊息 `'name' and 'level' attribute must be of type 'str'.` 的 `TypeError`。

# --hints--

你應該在你的 `__init__` 方法中有一個 `if` 述語。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0]`)) })
```

你的 `if` 述語應該檢查 `name` 和 `level` 是否有任一不是 `str` 的實體。

```js
({ test: () => runPython(`
  cond = _Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_conditions()[0]
  conditions = [
    "not (isinstance(name, str) and isinstance(level, str))",
    "not isinstance(name, str) or not isinstance(level, str)",
    "not (isinstance(level, str) and isinstance(name, str))",
    "not isinstance(level, str) or not isinstance(name, str)",
  ]
  assert any(cond.is_equivalent(c) for c in conditions)
`) })
```

你的 `if` 述語應該引發一個帶有訊息 `'name' and 'level' attribute must be of type 'str'.` 的 `TypeError`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError("\\'name\\' and \\'level\\' attribute must be of type \\'str\\'.")')`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
--fcc-editable-region--
    def __init__(self, name, level):
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
