---
id: 68c281cecda94ba2c6838e17
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

既然你已為 `name` 定義了取值函式，你就可以透過 `name` 屬性存取 `_name` 屬性。所以將 `charlie_brown.name` 列印到主控台。 

# --hints--

你應該將 `charlie_brown.name` 列印到主控台。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
