---
id: 68c2d2414f5e111c02f4db5e
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

當你呼叫 `repr` 函式時，`__repr__` 方法會在背後被呼叫。要看到它的運作，請在你的程式碼底部列印呼叫 `repr(charlie_brown)` 的結果。

# --hints--

你應該列印 `repr(charlie_brown)`。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(repr(charlie_brown))")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

--fcc-editable-region--
```
