---
id: 68c7f7012a700243eff1cbc0
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

用來實例化員工的層次應該從特定的層次中選擇。你將使用 `_base_salaries` 來驗證層次並為員工設定正確的薪資。

在你現有的 `if` 述語之後，創建另一個 `if`，用來檢查 `level` 是否不在 `Employee._base_salaries` 中。

在新的 `if` 述語中，引發一個 `ValueError`，並帶有訊息 `Invalid value '{level}' for 'level' attribute.`，其中 `{level}` 應該被 `level` 引數的值替換。

# --hints--

當 `level` 不在 `Employee._base_salaries` 中時，你應該引發一個 `ValueError`，其訊息為 `Invalid value '{level}' for 'level' attribute.`，其中 `{level}` 應該被 `level` 引數的值所取代。

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

當 `level` 在 `Employee._base_salaries` 中時，你不應該引發任何例外。

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
