---
id: 68c8002beb157fe47d2e7699
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

設值函式是一種用來設定屬性值的行為，允許進行驗證檢查和限制。你可以使用 `@propertyName.setter` 裝飾器來創建設值函式，其中 `propertyName` 應該與要設定的屬性名稱相符：

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, value):
        self._name = value

p = Person('Alice')
p.name = 'Abigail' # Calls the setter
print(p.name) # Abigail
```

在你的取值函式之後，創建一個帶有參數 `self` 和 `new_name` 的 `name` 方法。使用 `@name.setter` 裝飾該方法。在方法內，將 `self._name` 設定為 `new_name`。

# --hints--

你的 `Employee` 類別應該有一個帶有參數 `self` 和 `new_name` 的 `name` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_args("self, new_name")`)) })
```

你的 `name` 方法應該使用 `@name.setter` 裝飾。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_decorators("name.setter")`)) })
```

你的 `name` 方法應該將 `self._name` 設定為 `new_name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_stmt("self._name = new_name")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary


charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
