---
id: 68c9c3aa714bc326e026b826
title: 步驟 26
challengeType: 20
dashedName: step-26
---

# --description--

新的層級必須先檢查是否為有效層級才能設定。在你的設值函式開始處，創建一個 `if` 述語，當 `new_level` 不是 `Employee._base_salaries` 的鍵時，引發 `ValueError`。

對於錯誤訊息，請使用 `Invalid value '{new_level}' for 'level' attribute.`，其中 `{new_level}` 應該被傳遞給設值函式的引數所取代。

# --hints--

你應該在你的 `level` 設值函式內有一個 `if` 述語。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[0]`)) })
```

當 `new_level` 不是 `Employee._base_salaries` 的鍵時，你應該引發一個 `ValueError`，其訊息為 `Invalid value '{new_level}' for 'level' attribute.`，其中 `{new_level}` 應該被傳遞給設值函式的引數所替換。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "dreamer"
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

當 `new_level` 是 `Employee._base_salaries` 的鍵時，你不應該引發任何例外。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  for i in Employee._base_salaries:
    try:
      emp.level = i
    except Exception:
      assert False, "Expected not to raise any exception with valid new_level"    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level
--fcc-editable-region--
    @level.setter
    def level(self, new_level):
        
        self._level = new_level
    
--fcc-editable-region--
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
