---
id: 68c9a72c9dab42a15ce6972b
title: 第 25 步
challengeType: 20
dashedName: step-25
---

# --description--

現在你將為 `_level` 屬性創建一個設值函式。創建一個名為 `level` 的方法，帶有參數 `self` 和 `new_level`。使用 `@level.setter` 裝飾該方法。

在該行為內，將 `self._level` 設定為 `new_level`。

# --hints--

你的 `Employee` 類別應該有一個帶有參數 `self` 和 `new_level` 的 `level` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_args("self, new_level")`)) })
```

你的 `level` 方法應該使用 `@level.setter` 裝飾。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_decorators("level.setter")`)) })
```

你的 `level` 方法應該將 `self._level` 設定為 `new_level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_stmt("self._level = new_level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")
--fcc-editable-region--
    @property
    def level(self):
        return self._level

    
    
--fcc-editable-region--
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
