---
id: 68c9cab4b1118da59eecfc56
title: 第 31 步
challengeType: 20
dashedName: step-31
---

# --description--

在現有的 `if` 述語之後，創建另一個來引發 `ValueError`，當 `new_level` 已經是選定的層級時。

請注意，`_level` 在初始化時尚不存在，因此在比較之前使用 `hasattr(self, '_level')` 來檢查它是否存在。這樣可以避免在物件剛建立時引發 `AttributeError`。

訊息部分，使用 `'{level}' is already the selected level.`，其中 `{level}` 應該被當前層級取代。

# --hints--

你應該在 `if` 設值函式內有第三個 `level` 述語。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

當 `new_level` 與 `self.level` 相等時，你應該引發一個帶有訊息 `ValueError` 的 `'{level}' is already the selected level.`，其中 `{level}` 應該被當前層級取代。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "'trainee' is already the selected level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

你的 `if` 述語應該使用 `hasattr(self, '_level')` 來檢查 `_level` 是否存在，然後再進行比較。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2].find_conditions()[0].is_equivalent("hasattr(self, '_level') and new_level == self.level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
