---
id: 68c9a09163d4a42e9bc9b638
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

現在你有一個用於 `_name` 屬性的設值函式，該函式公開了 `name` 屬性。

所以嘗試像平常對任何屬性一樣更新 `charlie_brown` 的 `name` 為不同的名稱。這將在底層呼叫設值函式。

# --hints--

你應該將 `charlie_brown` 的 `name` 更新為不同的名稱。

```js
({ test: () => runPython(`
  assert _Node(_code).has_variable("charlie_brown.name")
  assert charlie_brown.name != 'Charlie Brown'
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
--fcc-editable-region--

--fcc-editable-region--
```
