---
id: 68caacb0f4311cc9be9a2132
title: 步驟 32
challengeType: 20
dashedName: step-32
---

# --description--

現在你將專注於編寫 `salary` 設值函式。在 `salary` 取值函式之後，為 `salary` 屬性創建一個簡單的設值函式，將 `self._salary` 設定為傳入該方法的引數值。之後，列印 `Salary updated to $`，接著是新的薪資，最後加上句號。

你將在接下來的幾個步驟中負責驗證新的薪資。

# --hints--

你的 `Employee` 類別應該有一個使用 `@salary.setter` 裝飾的 `salary` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("salary")[1].has_decorators("salary.setter")`)) })
```

你的 `salary` 設值函式應該將 `self._salary` 設定為傳入的引數值。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  built_in_print = print
  print = lambda x: x
  emp.salary += 100
  print = built_in_print
  assert emp.salary == Employee._base_salaries['trainee'] + 100
`) })
```

你的 `salary` 設值函式應該列印 `Salary updated to $`，後面接著新的薪資和句號。

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    emp = Employee('Frank', 'trainee')
    emp.salary += 100
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert f"Salary updated to \${Employee._base_salaries['trainee'] + 100}" in output
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
        if new_level == self.level:
            raise ValueError(f"'{self.level}' is already the selected level.")
        if Employee._base_salaries[new_level] < Employee._base_salaries[self.level]:
            raise ValueError(f"Cannot change to lower level.")
        print(f"'{self.name}' promoted to '{new_level}'.")
        self._salary = Employee._base_salaries[new_level]
        self._level = new_level
--fcc-editable-region--
    @property
    def salary(self):
        return self._salary

    
--fcc-editable-region--

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

charlie_brown.level = 'junior'

```
