---
id: 69b014adbc624861e6ff3dd1
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

現在 `name` 設值函式已經建立，請更新 `__init__`，改用 `self.name = name` 取代 `self._name = name`。這樣會透過設值函式委派 `name` 指派，並且在下一步會加入驗證。

同時，調整 `if` 條件只檢查 `level`，並相應更新 `TypeError` 訊息：

# --hints--

你的 `__init__` 方法應該使用 `self.name = name` 取代 `self._name = name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

你的 `__init__` 方法應該只檢查 `isinstance(level, str)`，並相應更新 `TypeError` 訊息。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_if("not isinstance(level, str)").find_body().has_stmt("raise TypeError(\\"'level' attribute must be of type 'str'.\\")")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
    def __init__(self, name, level):
--fcc-editable-region--
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]
--fcc-editable-region--

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        self._name = new_name
        
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary


charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
