---
id: 68de7360e1832ea180577a8a
title: N번째 피보나치 수 계산기 만들기
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과시켜 실습을 완료하세요.

**사용자 스토리:**

1. `fibonacci`라는 이름의 함수를 만들어야 합니다.
2. `sequence` 함수 내에 `fibonacci`라는 리스트를 정의하고, `[0, 1]` 값들로 초기화해야 합니다.
3. `fibonacci` 함수는 하나의 매개변수, 양의 정수 `n`를 받아야 합니다.  
4. `fibonacci(n)`를 호출하면 동적 프로그래밍 방식을 사용해 피보나치 수열에서 각 숫자가 바로 앞 두 숫자의 합인 `n`번째 수를 계산해 반환해야 합니다.
5. 피보나치 수열에서 `n` 위치에 있는 각 계산된 숫자는 `sequence` 리스트의 `n - 1` 인덱스에 저장해야 합니다.

# --hints--

`fibonacci`라는 이름의 함수가 있어야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

`fibonacci` 함수는 하나의 매개변수를 받아야 합니다.

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

`sequence` 함수 내에 `fibonacci`라는 리스트가 `[0, 1]`로 초기화되어 있어야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

`fibonacci(0)`는 `0`를 반환해야 합니다.

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

`fibonacci(1)`는 `1`를 반환해야 합니다.

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

`fibonacci(2)`는 `1`를 반환해야 합니다.

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

`fibonacci(3)`는 `2`를 반환해야 합니다.

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

`fibonacci(5)`는 `5`를 반환해야 합니다.

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

`fibonacci(10)`는 `55`를 반환해야 합니다.

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

`fibonacci(15)`는 `610`를 반환해야 합니다.

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

코드에서 재귀를 사용하지 않아야 합니다.

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```
