---
id: 5f3313e74582ad9d063e3a38
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

在 `html` 元素中添加一个 `head` 元素，接下来你可以添加 `title` 元素。 `title` 元素的文本应该是 `Cafe Menu`。

# --hints--

你应该有一个开启的 `<head>` 标签。

```js
assert(code.match(/<head>/i));
```

你应该有一个闭合的 `</head>` 标签。

```js
assert(code.match(/<head>/i));
```

你应该有一个开头的 `<title>` 标签。

```js
assert(code.match(/<title>/i));
```

你应该有一个闭合的 `</title>` 标签。

```js
assert(code.match(/<\/title>/i));
```

所有的 `title` 元素都应嵌套在 `head` 元素中。

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

你的 `title` 元素内应该有文本 `Cafe Menu`。 你可能需要检查一下拼写。

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
