---
id: 62a3a0a3c0a4b32915d26a6e
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Rimuovi la riga `console.log("Hello World");` per iniziare a scrivere il codice del progetto.

Nel precedente progetto, hai imparato come lavorare con le variabili e la parola chiave `let` in questo modo:

```js
let camperbot;
```

Hai imparato anche come inizializzare una variabile con un valore come questo:

```js
let age = 32;
```

Utilizza la parola chiave `let` per dichiarare una variabile chiamata `xp` e assegnagli il valore di `0`.

# --hints--

You should not have a `console.log("Hello World");` line in your code.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Dovresti usare la parola chiave `let` per dichiarare la variabile.

```js
assert.match(code, /let/);
```

You should declare a variable named `xp`.

```js
assert.match(code, /let\s+xp/);
```

`xp` dovrebbe avere il valore di `0`.

```js
assert.equal(xp, 0);
```

Dovresti inizializzare la variabile `xp` a `0`.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
