---
id: 62a3b41c9494f937560640ab
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Dichiara altre due variabili chiamate `monsterHealth` e `inventory`.

Per la variabile `inventory`, assegnagli il valore di un array contenente la stringa `"stick"`.

Ricorda che hai lavorato con gli array nel progetto precedente in questo modo:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

You should use `let` to declare your `monsterHealth` variable.

```js
assert.match(code, /let\s+monsterHealth/i);
```

You should use camelCase to name your `monsterHealth` variable.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` should not have a value.

```js
assert.isUndefined(monsterHealth);
```

You should not assign a value to your `monsterHealth` variable.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

You should use `let` to declare your `inventory` variable.

```js
assert.match(code, /let\s+inventory/i);
```

La variabile `inventory` dovrebbe essere un array.

```js
assert.isArray(inventory);
```

La variabile `inventory` dovrebbe contenere soltanto un valore.

```js
assert.lengthOf(inventory, 1);
```

La variabile `inventory` dovrebbe contenere la stringa `"stick"`.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
