---
id: 6716249b5405164036fd0b0d
title: بناء أداة تصور الفرز
challengeType: 25
dashedName: build-a-sorting-visualizer
demoType: onClick
---

# --description--

تُرتب خوارزمية الفرز الفقاعي متتالية من الأعداد الصحيحة بمقارنة أزواج من العناصر المجاورة بدءًا من بداية المتتالية. إذا كان العنصر الأول أكبر من الثاني، تقوم بتبديلهما. ثم تنتقل إلى الزوج التالي. عند الوصول إلى العنصر الأخير في المتتالية، تبدأ دورة جديدة من البداية وتكرر العملية حتى تُرتب العناصر. تتوقف الخوارزمية بعد إكمال دورة واحدة دون أي تبديلات.

في هذا المختبر، تم تزويدك بكل ملفات HTML وCSS. ستستخدم JavaScript لإكمال أداة تصور الفرز الفقاعي بحيث تعرض كل خطوة يحتاجها الفرز الفقاعي لترتيب مصفوفة من خمسة أعداد صحيحة.

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة باسم `generateElement` تُرجع عددًا صحيحًا عشوائيًا بين `1` و `100`، شاملاً الحدين.
1. يجب أن تملك دالة باسم `generateArray` تستخدم الدالة `generateElement` لإرجاع مصفوفة تحتوي على خمسة أعداد صحيحة عشوائية.
1. يجب أن تملك دالة باسم `generateContainer` تنشئ وتُرجع عنصر `div` فارغ.
1. يجب أن تملك دالة باسم `fillArrContainer` تأخذ عنصر HTML كمعلمة أولى ومصفوفة كمعلمة ثانية.
1. يجب أن تأخذ دالة `fillArrContainer` عنصرًا كمعامل أول ومصفوفة أعداد صحيحة كمعامل ثاني، ثم تملأ العنصر بخمسة عناصر `span`، يعرض كل `span` عددًا من المصفوفة.
1. يجب أن تملك دالة باسم `isOrdered` تأخذ عددين صحيحين وتُرجع قيمة منطقية تشير إلى ما إذا كان العدد الأول أقل من أو يساوي العدد الثاني.
1. يجب أن تملك دالة باسم `swapElements` تأخذ مصفوفة أعداد صحيحة وفهرسًا رقميًا.
1. يجب أن تعدل دالة `swapElements` المصفوفة في مكانها عن طريق تبديل العنصر عند الفهرس الممرر والعنصر الذي يليه إذا أعادت دالة `isOrdered` القيمة `false`.
1. يجب أن تملك دالة باسم `highlightCurrentEls` تأخذ عنصر HTML وفهرسًا رقميًا.
1. يجب أن تضبط دالة `highlightCurrentEls` حدود (`border`) العنصر الفرعي في الفهرس المعطى والعنصر الفرعي الذي يليه في العنصر المحدد لتكون بنمط `dashed`، ولون `red`، وعرض تختاره.
1. عند نقر `#generate-btn` يجب استخدام دالة `fillArrContainer` لملء `#starting-array` بخمسة عناصر `span`، كل منها يعرض عددًا عشوائيًا كنص. إذا وُجدت عناصر أخرى، يجب إزالتها من `#array-container`.
1. يجب تنفيذ خوارزمية Bubble Sort بحيث بعد نقر `#sort-btn` يحتوي `#array-container` على عنصر `div` لكل خطوة مطلوبة من الخوارزمية لفرز المصفوفة الابتدائية، بما في ذلك `div` يمثل المصفوفة الابتدائية و`div` يمثل المصفوفة المرتبة. الدوال التي أنشأتها حتى الآن يمكن أن تكون مفيدة هنا.
1. يجب أن يحتوي كل `div` على خمسة عناصر `span` تمثل المصفوفة في حالتها الحالية أثناء الفرز.
1. بعد نقر `#sort-btn` يجب أن يمثل `#starting-array` الخطوة الابتدائية مع المصفوفة الأصلية وأول عددين صحيحين مميزين.
1. في كل خطوة من خطوات الفرز، يجب استخدام `highlightCurrentEls` لتسليط الضوء على العددين اللذين تتم مقارنتهما، وتبديلهما في الخطوة التالية باستخدام `swapElements`.

# --hints--

يجب أن يكون لديك دالة باسم `generateElement`.

```js
assert.isFunction(generateElement);
```

يجب أن تُرجع دالة `generateElement` عددًا صحيحًا عشوائيًا بين `1` و`100` شاملاً.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    assert.strictEqual(generateElement(), 26);
    assert.strictEqual(generateElement(), 9);
} finally {
    randomMocker.restore();
}
```

يجب أن يكون لديك دالة باسم `generateArray`.

```js
assert.isFunction(generateArray)
```

يجب أن تستخدم دالة `generateArray` دالة `generateElement`.

```js
assert.match(generateArray.toString(), /\bgenerateElement\b/);
```

يجب أن تُرجع دالة `generateArray` مصفوفة تحتوي على خمسة أعداد صحيحة عشوائية بين `1` و`100`.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    assert.deepEqual(generateArray(), [26, 9, 58, 23, 38])
} finally {
    randomMocker.restore();
}
```

يجب أن يكون لديك دالة باسم `generateContainer`.

```js
assert.isFunction(generateContainer);
```

يجب أن تُرجع دالة `generateContainer` عنصر `div` فارغ.

```js
const div = generateContainer();
assert.equal(div.tagName, "DIV");
assert.isEmpty(div.children);
```

يجب أن يكون لديك دالة باسم `fillArrContainer`.

```js
assert.isFunction(fillArrContainer);
```

يجب أن تأخذ دالة `fillArrContainer()` عنصرًا كمعامل أول ومصفوفة من الأعداد الصحيحة كمعامل ثاني، ثم تملأ العنصر بخمسة عناصر `span`، بحيث يعرض كل `span` أحد الأعداد الصحيحة من المصفوفة.

```js
const testDiv = document.createElement("div");
const testArr = [15, 98, 17, 5, 63]
fillArrContainer(testDiv, testArr);
const children = testDiv.children;
assert.lengthOf(children, 5);
Array.from(children).forEach((el, i) => {
    assert.equal(el.tagName, "SPAN");
    assert.equal(el.innerText.trim(), testArr[i])
})
```

يجب أن يكون لديك دالة باسم `isOrdered`.

```js
assert.isFunction(isOrdered);
```

يجب أن تأخذ دالة `isOrdered` عددين صحيحين وتُرجع قيمة منطقية تشير إلى ما إذا كان العدد الأول أقل من أو يساوي الثاني.

```js
assert.lengthOf(isOrdered, 2);
assert.isTrue(isOrdered(2, 60));
assert.isFalse(isOrdered(10, 3));
assert.isTrue(isOrdered(5, 5));
```

يجب أن يكون لديك دالة باسم `swapElements`.

```js
assert.isFunction(swapElements);
```

يجب أن تأخذ دالة `swapElements` مصفوفة من الأعداد الصحيحة وفهرسًا رقميًا كمعاملين. يجب أن تعدل المصفوفة في مكانها بتبديل العنصر عند الفهرس المعطى والعنصر الذي يليه إذا كان العنصر الأول أكبر من الثاني.

```js
const testArr = [22, 4, 87, 47, 33];
swapElements(testArr, 0);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 1);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 2);
assert.deepEqual(testArr, [4, 22, 47, 87, 33]);
swapElements(testArr, 3);
assert.deepEqual(testArr, [4, 22, 47, 33, 87]);
```

يجب أن يكون لديك دالة باسم `highlightCurrentEls`.

```js
assert.isFunction(highlightCurrentEls);
```

يجب أن تعطي دالة `highlightCurrentEls` العناصر التابعة للعنصر المحدد، الواقعة عند الفهرس المعطى والفهرس التالي، حدودًا بنمط `dashed`، ولون `red`، وعرض تختاره.

```js
const testDiv = document.createElement("div");
document.querySelector("body").appendChild(testDiv)
for (let i = 0; i < 5; i++) {
    testDiv.appendChild(document.createElement("span"));
}
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|#F00|hsl\(0,\s*100%,\s*50%\))/;
const revertBorder = () => {
    for (const el of children) {
        el.style.border = "revert";
    }
}
const children = testDiv.children;

for (let i = 0; i < 3; i++) {
    highlightCurrentEls(testDiv, i);
    for (let j = 0; j < 5; j++) {
        let b = getComputedStyle(children[j]).border;
        if (j == i || j == i + 1) {            
            assert.match(b, redBorderRegex);
            assert.isAbove(parseFloat(b), 0);
        } else {
            assert.notMatch(b, redBorderRegex);
        }
    }
    revertBorder();
}
testDiv.remove();
```

عند النقر على `#generate-btn` يجب ملء `#starting-array` بخمسة عناصر `span`، كل منها يعرض رقمًا عشوائيًا بين `1` و`100` كنص.

```js
const genBtn = document.querySelector("#generate-btn");
genBtn.click();
const children = document.querySelector("#starting-array").querySelectorAll("span");
assert.lengthOf(children, 5);
Array.from(children).forEach(el => {
    assert.equal(el.tagName, "SPAN");
    const num = Number(el.innerText.trim());
    assert.isAtMost(num, 100);
    assert.isAtLeast(num, 1);
})
```

عندما يحتوي `#starting-array` بالفعل على مصفوفة مولدة، أو يحتوي `#array-container` على المصفوفة المرتبة، يجب أن يؤدي النقر على `#generate-btn` إلى إزالة العناصر الأخرى في `#array-container`، مع ترك `#starting-array` فقط بالأرقام الجديدة المولدة.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
genBtn.dispatchEvent(new Event("click"));

const prevNumbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));

sortBtn.dispatchEvent(new Event("click"));
genBtn.dispatchEvent(new Event("click"));

const container = document.querySelector("#array-container");
assert.lengthOf(container.children, 1);
const numbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));
assert.lengthOf(numbers, 5);
assert.isTrue(prevNumbers.some((num, index) => num !== numbers[index]))
```

بعد النقر على `#sort-btn`، يجب أن يحتوي `#array-container` على عدد من عناصر `div` يساوي عدد الخطوات المطلوبة من خوارزمية الفرز الفقاعي لفرز المصفوفة الابتدائية، بما في ذلك `div` يمثل المصفوفة الابتدائية و`div` يمثل المصفوفة المرتبة.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
// using randomMocker to be sure that the starting array requires 13 steps to be sorted
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.lengthOf(container.children, 13)
    Array.from(container.children).forEach(el => {assert.equal(el.tagName, "DIV")})
} finally {
    randomMocker.restore();
}
```

بعد النقر على `#sort-btn`، يجب أن يحتوي كل `div` داخل `#array-container` على خمسة عناصر `span`، كل منها يعرض رقمًا كنص، ومرتبة لتمثيل الخطوات المطلوبة من خوارزمية الفرز الفقاعي لفرز المصفوفة الابتدائية.

```js
const finalArr = [9, 23, 26, 38, 58];
const arrays = [
    [26, 9, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 23, 58, 38],
    [9, 26, 23, 38, 58],
    [9, 26, 23, 38, 58],
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr
]
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.isNotEmpty(container.children);
    Array.from(container.children).forEach((el, i) => {
        Array.from(el.children).forEach((j, k) => {
            assert.strictEqual(Number(j.innerText.trim()), arrays[i][k])
        })
    })
} finally {
    randomMocker.restore();
}
```

عند النقر على `#sort-btn`، يجب استخدام دالة `highlightCurrentEls` لتسليط الضوء على العناصر التي تتم مقارنتها في كل خطوة.

```js
const expectedHighlights = [
    [0, 1], [1, 2], [2, 3], [3, 4],
    [0, 1], [1, 2], [2, 3], [3, 4],
    [0, 1], [1, 2], [2, 3], [3, 4],
    null
];
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try{
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));

    const container = document.getElementById("array-container");
    const divs = Array.from(container.children);
    const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|hsl\(0,\s*100%,\s*50%\))/;

    divs.forEach((div, i) => {
        const children = Array.from(div.children);
        if(expectedHighlights[i]){
            const [index1, index2] = expectedHighlights[i];
            assert.match(getComputedStyle(children[index1]).border, redBorderRegex);
            assert.match(getComputedStyle(children[index2]).border, redBorderRegex);
        }
    })
} finally {
    randomMocker.restore();
}
```

بعد النقر على `#sort-btn`، يجب أن يمثل `#starting-array` الخطوة الابتدائية مع المصفوفة الأولية وأول عددين صحيحين مميزين باستخدام `highlightCurrentEls`.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");   
genBtn.dispatchEvent(new Event("click"));
sortBtn.dispatchEvent(new Event("click"));
const firstContainer = document.querySelector("#starting-array");
const children = firstContainer.children
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|hsl\(0,\s*100%,\s*50%\))/;
assert.match(getComputedStyle(children[0]).border, redBorderRegex);
assert.match(getComputedStyle(children[1]).border, redBorderRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

#sort-btn {
   display: none 
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js
const arrayContainer = document.getElementById("array-container");
const generateArrayBtn = document.getElementById("generate-btn");
const sortArrayBtn = document.getElementById("sort-btn");
const arrayLength = 5;
const minVal = 1;
const maxVal = 100;
const isStart = () => arrayContainer.children.length === 1;
const clearArrayContainer = () => { arrayContainer.innerHTML = '<div id="starting-array"></div>' };
const showSortBtn = () => {sortArrayBtn.style.display = "inline-block"};
const hideSortBtn = () => {sortArrayBtn.style.display = "none"};
const generateElement = () => Math.floor(Math.random() * maxVal + minVal);
const generateArray = () => Array.from({ length: arrayLength }, generateElement);
const fillArrContainer = (container, arr) => {
    container.innerHTML = "";
    arr.forEach(i => {
        const el = document.createElement("span");
        el.innerText = i;
        el.id = `number-${i}`
        container.appendChild(el);
    })
}
const isOrdered = (el1, el2) => el1 <= el2;

const swapElements = (arr, n = 0) => {
    if (n < arr.length - 1 && !isOrdered(arr[n], arr[n + 1])) {
        const temp = arr[n];
        arr[n] = arr[n + 1];
        arr[n + 1] = temp;
    }
}
const getLastChildren = () => arrayContainer.lastElementChild;
const getLastArr = () => {
    const els = Array.from(getLastChildren().children);
    const arr = els.map(el => Number(el.id.replace("number-", "")));
    return arr;
}
const generateContainer = () => {
    const container = document.createElement("div");
    arrayContainer.appendChild(container);
    return container;
}
const highlightCurrentEls = (container, n = 0) => {
    const children = container.children;
    children[n].style.border = "2px dashed red";
    children[n + 1].style.border = "2px dashed red";
}

const highlightSorted = () => {
    getLastChildren().style.border = "4px solid green";
}

const bubbleSort = () => {
    let swapped = true;
    while (swapped) {
        const startingArr = getLastArr();
        startingArr.forEach((_, i) => {
            if (i + 1 < arrayLength) {
                highlightCurrentEls(getLastChildren(), i);
                const arr = getLastArr();
                swapElements(arr, i);
                fillArrContainer(generateContainer(), arr);
            }
        })
        const lastArr = getLastArr()
        if (startingArr.every((el, i) => el === lastArr[i])) swapped = false;
    }
}


generateArrayBtn.addEventListener("click", () => {
    if (!isStart()) {
        clearArrayContainer();
    }
    fillArrContainer(document.getElementById("starting-array"), generateArray());
    showSortBtn();
});

sortArrayBtn.addEventListener("click", () => {
    bubbleSort();
    highlightSorted();
    hideSortBtn();
})
```
