---
id: 66a9ca099e3ecb1be10a2696
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Agora é hora de adicionar mais dois estudantes à tabela.

Seguindo o mesmo padrão dos passos anteriores, adicione uma quarta linha da coluna de estudante. Use os seguintes dados para os elementos da tabela de dados:

- Sobrenome: `Thompson`
- Nome: `Jane`
- Nota: `77`

Para a quinta linha da tabela de estudante, use os seguintes dados para os elementos da tabela de dados:

- Sobrenome: `Williams`
- Nome: `Natalie`
- Nota: `83`

# --hints--

Você deve ter uma quarta `table row` no elemento `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

Você deve ter um elemento `td` com o texto `Thompson` dentro do quarto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

Você deve ter um elemento `td` com o texto `Jane` dentro do quarto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

Você deve ter um elemento `td` com o texto `77` dentro do quarto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

Você deve ter um quinto `tr` dentro do seu elemento `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

Você deve ter um elemento `td` com o texto `Williams` dentro do quinto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

Você deve ter um elemento `td` com o texto `Natalie` dentro do quinto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

Você deve ter um elemento `td` com o texto `83` dentro do quinto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
