---
id: 66a9bb8578459a175432e7d0
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Para a primeira seção da tabela, agrupe o conteúdo do cabeçalho que representa os rótulos das colunas para o nome, sobrenome e nota final do exame do estudante.

O elemento <dfn>table head</dfn>, `thead`, é usado para agrupar o conteúdo do cabeçalho em uma tabela. 

Aqui está um exemplo usando o elemento `thead`:

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

Abaixo do seu elemento `caption`, adicione um elemento de cabeçalho de tabela.

# --hints--

Você deve ter uma tag de abertura para `thead`.

```js
assert.match(code, /<thead>/i);
```

Você deve ter uma tag de fechamento para `thead`.

```js
assert.match(code, /<\/thead>/i);
```

O elemento `thead` deve estar dentro do elemento `table`.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
