---
id: 66a9bcf00f13a418368a272e
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

O elemento head da tabela consiste em um elemento <dfn>table row</dfn>, `tr`, que contém os elementos <dfn>table header cell</dfn>, `th`.

Aqui está um exemplo usando os elementos `tr` e `th` para uma tabela de esportes:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

Dentro do elemento `thead`, adicione um elemento `tr`.

Dentro do elemento `tr`, adicione três elementos `th`. 

O primeiro elemento `th` deve conter o texto `Last Name` (Sobrenome).O segundo elemento `th` deve conter o texto `First Name` (Nome).O terceiro elemento `th` deve conter o texto `Grade` (Nota).

# --hints--

Você deve ter uma tag de abertura de `tr`.

```js
assert.match(code, /<tr>/i);
```

Você deve ter uma tag de fechamento de `tr`.

```js
assert.match(code, /<\/tr>/i);
```

O elemento `tr` deve estar dentro do elemento `thead`. 

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

O primeiro elemento `th` deve ter o texto `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

O segundo elemento `th` deve ter o texto `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

O terceiro elemento `th` deve ter o texto `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

Todos os elementos `th` devem estar dentro do elemento `tr`.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      --fcc-editable-region--
      <thead>     
        
      </thead>
      --fcc-editable-region--
    </table>
  </body>
</html>
```
