---
id: 66a9b48a211a73155621d0df
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Neste workshop, você praticará trabalhar com tabelas HTML construindo uma tabela de exame final para um grupo de estudantes. 

Nas lições anteriores, você aprendeu como trabalhar com o elemento `table` para representar dados tabulares.

Dentro do elemento `body`, coloque um elemento `table`.

# --hints--

Você deve ter uma tag de abertura de `table`.

```js
assert.match(code, /<table>/i);
```

Você deve ter uma tag de fechamento de `table`.

```js
assert.match(code, /<\/table>/i);
```

O elemento `table` deve estar dentro do elemento `body`.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
