---
id: 66a9b333487b9c14998539a5
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro do elemento `head`, aninhe um elemento `meta` com o atributo `charset` definido como `"UTF-8"`.

Abaixo desse elemento `meta`, adicione um elemento `title`. 

O texto do elemento `title` deve ser `Calculus Final Exams Table`.

Após completar o elemento `head`, você pode adicionar o elemento `body`.

# --hints--

Você deve ter um elemento `meta`.

```js
assert.isNotNull(document.querySelector("meta"));
```

O elemento `meta` é um elemento vazio, não deve ter uma tag de fechamento `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

A tag `meta` deve ter um atributo `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

O atributo `charset` deve ter um valor de `"UTF-8"`.

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

O elemento `meta` deve estar aninhado dentro do elemento `head`.

```js
const meta = document.querySelector('head > meta');
assert.strictEqual(meta?.parentElement?.tagName, 'HEAD');
```

Você deve acrescentar uma tag de abertura`title`.

```js
assert.match(code, /<title>/i);
```

Você deve acrescentar uma tag de fechamento para `title`.

```js
assert.match(code, /<\/title>/i);
```

O elemento `title` deve estar dentro do elemento `head`.

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

Seu elemento `title` deve conter o texto `Calculus Final Exams Table`. Você pode precisar verificar sua ortografia.

```js
const titleText = document.querySelector('title')?.innerText.trim();
assert.strictEqual(titleText, "Calculus Final Exams Table");
```

Você deve ter uma tag de abertura para `body`.

```js
assert.match(code, /<body>/i);
```

Você deve ter uma tag de fechamento para `body`.

```js
assert.match(code, /<\/body>/i);
```

Você não deve alterar seu elemento `head`. Certifique-se de que não deletou sua tag de fechamento.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

O elemento `body` deve estar depois do elemento `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
  --fcc-editable-region--
</html>
```
