---
id: 66a9c14e3b34c719e34bae20
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Para adicionar dados de estudantes à tabela, você precisará usar os elementos table row e table data.

O elemento <dfn>table data</dfn>, `td`, é usado para criar uma célula na tabela.

Aqui está um exemplo de uso do elemento `td` para uma tabela de jogadores:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

Dentro do elemento do corpo da sua tabela, adicione um elemento de linha da tabela, `tr`, com três elementos de dados da tabela, `td`. 

O primeiro elemento de dados da tabela deve conter o sobrenome de `Davis`. 

O segundo elemento de dados da tabela deve conter o primeiro nome de `Alex`. 

O terceiro elemento de dados da tabela deve conter a nota de `54`.

# --hints--

Você deve ter um elemento `tr` dentro do seu elemento `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

Você deve ter um elemento `td` com o texto `Davis` dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent, 'Davis');
```

Você deve ter um elemento `td` com o texto `Alex` dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent, 'Alex');
```

Você deve ter um elemento `td` com o texto `54` dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent, '54');
```

Você deve ter todos os três elementos `td` dentro do elemento `tr`.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
