---
id: 5a9d7286424fe3d0e10cad13
title: 給 CSS 變量設置備用值
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bDNfp'
forumTopicId: 301084
dashedName: attach-a-fallback-value-to-a-css-variable
---

# --description--

When using your variable as a CSS property value, you can attach a fallback value that your browser will revert to if the given variable is invalid.

**注意：** 備用值不是用於增強瀏覽器的兼容性，它也不適用於 IE 瀏覽器。 相反，它是用來讓瀏覽器在找不到你的變量時可以顯示一種顏色。

下面是操作方式：

```css
background: var(--penguin-skin, black);
```

如果你的變量沒有設置，這將會把背景設置爲 `black`。 提示：這對調試代碼也會很有幫助。

# --instructions--

`.penguin-top` 和 `.penguin-bottom` 類的變量看起來似乎有點問題。 請爲 class 爲 `penguin-top` 和 `penguin-bottom` 的元素的 `background` 屬性設置一個 `black` 的備用色。

# --hints--

class 爲 `penguin-top` 的元素的 `background` 屬性值應有 `black` 作爲備用顏色。

```js
assert.match(__helpers.removeCssComments(code), /.penguin-top\s*?{[\s\S]*background\s*?:\s*?var\(\s*?--pengiun-skin\s*?,\s*?black\s*?\)\s*?;[\s\S]*}[\s\S]*.penguin-bottom\s{/gi);
```

class 爲 `penguin-bottom` 的元素的 `background` 屬性值應有 `black` 作爲備用顏色。

```js
assert.match(__helpers.removeCssComments(code), /.penguin-bottom\s*?{[\s\S]*background\s*?:\s*?var\(\s*?--pengiun-skin\s*?,\s*?black\s*?\)\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
請爲 class 爲 <code>penguin-top</code> 和 <code>penguin-bottom</code> 的元素的 <code>background</code> 屬性設置一個 <code>black</code> 的備用色。
```
 和 penguin-bottom 的元素的 background 屬性設置一個 black 的備用色。
</code>

# --solutions--

```html
<style>
  .penguin-top {
    background: var(--pengiun-skin, black);
  }
  .penguin-bottom {
    background: var(--pengiun-skin, black);
  }
</style>
```
