---
id: 5b7d72c338cd7e35b63f3e14
title: 改善瀏覽器回退的相容性
challengeType: 0
forumTopicId: 301087
dashedName: improve-compatibility-with-browser-fallbacks
---

# --description--

當你使用 CSS 時，很可能會在某個時候遇到瀏覽器相容性問題。這就是為什麼提供瀏覽器備援以避免潛在問題很重要。

當你的瀏覽器解析網頁的 CSS 時，它會忽略任何它不認識或不支援的屬性。<dfn>例如</dfn>，如果你使用 CSS 變數為網站指定背景顏色，Internet Explorer 會忽略該背景顏色，因為它不支援 CSS 變數。在那種情況下，瀏覽器會使用它對該屬性所擁有的任何值。如果它找不到該屬性的其他設值，它會回復到預設值，而這通常不是理想的。

這表示如果你確實想要提供瀏覽器的備援，只要在你的宣告之前立即提供另一個更廣泛支援的值就很簡單。這樣舊版瀏覽器就會有備援的內容，而新版瀏覽器則會解讀層疊中較後出現的宣告。

# --instructions--

看起來正在使用變數來設定 `.red-box` 類別的背景顏色。讓我們透過在現有宣告之前添加另一個 `background` 宣告，並將其值設定為 `red`，來提升我們的瀏覽器相容性。

# --hints--

你的 `.red-box` 規則應該包含一個備用，該備用的 `background` 設為 `red`，並且緊接在現有的 `background` 宣告之前。

```js
const spacelessCode = __helpers.removeWhiteSpace(__helpers.removeCssComments(code));
assert.match(
  spacelessCode,
  /\.red-box{background:(red|#ff0000|#f00|rgb\(255,0,0\)|rgb\(100%,0%,0%\)|hsl\(0,100%,50%\));background:var\(--red-color\);height:200px;width:200px;}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  :root {
    --red-color: red;
  }
  .red-box {

    background: var(--red-color);
    height: 200px;
    width:200px;
  }
</style>
<div class="red-box"></div>
```

# --solutions--

```html
<style>
  :root {
    --red-color: red;
  }
  .red-box {
    background: red;
    background: var(--red-color);
    height: 200px;
    width:200px;
  }
</style>
<div class="red-box"></div>
```
