---
id: bad87fee1348bd9aedf07756
title: 使用 Important 改寫所有其他樣式
challengeType: 0
videoUrl: 'https://scrimba.com/c/cm24rcp'
forumTopicId: 18249
dashedName: override-all-other-styles-by-using-important
---

# --description--

太棒了！我們剛剛證明了行內樣式會覆寫你 `style` 元素中所有的 CSS 宣告。

但等等。還有最後一種改寫 CSS 的方法。這是所有方法中最強大的方法。但在我們這麼做之前，先來談談為什麼你會想要改寫 CSS。

在許多情況下，你會使用 CSS 函式庫。這些可能會意外地改寫你的 CSS。所以當你絕對需要確定一個元素具有特定的 CSS 時，你可以使用 `!important`。

讓我們回到 `pink-text` 類別宣告。請記得，我們的 `pink-text` 類別被後續的類別宣告、id 宣告和行內樣式覆寫了。

# --instructions--

讓我們為你的粉紅色文字元素的顏色宣告添加關鍵字 `!important`，以 100% 確保你的 `h1` 元素會是粉紅色。

這樣做的一個範例是：

```css
color: red !important;
```

# --hints--

你的 `h1` 元素應該有 `pink-text` 類別。

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

你的 `h1` 元素應該有 `blue-text` 類別。

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

你的 `h1` 元素應該有 `orange-text` 的 `id`。

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

你的 `h1` 元素應該有 `color: white` 的行內樣式。

```js
const commentessCode = __helpers.removeHtmlComments(code);
assert.match(commentessCode, /<h1.*style/gi);
assert.match(commentessCode, /<h1.*style.*color\s*?:/gi);
```

你的 `pink-text` 類別宣告應該包含 `!important` 關鍵字，以改寫所有其他宣告。

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*?\{[\s\S]*?color:.*pink.*!important\s*;?[^\.]*\}/g);
```

你的 `h1` 元素應該是粉紅色。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink !important;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
