---
id: bad87fee1348bd9aedf04756
title: 在後續的 CSS 中改寫樣式
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDQug'
forumTopicId: 18253
dashedName: override-styles-in-subsequent-css
---

# --description--

我們的 `pink-text` 類別覆寫了我們 `body` 元素的 CSS 宣告！

我們剛剛證明了我們的類別會覆寫 `body` 元素的 CSS。所以下一個合乎邏輯的問題是，我們可以做什麼來覆寫我們的 `pink-text` 類別？

# --instructions--

創建一個名為 `blue-text` 的額外 CSS 類別，賦予元素藍色。確保它位於你的 `pink-text` 類別宣告之下。

將 `blue-text` 類別添加到你的 `h1` 元素中，除了你的 `pink-text` 類別之外，讓我們看看哪一個會勝出。

將多個類別屬性套用到 HTML 元素時，類別之間以空格分隔，如下所示：

```html
class="class1 class2"
```

**注意：** 在 HTML 元素中列出類別的順序並不重要。

然而，`<style>` 區段中 `class` 宣告的順序才是重要的。第二個宣告式總是會比第一個具有更高的優先權。因為 `.blue-text` 是第二個宣告的，所以它會覆寫 `.pink-text` 的屬性。

# --hints--

你的 `h1` 元素應該有 `pink-text` 類別。

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

你的 `h1` 元素應該有 `blue-text` 類別。

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

`blue-text` 和 `pink-text` 都應該屬於同一個 `h1` 元素。

```js
assert.isTrue(document.querySelector('.pink-text').classList.contains('blue-text'));
```

你的 `h1` 元素應該是藍色的。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }

  .blue-text {
    color: blue;
  }  
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```
