---
id: bad82fee1322bd9aedf08721
title: 了解絕對單位與相對單位
challengeType: 0
videoUrl: 'https://scrimba.com/c/cN66JSL'
forumTopicId: 301089
dashedName: understand-absolute-versus-relative-units
---

# --description--

最後幾個挑戰都使用像素（`px`）來設定元素的外距或內距。像素是一種長度單位，用來告訴瀏覽器如何調整項目的大小或間距。除了 `px` 之外，CSS 還有許多不同的長度單位選項可供你使用。

長度單位的兩種主要型別是絕對和相對。絕對單位與實體的長度單位綁定。例如，`in` 和 `mm` 分別指涉英吋和毫米。絕對長度單位近似螢幕上的實際測量，但根據螢幕的決議程序會有一些差異。

相對單位，例如 `em` 或 `rem`，是相對於另一個長度值。 例如，`em` 是基於元素的字型大小。 如果你用它來設定 `font-size` 屬性本身，它是相對於父元素的 `font-size`。

**注意：** 有幾個相對單位的選項，可選方案與視窗大小相關。它們在響應式網頁設計原則章節中有涵蓋。

# --instructions--

為你的類別 `red-box` 的元素添加 `padding` 屬性並將其設為 `1.5em`。

# --hints--

你的 `red-box` 類別應該有一個 `padding` 屬性。

```js
const redBox =document.querySelector('.red-box'); 
const style = window.getComputedStyle(redBox); 

assert.notEqual(style['padding-top'], '0px');
assert.notEqual(style['padding-right'], '0px');
assert.notEqual(style['padding-bottom'], '0px');
assert.notEqual(style['padding-left'], '0px');
```

你的 `red-box` 類別應該給元素 1.5em 的 `padding`。

```js
assert.match(__helpers.removeCssComments(code), /\.red-box\s*?{[\s\S]*padding\s*:\s*?1\.5em/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;

  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;
    padding: 1.5em;
  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```
