---
id: bad87fee1348bd9aedf08736
title: 為 HTML Body 元素設定樣式
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB77PHW'
forumTopicId: 18313
dashedName: style-the-html-body-element
---

# --description--

現在讓我們重新開始，並討論 CSS 繼承。

每個 HTML 頁面都有一個 `body` 元素。

# --instructions--

我們可以透過給 `body` 元素設定黑色的 `background-color` 來證明它存在於此。

我們可以透過將以下內容添加到我們的 `style` 元素中來完成這個操作：

```css
body {
  background-color: black;
}
```

# --hints--

你的 `body` 元素應該有黑色的 `background-color`。

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color'];

assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

你的 CSS 規則應該正確格式化，包含開啟和關閉的大括號。

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

你的 CSS 規則應該以分號結尾。

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

# --seed--

## --seed-contents--

```html
<style>

</style>
```

# --solutions--

```html
<style>
body {
  background-color: black;
}
</style>
```
