---
id: bad87fee1348bd9aedf08808
title: 指定字型應如何降級
challengeType: 0
videoUrl: 'https://scrimba.com/c/cpVKBfQ'
forumTopicId: 18304
dashedName: specify-how-fonts-should-degrade
---

# --description--

所有瀏覽器中都有數個預設字型。這些泛型字型家族包括 `monospace`、`serif` 和 `sans-serif`。

當一個字型不可用時，你可以告訴瀏覽器「降級」到另一個字型。

例如，如果你想讓一個元素使用 `Helvetica` 字型，但當 `Helvetica` 不可用時退化為 `sans-serif` 字型，你會這樣指定：

```css
p {
  font-family: Helvetica, sans-serif;
}
```

泛型字型家族名稱不區分大小寫。此外，它們不需要引號，因為它們是 CSS 關鍵字。

# --instructions--

首先，將 `monospace` 字型應用到 `h2` 元素，使其現在擁有兩種字型 - `Lobster` 和 `monospace`。

在上一個挑戰中，你使用 `link` 頁籤匯入了 `Lobster` 字型。現在請將從 Google Fonts 匯入 `Lobster` 字型的程式碼（使用你之前學過的 HTML 註解）註解起來，使其不再可用。注意你的 `h2` 元素如何退化為 `monospace` 字型。

**注意：** 如果你的電腦已安裝 `Lobster` 字型，你將不會看到降級，因為你的瀏覽器能夠找到該字型。

# --hints--

你的 h2 元素應該使用字型 `Lobster`。

```js
const h2Element = document.querySelector('h2');
const fontFamily = window.getComputedStyle(h2Element)['font-family']; 
assert.match(fontFamily, /^"?lobster/i);
```

當 `Lobster` 不可用時，你的 h2 元素應該降級為字型 `monospace`。

```js
assert.match(__helpers.removeCssComments(code), /\s*h2\s*\{\s*font-family\s*\:\s*(\'|"|)Lobster\1\s*,\s*monospace\s*;?\s*\}/gi);
```

你應該在呼叫 Google 的 `Lobster` 字型前面加上 `<!--` 來註解掉它。

```js
assert.match(code, /<!--[^fc]/gi);
```

你應該透過添加 `-->` 來關閉你的註解。

```js
assert.match(code, /[^fc]-->/gi);
```

# --seed--

## --seed-contents--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<!--<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">-->
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster, monospace;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
