---
id: bad87fee1348bd9aede08718
title: 使用 RGB 值為元素上色
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkp2fr'
forumTopicId: 18369
dashedName: use-rgb-values-to-color-elements
---

# --description--

另一種你可以在 CSS 中表述顏色的方式是使用 `RGB` 值。

黑色的 `RGB` 值看起來像這樣：

```css
rgb(0, 0, 0)
```

白色的 `RGB` 值看起來像這樣：

```css
rgb(255, 255, 255)
```

與其像使用十六進位碼那樣使用六個十六進位數字，使用 `RGB` 時，你以介於 0 和 255 之間的數字來指定每種顏色的亮度。

如果你計算，單一顏色的兩位數等於 16 乘以 16，總共給我們 256 個值。因此從零開始計數的 `RGB`，擁有與十六進位碼完全相同數量的可能值。

以下是如何使用其 RGB 代碼將 `body` 背景更改為橘色的範例。

```css
body {
  background-color: rgb(255, 165, 0);
}
```

# --instructions--

讓我們將 `body` 元素背景色中的十六進位色碼取代為黑色的 RGB 值：`rgb(0, 0, 0)`

# --hints--

你的 `body` 元素應該有黑色背景。

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

你應該使用 `rgb` 為你的 `body` 元素添加黑色背景。

```js
assert.match(__helpers.removeCssComments(code), /rgb\s*\(\s*0\s*,\s*0\s*,\s*0\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #F00;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: rgb(0, 0, 0);
  }
</style>
```
