---
id: bad87fee1348bd9aedf08756
title: 優先考慮一種樣式勝過另一種樣式
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ8wnHv'
forumTopicId: 18258
dashedName: prioritize-one-style-over-another
---

# --description--

有時你的 HTML 元素會收到多個彼此衝突的樣式。

例如，你的 `h1` 元素不能同時是綠色和粉紅色。

讓我們看看當我們創建一個讓文字變成粉紅色的類別，然後將它套用到一個元素時會發生什麼。這個類別會改寫 `body` 元素的 `color: green;` CSS 屬性嗎？

# --instructions--

創建一個名為 `pink-text` 的 CSS 類別，為元素設定粉紅色。

為你的 `h1` 元素添加 `pink-text` 類別。

# --hints--

你的 `h1` 元素應該有 `pink-text` 類別。

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

你的 `<style>` 應該有一個 `pink-text` CSS 類別，用來改變 `color`。

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*\{\s*color\s*:\s*.+\s*;?\s*\}/g);
```

你的 `h1` 元素應該是粉紅色。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
</style>
<h1>Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```
