---
id: bad87fee1348bd9aedf08721
title: 使用十六進位碼混合顏色
challengeType: 0
videoUrl: 'https://scrimba.com/c/cK89PhP'
forumTopicId: 18359
dashedName: use-hex-code-to-mix-colors
---

# --description--

回顧一下，十六進位碼使用 6 個十六進位數字來表述顏色，每兩個數字分別代表紅色（R）、綠色（G）和藍色（B）組件。

從這三種純色（紅色、綠色和藍色）中，我們可以變化每種顏色的數量，來創建超過 1,600 萬種其他顏色！

例如，橙色是純紅色，混合一些綠色，且沒有藍色。在十六進位碼中，這表示為 `#FFA500`。

數字 `0` 是十六進位碼中最低的數字，並表述完全沒有顏色。

數字 `F` 是十六進位碼中最高的數字，並表述最大可能的亮度。

# --instructions--

將我們 `style` 元素中的顏色字取代為正確的十六進位色碼。

<table><tbody><tr><th>Color</th><th>Hex Code</th></tr><tr><td>Dodger Blue</td><td><code>#1E90FF</code></td></tr><tr><td>Green</td><td><code>#00FF00</code></td></tr><tr><td>Orange</td><td><code>#FFA500</code></td></tr><tr><td>Red</td><td><code>#FF0000</code></td></tr></tbody></table>

# --hints--

你的 `h1` 元素中帶有文字 `I am red!` 應該被賦予紅色的 `color`。

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

應使用顏色紅色的 `hex code`，而非字 `red`。

```js
assert.match(code, /\.red-text\s*?{\s*?color\s*:\s*?(#FF0000|#F00)\s*?;?\s*?}/gi);
```

你的 `h1` 元素中帶有文字 `I am green!` 應該被賦予 `color` green。

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

應該使用綠色的 `hex code`，而非字 `green`。

```js
assert.match(code, /\.green-text\s*?{\s*?color\s*:\s*?(#00FF00|#0F0)\s*?;?\s*?}/gi);
```

你的 `h1` 元素中帶有文字 `I am dodger blue!` 應該被賦予 `color` dodger blue。

```js
const blueText = document.querySelector('.dodger-blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(30, 144, 255)');
```

應該使用顏色 dodger blue 的 `hex code`，而非字 `dodgerblue`。

```js
assert.match(code, /\.dodger-blue-text\s*?{\s*?color\s*:\s*?#1E90FF\s*?;?\s*?}/gi);
```

你的 `h1` 元素中帶有文字 `I am orange!` 應該被賦予 `color` 為橙色。

```js
const orangeText = document.querySelector('.orange-text');
const color = window.getComputedStyle(orangeText)['color']; 

assert.strictEqual(color, 'rgb(255, 165, 0)');
```

應該使用顏色橘色的 `hex code`，而非字 `orange`。

```js
assert.match(code, /\.orange-text\s*?{\s*?color\s*:\s*?#FFA500\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: black;
  }
  .green-text {
    color: black;
  }
  .dodger-blue-text {
    color: black;
  }
  .orange-text {
    color: black;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #FF0000;
  }
  .green-text {
    color: #00FF00;
  }
  .dodger-blue-text {
    color: #1E90FF;
  }
  .orange-text {
    color: #FFA500;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```
