---
id: bad87fee1348bd9aedf06756
title: 使用行內樣式改寫類別宣告式
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDRha'
forumTopicId: 18252
dashedName: override-class-declarations-with-inline-styles
---

# --description--

所以我們已經證明了 id 宣告會覆寫類別宣告，無論它們在你的 `style` 元素 CSS 中宣告的位置。

還有其他方法可以改寫 CSS。你還記得行內樣式嗎？

# --instructions--

使用行內樣式嘗試將我們的 `h1` 元素設為白色。請記住，行內樣式看起來像這樣：

```html
<h1 style="color: green;">
```

在你的 `h1` 元素上保留 `blue-text` 和 `pink-text` 類別。

# --hints--

你的 `h1` 元素應該有 `pink-text` 類別。

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

你的 `h1` 元素應該有 `blue-text` 類別。

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

你的 `h1` 元素應該有 `orange-text` 的 id。

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

你的 `h1` 元素應該有行內樣式。

```js
assert.exists(document.querySelector('h1[style]'));
```

你的 `h1` 元素應該是白色的。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
