---
id: bad87fee1348bd9aedf08726
title: 使用十六進位碼指定顏色
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

你知道在 CSS 中還有其他表述顏色的方式嗎？其中一種方式稱為十六進位碼，簡稱 hex code。

我們通常使用 <dfn>十進位數</dfn>，或稱為基數 10 的數字，每個位數使用符號 0 到 9。<dfn>十六進位數</dfn>（或 <dfn>hex</dfn>）是基數 16 的數字。這表示它使用十六個不同的符號。像十進位數一樣，符號 0-9 表述數值零到九。接著 A、B、C、D、E、F 表述數值十到十五。總共，0 到 F 可以表述十六進位數中的一個位數，給我們 16 個可能的數值。你可以在<a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">這裡</a>找到更多關於十六進位數的資訊。

在 CSS 中，我們可以使用 6 個十六進位數字來表述顏色，每兩個數字分別代表紅色（R）、綠色（G）和藍色（B）組件。例如，`#000000` 是黑色，且也是最低可能的值。你可以在<a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">這裡</a>找到更多關於 RGB 色彩系統的資訊。

```css
body {
  color: #000000;
}
```

# --instructions--

將我們 `body` 元素的背景顏色中的字 `black` 取代為它的十六進位色碼表述 `#000000`。

# --hints--

你的 `body` 元素應該有黑色的 `background-color`。

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

應該使用代表黑色的十六進位色碼，而非字 `black`。

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
