---
id: bad87fee1348bd8aedf06756
title: 透過樣式化 ID 屬性來改寫類別宣告
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpDhB'
forumTopicId: 18251
dashedName: override-class-declarations-by-styling-id-attributes
---

# --description--

我們剛剛證明了瀏覽器會依照宣告式的順序，從上到下讀取 CSS。這表示在發生衝突時，瀏覽器會使用最後出現的 CSS 宣告式。請注意，即使我們將 `blue-text` 放在 `pink-text` 之前於 `h1` 元素的類別中，瀏覽器仍會依照宣告式的順序，而非使用順序來判斷！

但我們還沒完成。你還有其他方法可以改寫 CSS。你還記得 `id` 屬性嗎？

讓我們改寫你的 `pink-text` 和 `blue-text` 類別，並且透過給 `h1` 元素一個 id，然後為該 id 設定樣式，將你的 `h1` 元素變成橘色。

# --instructions--

給你的 `h1` 元素設定 `id` 屬性為 `orange-text`。請記住，id 樣式看起來像這樣：

```html
<h1 id="orange-text">
```

在你的 `h1` 元素上保留 `blue-text` 和 `pink-text` 類別。

為你的 `style` 元素中的 `orange-text` id 創建一個 CSS 宣告。以下是範例：

```css
#brown-text {
  color: brown;
}
```

**注意：** 無論你是在 `pink-text` 類別的上方還是下方宣告這個 CSS，都沒有關係，因為 `id` 屬性總是會擁有優先權。

# --hints--

你的 `h1` 元素應該有 `pink-text` 類別。

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

你的 `h1` 元素應該有 `blue-text` 類別。

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

你的 `h1` 元素應該有 `orange-text` 的 id。

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'),'orange-text');
```

應該只有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你的 `orange-text` id 應該有一個 CSS 宣告。

```js
assert.match(__helpers.removeCssComments(code), /#orange-text\s*{/gi);
```

你的 `h1` 不應該有任何 `style` 屬性。

```js
assert.notMatch(__helpers.removeHtmlComments(code), /<h1.*style.*>/gi);
```

你的 `h1` 元素應該是橘色。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 165, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
  #orange-text {
    color: orange;
  }  
</style>
<h1 id="orange-text"  class="pink-text blue-text">Hello World!</h1>
```
