---
id: bad87fee1348bd9aedf08746
title: 從 Body 元素繼承樣式
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bmdtR'
forumTopicId: 18204
dashedName: inherit-styles-from-the-body-element
---

# --description--

現在我們已經證明每個 HTML 頁面都有一個 `body` 元素，且其 `body` 元素也可以用 CSS 來設計樣式。

請記住，你可以像設定其他 HTML 元素一樣為你的 `body` 元素設計樣式，而你所有其他元素都會繼承你的 `body` 元素的樣式。

# --instructions--

首先，創建一個帶有文字 `Hello World` 的 `h1` 元素。

接著，為你的頁面上所有元素添加 `color: green;` 到你的 `body` 元素的樣式宣告，讓它們呈現 `green` 顏色。

最後，為你的 `body` 元素添加 `font-family: monospace;` 到你的 `body` 元素的樣式宣告中，將字型設定為 `monospace`。

# --hints--

你應該創建一個 `h1` 元素。

```js
assert.isNotEmpty(document.querySelectorAll('h1'));
```

你的 `h1` 元素應該有文字 `Hello World`。

```js
assert.match(
  document.querySelector('h1').textContent,
  /hello world/i
);
```

你的 `h1` 元素應該有一個結束標籤。

```js
const commentlessCode = __helpers.removeHtmlComments(code);
assert.match(commentlessCode, /<\/h1>/g);
assert.match(commentlessCode, /<h1/g);
assert.lengthOf(commentlessCode.match(/<\/h1>/g), commentlessCode.match(/<h1/g).length);
```

你的 `body` 元素應該有 `color` 屬性為 `green`。

```js
const bodyElement = document.querySelector('body');
const color = window.getComputedStyle(bodyElement)['color']; 
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

你的 `body` 元素應該具有 `monospace` 的 `font-family` 屬性。

```js
const bodyElement = document.querySelector('body');
const fontFamily = window.getComputedStyle(bodyElement)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

你的 `h1` 元素應該從你的 `body` 元素繼承字型 `monospace`。

```js
const h1Element = document.querySelector('h1');
const fontFamily = window.getComputedStyle(h1Element)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

你的 `h1` 元素應該從你的 `body` 元素繼承顏色 `green`。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }

</style>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }

</style>
<h1>Hello World!</h1>
```
