---
id: bad87fee1348bd9aedf08807
title: 匯入 Google 字型
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9MRsJ'
forumTopicId: 18200
dashedName: import-a-google-font
---

# --description--

除了指定大多數作業系統中常見的字型之外，我們也可以指定非標準的訂製網頁字型以用於我們的網站。網際網路上有許多網頁字型的來源。此範例中，我們將專注於 Google Fonts 函式庫。

Google Fonts 是一個免費的網頁字型函式庫，你可以透過取用該字型的 URL 在你的 CSS 中使用。

那麼，讓我們繼續匯入並套用 Google 字型（請注意，如果你的國家封鎖了 Google，你將需要跳過這個挑戰）。

要匯入 Google 字型，你可以從 Google Fonts 程式庫複製該字型的 URL，然後將它貼到你的 HTML 中。針對這個挑戰，我們將匯入 `Lobster` 字型。為此，請複製以下程式碼片段，並將它貼到你的程式碼編輯器頂端（在開啟的 `style` 元素之前）：

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
```

現在你可以在你的 CSS 中使用 `Lobster` 字型，方法是將 `Lobster` 作為 FAMILY_NAME，如以下範例所示：

```css
font-family: FAMILY_NAME, GENERIC_NAME;
```

GENERIC_NAME 是選用的，並且是在其他指定字型不可用時的備用字型。這在下一個挑戰中涵蓋。

字型名稱區分大小寫，且如果名稱中有空格，則需要用引號括起來。例如，使用 `"Open Sans"` 字型時需要引號，但使用 `Lobster` 字型時則不需要。

# --instructions--

將 `Lobster` 字型匯入你的網頁。然後，使用元素選擇器將 `Lobster` 設定為你的 `h2` 元素的 `font-family`。

# --hints--

你應該匯入 `Lobster` 字型。

```js
assert.exists(document.querySelector('link[href*="googleapis" i]'));
```

你的 `h2` 元素應該使用字型 `Lobster`。

```js
const h2 = document.querySelector('h2'); 
const fontFamily = window.getComputedStyle(h2)['font-family']; 
assert.match(fontFamily, /lobster/i);
```

你應該只使用 `h2` 元素選擇器來更改字型。

```js
assert.match(__helpers.removeHtmlComments(code), /\s*[^\.]h2\s*\{\s*font-family\s*:\s*('|"|)Lobster\1\s*(,\s*('|"|)[a-z -]+\3\s*)?(;\s*\}|\})/gi);
```

你的 `p` 元素仍應使用字型 `monospace`。

```js
const paragraphElement = document.querySelector('p');
const fontFamily = window.getComputedStyle(paragraphElement)['font-family']; 
assert.match(fontFamily, /monospace/i);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }

  h2 {
    font-family: Lobster;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
