---
id: bad87fee1348bd9aedf08719
title: 使用縮寫十六進位碼
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpKAm'
forumTopicId: 18338
dashedName: use-abbreviated-hex-code
---

# --description--

許多人會被超過 1600 萬種顏色的可能性弄得不知所措。而且很難記住十六進位碼。幸運的是，你可以將它縮短。

例如，紅色的十六進位碼 `#FF0000` 可以縮短為 `#F00`。這種縮短的形式為紅色、綠色和藍色各提供一個數字。

這會將可能的顏色總數減少到約 4,000 種。但瀏覽器會將 `#FF0000` 和 `#F00` 解讀為完全相同的顏色。

# --instructions--

繼續，嘗試使用縮寫的十六進位色碼為正確的元素上色。

<table><tbody><tr><th>Color</th><th>Short Hex Code</th></tr><tr><td>Cyan</td><td><code>#0FF</code></td></tr><tr><td>Green</td><td><code>#0F0</code></td></tr><tr><td>Red</td><td><code>#F00</code></td></tr><tr><td>Fuchsia</td><td><code>#F0F</code></td></tr></tbody></table>

# --hints--

你的 `h1` 元素中帶有文字 `I am red!` 應該被賦予紅色的 `color`。

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

應該使用紅色的縮寫十六進位碼，而非十六進位碼 `#FF0000`。

```js
assert.match(code,/\.red-text\s*?{\s*?color\s*:\s*?#F00\s*?;?\s*?}/gi);
```

你的 `h1` 元素中帶有文字 `I am green!` 應該被賦予 `color` green。

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

應該使用綠色的縮寫十六進位碼，而非十六進位碼 `#00FF00`。

```js
assert.match(__helpers.removeCssComments(code), /\.green-text\s*?{\s*?color\s*:\s*?#0F0\s*?;?\s*?}/gi);
```

你的 `h1` 元素中帶有文字 `I am cyan!` 應該被賦予 `color` 為 cyan。

```js
const cyanText = document.querySelector('.cyan-text');
const color = window.getComputedStyle(cyanText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 255)');
```

應該使用代表青色的縮寫十六進位色碼，而非十六進位色碼 `#00FFFF`。

```js
assert.match(__helpers.removeCssComments(code), /\.cyan-text\s*?{\s*?color\s*:\s*?#0FF\s*?;?\s*?}/gi);
```

你的 `h1` 元素中帶有文字 `I am fuchsia!`，應該設定 `color` 為 fuchsia。

```js
const fuchsiaText = document.querySelector('.fuchsia-text');
const color = window.getComputedStyle(fuchsiaText)['color']; 
assert.strictEqual(color,'rgb(255, 0, 255)');
```

應使用代表色彩 fuchsia 的縮寫十六進位碼，而非十六進位碼 `#FF00FF`。

```js
assert.match(__helpers.removeCssComments(code), /\.fuchsia-text\s*?{\s*?color\s*:\s*?#F0F\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: #000000;
  }
  .fuchsia-text {
    color: #000000;
  }
  .cyan-text {
    color: #000000;
  }
  .green-text {
    color: #000000;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #F00;
  }
  .fuchsia-text {
    color: #F0F;
  }
  .cyan-text {
    color: #0FF;
  }
  .green-text {
    color: #0F0;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```
