---
id: bad88fee1348bd9aedf08825
title: 調整元素的內距
challengeType: 0
videoUrl: 'https://scrimba.com/c/cED8ZC2'
forumTopicId: 301083
dashedName: adjust-the-padding-of-an-element
---

# --description--

現在讓我們先將 `Cat Photo App` 放一邊，暫時學習更多關於 HTML 的樣式設計。

你可能已經注意到，所有 HTML 元素本質上都是小矩形。

三個重要屬性控制每個 HTML 元素周圍的空間：`padding`、`border` 和 `margin`。

元素的 `padding` 控制元素內容與其 `border` 之間的間距。

在這裡，我們可以看到藍色框和紅色框巢狀於黃色框內。請注意，紅色框的 `padding` 比藍色框多。

當你增加藍色方塊的 `padding` 時，會增加文字與其周圍邊框之間的距離（`padding`）。

# --instructions--

將你的藍色框的 `padding` 更改為與紅色框相同。

# --hints--

你的 `blue-box` 類別應該給元素 `20px` 的 `padding`。

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
