---
id: bad87fee1348bd9aedf08805
title: 使用 CSS 選擇器來設定元素樣式
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKMBT2'
forumTopicId: 18349
dashedName: use-css-selectors-to-style-elements
---

# --description--

使用 CSS，有數百個 CSS 屬性可以用來改變頁面上元素的外觀。

當你輸入 `<h2 style="color: red;">CatPhotoApp</h2>` 時，你正在使用行內 CSS 為該個別的 `h2` 元素設定樣式，CSS 代表階層式樣式表。

這是一種指定元素樣式的方法，但有更好的方式來套用 CSS。

在你的程式碼頂部，創建一個像這樣的 `style` 區塊：

```html
<style>
</style>
```

在該樣式區塊內，你可以為所有 `h2` 元素創建一個 <dfn>CSS 選擇器</dfn>。例如，如果你想讓所有 `h2` 元素變成紅色，你會添加一條看起來像這樣的樣式規則：

```html
<style>
  h2 {
    color: red;
  }
</style>
```

請注意，每個元素的樣式規則必須同時有開啟和關閉的大括號（`{` 和 `}`）。你也需要確保元素的樣式定義位於開啟和關閉的 `style` 頁籤之間。最後，請務必在每個元素的樣式規則結尾加上分號。

# --instructions--

刪除你的 `h2` 元素的 style 屬性，並改為創建一個 CSS `style` 區塊。為你的所有 `h2` 元素添加必要的 CSS，使其變成藍色。

# --hints--

應該從你的 `h2` 元素中移除 `style` 屬性。

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

你應該創建一個 `style` 元素。

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
assert.isAtLeast(document.querySelectorAll('style:not(.fcc-hide-header)').length, 1);
```

你的 `h2` 元素應該是藍色的。

```js
const h2Element = document.querySelector('h2');
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

你的樣式表中 `h2` 的宣告應該是有效的，並且包含分號和結束大括號。

```js
assert.match(__helpers.removeCssComments(code), /h2\s*\{\s*color\s*:.*;\s*\}/g);
```

你所有的 `style` 元素都應該是有效的並且有結束標籤。

```js
assert.match(__helpers.removeHtmlComments(code), /<\/style>/g);

const closingElementLength = __helpers.removeHtmlComments(code).match(/<\/style>/g).length;
const openingElementsLength = __helpers.removeHtmlComments(code).match(/<style((\s)*((type|media|scoped|title|disabled)="[^"]*")?(\s)*)*>/g).length;
assert.strictEqual(closingElementLength, openingElementsLength);
```

# --seed--

## --seed-contents--

```html
<h2 style="color: red;">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  h2 {
    color: blue;
  }
</style>
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
