---
id: 66ad8294a0ad902f1b31b612
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop lernen Sie, wie Sie mit den Grundlagen von JavaScript arbeiten, indem Sie einen Begrüßungs-Bot erstellen.

In diesem ersten Schritt sollen Sie eine Nachricht von dem Begrüßungs-Bot in der Konsole ausgeben.

Denken Sie daran, dass Sie in vorherigen Lektionen etwas über `console.log()` und Strings gelernt haben.

Hier eine Erinnerung, wie man `console.log()` mit Strings verwendet:

```js
console.log("Hello, World!");
```

Fügen Sie eine `console.log()`-Anweisung hinzu, die den String `"Hi there!"` in der Konsole ausgibt. Vergessen Sie nicht die Anführungszeichen um die Nachricht!

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Ihr Code sollte eine `console.log()`-Anweisung enthalten.

```js
assert.isNotEmpty(logSpy.calls);
```

Ihre `console`-Anweisung sollte die Nachricht `"Hi there!"` enthalten. Vergessen Sie die Anführungszeichen nicht.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
