---
id: 66ad8d150264db3926eccfeb
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

In den vorherigen Lektionen haben Sie das Schlüsselwort `let` kennengelernt und wie man Variablen in JavaScript deklariert.

Hier eine Erinnerung daran, wie man eine Variable mit dem Schlüsselwort `let` deklariert:

```js
let greeting;
```

Verwenden Sie das Schlüsselwort `let`, um eine Variable namens `bot` zu deklarieren.

**HINWEIS**: Sie verwenden hier `let`, weil Sie später im Workshop den Wert der Variable `bot` ändern werden.

# --hints--

Sie sollten das Schlüsselwort `let` in Ihrem Code haben.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

Sie sollten eine Variable namens `bot` deklarieren.

```js
assert.isNotNull(bot);
```

Sie sollten das Schlüsselwort `let` verwenden, um die Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
