---
id: 66adbc8a9793d64250f7e609
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man Variablen Werte neu zuweist, zum Beispiel so:

```js
let name = "John";
console.log(name); // "John"

name = "Jane";
console.log(name); // "Jane"
```

Verwenden Sie die Zuweisung, um den String `"professorBot"` der Variablen `bot` zuzuweisen.

# --hints--

Sie sollten den String `"professorBot"` der Variablen `bot` zuweisen.

```js
assert.strictEqual(bot, "professorBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

--fcc-editable-region--

--fcc-editable-region--
```
