---
id: 66adb844118ba74107ce771f
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Die nächste Nachricht des Bots wird den Standort des Bots betreffen.

Erstellen Sie eine Variable namens `botLocationSentence`.

Verwenden Sie dann die Verkettung von Strings mit dem `+`-Operator, um den String `"I live in "` mit der Variablen `botLocation` gefolgt von einem Punkt (`.`) zu verbinden.

Weisen Sie diesen Wert der Variablen `botLocationSentence` zu.

Geben Sie anschließend den Wert von `botLocationSentence` in der Konsole aus.

# --hints--

Sie sollten eine Variable `botLocationSentence` haben.

```js
assert.isNotNull(botLocationSentence);
```

Ihre Variable `botLocationSentence` sollte den Wert eines Strings enthalten.

```js
assert.isString(botLocationSentence);
```

Sie sollten die Verkettung von Strings mit dem `+`-Operator verwenden, um den String `"I live in "` mit der Variablen `botLocation` gefolgt von einem Punkt (`.`) zu verbinden. Achten Sie auf Leerzeichen.

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

Ihr Code sollte eine `console`-Anweisung enthalten.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

Ihre `console`-Anweisung sollte die Variable `botLocationSentence` ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
