---
id: 66adcf383276814776aba3ca
title: Schritt 14
challengeType: 1
dashedName: step-14
---

# --description--

Erstellen Sie als Nächstes eine Variable namens `favoriteSubjectSentence`.

Verwenden Sie String-Verkettung, um den String `"My favorite subject is "` mit der Variablen `favoriteSubject` gefolgt von einem Punkt zu verbinden.

Weisen Sie das Ergebnis der Variablen `favoriteSubjectSentence` zu.

Geben Sie dann den Wert von `favoriteSubjectSentence` in der Konsole aus.

# --hints--

Sie sollten eine Variable namens `favoriteSubjectSentence` haben.

```js
assert.isNotNull(favoriteSubjectSentence);
```

Ihre Variable `favoriteSubjectSentence` sollte ein String sein.

```js
assert.isString(favoriteSubjectSentence);
```

Sie sollten String-Verkettung verwenden, um den String `"My favorite subject is "` mit der Variablen `favoriteSubject` gefolgt von einem Punkt zu verbinden. Achten Sie auf die Leerzeichen.

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

Ihr Code sollte eine `console`-Anweisung enthalten.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

Sie sollten den Wert von `favoriteSubjectSentence` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
