---
id: 66ad9c633fa26d3c1475eae3
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Jetzt ist es Zeit, eine weitere Nachricht vom Bot hinzuzufügen.

Fügen Sie dem Code eine weitere `console`-Anweisung hinzu, die die Nachricht `"Allow me to introduce myself."` protokolliert.

# --hints--

Sie sollten eine dritte `console.log()`-Anweisung in Ihrem Code haben.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

Ihre `console`-Anweisung sollte die Nachricht `"Allow me to introduce myself."` enthalten. Vergessen Sie nicht die Anführungszeichen.

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
