---
id: 66ad8ab945d266383f318cbf
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Jetzt sollten Sie die erste Nachricht vom Bot in der Konsole sehen.

Es ist Zeit, eine zweite Nachricht vom Bot hinzuzufügen.

Fügen Sie eine weitere `console.log`-Anweisung hinzu, um die Nachricht `"I am excited to talk to you."` in der Konsole auszugeben.

# --hints--

Sie sollten eine zweite `console.log()`-Anweisung in Ihrem Code haben.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

Ihre `console`-Anweisung sollte die Nachricht `"I am excited to talk to you."` enthalten. Vergessen Sie die Anführungszeichen nicht.

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
