---
id: 66adc42868cab843ccee87d9
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Jetzt ist es Zeit, den neuen Wert von `bot` zu sehen.

Beginnen Sie damit, eine neue Variable namens `nicknameIntroduction` zu erstellen.

Verwenden Sie String-Verkettung, um den String `"My nickname is "` mit der Variable `bot` zu verbinden, gefolgt von einem Punkt (`.`).

Weisen Sie den resultierenden String der Variable `nicknameIntroduction` zu.

Geben Sie dann den Wert von `nicknameIntroduction` in der Konsole aus.

# --hints--

Sie sollten eine Variable namens `nicknameIntroduction` haben.

```js
assert.isNotNull(nicknameIntroduction);
```

Ihre Variable `nicknameIntroduction` sollte ein String sein.

```js
assert.isString(nicknameIntroduction);
```

Sie sollten String-Verkettung verwenden, um den String `"My nickname is "` mit der Variable `bot` zu verbinden, gefolgt von einem Punkt (`.`). Achten Sie auf die Leerzeichen.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

Ihr Code sollte eine `console`-Anweisung enthalten.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

Sie sollten den Wert von `nicknameIntroduction` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
