---
id: 66ad984e92ce213a66f7e51c
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Jetzt ist es Zeit, Ihrer `bot`-Variablen einen Wert zuzuweisen.

In vorherigen Lektionen haben Sie gelernt, wie man Variablen Werte so zuweist:

```js
variableName = "Here is the value";
```

Denken Sie daran, dass das, was auf der rechten Seite des Zuweisungsoperators `=` steht, der Wert ist, den Sie der Variablen auf der linken Seite zuweisen.

Weisen Sie der `"teacherBot"`-Variablen den String `bot` zu.

# --hints--

Sie sollten Ihrer `bot`-Variablen einen String zuweisen.

```js
assert.isString(bot);
```

Ihrer `bot`-Variablen sollte der String `"teacherBot"` zugewiesen werden.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
