---
id: 66adc5919c1853448f119d11
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Jetzt sieht es so aus, als wolle der Bot seinen Spitznamen ändern.

Weisen Sie mit einer Zuweisung den String `"awesomeTeacherBot"` der Variablen `bot` zu.

# --hints--

Sie sollten den String `"awesomeTeacherBot"` der Variablen `bot` zuweisen.

```js
assert.strictEqual(bot, "awesomeTeacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
