---
id: 66adc82170acb1464bd348eb
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Die letzten Nachrichten des Bots werden sich auf das Lieblingsthema des Bots konzentrieren.

Beginnen Sie damit, eine neue Variable namens `favoriteSubject` zu erstellen und ihr den String `"Computer Science"` zuzuweisen.

# --hints--

Sie sollten eine Variable namens `favoriteSubject` haben.

```js
assert.isNotNull(favoriteSubject);
```

Ihre `favoriteSubject`-Variable sollte ein String sein.

```js
assert.isString(favoriteSubject);
```

Sie sollten der `"Computer Science"`-Variable den String `favoriteSubject` zuweisen.

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
