---
id: 66adc6046acc18453decf577
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Um den neuen Spitznamen des Bots zu sehen, müssen Sie eine neue Nachricht in der Konsole ausgeben.

Erstellen Sie eine neue Variable namens `newNicknameGreeting`.

Verwenden Sie dann String-Verkettung, um den String `"I love my nickname but I wish people would call me "` mit der Variable `bot` zu verbinden, gefolgt von einem Punkt.

Weisen Sie das Ergebnis der Variable `newNicknameGreeting` zu.

Geben Sie anschließend den Wert von `newNicknameGreeting` in der Konsole aus.

# --hints--

Sie sollten eine Variable namens `newNicknameGreeting` haben.

```js
assert.isNotNull(newNicknameGreeting);
```

Ihre Variable `newNicknameGreeting` sollte ein String sein.

```js
assert.isString(newNicknameGreeting);
```

Sie sollten String-Verkettung verwenden, um den String `"I love my nickname but I wish people would call me "` mit der Variable `bot` zu verbinden, gefolgt von einem Punkt.

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

Ihr Code sollte eine `console`-Anweisung enthalten.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

Sie sollten den Wert von `newNicknameGreeting` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
