---
id: 66ad998285df023b23dacdd3
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Jetzt ist es Zeit, die Variable `botLocation` zu initialisieren.

Wenn Sie Variablen mit mehreren Wörtern deklarieren müssen, können Sie die Benennungskonvention <dfn>CamelCase</dfn> verwenden.

Bei der Verwendung von `camelCase` wird das erste Wort komplett in Kleinbuchstaben geschrieben und der erste Buchstabe jedes folgenden Wortes wird großgeschrieben.

Hier ist ein Beispiel:

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

Deklarieren und weisen Sie den String `"the universe"` der Variable `botLocation` in derselben Zeile mit dem Schlüsselwort `let` zu.

# --hints--

Verwenden Sie das Schlüsselwort `let`, um eine Variable zu deklarieren.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Ihre Variable sollte `botLocation` heißen.

```js
assert.isNotNull(botLocation);
```

Sie sollten Ihrer Variable `botLocation` einen String zuweisen.

```js
assert.isString(botLocation);
```

Ihrer Variable `botLocation` sollte der String `"the universe"` zugewiesen werden.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
