---
id: 687ad7be1ff45032ccf1d92f
title: 建置主題切換器
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

在這個實驗中，你將建置一個可以切換不同主題的應用程式。當你切換到不同的主題時，背景顏色會改變，並且會顯示不同的訊息。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `button` 元素，裡面有文字 `Switch Theme`。
2. 你的 `button` 元素應該有以下屬性：
   - 設定為 `"theme-switcher-button"` 的 `id` 屬性。
   - 設定為 `"true"` 的 `aria-haspopup` 屬性。
   - 設定為 `"false"` 的 `aria-expanded` 屬性。
   - 設定為 `"theme-dropdown"` 的 `aria-controls` 屬性。
3. 你應該有一個帶有以下屬性的 `ul` 元素：
   - 一個 `id` 屬性設定為 `"theme-dropdown"`。
   - 設定為 `"menu"` 的 `role` 屬性。
   - 設定為 `"theme-switcher-button"` 的 `aria-labelledby` 屬性。
   - 一個 `hidden` 屬性。
4. 你的 `ul` 元素應該至少有兩個 `li` 項目，其 `role` 屬性設定為 `"menuitem"`，並且包含你選擇的表述不同主題的文字。
5. 你的每個 `li` 元素都應該有一個以 `theme-` 開頭並以你為該 `li` 元素文字設定的主題結尾的 `id` 屬性。例如，如果你的其中一個主題是 `Light`，那麼你的 `id` 應該是 `theme-light`。
6. 你應該有一個元素，其 `aria-live` 屬性設為 `"polite"`。
7. 你應該有一個 `themes` 陣列，裡面至少包含兩個物件，每個物件都包含 `name` 和 `message` 屬性。`name` 將表述不同的主題，而 `message` 將在主題切換時顯示。你可以自由設計 `name` 和 `message` 的值，但 `name` 的值應該與你在 `li` 項目中設定的選項相符。
8. 當使用者點擊 `#theme-switcher-button` 時，`#theme-dropdown` 元素應該顯示可用的主題。
9. 當 `#theme-dropdown` 元素顯示時，`button` 元素的 `aria-expanded` 屬性應設為 `"true"`。
10. 當使用者在 `#theme-dropdown` 元素顯示時點擊 `#theme-switcher-button`，應該關閉選單，並添加 `hidden` 屬性，且將 `aria-expanded` 屬性設為 `"false"`。
11. 當使用者點擊 `#theme-switcher-button` 並選擇主題時，應將對應的 `theme-<name>` 類別添加到 `body` 元素，並在 `aria-live="polite"` 元素中顯示你在 `themes` 陣列中設定的相關主題訊息。

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

你應該有一個 `button` 元素。

```js
assert.exists(document.querySelector("button"));
```

你的 `button` 元素應該有文字 `Switch Theme`。

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

你的 `button` 元素應該設定 `id` 屬性為 `"theme-switcher-button"`。

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

你的 `button` 元素應該有一個設定為 `"true"` 的 `aria-haspopup` 屬性。 

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

你的 `button` 元素應該有一個 `aria-expanded` 屬性設定為 `"false"`。 

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

你的 `button` 元素應該設定 `aria-controls` 屬性為 `"theme-dropdown"`。

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

你應該有一個 `ul` 元素。

```js
assert.exists(document.querySelector("ul"));
```

你的 `ul` 元素應該有一個 `id` 屬性設為 `"theme-dropdown"`。

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

你的 `ul` 元素應該有一個 `role` 屬性設定為 `"menu"`。

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

你的 `ul` 元素應該有一個 `aria-labelledby` 屬性，設定為 `"theme-switcher-button"`。

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

你的 `ul` 元素應該有一個 `hidden` 屬性。

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

你的 `ul` 元素應該至少有兩個 `li` 元素。

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

你的每個 `li` 項目都應該有一個 `role` 屬性設定為 `"menuitem"`。

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

你的每個 `li` 元素都應該有一段表述主題的文字。

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

你的每個 `li` 元素都應該有一個以 `theme-` 開頭並以該元素對應主題結尾的 `id` 屬性。確保你的 `id` 值全部為小寫。

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

你應該有一個元素，其 `aria-live` 屬性設為 `"polite"`。

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

你應該有一個 `themes` 陣列。

```js
assert.isArray(themes);
```

你的 `themes` 陣列應該至少有兩個物件。

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

你的 `themes` 陣列應該至少有兩個物件，每個物件都有 `name` 和 `message` 屬性。 

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

你的 `name` 和 `message` 屬性值應該是不為空的字串。

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

你的每個 `name` 屬性值都應該符合你為 `li` 元素提供的其中一個值。請確保這些值全部為小寫。

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

當使用者點擊 `#theme-switcher-button` 以顯示主題選項時，應該從 `#theme-dropdown` 元素中移除 `hidden` 屬性。

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

當 `#theme-dropdown` 元素顯示時，`button` 元素的 `aria-expanded` 屬性應設為 `"true"`。

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

當使用者在 `#theme-dropdown` 元素顯示時點擊 `#theme-switcher-button`，下拉選單應該被設回 `hidden`。

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

當使用者在 `#theme-dropdown` 元素顯示時點擊 `#theme-switcher-button`，`aria-expanded` 屬性應該設為 `"false"`。

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

當使用者點擊 `#theme-switcher-button` 並選擇主題時，應將對應的 `theme-<name>` 類別添加到 `body` 元素。

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1];
clickedItem?.dispatchEvent(new Event("click", {bubbles: true}));

const expectedClass = clickedItem?.id;
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

當使用者點擊 `#theme-switcher-button` 並選擇主題時，應該在 `aria-live="polite"` 元素中顯示來自 `themes` 陣列且與所選 `theme` 相關的 `message`。


```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", {bubbles: true}));

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>
  
  <div id="status" aria-live="polite"></div>
  
  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);
    
    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
