---
id: a3bfc1673c0526e06d3ac698
title: Jenga kalkuleta ya jumla ya namba msingi
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

Katika maabara hii, utajenga kalkuleta inayopokea namba na kurudisha jumla ya namba msingi zote ambazo ni ndogo au sawa na namba hiyo.

**Lengo**: Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi cha `sumPrimes` kinachokubali namba kama hoja.
1. Kitendakazi cha `sumPrimes` kinapaswa kurudisha jumla ya namba msingi zote ndogo au sawa na namba iliyotolewa.
1. Ikiwa namba ya ingizo ni ndogo kuliko `2`, kitendakazi kinapaswa kurudisha `0`.

# --hints--

Unapaswa kuwa na kitendakazi cha `sumPrimes`.

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` inapaswa kurudisha `17`.

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` inapaswa kurudisha `10`.

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` inapaswa kurudisha `2`.

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` inapaswa kurudisha `0`.

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` inapaswa kurudisha `73156`.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
