---
id: 6976db7ecfa770bf21307b20
title: Побудова движка ремікс-плейлиста
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

У цій лабораторній роботі ви створите програму, яка формує один ремікс-плейлист із кількох плейлистів, надісланих слухачами.

Кожен слухач надає список пісень, які він хоче почути. Деякі пісні можуть повторюватися, а деякі виконавці можуть з’являтися надто часто. Ваше завдання — крок за кроком обробити ці плейлисти: об’єднати їх в один список, оцінити кожну пісню, видалити дублікати, обмежити кількість появ одного виконавця, а потім створити фінальний порядок відтворення.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити функцію з назвою `flattenPlaylists`, яка приймає масив плейлистів, де кожен плейлист — це масив об’єктів із такими властивостями: `trackId`, `artist`, `title`, `votes`, `bpm`. Якщо вхідні дані не є масивом, `flattenPlaylists` має повертати порожній масив. Для вас надано приклад плейлиста, який можна використати для тестування функції.

2. `flattenPlaylists` має повертати плоский масив об’єктів треків, де кожен об’єкт містить усі оригінальні властивості треку плюс властивість `source`, що є масивом з індексом плейлиста та індексом треку, які вказують, звідки походить трек.

3. Ви повинні створити функцію з назвою `scoreTracks`, яка приймає масив об’єктів треків, як повертає `flattenPlaylists` (кожен із властивостями `trackId`, `artist`, `title`, `votes`, `bpm` та `source`) і повертає новий масив об’єктів треків, кожен із доданою властивістю `score`, обчисленою за формулою: `votes * 10 - Math.abs(bpm - 120)`.

4. Ви повинні створити функцію з назвою `dedupeTracks`, яка приймає масив об’єктів треків, як повертає `scoreTracks`, і повертає новий масив, у якому видалено дублікати за `trackId`, залишаючи лише першу появу кожного.

5. Ви повинні створити функцію з назвою `enforceArtistQuota`, яка приймає масив об’єктів треків, як повертає `dedupeTracks`, і число, що представляє максимальну дозволену кількість появ для кожного виконавця. Функція має повертати новий масив, у якому жоден виконавець не з’являється більше разів, ніж задано, залишаючи найраніші появи.

6. Ви повинні створити функцію з назвою `buildSchedule`, яка приймає масив об’єктів треків, як повертає `enforceArtistQuota`, і повертає новий масив об’єктів `{ slot, trackId }`, де `slot` — це індекс, що починається з 1, який представляє позицію кожного треку у порядку відтворення.

7. Ви повинні створити функцію з назвою `remixPlaylist`, яка приймає масив плейлистів і максимальну кількість дозволених появ для кожного виконавця. Функція має повертати фінальний розклад відтворення у вигляді масиву об’єктів `{ slot, trackId }`, викликаючи `flattenPlaylists`, `scoreTracks`, `dedupeTracks`, `enforceArtistQuota` і `buildSchedule` послідовно.

# --hints--

Ви повинні мати функцію з назвою `flattenPlaylists`.

```js
assert.isFunction(flattenPlaylists);
```

Ви повинні повертати порожній масив із `flattenPlaylists`, якщо вхідні дані не є масивом.

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

Кожен трек, який повертає `flattenPlaylists`, має містити поле `source`, яке є масивом із індексом плейлиста та індексом треку.

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

Ви повинні мати функцію з назвою `scoreTracks`.

```js
assert.isFunction(scoreTracks)
```

Кожен трек, який повертає `scoreTracks`, має містити числове поле `score`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

Ви повинні обчислити `score`, використовуючи цільовий BPM `120` і цю формулу: `votes * 10 - Math.abs(bpm - 120)`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

Ви повинні мати функцію з назвою `dedupeTracks`.

```js
assert.isFunction(dedupeTracks)
```

Якщо існують дублікати `trackId`, `dedupeTracks` має залишити лише першу появу треку.

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

Ви повинні мати функцію з назвою `enforceArtistQuota`.

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota` має гарантувати, що жоден виконавець не з’являється більше ніж `maxPerArtist` разів, видаляючи зайві треки та залишаючи найраніші.

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

Ви повинні мати функцію з назвою `buildSchedule`.

```js
assert.isFunction(buildSchedule);
```

`buildSchedule` має повертати масив об’єктів із формою `{ slot, trackId }`, де `slot` починається з `1`.

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

Ви повинні мати функцію з назвою `remixPlaylist`.

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist` має викликати допоміжні функції послідовно, щоб створити фінальний розклад.

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
