---
id: 689f7996426c1534fa5bea66
title: بناء مترجم نقاط الغولف
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

في لعبة الغولف، لكل حفرة قيمة `par`، وهي متوسط عدد `strokes` التي يُتوقع من لاعب الغولف أن يقوم بها لإدخال الكرة في الحفرة وإكمال اللعب. بناءً على مدى ارتفاع أو انخفاض `strokes` عن `par`، هناك لقب مختلف.

في هذا المختبر، ستكتب دالة تحول قيمة `par` و`strokes` إلى لقبها.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تمر لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `golfScore`.
1. يجب أن تأخذ `golfScore` معلمتين رقميتين، وهما قيمة par للمسار وعدد الضربات التي تم تنفيذها.
1. يجب أن تُرجع `golfScore` سلسلة نصية.
1. يجب أن تُرجع `golfScore` القيمة `"Hole-in-one!"` إذا كانت `strokes` تساوي `1`.
1. يجب أن تُرجع `golfScore` القيمة `"Eagle"` إذا كانت `strokes` أقل من أو تساوي `par` ناقص `2`.
1. يجب أن تُرجع `golfScore` القيمة `"Birdie"` إذا كانت `strokes` تساوي `par` ناقص `1`.
1. يجب أن تُرجع `golfScore` القيمة `"Par"` إذا كانت `strokes` تساوي `par`.
1. يجب أن تُرجع `golfScore` القيمة `"Bogey"` إذا كانت `strokes` تساوي `par` زائد `1`.
1. يجب أن تُرجع `golfScore` القيمة `"Double Bogey"` إذا كانت `strokes` تساوي `par` زائد `2`.
1. يجب أن تُرجع `golfScore` القيمة `"Go Home!"` إذا كانت `strokes` أكبر من أو تساوي `par` زائد `3`.

# --hints--

يجب أن تنشئ دالة باسم `golfScore`.

```js
assert.isFunction(golfScore);
```

يجب أن تأخذ `golfScore` معلمتين.

```js
assert.lengthOf(golfScore, 2);
```

يجب أن تُرجع `golfScore` سلسلة نصية

```js
assert.isString(golfScore(4,1)); 
```

يجب أن تُرجع `golfScore(1, 1)` السلسلة النصية `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

يجب أن تُرجع `golfScore(3, 1)` السلسلة النصية `Hole-in-one!`

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

يجب أن تُرجع `golfScore(4, 1)` السلسلة النصية `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

يجب أن تُرجع `golfScore(5, 1)` السلسلة النصية `Hole-in-one!`

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

يجب أن تُرجع `golfScore(4, 2)` السلسلة النصية `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

يجب أن تُرجع `golfScore(5, 2)` السلسلة النصية `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

يجب أن تُرجع `golfScore(3, 2)` السلسلة النصية `Birdie`

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

يجب أن تُرجع `golfScore(4, 3)` السلسلة النصية `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

يجب أن تُرجع `golfScore(5, 4)` السلسلة النصية `Birdie`

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

يجب أن تُرجع `golfScore(3, 3)` السلسلة النصية `Par`

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

يجب أن تُرجع `golfScore(4, 4)` السلسلة النصية `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

يجب أن تُرجع `golfScore(5, 5)` السلسلة النصية `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

يجب أن تُرجع `golfScore(3, 4)` السلسلة النصية `Bogey`

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

يجب أن تُرجع `golfScore(4, 5)` السلسلة النصية `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

يجب أن تُرجع `golfScore(5, 6)` السلسلة النصية `Bogey`

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

يجب أن تُرجع `golfScore(3, 5)` السلسلة النصية `Double Bogey`

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

يجب أن تُرجع `golfScore(4, 6)` السلسلة النصية `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

يجب أن تُرجع `golfScore(5, 7)` السلسلة النصية `Double Bogey`

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

يجب أن تُرجع `golfScore(3, 7)` السلسلة النصية `Go Home!`

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

يجب أن تُرجع `golfScore(4, 8)` السلسلة النصية `Go Home!`

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

يجب أن تُرجع `golfScore(5, 9)` السلسلة النصية `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
