---
id: 6995d0325fd7da40854fac40
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En una lección anterior, aprendiste sobre la búsqueda binaria y cómo se diferencia de la búsqueda lineal.

En este taller, pondrás eso en práctica construyendo un algoritmo de búsqueda binaria.

Para comenzar, crea una función `binarySearch` con parámetros `searchList` y `value`.

`searchList` es el arreglo de elementos dentro del cual la función buscará, y `value` es el objetivo que la función intentará encontrar dentro del `searchList`.

# --hints--

Deberías crear una función `binarySearch`.

```js
assert.isFunction(binarySearch);
```

Tu función `binarySearch` debe tener parámetros `searchList` y `value`.

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
