---
id: 6996ca21a74104721f2c7e5d
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Necesitas probar las cosas para que entiendas el flujo del algoritmo en esta etapa inicial.

Para hacer eso, primero necesitas salir del bucle. Eso es porque la implementación actual solo permitirá una iteración, así que si la condición no se cumple, habrá un bucle infinito.

Justo después de la instrucción `if`, usa la palabra clave `break` para salir del ciclo `while`. Luego, después del ciclo `while`, devuelve un arreglo vacío para indicar que no se encontró el valor.

# --hints--

Deberías usar la palabra clave `break` para salir del bucle `while` después de la declaración `if`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

Debes devolver un arreglo vacío después de tu ciclo `while`.

```js
assert.deepEqual(binarySearch([], NaN), []);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
