---
id: 6996d1032feed88e2957dc8e
title: Paso 15
challengeType: 1
dashedName: step-15
---

# --description--

Puedes ver que la última llamada imprime un arreglo vacío. Para especificar mejor que el valor no está en el arreglo buscado, actualiza la sentencia `return` para que devuelva un arreglo que contenga dos elementos: un arreglo vacío y la cadena `Value not found`.

# --hints--

Debes actualizar la sentencia `return` para que la función devuelva un arreglo con `[]` como primer elemento y `"Value not found"` como segundo elemento.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const validExpressions = [
  "[[], 'Value not found']",
  "[[], \"Value not found\"]",
  "[[], `Value not found`]"
];
assert.isTrue(validExpressions.some(expr => binarySearch?.hasReturn(expr) || binarySearch?.hasReturn("(" + expr + ")")));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
  return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
