---
id: 6996c60281feca6718cc15c1
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Dentro de tu función `binarySearch`, crea una variable `pathToTarget` y asígnale un arreglo vacío, para empezar.

# --hints--

Deberías tener una variable `pathToTarget` asignada a un arreglo vacío `[]` dentro de tu función `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
