---
id: 6996c95e5bede03bdf2937a2
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Lo siguiente es crear una condición que verifique si el `value` objetivo está en el medio.

Dentro del ciclo `while`, crea una sentencia `if` que compruebe si el `value` objetivo es igual a `valueAtMiddle`. Si lo es, devuelve la variable `pathToTarget`.

# --hints--

Debes crear una sentencia `if` con la condición `value === valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

Debes devolver la variable `pathToTarget` dentro de la sentencia `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
