---
id: 6996ccadf9048c79b857bb15
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

Puedes ver que la primera llamada devuelve `[3]`. Eso es porque `3` es el valor medio en `[1, 2, 3, 4, 5]`.

Por otro lado, la segunda llamada devuelve un arreglo vacío. Esto sucede porque, después de verificar el valor medio, el ciclo actualmente no tiene la lógica para reducir el rango de búsqueda actualizando las variables `low` o `high`. Solo actúa como una verificación única del punto medio y luego termina sin buscar más.

Para permitir que la búsqueda binaria siga reduciendo su búsqueda, agrega una sentencia `else if` que verifique si `value` es mayor que `valueAtMiddle`.

# --hints--

Debes agregar una sentencia `else if` con la condición `value > valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    if (value === valueAtMiddle) {
      return pathToTarget;
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
