---
id: 6996c7363a4a437839f1f15d
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Ahora que has definido los límites de tu búsqueda, es momento de crear el ciclo que ejecutará el algoritmo de búsqueda binaria.

La búsqueda binaria funciona reduciendo repetidamente el espacio de búsqueda. Este proceso continúa mientras haya un rango válido de elementos para revisar.

Puedes expresarlo con un ciclo `while` que continuará mientras tu puntero `low` sea menor o igual que tu puntero `high`.

Si `low` alguna vez se vuelve mayor que `high`, significa que el espacio de búsqueda se ha vaciado y que `value` no está en el arreglo.

Dentro de la función `binarySearch`, crea un ciclo `while` con una condición que verifique si `low` es menor o igual que high.

# --hints--

Debes crear un ciclo `while` que se ejecute mientras `low` sea menor o igual que `high`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
