---
id: 6996c8f94fe6b250f11a2a1e
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Ahora, incluye `valueAtMiddle` en el arreglo `pathToTarget` para registrar los pasos realizados durante la búsqueda, sin importar si es el objetivo. Puedes hacer esto con el método `push()`.

# --hints--

Debes usar el método `push()` para agregar `valueAtMiddle` al arreglo `pathToTarget`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
