---
id: 6996c675b599baa7b103192d
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

En la búsqueda binaria, existe el índice más bajo posible y el índice más alto posible.

El índice más bajo posible representa el límite más a la izquierda en la porción actual de los valores que se están buscando, en este caso, `searchList`. El índice más alto posible, por otro lado, representa el límite más a la derecha en los valores que se están buscando.

Para tener en cuenta esos dos, debes considerar todo el arreglo que se está buscando. Entonces, dentro de la función `binarySearch`, define una variable `low` con un valor de `0`, y una variable `high` con un valor que corresponda al último índice del arreglo que se está buscando.

# --hints--

Debes tener una variable llamada `low` asignada a `0` en tu función `binarySearch`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

Debes tener una variable llamada `high` asignada a `searchList.length - 1` en tu función `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
