---
id: 6996c8375f4301004459ee85
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Ahora que has calculado el índice `mid`, necesitas obtener el valor del `searchList` en ese índice. Este valor es el que compararás con tu parámetro `value`, que es el objetivo que estás buscando.

Crea una variable `valueAtMiddle` y asígnale el elemento de `searchList` que se encuentra en el índice `mid`.

# --hints--

Debes crear una variable `valueAtMiddle` asignada a `searchList[mid]`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
