---
id: 6996d18b31aab5cbbffbc731
title: Paso 16
challengeType: 1
dashedName: step-16
---

# --description--

Una última pequeña mejora que puedes añadir es mostrar también el índice en el cual se encuentra el `value` objetivo.

Entonces, en lugar de devolver solo `pathToTarget` en la sentencia `if`, devuelve un arreglo que contenga `pathToTarget` y un literal de plantilla con el mensaje `Value found at index ${mid}`.

¡Con eso, tu taller de algoritmo de búsqueda binaria está completo!

# --hints--

Debes actualizar la sentencia `return` en el `if` para que la función devuelva un arreglo con `pathToTarget` como el primer elemento y un literal de plantilla para `Value found at index ${mid}` como el segundo elemento. Asegúrate de usar backticks para el literal de plantilla.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch: binarySearchFn } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearchFn?.toString());
assert.match(
  codeWithoutComments,
  /return\s*\[\s*pathToTarget\s*,\s*(?:`Value found at index\s*\$\{\s*mid\s*\}`|['"]Value found at index\s*['"]\s*\.\s*concat\(\s*mid\s*\))\s*\]/
);
```

Tu función debe devolver los resultados esperados.

```js
assert.deepEqual(
  binarySearch([1, 2, 3, 4, 5], 3), [[3], "Value found at index 2"]
);
assert.deepEqual(
  binarySearch([1, 2, 3, 4, 5, 9], 4), [[3, 5, 4], "Value found at index 3"]
);
assert.deepEqual(
  binarySearch([1, 3, 5, 9, 14, 22], 10), [[], "Value not found"]
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
--fcc-editable-region--
      return pathToTarget;
--fcc-editable-region--
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```

# --solutions--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return [pathToTarget, `Value found at index ${mid}`];
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
