---
id: 6996c60281feca6718cc15c1
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Dentro de tu función `binarySearch`, crea una variable `pathToTarget` y asígnale un arreglo vacío, para empezar.

# --hints--

Debes declarar una variable `pathToTarget` con un arreglo vacío dentro de tu función `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
