---
id: 6996cfb6c01e9a802add38be
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

Finalmente, si el `value` no se encuentra en el medio y tampoco es mayor que `valueAtMiddle`, entonces el `value` debe ser menor que `valueAtMiddle` y debe estar a la izquierda. El bloque final `else` se encargará de esto.

Así que, termina el ciclo eliminando la palabra clave `break` y agregando un bloque `else`. Dentro de este, actualiza la variable `high` restando `1` de `mid`.

# --hints--

Debes reemplazar la palabra clave `break` con un bloque `else` que actualice la variable `high` a `mid - 1`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    }

--fcc-editable-region--      
    break;
--fcc-editable-region--

  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
