---
id: 6996ca21a74104721f2c7e5d
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Necesitas probar cosas para entender el flujo del algoritmo en esta etapa inicial.

Para eso, primero debes salir del ciclo. Eso es porque la implementación actual solo permite una iteración, así que si la condición no se cumple, habrá un ciclo infinito.

Justo después de la instrucción `if`, usa la palabra clave `break` para salir del ciclo `while`. Luego, después del ciclo `while`, devuelve un arreglo vacío para indicar que no se encontró el valor.

# --hints--

Debes usar la palabra clave `break` para salir del ciclo `while` después de la instrucción `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

Debes devolver un arreglo vacío después de tu ciclo `while`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
