---
id: 6996d0a2c2ac55041145a13f
title: Paso 14
challengeType: 1
dashedName: step-14
---

# --description--

¿Notaste que la segunda llamada de función ahora tiene los valores `3, 5, 4` en la lista? Eso indica que la búsqueda binaria está funcionando como se esperaba.

Así es como sucedió: El algoritmo primero verificó `3` como el medio de la lista inicial. Dado que `4` es mayor que `3`, la búsqueda se desplazó hacia la mitad derecha. Luego examinó `5` como el nuevo medio. Debido a que `4` es menor que `5`, la búsqueda se movió a la izquierda, identificando finalmente `4` como el medio del rango final.

Para probar la función de nuevo, llámala con `[1, 3, 5, 9, 14, 22], 10` y regístrala en la consola. Esta es una situación en la que el valor no se encontrará.

# --hints--

Debes llamar a la función con `binarySearch([1, 3, 5, 9, 14, 22], 10)` y registrarla en la consola.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,3,5,9,14,22],10))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
