---
id: 6996c675b599baa7b103192d
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

En la búsqueda binaria, hay el índice más bajo posible y el índice más alto posible.

El índice más bajo posible representa el límite más a la izquierda en la porción actual de los valores que se están buscando, en este caso, `searchList`. El índice más alto posible, por otro lado, representa el límite más a la derecha en los valores que se están buscando.

Para tener en cuenta esos dos, debes considerar todo el arreglo que se está buscando. Así que, dentro de la función `binarySearch`, declara una variable `low` con un valor de `0` y una variable `high` con un valor que represente el último índice del arreglo que se está buscando.

# --hints--

Debes declarar una variable llamada `low` con valor `0` en tu función `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

Debes declarar una variable llamada `high` con valor `searchList.length - 1` en tu función `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
