---
id: 6996c796385b6ea4e222d5e8
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Para comenzar con el ciclo, necesitas encontrar el índice medio del espacio de búsqueda actual.

Puedes obtener el índice medio calculando el promedio de las variables `low` y `high`. Tendrás que usar división entera para obtener el promedio después de sumar los dos valores, de modo que la respuesta siempre sea un número entero redondeado hacia abajo, ya que los índices deben ser enteros.

Dentro de tu ciclo `while`, crea una variable `mid` establecida como el promedio de `low` y `high`.

# --hints--

Dentro del ciclo while deberías crear una variable `mid` establecida como `Math.floor((low + high) / 2)`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
--fcc-editable-region--
      
--fcc-editable-region--
  }
}
```
