---
id: 6996cb5869acf58de751ae81
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Recuerda que la función recibe un arreglo llamado `searchList` y el `value` que estás buscando.

Ahora, llama a la función con `binarySearch([1, 2, 3, 4, 5], 3)` y `binarySearch([1, 2, 3, 4, 5, 9], 4)` y muestra ambas declaraciones en la consola.

# --hints--

Debes llamar a la función con `binarySearch([1, 2, 3, 4, 5], 3)` y mostrarlo en la consola.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*\]\s*,\s*3\s*\)\s*\)/
);
```

Debes llamar a la función con `binarySearch([1, 2, 3, 4, 5, 9], 4)` y mostrarlo en la consola.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*,\s*9\s*\]\s*,\s*4\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--  

--fcc-editable-region--  
```
