---
id: 641da5462576784453146ec2
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

欲しいデータが取得できたので、このデータを使って UI を表示していきましょう。 しかし、取得したデータには 26 人の著者データの配列が含まれています。これを一度に画面に表示しようとすると、パフォーマンスが低下する可能性があります。

代わりに、一度に 8 人ずつ表示するようにしましょう。また、次に表示するデータがある限り、次の 8 人を読み込むボタンを配置します。

`let` を使用して、`startingIndex` と `endingIndex` という 2 つの変数を作成してください。そして、それぞれ数値の `0` と `8` を代入してください。 また、`authorDataArr` という変数を `let` で作成し、空の配列を代入してください。

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'http://not-a-real-url.nowhere/no-image.jpg', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

`startingIndex` という変数の宣言には `let` を使用してください。

```js
assert.match(code, /let\s+startingIndex/)
```

変数 `startingIndex` を `0` に設定してください。

```js
assert.match(code, /let\s+startingIndex\s*=\s*0\s*;?/)
```

`endingIndex` という変数の宣言には `let` を使用してください。

```js
assert.match(code, /let\s+endingIndex/)
```

変数 `endingIndex` を `8` に設定してください。

```js
assert.match(code, /let\s+endingIndex\s*=\s*8\s*;?/)
```

`authorDataArr` という変数の宣言には `let` を使用してください。

```js
assert.match(code, /let\s+authorDataArr/)
```

変数 `authorDataArr` を空の配列 (`[]`) に設定してください。

```js
assert.match(code, /let\s+authorDataArr\s*=\s*\[\s*\]\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
