---
id: 5dc17d3bf86c76b9248c6eb4
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

The `p` element is used to create a paragraph of text on websites. Create a `p` element below your `h2` element and give it the following text:

`Everyone loves cute cats online!`

# --hints--

你的 `p` 元素应该有一个开始标签。 起始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('p'));
```

你的 `p` 元素应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/p\>/));
```

Your `p` element's text should be `Everyone loves cute cats online!` You have either omitted the text or have a typo.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

你的 `p` 元素应该在 `h2` 元素的下方。 你把顺序写错了。

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

