---
id: 5ddbd81294d8ddc1510a8e56
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

A link's text must be placed between the opening and closing tags of an anchor (`a`) element.

这是一个链接文本是 `click here to go to freeCodeCamp.org` 的示例：

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Add the anchor text `cat photos` to the anchor element. This will become the link's text.

# --hints--

你的锚点元素（`a`）应该有一个开始标签。 开始标签的语法是：`<elementName>`。

```js
assert(document.querySelector('a'));
```

你的锚元素（`a`）应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/a\>/));
```

Your anchor (`a`) element's text should be `cat photos`. Make sure to put the link text between the anchor (`a`) element's opening tag and closing tag.

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

