---
id: 5dfa2407b521be39a3de7be1
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

To open links in a new tab, you can use the `target` attribute on the anchor (`a`) element.

`target` 属性指定了在哪里打开链接文档。 `target="_blank"` 意思是在新标签页或窗口中打开链接的文档。

这是一个带有 `target` 属性的 `a` 元素的基本语法：

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Add a `target` attribute with the value `_blank` to the anchor (`a`) element's opening tag, so that the link opens in a new tab.

# --hints--

Your `p` element should have a nested anchor (`a`) element with the text `cat photos`. You may have deleted it or have a typo.

```js
const anchor = document.querySelectorAll('p > a');
assert(
  anchor.length &&
    anchor[0].innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Your anchor (`a`) element does not have a `target` attribute. 请检查在开始标签的名称后面要有一个空格，且/或所有的属性名称前面也要有一个空格。

```js
assert(document.querySelector('a').hasAttribute('target'));
```

The value of the `target` attribute should be `_blank`. 你可能落下了这个值或者有拼写错误。 所有属性的值都应该放在引号中。

```js
assert(document.querySelector('a').getAttribute('target') === '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

