---
id: 5ef9b03c81a63668521804d4
title: 步骤 34
challengeType: 0
dashedName: step-34
---

# --description--

`strong` 元素用于指示某些文本非常重要或紧急。

在你刚刚添加的 `figcaption` 中，通过将 `hate` 放在 `strong` 元素中来表明它非常重要。

# --hints--

你的 `strong` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('strong'));
```

你的 `strong` 元素应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/strong\>/));
```

你的 `strong` 元素应在文本 `Cats hate other cats.` 中将 `hate` 一词包裹起来。你要么忽略了文本，要么有错字。

```js
assert(
  document
    .querySelectorAll('figcaption')[1]
    .querySelector('strong')
    .innerText.trim().toLowerCase() === 'hate'
);
```

`figcaption` 的文本应该是 `Cats hate other cats.`，检查是否有错别字以及在 `strong` 元素的开头和结尾标签处是否存在必要的空格。

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert(
  secondFigCaption &&
    secondFigCaption.innerText
      .replace(/\s+/gi, ' ')
      .trim()
      .match(/cats hate other cats\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
--fcc-editable-region--

          <figcaption>Cats hate other cats.</figcaption>  

--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

