---
id: 5f07c98cdb9413cbd4b16750
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

It is time to add a new section. Add a second `section` element below the existing `section` element.

# --hints--

你的 `section` 元素应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert(document.querySelectorAll('section').length >= 2);
```

你应该只添加一个 `section` 开始标签。 请删除多余的。

```js
assert(document.querySelectorAll('section').length === 2);
```

你的 `section` 元素应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/section>/g).length >= 2);
```

你应该只添加一个 `section` 结束标签。 请删除多余的。

```js
assert(code.match(/<\/section>/g).length === 2);
```

第二个 `section` 元素不应嵌套在第一个 `section` 元素中。

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

两个 `section` 元素都应该在 `main` 元素的开始标签和结束标签之间。

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

