---
id: 61329d52e5010e08d9b9d66b
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

عنصر `meta` مهم آخر لـ إمكانية الوصول وتحسين محركات البحث هو تعريف `description`. تُستخدم قيمة الخاصية `content` من قبل محركات البحث لتوفير وصف لصفحتك.

أضف عنصر `meta` مع خاصية `name` مضبوطة على `description`، وامنحه خاصية `content` مفيدة.

# --hints--

يجب أن تضيف عنصر `meta` جديد إلى قسم `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

يجب أن تعطي عنصر `meta` خاصية `name` بقيمة `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

يجب أن تعطي عنصر `meta` خاصية `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

يجب ألا تتجاوز قيمة خاصية `content` 165 حرفًا. _هذا هو الحد الأقصى لطول الوصف حسب جوجل._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
