---
id: 61329b210dac0b08047fd6ab
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

بالاستمرار مع عناصر `meta`، تعريف `viewport` يُخبر المتصفح كيف يعرض الصفحة. تضمين واحد يحسن إمكانية الوصول البصري على الأجهزة المحمولة، ويُحسّن تحسين محركات البحث (_SEO_).

أضف تعريف `viewport` مع خاصية `content` تُفصل `width` و`initial-scale` للصفحة.

# --hints--

يجب إنشاء عنصر `meta` آخر داخل `head`.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

يجب إعطاء عنصر `meta` خاصية `name` بقيمة `viewport`.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

يجب إعطاء عنصر `meta` خاصية `content` بقيمة `width=device-width, initial-scale=1`.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
