---
id: 613e275749ebd008e74bb62e
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

حاليًا، يفترض عنصر `img` حجمه الافتراضي، وهو كبير جدًا. يحتوي CSS على دالة `max` التي تُرجع أكبر قيمة من مجموعة قيم مفصولة بفواصل. على سبيل المثال:

```css
img {
  width: max(250px, 25vw);
}
```

في المثال أعلاه، سيكون عرض الصورة 250px إذا كان عرض إطار العرض أقل من 1000 بكسل. إذا كان عرض إطار العرض أكبر من 1000 بكسل، سيكون عرض الصورة 25vw. ذلك لأن 25vw تساوي 25% من عرض إطار العرض.

قم بتغيير مقياس الصورة باستخدام `id` الخاص بها كمحدد، واضبط `width` ليكون أكبر قيمة بين `10rem` أو `18vw`.

# --hints--

يجب استخدام المحدد `#logo` لاستهداف عنصر `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

يجب أن تعطي عنصر الـ `img` عرضًا (`width`) بقيمة `max(10rem, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
