---
id: 61408f155e798909b6908712
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

لتمكين التنقل في الصفحة، أضف قائمة غير مرتبة تحتوي على ثلاثة عناصر قائمة التالية:

- `INFO`
- `HTML`
- `CSS`

يجب أن يُغلف نص عناصر القائمة بوسوم الرابط أو وسم الرابط للنص العالي.

# --hints--

يجب أن تُسكّن عنصر `ul` واحد داخل عنصر `nav`.

```js
assert.equal(document.querySelectorAll('nav > ul')?.length, 1);
```

يجب أن تُسكّن ثلاثة عناصر `li` داخل عنصر `ul`.

```js
assert.equal(document.querySelectorAll('nav > ul > li')?.length, 3);
```

يجب أن تُسكّن عنصر `a` واحد داخل كل عنصر `li`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.length, 3);
```

يجب أن يحتوي أول عنصر `a` على النص `INFO`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[0]?.textContent, 'INFO');
```

يجب أن يحتوي ثاني عنصر `a` على النص `HTML`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[1]?.textContent, 'HTML');
```

يجب أن يحتوي ثالث عنصر `a` على النص `CSS`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[2]?.textContent, 'CSS');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
--fcc-editable-region--
      <nav>

      </nav>
--fcc-editable-region--
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}
```
