---
id: 614ccc21ea91ef1736b9b578
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

مرحبًا بك في الجزء الأول من اختبار إمكانية الوصول. بما أنك أصبحت مطوِّرًا متمرسًا في HTML وCSS، بدأنا معك بالقالب الجاهز الأساسي.

ابدأ رحلة إمكانية الوصول هذه بإضافة خاصية `lang` إلى عنصر `html` الخاص بك. هذا يساعد قارئات الشاشة على تحديد لغة الصفحة.

# --hints--

يجب أن تضيف خاصية `lang` إلى عنصر `html`. _تلميح: يمكنك استخدام القيمة `en` للغة الإنجليزية._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
