---
id: 6690e10ebe2181212abc9652
title: Crie um cartão de visita
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve vincular a folha de estilo `styles.css` utilizando a tag `link` na seção `head` do arquivo HTML.
1. O plano de fundo da sua página deve ser `rosybrown`. A fonte geral deve ser `Arial`, com uma alternativa de `sans-serif`.   
1. Todos os elementos `p` devem ter uma `margin` superior e inferior de `5px`.
1. Todos os links na página não devem estar sublinhados.
1. Você deve ter um `div` com um atributo `class` com o valor `business-card` que conterá todos os outros elementos. 
1. O seletor `.business-card` deve ter propriedades para definir o elemento com largura de `300px` e uma cor de fundo de sua escolha. Também deve-se definir um `padding` de `20px` em todos os lados e uma `margin` superior de `100px`. O texto deve estar alinhado ao centro e o tamanho da fonte deve ser `16px`. As margens esquerda e direita devem ser definidas como `auto`.
1. Dentro do elemento `.business-card`, deve haver um elemento `img` com um atributo `class` com o valor `profile-image`. Você pode definir a fonte da imagem para `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg`, se desejar. Deve haver um atributo `alt` com uma descrição significativa.
1. O seletor `.profile-image` deve ter uma propriedade `max-width` com o valor `100%`.
1. Dentro do elemento `.business-card`, após o elemento `img`, você deve ter três elementos `p` com um atributo `class` que possui, respectivamente, os valores `full-name`, `designation` e `company`.
1. O primeiro elemento `p` deve conter seu nome.
1. O segundo elemento `p` deve conter sua designação.
1. O terceiro elemento `p` deve ter o nome da sua empresa. 
1. Deve haver um elemento `hr` abaixo do terceiro elemento `p`.
1. Após o elemento `hr`, devem haver dois elementos `p`. No primeiro elemento `p`, você deve ter um endereço de e-mail como texto. No segundo elemento `p`, o texto deve ser um número de telefone.
1. Após os dois elementos `p`, deve haver um elemento `a` com uma classe `portfolio-link`, o texto `Portfolio` e ele deve linkar para uma URL válida.
1. Deve haver um elemento `hr` após o elemento `a` contendo o link do portfólio.
1. Você deve ter outro elemento `div` com um atributo `class` com o valor `social-media`. Dentro deste elemento, deve haver um elemento `h2` com o texto `Conecte-se comigo`. 
1. Dentro dos elementos `.social-media`, devem haver três elementos `a`. Os elementos `a` devem conter os textos `Twitter`, `LinkedIn` e `GitHub`, respectivamente, com links para sites válidos.

# --hints--

O elemento `link` deve estar dentro do elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

Você deve definir a cor de fundo da página no seletor `body` para `rosybrown`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

Você deve definir a fonte da página como `Arial` no elemento `body` com uma alternativa de `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```


Seus elementos `p` devem ter uma `margin` superior e inferior de `5px`.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

Os links da página não devem ter sublinhado.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

Você deve ter um `div` principal com um atributo `class` com o valor `business-card`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

Seu seletor `.business-card` deve ter uma propriedade `width` com o valor `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

Seu seletor `.business-card` deve ter uma propriedade `background-color` com um valor de cor válido.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

Seu seletor `.business-card` deve ter uma propriedade `padding` com o valor `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

Seu seletor `.business-card` deve ter uma margem superior de `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

O texto dentro do elemento `.business-card` deve estar alinhado ao centro.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

O tamanho da fonte do texto dentro do elemento `.business-card` deve ser `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

As margens esquerda e direita do elemento `.business-card` devem ser definidas como `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

Dentro do elemento `.business-card`, deve haver um elemento `img` com um atributo `class` com o valor `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

A fonte da imagem deve ser definida para uma imagem válida.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

O atributo `alt` da imagem deve ser definido para uma descrição significativa.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

Seu seletor `.profile-image` deve ter uma propriedade `max-width` com o valor `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

Dentro do elemento `.business-card`, após o elemento `img`, deve haver um elemento `p` com um atributo `class` com o valor `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

Após o elemento `.full-name`, deve haver um elemento `p` com um atributo `class` com o valor `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

Após o elemento `.designation`, deve haver um elemento `p` com um atributo `class` com o valor `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

O primeiro elemento `p`, aquele com a classe `full-name`, deve conter seu nome.

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

O segundo elemento `p`, aquele com um atributo `class` de `designation`, deve conter sua designação.

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

O terceiro elemento `p`, aquele com o atributo `class` com valor `company`, deve conter o nome da sua empresa.

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

Após o terceiro elemento `p`, aquele com o atributo `class` com valor `company`, deve haver um elemento `hr`.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Após o primeiro elemento `hr`, deve haver um elemento `p` com um endereço de e-mail como seu texto.

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

Após o elemento `p` do e-mail, deve haver outro elemento `p` com um número de telefone como seu texto.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Após o elemento de número de telefone `p`, deve haver um elemento `a` com a classe `portfolio-link`. 

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

O elemento `.portfolio-link` deve ter o texto `Portfolio`.

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

Seu elemento `.portfolio-link` deve ter um `href` válido.

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

Após o elemento `a` contendo seu portfólio, deve haver outro elemento `hr` como um divisor.

```js
assert.exists(document.querySelector('a + hr'));
```

Após o segundo elemento `hr`, deve haver um elemento `div` com um atributo `class` com o valor `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

Dentro do elemento `.social-media`, deve haver um elemento `h2` com o texto `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

Dentro do elemento `.social-media` devem existir três elementos `a`.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

O primeiro elemento `a` deve conter o texto `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

O `href` do primeiro elemento `a` deve apontar para um link válido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

O segundo elemento `a` deve conter o texto `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

O `href` do segundo elemento `a` deve apontar para um link válido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

O terceiro elemento `a` deve conter o texto `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

O `href` do terceiro elemento `a` deve apontar para um link válido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
