---
id: 65f569725359e10d345bc52a
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

ستبني برنامجًا يحسب ويرسم مسار جسم مقذوف بناءً على زاوية ورمية وسرعة الارتقاء.

ابدأ بـاستيراد `math`، ستستخدمه كثيرًا في هذا المشروع لأنه يحتوي على طرق مفيدة مثل `math.radians` و`math.cos` و`math.sin` وغيرها.

أنشئ أيضًا هذه المتغيرات لتحمل قيمة تسارع الجاذبية وبعض الرموز الخاصة التي ستفيدك لاحقًا (انسخها والصقها).

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

استخدم `import math` لاستيراد وحدة `math`.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

يجب أن يكون لديك متغير `GRAVITATIONAL_ACCELERATION` بقيمة `9.81`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

يجب أن يكون لديك متغير `PROJECTILE` بقيمة `"∙"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

يجب أن يكون لديك متغير `x_axis_tick` بقيمة `"T"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

يجب أن يكون لديك متغير `y_axis_tick` بقيمة `"⊣"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
