---
id: 65fd4b3bedee044b4b957d33
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

أنشئ صنفًا باسم `Projectile` مع دالة `__init__` تهيئ الصنف باستخدام ثلاث معلمات: السرعة الابتدائية، والارتفاع الابتدائي، وزاوية رمي المقذوف، بهذا الترتيب.

داخل دالة `__init__`، عيّن هذه المعلمات إلى ثلاث خصائص خاصة: `__speed`، و`__height`، و`__angle`. الزاوية المعطاة ستكون بالدرجات؛ مع ذلك، يجب تخزينها داخليًا بالراديان. لتحقيق ذلك، استخدم الدالة `math.radians()` لتحويل الزاوية من درجات إلى راديان عند تعيينها إلى `__angle`.

استخدام شرطتين سفليتين قبل اسم الخاصية يُفعّل تشويش الأسماء في بايثون. هذا يعني أن الخصائص لا يمكن الوصول إليها مباشرة من خارج الصنف باستخدام أسمائها الأصلية، ويجب الوصول إليها بالأسماء المشوشة مثل `ball._Projectile__height` إذا دعت الحاجة من الخارج، مما يشير إلى أن هذه الخصائص مخصصة للاستخدام الداخلي فقط.

# --hints--

يجب أن تستخدم `class Projectile:` لإنشاء الصنف.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

يجب أن تحتوي صنف `Projectile` على دالة `__init__` بداخله.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

يجب أن تحتوي نسخة جديدة من صنف `Projectile` على قيم الخصائص الثلاثة مهيأة بشكل صحيح.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
