---
id: 65fd4bd84561a14d3e720061
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

لمتغير الصنف `__slots__` استخدام خاص في أصناف Python. إعلان `__slots__` وتعيينه متتالية من سلاسل نصية يقيّد إنشاء الخصائص لتلك المدرجة في تلك المتتالية. كما يمنع إنشاء الخاصية الخاصة `__dict__` ويسمح بوصول أكثر كفاءة إلى الخصائص.

يجب استخدام المتغير `__slots__` داخل الصنف لتعريف الخصائص التي يحتويها الصنف: عيّن إلى `__slots__` صف مرتب يحتوي على 3 سلاسل نصية، كل واحدة تساوي اسم خاصية من الخصائص المعرفة في الدالة `__init__`.

# --hints--

يجب تعيين إلى `__slots__` صف مرتب يحتوي على `'__height'` و `'__speed'` و `'__angle'`.

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
