---
id: 660400f02cf3a76795b82d4e
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

أول شيء يجب إعداده هو طريقة تحسب الإزاحة للمقذوف، وهي المسافة الأفقية التي يقطعها من وقت الرمي حتى يلمس المقذوف الأرض.

أنشئ طريقة `__calculate_displacement` التي تحتوي فقط على `self` كمعلمة، وأرجع الإزاحة للمقذوف.

استخدم الصيغة التالية لحساب إزاحة المقذوف:
\\[ d = \frac{v \cdot \cos\theta \cdot \left(v\cdot\sin\theta + \sqrt{v^2 \cdot \sin^2\theta + 2 \cdot g \cdot h}\right)}{g} \\]

حيث $d$ هي الإزاحة، و$v$ هي السرعة الابتدائية، و$\theta$ هي الزاوية، و$h$ هي ارتفاع بداية المقذوف.
بالنسبة لـ $g$ يمكنك استخدام المتغير `GRAVITATIONAL_ACCELERATION`.

يجب استخدام الطرق `math.cos()` و`math.sin()` للدوال المثلثية و`math.sqrt()` لحساب الجذر التربيعي. كما يجب أن تعرف أن $x^y$ يُكتب في بايثون كـ `x ** y`. أيضًا $\sin^2\theta$ تعني أن القيمة الناتجة من الجيب تُربع بعد ذلك.

تذكر أنه مع تشويش الأسماء تحتاج لاستدعاء الطريقة كـ `_Projectile__calculate_displacement` إذا أردت اختبارها أو استخدامها من خارج الصنف:

```py
ball = Projectile(10, 3, 45)
displacement_of_ball = ball._Projectile__calculate_displacement() # 12.6173996009878
```

# --hints--

يجب إعلان طريقة باسم `__calculate_displacement` باستخدام `def __calculate_displacement(self):`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").has_function("__calculate_displacement")`))})
```

يجب أن تحتوي طريقة `__calculate_displacement` على المعلمة `self` فقط.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_args('self')`))})
```

يجب أن تُرجع طريقة `__calculate_displacement` القيمة الصحيحة.

```js
({test: () => assert(runPython(`
p = Projectile(20, 21, 22)
disp = p._Projectile__calculate_displacement()
round(disp, 2) == 55.06 and round(disp, 2) != disp
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    
--fcc-editable-region--
```
