---
id: 660fcf3e1b9bb056b2edb567
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

في هذه المرحلة أنت جاهز لإنشاء التمثيل النصي.
ابدأ بإنشاء نسخة من صنف `Projectile`. أسفل تعريف الصنف، أنشئ متغيرًا باسم `ball` وعيّنه باستدعاء `Projectile` مع المعلمات `10, 3, 45`.

عرّف طريقة `__str__` داخل الصنف، لتعطي نسختك تمثيلًا نصيًا مناسبًا، واجعلها تُرجع سلسلة نصية بالتنسيق التالي:

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

يجب أن تبدأ وتنتهي بحرف سطر جديد، ويجب كتابة الزاوية كعدد صحيح بالدرجات، ويُطبع الإزاحة برقم عشري واحد.

ستجد `math.degrees` مفيدًا لتحويل الزاوية من راديان إلى درجات.

عندما تكون جاهزًا، يمكنك `print(ball)` لاختبار دالتك.

# --hints--

يجب أن يكون لديك متغير `ball`.

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

يجب أن تكون قيمة المتغير `ball` هي `Projectile(10, 3, 45)`.

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

يجب أن يكون التمثيل النصي لـ `Projectile(45, 45, 45)` صحيحًا.

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

يجب أن يكون التمثيل النصي صحيحًا أيضًا للحالات الأخرى.

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

يجب أن يكون لديك استدعاء `print(ball)`.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
