---
id: 6616ec0bf9ee8548ce6b0f08
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

حان الوقت الآن للعمل على حساب إحداثيات المسار! أنشئ طريقة من صنف `Projectile` باسم `__calculate_y_coordinate`، يجب أن تحتوي، بالإضافة إلى `self`، على معلمة `x`.

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

الصيغة أعلاه لحساب الموضع الرأسي $y$ لأي موضع أفقي $x$ معطى، مع زاوية البداية $\theta$، والسرعة $v_0$ والارتفاع $y_0$.

ستحتاج إلى استخدام `math.tan()` و `math.cos()` وتذكر أن `x ** y` هي الطريقة لكتابة $x^y$، وأن قيمة $g$ موجودة في المتغير `GRAVITATIONAL_ACCELERATION`.

نفذ الطريقة بحيث تُرجع إحداثي $y$.

إذا أردت اختبار الطريقة من خارج الصنف، يمكنك استدعاء `ball._Projectile__calculate_y_coordinate()` مع رقم كمعلمة بعد إعلان `ball`.

# --hints--

يجب أن يحتوي صنف `Projectile` على طريقة باسم `__calculate_y_coordinate`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

يجب أن تحتوي طريقة `__calculate_y_coordinate` على المعاملات `self, x`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

يجب أن تُرجع الطريقة القيمة الصحيحة. بالنسبة لـ `Projectile(12, 12, 12)`، مع 4 إحداثيات $x$ هي 0، 1، 2 و3، يجب أن يكون إحداثي y تقريبًا 12، 12.18، 12.28، 12.32.

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

يجب أن تُرجع الطريقة القيمة الصحيحة. بالنسبة لـ `Projectile(45, 12, 22)`، مع إحداثيات $x$ 6، 9، 10 و12 يجب أن يكون إحداثي y تقريبًا 14.32، 15.41، 15.76، 16.44.

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

يجب ألا تقوم طريقة `__calculate_y_coordinate` بتقريب النتيجة.

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
