---
id: 662fa453ec8033322b3ab2e3
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

حان الوقت للحديث عن التغليف وgetters. لقد كتبت ثلاث خصائص للحالة لتكون خاصة باستخدام شرطة سفلية مزدوجة في البداية. لاحظ أن هذه الخصائص تُسمى خاصة حسب الاصطلاح: رغم أنه يمكن الوصول إليها من الخارج، إلا أنه متفق على عدم القيام بذلك.

الـ getters هي التي تُستخدم للحصول على القيم من الخارج. لتعريف getter، عرّف طريقة تُرجع قيمة الخاصية المطلوبة وأضف لها مزين `@property`:

```py
class Nest:
    ...
    @property
    def number_of_eggs(self):
        return self.__number_of_eggs
```

المزين يحول الطريقة إلى خاصية، مما يعني أن الطريقة لا تُستدعى كطريقة عادية، بل تُستخدم كخاصية:

```py
n = Nest()
print(n.number_of_eggs)
```

في المثال أعلاه، الخاصية الخاصة `__number of eggs` تم الوصول إليها عبر الخاصية `number_of_eggs` للكائن `n`.

أنشئ getter باسم `speed` واجعله يُرجع قيمة الخاصية الخاصة `__speed`.

# --hints--

يجب أن تعرف طريقة باسم `speed`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('speed')
`))})
```

يجب أن تحتوي طريقة `speed` على معلمة واحدة فقط، وهي `self`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_args('self')
`))})
```

يجب أن تحتوي طريقة `speed` على مزين `@property`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_decorators('property')
`))})
```

يجب أن تُرجع طريقة `speed` قيمة الخاصية الخاصة `__speed`.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed == 22
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
