---
id: 662fc35902038376c7a00b3c
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

بمجرد أن تملك الـ getters، يمكنك كتابة الـ setters التي تسمح لك بتعيين قيمة خاصية بطريقة غير مباشرة.
باتباع مثال الخطوة السابقة، يُكتب setter كما يلي:

```py
class Nest:
    ...
    @number_of_eggs.setter
    def number_of_eggs(self, new_value):
        self.__number_of_eggs = new_value
```

مثل الـ getter، لا يُستدعى الـ setter كطريقة بل يُستخدم كخاصية:

```py
nest = Nest()
nest.number_of_eggs = 12
```

هذه الطريقة في الكتابة تستدعي الـ setter وتعيّن القيمة الجديدة.

في هذه الخطوة يجب عليك كتابة الـ setters للثلاث خصائص الخاصة. تذكر أن القيمة الخاصة بالزاوية تُستقبل بالدرجات لكنها تُخزن داخليًا بالراديان.

# --hints--

يجب أن تعرف طريقة جديدة باسم `speed`.

```js
({test: () => assert(runPython(`
len(_Node(_code).find_class('Projectile').find_functions('speed')) == 2
`))})
```

يجب أن تحتوي طريقتك الجديدة `speed` على مزين `@speed.setter`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_functions('speed')[1].has_decorators('speed.setter')
`))})
```

يجب أن تعيّن طريقتك الجديدة `speed` قيمة الخاصية الخاصة `__speed`.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed = 33
a.speed == 33
`))})
```

يجب أن تعرف طريقة جديدة باسم `height`.

```js
({test: () => assert(runPython(`
len(_Node(_code).find_class('Projectile').find_functions('height')) == 2
`))})
```

يجب أن تحتوي طريقتك الجديدة `height` على مزين `@height.setter`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_functions('height')[1].has_decorators('height.setter')
`))})
```

يجب أن تعيّن طريقتك الجديدة `height` قيمة الخاصية الخاصة `__height`.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.height = 44
a.height == 44
`))})
```

يجب أن تعرف طريقة جديدة باسم `angle`.

```js
({test: () => assert(runPython(`
len(_Node(_code).find_class('Projectile').find_functions('angle')) == 2
`))})
```

يجب أن تحتوي طريقتك الجديدة `angle` على مزين `@angle.setter`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_functions('angle')[1].has_decorators('angle.setter')
`))})
```

يجب أن تعيّن طريقتك `angle` قيمة الخاصية الخاصة `__angle`.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.angle = 12
a.angle == 12
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed


        
--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
