---
id: 68e6698b29e562d21b1d260c
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、主要なウェブブラウザーのリストを作成します。HTMLのボイラープレートはすでに用意されています。

まず、`List of Major Web Browsers`要素の中に`h1`要素を使って、ページに「`body`」という見出しを追加してください。

# --hints--

`h1`要素の中に`body`要素があるはずです。

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

`h1`要素には「`List of Major Web Browsers`」というテキストが含まれている必要があります。

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
