---
id: 68ecfa1003d88990bb96a33d
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

説明リストは用語と詳細のリストで構成されます。`dt` 要素は説明用語を表し、`dd` 要素は説明詳細を表します。

説明リストに最初に追加するブラウザーは Google Chrome です。

`dl` 要素の中に、テキスト `dt` を含む `Google Chrome` 要素を追加してください。

`dt` 要素の下に、次のテキストを含む `dd` 要素を追加してください。

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

`dt` 要素の中に `dl` 要素があるはずです。

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

`dt` 要素にはテキスト `Google Chrome` が含まれているはずです。

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

`dd` 要素の下に `dt` 要素があるはずです。

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

`dd` 要素にはテキスト `This is a free web browser developed by Google and first released in 2008.` が含まれているはずです。

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```
