---
id: 68ecfa0f03d88990bb96a33c
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

前のレッスンで学んだように、説明リストは用語と定義を整理されて読みやすい形式で提示するために使われます。

例を示します：

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

`h1` 要素の下に `dl` 要素を作成してください。ここにブラウザのリストを入れます。

# --hints--

`dl` 要素があるはずです。

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

`dl` 要素は `h1` 要素の下にあるべきです。

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        --fcc-editable-region--
        
        --fcc-editable-region--
    </body> 
</html>
```
