---
id: 68ecfa1003d88990bb96a33e
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

次に、2つ目のブラウザーをリストに追加しましょう。

`Google Chrome` のエントリーの下に、テキスト `dt` を含む別の `Firefox` 要素を追加してください。

`dt` 要素の下に、次のテキストを含む `dd` 要素を追加してください。

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

コード内に `dt` 要素が2つあるはずです。

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

2つ目の `dt` 要素にはテキスト `Firefox` が含まれているはずです。

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

コード内に `dd` 要素が2つあるはずです。

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

2つ目の `dd` 要素にはテキスト `This is a free web browser developed by the Mozilla Corporation and first created in 2004.` が含まれているはずです。

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
