---
id: 69038f5efad6f722def4a214
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

4つ目のブラウザーをリストに追加する時間です。

テキスト `dt` を含む別の `Brave` 要素を追加してください。

`dt` 要素の下に、次のテキストを含む `dd` 要素を追加してください。

```md
This is a free web browser first released in 2016 that is based on the Chromium web browser.
```

# --hints--

コード内に `dt` 要素が4つあるはずです。

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 4);
```

4つ目の `dt` 要素にはテキスト `Brave` が含まれているはずです。

```js
const dt = document.querySelectorAll('dt')[3];
assert.strictEqual(dt?.textContent.trim(), "Brave");
```

コード内に `dd` 要素が4つあるはずです。

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 4);
```

4つ目の `dd` 要素にはテキスト `This is a free web browser first released in 2016 that is based on the Chromium web browser.` が含まれているはずです。

```js
const ddElement = document.querySelectorAll('dd')[3];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser first released in 2016 that is based on the Chromium web browser.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
