---
id: 69038dd51a41ae215fcc030a
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

リストに追加する3つ目のブラウザーはSafariウェブブラウザー用です。

テキスト`dt`を含む別の`Safari`要素を追加してください。

`dt`要素の下に、次のテキストを含む`dd`要素を追加してください。

```md
This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.
```

# --hints--

コードには3つの`dt`要素があるはずです。

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 3);
```

3つ目の`dt`要素にはテキスト`Safari`が含まれているはずです。

```js
const dt = document.querySelectorAll('dt')[2];
assert.strictEqual(dt?.textContent.trim(), "Safari");
```

コードには3つの`dd`要素があるはずです。

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 3);
```

3つ目の`dd`要素にはテキスト`This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.`が含まれているはずです。

```js
const ddElement = document.querySelectorAll('dd')[2];
assert.strictEqual(ddElement?.textContent.trim(), "This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
