---
id: 67fe4adca39a9b224b72d642
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Ahora, la entrada de búsqueda debe funcionar como se espera y filtrar la lista de compras basándose en la entrada del usuario.

A continuación, agregaremos una forma de marcar los artículos en la lista para que los usuarios puedan llevar un registro de lo que ya han comprado.

Para hacer esto, necesitamos otra variable de estado y una función para establecerla. Use el `useState` hook para crear una variable de estado llamada `selectedItems` y una función para establecerla llamada `setSelectedItems`. Inicializar `selectedItems` como un array vacío. Este array almacenará los artículos que el usuario ha seleccionado, y aplicaremos algo de estilizado a esos artículos seleccionados más tarde.

# --hints--

Debe crear una variable de estado llamada `selectedItems` y una función para establecerla llamada `setSelectedItems`.

```js
assert.match(code, /(const|let|var)\s+\[\s*selectedItems\s*,\s*setSelectedItems\s*\]\s*=\s*useState\(/);
```

El valor inicial de `selectedItems` debe ser un array vacío.

```js
const selectedItemsValue = code.match(/(?:const|let|var)\s+\[\s*selectedItems\s*,\s*setSelectedItems\s*\]\s*=\s*useState\(\s*(\[\s*(.*?)\s*\])\s*\)/)?.[1] ?? null;

assert.isTrue(
  selectedItemsValue &&
  selectedItemsValue?.replace(/\s/g, "") === "[]"
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Shopping List</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main id="app-container"></main>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ShoppingList } from "./index.jsx";
      const appContainer = document.getElementById("app-container");
      const root = ReactDOM.createRoot(appContainer);
      root.render(<ShoppingList />);
    </script>
  </body>
</html>

```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

*,
*::before,
*::after {
  box-sizing: border-box;
}

body {
  font-family: Arial, sans-serif;
  line-height: 1.5;
  color: var(--dark-grey);
  background-color: var(--dark-grey);
}

main,
.container {
  display: flex;
  flex-direction: column;
  align-items: center;
}

.container {
  background-color: var(--light-grey);
  width: 90%;
  margin: 20px;
  padding: 10px;
}

h1 {
  color: var(--dark-grey);
}

form {
  text-align: center;
}

button {
  cursor: pointer;
}

button {
  cursor: pointer;
  width: 100px;
  margin: 3px;
  color: var(--dark-grey);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

input {
  color: var(--dark-grey);
  margin-left: 5px;
  padding: 3px;
}

li {
  text-align: left;
  margin: 10px 0;
}

.selected-item {
  font-weight: bold;
}

@media (min-width: 425px) {
  .container {
    width: 400px;
  }
}

```

```jsx
const { useState } = React;

const items = [
  "Apples",
  "Bananas",
  "Strawberries",
  "Blueberries",
  "Mangoes",
  "Pineapple",
  "Lettuce",
  "Broccoli",
  "Paper Towels",
  "Dish Soap",
];

export const ShoppingList = () => {
  const [query, setQuery] = useState("");
  --fcc-editable-region--

  --fcc-editable-region--

  const filteredItems = items.filter((item) => 
    item.toLowerCase().includes(query.toLowerCase())
  );

  return (
    <div className="container">
      <h1>Shopping List</h1>
      <form>
        <label htmlFor="search">Search for an item:</label>
        <input
          id="search"
          type="search"
          placeholder="Search..."
          aria-describedby="search-description"
          value={query}
          onChange={(e) => setQuery(e.target.value)}
        /> 
        <p id="search-description">Type to filter the list below:</p>
        <ul>
          {filteredItems.map((item) => 
            <li key={item}>
              {item}
            </li>
          )}
        </ul>
      </form>
    </div>
  );
};

```
