---
id: 56533eb9ac21ba0edf2244cf
title: Побудова колекції записів
challengeType: 26
dashedName: build-a-record-collection
---

# --description--

Ви створюєте функцію, яка допомагає в обслуговуванні колекції музичних альбомів. Колекція організована як об’єкт, що містить кілька альбомів, які також є об’єктами. Кожен альбом у колекції представлений унікальним `id` як назва властивості. Усередині кожного об’єкта альбому є різні властивості, що описують інформацію про альбом. Не всі альбоми мають повну інформацію.

Функція `updateRecords` приймає 4 аргументи, представлені такими параметрами функції:

-   `records` — об’єкт, що містить кілька окремих альбомів
-   `id` — число, яке представляє конкретний альбом в об’єкті `records`
-   `prop` — рядок, що представляє назву властивості альбому для оновлення
-   `value` — рядок, що містить інформацію для оновлення властивості альбому

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ваша функція завжди має повертати весь об’єкт `records`.
2. Якщо `value` є порожнім рядком, видаліть задану властивість `prop` з альбому.
3. Якщо `prop` не є `tracks` і `value` не є порожнім рядком, призначте `value` цій властивості `prop` альбому.
4. Якщо `prop` є `tracks` і `value` не є порожнім рядком, але альбом не має властивості `tracks`, створіть порожній масив і додайте до нього `value`.
5. Якщо проп є `tracks` і `value` не є порожнім рядком, додайте `value` у кінець існуючого масиву `tracks` альбому.

**Примітка:** Для тестів використовується копія об’єкта `recordCollection`. Ваша функція не повинна безпосередньо посилатися на об’єкт `recordCollection`, лише на параметр функції.

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

Ви маєте мати функцію `updateRecords`.

```js
assert.isFunction(updateRecords);
```

Після `updateRecords(recordCollection, 5439, "artist", "ABBA")` `artist` має бути рядком `ABBA`

```js
assert.equal(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'],
    'ABBA'
);
```

Після `updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")` `tracks` має мати рядок `Take a Chance on Me` як останній і єдиний елемент.

```js
updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me');
assert.lengthOf(_recordCollection[5439]['tracks'], 1);
assert.equal(_recordCollection[5439]['tracks'].pop(), 'Take a Chance on Me');
```

Після `updateRecords(recordCollection, 2548, "artist", "")` `artist` не має бути встановлено

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('artist'));
```

Після `updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")` `tracks` має мати рядок `Addicted to Love` як останній елемент.

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop(), 'Addicted to Love'
);
```

Після `updateRecords(recordCollection, 2468, "tracks", "Free")` `tracks` має мати рядок `1999` як перший елемент.

```js
assert.equal(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0], '1999'
);
```

Після `updateRecords(recordCollection, 2548, "tracks", "")` `tracks` не має бути встановлено

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('tracks'));
```

Після `updateRecords(recordCollection, 1245, "albumTitle", "Riptide")` `albumTitle` має бути рядком `Riptide`

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ], 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};



```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
