---
id: 68e6a99a0fd7386fb66e156f
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у стилізації прапорців, створюючи панель налаштувань гри.

Весь шаблонний код HTML (`DOCTYPE`, `html`, `head` та `body`) уже надано.

У межах `body` створіть елемент `div` з вкладеним елементом `h1`. Додайте елементу `h1` текст `Game Settings`.

# --hints--

У вас має бути елемент `div`.

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

У вас має бути елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Ваш елемент `h1` має бути всередині елемента `div`.

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

Ваш елемент `h1` має містити текст `Game Settings`.

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
