---
id: 68e9a3194160f8f5620ddf21
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Одразу після вашого елемента `h1` створіть чотири елементи `label`.

Всередині кожного з елементів `label` створіть елемент `input` з атрибутом `type`, встановленим у `"checkbox"`.

Після кожного елемента `input` додайте вашим елементам `label` такі тексти, у порядку: `Sound Effects`, `Background Music`, `Hard Mode` та `Haptic Feedback`.

# --hints--

У вас має бути чотири елементи `label`.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Перший ваш елемент `label` повинен містити елемент `input`.

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

Перший ваш елемент `input` повинен мати атрибут `type`, встановлений у `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

Перший ваш елемент `label` повинен містити текст `Sound Effects`.

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

Текст `Sound Effects` має йти після елемента `input`.

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

Другий ваш елемент `label` повинен містити елемент `input`.

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

Другий ваш елемент `input` повинен мати атрибут `type`, встановлений у `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

Другий ваш елемент `label` повинен містити текст `Background Music`.

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

Текст `Background Music` має йти після елемента `input`.

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

Третій ваш елемент `label` повинен містити елемент `input`.

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

Третій ваш елемент `input` повинен мати атрибут `type`, встановлений у `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

Третій ваш елемент `label` повинен містити текст `Hard Mode`.

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

Текст `Hard Mode` має йти після елемента `input`.

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

Четвертий ваш елемент `label` повинен містити елемент `input`.

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

Четвертий ваш елемент `input` повинен мати атрибут `type`, встановлений у `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

Четвертий ваш елемент `label` повинен містити текст `Haptic Feedback`.

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

Текст `Haptic Feedback` має йти після елемента `input`.

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
