---
id: 68e9a7bd60c971278d2d4ffc
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

У попередньому уроці ви дізналися про псевдокласи та псевдоелементи в CSS. Тепер ви застосуєте ці знання, створивши комбінований селектор типу з селектором псевдокласу.

Спочатку почніть із формату селектора типу, встановивши його як `input[type="checkbox"]` і додавши `:checked` в кінець.

Далі задайте йому `background-color` зі значенням `#f1be32`.

І нарешті, встановіть `border-color` зі значенням `#e2a60d`.

# --hints--

Ви повинні мати комбінований селектор типу і псевдокласу `input[type="checkbox"]:checked`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked'));
```

Ви повинні додати властивість `background-color` до вашого комбінованого селектора типу і псевдокласу `input[type="checkbox"]:checked`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor);
```

Ви повинні мати властивість `background-color` зі значенням `#f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor, "rgb(241, 190, 50)");
```

Ви повинні додати властивість `border-color` до вашого комбінованого селектора типу і псевдокласу `input[type="checkbox"]:checked`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor);
```

Ви повинні мати властивість `border-color` зі значенням `#e2a60d`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor, "rgb(226, 166, 13)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

--fcc-editable-region--

--fcc-editable-region--

```
