---
id: 68e9a834dfef752dc35f6f90
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Тепер, коли ваші прапорці змінюють колір, коли користувач натискає на них, ви застосуєте стилі, щоб зробити більш очевидним, що прапорець вибрано.

CSS-псевдокласи можна поєднувати. Ви можете додати один псевдоклас або псевдоелемент за іншим, щоб націлитися на елементи, які відповідають кільком умовам.

Створіть новий селектор для `input[type="checkbox"]`. Додайте псевдоклас `:checked`, а потім псевдоелемент `::after`. Це дозволить додати візуальний індикатор, коли прапорець вибрано.

Далі додайте `content` зі значенням `"✓"`.

# --hints--

Ви повинні мати комбінований селектор типу з псевдокласом і псевдоелементом `input[type="checkbox"]:checked::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

Ви повинні додати властивість `content` до вашого комбінованого селектора типу з псевдокласом і псевдоелементом `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

Ви повинні мати властивість `content` зі значенням `"✓"`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```
