---
id: 68ee86e3ee07b01626a392cc
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

У вашому множині селекторів `.settings-card` встановіть властивість `margin` зі значенням `auto`.

Встановлення властивості `margin` зі значенням `auto` автоматично регулює поля елемента, щоб рівномірно розподілити залишковий простір у його контейнері, що часто використовується для горизонтального центрування блочних елементів.

І нарешті, встановіть властивість `text-align` зі значенням `left`. Це вирівняє рядковий вміст, наприклад текст, по лівому краю його контейнера.

# --hints--

Ви повинні додати властивість `margin` до селектора класу `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

Ви повинні мати властивість `margin` зі значенням `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

Ви повинні додати властивість `text-align` до селектора класу `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

Ви повинні мати властивість `text-align` зі значенням `left`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
