---
id: 68f1560aff1214103ab6b59b
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Тепер ви почнете прикрашати сторінку за допомогою стилів CSS. Почніть зі створення селектора `body`.

Встановіть властивість `body` зі значенням `height` та властивість `100vh` зі значенням `background-color` для `#f0f0f0`.

Властивість `height` для `100vh` змушує `body` займати всю висоту вікна перегляду браузера, а світло-сірий фон створює ненав’язливу основу для сторінки.

Нарешті, встановіть властивість `text-align` зі значенням `center`. Це вирівняє по центру весь рядковий вміст, що міститься на сторінці, якщо тільки дочірній елемент не замістить це своїм власним вирівнюванням.

# --hints--

Ви повинні використати селектор типу `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Ви повинні додати властивість `height` у селектор типу `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

Ви повинні мати властивість `height` зі значенням `100vh`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

Ви повинні додати властивість `background-color` у селектор типу `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

Ви повинні мати властивість `background-color` зі значенням `#f0f0f0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

Ви повинні додати властивість `text-align` у селектор типу `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

Ви повинні мати властивість `text-align` зі значенням `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
