---
id: 68f15cd397503429f0345262
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

В оголошенні для селектора `input[type="checkbox"]:checked::after` встановіть властивість `display` у значення `block`. Встановлення властивості `display` як `block` робить елемент блочним, тобто він займає всю ширину свого контейнера і починається з нового рядка, що дозволяє легше керувати його шириною, висотою та відступами.

Далі встановіть `text-align` у значення `center`. Це по горизонталі вирівняє рядковий вміст (наприклад, текст або рядкові елементи) всередині блоку. Блок усе ще займатиме всю ширину контейнера, але все всередині нього буде вирівняне по центру.

# --hints--

Ви повинні додати властивість `display` до вашого комбінованого селектора типу з псевдокласом і псевдоелементом `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

Ви повинні мати властивість `display` зі значенням `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

Ви повинні додати властивість `text-align` до вашого комбінованого селектора типу з псевдокласом і псевдоелементом `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

Ви повинні мати властивість `text-align` зі значенням `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
