---
id: 68f15ef0446464325ce1aa45
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Потрібно додати відступи між прапорцями та мітками. Почніть із використання селектора для `label` у вашому CSS і встановіть властивість `display` зі значенням `block`.

Далі встановіть властивість `margin` зі значенням `8px auto`, щоб додати вертикальні відступи між елементами. Це створить однакові відступи зверху та знизу кожного елемента, допомагаючи відокремити вміст і покращити читабельність.

Нарешті, встановіть властивість `cursor` зі значенням `pointer` для елементів `label`. Це змінить курсор на іконку руки, коли користувач наводить курсор на `label`, сигналізуючи, що елемент активний, і покращуючи загальний користувацький досвід.

# --hints--

У вашому CSS має бути селектор типу для `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

Ви маєте додати властивість `display` до вашого селектора `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

Ви маєте встановити властивість `display` зі значенням `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

Ви маєте додати властивість `margin` до вашого селектора `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

Ви маєте встановити властивість `margin` зі значенням `8px auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

Ви маєте додати властивість `cursor` до вашого селектора `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

Ви маєте встановити властивість `cursor` зі значенням `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
