---
id: 68f29dbf9125476b4808d6fb
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Тепер настав час стилізувати контейнер `settings-card`. Створіть селектор класу для `settings-card`.

Саме тут буде розміщено все ваше форматування для контейнера. Встановіть `max-width` у `250px`, щоб задати загальний розмір контейнера.

Далі встановіть `padding` у `20px`, щоб між вмістом і кордоном контейнера був пробіл.

Після цього створіть заокруглені краї, встановивши `border-radius` у `10px`.

Потім додайте `box-shadow` зі значеннями `0 2px 6px rgba(0,0,0,0.2)`. Це створить легкий ефект "підняття", який додасть глибини контейнеру.

# --hints--

У вашому CSS має бути селектор класу для `settings-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

Ви повинні додати властивість `max-width` до селектора класу `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

У вас має бути властивість `max-width` зі значенням `250px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

Ви повинні додати властивість `padding` до селектора класу `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

У вас має бути властивість `padding` зі значенням `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

Ви повинні додати властивість `border-radius` до вашого селектора класу `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

У вас має бути властивість `border-radius` зі значенням `10px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

Ви повинні додати властивість `box-shadow` до селектора класу `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

У вас має бути `box-shadow` зі значенням `0 2px 6px rgba(0,0,0,0.2)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
