---
id: 587d7fa8367417b2b2512bcd
title: "Crea una barra per ogni punto dati nell'insieme"
challengeType: 6
forumTopicId: 301482
dashedName: create-a-bar-for-each-data-point-in-the-set
---

# --description--

L'ultima sfida ha aggiunto solo un rettangolo all'elemento `svg` per rappresentare una barra. Qui, combinerai ciò che hai imparato finora su `data()`, `enter()` e le forme SVG per creare e aggiungere in coda un rettangolo per ogni punto dati in `dataset`.

Una sfida precedente ha mostrato il formato per come creare e aggiungere in coda un `div` per ogni elemento in `dataset`:

```js
d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
```

Ci sono alcune differenze nel lavorare con elementi `rect` invece di elementi `div`. Gli elementi `rect` devono essere aggiunti in coda a un elemento `svg`, non direttamente al `body`. Inoltre, devi indicare a D3 dove posizionare ogni `rect` all'interno dell'area `svg`. Il posizionamento delle barre sarà trattato nella prossima sfida.

# --instructions--

Usa i metodi `data()`, `enter()` e `append()` per creare e aggiungere in coda un `rect` per ogni elemento in `dataset`. Le barre dovrebbero apparire tutte sovrapposte; questo sarà sistemato nella prossima sfida.

# --hints--

Il tuo documento dovrebbe avere 9 elementi `rect`.

```js
assert.lengthOf(document.querySelectorAll('rect'), 9);
```

Il tuo codice dovrebbe usare il metodo `data()`.

```js
assert.match(code, /\.data/g);
```

Il tuo codice dovrebbe usare il metodo `enter()`.

```js
assert.match(code, /\.enter/g);
```

Il tuo codice dovrebbe usare il metodo `append()`.

```js
assert.match(code, /\.append/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      // Add your code below this line

      // Add your code above this line
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
