---
id: 587d7fa8367417b2b2512bcc
title: Visualizzare forme con SVG
challengeType: 6
forumTopicId: 301485
dashedName: display-shapes-with-svg
---

# --description--

L'ultima sfida ha creato un elemento `svg` con una larghezza e un'altezza specificate, visibile perché aveva un `background-color` applicato nel tag `style`. Il codice ha riservato lo spazio per la larghezza e l'altezza indicate.

Il passaggio successivo è creare una forma da inserire nell'area `svg`. SVG supporta diverse forme, come rettangoli e cerchi, usate per visualizzare dati. Per esempio, una forma rettangolare (`<rect>`) SVG potrebbe creare una barra in un grafico a barre.

Quando posizioni una forma nell'area `svg`, puoi specificare la sua posizione con le coordinate `x` e `y`. Il punto di origine (0, 0) si trova nell'angolo in alto a sinistra. Valori positivi per `x` spostano la forma verso destra, e valori positivi per `y` la spostano verso il basso rispetto al punto di origine.

Per posizionare una forma al centro dell'area `svg` di 500 (larghezza) x 100 (altezza) della sfida precedente, la coordinata `x` sarebbe 250 e la coordinata `y` sarebbe 50.

Un `rect` SVG ha quattro attributi. Ci sono le coordinate `x` e `y` per indicare dove si trova nell'area `svg`. Ha anche un `height` e un `width` per specificarne la dimensione.

# --instructions--

Aggiungi una forma `rect` all'area `svg` usando `append()`, e assegna un attributo `width` con valore `25` e un attributo `height` con valore `100`. Inoltre, assegna agli attributi `rect` e `x` del `y` il valore `0`.

# --hints--

Il tuo documento dovrebbe contenere 1 elemento `rect`.

```js
assert.lengthOf(document.querySelectorAll('rect'), 1);
```

L'elemento `rect` dovrebbe avere un attributo `width` impostato su `25`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('width'), '25');
```

L'elemento `rect` dovrebbe avere un attributo `height` impostato su `100`.

```js
assert.strictEqual(
  document.querySelector('rect')?.getAttribute('height'),
  '100'
);
```

L'elemento `rect` dovrebbe avere un attributo `x` impostato su `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('x'), '0');
```

L'elemento `rect` dovrebbe avere un attributo `y` impostato su `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('y'), '0');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h)
      .append('rect')
      .attr('width', 25)
      .attr('height', 100)
      .attr('x', 0)
      .attr('y', 0);
  </script>
</body>
```
