---
id: 587d7fac367417b2b2512bdb
title: Imposta un dominio e un intervallo su una scala
challengeType: 6
forumTopicId: 301491
dashedName: set-a-domain-and-a-range-on-a-scale
---

# --description--

Per impostazione predefinita, le scale usano la relazione di identità. Questo significa che il valore di input corrisponde al valore di output. Tuttavia, le scale possono essere molto più flessibili e interessanti.

Supponiamo che un dataset abbia valori che vanno da 50 a 480. Queste sono le informazioni di input per una scala, conosciute anche come <dfn>dominio</dfn>.

Vuoi mappare quei punti lungo l'asse `x` nell'SVG, tra 10 unità e 500 unità. Queste sono le informazioni di output, conosciute anche come <dfn>intervallo</dfn>.

I metodi `domain()` e `range()` impostano questi valori per la scala. Entrambi i metodi prendono come argomento un array di almeno due elementi. Ecco un esempio:

```js
scale.domain([50, 480]);
scale.range([10, 500]);
scale(50);
scale(480);
scale(325);
scale(750);
d3.scaleLinear();
```

In ordine, i seguenti valori verrebbero mostrati nella console: `10`, `500`, `323.37` e `807.67`.

Nota che la scala usa la relazione lineare tra i valori del dominio e dell'intervallo per calcolare quale dovrebbe essere l'output per un dato numero. Il valore minimo nel dominio (50) corrisponde al valore minimo (10) nell'intervallo.

# --instructions--

Crea una scala e imposta il suo dominio su `[250, 500]` e l'intervallo su `[10, 150]`.

**Nota:** Puoi concatenare i metodi `domain()` e `range()` alla variabile `scale`.

# --hints--

Il tuo codice dovrebbe usare il metodo `domain()`.

```js
assert.match(code, /\.domain/g);
```

Il `domain()` del `scale` dovrebbe essere impostato su `[250, 500]`.

```js
assert.deepEqual(scale.domain(), [250, 500]);
```

Il tuo codice dovrebbe usare il metodo `range()`.

```js
assert.match(code, /\.range/g);
```

Il `range()` del `scale` dovrebbe essere impostato su `[10, 150]`.

```js
assert.deepEqual(scale.range(), [10, 150]);
```

Il testo nel `h2` dovrebbe essere `-102`.

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '-102');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line
    const scale = d3.scaleLinear();



    // Add your code above this line
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const scale = d3.scaleLinear();
    scale.domain([250, 500]);
    scale.range([10, 150]);
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
