---
id: 587d7fad367417b2b2512bdf
title: aggiungere assi a una visualizzazione
challengeType: 6
forumTopicId: 301472
dashedName: add-axes-to-a-visualization
---

# --description--

Un altro modo per migliorare il diagramma a dispersione è aggiungere un asse x e un asse y.

D3 ha due metodi, `axisLeft()` e `axisBottom()`, per renderizzare rispettivamente l'asse y e l'asse x. Ecco un esempio per creare l'asse x basato sul `xScale` nelle sfide precedenti:

```js
const xAxis = d3.axisBottom(xScale);
```

Il passaggio successivo è renderizzare l'asse sull'SVG. Per farlo, puoi usare un componente SVG generico, l'elemento `g`. `g` sta per gruppo. A differenza di `rect`, `circle` e `text`, un asse è solo una linea retta quando viene renderizzato. Poiché è una forma semplice, usare `g` funziona. L'ultimo passaggio è applicare un attributo `transform` per posizionare l'asse nel punto giusto sull'SVG. Altrimenti, la linea verrebbe renderizzata lungo il bordo dell'SVG e non sarebbe visibile. SVG supporta diversi tipi di `transforms`, ma per posizionare un asse serve `translate`. Quando viene applicato all'elemento `g`, sposta l'intero gruppo verso destra e verso il basso di quanto indicato. Ecco un esempio:

```js
const xAxis = d3.axisBottom(xScale);

svg
  .append('g')
  .attr('transform', 'translate(0, ' + (h - padding) + ')')
  .call(xAxis);
```

Il codice sopra posiziona l'asse x in fondo all'SVG. Poi viene passato come argomento al metodo `call()`. L'asse y funziona allo stesso modo, tranne che l'argomento `translate` ha la forma `(x, 0)`. Poiché `translate` è una stringa nel metodo `attr()` sopra, puoi usare la concatenazione per includere valori variabili nei suoi argomenti.

# --instructions--

Il diagramma a dispersione ora ha un asse x. Crea un asse y in una variabile chiamata `yAxis` usando il metodo `axisLeft()`. Poi renderizza l'asse usando un elemento `g`. Assicurati di usare un attributo `transform` per traslare l'asse di un numero di unità pari alla spaziatura interna verso destra e `0` unità verso il basso. Ricorda di `call()` l'asse.

# --hints--

Il tuo codice dovrebbe usare il metodo `axisLeft()` con `yScale` passato come argomento.

```js
assert.match(code, /\.axisLeft\(yScale\)/g);
```

L'elemento `g` dell'asse y dovrebbe avere un attributo `transform` per traslare l'asse di `(60, 0)`.

```js
const yAxis = document.querySelectorAll('g')[10];
assert.match(yAxis?.getAttribute('transform'), /translate\(60\s*,\s*0\)/g);
```

Il tuo codice dovrebbe chiamare `yAxis`.

```js
assert.match(code, /\.call\(\s*yAxis\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);
    // Add your code below this line
    const yAxis = undefined;
    // Add your code above this line

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);

    const yAxis = d3.axisLeft(yScale);

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    svg
      .append('g')
      .attr('transform', 'translate(' + padding + ',0)')
      .call(yAxis);
  </script>
</body>
```
