---
id: 587d7fa7367417b2b2512bc6
title: aggiungere definizioni di stile in linea agli elementi
challengeType: 6
forumTopicId: 301475
dashedName: add-inline-styling-to-elements
---

# --description--

D3 ti permette di aggiungere definizioni di stile CSS in linea su elementi dinamici con il metodo `style()`.

Il metodo `style()` prende come argomento una coppia chiave-valore separata da virgola. Ecco un esempio per impostare il colore del testo della selezione su blu:

```js
selection.style('color', 'blue');
```

# --instructions--

Aggiungi il metodo `style()` al codice nell'editor per fare in modo che tutto il testo mostrato abbia un `font-family` di `verdana`.

# --hints--

I tuoi elementi `h2` dovrebbero avere un `font-family` di `verdana`.

```js
const headingTwo = document.querySelector('h2');
assert.exists(headingTwo); 
const fontFamily = window.getComputedStyle(headingTwo)['font-family'];
assert.strictEqual(fontFamily, 'verdana');
```

Il tuo codice dovrebbe usare il metodo `style()`.

```js
assert.match(code, /\.style/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('font-family', 'verdana');
  </script>
</body>
```
