---
id: 587d7fac367417b2b2512bdd
title: Usa scale dinamiche
challengeType: 6
forumTopicId: 301495
dashedName: use-dynamic-scales
---

# --description--

I metodi D3 `min()` e `max()` sono utili per aiutare a impostare la scala.

Dato un set di dati complesso, una priorità è impostare la scala in modo che la visualizzazione si adatti alla larghezza e all'altezza del contenitore SVG. Vuoi che tutti i dati siano tracciati all'interno dell'SVG così da essere visibili sulla pagina web.

L'esempio qui sotto imposta la scala dell'asse x per i dati di uno scatter plot. Il metodo `domain()` passa informazioni alla scala sui valori grezzi dei dati per il grafico. Il metodo `range()` le fornisce informazioni sullo spazio effettivo sulla pagina web per la visualizzazione.

Nell'esempio, il dominio va da 0 al massimo nell'insieme. Usa il metodo `max()` con una funzione di callback basata sui valori x negli array. Il range usa la larghezza dell'SVG (`w`), ma include anche un po' di spaziatura interna. Questo crea spazio tra i punti dello scatter plot e il bordo dell'SVG.

```js
const dataset = [
  [34, 78],
  [109, 280],
  [310, 120],
  [79, 411],
  [420, 220],
  [233, 145],
  [333, 96],
  [222, 333],
  [78, 320],
  [21, 123]
];
const w = 500;
const h = 500;

const padding = 30;
const xScale = d3
  .scaleLinear()
  .domain([0, d3.max(dataset, d => d[0])])
  .range([padding, w - padding]);
```

La spaziatura interna può sembrare confusa all'inizio. Immagina l'asse x come una linea orizzontale da 0 a 500 (il valore di larghezza per l'SVG). Includere la spaziatura interna nel metodo `range()` forza il grafico a iniziare a 30 lungo quella linea (invece che a 0) e a finire a 470 (invece che a 500).

# --instructions--

Usa la variabile `yScale` per creare una scala lineare per l'asse y. Il dominio dovrebbe partire da zero e arrivare al valore massimo `y` nell'insieme. Il range dovrebbe usare l'altezza dell'SVG (`h`) e includere la spaziatura interna.

**Nota:** Ricorda di mantenere il grafico orientato correttamente. Quando imposti il range per le coordinate y, il valore più alto (altezza meno spaziatura interna) è il primo argomento, e il valore più basso è il secondo argomento.

# --hints--

Il testo in `h2` dovrebbe essere `30`.

```js
assert.strictEqual(output, 30);
assert.strictEqual(document.querySelector('h2')?.textContent, '30');
```

Il `domain()` di yScale dovrebbe essere equivalente a `[0, 411]`.

```js
assert.deepEqual(yScale.domain(), [0, 411]);
```

Il `range()` di yScale dovrebbe essere equivalente a `[470, 30]`.

```js
assert.deepEqual(yScale.range(), [470, 30]);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    // Padding between the SVG boundary and the plot
    const padding = 30;

    // Create an x and y scale

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    // Add your code below this line

    const yScale = undefined;


    // Add your code above this line

    const output = yScale(411); // Returns 30
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const padding = 30;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const output = yScale(411);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
