---
id: 587d7fa7367417b2b2512bc8
title: aggiungere classi con d3
challengeType: 6
forumTopicId: 301473
dashedName: add-classes-with-d3
---

# --description--

Usare molti stili in linea sugli elementi HTML diventa difficile da gestire, anche per app più piccole. È più semplice aggiungere una classe agli elementi e definire lo stile di quella classe una sola volta usando regole CSS. D3 ha il metodo `attr()` per aggiungere qualsiasi attributo HTML a un elemento, incluso un nome di classe.

Il metodo `attr()` funziona allo stesso modo di `style()`. Accetta valori separati da virgola e può usare una funzione di callback. Ecco un esempio per aggiungere una classe `container` a una selezione:

```js
selection.attr('class', 'container');
```

Nota che il parametro `class` rimarrà lo stesso ogni volta che devi aggiungere una classe e solo il parametro `container` cambierà.

# --instructions--

Aggiungi il metodo `attr()` al codice nell'editor e assegna una classe `bar` agli elementi `div`.

# --hints--

I tuoi elementi `div` dovrebbero avere una classe `bar`.

```js
assert.isTrue(document.querySelector('div')?.classList.contains('bar'));
```

Il tuo codice dovrebbe usare il metodo `attr()`.

```js
assert.match(code, /\.attr/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line
  </script>
</body>
```
