---
id: 587d7fa6367417b2b2512bc3
title: seleziona un gruppo di elementi con d3
challengeType: 6
forumTopicId: 301490
dashedName: select-a-group-of-elements-with-d3
---

# --description--

D3 ha anche il metodo `selectAll()` per selezionare un gruppo di elementi. Restituisce un array di nodi HTML per tutti gli elementi nel documento che corrispondono alla stringa di input. Ecco un esempio per selezionare tutti i tag àncora in un documento:

```js
const anchors = d3.selectAll('a');
```

Come il metodo `select()`, `selectAll()` supporta il concatenamento di metodi e puoi usarlo con altri metodi.

# --instructions--

Seleziona tutti i tag `li` nel documento e cambia il loro testo nella stringa `list item` concatenando il metodo `.text()`.

# --hints--

Dovrebbero esserci 3 elementi `li` nella pagina, e il testo in ciascuno dovrebbe essere `list item`. Maiuscole e spaziatura devono corrispondere esattamente.

```js
const listItems = document.querySelectorAll('li');
assert.lengthOf(listItems, 3);
for (let i = 0; i < listItems.length; i++) {
  assert.match(listItems[i]?.textContent, /list item/g);
}
```

Il tuo codice dovrebbe consultare l'oggetto `d3`.

```js
assert.match(code, /d3/g);
```

Il tuo codice dovrebbe usare il metodo `selectAll`.

```js
assert.match(code, /\.selectAll/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    d3.selectAll('li').text('list item');
  </script>
</body>
```
