---
id: 587d7fa8367417b2b2512bcb
title: Scopri SVG in D3
challengeType: 6
forumTopicId: 301489
dashedName: learn-about-svg-in-d3
---

# --description--

<dfn>SVG</dfn> sta per <dfn>Scalable Vector Graphics</dfn>.

Qui "scalable" significa che, se ingrandisci o riduci un oggetto, non apparirà pixelato. Si adatta al sistema di visualizzazione, sia che sia uno schermo mobile piccolo o un grande monitor TV.

SVG si usa per creare forme geometriche comuni. Poiché D3 mappa i dati in una rappresentazione visiva, usa SVG per creare le forme della visualizzazione. Le forme SVG per una pagina web devono essere inserite all’interno di un tag HTML `svg`.

Il CSS può essere scalabile quando gli stili usano unità relative (come `vh`, `vw` o percentuali), ma usare SVG è più flessibile per costruire visualizzazioni dei dati.

# --instructions--

Aggiungi un nodo `svg` al `body` usando `append()`. Dagli un attributo `width` impostato sulla costante `w` fornita e un attributo `height` impostato sulla costante `h` fornita usando i metodi `attr()` o `style()` per ciascuno. Lo vedrai nell’output perché ha un `background-color` rosa applicato nel tag `style`.

**Nota:** Quando usi gli attributi `attr()` width e height non hanno unità. Questo è il blocco base della scalabilità: l’elemento avrà sempre un rapporto larghezza-altezza di 5:1, indipendentemente dal livello di zoom.

# --hints--

Il tuo documento dovrebbe avere 1 elemento `svg`.

```js
assert.lengthOf(document.querySelectorAll('svg'), 1);
```

L’elemento `svg` dovrebbe avere un attributo `width` impostato su `500` o essere stilizzato con una larghezza di `500px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const width = window.getComputedStyle(svg)['width'];
assert.isTrue(svg?.getAttribute('width') === '500' || width === '500px');
```

L’elemento `svg` dovrebbe avere un attributo `height` impostato su `100` o essere stilizzato con un’altezza di `100px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const height = window.getComputedStyle(svg)['height'];
assert.isTrue(svg?.getAttribute('height') === '100' || height === '100px');
```

# --seed--

## --seed-contents--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3.select('body');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
  </script>
</body>
```
