---
id: 587d7fa9367417b2b2512bcf
title: "Modificare dinamicamente l'altezza di ogni barra"
challengeType: 6
forumTopicId: 301486
dashedName: dynamically-change-the-height-of-each-bar
---

# --description--

L'altezza di ogni barra può essere impostata al valore del punto dati nell'array, in modo simile a come è stato impostato dinamicamente il valore `x`.

```js
selection.attr('property', (d, i) => {});
```

Qui `d` sarebbe il valore del punto dati, e `i` sarebbe l'indice del punto dati nell'array.

# --instructions--

Modifica la funzione di callback per l'attributo `height` in modo che restituisca il valore del dato moltiplicato per 3.

**Nota:** Ricorda che moltiplicare tutti i punti dati per la stessa costante scala i dati (come uno zoom). Questo aiuta a vedere meglio le differenze tra i valori delle barre in questo esempio.

# --hints--

La prima `rect` dovrebbe avere un `height` di `36`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('height'),
  '36'
);
```

La seconda `rect` dovrebbe avere un `height` di `93`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('height'),
  '93'
);
```

La terza `rect` dovrebbe avere un `height` di `66`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('height'),
  '66'
);
```

La quarta `rect` dovrebbe avere un `height` di `51`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('height'),
  '51'
);
```

La quinta `rect` dovrebbe avere un `height` di `75`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('height'),
  '75'
);
```

La sesta `rect` dovrebbe avere un `height` di `54`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('height'),
  '54'
);
```

La settima `rect` dovrebbe avere un `height` di `87`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('height'),
  '87'
);
```

L'ottava `rect` dovrebbe avere un `height` di `42`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('height'),
  '42'
);
```

La nona `rect` dovrebbe avere un `height` di `27`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('height'),
  '27'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        // Add your code below this line


        
        
        // Add your code above this line
      });
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        return d * 3;
      });
  </script>
</body>
```
