---
id: 587d7fa7367417b2b2512bc5
title: Lavorare con dati dinamici in D3
challengeType: 6
forumTopicId: 301498
dashedName: work-with-dynamic-data-in-d3
---

# --description--

Le ultime due sfide coprono le basi per mostrare dati dinamicamente con D3 usando i metodi `data()` e `enter()`. Questi metodi prendono un insieme di dati e, insieme al metodo `append()`, creano un nuovo elemento DOM per ogni voce nell'insieme di dati.

Nella sfida precedente, hai creato un nuovo elemento `h2` per ogni elemento nell'array `dataset`, ma tutti contenevano lo stesso testo, `New Title`. Questo perché non hai ancora utilizzato i dati associati a ciascuno degli elementi `h2`.

Il metodo `text()` di D3 può prendere come argomento una stringa o una funzione di callback:

```js
selection.text(d => d);
```

Nell'esempio sopra, il parametro `d` si riferisce a una singola voce nell'insieme di dati a cui è associata una selezione.

Usando l'esempio attuale come contesto, il primo elemento `h2` è associato a 12, il secondo elemento `h2` è associato a 31, il terzo elemento `h2` è associato a 22, e così via.

# --instructions--

Modifica il metodo `text()` in modo che ogni elemento `h2` mostri il valore corrispondente dall'array `dataset` seguito da uno spazio e dalla stringa `USD`. Per esempio, la prima intestazione dovrebbe essere `12 USD`.

# --hints--

La prima `h2` dovrebbe avere il testo `12 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[0]?.textContent, '12 USD');
```

La seconda `h2` dovrebbe avere il testo `31 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[1]?.textContent, '31 USD');
```

La terza `h2` dovrebbe avere il testo `22 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[2]?.textContent, '22 USD');
```

La quarta `h2` dovrebbe avere il testo `17 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[3]?.textContent, '17 USD');
```

La quinta `h2` dovrebbe avere il testo `25 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[4]?.textContent, '25 USD');
```

La sesta `h2` dovrebbe avere il testo `18 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[5]?.textContent, '18 USD');
```

La settima `h2` dovrebbe avere il testo `29 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[6]?.textContent, '29 USD');
```

L'ottava `h2` dovrebbe avere il testo `14 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[7]?.textContent, '14 USD');
```

La nona `h2` dovrebbe avere il testo `9 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[8]?.textContent, '9 USD');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      // Add your code below this line

      .text('New Title');

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => `${d} USD`);
  </script>
</body>
```
