---
id: 587d7fa9367417b2b2512bce
title: Impostare dinamicamente le coordinate per ogni barra
challengeType: 6
forumTopicId: 301487
dashedName: dynamically-set-the-coordinates-for-each-bar
---

# --description--

L'ultima sfida ha creato e aggiunto in coda un rettangolo all'elemento `svg` per ogni punto in `dataset` per rappresentare una barra. Purtroppo, erano tutte sovrapposte.

La posizione di un rettangolo è gestita dagli attributi `x` e `y`. Questi indicano a D3 dove iniziare a disegnare la forma nell'area `svg`. L'ultima sfida li ha impostati entrambi a 0, quindi ogni barra era posizionata nell'angolo in alto a sinistra.

Per un grafico a barre, tutte le barre dovrebbero trovarsi allo stesso livello verticale, il che significa che il valore `y` rimane lo stesso (a 0) per tutte le barre. Il valore `x`, invece, deve cambiare man mano che aggiungi nuove barre. Ricorda che valori `x` più grandi spingono gli elementi più a destra. Man mano che scorri gli elementi dell'array in `dataset`, il valore `x` dovrebbe aumentare.

Il metodo `attr()` in D3 accetta una funzione di callback per impostare dinamicamente quell'attributo. La funzione di callback prende due argomenti, uno per il punto dati stesso (di solito `d`) e uno per l'indice del punto dati nell'array. Il secondo argomento per l'indice è opzionale. Ecco il formato:

```js
selection.attr('property', (d, i) => {});
```

È importante notare che NON devi scrivere un loop `for` o usare `forEach()` per iterare sugli elementi nel set di dati. Ricorda che il metodo `data()` analizza il set di dati, e qualsiasi metodo concatenato dopo `data()` viene eseguito una volta per ogni elemento nel set di dati.

# --instructions--

Modifica la funzione di callback dell'attributo `x` in modo che restituisca l'indice moltiplicato per 30.

**Nota:** Ogni barra ha una larghezza di 25, quindi aumentare ogni valore `x` di 30 aggiunge un po' di spazio tra le barre. Qualsiasi valore maggiore di 25 funzionerebbe in questo esempio.

# --hints--

La prima `rect` dovrebbe avere un valore `x` di `0`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('x'),
  '0'
);
```

La seconda `rect` dovrebbe avere un valore `x` di `30`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('x'),
  '30'
);
```

La terza `rect` dovrebbe avere un valore `x` di `60`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('x'),
  '60'
);
```

La quarta `rect` dovrebbe avere un valore `x` di `90`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('x'),
  '90'
);
```

La quinta `rect` dovrebbe avere un valore `x` di `120`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('x'),
  '120'
);
```

La sesta `rect` dovrebbe avere un valore `x` di `150`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('x'),
  '150'
);
```

La settima `rect` dovrebbe avere un valore `x` di `180`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('x'),
  '180'
);
```

L'ottava `rect` dovrebbe avere un valore `x` di `210`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('x'),
  '210'
);
```

La nona `rect` dovrebbe avere un valore `x` di `240`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('x'),
  '240'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        // Add your code below this line


      
        // Add your code above this line
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        return i * 30;
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
