---
id: 587d7fa8367417b2b2512bca
title: Cambia la presentazione di un grafico a barre
challengeType: 6
forumTopicId: 301481
dashedName: change-the-presentation-of-a-bar-chart
---

# --description--

L'ultima sfida ha creato un grafico a barre, ma ci sono un paio di modifiche di formattazione che potrebbero migliorarlo:

1. Aggiungi spazio tra ogni barra per separarle visivamente, cosa che si fa aggiungendo un margine al CSS per la classe `bar`

2. Aumenta l'altezza delle barre per mostrare meglio la differenza nei valori, cosa che si fa moltiplicando il valore per un numero per scalare l'altezza

# --instructions--

Per prima cosa, aggiungi un `margin` di `2px` alla classe `bar` nel tag `style`. Poi, modifica la funzione di callback nel metodo `style()` in modo che ritorni un valore `10` volte il valore originale dei dati (più il `px`).

**Nota:** Moltiplicare ogni punto dati per la _stessa_ costante modifica solo la scala. È come fare uno zoom, e non cambia il significato dei dati sottostanti.

# --hints--

Il primo `div` dovrebbe avere un `height` di `120` pixel e un `margin` di `2` pixel.

```js
const firstDiv = document.querySelector('div');
assert.exists(firstDiv); 
const height = window.getComputedStyle(firstDiv)['height'];
const margin = window.getComputedStyle(firstDiv)['margin-right'];
assert.strictEqual(height, '120px');
assert.strictEqual(margin, '2px');
```

Il secondo `div` dovrebbe avere un `height` di `310` pixel e un `margin` di `2` pixel.

```js
const secondDiv = document.querySelectorAll('div')[1];
assert.exists(secondDiv); 
const height = window.getComputedStyle(secondDiv)['height'];
const margin = window.getComputedStyle(secondDiv)['margin-right'];
assert.strictEqual(height, '310px');
assert.strictEqual(margin, '2px');
```

Il terzo `div` dovrebbe avere un `height` di `220` pixel e un `margin` di `2` pixel.

```js
const thirdDiv = document.querySelectorAll('div')[2];
assert.exists(thirdDiv); 
const height = window.getComputedStyle(thirdDiv)['height'];
const margin = window.getComputedStyle(thirdDiv)['margin-right'];
assert.strictEqual(height, '220px');
assert.strictEqual(margin, '2px');
```

Il quarto `div` dovrebbe avere un `height` di `170` pixel e un `margin` di `2` pixel.

```js
const fourthDiv = document.querySelectorAll('div')[3];
assert.exists(fourthDiv); 
const height = window.getComputedStyle(fourthDiv)['height'];
const margin = window.getComputedStyle(fourthDiv)['margin-right'];
assert.strictEqual(height, '170px');
assert.strictEqual(margin, '2px');
```

Il quinto `div` dovrebbe avere un `height` di `250` pixel e un `margin` di `2` pixel.

```js
const fifthDiv = document.querySelectorAll('div')[4];
assert.exists(fifthDiv); 
const height = window.getComputedStyle(fifthDiv)['height'];
const margin = window.getComputedStyle(fifthDiv)['margin-right'];
assert.strictEqual(height, '250px');
assert.strictEqual(margin, '2px');
```

Il sesto `div` dovrebbe avere un `height` di `180` pixel e un `margin` di `2` pixel.

```js
const sixthDiv = document.querySelectorAll('div')[5];
assert.exists(sixthDiv); 
const height = window.getComputedStyle(sixthDiv)['height'];
const margin = window.getComputedStyle(sixthDiv)['margin-right'];
assert.strictEqual(height, '180px');
assert.strictEqual(margin, '2px');
```

Il settimo `div` dovrebbe avere un `height` di `290` pixel e un `margin` di `2` pixel.

```js
const seventhDiv = document.querySelectorAll('div')[6];
assert.exists(seventhDiv); 
const height = window.getComputedStyle(seventhDiv)['height'];
const margin = window.getComputedStyle(seventhDiv)['margin-right'];
assert.strictEqual(height, '290px');
assert.strictEqual(margin, '2px');
```

L'ottavo `div` dovrebbe avere un`height` di `140` pixel e un `margin` di `2` pixel.

```js
const eighthDiv = document.querySelectorAll('div')[7];
assert.exists(eighthDiv); 
const height = window.getComputedStyle(eighthDiv)['height'];
const margin = window.getComputedStyle(eighthDiv)['margin-right'];
assert.strictEqual(height, '140px');
assert.strictEqual(margin, '2px');
```

Il nono `div` dovrebbe avere un `height` di `90` pixel e un `margin` di `2` pixel.

```js
const ninthDiv = document.querySelectorAll('div')[8];
assert.exists(ninthDiv); 
const height = window.getComputedStyle(ninthDiv)['height'];
const margin = window.getComputedStyle(ninthDiv)['margin-right'];
assert.strictEqual(height, '90px');
assert.strictEqual(margin, '2px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    /* Add your code below this line */

    
    /* Add your code above this line */
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d + 'px'); // Change this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    margin: 2px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d * 10 + 'px');
  </script>
</body>
```
