---
id: 587d7fa9367417b2b2512bd0
title: Invertire gli elementi SVG
challengeType: 6
forumTopicId: 301488
dashedName: invert-svg-elements
---

# --description--

Avrai notato che il grafico a barre sembrava capovolto o invertito. Questo dipende da come SVG usa le coordinate (x, y).

In SVG, il punto di origine per le coordinate si trova nell'angolo in alto a sinistra. Una coordinata `x` pari a 0 posiziona una forma sul bordo sinistro dell'area SVG. Una coordinata `y` pari a 0 posiziona una forma sul bordo superiore dell'area SVG. Valori `x` più alti spostano il rettangolo verso destra. Valori `y` più alti spostano il rettangolo verso il basso.

Per far sì che le barre siano orientate correttamente, devi cambiare il modo in cui viene calcolata la coordinata `y`. Deve tenere conto sia dell'altezza della barra sia dell'altezza totale dell'area SVG.

L'altezza dell'area SVG è 100. Se hai un punto dati pari a 0 nell'insieme, vorresti che la barra partisse dal fondo dell'area SVG (non dalla cima). Per farlo, la coordinata `y` deve avere un valore di 100. Se il valore del punto dati fosse 1, inizieresti con una coordinata `y` di 100 per posizionare la barra in basso. Poi devi considerare l'altezza della barra pari a 1, quindi la coordinata `y` finale sarebbe 99.

La coordinata `y` che è `y = heightOfSVG - heightOfBar` posizionerebbe le barre correttamente orientate.

# --instructions--

Cambia la funzione di callback per l'attributo `y` per posizionare le barre correttamente. Ricorda che l'`height` della barra è 3 volte il valore del dato `d`.

**Nota:** In generale, la relazione è `y = h - m * d`, dove `m` è la costante che scala i punti dati.

# --hints--

La prima `rect` dovrebbe avere un valore `y` pari a `64`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('y'),
  (h - dataset[0] * 3).toString()
);
```

La seconda `rect` dovrebbe avere un valore `y` pari a `7`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('y'),
  (h - dataset[1] * 3).toString()
);
```

La terza `rect` dovrebbe avere un valore `y` pari a `34`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('y'),
  (h - dataset[2] * 3).toString()
);
```

La quarta `rect` dovrebbe avere un valore `y` pari a `49`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('y'),
  (h - dataset[3] * 3).toString()
);
```

La quinta `rect` dovrebbe avere un valore `y` pari a `25`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('y'),
  (h - dataset[4] * 3).toString()
);
```

La sesta `rect` dovrebbe avere un valore `y` pari a `46`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('y'),
  (h - dataset[5] * 3).toString()
);
```

La settima `rect` dovrebbe avere un valore `y` pari a `13`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('y'),
  (h - dataset[6] * 3).toString()
);
```

L'ottava `rect` dovrebbe avere un valore `y` pari a `58`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('y'),
  (h - dataset[7] * 3).toString()
);
```

La nona `rect` dovrebbe avere un valore `y` pari a `73`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('y'),
  (h - dataset[8] * 3).toString()
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => {
        // Add your code below this line


        
        // Add your code above this line
      })
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```
