---
id: 587d7faa367417b2b2512bd4
title: aggiungere un effetto al passaggio del mouse a un elemento d3
challengeType: 6
forumTopicId: 301469
dashedName: add-a-hover-effect-to-a-d3-element
---

# --description--

È possibile aggiungere effetti che evidenziano una barra quando l'utente ci passa sopra con il mouse. Finora, la definizione dello stile per i rettangoli è stata applicata con i metodi integrati di D3 e SVG, ma puoi anche usare il CSS.

Imposti la classe CSS sugli elementi SVG con il metodo `attr()`. Poi la pseudo-classe `:hover` per la tua nuova classe contiene le regole di stile per qualsiasi effetto al passaggio del mouse.

# --instructions--

Usa il metodo `attr()` per aggiungere una classe `bar` a tutti gli elementi `rect`. Questo cambia il colore `fill` della barra in marrone quando ci passi sopra con il mouse.

# --hints--

I tuoi elementi `rect` dovrebbero avere una classe `bar`.

```js
assert.isTrue(document.querySelector('rect')?.classList.contains('bar'));
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy')
    // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body>
```
