---
id: 68af272f60e009597dd861e1
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste workshop, você vai criar um catálogo de mídia. Comece definindo uma classe chamada `Movie`.

Dê à sua classe um método `__init__` com os seguintes parâmetros: `self`, `title`, `year`, `director` e `duration`. Dentro do método `__init__`, atribua os parâmetros `title`, `year`, `director` e `duration` aos atributos da instância com o mesmo nome.

# --hints--

Você deve criar uma classe chamada `Movie`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

Sua classe `Movie` deve ter um método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

Seu método `__init__` deve ter os seguintes parâmetros: `self`, `title`, `year`, `director` e `duration`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

Você deve atribuir `title` a `self.title` dentro do seu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

Você deve atribuir `year` a `self.year` dentro do seu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

Você deve atribuir `director` a `self.director` dentro do seu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

Você deve atribuir `duration` a `self.duration` dentro do seu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
