---
id: 68af274560e009597dd861e3
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Agora, crie uma instância de `Movie` passando argumentos de sua escolha e atribua-a a uma variável chamada `movie1`. Em seguida, imprima `movie1` no console.

# --hints--

Você deve ter uma variável chamada `movie1`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

Você deve atribuir uma instância de `Movie` à variável `movie1`.

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

Você deve imprimir `movie1` no console.

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
```
