---
id: 68af273960e009597dd861e2
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Você quer uma exibição legível para depuração e impressão. Crie um método `__str__` que retorne uma string com o formato `title (year) - duration min, director` (onde `title`, `year`, `duration` e `director` devem ser substituídos pelos valores dos atributos correspondentes) para que imprimir um filme forneça um resumo agradável.

# --hints--

Sua classe `Movie` deve ter um método `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

Seu método `__str__` deve retornar uma string no formato `title (year) - duration min, director`, onde `title`, `year`, `duration` e `director` devem ser substituídos pelos valores dos atributos correspondentes.

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    
--fcc-editable-region--
```
