---
id: 68af277e60e009597dd861e9
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Agora crie uma nova classe chamada `MediaCatalogue` e dê a ela um método `__init__` que defina um atributo `items` como uma lista vazia. Isso representará a lista de itens no catálogo. 

# --hints--

Você deve ter uma classe chamada `MediaCatalogue`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("MediaCatalogue")`)
})
```

Sua classe `MediaCatalogue` deve ter um método `__init__`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_class("MediaCatalogue").has_function("__init__")
    `)
})
```

Seu método `__init__` deve definir o atributo `items` como uma lista vazia.

```js
({
    test: () => runPython(`
    _catalogue = MediaCatalogue()
    assert _catalogue.items == []
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
