---
id: 68af277560e009597dd861e8
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Agora delete sua linha `print(movie1)`. Depois, coloque sua atribuição `movie1` em um bloco `try` e crie uma cláusula `except` que capture `ValueError as e` e use uma f-string para imprimir `Validation Error:` seguido de um espaço e a mensagem de erro.

Sinta-se à vontade para alterar os argumentos passados para instanciar `movie1` para ver as diferentes mensagens de erro.

# --hints--

Você não deve ter `print(movie1)` no seu código.

```js
({
  test: () => runPython(`assert not _Node(_code).has_call("print(movie1)")`)
})
```

Você deve ter um bloco `try` com `movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)` dentro dele.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().has_stmt("movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)")`)
})
```

Você deve ter uma cláusula `except` que capture `ValueError as e`.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].has_except("ValueError", "e")`)
})
```

Você deve usar uma f-string para imprimir `Validation Error:` seguido de um espaço e a mensagem de erro dentro do bloco `except`.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_except("ValueError", "e").has_call("print(f'Validation Error: {e}')")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
--fcc-editable-region--
```
