---
id: 68af274e60e009597dd861e4
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Nos passos anteriores, você criou um objeto básico de filme. Agora, você vai adicionar algumas validações para garantir que filmes não possam ser instanciados com atributos inválidos.

No início do método `__init__`, crie uma declaração `if` que verifique se o `title` está vazio ou contém apenas espaços em branco (você pode usar o método `strip` para construir a condição).

Se for o caso, levante um `ValueError` com a mensagem `Title cannot be empty`.

# --hints--

Você deve lançar um `ValueError` com a mensagem `Title cannot be empty` quando `title` estiver vazio ou contiver apenas espaços em branco.

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Você não deve lançar um `ValueError` quando `title` for uma string válida.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
