---
id: 68b80a995417adac579f9545
title: Passo 19
challengeType: 20
dashedName: step-19
---

# --description--

Agora, complete seu método `__init__` validando `seasons` e `total_episodes`. Crie uma declaração `if` que verifique se `seasons` é menor que `1` e lance um `ValueError` com a mensagem `Seasons must be 1 or greater`.

Então, faça o mesmo para validar `total_episodes`. Desta vez use a mensagem `Total episodes must be 1 or greater`.

# --hints--

Você deve lançar um `ValueError` com a mensagem `Seasons must be 1 or greater` quando `seasons` for menor que `1`.

```js
({
  test: () => runPython(`
  _msg = 'Seasons must be 1 or greater'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 0, 182)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, -1, 182)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Você não deve gerar um `ValueError` quando `seasons` for um número válido.

```js
({
    test: () => runPython(`
    try:
      TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 1, 182)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid seasons'
    `)
})
```

Você deve levantar um `ValueError` com a mensagem `Total episodes must be 1 or greater` quando `total_episodes` for menor que `1`.

```js
({
  test: () => runPython(`
  _msg = 'Total episodes must be 1 or greater'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 0)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, -1)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Você não deve gerar um `ValueError` quando `total_episodes` for um número válido.

```js
({
    test: () => runPython(`
    try:
      TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 1)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid total_episodes'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)
--fcc-editable-region--
        

        self.seasons = seasons
        self.total_episodes = total_episodes
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
