---
id: 68e76b1ce8216249d04df670
title: 构建一个 Planet 类
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 `Planet` 的类。
1. `Planet` 类应具有一个 `__init__` 方法，该方法：
   - 有四个参数：`self`、`name`、`planet_type` 和 `star`。
   - 如果传入的任何参数不是字串类型，则引起带有消息 `name, planet type, and star must be strings` 的 `TypeError`。
   - 如果传入的任何参数是空字串，则引起带有消息 `name, planet_type, and star must be non-empty strings` 的 `ValueError`。
   - 将传入的值赋值给实例属性 `name`、`planet_type` 和 `star`。
1. `Planet` 类应具有一个 `orbit` 方法，该方法返回格式为 `{name} is orbiting around {star}...` 的字串。
1. `Planet` 类应具有一个 `__str__` 方法，该方法返回格式为 `Planet: {name} | Type: {planet_type} | Star: {star}` 的字串。
1. 你应该创建三个名为 `planet_1`、`planet_2` 和 `planet_3` 的 `Planet` 类实例。
1. 你应该打印每个行星对象以查看 `__str__` 方法的输出。
1. 你应该对每个行星对象调用 `orbit` 方法并打印结果。

# --hints--

你应该创建一个名为 `Planet` 的类。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

`Planet` 类应有一个 `__init__` 方法。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

`__init__` 方法应有四个参数：`self`、`name`、`planet_type` 和 `star`，按此顺序。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

如果传入的第一个参数不是字串类型，`__init__` 方法应引起带有消息 `name, planet type, and star must be strings` 的 `TypeError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

如果传入的第二个参数不是字串类型，`__init__` 方法应引起带有消息 `name, planet type, and star must be strings` 的 `TypeError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

如果传入的第三个参数不是字串类型，`__init__` 方法应引起带有消息 `name, planet type, and star must be strings` 的 `TypeError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

如果传入的第一个参数是空字串，`__init__` 方法应引起带有消息 `name, planet_type, and star must be non-empty strings` 的 `ValueError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

如果传入的第二个参数是空字串，`__init__` 方法应引起带有消息 `name, planet_type, and star must be non-empty strings` 的 `ValueError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

如果传入的第三个参数是空字串，`__init__` 方法应引起带有消息 `name, planet_type, and star must be non-empty strings` 的 `ValueError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

你应该在 `__init__` 方法中将 `name` 赋值给 `self.name`。

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

你应该在 `__init__` 方法中将 `planet_type` 赋值给 `self.planet_type`。

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

你应该在 `__init__` 方法中将 `star` 赋值给 `self.star`。

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

`Planet` 类应具有一个 `orbit` 方法。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

`orbit` 方法应返回格式为 `{name} is orbiting around {star}...` 的字串。

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

`Planet` 类应具有一个 `__str__` 方法。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

`__str__` 方法应返回格式为 `Planet: {name} | Type: {planet_type} | Star: {star}` 的字串。

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

你应该创建三个名为 `planet_1`、`planet_2` 和 `planet_3` 的 `Planet` 类实例。

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

你应该打印每个行星对象以查看 `__str__` 方法的输出。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

你应该对每个行星对象调用 `orbit` 方法并打印结果。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
