---
id: a39963a4c10bc8b4d4f06d7e
title: تنفيذ دالة لإزالة القيم
challengeType: 26
dashedName: implement-a-value-remover-function
---

# --description--

في هذا المختبر، ستنشئ دالة تأخذ مصفوفة أولية كمعلمتها الأولى، تليها معلمة أو أكثر تمثل القيم التي يجب إزالتها.

**الهدف**: نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تمتلك دالة `destroyer` تقبل مصفوفة ومعلمة أو أكثر إضافية.
1. يجب أن تُرجع دالة `destroyer` مصفوفة جديدة تستبعد كل العناصر من المعلمة الأولى التي تطابق أيًا من المعلمات التالية.
1. يجب أن تقبل الدالة عددًا غير محدد من المعلمات.

# --hints--

يجب أن تمتلك دالة `destroyer`.

```js
assert.isFunction(destroyer);
```

يجب أن تُرجع `destroyer([1, 2, 3, 1, 2, 3], 2, 3)` القيمة `[1, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

يجب أن تُرجع `destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` القيمة `[1, 5, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

يجب أن تُرجع `destroyer([3, 5, 1, 2, 2], 2, 3, 5)` القيمة `[1]`.

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

يجب أن تُرجع `destroyer([2, 3, 2, 3], 2, 3)` القيمة `[]`.

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

يجب أن تُرجع `destroyer(["tree", "hamburger", 53], "tree", 53)` القيمة `["hamburger"]`.

```js
assert.deepEqual(destroyer(["tree", "hamburger", 53], "tree", 53), ["hamburger"]);
```

يجب أن تُرجع `destroyer( ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan" )` القيمة `[12, 92, 65]`.

```js
assert.deepEqual(
  destroyer(
    ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"],
    "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan"
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function destroyer(arr, ...valsToRemove) {
  return arr.filter(item => !valsToRemove.includes(item));
}
```

