---
id: 587d8257367417b2b2512c7e
title: Використання пошуку у глибину в двійковому дереві пошуку
challengeType: 1
forumTopicId: 301719
dashedName: use-depth-first-search-in-a-binary-search-tree
---

# --description--

Ми знаємо, як шукати конкретне значення у двійковому дереві пошуку. Але що, якщо ми просто хочемо дослідити все дерево? Або якщо у нас немає впорядкованого дерева і нам потрібно просто знайти певне значення? Тут ми познайомимося з методами обходу дерева, які можна використовувати для дослідження структур даних у вигляді дерева. Перший — пошук у глибину. При пошуку у глибину задане піддерево досліджується якомога глибше, перш ніж пошук продовжиться в іншому піддереві. Існує три способи це зробити: у порядку зростання: почати пошук з найлівнішого вузла і закінчити на найправішому вузлі. Прямий порядок обходу: дослідити всі корені перед листками. Зворотній порядок обходу: дослідити всі листки перед коренями. Як ви можете здогадатися, ви можете обирати різні методи пошуку залежно від того, які дані зберігає ваше дерево і що саме ви шукаєте. Для двійкового дерева пошуку обхід у порядку зростання повертає вузли у відсортованому порядку.

# --instructions--

Тут ми створимо ці три методи пошуку для нашого двійкового дерева пошуку. Пошук у глибину — це за своєю суттю рекурсивна операція, яка продовжує досліджувати подальші піддерева, доки присутні дочірні вузли. Коли ви зрозумієте цю базову концепцію, ви зможете просто змінити порядок, у якому досліджуєте вузли та піддерева, щоб отримати будь-який із трьох описаних вище пошуків. Наприклад, при зворотньому порядку обходу ми хочемо рекурсивно дійти до листкового вузла, перш ніж почати повертати будь-які вузли, тоді як при прямому порядку обходу ми спочатку повертаємо вузли, а потім продовжуємо рекурсію вниз по дереву. Визначте методи `inorder`, `preorder` і `postorder` у нашому дереві. Кожен із цих методів має повертати масив елементів, які представляють обхід дерева. Обов’язково повертайте цілі числа, що містяться у кожному вузлі, а не самі вузли. Нарешті, повертайте `null`, якщо дерево порожнє.

# --hints--

Структура даних `BinarySearchTree` має існувати.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Двійкове дерево пошуку має мати метод під назвою `inorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.inorder == 'function';
  })()
);
```

Двійкове дерево пошуку має мати метод під назвою `preorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.preorder == 'function';
  })()
);
```

Двійкове дерево пошуку має мати метод під назвою `postorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.postorder == 'function';
  })()
);
```

Метод `inorder` має повертати масив значень вузлів, отриманих в результаті обходу у порядку зростання.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.inorder().join('') == '012345678910';
  })()
);
```

Метод `preorder` має повертати масив значень вузлів, отриманих в результаті прямого порядку обходу.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.preorder().join('') == '710325469810';
  })()
);
```

Метод `postorder` має повертати масив значень вузлів, отриманих в результаті зворотнього порядку обходу.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.postorder().join('') == '024653181097';
  })()
);
```

Метод `inorder` має повертати `null` для порожнього дерева.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    return test.inorder() == null;
  })()
);
```

Метод `preorder` має повертати `null` для порожнього дерева.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    return test.preorder() == null;
  })()
);
```

Метод `postorder` має повертати `null` для порожнього дерева.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    return test.postorder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.result = [];

  this.inorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.inorder(node.left);
    this.result.push(node.value);
    if (node.right) this.inorder(node.right);
    return this.result;
  };
  this.preorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    this.result.push(node.value);
    if (node.left) this.preorder(node.left);
    if (node.right) this.preorder(node.right);
    return this.result;
  };
  this.postorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.postorder(node.left);
    if (node.right) this.postorder(node.right);
    this.result.push(node.value);

    return this.result;
  };
}
```
