---
id: 587d8251367417b2b2512c65
title: Видалити елементи зі зв’язаного списку за індексом
challengeType: 1
forumTopicId: 301711
dashedName: remove-elements-from-a-linked-list-by-index
---

# --description--

Перш ніж перейти до іншої структури даних, давайте трохи попрактикуємося зі зв’язаними списками.

Напишемо метод `removeAt`, який видаляє `element` за заданим `index`. Метод має називатися `removeAt(index)`. Щоб видалити `element` за певним `index`, нам потрібно вести підрахунок кожного вузла під час руху по зв’язаному списку.

Загальна техніка для ітерування по елементах зв’язаного списку полягає у використанні <dfn>«бігуна»</dfn>, або сторожового вузла, який «вказує» на вузли, які порівнює ваш код. У нашому випадку, починаючи з `head` списку, ми починаємо з змінної `currentIndex`, яка починається з `0`. `currentIndex` має збільшуватися на одиницю для кожного пройденого вузла.

Як і в нашому методі `remove(element)`, який <a href="/learn/coding-interview-prep/data-structures/remove-elements-from-a-linked-list" target="_blank" rel="noopener noreferrer nofollow">ми розглядали в попередньому уроці</a>, потрібно бути обережними, щоб не залишити залишок списку без зв’язку, коли ми видаляємо вузол у методі `removeAt(index)`. Ми підтримуємо вузли суміжними, переконавшись, що вузол, який має референс на видалений вузол, має референс на наступний вузол.

# --instructions--

Напишіть метод `removeAt(index)`, який видаляє і повертає вузол за заданим `index`. Метод має повертати `null`, якщо заданий `index` є від’ємним або більшим за або рівним `length` зв’язаного списку.

**Примітка:** Не забувайте вести підрахунок `currentIndex`.

# --hints--

Ваш клас `LinkedList` має мати метод `removeAt`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.removeAt === 'function';
  })()
);
```

Ваш метод `removeAt` має зменшувати `length` зв’язаного списку на одиницю.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.removeAt(1);
    return test.size() === 2;
  })()
);
```

Ваш метод `removeAt` має видаляти елемент за вказаним індексом зі зв’язаного списку.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.add('bird');
    test.removeAt(1);
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"kitten","next":{"element":"bird","next":null}}}'
    );
  })()
);
```

Коли у зв’язаному списку є лише один елемент, ваш метод `removeAt` має видаляти і повертати елемент за вказаним індексом і зменшувати довжину списку.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    var removedItem = test.removeAt(0);
    return test.head() === null && test.size() === 0 && removedItem === 'cat';
  })()
);
```

Ваш метод `removeAt` має повертати елемент видаленого вузла.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.removeAt(1) === 'dog';
  })()
);
```

Ваш метод `removeAt` має повертати `null`, і зв’язаний список не має змінюватися, якщо заданий індекс менший за `0`.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(-1);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

Ваш метод `removeAt` має повертати `null`, і зв’язаний список не має змінюватися, якщо заданий індекс більший або рівний `length` списку.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(3);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
      head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function (element) {
    this.element = element;
    this.next = null;
  };

  this.size = function () {
    return length;
  };

  this.head = function () {
    return head;
  };

  this.add = function (element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.removeAt = function (index) {
    var currentNode = head;
    var previous = head;
    var count = 0;
    if (index >= length || index < 0) {
      return null;
    }
    if (index === 0) {
      var removed = head.element;
      head = currentNode.next;
    } else {
      while (count < index) {
        previous = currentNode;
        currentNode = currentNode.next;
        count++;
      }
      var removed = previous.next.element;
      previous.next = currentNode.next;
    }
    length--;
    return removed;
  };
}
```
