---
id: 587d825b367417b2b2512c8d
title: Створити ES6 JavaScript map
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--

Нова версія JavaScript надає нам вбудований об’єкт Map, який забезпечує більшість функціональності, яку ми писали вручну в останньому завданні. Цей об’єкт Map, хоча й схожий на звичайні об’єкти JavaScript, має корисні можливості, яких немає у звичайних об’єктів. Наприклад, ES6 Map відстежує порядок додавання елементів. Ось більш повний огляд його методів:

- `.has(key)` повертає true або false залежно від наявності ключа
- `.get(key)` повертає значення, пов’язане з ключем
- `.set(key, value)` призначає нову пару ключ, значення
- `.delete(key)` видаляє пару ключ, значення
- `.clear()` видаляє всі пари ключ, значення
- `.keys()` повертає новий ітератор map, який містить усі ключі в порядку додавання
- `.values()` повертає новий ітератор map, який містить усі значення в порядку додавання
- `.entries()` повертає новий ітератор map, який містить усі пари ключ, значення в порядку додавання

# --instructions--

Оголосіть об’єкт JavaScript Map і призначте його змінній з ім’ям myMap. Додайте до нього пару ключ, значення `freeCodeCamp`, `Awesome!`.

# --hints--

Об’єкт `myMap` має існувати.

```js
assert(typeof myMap === 'object');
```

`myMap` має містити пару ключ, значення `freeCodeCamp`, `Awesome!`.

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```
