---
id: 587d8255367417b2b2512c72
title: Використання .has і .size у множині ES6
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

Розглянемо методи .has і .size, доступні в об’єкті множини ES6.

Спочатку створіть множину ES6

```js
var set = new Set([1,2,3]);
```

Метод .has перевіряє, чи міститься значення в множині.

```js
var hasTwo = set.has(2);
```

Метод .size повертає ціле число, яке представляє розмір множини

```js
var howBig = set.size;
```

# --instructions--

У цьому завданні ми передамо масив і значення у функцію checkSet(). Ваша функція має створити множину ES6 з аргументу масиву. Знайдіть, чи містить множина аргумент значення. Знайдіть розмір множини. І поверніть ці два значення у масиві.

# --hints--

`checkSet([4, 5, 6], 3)` має повернути [ false, 3 ]

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
