---
id: 587d825a367417b2b2512c87
title: Create a Doubly Linked List
challengeType: 1
forumTopicId: 301626
dashedName: create-a-doubly-linked-list
---

# --description--

Усі зв’язані списки, які ми створювали досі, були однонапрямленими. Тут ми створимо <dfn>двонапрямлений зв’язаний список</dfn>. Як випливає з назви, вузли у двонапрямленому зв’язаному списку мають референси на наступний і попередній вузол у списку.

Це дозволяє проходити список у обох напрямках, але також вимагає більше пам’яті, оскільки кожен вузол повинен містити додатковий референс на попередній вузол у списку.

# --instructions--

Ми надали об’єкт `Node` і почали наш `DoublyLinkedList`. Додайте два методи до нашого двонапрямленого зв’язаного списку, які називаються `add` і `remove`. Метод `add` повинен додавати заданий елемент до списку, а метод `remove` повинен видаляти всі входження заданого елемента зі списку.

Будьте уважні, щоб опрацювати всі можливі крайні випадки при написанні цих методів, наприклад, видалення першого або останнього елемента. Також видалення будь-якого елемента зі списку, що порожній, має повертати `null`.

# --hints--

Структура даних `DoublyLinkedList` має існувати.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

`DoublyLinkedList` має мати метод під назвою `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.add == undefined) {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

`DoublyLinkedList` має мати метод під назвою `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.remove == undefined) {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Видалення елемента з порожнього списку має повертати `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.remove(100) == null;
  })()
);
```

Метод `add` має додавати елементи до списку.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(5);
    test.add(6);
    test.add(723);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '56723';
  })()
);
```

Кожен вузол має зберігати референс на попередній вузол.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(50);
    test.add(68);
    test.add(73);
    var result = [];
    var curr = test.tail;
    while (curr) {
      result.push(curr.data);
      curr = curr.prev;
    }
    return result.join('') == '736850';
  })()
);
```

Перший елемент має бути видаленим зі списку.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(25);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '3560';
  })()
);
```

Останній елемент має бути видаленим зі списку.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(60);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '2535';
  })()
);
```


# --seed--

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
