---
id: 587d8258367417b2b2512c80
title: Видалити листовий вузол у двійковому дереві пошуку
challengeType: 1
forumTopicId: 301637
dashedName: delete-a-leaf-node-in-a-binary-search-tree
---

# --description--

Це перше з трьох завдань, у яких ми реалізуємо складнішу операцію у двійкових деревах пошуку: видалення. Видалення складне, бо видалення вузлів порушує посилання в дереві. Ці посилання потрібно ретельно відновити, щоб зберегти структуру двійкового дерева. Для деяких видалень це означає, що дерево доведеться переставити. Загалом, ви зіткнетеся з одним із трьох випадків при спробі видалити вузол: Листовий вузол: ціль для видалення не має дочірніх елементів. Один дочірній елемент: ціль для видалення має лише одного дочірнього елемента. Два дочірні елементи: ціль для видалення має два дочірні вузли. Видалити листовий вузол легко — ми просто його видаляємо. Видалення вузла з одним дочірнім елементом також відносно просте: ми просто видаляємо його і з’єднуємо батьківський елемент з дочірнім вузлом видаленого вузла. Видалення вузла з двома дочірніми елементами складніше, бо це створює два дочірні вузли, які потрібно знову підключити до батьківського дерева. Ми розглянемо, як впоратися з цим випадком у третьому завданні. Крім того, потрібно враховувати деякі крайні випадки при обробці видалення. Що, якщо дерево порожнє? Що, якщо вузол для видалення — кореневий вузол? Що, якщо у дереві лише два елементи? Наразі давайте розглянемо перший випадок, коли ми видаляємо листовий вузол.

# --instructions--

Створіть метод у нашому двійковому дереві під назвою `remove`. Тут ми побудуємо логіку операції видалення. Спочатку потрібно створити функцію всередині remove, яка знаходить вузол, який ми хочемо видалити, у поточному дереві. Якщо вузол відсутній у дереві, `remove` має повернути `null`. Якщо цільовий вузол — листовий, без дочірніх елементів, тоді посилання батьківського елемента на нього має бути встановлене в `null`. Це ефективно видаляє вузол з дерева. Для цього вам також доведеться відслідковувати батьківський елемент вузла, який ми хочемо видалити. Корисно буде створити спосіб відслідковувати кількість дочірніх елементів цільового вузла, оскільки це визначить, до якого випадку належить наше видалення. Другий і третій випадки ми розглянемо у наступних завданнях. Успіхів!

# --hints--

Структура даних `BinarySearchTree` має існувати.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Двійкове дерево пошуку має мати метод під назвою `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Спроба видалити елемент з порожнього дерева має повертати `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    return test.remove(100) == null;
  })()
);
```

Спроба видалити елемент, якого немає, має повертати `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(15);
    test.add(30);
    return test.remove(100) == null;
  })()
);
```

Якщо кореневий вузол не має дочірніх елементів, його видалення має встановити корінь у `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(500);
    test.remove(500);
    return test.inorder() == null;
  })()
);
```

Метод `remove` має видаляти листові вузли з дерева.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(5);
    test.add(3);
    test.add(7);
    test.add(6);
    test.add(10);
    test.add(12);
    test.remove(3);
    test.remove(12);
    test.remove(10);
    return test.inorder().join('') == '567';
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    },
    inorder: function() {
      if (this.root == null) {
        return null;
      } else {
        var result = new Array();
        function traverseInOrder(node) {
          if (node.left != null) {
            traverseInOrder(node.left);
          }
          result.push(node.value);
          if (node.right != null) {
            traverseInOrder(node.right);
          }
        }
        traverseInOrder(this.root);
        return result;
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}

function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
}
```

# --solutions--

```js
// solution required
```
