---
id: 587d8251367417b2b2512c62
title: створити клас зв’язаного списку
challengeType: 1
forumTopicId: 301628
dashedName: create-a-linked-list-class
---

# --description--

Давайте створимо клас `linked list`. Кожен зв’язаний список повинен починатися з кількох базових властивостей: `head` (перший елемент у вашому списку) та `length` (кількість елементів у вашому списку). Іноді ви можете побачити реалізації зв’язаних списків, які включають `tail` для останнього елемента списку, але наразі ми обмежимося цими двома. Щоразу, коли ми додаємо елемент до зв’язаного списку, наша властивість `length` має збільшуватися на одиницю.

Нам знадобиться спосіб додавати елементи до нашого зв’язаного списку, тому першим методом, який ми хочемо створити, буде метод `add`.

Якщо наш список порожній, додати елемент до зв’язаного списку досить просто: ми просто обгортаємо цей елемент у клас `Node` і призначаємо цей вузол властивості `head` нашого зв’язаного списку.

Але що, якщо в нашому списку вже є один або більше елементів? Як додати елемент до списку? Пам’ятайте, що кожен вузол у зв’язаному списку має властивість `next`. Щоб додати вузол до списку, знайдіть останній вузол у списку і вкажіть властивість `next` цього останнього вузла на наш новий вузол. (Підказка: ви знаєте, що досягли кінця зв’язаного списку, коли властивість `next` вузла дорівнює `null`.)

# --instructions--

Напишіть метод add, який призначає перший вузол, який ви додаєте до зв’язаного списку, властивості `head`; після цього, при додаванні кожного нового вузла, кожен вузол має бути референсом у властивості `next` попереднього вузла.

Примітка

Властивість `length` вашого списку має збільшуватися на одиницю щоразу, коли елемент додається до зв’язаного списку.

# --hints--

Ваш клас `LinkedList` має мати метод `add`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.add === 'function';
  })()
);
```

Ваш клас `LinkedList` має призначати `head` першому доданому вузлу.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    return test.head().element === 'cat';
  })()
);
```

Попередній вузол `node` у вашому класі `LinkedList` має містити референс на найновіший створений вузол.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('fish');
    return test.head().next.element === 'dog' && test.head().next.next.element === 'fish';
  })()
);
```

Властивість `size` вашого класу `LinkedList` має дорівнювати кількості вузлів у зв’язаному списку.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.size() === 2;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() { 
  var length = 0; 
  var head = null; 

  var Node = function(element){
    this.element = element; 
    this.next = null; 
  }; 

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line
    if (head == null) {
      head = new Node(element);
    } 
    else {
      let currentNode = head;
      while (currentNode.next != null) {
        // currentNode.next will be last node of linked list after loop
        currentNode = currentNode.next;
      }
      currentNode.next = new Node(element);
    }
    length++;
    // Only change code above this line
  };
}
```
