---
id: 587d825d367417b2b2512c96
title: Пошук у глибину
challengeType: 1
forumTopicId: 301640
dashedName: depth-first-search
---

# --description--

Подібно до <dfn>пошуку у ширину</dfn>, тут ми вивчимо ще один алгоритм обходу графа, який називається <dfn>пошук у глибину</dfn>.

Якщо пошук у ширину досліджує ребра, поступово віддаляючись від початкового вузла, то <dfn>пошук у глибину</dfn> спочатку йде по шляху ребер якомога глибше.

Коли він досягає кінця шляху, пошук повертається назад до останнього вузла, у якого є неперевірений шлях, і продовжує пошук.

Анімація нижче показує, як працює алгоритм. Він починається з верхнього вузла і відвідує вузли у порядку їх номерів.

<img alt="animation showing the depth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/depth-first-search.gif'>

Зверніть увагу, що, на відміну від пошуку у ширину, кожного разу, коли відвідується вузол, він не відвідує всіх його сусідів. Замість цього спочатку відвідує одного з сусідів і продовжує рухатися цим шляхом, доки не залишиться вузлів для відвідування на цьому шляху.

Для реалізації цього алгоритму варто використати стек. Стек — це масив, у якому останній доданий елемент видаляється першим. Це також відома як структура даних <dfn>Last-In-First-Out</dfn>. Стек корисний у алгоритмах пошуку у глибину, бо, додаючи сусідів у стек, ми хочемо спочатку відвідати найновіших сусідів і видалити їх зі стеку.

Простим результатом цього алгоритму є список вузлів, до яких можна дістатися з заданого вузла. Тому також потрібно відслідковувати вузли, які ви відвідуєте.

# --instructions--

Напишіть функцію `dfs()`, яка приймає неорієнтовану матрицю суміжності `graph` і мітку вузла `root` як параметри. Мітка вузла — це просто числове значення вузла між `0` і `n - 1`, де `n` — загальна кількість вузлів у графі.

Ваша функція має повертати масив усіх вузлів, до яких можна дістатися з `root`.

# --hints--

Вхідний граф `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` із початковим вузлом `1` має повернути масив із `0`, `1`, `2` і `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })(),
  [0, 1, 2, 3]
);
```

Вхідний граф `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` із початковим вузлом `3` має повернути масив із `3`, `2`, `1` і `0`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3, 2, 1, 0]
);
```

Вхідний граф `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` із початковим вузлом `1` має повернути масив із чотирьох елементів.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })().length === 4
);
```

Вхідний граф `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` із початковим вузлом `3` має повернути масив із `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3]
);
```

Вхідний граф `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` із початковим вузлом `3` має повернути масив із одного елемента.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })().length === 1
);
```

Вхідний граф `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` із початковим вузлом `3` має повернути масив із `2` і `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [2, 3]
);
```

Вхідний граф `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` із початковим вузлом `3` має повернути масив із двох елементів.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })().length === 2
);
```

Вхідний граф `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` із початковим вузлом `0` має повернути масив із `0` і `1`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })(),
  [0, 1]
);
```

Вхідний граф `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` із початковим вузлом `0` має повернути масив із двох елементів.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })().length === 2
);
```

# --seed--

## --seed-contents--

```js
function dfs(graph, root) {

}

var exDFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(dfs(exDFSGraph, 3));
```

# --solutions--

```js
function dfs(graph, root) {
    var stack = [];
    var tempV;
    var visited = [];
    var tempVNeighbors = [];
    stack.push(root);
    while (stack.length > 0) {
        tempV = stack.pop();
        if (visited.indexOf(tempV) == -1) {
            visited.push(tempV);
            tempVNeighbors = graph[tempV];
            for (var i = 0; i < tempVNeighbors.length; i++) {
                if (tempVNeighbors[i] == 1) {
                    stack.push(i);
                }
            }
        }
    }
    return visited;
}
```
