---
id: 587d8254367417b2b2512c70
title: Створення та додавання до множин у ES6
challengeType: 1
forumTopicId: 301636
dashedName: create-and-add-to-sets-in-es6
---

# --description--

Тепер, коли ви опрацювали ES5, ви зробите щось подібне в ES6. Це буде значно простіше. ES6 містить вбудовану структуру даних `Set`, тому багато операцій, які ви писали вручну, тепер включені для вас. Давайте подивимося:

Щоб створити нову порожню множину:

```js
var set = new Set();
```

Ви можете створити множину з одним значенням:

```js
var set = new Set(1);
```

Ви можете створити множину з масиву:

```js
var set = new Set([1, 2, 3]);
```

Після створення множини ви можете додати потрібні значення за допомогою методу `add`:

```js
var set = new Set([1, 2, 3]);
set.add([4, 5, 6]);
```

Нагадаємо, множина — це структура даних, яка не може містити дублікати значень:

```js
var set = new Set([1, 2, 3, 1, 2, 3]);
// set contains [1, 2, 3] only
```

# --instructions--

Для цього завдання поверніть множину з такими значеннями: `1, 2, 3, 'Taco', 'Cat', 'Awesome'`

# --hints--

Ваша `Set` має містити лише значення `1, 2, 3, Taco, Cat, Awesome`.

```js
assert(
  (function () {
    var test = checkSet();
    return (
      test.size == 6 &&
      test.has(1) &&
      test.has(2) &&
      test.has(3) &&
      test.has('Taco') &&
      test.has('Cat') &&
      test.has('Awesome')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet() {
  var set = new Set([1, 2, 3, 3, 2, 1, 2, 3, 1]);
  // Only change code below this line

  // Only change code above this line
  console.log(Array.from(set));
  return set;
}

checkSet();
```

# --solutions--

```js
function checkSet(){var set = new Set([1,2,3,'Taco','Cat','Awesome']);
return set;}
```
