---
id: 587d8255367417b2b2512c73
title: Використання оператора розширення та нотаток для інтеграції ES5 Set()
challengeType: 1
forumTopicId: 301720
dashedName: use-spread-and-notes-for-es5-set-integration
---

# --description--

Ви пам’ятаєте оператор розширення ES6 `...`?

`...` може брати ітеровані об’єкти в ES6 і перетворювати їх на масиви.

Давайте створимо множину Set і подивимося на функцію розширення.

```js
var set = new Set([1,2,3]);
var setToArr = [...set]
console.log(setToArr) // returns [ 1, 2, 3 ]
```

# --instructions--

У цьому завданні ми передамо об’єкт множини до функції `checkSet`. Вона має повернути масив, що містить значення множини Set.

Тепер ви успішно навчилися використовувати об’єкт `Set()` ES6, гарна робота!

# --hints--

`checkSet(new Set([1,2,3,4,5,6,7])` має повертати `[1, 2, 3, 4, 5, 6, 7]`.

```js
var test = checkSet(new Set([1, 2, 3, 4, 5, 6, 7]));
assert.deepEqual(test, [1, 2, 3, 4, 5, 6, 7]);
```

# --seed--

## --seed-contents--

```js
function checkSet(set){
   // Only change code below this line

   // Only change code above this line
}
```

# --solutions--

```js
function checkSet(set){
return [...set];}
```
