---
id: 587d8255367417b2b2512c74
title: створити клас пріоритетної черги
challengeType: 1
forumTopicId: 301630
dashedName: create-a-priority-queue-class
---

# --description--

У цьому завданні ви створите пріоритетну чергу. Пріоритетна черга — це особливий тип черги, у якій елементи можуть мати додаткову інформацію, що вказує на їхній пріоритет. Це можна просто представити цілим числом. Пріоритет елемента заміщатиме порядок розміщення при визначенні послідовності вилучення елементів із черги. Якщо елемент із вищим пріоритетом буде доданий після елементів із нижчим пріоритетом, то елемент із вищим пріоритетом буде вилучений першим.

Наприклад, уявімо, що у нас є пріоритетна черга з трьома елементами:

```js
[['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Тут друге значення (ціле число) представляє пріоритет елемента. Якщо ми додамо `['human', 1]` з пріоритетом `1` (припускаючи, що нижчі пріоритети мають перевагу), він буде першим елементом, який буде вилучено. Колекція виглядатиме так:

```js
[['human', 1], ['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Ми почали писати `PriorityQueue` у редакторі коду. Вам потрібно додати метод `enqueue` для додавання елементів із пріоритетом, метод `dequeue` для вилучення та повернення елементів, метод `size` для повернення кількості елементів у черзі, метод `front` для повернення елемента на початку черги і нарешті метод `isEmpty`, який поверне `true`, якщо черга порожня, або `false`, якщо ні.

Метод `enqueue` повинен приймати елементи у форматі, показаному вище (`['human', 1]`), де `1` представляє пріоритет. Методи `dequeue` і `front` повинні повертати лише ім’я елемента, а не його пріоритет.

# --hints--

Ваш клас `PriorityQueue` повинен мати метод `enqueue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.enqueue === 'function';
  })()
);
```

Ваш клас `PriorityQueue` повинен мати метод `dequeue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.dequeue === 'function';
  })()
);
```

Ваш клас `PriorityQueue` повинен мати метод `size`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.size === 'function';
  })()
);
```

Ваш клас `PriorityQueue` повинен мати метод `front`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.front === 'function';
  })()
);
```

Ваш клас `PriorityQueue` повинен мати метод `isEmpty`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Ваш клас `PriorityQueue` повинен правильно відслідковувати поточну кількість елементів за допомогою методу `size` під час додавання та вилучення елементів.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    test.enqueue(['Jon Snow', 1]);
    var size1 = test.size();
    test.dequeue();
    var size2 = test.size();
    test.enqueue(['A', 3]);
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    return size1 === 2 && size2 === 1 && test.size() === 4;
  })()
);
```

Метод `front` повинен повертати правильний елемент на початку черги під час додавання та вилучення елементів.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    var front1 = test.front();
    test.enqueue(['Jon Snow', 1]);
    var front2 = test.front();
    test.dequeue();
    test.enqueue(['A', 3]);
    var front3 = test.front();
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    test.dequeue();
    var front4 = test.front();
    return (
      front1 === 'David Brown' &&
      front2 === 'Jon Snow' &&
      front3 === 'David Brown' &&
      front4 === 'A'
    );
  })()
);
```

Метод `isEmpty` повинен повертати `true`, коли черга порожня.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 1]);
    test.enqueue(['B', 1]);
    test.dequeue();
    var first = test.isEmpty();
    test.dequeue();
    return !first && test.isEmpty();
  })()
);
```

Пріоритетна черга повинна повертати елементи з вищим пріоритетом перед елементами з нижчим пріоритетом і повертати елементи у порядку першого додавання, якщо пріоритети однакові.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 5]);
    test.enqueue(['B', 5]);
    test.enqueue(['C', 5]);
    test.enqueue(['D', 3]);
    test.enqueue(['E', 1]);
    test.enqueue(['F', 7]);
    var result = [];
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    return result.join('') === 'EDABCF';
  })()
);
```

# --seed--

## --seed-contents--

```js
function PriorityQueue () {
  this.collection = [];
  this.printCollection = function() {
    console.log(this.collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function PriorityQueue() {
  this.collection = [];
  this.printCollection = function () {
    console.log(this.collection);
  };
  // Only change code below this line
  this.enqueue = function (newitem) {
    if (this.isEmpty()) {
      return this.collection.push(newitem);
    }

    this.collection = this.collection.reverse();
    var found_index = this.collection.findIndex(function (item) {
      return newitem[1] >= item[1];
    });
    if (found_index === -1) {
      this.collection.push(newitem);
    } else {
      this.collection.splice(found_index, 0, newitem);
    }
    this.collection = this.collection.reverse();
  };
  this.dequeue = function () {
    if (!this.isEmpty()) {
      return this.collection.shift()[0];
    } else {
      return "The queue is empty.";
    }
  };
  this.size = function () {
    return this.collection.length;
  };
  this.front = function () {
    return this.collection[0][0];
  };
  this.isEmpty = function () {
    return this.size() > 0 ? false : true;
  };
  // Only change code above this line
}
```
