---
id: 587d825c367417b2b2512c90
title: Пошук у ширину
challengeType: 1
forumTopicId: 301622
dashedName: breadth-first-search
---

# --description--

До цього моменту ми вивчили різні способи створення представлень графів. Що далі? Одне з природних питань — які відстані між будь-якими двома вузлами в графі? Тут на допомогу приходять <dfn>алгоритми обходу графа</dfn>.

<dfn>Алгоритми обходу</dfn> — це алгоритми для проходження або відвідування вузлів у графі. Один із типів алгоритмів обходу — це алгоритм пошуку у ширину.

Цей алгоритм починається з одного вузла і відвідує всіх його сусідів, які знаходяться на відстані одного ребра. Потім він переходить до відвідування кожного з їхніх сусідів і так далі, доки не будуть досягнуті всі вузли.

Важливою структурою даних, яка допоможе реалізувати алгоритм пошуку у ширину, є черга. Це масив, у який можна додавати елементи з одного кінця і видаляти з іншого. Це також відомо як <dfn>FIFO</dfn> або <dfn>перший прийшов — перший вийшов</dfn> структура даних.

Візуально алгоритм працює так:
<img alt="animation showing the breadth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/breadth-first-search.gif'>

Сіре затінення означає, що вузол додається до черги, а чорне — що вузол видаляється з черги. Зверніть увагу, що кожного разу, коли вузол видаляється з черги (вузол стає чорним), усі його сусіди додаються до черги (вузол стає сірим).

Щоб реалізувати цей алгоритм, вам потрібно буде ввести структуру графа і вузол, з якого ви хочете почати.

Спочатку потрібно знати відстані від початкового вузла або кількість ребер до нього. Всі відстані слід ініціалізувати великим числом, наприклад `Infinity`. Це запобігає проблемам підрахунку, якщо вузол недосяжний з початкового вузла. Далі потрібно перейти від початкового вузла до його сусідів. Ці сусіди знаходяться на відстані одного ребра, і на цьому етапі слід додати одиницю до відстаней, які ви відстежуєте.

# --instructions--

Напишіть функцію `bfs()`, яка приймає граф у вигляді матриці суміжності (двовимірний масив) та мітку вузла root як параметри. Мітка вузла — це просто ціле число вузла між `0` і `n - 1`, де `n` — загальна кількість вузлів у графі.

Ваша функція повинна повертати об’єкт JavaScript із парами ключ-значення, де ключ — це вузол, а значення — відстань від root. Якщо вузол недосяжний, його відстань має бути `Infinity`.

# --hints--

Вхідний граф `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` із початковим вузлом `1` має повернути `{0: 1, 1: 0, 2: 1, 3: 2}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 1);
assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: 2 });
```

Вхідний граф `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` із початковим вузлом `1` має повернути `{0: 1, 1: 0, 2: 1, 3: Infinity}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 0],
  [0, 0, 0, 0]
];
var results = bfs(graph, 1);
 assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: Infinity });
```

Вхідний граф `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` із початковим вузлом `0` має повернути `{0: 0, 1: 1, 2: 2, 3: 3}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 0);
 assert.deepEqual(results, { 0: 0, 1: 1, 2: 2, 3: 3 });
```

Вхідний граф `[[0, 1], [1, 0]]` із початковим вузлом `0` має повернути `{0: 0, 1: 1}`

```js
var graph = [
  [0, 1],
  [1, 0]
];
var results = bfs(graph, 0);
assert.deepEqual(results, { 0: 0, 1: 1 });
```

# --seed--

## --seed-contents--

```js
function bfs(graph, root) {
  var nodesLen = {};

  return nodesLen;
};

var exBFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(bfs(exBFSGraph, 3));
```

# --solutions--

```js
function bfs(graph, root) {
  var nodesLen = {};
  // Set all distances to infinity
  for (var i = 0; i < graph.length; i++) {
    nodesLen[i] = Infinity;
  }
  nodesLen[root] = 0; // ...except root node
  var queue = [root]; // Keep track of nodes to visit
  var current; // Current node traversing
  // Keep on going until no more nodes to traverse
  while (queue.length !== 0) {
    current = queue.shift();
    // Get adjacent nodes from current node
    var curConnected = graph[current]; // Get layer of edges from current
    var neighborIdx = []; // List of nodes with edges
    var idx = curConnected.indexOf(1); // Get first edge connection
    while (idx !== -1) {
      neighborIdx.push(idx); // Add to list of neighbors
      idx = curConnected.indexOf(1, idx + 1); // Keep on searching
    }
    // Loop through neighbors and get lengths
    for (var j = 0; j < neighborIdx.length; j++) {
      // Increment distance for nodes traversed
      if (nodesLen[neighborIdx[j]] === Infinity) {
        nodesLen[neighborIdx[j]] = nodesLen[current] + 1;
        queue.push(neighborIdx[j]); // Add new neighbors to queue
      }
    }
  }
  return nodesLen;
}
```
