---
id: 587d825a367417b2b2512c8a
title: Вставити елемент у максимальну купу
challengeType: 1
forumTopicId: 301703
dashedName: insert-an-element-into-a-max-heap
---

# --description--

Тепер ми перейдемо до іншої структури даних — двійкової купи. Двійкова купа — це частково впорядковане двійкове дерево, яке задовольняє властивість купи. Властивість купи визначає відношення між батьківськими та дочірніми вузлами. Можна мати максимальну купу, у якій усі батьківські вузли більші або рівні своїм дочірнім вузлам, або мінімальну купу, у якій навпаки. Двійкові купи також є повними двійковими деревами. Це означає, що всі рівні дерева заповнені повністю, а якщо останній рівень заповнений частково, то він заповнюється зліва направо.

Хоча двійкові купи можуть бути реалізовані як деревоподібні структури з вузлами, що містять посилання на лівий і правий дочірні вузли, часткове впорядкування згідно з властивістю купи дозволяє нам представляти купу у вигляді масиву. Нас цікавить відношення батько-дитина, і за допомогою простої арифметики ми можемо обчислити дочірні вузли будь-якого батьківського вузла та батьківський вузол будь-якого дочірнього.

Наприклад, розглянемо таке представлення двійкової мінімальної купи у вигляді масиву:

```js
[ 6, 22, 30, 37, 63, 48, 42, 76 ]
```

Кореневий вузол — це перший елемент, `6`. Його діти — `22` і `30`. Якщо подивитися на відношення між індексами масиву цих значень, для індексу `i` діти — `2 * i + 1` і `2 * i + 2`. Аналогічно, елемент з індексом `0` є батьком для двох дочірніх елементів з індексами `1` і `2`. Загалом, ми можемо знайти батька вузла з будь-яким індексом за формулою: `Math.floor((i - 1) / 2)`. Ці шаблони будуть справедливі, коли двійкове дерево зростатиме до будь-якого розміру. Нарешті, ми можемо зробити невелику правку, щоб спростити цю арифметику, пропускаючи перший елемент у масиві. Це створює таке відношення для будь-якого елемента з індексом `i`:

Приклад представлення масиву:

```js
[ null, 6, 22, 30, 37, 63, 48, 42, 76 ]
```

Лівий дочірній елемент елемента: `i * 2`

Правий дочірній елемент елемента: `i * 2 + 1`

Батьківський елемент елемента: `Math.floor(i / 2)`

Як тільки ви зрозумієте цю математику, використання представлення у вигляді масиву дуже корисне, бо розташування вузлів можна швидко визначити за допомогою цієї арифметики, а використання пам’яті зменшується, оскільки не потрібно зберігати посилання на дочірні вузли.

# --instructions--

Інструкції: Тут ми створимо максимальну купу. Почніть із створення методу `insert`, який додає елементи до нашої купи. Під час вставки важливо завжди підтримувати властивість купи. Для максимальної купи це означає, що кореневий елемент завжди повинен мати найбільше значення в дереві, а всі батьківські вузли мають бути більшими за своїх дітей. Для реалізації купи у вигляді масиву це зазвичай виконується у три кроки:

<ol>
  <li>Додайте новий елемент у кінець масиву.</li>
  <li>Якщо елемент більший за свого батька, поміняйте їх місцями.</li>
  <li>Продовжуйте міняти місцями, доки новий елемент не стане меншим за свого батька або доки не досягнете кореня дерева.</li>
</ol>

Нарешті, додайте метод `print`, який повертає масив усіх елементів, доданих до купи.

# --hints--

Структура даних `MaxHeap` має існувати.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap` має мати метод `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap` має мати метод `print`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

Метод `insert` має додавати елементи відповідно до властивості максимальної купи.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    test.insert(50);
    test.insert(100);
    test.insert(700);
    test.insert(32);
    test.insert(51);
    test.insert(800);
    const result = test.print();
    const solution = JSON.stringify([null,800,51,700,32,50,100]);
    const solutionWithoutNull = JSON.stringify([800,51,700,32,50,100]);

    return (result.length == 6) ? (JSON.stringify(result) == solutionWithoutNull) : (JSON.stringify(result) == solution);
  })()
);
```

# --seed--

## --seed-contents--

```js
var MaxHeap = function() {
  // Only change code below this line
    
  // Only change code above this line
};
```

# --solutions--

```js
var MaxHeap = function() {
    // Only change code below this line
    this.heap = [];
    this.parent = index => {
      return Math.floor((index - 1) / 2);
    }
    this.insert = element => {
      this.heap.push(element);
      this.heapifyUp(this.heap.length - 1);
    }
    this.heapifyUp = index => {
      let currentIndex = index,
      parentIndex = this.parent(currentIndex);
      while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
        this.swap(currentIndex, parentIndex);
        currentIndex = parentIndex;
        parentIndex = this.parent(parentIndex);
      }
    }
    this.swap = (index1, index2) => {
      [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
    }
    this.print = () => {
      return this.heap;
    }
    // Only change code above this line
};
```
