---
id: 587d825b367417b2b2512c8e
title: Створення хеш-таблиці
challengeType: 1
forumTopicId: 301627
dashedName: create-a-hash-table
---

# --description--

У цьому завданні ми познайомимося з хеш-таблицями. Хеш-таблиця використовується для реалізації асоціативних масивів або відображень пар ключ-значення, як-от об’єкти та Maps, які ми щойно вивчали. Наприклад, об’єкт JavaScript може бути реалізований як хеш-таблиця (його фактична реалізація залежить від середовища, в якому він працює). Принцип роботи хеш-таблиці полягає в тому, що вона приймає ключ як вхідні дані і детерміновано хешує цей ключ у деяке числове значення. Це числове значення потім використовується як фактичний ключ, за яким зберігається пов’язане значення. Якщо ви знову спробуєте отримати доступ до того самого ключа, хеш-функція обробить ключ, поверне той самий числовий результат, який потім використовується для пошуку пов’язаного значення. Це забезпечує дуже ефективний час пошуку в середньому O(1).

Хеш-таблиці можна реалізувати як масиви, де хеш-функції генерують індекси масиву в межах заданого діапазону. У цьому методі важливий вибір розміру масиву, а також хеш-функції. Наприклад, що станеться, якщо хеш-функція поверне однакове значення для двох різних ключів? Це називається колізією. Один зі способів обробки колізій — просто зберігати обидві пари ключ-значення за цим індексом. Тоді при пошуку будь-якого з них доведеться ітерувати по «відру» елементів, щоб знайти потрібний ключ. Хороша хеш-функція мінімізує колізії, щоб підтримувати ефективний час пошуку.

Тут ми не будемо заглиблюватися в деталі хешування чи реалізації хеш-таблиці, а просто спробуємо загалом зрозуміти, як вони працюють.

# --instructions--

Давайте створимо базову функціональність хеш-таблиці. Ми підготували для вас просту хеш-функцію, яку можна використовувати. Ви можете передати рядок у функцію `hash`, і вона поверне хешоване значення, яке можна використати як ключ для зберігання. Зберігайте елементи на основі цього хешованого значення в об’єкті `this.collection`. Створіть три методи: `add`, `remove` і `lookup`. Перший має приймати пару ключ-значення для додавання до хеш-таблиці. Другий має видаляти пару ключ-значення за переданим ключем. Третій має приймати ключ і повертати пов’язане значення або `null`, якщо ключ відсутній.

Обов’язково враховуйте колізії у вашому коді!

**Примітка:** Тести методу `remove` не пройдуть, доки методи `add` і `lookup` не будуть реалізовані правильно.

# --hints--

Функція `hash` має бути валідною.

```js
let calls = 0;
const ourHash = string => {
  calls++;
  let hashCode = 0;
  for (let i = 0; i < string.length; i++) {
    hashCode += string.charCodeAt(i);
  }
  return hashCode;
};

assert.strictEqual(hash('yek'),ourHash('yek'));

assert.strictEqual(hash('key'),ourHash('key'));
assert.strictEqual(hash('key1'),ourHash('key1'));
assert.strictEqual(hash('key2'),ourHash('key2'));
assert.strictEqual(hash('key3'),ourHash('key3'));

assert.strictEqual(hash('1key'),ourHash('1key'));
assert.strictEqual(hash('ke1y'),ourHash('ke1y'));
assert.strictEqual(hash('altKey'),ourHash('altKey'));

assert.strictEqual(called,calls); 
```

Структура даних `HashTable` має існувати.

```js
assert.isDefined(HashTable); 
let newHashTable = new HashTable(); 
assert.isObject(newHashTable); 
```

`HashTable` має мати метод `add`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.add); 
```

`HashTable` має мати метод `lookup`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.lookup); 
```

`HashTable` має мати метод `remove`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.remove); 
```

Метод `add` має додавати пари ключ-значення, а метод `lookup` має повертати значення, пов’язані з певним ключем.

```js
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
assert.strictEqual(newHashTable.lookup('key'),'value'); 
```

Метод `remove` має приймати ключ як вхідні дані і видаляти пов’язану пару ключ-значення.

```js
let hashValue = hash('key');
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
newHashTable.remove('key'); 
assert.notProperty(newHashTable.collection,hashValue); 
```

Метод `remove` має видаляти лише правильну пару ключ-значення.

```js
let newHashTable = new HashTable(); 
let hashValue = hash('key');
newHashTable.add('key','value');
newHashTable.add('yek','value');
newHashTable.add('altKey','value');
newHashTable.remove('yek');

assert.notExists(newHashTable.lookup('yek'));
assert.exists(newHashTable.lookup('altKey'));
assert.exists(newHashTable.lookup('key'));

newHashTable.remove('key');

assert.notProperty(newHashTable.collection,hashValue); 
assert.exists(newHashTable.lookup('altKey'));

```

Елементи мають додаватися з використанням хеш-функції.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1', 'value1')
newHashTable.add('key2', 'value2');
newHashTable.add('key3', 'value3');
assert.strictEqual(called,3); 
```

Хеш-таблиця має обробляти колізії.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1','value1');
newHashTable.add('1key', 'value2'); 
newHashTable.add('ke1y', 'value3'); 
assert.strictEqual(newHashTable.lookup('key1'),'value1');
assert.strictEqual(newHashTable.lookup('1key'),'value2');
assert.strictEqual(newHashTable.lookup('ke1y'),'value3');
```

# --seed--

## --seed-contents--

```js
var called = 0;
var hash = string => {
  called++;
  var hashed = 0;
  for (var i = 0; i < string.length; i++) {
    hashed += string.charCodeAt(i);
  }
  return hashed;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
var called = 0;
var hash = (string) => {
  called++;
  var hash = 0;
  for (var i = 0; i < string.length; i++) { hash += string.charCodeAt(i); }
  return hash;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line

  this.add = function(key, val) {
    var theHash = hash(key);
    if (!this.collection.hasOwnProperty(theHash)) {
      this.collection[theHash] = {};
    }
    this.collection[theHash][key] = val;
  }

  this.remove = function(key) {
    var theHash = hash(key);
    var hashedObj = this.collection[theHash];
    if (hashedObj.hasOwnProperty(key)) {
      delete hashedObj[key];
    }
    if (!Object.keys(hashedObj).length) {
      delete this.collection[theHash];
    }
  }

  this.lookup = function(key) {
    var theHash = hash(key);
    if (this.collection.hasOwnProperty(theHash)) {
      return this.collection[theHash][key];
    }
    return null
  }
  // Only change code above this line
};
```
