---
id: 587d8250367417b2b2512c5f
title: створити клас stack
challengeType: 1
forumTopicId: 301633
dashedName: create-a-stack-class
---

# --description--

У попередньому розділі ми говорили про те, що таке стек і як можна використати масив, щоб представити стек. У цьому розділі ми створимо власний клас стеку. Хоча можна створювати стеки за допомогою масивів, іноді краще обмежити рівень контролю над стеком. Окрім методів `push` і `pop`, у стеків є інші корисні методи. Давайте додамо методи `peek`, `isEmpty` і `clear` до нашого класу стеку.

# --instructions--

Напишіть метод `push`, який додає елемент на вершину стеку, метод `pop`, який видаляє і повертає елемент з вершини стеку, метод `peek`, який дивиться на верхній елемент стеку, метод `isEmpty`, який перевіряє, чи стек порожній, і метод `clear`, який видаляє всі елементи зі стеку. Зазвичай у стеків такого немає, але ми додали допоміжний метод `print`, який виводить колекцію в консоль.

# --hints--

Ваш клас `Stack` повинен мати метод `push`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.push === 'function';
  })()
);
```

Ваш клас `Stack` повинен мати метод `pop`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.pop === 'function';
  })()
);
```

Ваш клас `Stack` повинен мати метод `peek`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.peek === 'function';
  })()
);
```

Ваш клас `Stack` повинен мати метод `isEmpty`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.isEmpty === 'function';
  })()
);
```

Ваш клас `Stack` повинен мати метод `clear`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.clear === 'function';
  })()
);
```

Метод `peek` повинен повертати верхній елемент стеку

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.peek() === 'CS50' && test.peek() === 'CS50';
  })()
);
```

Метод `pop` повинен видаляти і повертати верхній елемент стеку

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.pop() === 'CS50' && test.pop() === 'CS61';
  })()
);
```

Метод `isEmpty` повинен повертати true, якщо стек не містить жодних елементів

```js
assert(
  (function () {
    var test = new Stack();
    return test.isEmpty();
  })()
);
```

Метод `clear` повинен видаляти всі елементи зі стеку

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    test.clear();
    return test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Stack() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
class Stack {
  constructor() {
    this.collection = [];
  }
  print() {
    console.log(this.collection);
  }
  push(val) {
    this.collection.push(val);
  }
  pop() {
    return this.collection.pop();
  }
  peek() {
    return this.collection[this.collection.length - 1];
  }
  isEmpty() {
    return this.collection.length === 0;
  }
  clear() {
    return (this.collection.length = 0);
  }
}
```
