---
id: 587d8259367417b2b2512c84
title: Створити префіксне дерево пошуку
challengeType: 1
forumTopicId: 301634
dashedName: create-a-trie-search-tree
---

# --description--

Тут ми перейдемо від двійкових дерев пошуку до іншого типу структури даних — префіксного дерева, або trie. Trie — це впорядковане дерево пошуку, яке зазвичай використовується для зберігання рядків або, більш загально, асоціативних масивів чи динамічних наборів даних, де ключі є рядками. Вони дуже добре підходять для зберігання множин даних, коли багато ключів мають перекривні префікси, наприклад, усі слова в словнику. На відміну від двійкового дерева, вузли не асоціюються з фактичними значеннями. Замість цього шлях до вузла представляє конкретний ключ. Наприклад, якщо ми хочемо зберегти рядок code у префіксному дереві, у нас буде чотири вузли, по одному на кожну літеру: c — o — d — e. Перейшовши цим шляхом через усі ці вузли, ми отримаємо рядок code — цей шлях і є ключем, який ми зберегли. Якщо ж ми захочемо додати рядок coding, він поділить перші три вузли з code, а потім відгалузиться після d. Таким чином великі набори даних можна зберігати дуже компактно. Крім того, пошук може бути дуже швидким, бо фактично обмежений довжиною рядка, який ви зберігаєте. Більше того, на відміну від двійкових дерев, вузол може зберігати будь-яку кількість дочірніх вузлів. Як ви могли здогадатися з наведеного прикладу, у вузлах, які позначають кінець ключа, зазвичай зберігаються метадані, щоб при наступних обходах цей ключ можна було відновити. Наприклад, якщо ми додамо codes у наш приклад, нам потрібен спосіб знати, що літера e у code позначає кінець ключа, який було введено раніше. Інакше ця інформація фактично загубиться, коли ми додамо codes.

# --instructions--

Давайте створимо префіксне дерево для зберігання слів. Воно повинно приймати слова через метод `add` і зберігати їх у структурі даних префіксного дерева. Також воно має дозволяти перевіряти, чи є заданий рядок словом, за допомогою методу `isWord`, і отримувати всі слова, додані до дерева, через метод `print`. Метод `isWord` має повертати булеве значення, а `print` — масив усіх цих слів у вигляді рядків. Щоб ми могли перевірити правильність реалізації цієї структури даних, ми надали структуру `Node` для кожного вузла дерева. Кожен вузол буде об’єктом із властивістю `keys`, яка є об’єктом JavaScript Map. Вона зберігатиме окремі літери, які є дійсними ключами кожного вузла. Також ми створили властивість `end` у вузлах, яку можна встановити в `true`, якщо вузол представляє кінець слова.

# --hints--

Метод `Trie` має мати метод `add`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

Метод `Trie` має мати метод `print`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

Метод `Trie` має мати метод `isWord`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.isWord == 'function';
  })()
);
```

Метод `print` має повертати всі елементи, додані до префіксного дерева, у вигляді рядків у масиві.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    test.add('jump');
    test.add('jumps');
    test.add('jumped');
    test.add('house');
    test.add('mouse');
    var added = test.print();
    return (
      added.indexOf('jump') != -1 &&
      added.indexOf('jumps') != -1 &&
      added.indexOf('jumped') != -1 &&
      added.indexOf('house') != -1 &&
      added.indexOf('mouse') != -1 &&
      added.length == 5
    );
  })()
);
```

Метод `isWord` має повертати `true` лише для слів, доданих до префіксного дерева, і `false` для всіх інших слів.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    test.add('hop');
    test.add('hops');
    test.add('hopped');
    test.add('hoppy');
    test.add('hope');
    return (
      test.isWord('hop') &&
      !test.isWord('ho') &&
      test.isWord('hopped') &&
      !test.isWord('hopp') &&
      test.isWord('hoppy') &&
      !test.isWord('hoping')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
var Node = function() {
  this.keys = new Map();
  this.end = false;
  this.setEnd = function() {
    this.end = true;
  };
  this.isEnd = function() {
    return this.end;
  };
};
var Trie = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
