---
id: 587d8255367417b2b2512c75
title: створити циклічну чергу
challengeType: 1
forumTopicId: 301625
dashedName: create-a-circular-queue
---

# --description--

У цьому завданні ви створите циклічну чергу. Циклічна черга — це черга, яка записує в кінець колекції, а потім починає перезаписувати себе на початку колекції. Така структура даних корисна в певних ситуаціях. Наприклад, циклічна черга може використовуватися для потокового медіаконтенту. Коли черга заповнюється, нові медіадані перезаписують старі.

Добре ілюструє цю концепцію масив довжиною `5`:

```js
[null, null, null, null, null]
 ^Read @ 0
 ^Write @ 0
```

Тут і позиції читання, і запису знаходяться на позиції `0`. Тепер у чергу додаються 3 нові записи `a`, `b` і `c`. Наша черга тепер виглядає так:

```js
[a, b, c, null, null]
 ^Read @ 0
          ^Write @ 3
```

Під час читання голівка читання може видаляти значення або зберігати їх:

```js
[null, null, null, null, null]
                   ^Read @ 3
                   ^Write @ 3
```

Тепер ми записуємо значення `d`, `e` і `f` у чергу. Коли запис досягає кінця масиву, він повертається на початок:

```js
[f, null, null, d, e]
                ^Read @ 3
    ^Write @ 1
```

Такий підхід вимагає постійного обсягу пам’яті, але дозволяє обробляти файли набагато більшого розміру.

# --instructions--

У цьому завданні ми реалізуємо циклічну чергу. Циклічна черга має надавати методи `enqueue` і `dequeue`, які дозволяють читати з черги та записувати в неї. Сам клас також має приймати ціле число як аргумент, щоб ви могли вказати розмір черги під час створення. Початкову версію цього класу ми вже написали для вас у редакторі коду.

Коли ви додаєте елементи в чергу, вказівник запису має рухатися вперед і повертатися на початок, коли досягає кінця черги. Метод `enqueue` має повертати елемент, який ви додали, якщо операція успішна; інакше він поверне `null`.

Аналогічно, вказівник читання має рухатися вперед, коли ви вилучаєте елементи з черги. Коли ви вилучаєте елемент, цей елемент має бути повернений. Якщо вилучити елемент не можна, потрібно повернути `null`.

Вказівнику запису не дозволяється рухатися далі за вказівник читання (наш клас не дозволить вам перезаписувати дані, які ви ще не прочитали), а вказівник читання не може рухатися далі за дані, які ви записали.

# --hints--

Метод `enqueue` має додавати елементи до циклічної черги.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Не слід додавати елементи далі за вказівник читання.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    test.enqueue(13);
    test.enqueue(25);
    test.enqueue(59);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Метод `dequeue` має вилучати елементи з черги.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 && test.dequeue() === 32 && test.dequeue() === 591
    );
  })()
);
```

Після вилучення елемента його позиція в черзі має бути скинута до `null`.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(672);
    test.dequeue();
    test.dequeue();
    var print = test.print();
    return print[0] === null && print[1] === null && print[2] === 672;
  })()
);
```

Спроба вилучити елемент далі за вказівник запису має повертати `null` і не рухати вказівник запису.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 &&
      test.dequeue() === 32 &&
      test.dequeue() === 591 &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.enqueue(100) === 100 &&
      test.dequeue() === 100
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
class CircularQueue {
  constructor(size) {

    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line

    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line

    // Only change code above this line
  }
}
```

# --solutions--

```js
class CircularQueue {
  constructor(size) {
    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line
    console.log(this.write, this.max);
    if (this.queue[this.write] === null) {
      this.queue[this.write++] = item;

      if (this.write > this.max) {
        this.write = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line
    if (this.queue[this.read] !== null) {
      let item = this.queue[this.read];
      this.queue[this.read++] = null;
      if (this.read > this.max) {
        this.read = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }
}
```
