---
id: 587d8257367417b2b2512c7b
title: Додати новий елемент до двійкового дерева пошуку
challengeType: 1
forumTopicId: 301618
dashedName: add-a-new-element-to-a-binary-search-tree
---

# --description--

Ця серія завдань познайомить вас зі структурою даних дерево. Дерева — це важлива і універсальна структура даних у комп’ютерних науках. Звісно, їхня назва походить від того, що при візуалізації вони дуже схожі на дерева, які ми знаємо з природного світу. Структура даних дерево починається з одного вузла, зазвичай його називають коренем, і звідси розгалужується на додаткові вузли, кожен з яких може мати дочірні вузли, і так далі. Зазвичай структуру даних візуалізують так, що кореневий вузол знаходиться зверху; можна уявити це як природне дерево, перевернуте догори корінням.

Спершу опишемо деяку загальну термінологію, з якою ми зустрінемося у деревах. Кореневий вузол — це вершина дерева. Точки даних у дереві називаються вузлами. Вузли з гілками, що ведуть до інших вузлів, називають батьківськими вузлами для вузлів, до яких ведуть ці гілки (дочірніх вузлів). Інші, більш складні сімейні терміни застосовуються, як ви можете очікувати. Піддерево — це всі нащадки певного вузла, гілки можуть називатися ребрами, а листові вузли — це вузли на кінцях дерева, які не мають дочірніх вузлів. Нарешті, зауважте, що дерева за своєю природою є рекурсивними структурами даних. Тобто будь-які діти вузла є батьками власного піддерева, і так далі. Рекурсивна природа дерев важлива для розуміння при розробці алгоритмів для поширених операцій над деревами.

Для початку ми розглянемо певний тип дерева — двійкове дерево. Насправді ми розглянемо конкретне двійкове дерево — двійкове дерево пошуку. Давайте пояснимо, що це означає. Хоча структура даних дерево може мати будь-яку кількість гілок у одному вузлі, двійкове дерево може мати лише дві гілки для кожного вузла. Крім того, двійкове дерево пошуку впорядковане відносно дочірніх піддерев так, що значення кожного вузла в лівому піддереві менше або дорівнює значенню батьківського вузла, а значення кожного вузла в правому піддереві більше або дорівнює значенню батьківського вузла. Дуже корисно візуалізувати це співвідношення, щоб краще його зрозуміти:

<div style='width: 100%; display: flex; justify-content: center; align-items: center;'><img alt="an example of a binary search tree" style='width: 100%; max-width: 350px; background-color: var(--gray-05);' src='https://user-images.githubusercontent.com/18563015/32136009-1e665d98-bbd6-11e7-9133-63184f9f9182.png'></div>

Тепер це впорядкування дуже легко побачити. Зверніть увагу, що кожне значення ліворуч від 8, кореневого вузла, менше за 8, а кожне значення праворуч — більше за 8. Також помітно, що це співвідношення застосовується до кожного з піддерев. Наприклад, перший лівий дочірній вузол — це піддерево. 3 — батьківський вузол, і він має рівно два дочірні вузли — за правилами двійкових дерев пошуку ми знаємо, навіть не дивлячись, що лівий дочірній вузол цього вузла (і всі його діти) будуть меншими за 3, а правий дочірній вузол (і всі його діти) будуть більшими за 3 (але також меншими за значення кореневого вузла структури), і так далі.

Двійкові дерева пошуку — це дуже поширені і корисні структури даних, бо вони забезпечують логарифмічний час у середньому випадку для кількох поширених операцій, таких як пошук, вставка та видалення.

# --instructions--

Почнемо з простого. Ми визначили каркас структури двійкового дерева пошуку, а також функцію для створення вузлів для нашого дерева. Зверніть увагу, що кожен вузол може мати ліве та праве значення. Вони будуть призначені дочірнім піддеревам, якщо вони існують. У нашому двійковому дереві пошуку ви створите метод для додавання нових значень до дерева. Метод має називатися `add` і приймати ціле число для додавання до дерева. Слідкуйте за тим, щоб підтримувати інваріант двійкового дерева пошуку: значення в кожному лівому дочірньому вузлі має бути меншим або рівним значенню батьківського вузла, а значення в кожному правому дочірньому вузлі — більшим або рівним значенню батьківського вузла. Тут зробимо так, щоб наше дерево не могло містити дублікати значень. Якщо ми спробуємо додати значення, яке вже існує, метод має повернути `null`. Інакше, якщо додавання пройшло успішно, має повернутися `undefined`.

**Підказка:** дерева за своєю природою є рекурсивними структурами даних!

# --hints--

Структура даних `BinarySearchTree` має існувати.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Двійкове дерево пошуку має мати метод з назвою `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

Метод додавання має додавати елементи відповідно до правил двійкового дерева пошуку.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    const expectedResult = [1, 4, 7, 8, 34, 45, 73, 87];
    const result = test.inOrder();
    return expectedResult.toString() === result.toString();
  })()
);
```

Додавання елемента, який вже існує, має повертати `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    return test.add(4) == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    inOrder() {
      if (!this.root) {
        return null;
      }
      var result = new Array();
      function traverseInOrder(node) {
        node.left && traverseInOrder(node.left);
        result.push(node.value);
        node.right && traverseInOrder(node.right);
      }
      traverseInOrder(this.root);
      return result;
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.add = function(element) {
    let current = this.root;
    if (!current) {
      this.root = new Node(element);
      return;
    } else {
      const searchTree = function(current) {
        if (current.value > element) {
          if (current.left) {
            return searchTree(current.left);
          } else {
            current.left = new Node(element);
            return;
          }
        } else if (current.value < element) {
          if (current.right) {
            return searchTree(current.right);
          } else {
            current.right = new Node(element);
            return;
          }
        } else {
          return null;
        }
      };
      return searchTree(current);
    }
  };
}
```
