---
id: 587d8250367417b2b2512c60
title: створити клас черги
challengeType: 1
forumTopicId: 301631
dashedName: create-a-queue-class
---

# --description--

Як і стеки, черги — це колекція елементів. Але на відміну від стеків, черги дотримуються принципу FIFO (перший прийшов — перший пішов). Елементи, додані до черги, додаються в хвіст, або кінець, черги, і лише елемент на початку черги можна видалити.

Ми могли б використати масив, щоб представити чергу, але так само, як і зі стеками, ми хочемо обмежити контроль над нашими чергами.

Два основні методи класу черги — це метод enqueue і метод dequeue. Метод enqueue додає елемент у хвіст черги, а метод dequeue видаляє і повертає елемент на початку черги. Інші корисні методи — це front, size і isEmpty.

# --instructions--

Напишіть метод `enqueue`, який додає елемент у хвіст черги, метод `dequeue`, який видаляє і повертає елемент на початку, метод `front`, який дозволяє побачити елемент на початку, метод `size`, який показує довжину, і метод `isEmpty`, щоб перевірити, чи порожня черга.

# --hints--

Ваш клас `Queue` повинен мати метод `enqueue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.enqueue === 'function';
  })()
);
```

Ваш клас `Queue` повинен мати метод `dequeue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.dequeue === 'function';
  })()
);
```

Ваш клас `Queue` повинен мати метод `front`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.front === 'function';
  })()
);
```

Ваш клас `Queue` повинен мати метод `size`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.size === 'function';
  })()
);
```

Ваш клас `Queue` повинен мати метод `isEmpty`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Метод `dequeue` повинен видаляти і повертати елемент на початку черги

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.dequeue() === 'Smith';
  })()
);
```

Метод `front` повинен повертати значення елемента на початку черги

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.front() === 'Smith';
  })()
);
```

Метод `size` повинен повертати довжину черги

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return test.size() === 1;
  })()
);
```

Метод `isEmpty` повинен повертати `false`, якщо в черзі є елементи

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return !test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Queue() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function Queue () { 
    var collection = [];
    this.print = function() {
        console.log(collection);
    };
    // Only change code below this line
    this.enqueue = function(item) {
        collection.push(item);
    }

    this.dequeue = function() {
        return collection.shift();
    }

    this.front = function() {
        return collection[0];
    }

    this.size = function(){
        return collection.length;
    }

    this.isEmpty = function() {
        return collection.length === 0 ? true : false;
    }
    // Only change code above this line
}
```
