---
id: 587d8251367417b2b2512c63
title: Видалити елементи зі зв’язаного списку
challengeType: 1
forumTopicId: 301712
dashedName: remove-elements-from-a-linked-list
---

# --description--

Наступний важливий метод, який потрібен будь-якій реалізації зв’язаного списку, — це метод `remove`. Цей метод має приймати елемент, який ми хочемо видалити, як аргумент, а потім шукати у списку та видаляти вузол, що містить цей елемент.

Коли ми видаляємо вузол зі зв’язаного списку, важливо не залишити випадково решту списку без зв’язку. Пам’ятайте, що властивість `next` кожного вузла вказує на вузол, який йде за ним у списку. Якщо ми видаляємо, наприклад, середній елемент, нам потрібно переконатися, що властивість `next` попереднього вузла цього елемента вказує на властивість `next` середнього елемента (яка є наступним вузлом у списку!)

Це може звучати дуже заплутано, тому повернемося до прикладу з конга-лінією, щоб мати хорошу концептуальну модель. Уявіть себе в конга-лінії, і людина безпосередньо перед вами виходить із лінії. Людина, яка щойно вийшла, більше не тримається за нікого в лінії — і ви більше не тримаєтеся за людину, що вийшла. Ви робите крок уперед і тримаєтеся за наступну людину, яку бачите.

Якщо елемент, який ми хочемо видалити, є `head` елементом, ми переназначаємо `head` на другий вузол зв’язаного списку.

# --instructions--

Напишіть метод `remove`, який приймає елемент і видаляє його зі зв’язаного списку.

**Примітка:** Властивість `length` списку має зменшуватися на одиницю щоразу, коли елемент видаляється зі зв’язаного списку.

# --hints--

Ваш клас `LinkedList` має мати метод `remove`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.remove === 'function';
  })()
);
```

Ваш метод `remove` має переназначати `head` на другий вузол, коли видаляється перший вузол.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.remove('cat');
    return test.head().element === 'dog';
  })()
);
```

Ваш метод `remove` має зменшувати `length` зв’язаного списку на одиницю за кожен видалений вузол.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('hamster');
    test.remove('cat');
    test.remove('fish');
    return test.size() === 2;
  })()
);
```

Ваш метод `remove` має переназначати референс попереднього вузла видаленого вузла на референс `next` видаленого вузла.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('snake');
    test.add('kitten');
    test.remove('snake');
    return test.head().next.next.element === 'kitten';
  })()
);
```

Ваш метод `remove` не має змінювати зв’язаний список, якщо елемент не існує у списку.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.remove('elephant');
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    if (head === null) {
      return;
    }
    var previous;
    var currentNode = head;

    while (currentNode.next !== null && currentNode.element !== element) {
      previous = currentNode;
      currentNode = currentNode.next;
    }
    
    if (currentNode.next === null && currentNode.element !== element) {
      return;
    }
    else if (previous) {
      previous.next = currentNode.next;
    } else {
      head = currentNode.next;
    }

    length--;
  };
} 
```
