---
id: 587d8251367417b2b2512c61
title: Робота з вузлами у зв’язаному списку
challengeType: 1
forumTopicId: 301721
dashedName: work-with-nodes-in-a-linked-list
---

# --description--

Ще одна загальна структура даних, з якою ви зіткнетеся в інформатиці, — це <dfn>зв’язаний список</dfn>. Зв’язаний список — це лінійна колекція елементів даних, які називаються «вузлами», кожен з яких вказує на наступний. Кожен <dfn>вузол</dfn> у зв’язаному списку містить дві ключові частини інформації: сам `element` і референс на наступний `node`.

Уявіть, що ви стоїте в конга-лінії. Ви тримаєтеся за людину попереду, а людина позаду тримається за вас. Ви бачите людину безпосередньо попереду, але вона закриває вид на інших людей попереду в лінії. Вузол — це як людина в конга-лінії: він знає, хто він, і бачить лише наступну людину в лінії, але не знає про інших людей попереду чи позаду.

# --instructions--

У нашому редакторі коду ми створили два вузли, `Kitten` і `Puppy`, і вручну з’єднали вузол `Kitten` з вузлом `Puppy`.

Створіть вузли `Cat` і `Dog` і вручну додайте їх до лінії.

# --hints--

Ваш вузол `Puppy` повинен містити референс на вузол `Cat`.

```js
assert(Puppy.next.element === 'Cat');
```

Ваш вузол `Cat` повинен містити референс на вузол `Dog`.

```js
assert(Cat.next.element === 'Dog');
```

# --seed--

## --seed-contents--

```js
var Node = function(element) {
  this.element = element;
  this.next = null;
};
var Kitten = new Node('Kitten');
var Puppy = new Node('Puppy');

Kitten.next = Puppy;
// Only change code below this line
```

# --solutions--

```js
// solution required
```
