---
id: 587d8257367417b2b2512c7c
title: Перевірка наявності елемента в двійковому дереві пошуку
challengeType: 1
forumTopicId: 301623
dashedName: check-if-an-element-is-present-in-a-binary-search-tree
---

# --description--

Тепер, коли ми маємо загальне уявлення про те, що таке двійкове дерево пошуку, давайте розглянемо його трохи детальніше. Двійкові дерева пошуку забезпечують логарифмічний час для загальних операцій пошуку, вставки та видалення у середньому випадку, і лінійний час у найгіршому випадку. Чому так? Кожна з цих базових операцій вимагає від нас знайти елемент у дереві (або у випадку вставки — знайти, куди його слід додати), і через структуру дерева на кожному батьківському вузлі ми розгалужуємося вліво або вправо, ефективно виключаючи половину розміру залишку дерева. Це робить пошук пропорційним логарифму кількості вузлів у дереві, що створює логарифмічний час для цих операцій у середньому випадку. Добре, а що з найгіршим випадком? Розглянемо побудову дерева з таких значень, додаючи їх зліва направо: `10`, `12`, `17`, `25`. Відповідно до правил двійкового дерева пошуку, ми додамо `12` праворуч від `10`, `17` праворуч від нього, і `25` праворуч від цього. Тепер наше дерево нагадує зв’язаний список, і проходження його для пошуку `25` вимагатиме послідовного проходження всіх елементів. Отже, лінійний час у найгіршому випадку. Проблема тут у тому, що дерево незбалансоване. Ми детальніше розглянемо, що це означає, у наступних завданнях.

# --instructions--

У цьому завданні ми створимо утиліту для нашого дерева. Напишіть метод `isPresent`, який приймає ціле число як вхідні дані і повертає булеве значення, що вказує на наявність або відсутність цього значення у двійковому дереві пошуку.

# --hints--

Структура даних `BinarySearchTree` має існувати.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Двійкове дерево пошуку має мати метод під назвою `isPresent`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isPresent == 'function';
  })()
);
```

Метод `isPresent` має правильно перевіряти наявність або відсутність елементів, доданих до дерева.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    test.add(4);
    test.add(7);
    test.add(411);
    test.add(452);
    return (
      test.isPresent(452) &&
      test.isPresent(411) &&
      test.isPresent(7) &&
      !test.isPresent(100)
    );
  })()
);
```

`isPresent` має обробляти випадки, коли дерево порожнє.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    return test.isPresent(5) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.isPresent = function (value) {
    var current = this.root
    while (current) {
      if (value === current.value) {
        return true;
      }
      current = value < current.value ? current.left : current.right;
    }
    return false;
  }
}
```
