---
id: 587d8256367417b2b2512c77
title: Список суміжності
challengeType: 1
forumTopicId: 301620
dashedName: adjacency-list
---

# --description--

Графи можна представити різними способами. Тут ми описуємо один із них, який називається <dfn>список суміжності</dfn>. Список суміжності — це по суті маркований список, де зліва знаходиться вузол, а справа перераховані всі інші вузли, з якими він з’єднаний. Нижче наведено приклад представлення списку суміжності.

<blockquote>Node1: Node2, Node3<br>Node2: Node1<br>Node3: Node1</blockquote>

Вищенаведений приклад — це неорієнтований граф, оскільки `Node1` з’єднаний з `Node2` і `Node3`, і ця інформація узгоджується з з’єднаннями, які показують `Node2` і `Node3`. Список суміжності для орієнтованого графа означає, що кожен рядок списку показує напрямок. Якби наведений вище граф був орієнтованим, то `Node2: Node1` означав би, що орієнтована ребро спрямоване від `Node2` до `Node1`. Ми можемо представити вищенаведений неорієнтований граф у вигляді списку суміжності, помістивши його всередину JavaScript-об’єкта.

```js
var undirectedG = {
  Node1: ["Node2", "Node3"],
  Node2: ["Node1"],
  Node3: ["Node1"]
};
```

Це також можна простіше представити у вигляді масиву, де вузли мають просто числа замість рядкових міток.

```js
var undirectedGArr = [
  [1, 2], // Node1
  [0],    // Node2
  [0]     // Node3
];
```

# --instructions--

Створіть соціальну мережу як неорієнтований граф із 4 вузлів/людей на ім’я `James`, `Jill`, `Jenny` і `Jeff`. Існують ребра/зв’язки між James і Jeff, Jill і Jenny, а також Jeff і Jenny.

# --hints--

`undirectedAdjList` має містити лише чотири вузли.

```js
assert(Object.keys(undirectedAdjList).length === 4);
```

Має бути ребро між `Jeff` і `James`.

```js
assert(
  undirectedAdjList.James.indexOf('Jeff') !== -1 &&
    undirectedAdjList.Jeff.indexOf('James') !== -1
);
```

Має бути ребро між `Jill` і `Jenny`.

```js
assert(
  undirectedAdjList.Jill.indexOf('Jenny') !== -1 &&
    undirectedAdjList.Jenny.indexOf('Jill') !== -1
);
```

Має бути ребро між `Jeff` і `Jenny`.

```js
assert(
  undirectedAdjList.Jeff.indexOf('Jenny') !== -1 &&
    undirectedAdjList.Jenny.indexOf('Jeff') !== -1
);
```

# --seed--

## --seed-contents--

```js
var undirectedAdjList = {};
```

# --solutions--

```js
var undirectedAdjList = {
  James: ['Jeff'],
  Jill: ['Jenny'],
  Jenny: ['Jill', 'Jeff'],
  Jeff: ['James', 'Jenny']
};
```
