---
id: 587d825b367417b2b2512c8c
title: Реалізувати сортування купою з мінімальною купою
challengeType: 1
forumTopicId: 301643
dashedName: implement-heap-sort-with-a-min-heap
---

# --description--

Тепер, коли ми можемо додавати та видаляти елементи, давайте розглянемо деякі застосування куп. Купи часто використовують для реалізації пріоритетних черг, оскільки вони завжди зберігають елемент із найбільшим або найменшим значенням на першій позиції. Крім того, їх використовують для реалізації алгоритму сортування, який називається сортуванням купою. Тут ми побачимо, як це зробити. Сортування купою використовує мінімальну купу, яка є протилежністю максимальної купи. Мінімальна купа завжди зберігає елемент із найменшим значенням у кореневій позиції.

Сортування купою працює так: береться несортований масив, кожен елемент масиву додається до мінімальної купи, а потім кожен елемент витягується з мінімальної купи у новий масив. Структура мінімальної купи гарантує, що новий масив міститиме оригінальні елементи у порядку від найменшого до найбільшого. Це один із найефективніших алгоритмів сортування зі середньою та гіршою складністю O(nlog(n)).

# --instructions--

Давайте реалізуємо сортування купою з мінімальною купою. Не соромтеся адаптувати свій код максимальної купи тут. Створіть об’єкт `MinHeap` з методами `insert`, `remove` та `sort`. Метод `sort` має повертати масив усіх елементів мінімальної купи, відсортованих від найменшого до найбільшого.

# --hints--

Структура даних `MinHeap` має існувати.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MinHeap` має мати метод `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MinHeap` має мати метод `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

`MinHeap` має мати метод `sort`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.sort == 'function';
  })()
);
```

Метод `sort` має повертати масив, що містить усі елементи, додані до мінімальної купи, у відсортованому порядку.

```js
assert(
  (() => {
    if (typeof MinHeap === 'undefined') {
      return false;
    }

    const heap = new MinHeap();
    const arr = createRandomArray(25);

    for (let i of arr) {
      heap.insert(i);
    }

    const result = heap.sort();
    arr.sort((a, b) => a - b);

    for (let i = 0; i < arr.length; i++) {
      if (arr[i] !== result[i]) {
        return false;
      }
    }
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
// Generate a randomly filled array
function createRandomArray(size = 5){
  let a = new Array(size);
  for(let i = 0; i < size; i++)
    a[i] = Math.floor(Math.random() * 100);
  
  return a;
}
const array = createRandomArray(25);

var MinHeap = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
