---
id: 587d8252367417b2b2512c67
title: додати елементи за певним індексом у зв’язаному списку
challengeType: 1
forumTopicId: 301619
dashedName: add-elements-at-a-specific-index-in-a-linked-list
---

# --description--

Давайте створимо метод addAt(index,element), який додає елемент за заданим індексом. Так само, як ми видаляємо елементи за певним індексом, нам потрібно відстежувати currentIndex під час проходження зв’язаного списку. Коли currentIndex співпаде з заданим індексом, потрібно переназначити властивість next попереднього вузла, щоб вона посилалась на новий доданий вузол. А новий вузол має посилатися на наступний вузол у currentIndex. Повертаючись до прикладу з конго-лінією, нова людина хоче приєднатися до лінії, але посередині. Ви стоїте посередині лінії, тому відпускаєте руки від людини попереду. Нова людина підходить і кладе руки на ту людину, на яку ви раніше тримали руки, а ви тепер тримаєте руки на новій людині.

# --instructions--

Створіть метод `addAt(index,element)`, який додає елемент за заданим індексом. Поверніть false, якщо елемент не вдалося додати. **Примітка:** Не забудьте перевірити, чи не є заданий індекс від’ємним або більшим за довжину зв’язаного списку.

# --hints--

Ваш метод `addAt` має переназначити `head` на новий вузол, якщо заданий індекс дорівнює 0.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'fish');
    return test.head().element === 'fish' && test.head().next.element === 'cat';
  })()
);
```

Ваш метод `addAt` має збільшувати довжину зв’язаного списку на одиницю для кожного нового вузла, доданого до списку.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'cat');
    return test.size() === 3;
  })()
);
```

Ваш метод `addAt` має повертати `false`, якщо вузол не вдалося додати.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.addAt(4, 'cat') === false;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element) {
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function() {
    return head;
  };

  this.add = function(element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if (head === null){
        head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };
  this.addAt = function (index, element) {
    if (index > length || index < 0) {
      return false;
    }
    var newNode = new Node(element);
    var currentNode = head;
    if (index === 0) {
      head = newNode;
    } else {
      var previousNode = null;
      var i = 0;
      while (currentNode && i < index) {
        previousNode = currentNode;
        currentNode = currentNode.next;
        i++;
      }
      previousNode.next = newNode;
    }
    newNode.next = currentNode;
    length++;
  }
}
```
