---
id: 587d825b367417b2b2512c8b
title: Видалити елемент із максимальної купи
challengeType: 1
forumTopicId: 301710
dashedName: remove-an-element-from-a-max-heap
---

# --description--

Тепер, коли ми можемо додавати елементи до нашої купи, давайте подивимося, як можна видаляти елементи. Видалення та вставка елементів вимагають схожої логіки. У максимальній купі зазвичай хочуть видалити найбільше значення, тому це просто означає вилучення його з кореня нашого дерева. Це порушить властивість купи нашого дерева, тому ми повинні відновити її якимось чином. Зазвичай для максимальної купи це робиться таким чином:

<ol>
  <li>Перемістіть останній елемент у купі на позицію кореня.</li>
  <li>Якщо будь-який дочірній елемент кореня більший за нього, поміняйте місцями корінь із дочірнім елементом із більшим значенням.</li>
  <li>Продовжуйте міняти місцями, доки батьківський елемент не буде більшим за обох дочірніх або доки не досягнете останнього рівня дерева.</li>
</ol>

# --instructions--

Інструкції: додайте метод до нашої максимальної купи з назвою `remove`. Цей метод має повертати найбільше значення, яке було додано до нашої максимальної купи, і видаляти його з купи. Також він повинен впорядкувати купу так, щоб властивість купи зберігалася. Після видалення елемента наступний найбільший елемент, що залишився в купі, має стати коренем.

# --hints--

Структура даних `MaxHeap` має існувати.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap` має мати метод з назвою `print`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

`MaxHeap` має мати метод з назвою `insert`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap` має мати метод з назвою `remove`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Метод `remove` має видаляти найбільший елемент із максимальної купи, зберігаючи властивість максимальної купи.

```js
function isHeap(arr, i, n) {
    if( arr[i] < arr[2 * i + 1] || arr[i] < arr[2 * i + 2] ){
        return false;
    }
    if (i > (n - 1) / 2) {
        return true;
    }
    if (isHeap(arr, 2 * i + 1, n) && isHeap(arr, 2 * i + 2, n)) {
        return true;
    }
    return false;
}


assert(
  (function () {
    let test = false;

    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    
    let max = Infinity;
    const [result, vals] = [[], [9, 3, 5, 2, 15, 3, 7, 12, 7, 10, 90]];

    vals.forEach((val) => test.insert(val));

    for (let i = 0; i < vals.length; i++) {
      const curHeap = test.print();
      const arr = curHeap[0] === null ? curHeap.slice(1) : curHeap;
      
      if (!isHeap(arr, 0, arr.length - 1)) {
        return false;
      }

      const removed = test.remove();
      if (!vals.includes(removed)) return false;
      if (removed > max) return false
      max = removed;
      result.push(removed);
    }
    
    for (let i = 0; i < vals.length; i++) {
      if (!result.includes(vals[i])) {
        return false;
      }
    }
    
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
const MaxHeap = function () {
  this.heap = [];
  this.parent = index => {
    return Math.floor((index - 1) / 2);
  }
  this.insert = element => {
    this.heap.push(element);
    this.heapifyUp(this.heap.length - 1);
  }
  this.heapifyUp = index => {
    let currentIndex = index,
    parentIndex = this.parent(currentIndex);
    while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
      this.swap(currentIndex, parentIndex);
      currentIndex = parentIndex;
      parentIndex = this.parent(parentIndex);
    }
  }
  this.swap = (index1, index2) => {
    [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
  }
  this.print = () => {
    return this.heap;
  }
  // Only change code below this line

  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
