---
id: 62cc5b1779e4d313466f73c5
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `head` 元素中嵌套一个 `link` 元素。 设置其 `rel` 属性值为 `stylesheet`，`href` 属性值为 `styles.css`。

# --hints--

代码应该有一个 `link` 元素。

```js
assert.isNotNull(document.querySelector('link'));
```

不应更改现有的 `head` 标签。 确保没有删除结束标签。

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

你应该有一个 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

你的 `link` 元素应该是一个空元素，它不应该有结束标签 `</link>`。

```js
assert.notMatch(code, /<\/link>/);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
