---
id: 65f055b9190fc41ca35549b8
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Python ofrece varios métodos que incluyen tanto un doble guion bajo al inicio como al final en sus nombres. Puede que ya estés familiarizado con algunos, como `__init__` y `__str__`. Estos métodos, que siguen el patrón de nombres `__<name>__`, se conocen como métodos especiales, métodos mágicos o métodos *dunder* (que significa doble guion bajo).

Definir métodos especiales en una clase afecta el comportamiento de esa clase. Se llaman internamente en situaciones específicas (por ejemplo, `__init__` durante la instanciación, `__str__` cuando el objeto se imprime o se pasa a `str()`). En este proyecto, vas a aprender algunos de los métodos especiales más usados.

Por ahora, asigna `x` al atributo `x` del objeto `Vector`.

# --hints--

Debes asignar `x` a `self.x`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
