---
id: 65f03d9f92eac9183a4d3281
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Un vector puede definirse por dos coordenadas, `x` y `y`, en el plano euclidiano. La distancia entre el origen de los ejes y el punto `(x, y)` será su longitud, o norma. Y la dirección del vector apuntará hacia `(x, y)`.

<img alt="a 2-dimensional vector of coordinates (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Dentro de la clase `Vector`, crea un método `__init__` y dale tres parámetros, `self`, `x` y `y`.

# --hints--

Debes definir un método `__init__` dentro de la clase `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

Tu método `__init__` debe tomar tres parámetros: `self`, `x` y `y`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
