---
id: 65f058e06f34fd1f0ee6e55d
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

A continuación, dentro de la clase `Vector`, crea un método `norm` vacío y dale un`self` parámetro.

# --hints--

Su clase `Vector` debe tener un método `norm`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

El método `norm` debe tener un parámetro `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
