---
id: 65f0694fb2296f3caadf8347
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

La longitud de un vector $\mathbf{a}$, o norma, se indica típicamente como $\| \mathbf{a} \|$. Se puede calcular como la raíz cuadrada de la suma de sus componentes al cuadrado:
\[ \| \mathbf{a} \| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2}
\]

Calcula la norma del vector y devuelve el resultado desde tu método `norm`.

# --hints--

Tu método `norm` debe usar la fórmula proporcionada para calcular y devolver un número que represente la norma de la instancia actual del vector. No aproximar el valor.

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
