---
id: 65f06a8e5a57673d700c79c3
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Fuera de la clase `Vector`, crea una instancia de `Vector` pasando los enteros `2` y `3` como los valores `x` y `y` y asígnala a una variable llamada `v1`.

# --hints--

Debes declarar una variable `v1` y asignarle `Vector(2, 3)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
