---
id: 65f01969115f933073b6be03
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

Un vector es un objeto que tiene una longitud (o magnitud) y una dirección, y no puede ser expresado por un solo número. En física, los vectores se usan comúnmente para representar fuerzas, velocidades, aceleraciones y otras cantidades.

En este proyecto vas a construir un espacio vectorial, un conjunto en el que se define una serie de operaciones entre los elementos de ese conjunto. Pronto aprenderás todos los detalles.

Por ahora, comienza el proyecto declarando una clase vacía llamada `Vector`.

# --hints--

Debes declarar una clase vacía llamada `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
