---
id: 65f3fa60a93b84110b3f1708
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Quieres devolver una representación en cadena significativa. Por ejemplo, en el caso de `v1`, quieres devolver una cadena que contenga la tupla con los valores de los componentes del vector: `(2, 3)`.

Desde el método `__str__`, devuelve una cadena que represente el vector como una tupla que contenga los componentes del vector en orden. Luego, sal de la clase `Vector` e imprime `v1` para verificar el resultado.

# --hints--

Tu método `__str__` debe devolver una cadena que represente el vector como una tupla que contenga los componentes del vector en orden.

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

Debes imprimir `v1`.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
