---
id: 65f3fe07cc763212efe91285
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Un vector puede tener un número `n` de dimensiones (componentes). Aquí tienes una representación de un vector tridimensional:

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Hasta ahora, creaste un vector de 2 dimensiones. Quieres poder representar vectores con un número diferente de dimensiones sin reescribir el código necesario para cada caso específico. Para eso, usarás herencia.

Comienza renombrando la clase `Vector` a `R2Vector` para especificar que esta clase va a manejar vectores bidimensionales. Recuerda modificar también la instanciación de `v1`.

# --hints--

Debes renombrar la clase `Vector` a `R2Vector`. Recuerda modificar también la instanciación de `v1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
