---
id: 65f3fa097f26b510db6c710b
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Genial, `norm` está funcionando como se esperaba. Ahora, si intentas imprimir `v1`, obtendrás la representación en cadena predeterminada de un objeto (algo como `<__main__.Vector object at 0x11eb778>`).

Dentro de la clase `Vector`, declara un método vacío `__str__` para implementar una representación en cadena legible. Recuerda darle un parámetro `self`.

Presta atención a no imprimir `v1` hasta que `__str__` devuelva una cadena, de lo contrario obtendrás un `TypeError`, porque `__str__` siempre debe devolver una cadena.

# --hints--

Debes definir un método `__str__` dentro de la clase `Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

Tu método `__str__` debe tener un parámetro `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
