---
id: 65f4535bbdb28d436ff3ddc9
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Como puedes ver en la salida, `__dict__` contiene los valores de los atributos de tu instancia. En lugar de agregar explícitamente los cuadrados de `self.x` y `self.y`, vas a iterar sobre los valores almacenados en `__dict__` para calcular el valor de la norma.

Dentro del cuerpo del método `norm`, reemplaza el contenido de los paréntesis con una expresión generadora que eleve cada valor `val` en `self.__dict__.values()` a la potencia de `2`. Además, pasa esa expresión generadora como argumento a la función `sum`, para que todos los cuadrados se sumen antes de calcular la raíz cuadrada.

# --hints--

Debes retornar `sum(val**2 for val in self.__dict__.values())**0.5` desde el método `norm()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in self.__dict__.values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)

```
