---
id: 65f40401e6ef53173c04e27d
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Dentro de tu nueva clase, declara un método `__init__`. Dale `self`, `x`, `y` y `z` como parámetros.

# --hints--

Debes definir un método `__init__` dentro de tu clase `R3Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").has_function("__init__")
    `))
})
```

Tu método `__init__` debe tener cuatro parámetros `self`, `x`, `y` y `z`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, x, y, z")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
