---
id: 65f4520e363e2642f8112e33
title: Paso 19
challengeType: 20
dashedName: step-19
---

# --description--

Como puedes ver, algo no está del todo bien. Los métodos `norm` y `__str__` heredados de `R2Vector` no pueden adaptarse a un vector de 3 dimensiones. Su implementación tiene que ser más flexible.

Cada objeto en Python tiene un atributo especial llamado `__dict__`, que es un diccionario que almacena los atributos del objeto.

Elimina las llamadas existentes a `print`. Luego, imprime el atributo `__dict__` de tus vectores `v1` y `v2` para ver cómo se ven.

# --hints--

No deberías tener `print(v2)` ni `print(v2.norm())` en tu código.

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

Debes imprimir el atributo `__dict__` de `v1` y `v2`.

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
