---
id: 65f412c208c3791fee305acf
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

En Python, puedes forzar el uso de argumentos solo por palabra clave agregando un `*` como argumento adicional en la firma de la función o método. Modifica ambos métodos `__init__` agregando un `*` como segundo parámetro (después de `self`). Cada parámetro colocado después de eso requerirá el uso de un argumento por palabra clave en la llamada a la función/método.

Esto significa que también necesitas modificar la llamada `super().__init__(x, y)`. Hazlo asignando a `x` el valor `x`, y a `y` el valor `y`.

Finalmente, modifica la instanciación de `v1` y `v2` usando argumentos con nombre.

# --hints--

Debes forzar los argumentos por palabra clave agregando un `*` como segundo parámetro en ambos métodos `__init__` de tu código.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

Debes modificar tu llamada a `super.__init__(x, y)` para usar argumentos con nombre.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

Debes modificar la asignación de `v1` para que sea `R2Vector(x=2, y=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

Debes modificar la asignación de `v2` para que sea `R3Vector(x=2, y=2, z=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
