---
id: 65f407ea37ad6e181b90462e
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

Podrías asignar cada parámetro `i` a `self.i` como hiciste antes. Aunque en este caso tienes pocas líneas para repetir, una forma de evitar esta repetición es usando la función `super()`. `super()` te permite referirte implícitamente a la clase padre: `super().__init__(x, y)` llama al método `__init__` de la clase padre.

Agrega una llamada a `super().__init__(x, y)` en tu método `__init__`.

# --hints--

Debes tener una llamada a `super().__init__(x, y)` en tu método `__init__`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
