---
id: 65f8709620b2ce1a62608f5a
title: Pasos 23
challengeType: 20
dashedName: step-23
---

# --description--

Cuando necesites acceder dinámicamente a algunos atributos a partir de una entrada de cadena, la función incorporada `getattr()` es lo que necesitas. Toma un objeto como su primer argumento y una cadena que contiene el nombre del atributo como su segundo argumento.

Comienza a corregir el método `__str__` reemplazando la cadena que devuelve `__str__()` con una expresión generadora que itere a través de los atributos del objeto y llame a `getattr()` para cada atributo `i`.

# --hints--

Debes devolver una expresión generadora que itere sobre `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

Debes devolver una expresión generadora que use `i` como la variable de iteración.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

Debes devolver una expresión generadora que llame a `getattr(self, i)` para cada elemento `i` en `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
