---
id: 65f8749b13774b1d2e4a7fba
title: Paso 27
challengeType: 20
dashedName: step-27
---

# --description--

Mientras que el método `__str__` devuelve una representación en cadena legible para humanos de un objeto, el método `__repr__` se supone que debe devolver la cadena necesaria para instanciar el objeto.

Agrega un método `__repr__` con un parámetro `self` dentro de la clase `R2Vector`.

# --hints--

Tu clase `R2Vector` debe tener un método `__repr__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__repr__")`)) })
```

Tu método `__repr__` debe tener un parámetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
