---
id: 65f9bc10a9fb1612e066e9e8
title: Paso 42
challengeType: 20
dashedName: step-42
---

# --description--

Al sumar dos vectores, cada componente de un vector se suma al mismo componente del otro vector. Por ejemplo, sumar `(1, 2)` y `(2, 4)` genera un tercer vector `(3, 6)`, donde `3` es la suma de `1` y `2` y `6` es la suma de `2` y `4`.

El diccionario `kwargs` contendrá los pares clave-valor necesarios para instanciar un nuevo vector de la misma clase de los dos vectores sumados.

Convierte el diccionario vacío en una comprensión de diccionario que itere a través de `vars(self)` y para cada clave (`i`) cree un par clave-valor, donde la clave es la misma clave de la iteración actual y su valor es la suma de los valores de los atributos `i` de los dos operandos.

# --hints--

Debes convertir el diccionario vacío asignado a `kwargs` en una comprensión de diccionario que itere sobre `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

La comprensión de diccionario asignada a `kwargs` debe usar la variable `i` para iterar sobre `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

La comprensión de diccionario asignada a `kwargs` debe usar `i` como clave.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

La comprensión de diccionario asignada a `kwargs` debe asignar la suma del atributo `i` de `self` y el atributo `i` de `other` a la clave `i` para cada `i` en `vars(self)`. Usa la función `getattr()` para eso.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
