---
id: 65f9b9710cca621244d3bde1
title: Paso 40
challengeType: 20
dashedName: step-40
---

# --description--

En Python, `NotImplemented` es un valor especial usado para indicar que una operación no está implementada para un caso específico.

`NotImplemented` no genera una excepción inmediatamente. En cambio, indica al otro operando cómo realizar la operación. Si la solicitud no puede ser satisfecha, se devuelve un `TypeError` por defecto.

Porque quieres poder sumar dos vectores solo si pertenecen a la misma clase, devuelve `NotImplemented` desde la estructura `if` que creaste en el paso anterior.

# --hints--

Debes devolver `NotImplemented` desde dentro del cuerpo del `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
