---
id: 65f9b17a638f0d0dcce8c354
title: Paso 38
challengeType: 20
dashedName: step-38
---

# --description--

Para crear un espacio vectorial, necesitas definir cómo deben comportarse los vectores en varios casos. Los vectores se pueden sumar, formando un nuevo vector.

El método especial `__add__` puede ser implementado para sobrescribir lo que sucede por defecto cuando dos objetos se suman usando el operador `+`.

En este momento, intentar sumar dos instancias de `R2Vector` o `R3Vector` generaría una excepción. Crea un método vacío `__add__` dentro de la clase `R2Vector` y dale dos parámetros: `self` y `other`.

# --hints--

Debes definir un método `__add__` dentro de la clase `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

Tu método `__add__` debe tener dos parámetros, `self` y `other`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
