---
id: 65f9dd6e5a08af19c196c2df
title: Paso 45
challengeType: 20
dashedName: step-45
---

# --description--

El vector resultante de la resta de un vector a otro se obtiene calculando la diferencia de cada una de sus componentes. Por ejemplo, restar `(2, 4)` de `(7, 3)` genera un tercer vector `(5, -1)`, donde `5` es la diferencia entre `7` y `2` y `-1` es la diferencia entre `3` y `4`.

Como hiciste antes dentro del método `__add__`, declara una variable `kwargs` después de la sentencia `if`. Asígnale una comprensión de diccionario que itere a través de los atributos del objeto y cree un par clave-valor para cada clave `i`, donde la clave es la misma clave de la iteración actual y su valor es la diferencia entre los valores de los atributos `i` de los dos operandos.

# --hints--

Debes declarar una variable `kwargs` y asignarle una comprensión de diccionario que itere sobre los atributos del objeto.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

La comprensión de diccionario asignada a `kwargs` debe usar la variable `i` para iterar sobre los atributos del objeto.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

La comprensión de diccionario asignada a `kwargs` debe usar `i` como clave.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

La comprensión de diccionario asignada a `kwargs` debe asignar la diferencia entre el atributo `i` de `self` y el atributo `i` de `other` a la clave `i` para cada `i`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
