---
id: 66682150151af29efec9727d
title: Paso 33
challengeType: 20
dashedName: step-33
---

# --description--

El método `__getattribute__` se llama internamente cada vez que intentas acceder a un atributo de instancia. Si el atributo no se encuentra a nivel de instancia, el método lo buscará a nivel de clase y, eventualmente, en las clases padre.

Dentro de tu clase, define un método `__getattribute__` con dos parámetros, `self` y `attr`, y haz que retorne la cadena `'calling __getattribute__'`. En breve sobreescribirás la implementación predeterminada solo para ver cómo funciona el acceso a atributos.

# --hints--

Debes definir un método llamado `__getattribute__` con dos parámetros, `self` y `attr`, dentro de la clase `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_args("self, attr")`)) })
```

Tu método `__getattribute__` debería devolver la cadena `'calling __getattribute__'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_return("'calling __getattribute__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

```
