---
id: 666832e427d70bc5219dc62a
title: Paso 36
challengeType: 20
dashedName: step-36
---

# --description--

Como puedes ver en la salida, aunque definiste `__getattr__` en tu clase, este método aún no se llama. Esto sucede porque el acceso predeterminado a atributos ocurre a través de `__getattribute__`. Por lo tanto, puedes ver el valor del atributo impreso en el terminal.

Ahora modifica las dos últimas líneas de código para acceder al atributo `z` de `v1` en ambas llamadas a `print()`. Esta vez, vas a acceder a un atributo que `v1` no tiene. Como resultado, `__getattribute__` generará un error y se llamará a `__getattr__`.

# --hints--

Debes modificar tu código para imprimir el atributo `z` de `v1` usando tanto el operador punto como la función `getattr()`.

```js
({ test: () => assert(runPython(`
n = _Node(_code)
(n.has_call("print(v1.z)") and n.has_call("print(getattr(v1, 'z'))")) or n.has_call("print(v1.z, getattr(v1, 'z'))") or n.has_call("print(getattr(v1, 'z'), v1.z)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __getattr__(self, attr):                
        return 'calling __getattr__'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')
--fcc-editable-region--
print(v1.x)
print(getattr(v1, 'x'))
--fcc-editable-region--
```
