---
id: 666827a6fd0dbaafe8330ea6
title: Paso 35
challengeType: 20
dashedName: step-35
---

# --description--

`__getattr__` es otro método especial que juega un papel en el acceso a atributos.

La implementación predeterminada de `__getattribute__` es lanzar un `AttributeError` cuando el atributo solicitado no es un atributo de instancia o no está presente en el árbol de la clase.

En ese caso, se llama a `__getattr__` si está definido por la clase. Puedes considerarlo como una especie de recurso alternativo cuando el procedimiento habitual de acceso a atributos falla.

Convierte el método `__getattribute__` en `__getattr__` y modifica la cadena que devuelve el método a `'calling __getattr__'`.

# --hints--

Debes definir un método llamado `__getattr__` con dos parámetros, `self` y `attr`, dentro de la clase `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_args("self, attr")`)) })
```

Tu método `__getattr__` debería devolver la cadena `'calling __getattr__'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_return("'calling __getattr__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __getattribute__(self, attr):                
        return 'calling __getattribute__'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

print(v1.x)
print(getattr(v1, 'x'))

```
