---
id: 6565a536ba1f9f25bd30e88b
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo progetto implementerai l'algoritmo di Luhn. Questo algoritmo è una formula per convalidare diversi numeri di identificazione.

Inizia dichiarando una funzione chiamata `main`, che servirà come punto di ingresso del programma. Usa la parola chiave `pass` per evitare un errore.

# --hints--

Dovresti avere `def main():` nel tuo codice.

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

Dovresti avere `pass` nel tuo codice.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
