---
id: 65687db6a9fcf3c548a6c1b7
title: Passaggio 28
challengeType: 20
dashedName: step-28
---

# --description--

Parte dell'algoritmo consiste nel raddoppiare ogni secondo numero, partendo da destra. Se il risultato del raddoppio è maggiore o uguale a `10`, somma insieme le due cifre. Per esempio, se la cifra è `6`, raddoppiala per ottenere `12`. Somma `1` e `2` per ottenere `3`. Puoi farlo usando la _divisione intera_ per ottenere la prima cifra e l'operatore modulo (`%`) per ottenere la seconda cifra:

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

La divisione intera restituisce il quoziente della divisione, arrotondato per difetto al numero intero più vicino.

All'interno dell'istruzione `if`, assegna a `number` il risultato di `number // 10` (divisione intera) più il modulo di `number` e `10`.

# --hints--

Dovresti avere `number = number // 10 + number % 10` all'interno dell'istruzione `if`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
