---
id: 656877f71bba2b97acedf9af
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Python include classi integrate che possono aiutarci nella manipolazione delle stringhe. Una di queste è la classe `str`. Ha un metodo chiamato `maketrans` che può aiutarci a creare una tabella di traduzione. Questa tabella può essere usata per sostituire caratteri in una stringa:

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

Il codice sopra, quando chiamato su una stringa, sostituirà tutti i caratteri `t` con `c` e tutti i caratteri `l` con `b`.

Crea una variabile chiamata `card_translation` e assegnale una tabella di traduzione per sostituire tutti i caratteri `-` e ` ` con una stringa vuota.

# --hints--

Dovresti creare una variabile `card_translation` all'interno di `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

Dovresti assegnare a `card_translation` il valore `str.maketrans({'-': '', ' ': ''})`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
