---
id: 656879a66338b2a461d5d307
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

L'algoritmo di Luhn è il seguente:

1. Da destra a sinistra, raddoppia il valore di ogni secondo numero; se il prodotto è maggiore di 9, somma le cifre del prodotto.
2. Calcola la somma di tutte le cifre.
3. Se la somma di tutte le cifre è un multiplo di 10, allora il numero è valido; altrimenti non è valido.

Prendi come esempio un numero di conto "7992739871" a cui verrà aggiunta una cifra di controllo, diventando della forma 7992739871x:

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

Sostituisci l'istruzione `pass` con una variabile `sum_of_odd_digits` e un valore di `0`.

# --hints--

Dovresti avere `sum_of_odd_digits = 0` all'interno della funzione `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

Non dovresti avere `pass` all'interno della funzione `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
