---
id: 656878585631369a6b2d2191
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Definire la traduzione non traduce di per sé la stringa. Il metodo `translate` deve essere chiamato sulla stringa da tradurre con la tabella di traduzione come argomento:

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

Crea una variabile chiamata `translated_card_number` e assegnale il risultato della chiamata al metodo `translate` su `card_number` con `card_translation` come argomento.

# --hints--

Dovresti creare una variabile `translated_card_number` all’interno di `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

Dovresti assegnare a `translated_card_number` il valore `card_number.translate(card_translation)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
