---
id: 656879c6f35a85a59c06b3a7
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Hai già consultato elementi (caratteri) di una stringa usando l’operatore indice `[]`. Puoi anche usare l’operatore indice per consultare un intervallo di caratteri in una stringa con `string[start:stop:step]`:

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

Dove `start` è l’indice iniziale (inclusivo), `stop` è l’indice finale (esclusivo) e `step` è il numero di caratteri da saltare. Se non specificato, `step` è predefinito a 1.

Crea una variabile chiamata `card_number_reversed` e assegnale il valore dei primi 4 caratteri di `card_number`.

# --hints--

Dovresti avere `card_number_reversed = card_number[0:4]` all’interno della funzione `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
