---
id: 656b481ec4976439565a78b2
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Proprio come il passaggio è opzionale, anche un inizio a `0` e una fine alla fine della slice sono opzionali:

```python
my_string = 'camperbot'
camperbot = my_string[::]
```

Assegna il valore invertito dell'intera stringa `card_number` alla variabile `card_number_reversed`.

# --hints--

Dovresti avere `card_number_reversed = card_number[::-1]` all'interno della funzione `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*:\s*:\s*-1\s*\]/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[-1:-5:-1]
    print(card_number_reversed)
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
