---
id: 656494269ccce754411a2112
title: الخطوة 27
challengeType: 20
dashedName: step-27
---

# --description--

تسمح لك وحدة `re` باستخدام *التعبيرات النمطية* في شفرتك. ستتعلم المزيد عن التعبيرات النمطية قريبًا جدًا.

أضف الآن عبارة `import` في أعلى شفرتك لاستيراد وحدة `re`.

# --hints--

يجب استيراد وحدة `re`.

```js
({ test: () => assert.match(code, /^import\s+re/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

import secrets
import string
--fcc-editable-region--

def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
