---
id: 6564de10a0887f882b0012d3
title: الخطوة 48
challengeType: 20
dashedName: step-48
---

# --description--

بنفس الطريقة التي يكون فيها الصنف `[0-9]` مكافئًا لـ `\d`، يكون الصنف `[^0-9]` مكافئًا لـ `\D`. يمكن مطابقة الأحرف الأبجدية الرقمية باستخدام `\w` والأحرف غير الأبجدية الرقمية باستخدام `\W`.

استبدل صنف الأحرف `[^a-zA-Z0-9]` بـ `\W`.

# --hints--

يجب استبدال صنف الأحرف `[^a-zA-Z0-9]` بـ `\W`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\\W\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'[^a-zA-Z0-9]')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
