---
id: 656474314bf37d3c83a85143
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

تأخذ الدالة `choice()` من الوحدة `random` متتالية وتُرجع عنصرًا عشوائيًا من المتتالية.

عدّل استدعاء `print()` الخاص بك لاستخدام الدالة `choice()` ومرّر `all_characters` كمعلمة.

# --hints--

يجب تعديل استدعاء `print()` الحالي لطباعة `random.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
