---
id: 656469e62377ab34a5bcb8d1
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

أعلن عن متغيرين جديدين باسم `digits` و `symbols` وقم بتعيين `string.digits` و `string.punctuation` لهما، على التوالي.

# --hints--

يجب إعلان متغير باسم `digits`.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

يجب تعيين `string.digits` إلى متغير `digits` الخاص بك.

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

يجب إعلان متغير باسم `symbols`.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

يجب تعيين `string.punctuation` إلى متغير `symbols` الخاص بك.

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
