---
id: 6564c67db26c417561ab510d
title: الخطوة 31
challengeType: 20
dashedName: step-31
---

# --description--

كما ترى من الناتج، الآن التعبير النمطي الخاص بك يطابق أول `l` داخل السلسلة النصية.

في نمطك، يمكنك إضافة المحدد الكمي بعد الحرف لتحديد عدد مرات تكرار هذا الحرف. على سبيل المثال، المحدد الكمي `+` يعني أن الحرف يجب أن يتكرر مرة واحدة أو أكثر.

أضف المحدد الكمي `+` إلى نمطك.

# --hints--

يجب تعديل متغير `pattern` إلى `re.compile('l+')`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
