---
id: 65648f4b2281ba50051ae39c
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

بعد ذلك، اكتب حلقة `for` مع المتغير `i` كمتغير للحلقة. استخدم الدالة `range()` للتكرار حتى قيمة `length`.

داخل الحلقة، استخدم معامل التعيين بالإضافة لإضافة حرف عشوائي من `all_characters` إلى القيمة الحالية لـ `password`. استخدم الدالة `choice()` من وحدة `secrets` لذلك.

# --hints--

يجب أن تكتب حلقة `for` تتكرر عبر `range(length)`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

يجب أن تستخدم معامل `+=` لإضافة حرف عشوائي من `all_characters` إلى القيمة الحالية لـ `password`.

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password
    
--fcc-editable-region--
```
