---
id: 656485a8a3496d4a36b1496a
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

تشكل هذه المتغيرات الثلاثة الأحرف الممكنة للاختيار منها عند توليد كلمة المرور.

قبلها مباشرة، أضف تعليقًا يقول `Define the possible characters for the password`.

# --hints--

يجب أن تضيف التعليق فوق المتغير `letters` مباشرة.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
