---
id: 6564a767a7241362f7d8d664
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

يُستخدم الشرطة السفلية المفردة المستقلة لتمثيل قيمة لا تهتم بها أو لن تُستخدم في شفرتك.
المتغير الخاص بالتكرار لديك غير مستخدم فعليًا.

غيّر المتغير `i` إلى الشرطة السفلية المفردة.

# --hints--

يجب أن تغيّر المتغير `i` إلى `_`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/for\s+_\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--
    password = ''
    # Generate password
    for i in range(length):
        password += secrets.choice(all_characters)
    
--fcc-editable-region--
```
