---
id: 6564d68c34027a8072a704f4
title: الخطوة 42
challengeType: 20
dashedName: step-42
---

# --description--

حرف النقطة هو رمز بدل يطابق أي حرف في سلسلة نصية — باستثناء حرف السطر الجديد افتراضيًا. عدل `pattern` ليطابق السلسلة النصية كاملةً عن طريق استبدال النمط الحالي بـ `.` متبوعًا بـ المحدد الكمي `+`.

# --hints--

يجب تعديل متغير `pattern` ليطابق سلسلة `quote` كاملة.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), ".+") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '[^a-z]t'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
