---
id: 656475bbf1c2573de1d2c69c
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

في كل مرة تُنفذ فيها الشفرة، يجب أن ترى حرفًا عشوائيًا من سلسلة النص `all_characters`. هذا بالضبط ما تريد تحقيقه لإنشاء كلمة مرور عشوائية.

مع ذلك، الخوارزمية التي يعتمد عليها `random` تجعل الأعداد الزائفة العشوائية المولدة قابلة للتنبؤ. لذلك، رغم أن وحدة `random` مناسبة لأكثر التطبيقات شيوعًا، لا يمكن استخدامها لأغراض التشفير بسبب طبيعتها الحتمية.

بدلًا من استيراد `random`، استورد وحدة `secrets`. ثم غيّر استدعاء `print()` ليستخدم `secrets.choice(all_characters)`.

# --hints--

يجب استيراد وحدة `secrets` بدلًا من وحدة `random`.

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

يجب تعديل استدعاء `print()` الحالي ليطبع `secrets.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
