---
id: 6564cbb6311a40783b3f5de6
title: الخطوة 36
challengeType: 20
dashedName: step-36
---

# --description--

تسمح لك أصناف الحروف أيضًا بتحديد نطاق من الحروف للمطابقة. يجب أن تحدد الحرفين البداية والنهاية مفصولين بشرطة، `-`. تتبع الحروف ترتيب Unicode.

عدّل المتغير `pattern` ليطابق أي حرف `t` يسبقه حرف أحرف صغيرة في المتغير `quote`. استخدم نطاق الحروف من `a` إلى `z` لذلك.

# --hints--

يجب تعديل المتغير `pattern` ليطابق أي حرف `t` في `quote` يسبقه حرف أحرف صغيرة باستخدام الصنف `[a-z]`.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
