---
id: 6564642ce403fe29cc23fee7
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

الوحدة في Python هي ملف يحتوي على شفرة مصممة لأداء مهام محددة. تحتوي مكتبة Python القياسية على العديد من الوحدات التي يمكنك استيرادها واستخدامها في شفرتك. يمكنك تحقيق ذلك باستخدام جملة `import` متبوعة باسم الوحدة.

ابدأ هذا المشروع باستيراد الوحدة `string`.

# --hints--

يجب عليك استيراد الوحدة `string`.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
