---
id: 6564dd65c3c2fa873a83d213
title: خطوة 47
challengeType: 20
dashedName: step-47
---

# --description--

في صنف الحروف، يمكنك دمج عدة نطاقات بكتابة نطاق واحد بعد الآخر داخل الأقواس المربعة (دون أي أحرف إضافية). على سبيل المثال: `[a-d3-6]` هو دمج بين `[a-d]` و `[3-6]`.

الآن، عدل التعبير النمطي الأخير ليطابق أي حرف غير أبجدي رقمي. اجمع نطاقات `a-z` و `A-Z` و `0-9` في صنف حروف واحد وأضف `^` كأول حرف لعكس النمط.

# --hints--

يجب أن تكمل نمطك الرابع ليطابق أي حرف غير أبجدي رقمي.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\[\^a-zA-Z0-9\]\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
