---
id: 65646ffeaed2d238c562a014
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

من الاصطلاحات الشائعة وضع عبارات `import` في أعلى شفرتك. بالإضافة إلى ذلك، في حالة وجود عدة عبارات `import`، رتبها أبجديًا لتحسين قابلية القراءة.

في أعلى شفرتك، استورد الوحدة `random`.

# --hints--

يجب استيراد الوحدة `random`.

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
