---
id: 6564883669b5af4b69f794cc
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

المتغير `all_characters` هو سلسلة نصية تتكون من جميع الأحرف الصغيرة والكبيرة، وكل الأرقام العشرة، وعدد من الأحرف الخاصة.

قبل هذا المتغير مباشرة، أضف تعليقًا يقول `Combine all characters`.

# --hints--

يجب أن تضيف التعليق فوق متغير `all_characters` مباشرة.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
