---
id: 65649a631440b757f9f7b428
title: الخطوة 28
challengeType: 20
dashedName: step-28
---

# --description--

التعبير النمطي، أو التعبير النمطي، هو نمط يُستخدم لمطابقة تركيبة محددة من الأحرف داخل سلسلة نصية. هو بديل صالح لعبارات الشرط `if`/`else` عندما تحتاج إلى مطابقة أو إيجاد أنماط داخل سلسلة نصية لأغراض التحقق، استبدال الأحرف، وغيرها.

الدالة `compile()` من وحدة `re` تُحوّل السلسلة النصية الممررة كمعلمة إلى كائن تعبير نمطي يمكن استخدامه بواسطة طرق `re` الأخرى.

أعلن عن متغير جديد باسم `pattern` وقم بتعيين قيمة `re.compile('i')` لهذا المتغير.

# --hints--

يجب إعلان متغير `pattern`.

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

يجب تعيين `re.compile('i')` إلى متغير `pattern` الخاص بك.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
