---
id: 6564b8c9349bd76dc037967b
title: الخطوة 29
challengeType: 20
dashedName: step-29
---

# --description--

تُحلل الدالة `search()` من وحدة `re` السلسلة النصية الممررة كمعلمة بحثًا عن أول مكان يتطابق فيه التعبير النمطي مع السلسلة.

أعلن عن متغير باسم `quote` وعيّن له السلسلة النصية `'Not all those who wander are lost.'`. ثم اطبع نتيجة `pattern.search(quote)`.

# --hints--

يجب أن يكون لديك متغير `quote`.

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

يجب تعيين السلسلة النصية المقدمة إلى متغير `quote` الجديد.

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

يجب طباعة `pattern.search(quote)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
