---
id: 656470d517833a39bb8b5608
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

تحتوي وحدة `random` على مولد أرقام عشوائية زائفة. تعتمد معظم وظائفها على الدالة `random()`، التي تُرجع رقمًا عشريًا في النطاق بين `0.0` (شامل) و `1.0` (حصري).

استدعِ الدالة `random()` من وحدة `random` واطبع النتيجة.

# --hints--

يجب أن تطبع نتيجة استدعاء `random.random()`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
