---
id: 6564683821b2ee3174e7250e
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

يمكنك الوصول إلى الأدوات المعرفة داخل الوحدة المستوردة عبر ترميز النقطة. يعمل ترميز النقطة بإضافة نقطة متبوعة باسم الأداة إلى اسم الوحدة. على سبيل المثال، هكذا تصل إلى الثابت `ascii_lowercase`:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

في هذا المشروع، ستستخدم ثوابت مختلفة من وحدة `string`.

أعلن عن متغير جديد اسمه `letters` وعيّن له القيمة `string.ascii_letters`.

# --hints--

يجب إعلان متغير باسم `letters`.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

يجب تعيين `string.ascii_letters` إلى متغير `letters`.

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
