---
id: 656477845006313fbfea0ad1
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

رغم أن التأثير قد يبدو مساويًا لـ `random.choice()`، يضمن لك `secrets` أكثر مصادر العشوائية أمانًا التي يمكن لنظام التشغيل الخاص بك توفيرها.

الآن، احذف استدعائي `print()` الخاصين بك.

# --hints--

يجب حذف استدعائي `print()` الخاصين بك.

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
