---
id: 6564f32b18480893cf7799fd
title: الخطوة 52
challengeType: 20
dashedName: step-52
---

# --description--

الآن، اجمع سلسلتك النصية الخام مع سلسلة نصية f واستخدم المتغير `symbols` داخل فئة الحروف. تذكر أنه يمكنك إدخال متغير داخل سلسلة نصية f باستخدام الأقواس المعقوفة `{ }`.

# --hints--

العنصر الثاني في صف المرتب الرابع من القيود يجب أن يكون السلسلة النصية `fr'[{symbols}]'`.

```js
({ test: () => assert.match(code, /\(\s*special_chars\s*,\s*(fr|rf)("|')\[\{\s*symbols\s*\}\]\2/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
--fcc-editable-region--            
            (special_chars, r'\W')
--fcc-editable-region--            
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
