---
id: 656ef54070c72161e6feb90f
title: الخطوة 22
challengeType: 20
dashedName: step-22
---

# --description--

يبدو كل شيء جيدًا، لكن سيكون من المفيد وجود طريقة للتحقق من أن كلمة المرور المُولدة تلتزم بميزات محددة. على سبيل المثال، الحد الأدنى من الأحرف الخاصة، أو الأرقام، أو الأحرف الكبيرة/الصغيرة. ستتولى ذلك قريبًا جدًا.

في الوقت الحالي، أضف تعليقًا للكود على آخر سطرين من شفرتك.

# --hints--

يجب أن تحول آخر سطرين من شفرتك إلى تعليقات.

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
