---
id: 656ef89b4b486b7a16077864
title: الخطوة 26
challengeType: 20
dashedName: step-26
---

# --description--

الصف المرتب هو هيكل بيانات مدمج آخر في Python. الصفوف المرتبة تشبه القوائم كثيرًا، لكنها تُعرّف باستخدام أقواس عادية `()` بدلاً من الأقواس المربعة. أيضًا، الصفوف المرتبة غير قابلة للتغيير، على عكس القوائم.

```py
my_tuple = ('larch', 1, True)
```

قائمة `constraints` الخاصة بك ستخزن صفوفًا مرتبة. العنصر الأول من كل صف مرتب سيكون معلمة قيد.

عدّل تعيين قائمة `constraints` بإضافة صف مرتب إلى قائمتك. استخدم `nums` كالعنصر الأول وسلسلة نصية فارغة كالعنصر الثاني.

# --hints--

يجب أن تضيف صفًا مرتبًا إلى قائمة `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

يجب أن تضيف `(nums, '')` إلى قائمة `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
