---
id: 656f0bc259c1f6b57486ed68
title: الخطوة 33
challengeType: 20
dashedName: step-33
---

# --description--

للتحقق من أن كلمة المرور المُولدة تفي بالميزات المطلوبة، ستستخدم الدالة `findall()` من الوحدة `re`. هي مشابهة للدالة `search` لكنها تُرجع قائمة بكل حالات النمط المطابق.

استبدل استدعاء `search()` بـ `findall()` وتحقق من الناتج.

# --hints--

يجب تعديل استدعاء `print()` الحالي لديك باستبدال `search()` بـ `findall()`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.search(pattern, quote))
--fcc-editable-region--
```
