---
id: 6577333feab1e8e927014f03
title: الخطوة 63
challengeType: 20
dashedName: step-63
---

# --description--

وجود `all([expression for i in iterable])` يعني أن قائمة جديدة تُنشأ بتقييم `expression` لكل `i` في `iterable`. بعد أن تُكرّر الدالة `all()` على القائمة الجديدة، تُحذف القائمة تلقائيًا لأنها لم تعد ضرورية.

يمكن توفير الذاكرة باستخدام تعبير مولِّد. تتبع تعبيرات المولِّد صيغة نحوية مشابهة لتعبيرات القوائم لكنها تستخدم الأقواس العادية بدلًا من الأقواس المربعة.

حوّل تعبير القائمة الخاص بك إلى تعبير مولِّد بحذف الأقواس المربعة.

# --hints--

يجب أن تحوّل تعبير القائمة الخاص بك إلى تعبير مولِّد بحذف الأقواس المربعة.

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        if all(
            [
                constraint <= len(re.findall(pattern, password))
                for constraint, pattern in constraints
            ]
        ):
--fcc-editable-region--
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
