---
id: 6577319039f4f7de9251b822
title: الخطوة 61
challengeType: 20
dashedName: step-61
---

# --description--

بدلًا من استخدام حلقة ومتغير عداد، يمكنك تحقيق نفس النتيجة بطريقة مختلفة ستنفذها في الخطوات التالية.

`all()` دالة مدمجة في Python تُرجع `True` إذا كانت كل العناصر داخل قابل للتكرار المعطى تُقيّم إلى `True`. وإلا، تُرجع `False`.

استبدل حلقة `for` الحالية وعبارتي `if` بعبارة `if` واحدة. لاستخدام شرط `if`، استدعِ الدالة `all()` ومرر قائمة فارغة كمعلمة لاستدعاء الدالة.

# --hints--

يجب أن تستبدل حلقة `for` الحالية وعبارتي `if` بعبارة `if` واحدة.

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

يجب أن يكون شرط `if` الجديد هو `all([])`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

يجب أن يحتوي جسم عبارة `if` الجديدة على `break`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
            
        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
