---
id: 657732654845d8e2fb1217e6
title: الخطوة 62
challengeType: 20
dashedName: step-62
---

# --description--

في الوقت الحالي، تأخذ الدالة `all()` قائمة فارغة كمعلمة. املأ تلك القائمة الفارغة باستخدام صيغة الفهم بحيث يخزن المخزن نتائج تقييم التعبير `constraint <= len(re.findall(pattern, password))` لكل صف مرتب `constraint`-`pattern` في قائمة `constraints`.

بهذه الطريقة، ستخرج من حلقة `while` فقط بعد استيفاء جميع المتطلبات.

# --hints--

يجب أن تمرر `[constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints]` إلى الدالة `all()`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("all([constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints])")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        if all([]):
--fcc-editable-region--        
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
