---
id: 66eaddd04a9e533fba689001
title: 設計一個部落格文章卡片
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

在本實驗中，你將練習如何透過建立部落格文章卡片來設定背景和邊框的樣式。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `div` 元素，其類別為 `blog-post-card`，用來容納你所有的卡片元素。
2. 在 `.blog-post-card` 元素內，你應該有一個具有有效 `alt` 屬性和文字的影像元素，並且該元素的類別為 `post-img`。你可以將 `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` 用於影像的 `src` 屬性。
3. 你應該在 `.blog-post-card` 元素內有一個類別為 `post-content` 的 `div`，內容如下：
    - 一個具有 `post-title` 類別和用於你的部落格文章標題的文字的 `h2` 元素。
    - 一個具有 `post-excerpt` 類別和用於摘要你的部落格文章的文字的 `p` 元素。
    - 一個具有 `read-more` 類別且文字為 `Read More` 的 `a` 元素。
4. 你應該將 `body` 元素的背景顏色設為非白色的值。
5. 你應該將以下樣式套用到 `.blog-post-card` 元素：
    - 一個白色背景。
    - 圓角。
    - 你選擇的寬度。
    - 你選擇的文字對齊方式。
6. `.post-img` 元素應該設計成讓影像填滿卡片的整個寬度，並且有一個底部邊框。
7. `.post-content` 元素應該設計成在卡片內有內距。
8. `.post-title` 和 `.post-excerpt` 元素的文字顏色應該不是預設值，且四邊都應有邊距。
9. `.read-more` 元素應該被設計成按鈕樣式並具有：
    - 非預設的文字顏色。
    - 背景顏色。
    - 四邊的邊界。
    - 四邊的內距。
    - 圓角。
    - 設定為 `inline-block` 的 `display` 屬性。
    - 一個會改變其背景顏色的暫留效果。

**注意：** 請務必在你的 HTML 中聯結你的樣式表並套用你的 CSS。

# --hints--

你應該有一個帶有 `blog-post-card` 類別的 `div`。

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

你應該有一個 `img` 元素，其 `class` 為 `post-img`。請確保你的影像有一個帶有文字的 `alt` 屬性和一個有值的 `src` 屬性。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

你應該有一個 `div`，其類別為 `post-content`。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

你應該有一個帶有 `post-title` 類別的 `h2` 元素。請確保它有一些文字內容作為你的部落格文章標題。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

你應該有一個帶有 `post-excerpt` 類別的 `p`。確保它有一些文字內容來摘要你的部落格文章。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

你應該有一個 `a` 元素，其 `class` 為 `read-more`。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

你的 `.read-more` 元素應該有文字 `Read More`。

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

你的 `.blog-post-card` 元素應該有一個 `border-radius` 屬性且其值（不應該是 `0` 或負值）。

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

你的 `.blog-post-card` 元素應該有白色背景。

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

你的 `body` 元素應該有非白色的背景色。

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

你應該標的 `.blog-post-card` 並設定它的 `width` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

你應該標的 `.blog-post-card` 並設定它的 `text-align` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

你應該標的 `.post-content` 並設定它的 `padding` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

你的 `.read-more` 元素應該有暫留效果。

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

你應該標的 `.read-more` 並設定它的 `color` 屬性。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

你應該標的 `.read-more` 並設定它的 `background-color` 屬性。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

你應該標的 `.read-more` 並設定它的 `margin` 屬性。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

你應該標的 `.read-more` 並設定它的 `display` 屬性。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

你應該標的 `.read-more` 並設定它的 `border-radius` 屬性。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

你應該標的 `.read-more` 並設定它的 `padding` 屬性。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

你的 `.post-img` 元素應該填滿卡片的 `width`。

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

你的 `.post-img` 元素應該有一個 `border-bottom` 值。

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

你的 `.post-title` 和 `.post-excerpt` 元素應該有邊距和非預設值的文字顏色。

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
