---
id: a77dbc43c33f39daa4429b4f
title: Eine Funktion zum Prüfen von Boolean-Werten erstellen
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

In diesem Labor erstellen Sie eine Funktion, die überprüft, ob ein Wert als Boolean-Primitiv klassifiziert ist.

Boolean-Primitiven sind `true` und `false`.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `booWho` haben, die ein Argument entgegennimmt.
1. Wenn das übergebene Argument ein Boolean-Primitiv ist, sollte die Funktion `true` zurückgeben.
1. Wenn das Argument ein anderer Wert ist, sollte die Funktion `false` zurückgeben.

# --hints--

Sie sollten eine `booWho`-Funktion haben.

```js
assert.isFunction(booWho);
```

`booWho(true)` sollte `true` zurückgeben.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` sollte `true` zurückgeben.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` sollte `false` zurückgeben.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` sollte `false` zurückgeben.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` sollte `false` zurückgeben.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` sollte `false` zurückgeben.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` sollte `false` zurückgeben.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` sollte `false` zurückgeben.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` sollte `false` zurückgeben.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` sollte `false` zurückgeben.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
