---
id: 6814d8e1516e86b171929de4
title: "Zoezi la 1: Mizani ya Vowels"
challengeType: 28
dashedName: challenge-1
---

# --description--

Kutolewa mfuatano wa herufi, amua kama idadi ya vokali katika nusu ya kwanza ya mfuatano huo ni sawa na idadi ya vokali katika nusu ya pili.

- Mfuatano unaweza kuwa na herufi yoyote.
- Herufi `a`, `e`, `i`, `o`, na `u`, iwe ni herufi kubwa au herufi ndogo, zinachukuliwa kuwa vokali.
- Ikiwa kuna idadi isiyo ya kawaida ya herufi katika mfuatano, puuza herufi ya katikati.

# --hints--

`isBalanced("racecar")` inapaswa kurudisha `true`.

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")` inapaswa kurudisha `true`.

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")` inapaswa kurudisha `false`.

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")` inapaswa kurudisha `false`.

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")` inapaswa kurudisha `true`.

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")` inapaswa kurudisha `false`.

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")` inapaswa kurudisha `true`.

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

`isBalanced("a")` inapaswa kurudisha `true`.

```js
assert.isTrue(isBalanced("a"));
```

`isBalanced("E")` inapaswa kurudisha `true`.

```js
assert.isTrue(isBalanced("E"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const chars = [...s];
  const half = Math.floor(chars.length / 2);

  const firstHalf = chars.slice(0, half);
  const secondHalf =
    chars.length % 2 === 0 ? chars.slice(half) : chars.slice(half + 1);

  const countVowels = arr =>
    arr.filter(c => vowels.includes(c.toLowerCase())).length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
