---
id: 681cb1a2dab50c87ddb2e514
title: "Zoezi la 3: Mfululizo wa Fibonacci"
challengeType: 28
dashedName: challenge-3
---

# --description--

Mfululizo wa Fibonacci ni mfululizo wa nambari ambapo kila nambari ni jumla ya nambari mbili zilizotangulia. Unapoanza na `0` na `1`, nambari kumi za kwanza katika mfululizo ni `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

Ukitoa safu ya data yenye nambari mbili za kwanza za mfululizo wa Fibonacci, na nambari kamili inayowakilisha urefu wa mfululizo, rudisha safu ya data yenye mfululizo wa urefu uliotolewa.

- Kitendakazi chako kinapaswa kushughulikia mfululizo wa urefu wowote mkubwa au sawa na sifuri.
- Ikiwa urefu ni sifuri, rudisha safu ya data tupu.
- Kumbuka kuwa nambari za kuanzia ni sehemu ya mfululizo.

# --hints--

`fibonacciSequence([0, 1], 20)` inapaswa kurudisha `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]);
```

`fibonacciSequence([21, 32], 1)` inapaswa kurudisha `[21]`.

```js
assert.deepEqual(fibonacciSequence([21, 32], 1), [21]);
```

`fibonacciSequence([0, 1], 0)` inapaswa kurudisha `[]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 0), []);
```

`fibonacciSequence([10, 20], 2)` inapaswa kurudisha `[10, 20]`.

```js
assert.deepEqual(fibonacciSequence([10, 20], 2), [10, 20]);
```

`fibonacciSequence([123456789, 987654321], 5)` inapaswa kurudisha `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
assert.deepEqual(fibonacciSequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541]);
```

# --seed--

## --seed-contents--

```js
function fibonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function fibonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2];
    sequence.push(nextValue);
  }
  return sequence;
}
```
