---
id: 681cb1afdab50c87ddb2e516
title: "Zoezi la 5: Mfuatano wa herufi wa Jbelmud"
challengeType: 28
dashedName: challenge-5
---

# --description--

Ukipata mfuatano wa herufi, rudisha toleo lililochanganywa la mfuatano huo ambapo kila neno hubadilishwa kwa kutumia vizingiti vifuatavyo:

- Herufi ya kwanza na ya mwisho ya maneno hubaki mahali pake
- Herufi zote kati ya herufi ya kwanza na ya mwisho hupangwa kwa mpangilio wa alfabeti.
- Mfuatano wa herufi wa ingizo hautakuwa na alama za uandishi, na utakuwa mzima kwa herufi ndogo.

# --hints--

`jbelmu("hello world")` inapaswa kurudisha `"hello wlord"`.

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` inapaswa kurudisha `"i love jbelmud text"`.

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` inapaswa kurudisha `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` inapaswa kurudisha `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
