---
id: 681cb05adab50c87ddb2e513
title: "Zoezi la 2: Ukaguzi wa Msingi"
challengeType: 28
dashedName: challenge-2
---

# --description--

Kutolewa mfuatano wa herufi unaowakilisha nambari, na msingi wa nambari kamili kutoka 2 hadi 36, amua kama nambari hiyo ni halali katika msingi huo.

- Mfuatano wa herufi unaweza kuwa na nambari kamili, na herufi kubwa au ndogo.
- Ukaguzi unapaswa kuwa usiozingatia tofauti ya herufi kubwa na ndogo.
- Msingi unaweza kuwa nambari yoyote kati ya 2-36.
- Nambari ni halali ikiwa kila herufi ni tarakimu halali katika msingi uliotolewa.
- Mfano wa tarakimu halali kwa misingi:
  - Msingi 2: 0-1
  - Msingi 8: 0-7
  - Msingi 10: 0-9
  - Msingi 16: 0-9 na A-F
  - Msingi 36: 0-9 na A-Z

# --hints--

`isValidNumber("10101", 2)` should return `true`.

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` should return `false`.

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` should return `true`.

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` should return `false`.

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` should return `true`.

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` should return `false`.

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` should return `true`.

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` should return `true`.

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` should return `true`.

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` should return `true`.

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` should return `false`.

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` should return `true`.

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` should return `false`.

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` should return `true`.

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` should return `true`.

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` should return `true`.

```js
assert.isTrue(isValidNumber("z", 36))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
