---
id: 681cb1b0dab50c87ddb2e51b
title: "Changamoto 10: Mikwaju 3"
challengeType: 28
dashedName: challenge-10
---

# --description--

Ikiwa umepewa nambari kamili kati ya 1 na 10,000, rudisha hesabu ya ni nambari ngapi kutoka 1 hadi nambari hiyo ambayo mraba wake una angalau tarakimu moja ya 3.

# --hints--

`squaresWithThree(1)` inapaswa kurudisha `0`.

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` inapaswa kurudisha `1`.

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` inapaswa kurudisha `19`.

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` inapaswa kurudisha `326`.

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` inapaswa kurudisha `4531`.

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
