---
id: 6821ebf3237de8297eaee797
title: "Zoezi la 21: Kizalishaji cha Hex"
challengeType: 28
dashedName: challenge-21
---

# --description--

Ukipokea mfuatano wa herufi wa rangi ya CSS yenye jina, zalishe msimbo wa rangi ya hexadecimal (hex) wa nasibu ambao una rangi kuu katika rangi iliyotolewa.

- Kitendakazi kinapaswa kushughulikia `"red"`, `"green"`, au `"blue"` kama hoja ya ingizo.
- Ikiwa ingizo si mojawapo ya hayo, kitendakazi kinapaswa kurudisha `"Invalid color"`.
- Kitendakazi kinapaswa kurudisha msimbo wa rangi ya hex wenye herufi sita wa nasibu ambapo thamani ya rangi ya ingizo ni kubwa kuliko nyingine yoyote.
- Mfano wa matokeo halali kwa ingizo fulani:

| Ingizo   | Matokeo   |
|---------|----------|
| `"red"`   | `"FF0000"` |
| `"green"` | `"00FF00"` |
| `"blue"`  | `"0000FF"` |

# --hints--

`generateHex("yellow")` inapaswa kurudisha `"Invalid color"`.

```js
assert.equal(generateHex("yellow"), "Invalid color");
```

`generateHex("red")` inapaswa kurudisha mfuatano wa herufi sita.

```js
assert.lengthOf(generateHex("red"), 6);
```

`generateHex("red")` inapaswa kurudisha msimbo halali wa rangi ya hex wenye herufi sita.

```js
const hex = generateHex("red").toUpperCase();
const isValidHex = /^[0-9A-F]{6}$/.test(hex);
assert.isTrue(isValidHex);
```

`generateHex("red")` inapaswa kurudisha rangi halali ya hex yenye thamani ya rangi nyekundu kubwa kuliko rangi nyingine.

```js
const hex = generateHex("red").toUpperCase();
const isValidHex = /^[0-9A-F]{6}$/.test(hex);
assert.isTrue(isValidHex);

const r = parseInt(hex.slice(0, 2), 16);
const g = parseInt(hex.slice(2, 4), 16);
const b = parseInt(hex.slice(4, 6), 16);

assert.isAbove(r, g);
assert.isAbove(r, b);
```

Kupiga wito `generateHex("red")` mara mbili kunapaswa kurudisha thamani mbili tofauti za rangi za hex ambapo nyekundu ni kuu.

```js
const hex1 = generateHex("red").toUpperCase();
const isValidHex1 = /^[0-9A-F]{6}$/.test(hex1)
assert.isTrue(isValidHex1);

const r1 = parseInt(hex1.slice(0, 2), 16);
const g1 = parseInt(hex1.slice(2, 4), 16);
const b1 = parseInt(hex1.slice(4, 6), 16);

assert.isAbove(r1, g1);
assert.isAbove(r1, b1);

const hex2 = generateHex("red").toUpperCase();
const isValidHex2 = /^[0-9A-F]{6}$/.test(hex2);
assert.isTrue(isValidHex2);

const r2 = parseInt(hex2.slice(0, 2), 16);
const g2 = parseInt(hex2.slice(2, 4), 16);
const b2 = parseInt(hex2.slice(4, 6), 16);

assert.isAbove(r2, g2);
assert.isAbove(r2, b2);
assert.notEqual(hex1, hex2);
```

Kupiga wito `generateHex("green")` mara mbili kunapaswa kurudisha thamani mbili tofauti za rangi za hex ambapo kijani ni kuu.

```js
const hex1 = generateHex("green").toUpperCase();
const isValidHex1 = /^[0-9A-F]{6}$/.test(hex1)
assert.isTrue(isValidHex1);

const r1 = parseInt(hex1.slice(0, 2), 16);
const g1 = parseInt(hex1.slice(2, 4), 16);
const b1 = parseInt(hex1.slice(4, 6), 16);

assert.isAbove(g1, r1);
assert.isAbove(g1, b1);

const hex2 = generateHex("green").toUpperCase();
const isValidHex2 = /^[0-9A-F]{6}$/.test(hex2);
assert.isTrue(isValidHex2);

const r2 = parseInt(hex2.slice(0, 2), 16);
const g2 = parseInt(hex2.slice(2, 4), 16);
const b2 = parseInt(hex2.slice(4, 6), 16);

assert.isAbove(g2, r2);
assert.isAbove(g2, b2);
assert.notEqual(hex1, hex2);
```

Kupiga wito `generateHex("blue")` mara mbili kunapaswa kurudisha thamani mbili tofauti za rangi za hex ambapo buluu ni kuu.

```js
const hex1 = generateHex("blue").toUpperCase();
const isValidHex1 = /^[0-9A-F]{6}$/.test(hex1)
assert.isTrue(isValidHex1);

const r1 = parseInt(hex1.slice(0, 2), 16);
const g1 = parseInt(hex1.slice(2, 4), 16);
const b1 = parseInt(hex1.slice(4, 6), 16);

assert.isAbove(b1, r1);
assert.isAbove(b1, g1);

const hex2 = generateHex("blue").toUpperCase();
const isValidHex2 = /^[0-9A-F]{6}$/.test(hex2);
assert.isTrue(isValidHex2);

const r2 = parseInt(hex2.slice(0, 2), 16);
const g2 = parseInt(hex2.slice(2, 4), 16);
const b2 = parseInt(hex2.slice(4, 6), 16);

assert.isAbove(b2, r2);
assert.isAbove(b2, g2);
assert.notEqual(hex1, hex2);
```

# --seed--

## --seed-contents--

```js
function generateHex(color) {

  return color;
}
```

# --solutions--

```js
function generateHex(color) {
  const toHex = n => n.toString(16).padStart(2, "0").toUpperCase();

  const dominant = Math.floor(Math.random() * 86) + 170;
  const weak1 = Math.floor(Math.random() * 170);
  const weak2 = Math.floor(Math.random() * 170);

  let r, g, b;

  switch (color) {
    case "red":
      r = dominant;
      g = weak1;
      b = weak2;
      break;
    case "green":
      r = weak1;
      g = dominant;
      b = weak2;
      break;
    case "blue":
      r = weak1;
      g = weak2;
      b = dominant;
      break;
    default:
      return "Invalid color";
  }

  return `${toHex(r)}${toHex(g)}${toHex(b)}`;
}
```
