---
id: 6821ebee237de8297eaee796
title: "Zoezi la 20: Nakala za safu ya data"
challengeType: 28
dashedName: challenge-20
---

# --description--

Ukitoa safu ya data ya nambari kamili, rudisha safu ya data ya nambari kamili zinazojirudia zaidi ya mara moja katika safu ya awali, zikiwa zimepangwa kwa mpangilio wa kuongezeka. Ikiwa hakuna thamani inayojirudia zaidi ya mara moja, rudisha safu ya data tupu.

- Jumuisha tu mfano mmoja wa kila thamani katika safu ya data inayorudishwa.

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` inapaswa kurudisha `[]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` inapaswa kurudisha `[1, 2]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` inapaswa kurudisha `[-6, 0, 2, 4, 5, 23]`.

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
