---
id: 6821ebe4237de8297eaee794
title: "Zoezi la 18: Bora ya Pili"
challengeType: 28
dashedName: challenge-18
---

# --description--

Ikiwa umepewa safu ya data ya nambari kamili zinazowakilisha bei za kompyuta mpakato tofauti, na nambari kamili inayowakilisha bajeti yako, rudisha:

1. Kompyuta mpakato ya pili ghali zaidi ikiwa iko ndani ya bajeti yako, au
2. Kompyuta mpakato ghali zaidi ambayo iko ndani ya bajeti yako, au
3. `0` ikiwa hakuna kompyuta mpakato yoyote ndani ya bajeti yako.

- Bei zinazojirudia zinapaswa kupuuzwa.

# --hints--

`getLaptopCost([1500, 2000, 1800, 1400], 1900)` inapaswa kurudisha `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900)` inapaswa kurudisha `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([2099, 1599, 1899, 1499], 2200)` inapaswa kurudisha `1899`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 2200), 1899);
```

`getLaptopCost([2099, 1599, 1899, 1499], 1000)` inapaswa kurudisha `0`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 1000), 0);
```

`getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` inapaswa kurudisha `1400`

```js
assert.equal(getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400);
```

# --seed--

## --seed-contents--

```js
function getLaptopCost(laptops, budget) {

  return laptops;
}
```

# --solutions--

```js
function getLaptopCost(laptops, budget) {
  laptops = [...new Set(laptops)].sort((a, b) => b - a);

  if (budget >= laptops[1]) return laptops[1];
  if (budget < laptops[laptops.length - 1]) return 0;

  for (let i = 2; i < laptops.length; i++) {
    if (budget >= laptops[i]) return laptops[i];
  }
}
```
