---
id: 6821ebc9237de8297eaee78f
title: "Zoezi la 13: Namba msingi isiyo ya kawaida"
challengeType: 28
dashedName: challenge-13
---

# --description--

Ukichukulia nambari kamili, amua kama nambari hiyo ni namba msingi au namba msingi hasi.

- Namba msingi ni nambari kamili chanya kubwa kuliko 1 ambayo inaweza kugawanywa tu na 1 na yenyewe.
- Namba msingi hasi ni toleo hasi la namba msingi chanya.
- `1` na `0` hazihesabiwi kuwa namba msingi.

# --hints--

`isUnnaturalPrime(1)` inapaswa kurudisha `false`.

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)` inapaswa kurudisha `false`.

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)` inapaswa kurudisha `true`.

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)` inapaswa kurudisha `true`.

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)` inapaswa kurudisha `false`.

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)` inapaswa kurudisha `true`.

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)` inapaswa kurudisha `true`.

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)` inapaswa kurudisha `false`.

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)` inapaswa kurudisha `false`.

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
