---
id: 6821ebf8237de8297eaee798
title: "Zoezi la 22: mfululizo wa Tribonacci"
challengeType: 28
dashedName: challenge-22
---

# --description--

Mfululizo wa Tribonacci ni mfululizo wa nambari ambapo kila nambari ni jumla ya nambari tatu zilizotangulia. Unapoanza na `0`, `0` na `1`, nambari kumi za kwanza katika mfululizo ni `0`, `0`, `1`, `1`, `2`, `4`, `7`, `13`, `24`, `44`.

Ukitoa safu ya data yenye nambari tatu za kwanza za mfululizo wa Tribonacci, na nambari kamili inayowakilisha urefu wa mfululizo, rudisha safu ya data yenye mfululizo wa urefu uliotolewa.

- Kitendakazi chako kinapaswa kushughulikia mfululizo wa urefu wowote mkubwa au sawa na sifuri.
- Ikiwa urefu ni sifuri, rudisha safu ya data tupu.
- Kumbuka kuwa nambari za kuanzia ni sehemu ya mfululizo.

# --hints--

`tribonacciSequence([0, 0, 1], 20)` inapaswa kurudisha `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)` inapaswa kurudisha `[21]`.


```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)` inapaswa kurudisha `[]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)` inapaswa kurudisha `[10, 20]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)` inapaswa kurudisha `[10, 20, 30]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)` inapaswa kurudisha `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```
