---
id: 6821ebce237de8297eaee790
title: "Zoezi la 14: Mapigano ya herufi"
challengeType: 28
dashedName: challenge-14
---

# --description--

Kutolewa mfuatano wa herufi miwili unaowakilisha jeshi lako na jeshi la mpinzani, kila herufi kutoka jeshi lako inapigana na herufi iliyopo katika nafasi sawa kutoka jeshi la mpinzani kwa kutumia sheria zifuatazo:

- Herufi `a-z` zina nguvu ya `1-26`, mtawaliwa.
- Herufi `A-Z` zina nguvu ya `27-52`, mtawaliwa.
- Nambari `0-9` zina nguvu sawa na thamani yao ya uso.
- Herufi zote nyingine zina thamani ya sifuri.
- Kila herufi inaweza kupigana vita moja tu.

Kwa kila vita, herufi yenye nguvu zaidi hushinda. Jeshi lenye ushindi mwingi zaidi hushinda vita. Rudisha thamani zifuatazo:

- `"Opponent retreated"` ikiwa jeshi lako lina herufi nyingi zaidi kuliko jeshi la mpinzani.
- `"We retreated"` ikiwa jeshi la mpinzani lina herufi nyingi zaidi kuliko lako.
- `"We won"` ikiwa jeshi lako limeshinda vita nyingi zaidi.
- `"We lost"` ikiwa jeshi la mpinzani limeshinda vita nyingi zaidi.
- `"It was a tie"` ikiwa majeshi yote mawili yameshinda vita idadi sawa.

# --hints--

`battle("Hello", "World")` inapaswa kurudisha `"We lost"`.

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")` inapaswa kurudisha `"We won"`.

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")` inapaswa kurudisha `"We won"`.

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")` inapaswa kurudisha `"Opponent retreated"`.

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")` inapaswa kurudisha `"We retreated"`.

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")` inapaswa kurudisha `"It was a tie"`.

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")` inapaswa kurudisha `"It was a tie"`.

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```
