---
id: 6821ebdf237de8297eaee793
title: "Zoezi la 17: Kutotii mpangilio wa operesheni"
challengeType: 28
dashedName: challenge-17
---

# --description--

Ukizingatia safu ya data ya nambari kamili na safu ya data ya opereta za mfuatano wa herufi, tumia operesheni hizo kwa nambari moja baada ya nyingine kutoka kushoto kwenda kulia. Rudia operesheni hizo kadri inavyohitajika hadi nambari zote zitumike. Rudisha matokeo ya mwisho.

Kwa mfano, ukizingatia `[1, 2, 3, 4, 5]` na `['+', '*']`, rudisha matokeo ya kutathmini `1 + 2 * 3 + 4 * 5` kutoka kushoto kwenda kulia bila kuzingatia mpangilio wa kawaida wa operesheni.

- Opereta halali ni `+`, `-`, `*`, `/`, na `%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` inapaswa kurudisha `3`

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` inapaswa kurudisha `38`

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` inapaswa kurudisha `60`

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` inapaswa kurudisha `30`

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])` inapaswa kurudisha `-2`

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```
