---
id: 68adce01c0e1144d0a90295c
title: "Zoezi la 28: Mchambuzi wa Nambari za Kirumi"
challengeType: 28
dashedName: challenge-28
---

# --description--

Kutolewa mfuatano wa herufi unaowakilisha nambari ya Kirumi, rudisha thamani yake ya nambari kamili.

Nambari za Kirumi zinajumuisha alama na thamani zifuatazo:

| Alama | Thamani |
|--------|---------|
| I      | 1       |
| V      | 5       |
| X      | 10      |
| L      | 50      |
| C      | 100     |
| D      | 500     |
| M      | 1000    |

- Nambari husomwa kutoka kushoto kwenda kulia. Ikiwa nambari ndogo inaonekana kabla ya nambari kubwa, thamani yake hutolewa. Vinginevyo, thamani huongezwa.

# --hints--

`parseRomanNumeral("III")` inapaswa kurudisha `3`.

```js
assert.equal(parseRomanNumeral("III"), 3);
```

`parseRomanNumeral("IV")` inapaswa kurudisha `4`.

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

`parseRomanNumeral("XXVI")` inapaswa kurudisha `26`.

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

`parseRomanNumeral("XCIX")` inapaswa kurudisha `99`.

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

`parseRomanNumeral("CDLX")` inapaswa kurudisha `460`.

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

`parseRomanNumeral("DIV")` inapaswa kurudisha `504`.

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

`parseRomanNumeral("MMXXV")` inapaswa kurudisha `2025`.

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
