---
id: 68adce01c0e1144d0a90295e
title: "Zoezi la 29: Mtengenezaji wa herufi za awali"
challengeType: 28
dashedName: challenge-29
---

# --description--

Ikiwa umepewa mfuatano wa herufi unao kuwa na neno moja au zaidi, rudisha herufi za awali za maneno hayo kwa kutumia masharti yafuatayo:

- Herufi za awali zinapaswa kuwa herufi kubwa, isipokuwa kama imetajwa vinginevyo.
- Herufi za awali za maneno haya yafuatayo zisitiliwe isipokuwa kama ni neno la kwanza katika mfuatano uliotolewa: `a`, `for`, `an`, `and`, `by`, na `of`.
- Herufi za awali zinapaswa kurudishwa kwa mpangilio zilivyopewa.
- Herufi za awali hazipaswi kuwa na nafasi yoyote kati yao.

# --hints--

`buildAcronym("Search Engine Optimization")` inapaswa kurudisha `"SEO"`.

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")` inapaswa kurudisha `"FAQ"`.

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")` inapaswa kurudisha `"NASA"`.

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` inapaswa kurudisha `"FBI"`.

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")` inapaswa kurudisha `"FYI"`.

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("By the way")` inapaswa kurudisha `"BTW"`.

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` inapaswa kurudisha `"AUHWPOTIMSH"`.

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
