---
id: 68adce01c0e1144d0a902956
title: "Zoezi la 25: Kurudia vokali"
challengeType: 28
dashedName: challenge-25
---

# --description--

Ukitoa mfuatano wa herufi, rudisha toleo jipya la mfuatano huo ambapo kila vokali inarudiwa mara moja zaidi kuliko vokali iliyotangulia uliyokutana nayo. Kwa mfano, vokali ya kwanza katika sentensi haibadiliki. Vokali ya pili inaonekana mara mbili mfululizo. Vokali ya tatu inaonekana mara tatu mfululizo, na kadhalika.

- Herufi `a`, `e`, `i`, `o`, na `u`, iwe ni herufi kubwa au herufi ndogo, zinachukuliwa kuwa vokali.
- Vokali ya asili inapaswa kuhifadhi aina yake ya herufi.
- Vokali zinazorudiwa zinapaswa kuwa herufi ndogo.
- Herufi zote zisizo za vokali zinapaswa kuhifadhi aina yao ya herufi kama zilivyo asili.

# --hints--

`repeatVowels("hello world")` inapaswa kurudisha `"helloo wooorld"`.

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` inapaswa kurudisha `"freeeCooodeeeeCaaaaamp"`.

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` inapaswa kurudisha `"AEeIiiOoooUuuuu"`.

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` inapaswa kurudisha `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
