---
id: 68adce01c0e1144d0a90295a
title: "Zoezi la 27: Mzunguko wa matrix"
challengeType: 28
dashedName: challenge-27
---

# --description--

Ikiwa umepewa matrix (safu ya data ya safu za data), zungusha matrix hiyo kwa digrii 90 kwa mwelekeo wa saa na urudishe. Kwa mfano, ukipata `[[1, 2], [3, 4]]`, inayofanana na hii:

| 1 | 2 |
|---|---|
| 3 | 4 |

Unapaswa kurudisha `[[3, 1], [4, 2]]`, inayofanana na hii:

| 3 | 1 |
|---|---|
| 4 | 2 |

# --hints--

`rotate([[1]])` inapaswa kurudisha `[[1]]`.

```js
assert.deepEqual(rotate([[1]]), [[1]]);
```

`rotate([[1, 2], [3, 4]])` inapaswa kurudisha `[[3, 1], [4, 2]]`.

```js
assert.deepEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]]);
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` inapaswa kurudisha `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`.

```js
assert.deepEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]]);
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` inapaswa kurudisha `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`.

```js
assert.deepEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function rotate(matrix) {

  return matrix;
}
```

# --solutions--

```js
function rotate(matrix) {
  const n = matrix.length;
  const result = Array.from({ length: n }, () => Array(n).fill(0));

  for (let i = 0; i < n; i++) {
    for (let j = 0; j < n; j++) {
      result[j][n - 1 - i] = matrix[i][j];
    }
  }

  return result;
}
```
