---
id: 68af0687ef34c76c28ffa54b
title: "Zoezi la 32: Geuza Sentensi"
challengeType: 28
dashedName: challenge-32
---

# --description--

Kutolewa mfuatano wa herufi wa maneno, rudisha mfuatano mpya wa herufi wenye maneno yakiwa katika mpangilio wa kinyume. Kwa mfano, neno la kwanza linapaswa kuwa mwishoni mwa mfuatano uliorudishwa, na neno la mwisho linapaswa kuwa mwanzoni mwa mfuatano uliorudishwa.

- Katika mfuatano uliotolewa, maneno yanaweza kutenganishwa na nafasi moja au zaidi.
- Mfuatano uliorudishwa unapaswa kuwa na nafasi moja tu kati ya maneno.

# --hints--

`reverseSentence("world hello")` inapaswa kurudisha `"hello world"`.

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` inapaswa kurudisha `"git commit push"`.

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` inapaswa kurudisha `"sudo install npm"`.

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` inapaswa kurudisha `"export function default import"`.

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
