---
id: 68af0687ef34c76c28ffa54f
title: "Zoezi la 34: Nambari Zinazokosekana"
challengeType: 28
dashedName: challenge-34
---

# --description--

Ikiwa umepewa safu ya data ya nambari kamili kutoka 1 hadi `n`, ikijumuisha, rudisha safu ya data ya nambari zote kamili zinazokosekana kati ya 1 na `n` (ambapo `n` ni nambari kubwa zaidi katika safu ya data iliyotolewa).

- Safu ya data iliyotolewa inaweza isipangwe na inaweza kuwa na nakala rudufu.
- Safu ya data itakayorejeshwa inapaswa kuwa katika mpangilio wa kuongezeka.
- Ikiwa hakuna nambari kamili zinazokosekana, rudisha safu ya data tupu.

# --hints--

`findMissingNumbers([1, 3, 5])` inapaswa kurudisha `[2, 4]`.

```js
assert.deepEqual(findMissingNumbers([1, 3, 5]), [2, 4]);
```

`findMissingNumbers([1, 2, 3, 4, 5])` inapaswa kurudisha `[]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 10])` inapaswa kurudisha `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
assert.deepEqual(findMissingNumbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([10, 1, 10, 1, 10, 1])` inapaswa kurudisha `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
assert.deepEqual(findMissingNumbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([3, 1, 4, 1, 5, 9])` inapaswa kurudisha `[2, 6, 7, 8]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])` inapaswa kurudisha `[11]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11]);
```

# --seed--

## --seed-contents--

```js
function findMissingNumbers(arr) {

  return arr;
}
```

# --solutions--

```js
function findMissingNumbers(arr) {
  if (arr.length === 0) return [];
  
  const n = Math.max(...arr);
  const set = new Set(arr);
  const result = [];

  for (let i = 1; i <= n; i++) {
    if (!set.has(i)) {
      result.push(i);
    }
  }
  return result;
}
```
