---
id: 68af0687ef34c76c28ffa547
title: "Zoezi la 30: Herufi za kipekee"
challengeType: 28
dashedName: challenge-30
---

# --description--

Kutolewa mfuatano wa herufi, amua kama herufi zote katika mfuatano huo ni za kipekee.

- Herufi kubwa na herufi ndogo zinapaswa kuchukuliwa kuwa herufi tofauti.

# --hints--

`allUnique("abc")` inapaswa kurudisha `true`.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` inapaswa kurudisha `true`.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` inapaswa kurudisha `true`.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` inapaswa kurudisha `true`.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` inapaswa kurudisha `false`.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` inapaswa kurudisha `false`.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` inapaswa kurudisha `false`.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
