---
id: 68af0687ef34c76c28ffa54d
title: "Zoezi la 33: Muda wa Kuangalia Skrini"
challengeType: 28
dashedName: challenge-33
---

# --description--

Kutolewa safu ya data ya nambari kamili saba, zinazowakilisha muda wa wiki moja, ambapo kila nambari kamili ni idadi ya saa ulizotumia kwenye simu yako siku hiyo, amua kama muda huo wa kuangalia skrini ni mwingi kulingana na vizingiti hivi:

- Ikiwa siku yoyote moja ina saa 10 au zaidi, ni muda mwingi.
- Ikiwa wastani wa siku tatu mfululizo ni zaidi au sawa na saa 8, ni muda mwingi.
- Ikiwa wastani wa siku saba ni zaidi au sawa na saa 6, ni muda mwingi.

# --hints--

`tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7])` inapaswa kurudisha `false`.

```js
assert.isFalse(tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7]));
```

`tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3])` inapaswa kurudisha `false`.

```js
assert.isFalse(tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3]));
```

`tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6])` inapaswa kurudisha `false`.

```js
assert.isFalse(tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6]));
```

`tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4])` inapaswa kurudisha `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4]));
```

`tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0])` inapaswa kurudisha `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0]));
```

`tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4])` inapaswa kurudisha `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4]));
```

`tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6])` inapaswa kurudisha `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6]));
```

# --seed--

## --seed-contents--

```js
function tooMuchScreenTime(hours) {

  return hours;
}
```

# --solutions--

```js
function tooMuchScreenTime(hours) {
  for (let h of hours) {
    if (h >= 10) return true;
  }

  for (let i = 0; i <= 4; i++) {
    const avg3 = (hours[i] + hours[i+1] + hours[i+2]) / 3;
    if (avg3 >= 8) return true;
  }

  const weeklyAvg = hours.reduce((sum, h) => sum + h, 0) / 7;
  if (weeklyAvg >= 6) return true;

  return false;
}
```
