---
id: 68b06e589bf2273243814777
title: "Zoezi la 39: Jaza Tanki"
challengeType: 28
dashedName: challenge-39
---

# --description--

Kutolewa ukubwa wa tanki la mafuta, kiwango cha sasa cha mafuta, na bei kwa galoni, rudisha gharama ya kujaza tanki lote.

- `tankSize` ni uwezo wa jumla wa tanki kwa galoni.
- `fuelLevel` ni kiasi cha sasa cha mafuta kwenye tanki kwa galoni.
- `pricePerGallon` ni gharama ya galoni moja ya mafuta.
- Thamani inayorudishwa inapaswa kuzungushwa hadi sehemu mbili za desimali kwa muundo: `"$d.dd"`.

# --hints--

`costToFill(20, 0, 4.00)` inapaswa kurudisha `"$80.00"`.

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` inapaswa kurudisha `"$17.50"`.

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` inapaswa kurudisha `"$29.25"`.

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` inapaswa kurudisha `"$0.00"`.

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` inapaswa kurudisha `"$21.89"`.

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
