---
id: 68b06e589bf2273243814773
title: "Zoezi la 37: Kuweka herufi kubwa mwanzoni mwa sentensi"
challengeType: 28
dashedName: challenge-37
---

# --description--

Ukipata aya, rudisha aya mpya ambapo herufi ya kwanza ya kila sentensi imeandikwa kwa herufi kubwa.

- Herufi nyingine zote ziendelee kama zilivyo.
- Sentensi zinaweza kumalizika kwa nukta (`.`), alama moja au zaidi za kuuliza (`?`), au alama moja au zaidi za mshangao (`!`).

# --hints--

`capitalize("this is a simple sentence.")` inapaswa kurudisha `"This is a simple sentence."`.

```js
assert.equal(capitalize("this is a simple sentence."), "This is a simple sentence.");
```

`capitalize("hello world. how are you?")` inapaswa kurudisha `"Hello world. How are you?"`.

```js
assert.equal(capitalize("hello world. how are you?"), "Hello world. How are you?");
```

`capitalize("i did today's coding challenge... it was fun!!")` inapaswa kurudisha `"I did today's coding challenge... It was fun!!"`.

```js
assert.equal(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!");
```

`capitalize("crazy!!!strange???unconventional...sentences.")` inapaswa kurudisha `"Crazy!!!Strange???Unconventional...Sentences."`.

```js
assert.equal(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.");
```

`capitalize("there's a space before this period . why is there a space before that period ?")` inapaswa kurudisha `"There's a space before this period . Why is there a space before that period ?"`.

```js
assert.equal(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?");
```

# --seed--

## --seed-contents--

```js
function capitalize(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function capitalize(paragraph) {
  return paragraph.replace(/([.!?]+)(\s*)([a-z])/g, (match, punc, spaces, char) => {
    return punc + spaces + char.toUpperCase();
  })
  .replace(/^([a-z])/, (match, char) => char.toUpperCase());
}
```
