---
id: 68b06e589bf2273243814775
title: "Zoezi la 38: Kizalishaji cha slug"
challengeType: 28
dashedName: challenge-38
---

# --description--

Ukipata mfuatano wa herufi, rudisha toleo linalofaa kwa URL la mfuatano huo kwa kutumia masharti yafuatayo:

- Herufi zote ziwe herufi ndogo.
- Herufi, nambari, au nafasi zisizo za herufi, nambari, au nafasi zifutwe.
- Nafasi zote zibadilishwe na msimbo wa nafasi uliokodishwa kwa URL `%20`.
- Nafasi zinazojirudia zifanyike badiliko la `%20` moja tu.
- Mfuatano uliorudishwa usiwe na `%20` mwanzoni au mwishoni.

# --hints--

`generateSlug("helloWorld")` inapaswa kurudisha `"helloworld"`.

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` inapaswa kurudisha `"hello%20world"`.

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` inapaswa kurudisha `"helloworld"`.

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` inapaswa kurudisha `"hello%20world"`.

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` inapaswa kurudisha `"h3110%20w0r1d"`.

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
