---
id: 68b06e589bf2273243814771
title: "Zoezi la 36: Mrekebishaji wa Thermostat"
challengeType: 28
dashedName: challenge-36
---

# --description--

Kutolewa joto la sasa la chumba na joto lengwa, rudisha mfuatano wa herufi unaoonyesha jinsi ya kurekebisha joto la chumba kulingana na vizingiti hivi:

- Rudisha `"heat"` ikiwa joto la sasa liko chini ya joto lengwa.
- Rudisha `"cool"` ikiwa joto la sasa liko juu ya joto lengwa.
- Rudisha `"hold"` ikiwa joto la sasa ni sawa na joto lengwa.

# --hints--

`adjustThermostat(68, 72)` inapaswa kurudisha `"heat"`.

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` inapaswa kurudisha `"cool"`.

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` inapaswa kurudisha `"hold"`.

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` inapaswa kurudisha `"heat"`.

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` inapaswa kurudisha `"cool"`.

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` inapaswa kurudisha `"hold"`.

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
