---
id: 68b1f72371a5ac895ac70a06
title: "Zoezi la 42: Hifadhi ya video"
challengeType: 28
dashedName: challenge-42
---

# --description--

Ukizingatia ukubwa wa video, kitengo cha ukubwa wa video, uwezo wa diski kuu, na kitengo cha diski kuu, rudisha idadi ya video ambazo diski kuu inaweza kuhifadhi kwa kutumia vizingiti vifuatavyo:

- Kitengo cha ukubwa wa video kinaweza kuwa bytes (`"B"`), kilobytes (`"KB"`), megabytes (`"MB"`), au gigabytes (`"GB"`).
- Ikiwa hakutolewa mojawapo ya vitengo vya video vilivyotajwa hapo juu, rudisha `"Invalid video unit"`.
- Kitengo cha uwezo wa diski kuu kinaweza kuwa gigabytes (`"GB"`) au terabytes (`"TB"`).
- Ikiwa hakutolewa mojawapo ya vitengo vya diski kuu vilivyotajwa hapo juu, rudisha `"Invalid drive unit"`.
- Rudisha idadi ya video kamili ambazo diski inaweza kuhifadhi.
- Tumia mabadiliko yafuatayo:

| Kitengo | Sawia na |
|:-------:|:--------:|
| 1 B     |   1 B    |
| 1 KB    |  1000 B  |
| 1 MB    |  1000 KB |
| 1 GB    |  1000 MB |
| 1 TB    |  1000 GB |

Kwa mfano, ukipokea `500`, `"MB"`, `100`, na `"GB"` kama hoja, amua ni video ngapi za 500 MB zinaweza kuhifadhiwa kwenye diski kuu ya 100 GB.

# --hints--

`numberOfVideos(500, "MB", 100, "GB")` inapaswa kurudisha `200`.

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")` inapaswa kurudisha `"Invalid video unit"`.

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")` inapaswa kurudisha `"Invalid drive unit"`.

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")` inapaswa kurudisha `4000`.

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")` inapaswa kurudisha `1466`.

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```
