---
id: 68b1f72371a5ac895ac70a0a
title: "Zoezi la 44: Kioo cha mfuatano wa herufi"
challengeType: 28
dashedName: challenge-44
---

# --description--

Ukizingatia mfuatano wa herufi miwili, amua kama mfuatano wa pili ni kioo cha wa kwanza.

- Mfuatano wa herufi huchukuliwa kuwa kioo ikiwa una herufi sawa kwa mpangilio wa mviringo.
- Tazama herufi kubwa na herufi ndogo kama tofauti.
- Puuza herufi zisizo za alfabeti zote.

# --hints--

`isMirror("helloworld", "helloworld")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
