---
id: 68b1f72371a5ac895ac70a02
title: "Zoezi la 40: Hifadhi ya Picha"
challengeType: 28
dashedName: challenge-40
---

# --description--

Ukizingatia ukubwa wa picha kwa megabaiti (MB), na uwezo wa diski kuu kwa gigabaiti (GB), rudisha idadi ya picha ambazo diski kuu inaweza kuhifadhi kwa kutumia masharti yafuatayo:

- 1 gigabaiti ni sawa na megabaiti 1000.
- Rudisha idadi kamili ya picha ambazo diski inaweza kuhifadhi.

# --hints--

`numberOfPhotos(1, 1)` inapaswa kurudisha `1000`.

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

`numberOfPhotos(2, 1)` inapaswa kurudisha `500`.

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

`numberOfPhotos(4, 256)` inapaswa kurudisha `64000`.

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

`numberOfPhotos(3.5, 750)` inapaswa kurudisha `214285`.

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

`numberOfPhotos(3.5, 5.5)` inapaswa kurudisha `1571`.

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```
