---
id: 68b1f72371a5ac895ac70a04
title: "Zoezi la 41: Hifadhi ya faili"
challengeType: 28
dashedName: challenge-41
---

# --description--

Ukizingatia ukubwa wa faili, kitengo cha ukubwa wa faili, na uwezo wa diski kuu kwa gigabaiti (GB), rudisha idadi ya mafaili ambayo diski kuu inaweza kuhifadhi kwa kutumia vizingiti vifuatavyo:

- Kitengo cha ukubwa wa faili kinaweza kuwa bait (`"B"`), kilobaiti (`"KB"`), au megabaiti (`"MB"`).
- Rudisha idadi ya mafaili kamili ambayo diski inaweza kuingiza.
- Tumia mabadiliko yafuatayo:

| Kitengo | Sawia na |
|:-------:|:--------:|
| 1 B     |   1 B    |
| 1 KB    |  1000 B  |
| 1 MB    |  1000 KB |
| 1 GB    |  1000 MB |

Kwa mfano, ukipata `500`, `"KB"`, na `1` kama hoja, amua ni mafaili mangapi ya 500 KB yanaweza kuhifadhiwa kwenye diski kuu ya 1 GB.

# --hints--

`numberOfFiles(500, "KB", 1)` inapaswa kurudisha `2000`.

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)` inapaswa kurudisha `20000`.

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)` inapaswa kurudisha `200`.

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)` inapaswa kurudisha `366210`.

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)` inapaswa kurudisha `453514`.

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)` inapaswa kurudisha `166666`.

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
