---
id: 68b7687dded630607aceccb3
title: "Zoezi la 49: Mchambuzi wa kichwa cha CSV"
challengeType: 28
dashedName: challenge-49
---

# --description--

Ukichukulia mstari wa kwanza wa faili ya thamani zilizoachwa kwa koma (CSV), rudisha safu ya data inayojumuisha vichwa vya sehemu.

- Mstari wa kwanza wa faili ya CSV una vichwa vya sehemu vilivyoachwa kwa koma.
- Ondoa nafasi zozote za mwanzoni au mwishoni kutoka kila kichwa cha sehemu.

# --hints--

`getHeadings("name,age,city")` inapaswa kurudisha `["name", "age", "city"]`.

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` inapaswa kurudisha `["first name", "last name", "phone"]`.

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` inapaswa kurudisha `["username", "email", "signup date"]`.

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
