---
id: 68b7cadffed0e75a517da671
title: "Zoezi la 51: Mpangilio wa nambari ya simu"
challengeType: 28
dashedName: challenge-51
---

# --description--

Ukipata mfuatano wa herufi wenye tarakimu kumi na moja, rudisha mfuatano huo kama nambari ya simu kwa muundo huu: `"+D (DDD) DDD-DDDD"`.

# --hints--

`formatNumber("05552340182")` inapaswa kurudisha `"+0 (555) 234-0182"`.

```js
assert.equal(formatNumber("05552340182"), "+0 (555) 234-0182");
```

`formatNumber("15554354792")` inapaswa kurudisha `"+1 (555) 435-4792"`.

```js
assert.equal(formatNumber("15554354792"), "+1 (555) 435-4792");
```

# --seed--

## --seed-contents--

```js
function formatNumber(number) {

  return number;
}
```

# --solutions--

```js
function formatNumber(number) {
  const country = number[0];
  const area = number.slice(1, 4);
  const prefix = number.slice(4, 7);
  const line = number.slice(7);

  return `+${country} (${area}) ${prefix}-${line}`;
}
```
