---
id: 68b7cadffed0e75a517da675
title: "Zoezi la 53: Kutumia nambari kamili chanya kubadilisha kutoka desimali hadi binary"
challengeType: 28
dashedName: challenge-53
---

# --description--

Ikiwa umepewa nambari kamili isiyo na hasi, rudisha uwakilishi wake wa binary kama mfuatano wa herufi.

Nambari ya binary hutumia tu tarakimu `0` na `1` kuwakilisha nambari yoyote. Ili kubadilisha nambari ya desimali kuwa binary, gawanya nambari hiyo kwa `2` mara kwa mara na rekodi mabaki. Rudia hadi nambari iwe sifuri. Soma mabaki yaliyorekodiwa kuanzia ya mwisho hadi ya kwanza. Kwa mfano, kubadilisha `12` kuwa binary:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` katika binary ni `1100`.

# --hints--

`toBinary(5)` inapaswa kurudisha `"101"`.

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)` inapaswa kurudisha `"1100"`.

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)` inapaswa kurudisha `"110010"`.

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)` inapaswa kurudisha `"1100011"`.

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
