---
id: 68b7687dded630607aceccb1
title: "Zoezi la 48: Kigeuzi cha spam"
challengeType: 28
dashedName: challenge-48
---

# --description--

Ikiwa nambari ya simu iko katika muundo wa `"+A (BBB) CCC-DDDD"`, ambapo kila herufi inawakilisha nambari kama ifuatavyo:

- `A` inawakilisha msimbo wa nchi na inaweza kuwa na nambari yoyote ya tarakimu.
- `BBB` inawakilisha msimbo wa eneo na daima itakuwa na tarakimu tatu.
- `CCC` na `DDDD` zinawakilisha nambari ya ndani na daima zitakuwa na tarakimu tatu na nne mtawalia.

Amua kama ni nambari ya spam kulingana na vigezo vifuatavyo:

- Msimbo wa nchi ni wa tarakimu zaidi ya 2 au hauanza na sifuri (`0`).
- Msimbo wa eneo ni zaidi ya 900 au chini ya 200.
- Jumla ya tarakimu tatu za kwanza za nambari ya ndani inaonekana ndani ya tarakimu nne za mwisho za nambari ya ndani.
- Nambari ina tarakimu ile ile mara nne au zaidi mfululizo (bila kuzingatia herufi za muundo).

# --hints--

`isSpam("+0 (200) 234-0182")` inapaswa kurudisha `false`.

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` inapaswa kurudisha `false`.

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
