---
id: 68b7cadffed0e75a517da673
title: "Zoezi la 52: Nambari ya binary kuwa desimali"
challengeType: 28
dashedName: challenge-52
---

# --description--

Ikiwa umepewa mfuatano wa herufi unaowakilisha nambari ya binary, rudisha sawa yake katika mfumo wa desimali kama nambari.

Nambari ya binary hutumia tu tarakimu `0` na `1` kuwakilisha nambari yoyote. Ili kubadilisha binary kuwa desimali, zidisha kila tarakimu kwa nguvu ya `2` kisha ziongeze pamoja. Anza kwa kuzidisha tarakimu ya kulia kabisa kwa `2^0`, tarakimu inayofuata kushoto kwa `2^1`, na kadhalika. Mara zote tarakimu zote zitakapozidishwa kwa nguvu ya `2`, ongeza matokeo hayo pamoja.

Kwa mfano, nambari ya binary `101` ni sawa na `5` katika mfumo wa desimali kwa sababu:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")` inapaswa kurudisha `5`.

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")` inapaswa kurudisha `10`.

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")` inapaswa kurudisha `18`.

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")` inapaswa kurudisha `85`.

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
