---
id: 68c1a929005bf54d342aa8d4
title: "Zoezi la 57: Wiki ya Anga siku ya 3: Simu Nyumbani"
challengeType: 28
dashedName: challenge-57
---

# --description--

Kwa siku ya tatu ya Wiki ya Anga, unapewa safu ya data ya nambari zinazowakilisha umbali (kwa kilomita) kati yako, satelaiti, na sayari yako ya nyumbani katika njia ya data ya mawasiliano. Tambua ni muda gani ujumbe utakaopelekwa kupitia njia hiyo utachukua kufikia sayari lengwa kwa kutumia vizingiti vifuatavyo:

- Thamani ya kwanza katika safu ya data ni umbali kutoka mahali ulipo hadi satelaiti ya kwanza.
- Kila thamani inayofuata, isipokuwa ya mwisho, ni umbali hadi satelaiti inayofuata.
- Thamani ya mwisho katika safu ya data ni umbali kutoka satelaiti ya mwisho hadi sayari yako ya nyumbani.
- Ujumbe husafiri kwa kasi ya 300,000 km/s.
- Kila satelaiti ambayo ujumbe unapitisha huongeza ucheleweshaji wa usambazaji wa sekunde 0.5.
- Rudisha nambari iliyozungushwa hadi sehemu 4 za desimali, ukiondoa sifuri za mwisho.

# --hints--

`sendMessage([300000, 300000])` inapaswa kurudisha `2.5`.

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` inapaswa kurudisha `3.0627`.

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` inapaswa kurudisha `364.5`.

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` inapaswa kurudisha `1674.3333`.

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` inapaswa kurudisha `2.4086`.

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` inapaswa kurudisha `21.1597`.

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
