---
id: 68c1a929005bf54d342aa8d6
title: "Zoezi la 59: Wiki ya Anga Siku ya 5: Eneo la Goldilocks"
challengeType: 28
dashedName: challenge-59
---

# --description--

Kwa siku ya tano ya Wiki ya Anga, utahesabu "eneo la Goldilocks" la nyota - eneo linalozunguka nyota ambapo hali ni "sawa kabisa" kwa maji ya kioevu kuwepo.

Ukizingatia uzito wa nyota, rudisha safu ya data yenye umbali wa kuanzia na wa mwisho wa Eneo lake la Goldilocks kwa Vipimo vya Anga.

Ili kuhesabu Eneo la Goldilocks:

1. Tafuta mwangaza wa nyota kwa kuinua uzito wake kwa nguvu ya 3.5.
2. Mwanzo wa eneo ni mara 0.95 ya mzizi wa mraba wa mwangaza wake.
3. Mwisho wa eneo ni mara 1.37 ya mzizi wa mraba wa mwangaza wake.

- Rudisha umbali uliopunguzwa hadi sehemu mbili za desimali.

Kwa mfano, ukipata `1` kama uzito, rudisha `[0.95, 1.37]`.

# --hints--

`goldilocksZone(1)` inapaswa kurudisha `[0.95, 1.37]`.

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)` inapaswa kurudisha `[0.28, 0.41]`.

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)` inapaswa kurudisha `[21.85, 31.51]`.

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)` inapaswa kurudisha `[9.38, 13.52]`.

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)` inapaswa kurudisha `[179.69, 259.13]`.

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
