---
id: 68c1a929005bf54d342aa8d5
title: "Zoezi la 58: Wiki ya Anga Day 4: Sehemu ya Kutua"
challengeType: 28
dashedName: challenge-58
---

# --description--

Katika siku ya nne ya Wiki ya Anga, unapewa matriki ya nambari (safu ya data ya safu za data), zinazowakilisha maeneo yanayoweza kutumiwa kutua kwa rover yako. Tafuta sehemu salama zaidi ya kutua kulingana na sheria zifuatazo:

- Kila sehemu katika matriki itakuwa na nambari kutoka `0-9`, ikiwa ni pamoja.
- Kila `0` inawakilisha sehemu inayowezekana ya kutua.
- Kila nambari isipokuwa `0` ni hatari sana kutua. Kadri nambari inavyoongezeka, ndivyo hatari inavyoongezeka.
- Sehemu salama zaidi inafafanuliwa kama seli ya `0` ambayo seli zinazozunguka (hadi majirani 4, usizingatie pembe) zina jumla ya hatari ndogo zaidi.
- Puuza majirani walioko nje ya mipaka (kona na kingo zina majirani wachache tu).
- Rudisha viashiria vya sehemu salama zaidi ya kutua. Daima kutakuwa na sehemu moja tu salama zaidi.

Kwa mfano, ukipata:

```js
[
  [1, 0],
  [2, 0]
]
```

Rudisha `[0, 1]`, viashiria vya `0` katika safu ya kwanza.

# --hints--

`findLandingSpot([[1, 0], [2, 0]])` inapaswa kurudisha `[0, 1]`.

```js
assert.deepEqual(findLandingSpot([[1, 0], [2, 0]]), [0, 1]);
```

`findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` inapaswa kurudisha `[1, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1]);
```

`findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` inapaswa kurudisha `[2, 2]`.

```js
assert.deepEqual(findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2]);
```

`findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` inapaswa kurudisha `[2, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1]);
```

# --seed--

## --seed-contents--

```js
function findLandingSpot(matrix) {

  return matrix;
}
```

# --solutions--

```js
function findLandingSpot(matrix) {
  let bestSpot = null;
  let lowestNeighborSum = Infinity;

  for (let i = 0; i < matrix.length; i++) {
    for (let j = 0; j < matrix[i].length; j++) {

      if (matrix[i][j] === 0) {
        let currentNeighborSum = 0;

        if (i > 0) currentNeighborSum += matrix[i - 1][j];
        if (j < matrix[i].length - 1) currentNeighborSum += matrix[i][j + 1];
        if (i < matrix.length - 1) currentNeighborSum += matrix[i + 1][j];
        if (j > 0) currentNeighborSum += matrix[i][j - 1];

        if (currentNeighborSum < lowestNeighborSum) {
          lowestNeighborSum = currentNeighborSum;
          bestSpot = [i, j];
        }
      }
    }
  }

  return bestSpot;
}
```
