---
id: 68c1a929005bf54d342aa8d2
title: "Zoezi la 55: Wiki ya Anga Day 1: Uainishaji wa Nyota"
challengeType: 28
dashedName: challenge-55
---

# --description--

Tarehe 4 Oktoba inaashiria kuanza kwa Wiki ya Anga Duniani. Siku saba zijazo zitakuletea changamoto za uandishi wa programu zenye mandhari ya astronomia.

Kwa changamoto ya leo, unapewa joto la uso wa nyota kwa Kelvin (K) na unahitaji kubaini uainishaji wake wa nyota kulingana na viwango vifuatavyo:

- `"O"`: 30,000 K au zaidi
- `"B"`: 10,000 K - 29,999 K
- `"A"`: 7,500 K - 9,999 K
- `"F"`: 6,000 K - 7,499 K
- `"G"`: 5,200 K - 5,999 K
- `"K"`: 3,700 K - 5,199 K
- `"M"`: 0 K - 3,699 K

- Rudisha uainishaji wa nyota uliotolewa.

# --hints--

`classification(5778)` inapaswa kurudisha `"G"`.

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` inapaswa kurudisha `"M"`.

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` inapaswa kurudisha `"A"`.

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` inapaswa kurudisha `"K"`.

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` inapaswa kurudisha `"M"`.

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` inapaswa kurudisha `"O"`.

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` inapaswa kurudisha `"F"`.

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` inapaswa kurudisha `"B"`.

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
