---
id: 68c497f3aaefc9fd9f1b0e24
title: "Zoezi la 60: Wiki ya Anga Day 6: Awamu ya Mwezi"
challengeType: 28
dashedName: challenge-60
---

# --description--

Kwa siku ya sita ya Wiki ya Anga, utapewa tarehe katika muundo `"YYYY-MM-DD"` na unahitaji kubaini awamu ya mwezi kwa siku hiyo kwa kutumia sheria zifuatazo:

Tumia mzunguko rahisi wa mwezi wa siku 28, uliogawanywa katika awamu nne sawa:

- `"New"`: siku 1 - 7
- `"Waxing"`: siku 8 - 14
- `"Full"`: siku 15 - 21
- `"Waning"`: siku 22 - 28

Baada ya siku 28, mzunguko unarudi na siku 1, mwezi mpya.

- Tumia `"2000-01-06"` kama rejea ya mwezi mpya (siku 1 ya mzunguko) kubaini awamu ya siku iliyotolewa.
- Hautapewa tarehe yoyote kabla ya tarehe ya rejea.
- Rudisha awamu sahihi kama mfuatano wa herufi.

**Kumbuka:** Siku 1 inawakilisha siku ya mwezi mpya, ikimaanisha siku 0 zimepita tangu mwezi mpya uliopita.

# --hints--

`moonPhase("2000-01-12")` inapaswa kurudisha `"New"`.

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` inapaswa kurudisha `"Waxing"`.

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` inapaswa kurudisha `"Full"`.

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` inapaswa kurudisha `"Waning"`.

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` inapaswa kurudisha `"New"`.

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
