---
id: 68c497f3aaefc9fd9f1b0e26
title: "Zoezi la 62: Hex hadi Desimali"
challengeType: 28
dashedName: challenge-62
---

# --description--

Imetolewa mfuatano wa herufi unaowakilisha nambari ya hexadecimal (msingi 16), rudisha thamani yake ya desimali (msingi 10) kama nambari kamili.

Hexadecimal ni mfumo wa nambari unaotumia tarakimu 16:

- `0-9` wakilisha thamani za `0` hadi `9`.
- `A-F` wakilisha thamani za `10` hadi `15`.

Hapa kuna jedwali la mabadiliko ya sehemu:

| Hexadecimal | Desimali |
|:-----------:|:--------:|
| 0           | 0        |
| 1           | 1        |
| ...         | ...      |
| 9           | 9        |
| A           | 10       |
| ...         | ...      |
| F           | 15       |
| 10          | 16       |
| ...         | ...      |
| 9F          | 159      |
| A0          | 160      |
| ...         | ...      |
| FF          | 255      |
| 100         | 256      |

- Mfuatano wa herufi utakuwa na herufi `0–9` na `A–F` tu.

# --hints--

`hexToDecimal("A")` inapaswa kurudisha `10`.

```js
assert.equal(hexToDecimal("A"), 10);
```

`hexToDecimal("15")` inapaswa kurudisha `21`.

```js
assert.equal(hexToDecimal("15"), 21);
```

`hexToDecimal("2E")` inapaswa kurudisha `46`.

```js
assert.equal(hexToDecimal("2E"), 46);
```

`hexToDecimal("FF")` inapaswa kurudisha `255`.

```js
assert.equal(hexToDecimal("FF"), 255);
```

`hexToDecimal("A3F")` inapaswa kurudisha `2623`.

```js
assert.equal(hexToDecimal("A3F"), 2623);
```

# --seed--

## --seed-contents--

```js
function hexToDecimal(hex) {

  return hex;
}
```

# --solutions--

```js
function hexToDecimal(hex) {

  return parseInt(hex, 16);
}
```
