---
id: 68cae5b538ff798bbd4da008
title: "Zoezi la 70: Kichujio cha sifa ya kipengele cha HTML"
challengeType: 28
dashedName: challenge-70
---

# --description--

Ukipata mfuatano wa herufi wa kipengele halali cha HTML, rudisha sifa za kipengele hicho kwa kutumia vigezo vifuatavyo:

- Utapewa kipengele kimoja tu.
- Sifa zitakuwa katika muundo: `attribute="value"`.
- Rudisha safu ya data ya mfuatano wa herufi yenye kila sifa ya kipengele na thamani yake, zikitenganishwa kwa koma, katika muundo huu: `["attribute1, value1", "attribute2, value2"]`.
- Rudisha sifa kwa mpangilio walivyopewa.
- Ikiwa hakuna sifa zitakazopatikana, rudisha safu ya data tupu.

# --hints--

`extractAttributes('<span class="red"></span>')` inapaswa kurudisha `["class, red"]`.

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` inapaswa kurudisha `["charset, UTF-8"]`.

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` inapaswa kurudisha `[]`.

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` inapaswa kurudisha `["name, email", "type, email", "required, true"]`.

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` inapaswa kurudisha `["id, submit", "class, btn btn-primary"]`.

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
