---
id: 68cae5b538ff798bbd4da003
title: "Zoezi la 65: Hesabu ya mfuatano wa herufi"
challengeType: 28
dashedName: challenge-65
---

# --description--

Ukizingatia mfuatano wa herufi miwili, amua ni mara ngapi mfuatano wa herufi wa pili unaonekana katika wa kwanza.

- Mfuatano wa mfano unaweza kuingiliana katika mfuatano wa herufi wa kwanza. Kwa mfano, `"aaa"` ina `"aa"` mara mbili. `a` mbili za kwanza na mbili za pili.

# --hints--

`count('abcdefg', 'def')` inapaswa kurudisha `1`.

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` inapaswa kurudisha `0`.

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` inapaswa kurudisha `2`.

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` inapaswa kurudisha `3`.

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` inapaswa kurudisha `10`.

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
