---
id: 68cae5b538ff798bbd4da002
title: "Zoezi la 64: 24 hadi 12"
challengeType: 28
dashedName: challenge-64
---

# --description--

Imetolewa mfuatano wa herufi unaowakilisha saa za mchana katika mfumo wa saa 24 wa `"HHMM"`, rudisha saa hiyo katika mfumo wake sawa wa saa 12 wa `"H:MM AM"` au `"H:MM PM"`.

- Ingizo lililotolewa daima litakuwa mfuatano wa herufi wenye tarakimu nne katika mfumo wa saa 24, kuanzia `"0000"` hadi `"2359"`.

# --hints--

`to12("1124")` inapaswa kurudisha `"11:24 AM"`.

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")` inapaswa kurudisha `"9:00 AM"`.

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")` inapaswa kurudisha `"2:55 PM"`.

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")` inapaswa kurudisha `"11:46 PM"`.

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")` inapaswa kurudisha `"12:30 AM"`.

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```
