---
id: 68d2ba1468508398389487cf
title: "Zoezi la 74: Nyimbo unazopenda zaidi"
challengeType: 28
dashedName: challenge-74
---

# --description--

Je, unakumbuka iPods? Mfano wa kwanza ulitoka miaka 24 iliyopita leo, Oktoba 23, 2001.

Ukizingatia safu ya data ya vitu vya nyimbo vinavyowakilisha orodha yako ya nyimbo ya iPod, rudisha safu ya data yenye vichwa vya nyimbo mbili zinazochezwa zaidi, ukiwa na wimbo unaochezwa zaidi kwanza.

- Kila kitu kitakuwa na sifa ya `"title"` (mfuatano wa herufi), na sifa ya `"plays"` (nambari kamili).

# --hints--

`favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` inapaswa kurudisha `["Sync or Swim", "Earbud Blues"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"]);
```

`favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` inapaswa kurudisha `["Clickwheel Love", "99 Downloads"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"]);
```

`favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` inapaswa kurudisha `["Song B", "Song C"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"]);
```

# --seed--

## --seed-contents--

```js
function favoriteSongs(playlist) {

  return playlist;
}
```

# --solutions--

```js
function favoriteSongs(playlist) {
  const sorted = [...playlist].sort((a, b) => b.plays - a.plays);
  return sorted.slice(0, 2).map(song => song.title);
}
```
