---
id: 68d2ba1468508398389487d0
title: "Zoezi la 75: Hazina iliyofichwa"
challengeType: 28
dashedName: challenge-75
---

# --description--

Imetolewa safu ya data ya 2D inayowakilisha ramani ya sakafu ya bahari ambayo ina hazina iliyofichwa, na safu ya data yenye kuratibu (`[row, column]`) kwa ajili ya kuogelea kwa utafutaji wako wa hazina, rudisha `"Empty"`, `"Found"`, au `"Recovered"` kwa kutumia sheria zifuatazo:

- Safu ya data ya 2D iliyotolewa itakuwa na hazina moja tu ambayo haijarekebishwa, ambayo itachukua seli nyingi.
- Kila seli katika safu ya data ya 2D itakuwa na moja ya thamani zifuatazo:
  - `"-"`: Hakuna hazina.
  - `"O"`: Sehemu ya hazina ambayo haijapatikana.
  - `"X"`: Sehemu ya hazina ambayo tayari imepatikana.

- Ikiwa eneo la kuogelea halina hazina, rudisha `"Empty"`.
- Ikiwa eneo la kuogelea linapata hazina, lakini angalau sehemu moja ya hazina bado haijapatikana, rudisha `"Found"`.
- Ikiwa eneo la kuogelea linapata sehemu ya mwisho ya hazina ambayo haijapatikana, rudisha `"Recovered"`.

Kwa mfano, ikitolewa:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

Na `[2, 1]` kwa kuratibu za eneo la kuogelea, rudisha `"Recovered"` kwa sababu kuogelea kulipata sehemu ya mwisho ya hazina ambayo haijapatikana.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` inapaswa kurudisha `"Recovered"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` inapaswa kurudisha `"Empty"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` inapaswa kurudisha `"Found"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` inapaswa kurudisha `"Found"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` inapaswa kurudisha `"Recovered"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` inapaswa kurudisha `"Empty"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```
