---
id: 68d30845cc08266018fc46be
title: "Zoezi la 78: Mfululizo wa nambari kamili"
challengeType: 28
dashedName: challenge-78
---

# --description--

Ukipata nambari kamili chanya, rudisha mfuatano wa herufi unaoonyesha nambari kamili zote kuanzia `1` hadi, na ikiwa ni pamoja na, nambari hiyo, kwa mpangilio wa nambari.

Kwa mfano, ukipata `5`, rudisha `"12345"`.

# --hints--

`sequence(5)` inapaswa kurudisha `"12345"`.

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)` inapaswa kurudisha `"12345678910"`.

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)` inapaswa kurudisha `"1"`.

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)` inapaswa kurudisha `"123456789101112131415161718192021222324252627"`.

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
