---
id: 68d30845cc08266018fc46bc
title: "Zoezi la 76: DNA ya nyongeza"
challengeType: 28
dashedName: challenge-76
---

# --description--

Kutolewa mfuatano wa herufi unaowakilisha mfululizo wa DNA, rudisha mfuatano wake wa nyongeza ukitumia sheria zifuatazo:

- DNA inajumuisha herufi `"A"`, `"C"`, `"G"`, na `"T"`.
- Herufi `"A"` na `"T"` ni za nyongeza kwa kila mmoja.
- Herufi `"C"` na `"G"` ni za nyongeza kwa kila mmoja.

Kwa mfano, ukipata `"ACGT"`, rudisha `"TGCA"`.

# --hints--

`complementaryDNA("ACGT")` inapaswa kurudisha `"TGCA"`.

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` inapaswa kurudisha `"TACGCATGCAATCG"`.

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` inapaswa kurudisha `"CCGAATGCTAGCTTC"`.

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` inapaswa kurudisha `"CTAGATCGATCCGATCGATC"`.

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
