---
id: 68e39ed6106dac2f0a98fd63
title: "Zoezi la 81: namba msingi ya n"
challengeType: 28
dashedName: challenge-81
---

# --description--

Namba msingi ni nambari kamili chanya kubwa kuliko 1 ambayo hugawanyika tu na 1 na yenyewe. Nambari tano za msingi za kwanza ni `2`, `3`, `5`, `7`, na `11`.

Ukitoa nambari kamili chanya `n`, rudisha namba msingi ya `n`. Kwa mfano, ukitoa `5` rudisha namba ya msingi ya 5: `11`.

# --hints--

`nthPrime(5)` inapaswa kurudisha `11`.

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)` inapaswa kurudisha `29`.

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)` inapaswa kurudisha `53`.

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)` inapaswa kurudisha `523`.

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)` inapaswa kurudisha `7919`.

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
