---
id: 68e39ed6106dac2f0a98fd65
title: "Zoezi la 83: Uthibitisho wa Sahihi"
challengeType: 28
dashedName: challenge-83
---

# --description--

Kutolewa mfuatano wa herufi wa ujumbe, mfuatano wa herufi wa funguo ya siri, na nambari ya sahihi, amua kama sahihi ni halali kwa kutumia njia hii ya usimbaji:

- Herufi katika ujumbe na funguo ya siri zina thamani hizi:
  - `a` hadi `z` zina thamani `1` hadi `26` mtawalia.
  - `A` hadi `Z` zina thamani `27` hadi `52` mtawalia.
- Herufi nyingine zote hazina thamani.
- Hesabu sahihi kwa kuchukua jumla ya ujumbe pamoja na jumla ya funguo ya siri.

Kwa mfano, ikitolewa ujumbe `"foo"` na funguo ya siri `"bar"`, sahihi itakuwa `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

Kagua kama sahihi iliyohesabiwa inalingana na sahihi iliyotolewa.

# --hints--

`verify("foo", "bar", 57)` inapaswa kurudisha `true`.

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` inapaswa kurudisha `false`.

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` inapaswa kurudisha `true`.

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` inapaswa kurudisha `false`.

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` inapaswa kurudisha `true`.

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` inapaswa kurudisha `true`.

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
