---
id: 68e39ed6106dac2f0a98fd62
title: "Zoezi la 80: Mpangaji wa barua pepe"
challengeType: 28
dashedName: challenge-80
---

# --description--

Mnamo Oktoba 29, 1971, barua pepe ya kwanza kabisa ilitumwa, ikianzisha muundo wa `username@domain` ambao bado tunautumia. Sasa, kuna mabilioni ya anwani za barua pepe.

Katika changamoto hii, unapewa orodha ya anwani za barua pepe na unahitaji kuzipanga kwa mpangilio wa alfabeti kwa jina la kikoa kwanza (sehemu baada ya `@`), na jina la mtumiaji pili (sehemu kabla ya `@`).

- Upangaji unapaswa kuwa usiozingatia herufi kubwa au ndogo.
- Ikiwa barua pepe zaidi ya moja zina jina la kikoa sawa, zipange kwa jina la mtumiaji.
- Rudisha safu ya anwani zilizopangwa.
- Anwani zilizorejeshwa zinapaswa kuhifadhi herufi zao za awali.

Kwa mfano, ukipata `["jill@mail.com", "john@example.com", "jane@example.com"]`, rudisha `["jane@example.com", "john@example.com", "jill@mail.com"]`.

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` inapaswa kurudisha `["jane@example.com", "john@example.com", "jill@mail.com"]`.

```js
assert.deepEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"]);
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` inapaswa kurudisha `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`.

```js
assert.deepEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"]);
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` inapaswa kurudisha `["user@x.com", "user@y.com", "user@z.com"]`.

```js
assert.deepEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"]);
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` inapaswa kurudisha `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`.

```js
assert.deepEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]);
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` inapaswa kurudisha `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`.

```js
assert.deepEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]);
```

# --seed--

## --seed-contents--

```js
function sort(emails) {

  return emails;
}
```

# --solutions--

```js
function sort(emails) {
  return emails.slice().sort((a, b) => {
    const [userA, domainA] = a.split('@');
    const [userB, domainB] = b.split('@');

    const domainCompare = domainA.toLowerCase().localeCompare(domainB.toLowerCase());
    if (domainCompare !== 0) return domainCompare;

    return userA.toLowerCase().localeCompare(userB.toLowerCase());
  });
}
```
