---
id: 68ee9e3066cfd4eb2328e8a8
title: "Zoezi la 89: Kuhesabu Kadi"
challengeType: 28
dashedName: challenge-89
---

# --description--

Kifurushi cha kawaida cha kadi za kuchezea kina kadi 13 za kipekee katika kila rangi. Ukipata nambari kamili inayowakilisha idadi ya kadi za kuchagua kutoka kwenye kifurushi, rudisha idadi ya mchanganyiko wa kipekee wa kadi unazoweza kuchagua.

- Mpangilio hauna umuhimu. Kuchagua kadi A kisha kadi B ni sawa na kuchagua kadi B kisha kadi A.

Kwa mfano, ukipata `52`, rudisha `1`. Kuna mchanganyiko mmoja tu wa kadi 52 kuchagua kutoka kwenye kifurushi cha kadi 52. Na ukipata `2`, rudisha `1326`. Kuna mchanganyiko 1326 ya kadi unayoweza kupata unapochagua kadi 2 kutoka kwenye kifurushi.

# --hints--

`combinations(52)` inapaswa kurudisha `1`.

```js
assert.equal(combinations(52), 1);
```

`combinations(1)` inapaswa kurudisha `52`.

```js
assert.equal(combinations(1), 52);
```

`combinations(2)` inapaswa kurudisha `1326`.

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)` inapaswa kurudisha `2598960`.

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)` inapaswa kurudisha `15820024220`.

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)` inapaswa kurudisha `1326`.

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```
