---
id: 68ee9e3066cfd4eb2328e8a5
title: "Zoezi la 86: Utafutaji wa Picha"
challengeType: 28
dashedName: challenge-86
---

# --description--

Mnamo tarehe 4 Novemba, 2001, Google ilizindua utafutaji wake wa picha, ukiruhusu watu kupata picha kwa kutumia maneno ya utafutaji. Katika changamoto hii, utaiga utafutaji wa picha.

Ukitoa safu ya data ya majina ya picha na neno la utafutaji, rudisha safu ya data ya majina ya picha zinazojumuisha neno la utafutaji.

- Puuza tofauti za herufi kubwa na ndogo wakati wa kulinganisha maneno ya utafutaji.
- Rudisha picha kwa mpangilio ule ule zinavyoonekana katika safu ya data ya ingizo.

# --hints--

`imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` inapaswa kurudisha `["dog.png"]`.

```js
assert.deepEqual(imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"]);
```

`imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` inapaswa kurudisha `["Sunset.jpg", "sunflower.jpeg"]`.

```js
assert.deepEqual(imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"]);
```

`imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` inapaswa kurudisha `["Moon.png", "stars.png"]`.

```js
assert.deepEqual(imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"]);
```

`imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` inapaswa kurudisha `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`.

```js
assert.deepEqual(imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]);
```

# --seed--

## --seed-contents--

```js
function imageSearch(images, term) {

  return images;
}
```

# --solutions--

```js
function imageSearch(images, term) {
  const lowerTerm = term.toLowerCase();
  return images.filter(img => img.toLowerCase().includes(lowerTerm));
}
```
