---
id: 68ee9e3066cfd4eb2328e8a7
title: "Zoezi la 88: Mtafuta wa siku za wiki"
challengeType: 28
dashedName: challenge-88
---

# --description--

Ikiwa umepewa mfuatano wa herufi wa tarehe katika muundo `YYYY-MM-DD`, rudisha siku ya wiki.

Siku halali za kurudisha ni:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Hakikisha usizingatie maeneo ya saa.

# --hints--

`getWeekday("2025-11-06")` inapaswa kurudisha `Thursday`.

```js
assert.equal(getWeekday("2025-11-06"), "Thursday");
```

`getWeekday("1999-12-31")` inapaswa kurudisha `Friday`.

```js
assert.equal(getWeekday("1999-12-31"), "Friday");
```

`getWeekday("1111-11-11")` inapaswa kurudisha `Saturday`.

```js
assert.equal(getWeekday("1111-11-11"), "Saturday");
```

`getWeekday("2112-12-21")` inapaswa kurudisha `Wednesday`.

```js
assert.equal(getWeekday("2112-12-21"), "Wednesday");
```

`getWeekday("2345-10-01")` inapaswa kurudisha `Monday`.

```js
assert.equal(getWeekday("2345-10-01"), "Monday");
```

# --seed--

## --seed-contents--

```js
function getWeekday(dateString) {

  return dateString;
}
```

# --solutions--

```js
function getWeekday(dateString) {
  const days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"];
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  return days[date.getUTCDay()];
}
```
