---
id: 68f6587287ad1f4ad39b0c7d
title: "Zoezi la 91: Tafuta neno"
challengeType: 28
dashedName: challenge-91
---

# --description--

Ikiwa umepewa matriki (safu ya safu za herufi moja moja) na neno la kutafuta, rudisha nambari za kuanzia na kumalizika kwa neno hilo katika matriki.

- Matriki iliyotolewa itajazwa na herufi ndogo zote (`a-z`).
- Neno la kutafuta litaonekana mara moja tu katika matriki.
- Neno la kutafuta litaonekana katika mstari wa moja kwa moja katika mwelekeo mmoja kati ya haya:
  - kutoka kushoto kwenda kulia
  - kutoka kulia kwenda kushoto
  - kutoka juu kwenda chini
  - kutoka chini kwenda juu

Kwa mfano, ukipata matriki:

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

Na neno `"cat"`, rudisha:

```md
[[0, 1], [2, 1]]
```

Ambapo `[0, 1]` ni nambari za kuanzia kwa `"c"` (mwanzoni mwa neno), na `[2, 1]` ni nambari za mwisho kwa `"t"` (mwishoni mwa neno).

# --hints--

`findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")` inapaswa kurudisha `[[0, 1], [2, 1]]`.

```js
assert.deepEqual(findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]]);
```

`findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")` inapaswa kurudisha `[[0, 0], [0, 2]]`.

```js
assert.deepEqual(findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]]);
```

`findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")` inapaswa kurudisha `[[3, 3], [0, 3]]`.

```js
assert.deepEqual(findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]]);
```

`findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")` inapaswa kurudisha `[[1, 3], [1, 1]]`.

```js
assert.deepEqual(findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]]);
```

# --seed--

## --seed-contents--

```js
function findWord(matrix, word) {

  return matrix;
}
```

# --solutions--

```js
function findWord(matrix, word) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const len = word.length;

  const directions = [
    [0, 1],
    [0, -1],
    [1, 0],
    [-1, 0]
  ];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      for (let [dr, dc] of directions) {
        let match = true;
        for (let i = 0; i < len; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols || matrix[nr][nc] !== word[i]) {
            match = false;
            break;
          }
        }
        if (match) {
          return [
            [r, c],
            [r + dr * (len - 1), c + dc * (len - 1)]
          ];
        }
      }
    }
  }
}
```
