---
id: 68f6587287ad1f4ad39b0c80
title: "Zoezi la 94: Kizalishaji cha Sahihi ya Barua Pepe"
challengeType: 28
dashedName: challenge-94
---

# --description--

Kutolewa mfuatano wa herufi wa jina la mtu, cheo, na kampuni, rudisha sahihi ya barua pepe kama mfuatano mmoja wa herufi ukitumia kanuni zifuatazo:

- Jina linapaswa kuonekana kwanza, likifuatiwa na kiambishi awali kinachotegemea herufi ya kwanza ya jina. Kwa majina yanayoanza na (bila kujali herufi kubwa au ndogo):
  - `A-I`: Tumia `>>` kama kiambishi awali.
  - `J-R`: Tumia `--` kama kiambishi awali.
  - `S-Z`: Tumia `::` kama kiambishi awali.
- Koma na nafasi (`, `) inapaswa kufuatia jina.
- Cheo na kampuni vinapaswa kufuatia koma na nafasi, vikitenganishwa na `" at "` (ikiwa na nafasi pande zote).

Kwa mfano, ukitolewa `"Quinn Waverly"`, `"Founder and CEO"`, na `"TechCo"` rudisha `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")` inapaswa kurudisha `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")` inapaswa kurudisha `">>Alice Reed, Engineer at TechCo"`.

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")` inapaswa kurudisha `"::Tina Vaughn, Developer at example.com"`.

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")` inapaswa kurudisha `">>B. B., Product Tester at AcmeCorp"`.

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")` inapaswa kurudisha `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
