---
id: 68f6587287ad1f4ad39b0c81
title: "Zoezi la 95: Shift ya safu ya data"
challengeType: 28
dashedName: challenge-95
---

# --description--

Ikiwa umepewa safu ya data na nambari kamili inayowakilisha ni nafasi ngapi za kusogeza safu hiyo, rudisha safu ya data iliyosogezwa.

- Nambari kamili chanya husogeza safu ya data kushoto.
- Nambari kamili hasi husogeza safu ya data kulia.
- Kusogeza kunazunguka ndani ya safu ya data.

Kwa mfano, ukipata `[1, 2, 3]` na `1`, sogeza safu ya data mara 1 kushoto, ukirudisha `[2, 3, 1]`.

# --hints--

`shiftArray([1, 2, 3], 1)` inapaswa kurudisha `[2, 3, 1]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` inapaswa kurudisha `[3, 1, 2]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` inapaswa kurudisha `["charlie", "alpha", "bravo"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` inapaswa kurudisha `["bravo", "charlie", "alpha"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` inapaswa kurudisha `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
