---
id: 68f6587287ad1f4ad39b0c85
title: "Zoezi la 99: Jaribio la Alama za Vidole"
challengeType: 28
dashedName: challenge-99
---

# --description--

Kutolewa mfuatano wa herufi mbili zinazowakilisha alama za vidole, amua kama ni mechi kwa kutumia sheria zifuatazo:

- Kila alama ya kidole itakuwa na herufi ndogo tu (`a-z`).
- Alama mbili za vidole zinachukuliwa kuwa mechi ikiwa:
  - Zina urefu sawa.
  - Idadi ya herufi zinazotofautiana haizidi 10% ya urefu wa alama ya kidole.

# --hints--

`isMatch("helloworld", "helloworld")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMatch("helloworld", "helloworld"));
```

`isMatch("helloworld", "helloworlds")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMatch("helloworld", "helloworlds"));
```

`isMatch("helloworld", "jelloworld")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMatch("helloworld", "jelloworld"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog"));
```

`isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat"));
```

# --seed--

## --seed-contents--

```js
function isMatch(fingerprintA, fingerprintB) {

  return fingerprintA;
}
```

# --solutions--

```js
function isMatch(fingerprintA, fingerprintB) {
  if (fingerprintA.length !== fingerprintB.length) return false;

  const length = fingerprintA.length;
  let mismatches = 0;

  for (let i = 0; i < length; i++) {
    if (fingerprintA[i] !== fingerprintB[i]) {
      mismatches++;
      if (mismatches > length * 0.1) return false;
    }
  }

  return true;
}
```
