---
id: 68f6587287ad1f4ad39b0c83
title: "Zoezi la 97: GCD"
challengeType: 28
dashedName: challenge-97
---

# --description--

Ukizingatia nambari kamili chanya mbili, rudisha mgawanyiko wao mkubwa zaidi wa pamoja (GCD).

- GCD ya nambari kamili mbili ni nambari kubwa zaidi inayogawanya nambari zote mbili kwa usawa bila kuacha mabaki.

Kwa mfano, wagawanyiko wa `4` ni `1`, `2`, na `4`. Wagawanyiko wa `6` ni `1`, `2`, `3`, na `6`. Kwa hivyo ukizingatia `4` na `6`, rudisha `2`, nambari kubwa zaidi inayojitokeza katika seti zote mbili za wagawanyiko.

# --hints--

`gcd(4, 6)` inapaswa kurudisha `2`.

```js
assert.equal(gcd(4, 6), 2);
```

`gcd(20, 15)` inapaswa kurudisha `5`.

```js
assert.equal(gcd(20, 15), 5);
```

`gcd(13, 17)` inapaswa kurudisha `1`.

```js
assert.equal(gcd(13, 17), 1);
```

`gcd(654, 456)` inapaswa kurudisha `6`.

```js
assert.equal(gcd(654, 456), 6);
```

`gcd(3456, 4320)` inapaswa kurudisha `864`.

```js
assert.equal(gcd(3456, 4320), 864);
```

# --seed--

## --seed-contents--

```js
function gcd(x, y) {

  return x;
}
```

# --solutions--

```js
function gcd(x, y) {
  while (y !== 0) {
    const temp = y;
    y = x % y;
    x = temp;
  }
  return x;
}
```
