---
id: 68f6587287ad1f4ad39b0c7c
title: "Zoezi la 90: Kiwango cha herufi"
challengeType: 28
dashedName: challenge-90
---

# --description--

Katika changamoto hii, unapewa mfuatano wa herufi na unahitaji kubaini kama unaendana na chapisho la mitandao ya kijamii. Rudisha mfuatano wa herufi ufuatao kulingana na sheria zilizotolewa:

- `"short post"` ikiwa unaendana na kiwango cha herufi 40.
- `"long post"` ikiwa ni zaidi ya herufi 40 na unaendana na kiwango cha herufi 80.
- `"invalid post"` ikiwa ni mrefu mno kuendana na viwango vyovyote viwili.

# --hints--

`canPost("Hello world")` inapaswa kurudisha `"short post"`.

```js
assert.equal(canPost("Hello world"), "short post");
```

`canPost("This is a longer message but still under eighty characters.")` inapaswa kurudisha `"long post"`.

```js
assert.equal(canPost("This is a longer message but still under eighty characters."), "long post");
```

`canPost("This message is too long to fit into either of the character limits for a social media post.")` inapaswa kurudisha `"invalid post"`.

```js
assert.equal(canPost("This message is too long to fit into either of the character limits for a social media post."), "invalid post");
```

# --seed--

## --seed-contents--

```js
function canPost(message) {

  return message;
}
```

# --solutions--

```js
function canPost(message) {
  if (message.length <= 40) {
    return "short post"
  } else if (message.length <= 80) {
    return "long post"
  } else {
    return "invalid post"
  }
}
```
