---
id: 68ffb91507a5b645769328c9
title: "Zoezi la 106: Mthibitishaji wa ujumbe"
challengeType: 28
dashedName: challenge-106
---

# --description--

Ukipata mfuatano wa herufi wa ujumbe na mfuatano wa herufi wa uthibitisho, amua kama ujumbe ni halali.

- Ujumbe ni halali ikiwa kila neno katika ujumbe huanza na herufi inayolingana katika mfuatano wa herufi wa uthibitisho, kwa mpangilio.
- Herufi hazizingatii tofauti ya herufi kubwa na ndogo.
- Maneno katika ujumbe yamegawanyika kwa nafasi moja moja.

# --hints--

`isValidMessage("hello world", "hw")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
