---
id: 68ffb91507a5b645769328c6
title: "Zoezi la 103: LCM"
challengeType: 28
dashedName: challenge-103
---

# --description--

Ukizingatia nambari kamili mbili, rudisha mzalishaji mdogo wa kawaida (LCM) wa nambari hizo mbili.

LCM ya nambari mbili ni nambari kamili chanya ndogo kabisa ambayo ni mzalishaji wa nambari zote mbili. Kwa mfano, ukizingatia `4` na `6`, rudisha `12` kwa sababu:

- Mzalishaji wa `4` ni `4`, `8`, `12` na kadhalika.
- Mzalishaji wa `6` ni `6`, `12`, `18` na kadhalika.
- `12` ni nambari ndogo kabisa ambayo ni mzalishaji wa zote mbili.

# --hints--

`lcm(4, 6)` inapaswa kurudisha `12`.

```js
assert.equal(lcm(4, 6), 12);
```

`lcm(9, 6)` inapaswa kurudisha `18`.

```js
assert.equal(lcm(9, 6), 18);
```

`lcm(10, 100)` inapaswa kurudisha `100`.

```js
assert.equal(lcm(10, 100), 100);
```

`lcm(13, 17)` inapaswa kurudisha `221`.

```js
assert.equal(lcm(13, 17), 221);
```

`lcm(45, 70)` inapaswa kurudisha `630`.

```js
assert.equal(lcm(45, 70), 630);
```

# --seed--

## --seed-contents--

```js
function lcm(a, b) {

  return a;
}
```

# --solutions--

```js
function lcm(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  return Math.abs(a * b) / gcd(a, b);
}
```
