---
id: 68ffb91507a5b645769328c5
title: "Zoezi la 102: Neno refu zaidi"
challengeType: 28
dashedName: challenge-102
---

# --description--

Ukitoa mfuatano wa herufi wa sentensi, rudisha neno refu zaidi katika sentensi hiyo.

- Maneno yamegawanyika kwa nafasi moja tu.
- Herufi pekee (`a-z`, isiyoegemea herufi kubwa au ndogo) ndizo zinahesabiwa kwa urefu wa neno.
- Ikiwa kuna maneno mengi yenye urefu sawa, rudisha lile la kwanza lililoonekana.
- Rudisha neno kama lilivyo katika mfuatano uliotolewa, bila alama za uandishi.

# --hints--

`longestWord("The quick red fox")` inapaswa kurudisha `"quick"`.

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` inapaswa kurudisha `"challenge"`.

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` inapaswa kurudisha `"This"`.

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` inapaswa kurudisha `"exclamation"`.

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` inapaswa kurudisha `"tie"`.

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` inapaswa kurudisha `"Wouldnt"`.

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
