---
id: 68ffb91507a5b645769328c4
title: "Zoezi la 101: Mbadalishaji wa kichwa cha sehemu cha Markdown"
challengeType: 28
dashedName: challenge-101
---

# --description--

Ikiwa umepewa mfuatano wa herufi unaowakilisha kichwa cha sehemu cha Markdown, rudisha kichwa cha sehemu kinacholingana cha HTML.

Kichwa cha sehemu halali cha Markdown lazima:

- Kianze na nafasi sifuri au zaidi, ikifuatiwa na
- Herufi nambari 1 hadi 6 za hash (`#`) mfululizo, kisha
- Angalau nafasi moja. Na hatimaye,
- Maandishi ya kichwa cha sehemu.

Idadi ya alama za hash huamua kiwango cha kichwa cha sehemu. Kwa mfano, alama moja ya hash inalingana na lebo ya `h1`, na alama sita za hash zinalingana na lebo ya `h6`.

Ikiwa mfuatano wa herufi uliotolewa hauendani kabisa na muundo huo hapo juu, rudisha `"Invalid format"`.

Kwa mfano, ukipokea `"# My level 1 heading"`, rudisha `"<h1>My level 1 heading</h1>"`.

Kumbuka: Konsoli inaweza isionyeshe lebo za HTML ndani ya mfuatano wa herufi wakati wa kurekodi ujumbe. Angalia konsoli ya kivinjari kuona rekodi zenye lebo zilizoingizwa.

# --hints--

`convert("# My level 1 heading")` inapaswa kurudisha `"<h1>My level 1 heading</h1>"`.

```js
assert.equal(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>");
```

`convert("My heading")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convert("My heading"), "Invalid format");
```

`convert("##### My level 5 heading")` inapaswa kurudisha `"<h5>My level 5 heading</h5>"`.

```js
assert.equal(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>");
```

`convert("#My heading")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convert("#My heading"), "Invalid format");
```

`convert("  ###  My level 3 heading")` inapaswa kurudisha `"<h3>My level 3 heading</h3>"`.

```js
assert.equal(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>");
```

`convert("####### My level 7 heading")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convert("####### My level 7 heading"), "Invalid format");
```

`convert("## My #2 heading")` inapaswa kurudisha `"<h2>My #2 heading</h2>"`.

```js
assert.equal(convert("## My #2 heading"), "<h2>My #2 heading</h2>");
```

# --seed--

## --seed-contents--

```js
function convert(heading) {

  return heading;
}
```

# --solutions--

```js
function convert(heading) {
  const trimmed = heading.trim()

  if (!/^#{1,6}\s+\S+/.test(trimmed)) return "Invalid format"

  const hashCount = trimmed.indexOf(' ');
  const text = trimmed.slice(hashCount).trimStart();

  return `<h${hashCount}>${text}</h${hashCount}>`;
}
```
