---
id: 68ffb91507a5b645769328c3
title: "Zoezi la 100: herufi 100"
challengeType: 28
dashedName: challenge-100
---

# --description--

Karibu kwenye changamoto ya 100 ya kila siku ya uandishi wa programu!

Ukipata mfuatano wa herufi, rudia herufi zake hadi matokeo yawe na herufi 100 kamili. Ikiwa rudufu zako zitazidi herufi 100, kata ziada ili ziwe herufi 100 tu.

# --hints--

`oneHundred("One hundred ")` inapaswa kurudisha `"One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One "`.

```js
assert.equal(oneHundred("One hundred "), "One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One ");
```

`oneHundred("freeCodeCamp ")` inapaswa kurudisha `"freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC"`.

```js
assert.equal(oneHundred("freeCodeCamp "), "freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC");
```

`oneHundred("daily challenges ")` inapaswa kurudisha `"daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge"`.

```js
assert.equal(oneHundred("daily challenges "), "daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge");
```

`oneHundred("!")` inapaswa kurudisha `"!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!"`.

```js
assert.equal(oneHundred("!"), "!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!");
```

# --seed--

## --seed-contents--

```js
function oneHundred(chars) {

  return chars;
}
```

# --solutions--

```js
function oneHundred(chars) {
  let result = "";
  let i = 0;

  while (result.length < 100) {
    result += chars[i % chars.length];
    i++;
  }

  return result.slice(0, 100);
}
```
