---
id: 68ffb91507a5b645769328cc
title: "Zoezi la 109: Nina umri gani tena?"
challengeType: 28
dashedName: challenge-109
---

# --description--

Ukizingatia tarehe ya siku ya kuzaliwa ya mtu kwa muundo `YYYY-MM-DD`, rudisha umri wa mtu huyo kufikia tarehe 27 Novemba, 2025.

- Kubali kuwa siku zote za kuzaliwa ni tarehe halali kabla ya 27 Novemba, 2025.
- Rudisha umri kama nambari kamili.
- Hakikisha unazingatia kama mtu huyo tayari ameshafikia siku yake ya kuzaliwa mwaka 2025.

# --hints--

`calculateAge("2000-11-20")` inapaswa kurudisha `25`.

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` inapaswa kurudisha `24`.

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` inapaswa kurudisha `11`.

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` inapaswa kurudisha `31`.

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` inapaswa kurudisha `30`.

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
