---
id: 69162d64f96574d9bb629efc
title: "Zoezi la 111: Mwelekeo wa Mpira"
challengeType: 28
dashedName: challenge-111
---

# --description--

Changamoto ya leo imetokana na mchezo wa video wa Pong, uliozinduliwa tarehe 29 Novemba, 1972.

Ukizingatia matriki (safu ya data ya safu za data) inayojumuisha mahali pa mpira (`2`), na mahali pa mpira kabla yake (`1`), rudisha nambari za matriki za mahali pa mpira litakalofuata.

- Mpira kila mara husogea kwa mstari wa moja kwa moja.
- Mwelekeo wa mwendo unategemea jinsi mpira ulivyohamia kutoka `1` hadi `2`.
- Mipaka ya matriki inachukuliwa kuwa kuta. Ikiwa mpira unagonga:
  - kuta ya juu au chini, unaruka kwa kubadilisha mwelekeo wake wima.
  - kuta ya kushoto au kulia, unaruka kwa kubadilisha mwelekeo wake wima.
  - kona, unaruka kwa kubadilisha mwelekeo yote mawili.

# --hints--

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]])` inapaswa kurudisha `[2, 3]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]]), [2, 3]);
```

`getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]])` inapaswa kurudisha `[3, 0]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]]), [3, 0]);
```

`getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]])` inapaswa kurudisha `[1, 2]`.

```js
assert.deepEqual(getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]]), [1, 2]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]])` inapaswa kurudisha `[1, 1]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]]), [1, 1]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]])` inapaswa kurudisha `[2, 2]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]]), [2, 2]);
```

# --seed--

## --seed-contents--

```js
function getNextLocation(matrix) {

  return matrix;
}
```

# --solutions--

```js
function getNextLocation(matrix) {
  let prev = null;
  let curr = null;

  for (let row = 0; row < matrix.length; row++) {
    for (let col = 0; col < matrix[0].length; col++) {
      if (matrix[row][col] === 1) prev = [row, col];
      if (matrix[row][col] === 2) curr = [row, col];
    }
  }

  const [prevRow, prevCol] = prev;
  const [currRow, currCol] = curr;

  let dirX = currCol - prevCol;
  let dirY = currRow - prevRow;

  let nextRow = currRow + dirY;
  let nextCol = currCol + dirX;

  const maxRow = matrix.length - 1;
  const maxCol = matrix[0].length - 1;


  if (nextCol < 0 || nextCol > maxCol) dirX *= -1;
  if (nextRow < 0 || nextRow > maxRow) dirY *= -1;

  nextRow = currRow + dirY;
  nextCol = currCol + dirX;

  return [nextRow, nextCol];
}
```
