---
id: 69162d64f96574d9bb629efd
title: "Zoezi la 112: Kigeuzi cha AI"
challengeType: 28
dashedName: challenge-112
---

# --description--

Changamoto ya leo imetokana na uzinduzi wa ChatGPT tarehe 30 Novemba 2022.

Ukizingatia mfuatano wa herufi wa sentensi moja au zaidi, amua kama huenda ulitengenezwa na AI kwa kutumia sheria zifuatazo:

- Una dash mbili au zaidi (`-`).
- Una seti mbili au zaidi za mabano (`()`). Maandishi yanaweza kuwa ndani ya mabano.
- Una maneno matatu au zaidi yenye herufi 7 au zaidi.

- Maneno yamegawanyika kwa nafasi moja tu na yanajumuisha herufi pekee (`A-Z`). Usijumuishe alama za uandishi au herufi zisizo za herufi kama sehemu ya neno.

Ikiwa sentensi iliyotolewa inakidhi sheria yoyote kati ya hizi, rudisha `"AI"`, vinginevyo, rudisha `"Human"`.

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` inapaswa kurudisha `"Human"`.

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` inapaswa kurudisha `"Human"`.

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` inapaswa kurudisha `"AI"`.

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` inapaswa kurudisha `"AI"`.

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` inapaswa kurudisha `"AI"`.

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
