---
id: 69162d64f96574d9bb629efb
title: "Changamoto 110: Mtabiri wa Neno"
challengeType: 28
dashedName: challenge-110
---

# --description--

Ukizingatia mfuatano wa herufi mbili zenye urefu sawa, neno la siri na neno la kubahatisha, linganisha neno la kubahatisha na neno la siri kwa kutumia sheria zifuatazo:

- Neno la siri na neno la kubahatisha litakuwa na herufi kubwa tu (`"A"` hadi `"Z"`);
- Kwa kila herufi katika neno la kubahatisha, ibadilishe na nambari kulingana na jinsi inavyolingana na neno la siri:
  - `"2"` ikiwa herufi iko katika neno la siri na ipo mahali sahihi.
  - `"1"` ikiwa herufi iko katika neno la siri lakini ipo mahali potofu.
  - `"0"` ikiwa herufi haipo katika neno la siri.
- Kila herufi katika neno la siri inaweza kutumika mara moja tu.
- Mlingano kamili (`"2"`) unapelekwa kwanza, kisha mlingano wa sehemu (`"1"`) unapelekwa kutoka kushoto kwenda kulia kwa herufi zilizobaki.
- Ikiwa herufi inatokea mara nyingi katika neno la kubahatisha, inaweza kulingana mara ngapi tu inavyotokea katika neno la siri.

Kwa mfano, ukizingatia neno la siri la `"APPLE"` na neno la kubahatisha la `"POPPA"`, rudisha `"10201"`:

Herufi ya kwanza ya `"P"` haipo mahali sahihi (`"1"`), `"O"` haipo katika neno la siri (`"0"`), herufi ya pili ya `"P"` iko mahali sahihi (`"2"`), herufi ya tatu ya `"P"` ni sifuri (`"0"`) kwa sababu herufi mbili za `"P"` katika neno la siri zimetumika, na `"A"` haipo mahali sahihi (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` inapaswa kurudisha `"10201"`.

```js
assert.equal(compare("APPLE", "POPPA"), "10201");
```

`compare("REACT", "TRACE")` inapaswa kurudisha `"11221"`.

```js
assert.equal(compare("REACT", "TRACE"), "11221");
```

`compare("DEBUGS", "PYTHON")` inapaswa kurudisha `"000000"`.

```js
assert.equal(compare("DEBUGS", "PYTHON"), "000000");
```

`compare("JAVASCRIPT", "TYPESCRIPT")` inapaswa kurudisha `"0000222222"`.

```js
assert.equal(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222");
```

`compare("ORANGE", "ROUNDS")` inapaswa kurudisha `"110200"`.

```js
assert.equal(compare("ORANGE", "ROUNDS"), "110200");
```

`compare("WIRELESS", "ETHERNET")` inapaswa kurudisha `"10021000"`.

```js
assert.equal(compare("WIRELESS", "ETHERNET"), "10021000");
```

# --seed--

## --seed-contents--

```js
function compare(word, guess) {

  return word;
}
```

# --solutions--

```js
function compare(word, guess) {
  const n = word.length;
  let result = Array(n).fill("0");
  let used = Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    if (guess[i] === word[i]) {
      result[i] = "2";
      used[i] = true;
    }
  }

  for (let i = 0; i < n; i++) {
    if (result[i] === "0") {
      for (let j = 0; j < n; j++) {
        if (!used[j] && guess[i] === word[j]) {
          result[i] = "1";
          used[j] = true;
          break;
        }
      }
    }
  }

  return result.join("");
}
```
