---
id: 69162d64f96574d9bb629eff
title: "Zoezi la 114: Kutoka camel hadi snake"
challengeType: 28
dashedName: challenge-114
---

# --description--

Ikiwa umepewa mfuatano wa herufi katika camel case, rudisha toleo la snake case la mfuatano huo kwa kutumia sheria zifuatazo:

- Mfuatano wa herufi wa ingizo utakuwa na herufi tu (`A-Z` na `a-z`) na daima utaanza na herufi ndogo.
- Kila herufi kubwa katika mfuatano wa camel case huanzisha neno jipya.
- Badilisha herufi zote kuwa herufi ndogo.
- Tenganisha maneno kwa mstari chini (underscore) (`_`).

# --hints--

`toSnake("helloWorld")` inapaswa kurudisha `"hello_world"`.

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` inapaswa kurudisha `"my_variable_name"`.

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` inapaswa kurudisha `"freecodecamp_daily_challenges"`.

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
