---
id: 69162d64f96574d9bb629f00
title: "Zoezi la 115: Mbadalishaji wa Kipengee cha Orodha Iliyo Pangwa ya Markdown"
challengeType: 28
dashedName: challenge-115
---

# --description--

Ikiwa umepewa mfuatano wa herufi unaowakilisha kipengee cha orodha iliyopangwa katika Markdown, rudisha mfuatano wa herufi wa HTML unaolingana.

Kipengee halali cha orodha iliyopangwa katika Markdown lazima:

- Kuanza na nafasi sifuri au zaidi, ikifuatiwa na
- Nambari (1 au zaidi) na nukta (`.`), ikifuatiwa na
- Angalau nafasi moja, kisha
- Maandishi ya kipengee cha orodha.

Kama mfuatano wa herufi hauendani kabisa na muundo huo hapo juu, rudisha `"Invalid format"`. Vinginevyo, zungusha maandishi ya kipengee cha orodha ndani ya lebo za `li` na rudisha mfuatano wa herufi.

Kwa mfano, ukipokea `"1. My item"`, rudisha `"<li>My item</li>"`.

Kumbuka: Konsoli huenda isionyeshe lebo za HTML katika mfuatano wa herufi wakati wa kurekodi ujumbe. Angalia konsoli ya kivinjari kuona rekodi zenye lebo zikiwemo.

# --hints--

`convertListItem("1. My item")` inapaswa kurudisha `"<li>My item</li>"`.

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` inapaswa kurudisha `"<li>Another item</li>"`.

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` inapaswa kurudisha `"<li>list item text</li>"`.

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
