---
id: 69162d64f96574d9bb629f02
title: "Zoezi la 117: Tofauti ya Usawa"
challengeType: 28
dashedName: challenge-117
---

# --description--

Ukitoa safu mbili za data, rudisha safu mpya inayojumuisha tofauti ya usawa kati yao.

- Tofauti ya usawa kati ya seti mbili ni seti ya thamani zinazojitokeza katika seti moja au nyingine, lakini si zote mbili.
- Rudisha thamani kwa mpangilio walipojitokeza kwa mara ya kwanza katika safu za data za ingizo.

# --hints--

`difference([1, 2, 3], [3, 4, 5])` inapaswa kurudisha `[1, 2, 4, 5]`.

```js
assert.deepEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5]);
```

`difference(["a", "b"], ["c", "b"])` inapaswa kurudisha `["a", "c"]`.

```js
assert.deepEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"]);
```

`difference([1, "a", 2], [2, "b", "a"])` inapaswa kurudisha `[1, "b"]`.

```js
assert.deepEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"]);
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` inapaswa kurudisha `[2, 4, 6, 8]`.

```js
assert.deepEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8]);
```

# --seed--

## --seed-contents--

```js
function difference(arr1, arr2) {

  return arr1;
}
```

# --solutions--

```js
function difference(arr1, arr2) {
  const diff = [];

  for (const v of arr1) {

    if (!diff.includes(v) && !arr2.includes(v)) diff.push(v)
  }
  
  for (const v of arr2) {
    if (!diff.includes(v) && !arr1.includes(v)) diff.push(v)
  }

  return diff;
}
```
