---
id: 69162d64f96574d9bb629f01
title: "Zoezi la 116: Hesabu ya permutation"
challengeType: 28
dashedName: challenge-116
---

# --description--

Ikiwa umepewa mfuatano wa herufi, rudisha idadi ya permutation tofauti zinazoweza kuundwa kutoka kwa herufi zake.

- Permutation ni mpangilio upya wowote wa herufi katika mfuatano.
- Usihesabu permutation zinazojirudia.
- Ikiwa mfuatano una herufi zinazojirudia, mipangilio inayojirudia inapaswa kuhesabiwa mara moja tu.
- Mfuatano utakuwa na herufi tu (`A-Z`, `a-z`).

Kwa mfano, ukipata `"abb"`, rudisha `3` kwa sababu kuna njia tatu za kipekee za kupanga herufi: `"abb"`, `"bab"`, na `"bba"`.

# --hints--

`countPermutations("abb")` inapaswa kurudisha `3`.

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")` inapaswa kurudisha `6`.

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")` inapaswa kurudisha `630`.

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")` inapaswa kurudisha `39916800`.

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
