---
id: 691b559495c5cb5a37b9b482
title: "Zoezi la 122: Parser ya herufi nzito ya Markdown"
challengeType: 28
dashedName: challenge-122
---

# --description--

Ikiwa umepewa mfuatano wa herufi ambao unaweza kuwa na maandishi ya herufi nzito katika Markdown, rudisha mfuatano wa herufi sawa wa HTML.

- Maandishi ya herufi nzito katika Markdown ni maandishi yoyote yanayoanza na kumalizika na nyota mbili (`**`) au mstari chini miwili (`__`).
- Hakuna nafasi yoyote kati ya maandishi na nyota au mistari chini, lakini nafasi zinaweza kuwepo ndani ya maandishi yenyewe.
- Badilisha matukio yote ya herufi nzito kuwa lebo za HTML `b` na rudisha mfuatano wa herufi.

Kwa mfano, ukipokea `"**This is bold**"`, rudisha `"<b>This is bold</b>"`.

Kumbuka: Konsoli inaweza isionyeshe lebo za HTML katika mfuatano wa herufi wakati wa kurekodi ujumbe. Angalia konsoli ya kivinjari kuona rekodi zenye lebo zikiambatanishwa.

# --hints--

`parseBold("**This is bold**")` inapaswa kurudisha `"<b>This is bold</b>"`.

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")` inapaswa kurudisha `"<b>This is also bold</b>"`.

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")` inapaswa kurudisha `"**This is not bold **"`.

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")` inapaswa kurudisha `"__ This is also not bold__"`.

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")` inapaswa kurudisha `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```
