---
id: 691b559495c5cb5a37b9b488
title: "Zoezi la 128: Hesabu ya Konsonanti"
challengeType: 28
dashedName: challenge-128
---

# --description--

Ukizingatia mfuatano wa herufi na nambari lengwa, amua kama mfuatano huo una idadi halisi ya konsonanti kama nambari lengwa.

- Konsonanti ni herufi zote za alfabeti isipokuwa `"a"`, `"e"`, `"i"`, `"o"`, na `"u"` katika hali yoyote ya herufi.
- Puuza nambari, alama za uandishi, nafasi, na herufi nyingine zisizo za herufi wakati wa kuhesabu.

# --hints--

`hasConsonantCount("helloworld", 7)` inapaswa kurudisha `true`.

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)` inapaswa kurudisha `false`.

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)` inapaswa kurudisha `true`.

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` inapaswa kurudisha `false`.

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` inapaswa kurudisha `true`.

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
