---
id: 691b559495c5cb5a37b9b486
title: "Zoezi la 126: Andika kwa herufi kubwa"
challengeType: 28
dashedName: challenge-126
---

# --description--

Ukipata mfuatano wa herufi unaoitwa title, rudisha mfuatano mpya uliopangwa kwa mtindo wa herufi kubwa kwa kutumia sheria zifuatazo:

- Andika herufi ya kwanza ya kila neno kwa herufi kubwa.
- Fanya herufi zote nyingine katika kila neno kuwa herufi ndogo.
- Maneno hubaki kutenganishwa na nafasi moja tu kila mara.

# --hints--

`titleCase("hello world")` inapaswa kurudisha `"Hello World"`.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` inapaswa kurudisha `"The Quick Brown Fox"`.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` inapaswa kurudisha `"Javascript And Python"`.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` inapaswa kurudisha `"Avocado Toast For Breakfast"`.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
