---
id: 691b559495c5cb5a37b9b485
title: "Changamoto 125: Mchezo wa Maisha"
challengeType: 28
dashedName: challenge-125
---

# --description--

Ikiwa umepewa matriki (safu ya safu) inayowakilisha hali ya sasa katika Mchezo wa Maisha wa Conway, rudisha hali inayofuata ya matriki kwa kutumia sheria hizi:

- Kila seli ni `1` (hai) au `0` (kifo).
- Majirani wa seli ni seli hadi nane zinazozunguka (wima, mlalo, na mwinuko).
- Seli zilizokaribu na mipaka zina majirani wachache kuliko nane.

Sheria za kusasisha kila seli:

- Seli yoyote hai yenye majirani hai chini ya wawili hufa (ukosefu wa watu).
- Seli yoyote hai yenye majirani hai wawili au watatu huendelea kuishi.
- Seli yoyote hai yenye majirani hai zaidi ya watatu hufa (watu wengi sana).
- Seli yoyote iliyokufa yenye majirani hai watatu hasa huamka (uzazi).

Kwa mfano, ikitolewa:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

rudisha:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Kila seli husasishwa kulingana na idadi ya majirani hai. Kwa mfano, `[0][0]` hubaki amekufa (majirani hai 2), `[0][1]` hubaki hai (majirani hai 2), `[0][2]` hufa (majirani hai 3), na kadhalika.

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` inapaswa kurudisha `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` inapaswa kurudisha `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` inapaswa kurudisha `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` inapaswa kurudisha `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
