---
id: 691b5597f0f3e85a588a5c94
title: "Zoezi la 130: Bodi ya mcheza-damu"
challengeType: 28
dashedName: challenge-130
---

# --description--

Ikiwa umepewa safu ya data yenye nambari mbili, ya kwanza ikiwa ni idadi ya safu na ya pili ikiwa ni idadi ya nguzo, rudisha matriki (safu ya safu za data) iliyojaa herufi za `"X"` na `"O"` za ukubwa uliotolewa.

- Herufi hizo zinapaswa kubadilishana kama bodi ya mcheza-damu.
- Seli ya juu-kushoto lazima iwe daima `"X"`.

Kwa mfano, ukipata `[3, 3]`, rudisha:

```json
[
  ["X", "O", "X"],
  ["O", "X", "O"],
  ["X", "O", "X"]
]
```

# --hints--

`createBoard([3, 3])` inapaswa kurudisha `[["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]`.

```js
assert.deepEqual(createBoard([3, 3]), [["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]);
```

`createBoard([6, 1])` inapaswa kurudisha `[["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]`.

```js
assert.deepEqual(createBoard([6, 1]), [["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]);
```

`createBoard([2, 10])` inapaswa kurudisha `[["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]`.

```js
assert.deepEqual(createBoard([2, 10]), [["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]);
```

`createBoard([5, 4])` inapaswa kurudisha `[["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]`.

```js
assert.deepEqual(createBoard([5, 4]), [["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]);
```

# --seed--

## --seed-contents--

```js
function createBoard(dimensions) {
  return dimensions;
}
```

# --solutions--

```js
function createBoard(dimensions) {
  const [rows, cols] = dimensions;
  const board = [];

  for (let r = 0; r < rows; r++) {
    const row = [];
    for (let c = 0; c < cols; c++) {
      row.push((r + c) % 2 === 0 ? "X" : "O");
    }
    board.push(row);
  }

  return board;
}
```
