---
id: 691b559495c5cb5a37b9b483
title: "Zoezi la 123: Mtengenezaji wa Nambari za Kirumi"
challengeType: 28
dashedName: challenge-123
---

# --description--

Ukipata nambari kamili, rudisha thamani yake sawa katika nambari za Kirumi.

Nambari za Kirumi hutumia alama hizi:

| Alama | Thamani |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

Nambari za Kirumi zinaandikwa kutoka kubwa kwenda ndogo, kutoka kushoto kwenda kulia kwa kutumia sheria zifuatazo:

- Kuongeza hutumika wakati alama inafuatiwa na nyingine yenye thamani sawa au ndogo. Kwa mfano, `18` inaandikwa kama `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- Kutoa hutumika wakati alama ndogo inaonekana kabla ya kubwa, kuwakilisha 4 au 9 katika thamani yoyote ya nafasi. Kwa mfano, 19 inaandikwa kama `XIX` (10 + (10 - 1)).
- Hakuna alama inayorudiwa zaidi ya mara tatu mfululizo. Kutoa hutumika pale unapotakiwa kuandika alama zaidi ya mara tatu mfululizo. Hivyo nambari kubwa zaidi unayoweza kuandika ni 3999.

Hapa kuna mfano mwingine: ukipata `1464`, rudisha `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`toRoman(18)` inapaswa kurudisha `"XVIII"`.

```js
assert.equal(toRoman(18), "XVIII");
```

`toRoman(19)` inapaswa kurudisha `"XIX"`.

```js
assert.equal(toRoman(19), "XIX");
```

`toRoman(1464)` inapaswa kurudisha `"MCDLXIV"`.

```js
assert.equal(toRoman(1464), "MCDLXIV");
```

`toRoman(2025)` inapaswa kurudisha `"MMXXV"`.

```js
assert.equal(toRoman(2025), "MMXXV");
```

`toRoman(3999)` inapaswa kurudisha `"MMMCMXCIX"`.

```js
assert.equal(toRoman(3999), "MMMCMXCIX");
```

# --seed--

## --seed-contents--

```js
function toRoman(num) {
  return num;
}
```

# --solutions--

```js
function toRoman(num) {
  const symbols = [
    ["M", 1000],
    ["CM", 900],
    ["D", 500],
    ["CD", 400],
    ["C", 100],
    ["XC", 90],
    ["L", 50],
    ["XL", 40],
    ["X", 10],
    ["IX", 9],
    ["V", 5],
    ["IV", 4],
    ["I", 1]
  ];

  let result = "";
  for (const [sym, val] of symbols) {
    while (num >= val) {
      result += sym;
      num -= val;
    }
  }
  return result;
}
```
