---
id: 6925e2068081f40f549ced1d
title: "Zoezi la 139: Jiwe, Karatasi, Wembe"
challengeType: 28
dashedName: challenge-139
---

# --description--

Kutolewa mfuatano wa herufi miwili, wa kwanza ukiwakilisha Mchezaji 1 na wa pili ukiwakilisha Mchezaji 2, amua mshindi wa mechi ya Jiwe, Karatasi, Wembe.

- Mfuatano wa herufi wa ingizo daima utakuwa `"Rock"`, `"Paper"`, au `"Scissors"`.
- `"Rock"` hushinda `"Scissors"`.
- `"Paper"` hushinda `"Rock"`.
- `"Scissors"` hushinda `"Paper"`.

Rudisha:

- `"Player 1 wins"` ikiwa Mchezaji 1 atashinda.
- `"Player 2 wins"` ikiwa Mchezaji 2 atashinda.
- `"Tie"` ikiwa wachezaji wote wawili watateua chaguo sawa.

# --hints--

`rockPaperScissors("Rock", "Rock")` inapaswa kurudisha `"Tie"`.

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")` inapaswa kurudisha `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")` inapaswa kurudisha `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")` inapaswa kurudisha `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")` inapaswa kurudisha `"Tie"`.

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")` inapaswa kurudisha `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
