---
id: 69272dcf1c24b44fd79137c5
title: "Zoezi la 142: Jumlisha mfuatano wa herufi"
challengeType: 28
dashedName: challenge-142
---

# --description--

Ikiwa umepewa mfuatano wa herufi unaohusisha nambari na herufi nyingine, rudisha jumla ya nambari zote zilizopo katika mfuatano huo.

- Tazama nambari zinazofuata moja baada ya nyingine kama nambari moja kamili. Kwa mfano, `"13"` inahesabiwa kama 13, si 1 + 3.
- Puuza herufi zozote ambazo si nambari.

# --hints--

`stringSum("3apples2bananas")` inapaswa kurudisha `5`.

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` inapaswa kurudisha `17`.

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` inapaswa kurudisha `125344`.

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` inapaswa kurudisha `321`.

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` inapaswa kurudisha `1653`.

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
