---
id: 69272dcf1c24b44fd79137c3
title: "Zoezi la 140: SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

Kutolewa mfuatano wa herufi unaowakilisha jina la kigezo, rudisha jina la kigezo lililobadilishwa kuwa SCREAMING_SNAKE_CASE.

Majina ya vigezo yaliyotolewa yataandikwa kwa mojawapo ya miundo ifuatayo:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

Katika miundo hapo juu, maneno yamegawanywa kwa mstari wa chini (`_`), mstari wa kati (`-`), au neno jipya linaanza kwa herufi kubwa.

Ili kubadilisha kuwa SCREAMING_SNAKE_CASE:

- Fanya herufi zote kuwa kubwa
- Gawanya maneno kwa mstari wa chini (`_`)

# --hints--

`toScreamingSnakeCase("userEmail")` inapaswa kurudisha `"USER_EMAIL"`.

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` inapaswa kurudisha `"USER_PASSWORD"`.

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` inapaswa kurudisha `"USER_ID"`.

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` inapaswa kurudisha `"USER_ADDRESS"`.

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` inapaswa kurudisha `"USERNAME"`.

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

`toScreamingSnakeCase("my_variable_name")` inapaswa kurudisha `"MY_VARIABLE_NAME"`.

```js
assert.equal(toScreamingSnakeCase("my_variable_name"), "MY_VARIABLE_NAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();
}
```
