---
id: 69272dcf1c24b44fd79137c6
title: "Zoezi la 143: Parser ya Markdown ya Kutilia Mkazo kwa Italiki"
challengeType: 28
dashedName: challenge-143
---

# --description--

Ikiwa umepewa mfuatano wa herufi ambao unaweza kuwa na maandishi yaliyo katika mtindo wa italiki katika Markdown, rudisha mfuatano wa herufi wa HTML unaolingana.

- Maandishi ya italiki katika Markdown ni maandishi yoyote yanayoanza na kumalizika na nyota moja (`*`) au mstari chini mmoja (`_`).
- Hakuna nafasi zinazoruhusiwa kati ya maandishi na nyota au mstari chini, lakini nafasi zinaweza kuwepo ndani ya maandishi yenyewe.
- Badilisha matukio yote ya italiki kuwa lebo za HTML `i` na rudisha mfuatano wa herufi.

Kwa mfano, ukipokea `"*This is italic*"`, rudisha `"<i>This is italic</i>"`.

Kumbuka: Konsoli inaweza isionyeshe lebo za HTML ndani ya mfuatano wa herufi wakati wa kurekodi ujumbe. Angalia konsoli ya kivinjari ili kuona rekodi zenye lebo zikiwemo.

# --hints--

`parseItalics("*This is italic*")` inapaswa kurudisha `"<i>This is italic</i>"`.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` inapaswa kurudisha `"<i>This is also italic</i>"`.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` inapaswa kurudisha `"*This is not italic *"`.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` inapaswa kurudisha `"_ This is also not italic_"`.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` inapaswa kurudisha `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
