---
id: 69272dcf1c24b44fd79137c4
title: "Zoezi la 141: Mafuta ya Kuondoka"
challengeType: 28
dashedName: challenge-141
---

# --description--

Kutokana na idadi ya galoni za mafuta zilizopo sasa katika ndege yako, na idadi inayohitajika ya lita za mafuta kufikia mahali unapotaka kwenda, amua ni galoni ngapi za mafuta za ziada unapaswa kuongeza.

- Galoni 1 ni sawa na lita 3.78541.
- Ikiwa ndege tayari ina mafuta ya kutosha, rudisha `0`.
- Unaweza kuongeza galoni kamili tu.
- Usijumlishe desimali katika nambari ya kurudisha.

# --hints--

`fuelToAdd(0, 1)` inapaswa kurudisha `1`.

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)` inapaswa kurudisha `6`.

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)` inapaswa kurudisha `0`.

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)` inapaswa kurudisha `4520`.

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)` inapaswa kurudisha `12209`.

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
