---
id: 69306364df283fcaff2e1ad8
title: "Zoezi la 147: Kihesabu cha Mwaka wa Kuruka"
challengeType: 28
dashedName: challenge-147
---

# --description--

Ikiwa umepata mwaka wa nambari kamili, amua kama ni mwaka wa kuruka.

Mwaka ni mwaka wa kuruka ikiwa unakidhi sheria zifuatazo:

- Mwaka unagawika kwa usawa na 4, na
- Mwaka haupagawiki kwa usawa na 100, isipokuwa
- Mwaka unagawika kwa usawa na 400.

# --hints--

`isLeapYear(2024)` inapaswa kurudisha `true`.

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)` inapaswa kurudisha `false`.

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)` inapaswa kurudisha `false`.

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)` inapaswa kurudisha `true`.

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)` inapaswa kurudisha `false`.

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)` inapaswa kurudisha `true`.

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)` inapaswa kurudisha `false`.

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
