---
id: 69306364df283fcaff2e1ad9
title: "Zoezi la 148: Shinikizo la Taya"
challengeType: 28
dashedName: challenge-148
---

# --description--

Imetolewa safu ya data yenye nambari nne zinazowakilisha shinikizo la taya kwa psi kwa matairi manne ya gari lako, na safu nyingine ya nambari mbili zinazowakilisha kiwango cha chini na cha juu cha shinikizo kwa matairi yako kwa bar, rudisha safu ya mfuatano wa herufi nne inayobainisha hali ya kila taya.

- 1 bar ni sawa na 14.5038 psi.

Rudisha safu ya data yenye thamani zifuatazo kwa kila taya:

- `"Low"` ikiwa shinikizo la taya ni chini ya kiwango cha chini kilichoruhusiwa.
- `"Good"` ikiwa iko kati ya kiwango cha chini na cha juu kilichoruhusiwa.
- `"High"` ikiwa iko juu ya kiwango cha juu kilichoruhusiwa.

# --hints--

`tireStatus([32, 28, 35, 29], [2, 3])` inapaswa kurudisha `["Good", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"]);
```

`tireStatus([32, 28, 35, 30], [2, 2.3])` inapaswa kurudisha `["Good", "Low", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"]);
```

`tireStatus([29, 26, 31, 28], [2.1, 2.5])` inapaswa kurudisha `["Low", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"]);
```

`tireStatus([31, 31, 30, 29], [1.5, 2])` inapaswa kurudisha `["High", "High", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"]);
```

`tireStatus([30, 28, 30, 29], [1.9, 2.1])` inapaswa kurudisha `["Good", "Good", "Good", "Good"]`.

```js
assert.deepEqual(tireStatus([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"]);
```

# --seed--

## --seed-contents--

```js
function tireStatus(pressuresPSI, rangeBar) {

  return pressuresPSI;
}
```

# --solutions--

```js
function tireStatus(pressuresPSI, rangeBar) {
  const psiToBar = 1 / 14.5038;
  const [minBar, maxBar] = rangeBar;

  return pressuresPSI.map(psi => {
    const bar = psi * psiToBar;
    if (bar < minBar) return "Low";
    if (bar > maxBar) return "High";
    return "Good";
  });
}
```
