---
id: 69306364df283fcaff2e1ad5
title: "Zoezi la 144: Mfululizo wa azimio"
challengeType: 28
dashedName: challenge-144
---

# --description--

Ukipata safu ya data za safu ndogo, ambapo kila safu ndogo inawakilisha siku ya shughuli zako za azimio na ina nambari kamili tatu: idadi ya hatua ulizotembea siku hiyo, dakika za kutumia skrini siku hiyo, na idadi ya kurasa ulizosoma siku hiyo; amua kama unazingatia azimio lako.

- Safu ndogo ya kwanza ni siku ya 1, na ya pili ni siku ya 2, na kadhalika.

Siku inachukuliwa kuwa yenye mafanikio ikiwa malengo matatu yafuatayo yamefikiwa:

- Umetembea angalau hatua 10,000.
- Hukupita dakika 120 za kutumia skrini.
- Umesoma angalau kurasa tano.

Kama siku zote zilizotolewa ni za mafanikio, rudisha `"Resolution on track: N day streak."` Ambapo `N` ni idadi ya siku za mafanikio.

Kama siku moja au zaidi hazijafanikiwa, rudisha `"Resolution failed on day X: N day streak."`. Ambapo `X` ni nambari ya siku ya kwanza isiyofanikiwa, na `N` ni idadi ya siku za mafanikio kabla ya siku isiyofanikiwa ya kwanza.

# --hints--

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` inapaswa kurudisha `"Resolution on track: 3 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.");
```

`resolutionStreak([[10000, 120, 5], [10950, 121, 11]])` inapaswa kurudisha `"Resolution failed on day 2: 1 day streak."`

```js
assert.equal(resolutionStreak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.");
```

`resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` inapaswa kurudisha `"Resolution failed on day 3: 2 day streak."`

```js
assert.equal(resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.");
```

`resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` inapaswa kurudisha `"Resolution failed on day 4: 3 day streak."`

```js
assert.equal(resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.");
```

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` inapaswa kurudisha `"Resolution on track: 10 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.");
```

# --seed--

## --seed-contents--

```js
function resolutionStreak(days) {

  return days;
}
```

# --solutions--

```js
function resolutionStreak(days) {
  let streak = 0;

  for (let i = 0; i < days.length; i++) {
    const [steps, screen, pages] = days[i];
    if (steps >= 10000 && screen <= 120 && pages >= 5) {
      streak++;
    } else {
      return `Resolution failed on day ${i + 1}: ${streak} day streak.`;
    }
  }

  return `Resolution on track: ${streak} day streak.`;
}
```
