---
id: 69306364df283fcaff2e1ad6
title: "Zoezi la 145: Nambari ya Nth ya Fibonacci"
challengeType: 28
dashedName: challenge-145
---

# --description--

Ikiwa umepewa nambari kamili `n`, rudisha nambari ya `n` katika mfululizo wa fibonacci.

Mfululizo wa Fibonacci ni mfululizo wa nambari ambapo kila nambari ni jumla ya nambari mbili zilizotangulia. Nambari kumi za kwanza katika mfululizo ni `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

# --hints--

`nthFibonacci(4)` inapaswa kurudisha `2`.

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)` inapaswa kurudisha `34`.

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)` inapaswa kurudisha `377`.

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)` inapaswa kurudisha `63245986`.

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)` inapaswa kurudisha `1304969544928657`.

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
