---
id: 69373793f5a867f769cde137
title: "Zoezi la 152: Namba msingi mduara"
challengeType: 28
dashedName: challenge-152
---

# --description--

Ikiwa namba kamili, amua kama ni namba msingi mduara.

Namba msingi mduara ni namba kamili ambapo mizunguko yote ya tarakimu zake ni nambari msingi pia.

Kwa mfano, `197` ni namba msingi mduara kwa sababu mizunguko yote ya tarakimu zake: `197`, `971`, na `719`, ni namba msingi.

# --hints--

`isCircularPrime(197)` inapaswa kurudisha `true`.

```js
assert.isTrue(isCircularPrime(197));
```

`isCircularPrime(23)` inapaswa kurudisha `false`.

```js
assert.isFalse(isCircularPrime(23));
```

`isCircularPrime(13)` inapaswa kurudisha `true`.

```js
assert.isTrue(isCircularPrime(13));
```

`isCircularPrime(89)` inapaswa kurudisha `false`.

```js
assert.isFalse(isCircularPrime(89));
```

`isCircularPrime(1193)` inapaswa kurudisha `true`.

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
