---
id: 69373793f5a867f769cde138
title: "Zoezi la 153: Tic-Tac-Toe"
challengeType: 28
dashedName: challenge-153
---

# --description--

Ikiwa umepewa matriki ya 3×3 (safu ya data ya safu za data) inayowakilisha mchezo wa Tic-Tac-Toe uliokamilika, amua mshindi.

- Kila kipengele katika matriki iliyotolewa ni `"X"` au `"O"`.

Mchezaji hushinda ikiwa ana wahusika wake watatu mfululizo - kwa usawa wa mstari, wima, au mwinuko.

Rudisha:

- `"X wins"` ikiwa mchezaji X ana wahusika watatu mfululizo.
- `"O wins"` ikiwa mchezaji O ana wahusika watatu mfululizo.
- `"Draw"` ikiwa hakuna mchezaji aliye na wahusika watatu mfululizo.

# --hints--

`ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` inapaswa kurudisha `"X wins"`.

```js
assert.equal(ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins");
```

`ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]])` inapaswa kurudisha `"X wins"`.

```js
assert.equal(ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` inapaswa kurudisha `"Draw"`.

```js
assert.equal(ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw");
```

`ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]])` inapaswa kurudisha `"O wins"`.

```js
assert.equal(ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` inapaswa kurudisha `"X wins"`.

```js
assert.equal(ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` inapaswa kurudisha `"Draw"`.

```js
assert.equal(ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw");
```

# --seed--

## --seed-contents--

```js
function ticTacToe(board) {

  return board;
}
```

# --solutions--

```js
function ticTacToe(board) {
  const lines = [];

  for (let row of board) {
    lines.push(row);
  }

  for (let c = 0; c < 3; c++) {
    lines.push([board[0][c], board[1][c], board[2][c]]);
  }

  lines.push([board[0][0], board[1][1], board[2][2]]);
  lines.push([board[0][2], board[1][1], board[2][0]]);

  for (let line of lines) {
    if (line.every(cell => cell === "X")) return "X wins";
    if (line.every(cell => cell === "O")) return "O wins";
  }

  return "Draw";
}
```
