---
id: 6939b873185d8e00d453563e
title: "Zoezi la 159: Hypotenuse ya nambari kamili"
challengeType: 28
dashedName: challenge-159
---

# --description--

Ukizingatia nambari kamili chanya mbili zinazowakilisha urefu wa miguu miwili (pande fupi mbili) za pembetatu ya haki, amua kama hypotenuse ni nambari kamili.

Urefu wa hypotenuse huhesabiwa kwa kuongeza mraba wa urefu wa miguu miwili kisha kuchukua mzizi wa mraba wa jumla hiyo (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

`isIntegerHypotenuse(3, 4)` inapaswa kurudisha `true`.

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)` inapaswa kurudisha `false`.

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)` inapaswa kurudisha `true`.

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)` inapaswa kurudisha `false`.

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)` inapaswa kurudisha `true`.

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)` inapaswa kurudisha `false`.

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
