---
id: 6939b873185d8e00d453563f
title: "Zoezi la 160: Miondoko ya Knight"
challengeType: 28
dashedName: challenge-160
---

# --description--

Kutokana na nafasi ya knight kwenye ubao wa chess, rudisha idadi ya mstatili halali knight anaweza kuhamia.

Ubao wa chess wa kawaida ni 8x8, na safu zimepewa lebo `A` hadi `H` (kutoka kushoto kwenda kulia) na mistari imepewa lebo `1` hadi `8` (kutoka chini kwenda juu). Inaonekana hivi:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Knight huenda kwa mwelekeo wa "L": mstatili wawili katika mwelekeo mmoja (wima au usawa), na mstatili mmoja katika mwelekeo wa wima.

Hii inamaanisha knight anaweza kuhamia hadi nafasi nane zinazowezekana, lakini chache zaidi anapokaribia mipaka ya ubao. Kwa mfano, kama knight alikuwa kwenye `A1`, angeweza kuhamia tu `B3` au `C2`.

# --hints--

`knightMoves("A1")` inapaswa kurudisha `2`.

```js
assert.equal(knightMoves("A1"), 2);
```

`knightMoves("D4")` inapaswa kurudisha `8`.

```js
assert.equal(knightMoves("D4"), 8);
```

`knightMoves("G6")` inapaswa kurudisha `6`.

```js
assert.equal(knightMoves("G6"), 6);
```

`knightMoves("B8")` inapaswa kurudisha `3`.

```js
assert.equal(knightMoves("B8"), 3);
```

`knightMoves("H3")` inapaswa kurudisha `4`.

```js
assert.equal(knightMoves("H3"), 4);
```

# --seed--

## --seed-contents--

```js
function knightMoves(position) {

  return position;
}
```

# --solutions--

```js
function knightMoves(position) {
  const col = position[0].charCodeAt(0) - 65;
  const row = parseInt(position[1]) - 1;

  const moves = [
    [2, 1], [2, -1],
    [-2, 1], [-2, -1],
    [1, 2], [1, -2],
    [-1, 2], [-1, -2]
  ];

  let valid = 0;

  for (const [dx, dy] of moves) {
    const newCol = col + dx;
    const newRow = row + dy;

    if (newCol >= 0 && newCol < 8 && newRow >= 0 && newRow < 8) {
      valid++;
    }
  }

  return valid;
}
```
