---
id: 6939b873185d8e00d453563d
title: "Zoezi la 158: Kubadilisha Safu ya data"
challengeType: 28
dashedName: challenge-158
---

# --description--

Ikiwa umepewa safu ya data yenye thamani mbili, rudisha safu ya data yenye thamani hizo zimeshushushwa.

Kwa mfano, ikiwa umepewa `["A", "B"]` rudisha `["B", "A"]`.

# --hints--

`arraySwap(["A", "B"])` inapaswa kurudisha `["B", "A"]`.

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` inapaswa kurudisha `[20, 25]`.

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` inapaswa kurudisha `[false, true]`.

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` inapaswa kurudisha `[1, "1"]`.

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
