---
id: 6939b873185d8e00d453563b
title: "Zoezi la 156: Nambari isiyo sawa au sawa?"
challengeType: 28
dashedName: challenge-156
---

# --description--

Ikiwa umepewa nambari kamili chanya, rudisha `"Odd"` ikiwa ni nambari isiyo sawa, na `"Even"` ikiwa ni nambari sawa.

# --hints--

`oddOrEven(1)` inapaswa kurudisha `"Odd"`.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` inapaswa kurudisha `"Even"`.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` inapaswa kurudisha `"Odd"`.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` inapaswa kurudisha `"Even"`.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` inapaswa kurudisha `"Odd"`.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
