---
id: 694596b0585c11170ac7c7f9
title: "Zoezi la 161: Usafirishaji wa bure"
challengeType: 28
dashedName: challenge-161
---

# --description--

Ikiwa umepewa safu ya data ya mfuatano wa herufi unaowakilisha vitu vilivyoko kwenye kikapu chako cha ununuzi, na nambari ya kiwango cha chini cha agizo ili kustahili usafirishaji wa bure, amua kama vitu vilivyoko kwenye kikapu chako cha ununuzi vinastahili usafirishaji wa bure.

Safu ya data iliyotolewa itakuwa na vitu kutoka kwenye orodha ifuatayo:

| Kitu | Bei |
| - | - |
| `"shirt"`	| 34.25 |
| `"jeans"`	| 48.50 |
| `"shoes"` |	75.00 |
| `"hat"`	| 19.95 |
| `"socks"` | 15.00 |
| `"jacket"` | 109.95 |

# --hints--

`getsFreeShipping(["shoes"], 50)` inapaswa kurudisha `true`.

```js
assert.isTrue(getsFreeShipping(["shoes"], 50));
```

`getsFreeShipping(["hat", "socks"], 50)` inapaswa kurudisha `false`.

```js
assert.isFalse(getsFreeShipping(["hat", "socks"], 50));
```

`getsFreeShipping(["jeans", "shirt", "jacket"], 75)` inapaswa kurudisha `true`.

```js
assert.isTrue(getsFreeShipping(["jeans", "shirt", "jacket"], 75));
```

`getsFreeShipping(["socks", "socks", "hat"], 75)` inapaswa kurudisha `false`.

```js
assert.isFalse(getsFreeShipping(["socks", "socks", "hat"], 75));
```

`getsFreeShipping(["shirt", "shirt", "jeans", "socks"], 100)` inapaswa kurudisha `true`.

```js
assert.isTrue(getsFreeShipping(["shirt", "shirt", "jeans", "socks"], 100));
```

`getsFreeShipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200)` inapaswa kurudisha `false`.

```js
assert.isFalse(getsFreeShipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200));
```

# --seed--

## --seed-contents--

```js
function getsFreeShipping(cart, minimum) {

  return cart;
}
```

# --solutions--

```js
function getsFreeShipping(cart, minimum) {
  const prices = {
    shirt: 34.25,
    jeans: 48.50,
    shoes: 75.00,
    hat: 19.95,
    socks: 15.00,
    jacket: 109.95
  };

  let total = 0;

  for (const item of cart) {
    total += prices[item];
  }

  return total >= minimum;
}
```
