---
id: 694596b0585c11170ac7c7fc
title: "Zoezi la 164: Kuchambua msimbo ndani ya mstari wa Markdown"
challengeType: 28
dashedName: challenge-164
---

# --description--

Ikiwa umepewa mfuatano wa herufi wa Markdown unaojumuisha sehemu moja au zaidi za msimbo ndani ya mstari, rudisha mfuatano wa herufi wa HTML unaolingana.

Sehemu za msimbo ndani ya mstari katika Markdown hutumia alama moja ya backtick (``` ` ```) mwanzoni na mwishoni mwa maandishi ya sehemu ya msimbo.

Rudisha mfuatano uliotolewa ukiwa na sehemu zote za msimbo zimebadilishwa kuwa lebo za HTML `code`.

Kwa mfano, ukipokea mfuatano ```"Use `let` to declare the variable."```, rudisha `"Use <code>let</code> to declare the variable."`.

Kumbuka: Konsoli huenda isionyeshe lebo za HTML ndani ya mfuatano wakati wa kurekodi ujumbe. Angalia konsoli ya kivinjari kuona rekodi zikiwa na lebo zilizoingizwa.

# --hints--

```parseInlineCode("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
assert.equal(parseInlineCode("Tumia `let` kutangaza kigezo."), "Tumia <code>let</code> kutangaza kigezo.");
```

```parseInlineCode("Tumia `let` au `const` kutangaza kigezo.")``` should return `"Tumia <code>let</code> au <code>const</code> kutangaza kigezo."`.

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

```parseInlineCode("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
assert.equal(parseInlineCode("Endesha `npm install` kisha `npm start`."), "Endesha <code>npm install</code> kisha <code>npm start</code>.");
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```
