---
id: 694596b0585c11170ac7c7fa
title: "Zoezi la 162: Matumizi ya nishati"
challengeType: 28
dashedName: challenge-162
---

# --description--

Kutokana na idadi ya Kalori zilizochomwa wakati wa mazoezi, na idadi ya watt-saa zilizotumika na vifaa vyako vya elektroniki wakati wa mazoezi hayo, amua ni ipi ilitumia nishati zaidi.

Ili kuzilinganisha, badilisha thamani zote mbili kuwa jouli kwa kutumia mabadiliko yafuatayo:

- 1 Kalori ni sawa na jouli 4184.
- 1 watt-saa ni sawa na jouli 3600.

Rudisha:

- `"Workout"` ikiwa mazoezi yalitumia nishati zaidi.
- `"Devices"` ikiwa kifaa kilitumia nishati zaidi.
- `"Equal"` ikiwa vyote vilitumia kiasi sawa cha nishati.

# --hints--

`compareEnergy(250, 50)` inapaswa kurudisha `"Workout"`.

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` inapaswa kurudisha `"Devices"`.

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` inapaswa kurudisha `"Equal"`.

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` inapaswa kurudisha `"Workout"`.

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` inapaswa kurudisha `"Devices"`.

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` inapaswa kurudisha `"Equal"`.

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
