---
id: 696655d24b614176d4c9b78d
title: "Zoezi la 170: Siku isiyo sawa au sawa"
challengeType: 28
dashedName: challenge-170
---

# --description--

Ukitoa timestamp (idadi ya millisekunde tangu mwanzo wa Unix epoch), rudisha:

- `"odd"` ikiwa nambari ya siku ya mwezi kwa timestamp hiyo ni isiyo sawa.
- `"even"` ikiwa nambari ya siku ya mwezi kwa timestamp hiyo ni sawa.

Kwa mfano, ukitoa `1769472000000`, timestamp ya tarehe 27 Januari, 2026, rudisha `"odd"` kwa sababu nambari ya siku (`27`) ni nambari isiyo sawa.

Kumbuka: Timestamp iko katika millisekunde na unapaswa kutumia tarehe katika saa za UTC, si saa za eneo lako la wakati.

# --hints--

`oddOrEvenDay(1769472000000)` inapaswa kurudisha `"odd"`.

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` inapaswa kurudisha `"even"`.

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` inapaswa kurudisha `"odd"`.

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` inapaswa kurudisha `"odd"`.

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` inapaswa kurudisha `"even"`.

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
