---
id: 696655d24b614176d4c9b789
title: "Zoezi la 166: Kagua mfuatano wa herufi wa hex"
challengeType: 28
dashedName: challenge-166
---

# --description--

Ukipata mfuatano wa herufi, thibitisha kama ni rangi halali ya hex ya CSS. Rangi halali ya hex ya CSS lazima:

- Ianze na `#`, na
- ifuatwe na herufi za hex 3 au 6.

Herufi za hex ni nambari `0` hadi `9` na herufi `a` hadi `f` (haina tofauti kati ya herufi kubwa na ndogo).

# --hints--

`isValidHex("#123")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
