---
id: 696655d24b614176d4c9b78c
title: "Zoezi la 169: FizzBuzz ndogo"
challengeType: 28
dashedName: challenge-169
---

# --description--

Ukipata nambari kamili, rudisha mfuatano wa herufi kulingana na sheria zifuatazo:

- Ikiwa nambari inagawika kwa 3, rudisha `"Fizz"`.
- Ikiwa nambari inagawika kwa 5, rudisha `"Buzz"`.
- Ikiwa nambari inagawika kwa 3 na 5, rudisha `"FizzBuzz"`.
- Vinginevyo, rudisha nambari hiyo kama mfuatano wa herufi.

# --hints--

`fizzBuzzMini(3)` inapaswa kurudisha `"Fizz"`.

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` inapaswa kurudisha `"4"`.

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` inapaswa kurudisha `"Buzz"`.

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` inapaswa kurudisha `"FizzBuzz"`.

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` inapaswa kurudisha `"98"`.

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
