---
id: 69738771fb5a7b8b24cca29d
title: "Zoezi la 171: Panga Safu ya Data Kuwa Moja"
challengeType: 28
dashedName: challenge-171
---

# --description--

Ikiwa umepewa safu ya data inayojumuisha safu zilizopangwa ndani, rudisha safu mpya yenye thamani zote zilizopangwa kuwa safu moja ya kipimo kimoja. Hifadhi mpangilio wa awali wa vitu katika safu hizo.

# --hints--

`flatten([1, [2, 3], 4])` inapaswa kurudisha `[1, 2, 3, 4]`.

```js
assert.deepEqual(flatten([1, [2, 3], 4]), [1, 2, 3, 4]);
```

`flatten([5, [4, [3, 2]], 1])` inapaswa kurudisha `[5, 4, 3, 2, 1]`.

```js
assert.deepEqual(flatten([5, [4, [3, 2]], 1]), [5, 4, 3, 2, 1]);
```

`flatten(["A", [[[["B"]]]], "C"])` inapaswa kurudisha `["A", "B", "C"]`.

```js
assert.deepEqual(flatten(["A", [[[["B"]]]], "C"]), ["A", "B", "C"]);
```

`flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]])` inapaswa kurudisha `["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]`.

```js
assert.deepEqual(flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]]), ["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]);
```

`flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]])` inapaswa kurudisha `["red","blue","green","yellow","purple","orange","pink","brown"]`.

```js
assert.deepEqual(flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]]), ["red","blue","green","yellow","purple","orange","pink","brown"]);
```

# --seed--

## --seed-contents--

```js
function flatten(arr) {

  return arr;
}
```

# --solutions--

```js
function flatten(arr) {
  let result = [];

  for (const item of arr) {
    if (Array.isArray(item)) {
      result = result.concat(flatten(item));
    } else {
      result.push(item);
    }
  }

  return result;
}
```
