---
id: 69738771fb5a7b8b24cca2a5
title: "Zoezi la 179: Mabadiliko ya Mfukoni"
challengeType: 28
dashedName: challenge-179
---

# --description--

Imetolewa safu ya data ya nambari kamili zinazowakilisha sarafu zilizoko mfukoni mwako, ambapo kila nambari kamili ni thamani ya sarafu kwa senti, rudisha jumla ya kiasi kwa muundo `"$D.CC"`.

- Senti 100 ni sawa na dola 1.
- Katika thamani ya return, jumuisha sifuri ya kuongoza kwa kiasi chini ya dola moja na daima nambari mbili tu kwa senti.

# --hints--

`countChange([25, 10, 5, 1])` inapaswa kurudisha `"$0.41"`.

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` inapaswa kurudisha `"$1.43"`.

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])` inapaswa kurudisha `"$37.55"`.

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` inapaswa kurudisha `"$0.70"`.

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])` inapaswa kurudisha `"$0.01"`.

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])` inapaswa kurudisha `"$1.00"`.

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```
