---
id: 69738771fb5a7b8b24cca2a2
title: "Zoezi la 176: Siku ya Groundhog"
challengeType: 28
dashedName: challenge-176
---

# --description--

Leo ni Siku ya Groundhog, ambapo groundhog hutabiri hali ya hewa kulingana na kama inaona kivuli chake au la.

Kutolewa thamani inayowakilisha muonekano wa groundhog, rudisha utabiri sahihi:

- Ikiwa thamani iliyotolewa ni thamani ya kweli au siyo kweli `true` (groundhog aliona kivuli chake), rudisha `"Looks like we'll have six more weeks of winter."`.
- Ikiwa thamani ni thamani ya kweli au siyo kweli `false` (groundhog hakuona kivuli chake), rudisha `"It's going to be an early spring."`.
- Ikiwa thamani ni kitu kingine chochote (groundhog hakuonekana), rudisha `"No prediction this year."`.

# --hints--

`groundhogDayPrediction(true)` inapaswa kurudisha `"Looks like we'll have six more weeks of winter."`.

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

`groundhogDayPrediction(false)` inapaswa kurudisha `"It's going to be an early spring."`.

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

`groundhogDayPrediction(null)` inapaswa kurudisha `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")` inapaswa kurudisha `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")` inapaswa kurudisha `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
