---
id: 69738771fb5a7b8b24cca2a4
title: "Zoezi la 178: Punguza mfuatano wa herufi"
challengeType: 28
dashedName: challenge-178
---

# --description--

Ikiwa umepewa mfuatano wa herufi, rudisha kama ulivyo ikiwa una herufi 20 au chini yake. Ikiwa una herufi zaidi ya 20, punguza hadi herufi 17 za kwanza kisha ongeza `"..."` mwishoni (ili jumla iwe herufi 20) na rudisha matokeo.

# --hints--

`truncateText("Hello, world!")` inapaswa kurudisha `"Hello, world!"`.

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` inapaswa kurudisha `"This string shoul..."`.

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` inapaswa kurudisha `"Exactly twenty chars"`.

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` inapaswa kurudisha `"...................."`.

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
