---
id: 69738771fb5a7b8b24cca2a0
title: "Zoezi la 174: Mtafutaji wa Zodiac"
challengeType: 28
dashedName: challenge-174
---

# --description--

Kutolewa mfuatano wa herufi wa tarehe katika muundo `"YYYY-MM-DD"`, rudisha alama ya zodiac kwa tarehe hiyo ukitumia chati ifuatayo:

| Kipindi cha tarehe | Alama ya zodiac |
| - | - |
| Machi 21 - Aprili 19 | `"Aries"` |
| Aprili 20 - Mei 20 | `"Taurus"` |
| Mei 21 - Juni 20 | `"Gemini"` |
| Juni 21 - Julai 22 | `"Cancer"` |
| Julai 23 - Agosti 22 | `"Leo"` |
| Agosti 23 - Septemba 22 | `"Virgo"` |
| Septemba 23 - Oktoba 22 | `"Libra"` |
| Oktoba 23 - Novemba 21 | `"Scorpio"` |
| Novemba 22 - Desemba 21 | `"Sagittarius"` |
| Desemba 22 - Januari 19 | `"Capricorn"` |
| Januari 20 - Februari 18 | `"Aquarius"` |
| Februari 19 - Machi 20 | `"Pisces"` |

- Alama za zodiac zinategemea tu mwezi na siku, unaweza kupuuza mwaka.

# --hints--

`getSign("2026-01-31")` inapaswa kurudisha `"Aquarius"`.

```js
assert.equal(getSign("2026-01-31"), "Aquarius");
```

`getSign("2001-06-10")` inapaswa kurudisha `"Gemini"`.

```js
assert.equal(getSign("2001-06-10"), "Gemini");
```

`getSign("1985-09-07")` inapaswa kurudisha `"Virgo"`.

```js
assert.equal(getSign("1985-09-07"), "Virgo");
```

`getSign("2023-03-19")` inapaswa kurudisha `"Pisces"`.

```js
assert.equal(getSign("2023-03-19"), "Pisces");
```

`getSign("2045-11-05")` inapaswa kurudisha `"Scorpio"`.

```js
assert.equal(getSign("2045-11-05"), "Scorpio");
```

`getSign("1985-12-06")` inapaswa kurudisha `"Sagittarius"`.

```js
assert.equal(getSign("1985-12-06"), "Sagittarius");
```

`getSign("2025-12-30")` inapaswa kurudisha `"Capricorn"`.

```js
assert.equal(getSign("2025-12-30"), "Capricorn");
```

`getSign("2018-10-08")` inapaswa kurudisha `"Libra"`.

```js
assert.equal(getSign("2018-10-08"), "Libra");
```

`getSign("1958-05-04")` inapaswa kurudisha `"Taurus"`.

```js
assert.equal(getSign("1958-05-04"), "Taurus");
```

# --seed--

## --seed-contents--

```js
function getSign(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getSign(dateStr) {
  const [, monthStr, dayStr] = dateStr.split("-");
  const month = Number(monthStr);
  const day = Number(dayStr);

  if ((month === 3 && day >= 21) || (month === 4 && day <= 19)) return "Aries";
  if ((month === 4 && day >= 20) || (month === 5 && day <= 20)) return "Taurus";
  if ((month === 5 && day >= 21) || (month === 6 && day <= 20)) return "Gemini";
  if ((month === 6 && day >= 21) || (month === 7 && day <= 22)) return "Cancer";
  if ((month === 7 && day >= 23) || (month === 8 && day <= 22)) return "Leo";
  if ((month === 8 && day >= 23) || (month === 9 && day <= 22)) return "Virgo";
  if ((month === 9 && day >= 23) || (month === 10 && day <= 22)) return "Libra";
  if ((month === 10 && day >= 23) || (month === 11 && day <= 21)) return "Scorpio";
  if ((month === 11 && day >= 22) || (month === 12 && day <= 21)) return "Sagittarius";
  if ((month === 12 && day >= 22) || (month === 1 && day <= 19)) return "Capricorn";
  if ((month === 1 && day >= 20) || (month === 2 && day <= 18)) return "Aquarius";
  if ((month === 2 && day >= 19) || (month === 3 && day <= 20)) return "Pisces";
}
```
