---
id: 69738771fb5a7b8b24cca29f
title: "Zoezi la 173: Mchezo halali wa pawn"
challengeType: 28
dashedName: challenge-173
---

# --description--

Kutokana na nafasi ya moja ya pawns zako kwenye ubao wa chess, rudisha safu ya data ya pande zote halali ambazo inaweza kusogea kwa mpangilio wa kuongezeka.

Ubao wa chess wa kawaida ni 8x8, na safu zimepewa lebo `A` hadi `H` (kutoka kushoto kwenda kulia) na mistari imepewa lebo `1` hadi `8` (kutoka chini kwenda juu). Inaonekana kama ifuatavyo:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Kwa changamoto hii:

- Wewe ni mchezaji aliye chini ya ubao.
- Pawns zinaweza kusogea tu moja kwa moja "juu".
- Isipokuwa pawn iko kwenye mstari wa kuanzia (mstari wa 2), basi inaweza kusogea moja au mbili juu.

Kwa mfano, ukipata `"D4"`, rudisha `["D5"]`, ni pande pekee ambayo pawn yako inaweza kusogea.
Ukipata `"B2"`, rudisha `["B3", "B4"]`, kwa sababu iko kwenye mstari wa kuanzia na inahitaji kuwa kwa mpangilio wa kuongezeka.

# --hints--

`findPawnMoves("D4")` inapaswa kurudisha `["D5"]`.

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")` inapaswa kurudisha `["B3", "B4"]`.

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")` inapaswa kurudisha `["A8"]`.

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")` inapaswa kurudisha `["G3", "G4"]`.

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")` inapaswa kurudisha `["E4"]`.

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```
