---
id: 697a49e9860d24853adef67f
title: "Zoezi la 194: Michezo ya Majira ya Baridi 2026 Siku ya 15: Kuteleza kwa Freestyle"
challengeType: 28
dashedName: challenge-194
---

# --description--

Ikiwa jina la mbinu lina maneno mawili, thibitisha kama ni jina halali la mbinu ya kuteleza kwa freestyle.

Mbinu ni halali ikiwa neno la kwanza liko kwenye orodha ya maneno ya kwanza halali, na neno la pili liko kwenye orodha ya maneno ya pili halali.

- Maneno mawili yatatenganishwa na nafasi moja tu.

Maneno ya kwanza halali:

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

Maneno ya pili halali:

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`isValidTrick("Polar Vortex")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidTrick("Polar Vortex"));
```

`isValidTrick("Solar Icequake")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidTrick("Solar Icequake"));
```

`isValidTrick("Thunder Blizzard")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidTrick("Thunder Blizzard"));
```

`isValidTrick("Phantom Frostbite")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidTrick("Phantom Frostbite"));
```

`isValidTrick("Ghost Avalanche")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidTrick("Ghost Avalanche"));
```

`isValidTrick("Snowstorm Shadow")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidTrick("Snowstorm Shadow"));
```

`isValidTrick("Solar Sky")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidTrick("Solar Sky"));
```

# --seed--

## --seed-contents--

```js
function isValidTrick(trickName) {

  return trickName;
}
```

# --solutions--

```js
function isValidTrick(trickName) {
  const validFirst = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"];
  const validSecond = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"];

  const words = trickName.split(" ");
  const [first, second] = words;
  return validFirst.includes(first) && validSecond.includes(second);
}
```
