---
id: 697a49e6ff50d756c9b69360
title: "Zoezi la 183: Michezo ya Majira ya Baridi 2026 Siku ya 4: Kuruka kwa Ski"
challengeType: 28
dashedName: challenge-183
---

# --description--

Ukizingatia pointi za umbali, pointi za mtindo, thamani ya fidia ya upepo, na bonasi ya K-point, hesabu alama zako za kuruka kwa ski na uamue kama umeweza kushinda medali au la.

- Alama zako zinahesabiwa kwa jumlisha pointi hizo nne zilizo hapo juu.

Alama za jumla za wachezaji wengine wa kuruka kwa sasa ni:

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- Ikiwa alama zako ndizo bora zaidi, rudisha `"Gold"`
- Ikiwa ni za pili kwa ubora, rudisha `"Silver"`
- Ikiwa ni za tatu kwa ubora, rudisha `"Bronze"`
- Vinginevyo, rudisha `"No Medal"`

# --hints--

`skiJumpMedal(125.0, 58.0, 0.0, 6.0)` inapaswa kurudisha `"Gold"`.

```js
assert.equal(skiJumpMedal(125.0, 58.0, 0.0, 6.0), "Gold");
```

`skiJumpMedal(119.0, 50.0, 1.0, 4.0)` inapaswa kurudisha `"Bronze"`.

```js
assert.equal(skiJumpMedal(119.0, 50.0, 1.0, 4.0), "Bronze");
```

`skiJumpMedal(122.0, 52.0, -1.0, 4.0)` inapaswa kurudisha `"Silver"`.

```js
assert.equal(skiJumpMedal(122.0, 52.0, -1.0, 4.0), "Silver");
```

`skiJumpMedal(118.0, 50.5, -1.5, 4.0)` inapaswa kurudisha `"No Medal"`.

```js
assert.equal(skiJumpMedal(118.0, 50.5, -1.5, 4.0), "No Medal");
```

`skiJumpMedal(124.0, 50.5, 2.0, 5.0)` inapaswa kurudisha `"Gold"`.

```js
assert.equal(skiJumpMedal(124.0, 50.5, 2.0, 5.0), "Gold");
```

`skiJumpMedal(119.0, 49.5, 0.0, 3.0)` inapaswa kurudisha `"No Medal"`.

```js
assert.equal(skiJumpMedal(119.0, 49.5, 0.0, 3.0), "No Medal");
```

# --seed--

## --seed-contents--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {

  return distancePoints;
}
```

# --solutions--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {
  const myScore = distancePoints + stylePoints + windComp + kPointBonus;
  const otherScores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0];
  const allScores = [...otherScores, myScore];
  allScores.sort((a, b) => b - a);
  const rank = allScores.indexOf(myScore) + 1;

  if (rank === 1) return "Gold";
  if (rank === 2) return "Silver";
  if (rank === 3) return "Bronze";
  return "No Medal";
}
```
