---
id: 697a49e9860d24853adef67d
title: "Zoezi la 192: Michezo ya Majira ya Baridi 2026 Siku ya 13: Mchanganyiko wa Nordic"
challengeType: 28
dashedName: challenge-192
---

# --description--

Kutolewa safu ya data ya alama za kuruka kwa wanariadha, hesabu nyakati zao za kuchelewa kuanza kwa sehemu ya mbio za msalaba wa Nordic Combined.

Mwanariadha mwenye alama ya juu zaidi ya kuruka huanza kwanza (hakuna kuchelewa kwa sekunde 0). Wanariadha wengine wote huanza baadaye kulingana na umbali wa alama zao za kuruka kutoka kwa alama bora zaidi.

Ili kuhesabu kuchelewa kwa kila mwanariadha, toa alama ya kuruka ya mwanariadha kutoka kwa alama bora kabisa ya kuruka kisha zidisha matokeo kwa 1.5. Punguza kuchelewa hadi nambari kamili inayofuata.

# --hints--

`calculateStartDelays([120, 110, 125])` inapaswa kurudisha `[8, 23, 0]`.

```js
assert.deepEqual(calculateStartDelays([120, 110, 125]), [8, 23, 0]);
```

`calculateStartDelays([118, 125, 122, 120])` inapaswa kurudisha `[11, 0, 5, 8]`.

```js
assert.deepEqual(calculateStartDelays([118, 125, 122, 120]), [11, 0, 5, 8]);
```

`calculateStartDelays([100, 105, 95, 110, 120, 115, 108])` inapaswa kurudisha `[30, 23, 38, 15, 0, 8, 18]`.

```js
assert.deepEqual(calculateStartDelays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18]);
```

`calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` inapaswa kurudisha `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
assert.deepEqual(calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12]);
```

# --seed--

## --seed-contents--

```js
function calculateStartDelays(jumpScores) {

  return jumpScores
}
```

# --solutions--

```js
function calculateStartDelays(jumpScores) {
  const bestJump = Math.max(...jumpScores);
  return jumpScores.map(score => Math.ceil((bestJump - score) * 1.5));
}
```
