---
id: 697a49e6ff50d756c9b69363
title: "Zoezi la 186: Michezo ya Majira ya Baridi 2026 Siku ya 7: Kuendesha Sketi kwa Kasi"
challengeType: 28
dashedName: challenge-186
---

# --description--

Kutolewa safu mbili za data zinazowakilisha nyakati za mizunguko (katika sekunde) kwa waendesha sketi wawili wa kasi, rudisha nambari ya mzunguko ambapo tofauti ya nyakati za mizunguko ni kubwa zaidi.

Kipengele cha kwanza cha kila safu ya data kinahusiana na mzunguko wa 1, cha pili na mzunguko wa 2, na kadhalika.

# --hints--

`largestDifference([26.11, 25.80, 25.92, 26.23, 26.07], [25.93, 25.74, 26.53, 26.11, 26.30])` inapaswa kurudisha `3`.

```js
assert.equal(largestDifference([26.11, 25.80, 25.92, 26.23, 26.07], [25.93, 25.74, 26.53, 26.11, 26.30]), 3);
```

`largestDifference([27.04, 25.94, 26.22, 26.07, 26.18], [25.59, 25.80, 26.11, 26.01, 26.23])` inapaswa kurudisha `1`.

```js
assert.equal(largestDifference([27.04, 25.94, 26.22, 26.07, 26.18], [25.59, 25.80, 26.11, 26.01, 26.23]), 1);
```

`largestDifference([25.82, 25.90, 26.05, 26.00, 26.48], [25.85, 25.92, 26.07, 25.98, 25.95])` inapaswa kurudisha `5`.

```js
assert.equal(largestDifference([25.82, 25.90, 26.05, 26.00, 26.48], [25.85, 25.92, 26.07, 25.98, 25.95]), 5);
```

`largestDifference([25.88, 26.10, 25.95, 26.05, 26.00], [25.90, 26.55, 25.92, 26.03, 26.01])` inapaswa kurudisha `2`.

```js
assert.equal(largestDifference([25.88, 26.10, 25.95, 26.05, 26.00], [25.90, 26.55, 25.92, 26.03, 26.01]), 2);
```

`largestDifference([25.92, 26.01, 26.05, 25.88, 26.12], [25.95, 26.00, 26.03, 26.45, 26.10])` inapaswa kurudisha `4`.

```js
assert.equal(largestDifference([25.92, 26.01, 26.05, 25.88, 26.12], [25.95, 26.00, 26.03, 26.45, 26.10]), 4);
```

# --seed--

## --seed-contents--

```js
function largestDifference(skater1, skater2) {

  return skater1;
}
```

# --solutions--

```js
function largestDifference(skater1, skater2) {
  let maxDiff = -1;
  let lapNumber = 1;

  for (let i = 0; i < skater1.length; i++) {
    const diff = Math.abs(skater1[i] - skater2[i]);
    if (diff > maxDiff) {
      maxDiff = diff;
      lapNumber = i + 1;
    }
  }

  return lapNumber;
}
```
