---
id: 697a49e6ff50d756c9b6935f
title: "Changamoto 182: Michezo ya Majira ya Baridi 2026 Siku ya 3: Biathlon"
challengeType: 28
dashedName: challenge-182
---

# --description--

Imetolewa safu ya data ya nambari kamili, ambapo kila thamani inawakilisha idadi ya malengo yaliyopigwa katika mzunguko mmoja wa biathlon, rudisha jumla ya umbali wa adhabu ambayo mwanariadha lazima apige ski.

- Kila mzunguko una malengo 5.
- Kila lengo lililokosewa husababisha mzunguko wa adhabu wa mita 150.

# --hints--

`calculatePenaltyDistance([4, 4])` inapaswa kurudisha `300`.

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])` inapaswa kurudisha `0`.

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])` inapaswa kurudisha `450`.

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])` inapaswa kurudisha `150`.

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])` inapaswa kurudisha `1500`.

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```
