---
id: 697a49e6ff50d756c9b6935d
title: "Zoezi la 180: Michezo ya Majira ya Baridi 2026 Siku ya 1: Siku ya Ufunguzi"
challengeType: 28
dashedName: challenge-180
---

# --description--

Leo ni mwanzo wa Michezo ya Majira ya Baridi ya 2026. Siku 17 zijazo zitakuletea changamoto za msimbo zilizoongozwa na michezo hiyo.

Kwa changamoto ya kwanza, unapewa msimbo wa nchi wenye herufi mbili na unahitaji kurudisha emoji ya bendera ya nchi hiyo.

Tumia orodha hii:

| Nchi | Msimbo | Bendera |
| - | -: | :--: |
| Albania                    |            `"AL"` | `"馃嚘馃嚤"` |
| Andorra                    |            `"AD"` | `"馃嚘馃嚛"` |
| Argentina                  |            `"AR"` | `"馃嚘馃嚪"` |
| Armenia                    |            `"AM"` | `"馃嚘馃嚥"` |
| Australia                  |            `"AU"` | `"馃嚘馃嚭"` |
| Austria                    |            `"AT"` | `"馃嚘馃嚬"` |
| Azerbaijan                 |            `"AZ"` | `"馃嚘馃嚳"` |
| Belgium                    |            `"BE"` | `"馃嚙馃嚜"` |
| Benin                      |            `"BJ"` | `"馃嚙馃嚡"` |
| Bolivia                    |            `"BO"` | `"馃嚙馃嚧"` |
| Bosnia and Herzegovina     |            `"BA"` | `"馃嚙馃嚘"` |
| Brazil                     |            `"BR"` | `"馃嚙馃嚪"` |
| Bulgaria                   |            `"BG"` | `"馃嚙馃嚞"` |
| Canada                     |            `"CA"` | `"馃嚚馃嚘"` |
| Chile                      |            `"CL"` | `"馃嚚馃嚤"` |
| China                      |            `"CN"` | `"馃嚚馃嚦"` |
| Colombia                   |            `"CO"` | `"馃嚚馃嚧"` |
| Croatia                    |            `"HR"` | `"馃嚟馃嚪"` |
| Cyprus                     |            `"CY"` | `"馃嚚馃嚲"` |
| Czech Republic             |            `"CZ"` | `"馃嚚馃嚳"` |
| Denmark                    |            `"DK"` | `"馃嚛馃嚢"` |
| Ecuador                    |            `"EC"` | `"馃嚜馃嚚"` |
| Eritrea                    |            `"ER"` | `"馃嚜馃嚪"` |
| Estonia                    |            `"EE"` | `"馃嚜馃嚜"` |
| Finland                    |            `"FI"` | `"馃嚝馃嚠"` |
| France                     |            `"FR"` | `"馃嚝馃嚪"` |
| Georgia                    |            `"GE"` | `"馃嚞馃嚜"` |
| Germany                    |            `"DE"` | `"馃嚛馃嚜"` |
| Great Britain              |            `"GB"` | `"馃嚞馃嚙"` |
| Greece                     |            `"GR"` | `"馃嚞馃嚪"` |
| Guinea-Bissau              |            `"GW"` | `"馃嚞馃嚰"` |
| Haiti                      |            `"HT"` | `"馃嚟馃嚬"` |
| Hong Kong                  |            `"HK"` | `"馃嚟馃嚢"` |
| Hungary                    |            `"HU"` | `"馃嚟馃嚭"` |
| Iceland                    |            `"IS"` | `"馃嚠馃嚫"` |
| India                      |            `"IN"` | `"馃嚠馃嚦"` |
| Iran                       |            `"IR"` | `"馃嚠馃嚪"` |
| Ireland                    |            `"IE"` | `"馃嚠馃嚜"` |
| Israel                     |            `"IL"` | `"馃嚠馃嚤"` |
| Italy                      |            `"IT"` | `"馃嚠馃嚬"` |
| Jamaica                    |            `"JM"` | `"馃嚡馃嚥"` |
| Japan                      |            `"JP"` | `"馃嚡馃嚨"` |
| Kazakhstan                 |            `"KZ"` | `"馃嚢馃嚳"` |
| Kenya                      |            `"KE"` | `"馃嚢馃嚜"` |
| Kosovo                     |            `"XK"` | `"馃嚱馃嚢"` |
| Kyrgyzstan                 |            `"KG"` | `"馃嚢馃嚞"` |
| Latvia                     |            `"LV"` | `"馃嚤馃嚮"` |
| Lebanon                    |            `"LB"` | `"馃嚤馃嚙"` |
| Liechtenstein              |            `"LI"` | `"馃嚤馃嚠"` |
| Lithuania                  |            `"LT"` | `"馃嚤馃嚬"` |
| Luxembourg                 |            `"LU"` | `"馃嚤馃嚭"` |
| Madagascar                 |            `"MG"` | `"馃嚥馃嚞"` |
| Malaysia                   |            `"MY"` | `"馃嚥馃嚲"` |
| Malta                      |            `"MT"` | `"馃嚥馃嚬"` |
| Mexico                     |            `"MX"` | `"馃嚥馃嚱"` |
| Moldova                    |            `"MD"` | `"馃嚥馃嚛"` |
| Monaco                     |            `"MC"` | `"馃嚥馃嚚"` |
| Mongolia                   |            `"MN"` | `"馃嚥馃嚦"` |
| Montenegro                 |            `"ME"` | `"馃嚥馃嚜"` |
| Morocco                    |            `"MA"` | `"馃嚥馃嚘"` |
| Netherlands                |            `"NL"` | `"馃嚦馃嚤"` |
| New Zealand                |            `"NZ"` | `"馃嚦馃嚳"` |
| Nigeria                    |            `"NG"` | `"馃嚦馃嚞"` |
| North Macedonia            |            `"MK"` | `"馃嚥馃嚢"` |
| Norway                     |            `"NO"` | `"馃嚦馃嚧"` |
| Pakistan                   |            `"PK"` | `"馃嚨馃嚢"` |
| Philippines                |            `"PH"` | `"馃嚨馃嚟"` |
| Poland                     |            `"PL"` | `"馃嚨馃嚤"` |
| Portugal                   |            `"PT"` | `"馃嚨馃嚬"` |
| Puerto Rico                |            `"PR"` | `"馃嚨馃嚪"` |
| Romania                    |            `"RO"` | `"馃嚪馃嚧"` |
| San Marino                 |            `"SM"` | `"馃嚫馃嚥"` |
| Saudi Arabia               |            `"SA"` | `"馃嚫馃嚘"` |
| Serbia                     |            `"RS"` | `"馃嚪馃嚫"` |
| Singapore                  |            `"SG"` | `"馃嚫馃嚞"` |
| Slovakia                   |            `"SK"` | `"馃嚫馃嚢"` |
| Slovenia                   |            `"SI"` | `"馃嚫馃嚠"` |
| South Africa               |            `"ZA"` | `"馃嚳馃嚘"` |
| South Korea                |            `"KR"` | `"馃嚢馃嚪"` |
| Spain                      |            `"ES"` | `"馃嚜馃嚫"` |
| Sweden                     |            `"SE"` | `"馃嚫馃嚜"` |
| Switzerland                |            `"CH"` | `"馃嚚馃嚟"` |
| Thailand                   |            `"TH"` | `"馃嚬馃嚟"` |
| Trinidad & Tobago          |            `"TT"` | `"馃嚬馃嚬"` |
| Turkey                     |            `"TR"` | `"馃嚬馃嚪"` |
| Ukraine                    |            `"UA"` | `"馃嚭馃嚘"` |
| United Arab Emirates       |            `"AE"` | `"馃嚘馃嚜"` |
| United States              |            `"US"` | `"馃嚭馃嚫"` |
| Uruguay                    |            `"UY"` | `"馃嚭馃嚲"` |
| Uzbekistan                 |            `"UZ"` | `"馃嚭馃嚳"` |
| Venezuela                  |            `"VE"` | `"馃嚮馃嚜"` |

# --hints--

`getFlag("AL")` inapaswa kurudisha `"馃嚘馃嚤"`.

```js
assert.equal(getFlag("AL"), "馃嚘馃嚤");
```

`getFlag("AD")` inapaswa kurudisha `"馃嚘馃嚛"`.

```js
assert.equal(getFlag("AD"), "馃嚘馃嚛");
```

`getFlag("AR")` inapaswa kurudisha `"馃嚘馃嚪"`.

```js
assert.equal(getFlag("AR"), "馃嚘馃嚪")
```

`getFlag("AM")` inapaswa kurudisha `"馃嚘馃嚥"`.

```js
assert.equal(getFlag("AM"), "馃嚘馃嚥")
```

`getFlag("AU")` inapaswa kurudisha `"馃嚘馃嚭"`.

```js
assert.equal(getFlag("AU"), "馃嚘馃嚭")
```

`getFlag("AT")` inapaswa kurudisha `"馃嚘馃嚬"`.

```js
assert.equal(getFlag("AT"), "馃嚘馃嚬")
```

`getFlag("AZ")` inapaswa kurudisha `"馃嚘馃嚳"`.

```js
assert.equal(getFlag("AZ"), "馃嚘馃嚳")
```

`getFlag("BE")` inapaswa kurudisha `"馃嚙馃嚜"`.

```js
assert.equal(getFlag("BE"), "馃嚙馃嚜")
```

`getFlag("BJ")` inapaswa kurudisha `"馃嚙馃嚡"`.

```js
assert.equal(getFlag("BJ"), "馃嚙馃嚡")
```

`getFlag("BO")` inapaswa kurudisha `"馃嚙馃嚧"`.

```js
assert.equal(getFlag("BO"), "馃嚙馃嚧")
```

`getFlag("BA")` inapaswa kurudisha `"馃嚙馃嚘"`.

```js
assert.equal(getFlag("BA"), "馃嚙馃嚘")
```

`getFlag("BR")` inapaswa kurudisha `"馃嚙馃嚪"`.

```js
assert.equal(getFlag("BR"), "馃嚙馃嚪")
```

`getFlag("BG")` inapaswa kurudisha `"馃嚙馃嚞"`.

```js
assert.equal(getFlag("BG"), "馃嚙馃嚞")
```

`getFlag("CA")` inapaswa kurudisha `"馃嚚馃嚘"`.

```js
assert.equal(getFlag("CA"), "馃嚚馃嚘")
```

`getFlag("CL")` inapaswa kurudisha `"馃嚚馃嚤"`.

```js
assert.equal(getFlag("CL"), "馃嚚馃嚤")
```

`getFlag("CN")` inapaswa kurudisha `"馃嚚馃嚦"`.

```js
assert.equal(getFlag("CN"), "馃嚚馃嚦")
```

`getFlag("CO")` inapaswa kurudisha `"馃嚚馃嚧"`.

```js
assert.equal(getFlag("CO"), "馃嚚馃嚧")
```

`getFlag("HR")` inapaswa kurudisha `"馃嚟馃嚪"`.

```js
assert.equal(getFlag("HR"), "馃嚟馃嚪")
```

`getFlag("CY")` inapaswa kurudisha `"馃嚚馃嚲"`.

```js
assert.equal(getFlag("CY"), "馃嚚馃嚲")
```

`getFlag("CZ")` inapaswa kurudisha `"馃嚚馃嚳"`.

```js
assert.equal(getFlag("CZ"), "馃嚚馃嚳")
```

`getFlag("DK")` inapaswa kurudisha `"馃嚛馃嚢"`.

```js
assert.equal(getFlag("DK"), "馃嚛馃嚢")
```

`getFlag("EC")` inapaswa kurudisha `"馃嚜馃嚚"`.

```js
assert.equal(getFlag("EC"), "馃嚜馃嚚")
```

`getFlag("ER")` inapaswa kurudisha `"馃嚜馃嚪"`.

```js
assert.equal(getFlag("ER"), "馃嚜馃嚪")
```

`getFlag("EE")` inapaswa kurudisha `"馃嚜馃嚜"`.

```js
assert.equal(getFlag("EE"), "馃嚜馃嚜")
```

`getFlag("FI")` inapaswa kurudisha `"馃嚝馃嚠"`.

```js
assert.equal(getFlag("FI"), "馃嚝馃嚠")
```

`getFlag("FR")` inapaswa kurudisha `"馃嚝馃嚪"`.

```js
assert.equal(getFlag("FR"), "馃嚝馃嚪")
```

`getFlag("GE")` inapaswa kurudisha `"馃嚞馃嚜"`.

```js
assert.equal(getFlag("GE"), "馃嚞馃嚜")
```

`getFlag("DE")` inapaswa kurudisha `"馃嚛馃嚜"`.

```js
assert.equal(getFlag("DE"), "馃嚛馃嚜")
```

`getFlag("GB")` inapaswa kurudisha `"馃嚞馃嚙"`.

```js
assert.equal(getFlag("GB"), "馃嚞馃嚙")
```

`getFlag("GR")` inapaswa kurudisha `"馃嚞馃嚪"`.

```js
assert.equal(getFlag("GR"), "馃嚞馃嚪")
```

`getFlag("GW")` inapaswa kurudisha `"馃嚞馃嚰"`.

```js
assert.equal(getFlag("GW"), "馃嚞馃嚰")
```

`getFlag("HT")` inapaswa kurudisha `"馃嚟馃嚬"`.

```js
assert.equal(getFlag("HT"), "馃嚟馃嚬")
```

`getFlag("HK")` inapaswa kurudisha `"馃嚟馃嚢"`.

```js
assert.equal(getFlag("HK"), "馃嚟馃嚢")
```

`getFlag("HU")` inapaswa kurudisha `"馃嚟馃嚭"`.

```js
assert.equal(getFlag("HU"), "馃嚟馃嚭")
```

`getFlag("IS")` inapaswa kurudisha `"馃嚠馃嚫"`.

```js
assert.equal(getFlag("IS"), "馃嚠馃嚫")
```

`getFlag("IN")` inapaswa kurudisha `"馃嚠馃嚦"`.

```js
assert.equal(getFlag("IN"), "馃嚠馃嚦")
```

`getFlag("IR")` inapaswa kurudisha `"馃嚠馃嚪"`.

```js
assert.equal(getFlag("IR"), "馃嚠馃嚪")
```

`getFlag("IE")` inapaswa kurudisha `"馃嚠馃嚜"`.

```js
assert.equal(getFlag("IE"), "馃嚠馃嚜")
```

`getFlag("IL")` inapaswa kurudisha `"馃嚠馃嚤"`.

```js
assert.equal(getFlag("IL"), "馃嚠馃嚤")
```

`getFlag("IT")` inapaswa kurudisha `"馃嚠馃嚬"`.

```js
assert.equal(getFlag("IT"), "馃嚠馃嚬")
```

`getFlag("JM")` inapaswa kurudisha `"馃嚡馃嚥"`.

```js
assert.equal(getFlag("JM"), "馃嚡馃嚥")
```

`getFlag("JP")` inapaswa kurudisha `"馃嚡馃嚨"`.

```js
assert.equal(getFlag("JP"), "馃嚡馃嚨")
```

`getFlag("KZ")` inapaswa kurudisha `"馃嚢馃嚳"`.

```js
assert.equal(getFlag("KZ"), "馃嚢馃嚳")
```

`getFlag("KE")` inapaswa kurudisha `"馃嚢馃嚜"`.

```js
assert.equal(getFlag("KE"), "馃嚢馃嚜")
```

`getFlag("XK")` inapaswa kurudisha `"馃嚱馃嚢"`.

```js
assert.equal(getFlag("XK"), "馃嚱馃嚢")
```

`getFlag("KG")` inapaswa kurudisha `"馃嚢馃嚞"`.

```js
assert.equal(getFlag("KG"), "馃嚢馃嚞")
```

`getFlag("LV")` inapaswa kurudisha `"馃嚤馃嚮"`.

```js
assert.equal(getFlag("LV"), "馃嚤馃嚮")
```

`getFlag("LB")` inapaswa kurudisha `"馃嚤馃嚙"`.

```js
assert.equal(getFlag("LB"), "馃嚤馃嚙")
```

`getFlag("LI")` inapaswa kurudisha `"馃嚤馃嚠"`.

```js
assert.equal(getFlag("LI"), "馃嚤馃嚠")
```

`getFlag("LT")` inapaswa kurudisha `"馃嚤馃嚬"`.

```js
assert.equal(getFlag("LT"), "馃嚤馃嚬")
```

`getFlag("LU")` inapaswa kurudisha `"馃嚤馃嚭"`.

```js
assert.equal(getFlag("LU"), "馃嚤馃嚭")
```

`getFlag("MG")` inapaswa kurudisha `"馃嚥馃嚞"`.

```js
assert.equal(getFlag("MG"), "馃嚥馃嚞")
```

`getFlag("MY")` inapaswa kurudisha `"馃嚥馃嚲"`.

```js
assert.equal(getFlag("MY"), "馃嚥馃嚲")
```

`getFlag("MT")` inapaswa kurudisha `"馃嚥馃嚬"`.

```js
assert.equal(getFlag("MT"), "馃嚥馃嚬")
```

`getFlag("MX")` inapaswa kurudisha `"馃嚥馃嚱"`.

```js
assert.equal(getFlag("MX"), "馃嚥馃嚱")
```

`getFlag("MD")` inapaswa kurudisha `"馃嚥馃嚛"`.

```js
assert.equal(getFlag("MD"), "馃嚥馃嚛")
```

`getFlag("MC")` inapaswa kurudisha `"馃嚥馃嚚"`.

```js
assert.equal(getFlag("MC"), "馃嚥馃嚚")
```

`getFlag("MN")` inapaswa kurudisha `"馃嚥馃嚦"`.

```js
assert.equal(getFlag("MN"), "馃嚥馃嚦")
```

`getFlag("ME")` inapaswa kurudisha `"馃嚥馃嚜"`.

```js
assert.equal(getFlag("ME"), "馃嚥馃嚜")
```

`getFlag("MA")` inapaswa kurudisha `"馃嚥馃嚘"`.

```js
assert.equal(getFlag("MA"), "馃嚥馃嚘")
```

`getFlag("NL")` inapaswa kurudisha `"馃嚦馃嚤"`.

```js
assert.equal(getFlag("NL"), "馃嚦馃嚤")
```

`getFlag("NZ")` inapaswa kurudisha `"馃嚦馃嚳"`.

```js
assert.equal(getFlag("NZ"), "馃嚦馃嚳")
```

`getFlag("NG")` inapaswa kurudisha `"馃嚦馃嚞"`.

```js
assert.equal(getFlag("NG"), "馃嚦馃嚞")
```

`getFlag("MK")` inapaswa kurudisha `"馃嚥馃嚢"`.

```js
assert.equal(getFlag("MK"), "馃嚥馃嚢")
```

`getFlag("NO")` inapaswa kurudisha `"馃嚦馃嚧"`.

```js
assert.equal(getFlag("NO"), "馃嚦馃嚧")
```

`getFlag("PK")` inapaswa kurudisha `"馃嚨馃嚢"`.

```js
assert.equal(getFlag("PK"), "馃嚨馃嚢")
```

`getFlag("PH")` inapaswa kurudisha `"馃嚨馃嚟"`.

```js
assert.equal(getFlag("PH"), "馃嚨馃嚟")
```

`getFlag("PL")` inapaswa kurudisha `"馃嚨馃嚤"`.

```js
assert.equal(getFlag("PL"), "馃嚨馃嚤")
```

`getFlag("PT")` inapaswa kurudisha `"馃嚨馃嚬"`.

```js
assert.equal(getFlag("PT"), "馃嚨馃嚬")
```

`getFlag("PR")` inapaswa kurudisha `"馃嚨馃嚪"`.

```js
assert.equal(getFlag("PR"), "馃嚨馃嚪")
```

`getFlag("RO")` inapaswa kurudisha `"馃嚪馃嚧"`.

```js
assert.equal(getFlag("RO"), "馃嚪馃嚧")
```

`getFlag("SM")` inapaswa kurudisha `"馃嚫馃嚥"`.

```js
assert.equal(getFlag("SM"), "馃嚫馃嚥")
```

`getFlag("SA")` inapaswa kurudisha `"馃嚫馃嚘"`.

```js
assert.equal(getFlag("SA"), "馃嚫馃嚘")
```

`getFlag("RS")` inapaswa kurudisha `"馃嚪馃嚫"`.

```js
assert.equal(getFlag("RS"), "馃嚪馃嚫")
```

`getFlag("SG")` inapaswa kurudisha `"馃嚫馃嚞"`.

```js
assert.equal(getFlag("SG"), "馃嚫馃嚞")
```

`getFlag("SK")` inapaswa kurudisha `"馃嚫馃嚢"`.

```js
assert.equal(getFlag("SK"), "馃嚫馃嚢")
```

`getFlag("SI")` inapaswa kurudisha `"馃嚫馃嚠"`.

```js
assert.equal(getFlag("SI"), "馃嚫馃嚠")
```

`getFlag("ZA")` inapaswa kurudisha `"馃嚳馃嚘"`.

```js
assert.equal(getFlag("ZA"), "馃嚳馃嚘")
```

`getFlag("KR")` inapaswa kurudisha `"馃嚢馃嚪"`.

```js
assert.equal(getFlag("KR"), "馃嚢馃嚪")
```

`getFlag("ES")` inapaswa kurudisha `"馃嚜馃嚫"`.

```js
assert.equal(getFlag("ES"), "馃嚜馃嚫")
```

`getFlag("SE")` inapaswa kurudisha `"馃嚫馃嚜"`.

```js
assert.equal(getFlag("SE"), "馃嚫馃嚜")
```

`getFlag("CH")` inapaswa kurudisha `"馃嚚馃嚟"`.

```js
assert.equal(getFlag("CH"), "馃嚚馃嚟")
```

`getFlag("TH")` inapaswa kurudisha `"馃嚬馃嚟"`.

```js
assert.equal(getFlag("TH"), "馃嚬馃嚟")
```

`getFlag("TT")` inapaswa kurudisha `"馃嚬馃嚬"`.

```js
assert.equal(getFlag("TT"), "馃嚬馃嚬")
```

`getFlag("TR")` inapaswa kurudisha `"馃嚬馃嚪"`.

```js
assert.equal(getFlag("TR"), "馃嚬馃嚪")
```

`getFlag("UA")` inapaswa kurudisha `"馃嚭馃嚘"`.

```js
assert.equal(getFlag("UA"), "馃嚭馃嚘")
```

`getFlag("AE")` inapaswa kurudisha `"馃嚘馃嚜"`.

```js
assert.equal(getFlag("AE"), "馃嚘馃嚜")
```

`getFlag("US")` inapaswa kurudisha `"馃嚭馃嚫"`.

```js
assert.equal(getFlag("US"), "馃嚭馃嚫")
```

`getFlag("UY")` inapaswa kurudisha `"馃嚭馃嚲"`.

```js
assert.equal(getFlag("UY"), "馃嚭馃嚲")
```

`getFlag("UZ")` inapaswa kurudisha `"馃嚭馃嚳"`.

```js
assert.equal(getFlag("UZ"), "馃嚭馃嚳")
```

`getFlag("VE")` inapaswa kurudisha `"馃嚮馃嚜"`.

```js
assert.equal(getFlag("VE"), "馃嚮馃嚜")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
