---
id: 697a49e6ff50d756c9b69365
title: "Zoezi la 188: Michezo ya Majira ya Baridi 2026 Siku ya 9: Skeleton"
challengeType: 28
dashedName: challenge-188
---

# --description--

Kutolewa mfuatano wa herufi unaowakilisha mizunguko kwenye njia ya skeleton, amua ugumu wa njia hiyo.

- Mfuatano wa herufi uliotolewa utakuwa na herufi tu:
  - `"L"` kwa mzunguko wa kushoto
  - `"R"` kwa mzunguko wa kulia
  - `"S"` kwa sehemu ya moja kwa moja

- Kila mabadiliko ya mwelekeo huongeza alama 15 (`"L"` ikifuatiwa na `"R"` au kinyume chake).
- Mizunguko mingine yote huongeza alama 5 kila moja (herufi nyingine zote za `"L"` au `"R"`).
- Sehemu za moja kwa moja haziongezi alama.

Ugumu wa njia unategemea jumla ya alama. Rudisha:

- `"Easy"` ikiwa jumla ni 0 - 100
- `"Medium"` ikiwa jumla ni 101-200
- `"Hard"` ikiwa jumla ni zaidi ya 200

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")` inapaswa kurudisha `"Easy"`.

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` inapaswa kurudisha `"Hard"`.

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")` inapaswa kurudisha `"Medium"`.

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` inapaswa kurudisha `"Hard"`.

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")` inapaswa kurudisha `"Medium"`.

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` inapaswa kurudisha `"Easy"`.

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
