---
id: 698a1a73ade5ac0e19180fa7
title: "Zoezi la 204: Jumlisha herufi"
challengeType: 28
dashedName: challenge-204
---

# --description--

Ukitolewa mfuatano wa herufi, rudisha jumla ya herufi zake.

- Herufi ni A-Z kwa herufi kubwa au herufi ndogo
- Thamani za herufi ni: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26
- Herufi kubwa na herufi ndogo zina thamani sawa.
- Puuza herufi zisizo herufi zote.

# --hints--

`sumLetters("Hello")` inapaswa kurudisha `52`.

```js
assert.equal(sumLetters("Hello"), 52);
```

`sumLetters("freeCodeCamp")` inapaswa kurudisha `94`.

```js
assert.equal(sumLetters("freeCodeCamp"), 94);
```

`sumLetters("The quick brown fox jumps over the lazy dog.")` inapaswa kurudisha `473`.

```js
assert.equal(sumLetters("The quick brown fox jumps over the lazy dog."), 473);
```

`sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` inapaswa kurudisha `1681`.

```js
assert.equal(sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681);
```

`sumLetters("</404>")` inapaswa kurudisha `0`.

```js
assert.equal(sumLetters("</404>"), 0);
```

# --seed--

## --seed-contents--

```js
function sumLetters(str) {

  return str;
}
```

# --solutions--

```js
function sumLetters(str) {
  let total = 0;
  for (let char of str) {
    let upper = char.toUpperCase();
    if (upper >= "A" && upper <= "Z") {
      total += upper.charCodeAt(0) - "A".charCodeAt(0) + 1;
    }
  }
  return total;
}
```
