---
id: 698a1a73ade5ac0e19180fa6
title: "Zoezi la 203: Safu ya data isiyo na safu ndogo"
challengeType: 28
dashedName: challenge-203
---

# --description--

Ikiwa umepewa safu ya data, thibitisha kama ni safu isiyo na safu ndogo.

- Safu ya data ni isiyo na safu ndogo ikiwa hakuna kipengele chake chochote kilicho safu ya data.

# --hints--

`isFlat([1, 2, 3, 4])` inapaswa kurudisha `true`.

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])` inapaswa kurudisha `false`.

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])` inapaswa kurudisha `true`.

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])` inapaswa kurudisha `false`.

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])` inapaswa kurudisha `false`.

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```
