---
id: 698a1a863194f1f4e63f645e
title: "Zoezi la 207: Pengo Ndogo Zaidi"
challengeType: 28
dashedName: challenge-207
---

# --description--

Ukitolewa mfuatano wa herufi, rudisha sehemu ndogo ya mfuatano kati ya herufi mbili zinazofanana ambazo zina pengo ndogo zaidi ya herufi kati yao (pengo ndogo zaidi).

- Daima kutakuwa na angalau jozi moja ya herufi zinazolingana.
- Sehemu ndogo ya mfuatano inayorudishwa haipaswi kujumuisha herufi zinazolingana.
- Ikiwa pengo mbili au zaidi ni za urefu sawa, rudisha herufi kutoka kwa pengo la kwanza.

Kwa mfano, ukitolewa `"ABCDAC"`, rudisha `"DA"` kwa sababu:

- Ni `"A"` na `"C"` pekee zinazorudiwa katika mfuatano.
- Idadi ya herufi kati ya herufi mbili za `"A"` ni 3, na kati ya herufi za `"C"` ni 2.
- Hivyo rudisha mfuatano kati ya herufi mbili za `"C"`.

# --hints--

`smallestGap("ABCDAC")` inapaswa kurudisha `"DA"`.

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` inapaswa kurudisha `"e"`.

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` inapaswa kurudisha `"#q6e&rkf(p"`.

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` inapaswa kurudisha `""`.

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` inapaswa kurudisha `"fox"`.

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
