---
id: 698a1a73ade5ac0e19180fa8
title: "Zoezi la 205: Hesabu ya Kyuubu Kamili"
challengeType: 28
dashedName: challenge-205
---

# --description--

Ukizingatia nambari kamili mbili, amua ni kyuubu kamili ngapi zipo katika safu kati ya na pamoja na nambari hizo mbili.

- Nambari ya chini haijahakikishiwa kuwa hoja ya kwanza.
- Nambari ni kyuubu kamili ikiwa ipo nambari kamili (`n`) ambapo `n * n * n = number`. Kwa mfano, 27 ni kyuubu kamili kwa sababu `3 * 3 * 3 = 27`.

# --hints--

`countPerfectCubes(3, 30)` inapaswa kurudisha `2`.

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)` inapaswa kurudisha `3`.

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)` inapaswa kurudisha `4`.

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)` inapaswa kurudisha `9`.

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)` inapaswa kurudisha `41`.

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
