---
id: 698a1a73ade5ac0e19180fa9
title: "Zoezi la 206: Thamani za Kadi za Mchezo"
challengeType: 28
dashedName: challenge-206
---

# --description--

Ukitoa safu ya data ya kadi za mchezo, rudisha safu mpya ya data yenye thamani ya nambari ya kila kadi.

Thamani za kadi:

- Ace (`"A"`) ina thamani ya 1.
- Kadi zilizo na nambari (`"2"` - `"10"`) zina thamani sawa na nambari zilizo mbele yao: 2 - 10, mtawalia.
- Kadi za uso: Jack (`"J"`), Queen (`"Q"`), na King (`"K"`) kila moja ina thamani ya 10.

Aina za kadi:

- Kila kadi ina aina: Spades (`"S"`), Clubs (`"C"`), Diamonds (`"D"`), au Hearts (`"H"`).

Muundo wa kadi:

- Kila kadi inawakilishwa kama mfuatano wa herufi: `"valueSuit"`. Kwa mfano: `"AS"` ni Ace wa Spades, `"10H"` ni Kumi wa Hearts, na `"QC"` ni Queen wa Clubs.

# --hints--

`cardValues(["3H", "4D", "5S"])` inapaswa kurudisha `[3, 4, 5]`.

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` inapaswa kurudisha `[1, 10, 10, 6, 7]`.

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` inapaswa kurudisha `[8, 10, 2, 10, 9]`.

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` inapaswa kurudisha `[1, 10]`.

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` inapaswa kurudisha `[10, 10, 10, 10, 1]`.

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
