---
id: 6994cff2290543b3aec9f50f
title: "Zoezi la 210: Kagua halali ya HSL"
challengeType: 28
dashedName: challenge-210
---

# --description--

Kutolewa mfuatano wa herufi, amua kama ni thamani halali ya rangi ya CSS `hsl()`.

- Thamani halali ya HSL lazima iwe katika muundo `"hsl(h, s%, l%)"`, ambapo:
  - `h` (hue) lazima iwe nambari kati ya 0 na 360 (ikiwa ni pamoja).
  - `s` (saturation) lazima iwe asilimia kati ya 0% na 100%.
  - `l` (lightness) lazima iwe asilimia kati ya 0% na 100%.

- Nafasi zinaruhusiwa tu:
  - Baada ya mabano ya kufungua
  - Kabla na/au baada ya koma
  - Kabla na/au baada ya mabano ya kufunga

- Hiari, thamani inaweza kumalizika na nusu nukta (`";"`).

Kwa mfano, `"hsl(240, 50%, 50%)"` ni thamani halali ya HSL.

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
