---
id: 6994cff2290543b3aec9f510
title: "Zoezi la 211: Jumla ya safu ya data"
challengeType: 28
dashedName: challenge-211
---

# --description--

Ukipata safu ya data ya nambari, rudisha jumla ya nambari zote.

# --hints--

`sumArray([1, 2, 3, 4, 5])` inapaswa kurudisha `15`.

```js
assert.equal(sumArray([1, 2, 3, 4, 5]), 15);
```

`sumArray([42])` inapaswa kurudisha `42`.

```js
assert.equal(sumArray([42]), 42);
```

`sumArray([5, -2, 7, -3])` inapaswa kurudisha `7`.

```js
assert.equal(sumArray([5, -2, 7, -3]), 7);
```

`sumArray([203, 145, -129, 6293, 523, -919, 845, 2434])` inapaswa kurudisha `9395`.

```js
assert.equal(sumArray([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395);
```

`sumArray([0, 0])` inapaswa kurudisha `0`.

```js
assert.equal(sumArray([0, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function sumArray(numbers) {

  return numbers;
}
```

# --solutions--

```js
function sumArray(numbers) {

  return numbers.reduce((a, c) => c + a, 0);
}
```
