---
id: 699c8e045ee7cb94ed2322da
title: "Zoezi la 219: Vikumbusho vya Miaka ya Harusi"
challengeType: 28
dashedName: challenge-219
---

# --description--

Ikiwa umepewa nambari kamili inayowakilisha idadi ya miaka mume na mke wamekuwa katika ndoa, rudisha kumbu kumbu yao ya hivi karibuni ya maadhimisho kulingana na chati hii:

| Miaka ya Ndoa | Kumbu kumbu  |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- Ikiwa hawajafikia kumbu kumbu ya kwanza, rudisha `"Newlyweds"`.

# --hints--

`getMilestone(0)` inapaswa kurudisha `"Newlyweds"`.

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` inapaswa kurudisha `"Paper"`.

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` inapaswa kurudisha `"Wood"`.

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` inapaswa kurudisha `"Tin"`.

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` inapaswa kurudisha `"Silver"`.

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` inapaswa kurudisha `"Ruby"`.

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` inapaswa kurudisha `"Gold"`.

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` inapaswa kurudisha `"Diamond"`.

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` inapaswa kurudisha `"Platinum"`.

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
