---
id: 699c8e045ee7cb94ed2322d5
title: "Zoezi la 214: Kagua mnyororo wa domino"
challengeType: 28
dashedName: challenge-214
---

# --description--

Ikiwa umepewa safu ya data ya mfululizo wa domino, amua kama inaunda mnyororo halali.

- Kila kipengele katika safu ya data kinawakilisha domino na kitakuwa safu ya nambari mbili kutoka 1 hadi 6, (ikiwa ni pamoja).
- Ili mnyororo uwe halali, nambari ya pili ya kila domino lazima iendane na nambari ya kwanza ya domino inayofuata.
- Nambari ya kwanza ya domino ya kwanza na nambari ya mwisho ya domino ya mwisho hazihitaji kuendana na kitu chochote.

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```
