---
id: 699c8e045ee7cb94ed2322d4
title: "Zoezi la 213: Mbadalishaji wa urefu wa neno"
challengeType: 28
dashedName: challenge-213
---

# --description--

Ukitoa mfuatano wa herufi wa maneno, rudisha mfuatano mpya ambapo kila neno limebadilishwa na urefu wake.

- Maneno katika mfuatano uliotolewa yataachwa na nafasi moja kati yao
- Hifadhi nafasi hizo katika mfuatano uliorudishwa.

Kwa mfano, ukitoa `"hello world"`, rudisha `"5 5"`.

# --hints--

`convertWords("hello world")` inapaswa kurudisha `"5 5"`.

```js
assert.equal(convertWords("hello world"), "5 5");
```

`convertWords("Thanks and happy coding")` inapaswa kurudisha `"6 3 5 6"`.

```js
assert.equal(convertWords("Thanks and happy coding"), "6 3 5 6");
```

`convertWords("The quick brown fox jumps over the lazy dog")` inapaswa kurudisha `"3 5 5 3 5 4 3 4 3"`.

```js
assert.equal(convertWords("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3");
```

`convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` inapaswa kurudisha `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`.

```js
assert.equal(convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4");
```

# --seed--

## --seed-contents--

```js
function convertWords(str) {

  return str;
}
```

# --solutions--

```js
function convertWords(str) {
  return str
    .split(" ")
    .map(word => word.length)
    .join(" ");
}
```
