---
id: 699c8e045ee7cb94ed2322d6
title: "Zoezi la 215: Kihesabu cha Ada ya Maegesho"
challengeType: 28
dashedName: challenge-215
---

# --description--

Kutolewa mfuatano wa herufi miwili unaowakilisha wakati uliweka gari lako maegesho na wakati uliuchukua, hesabu ada ya maegesho.

- Mfuatano wa herufi uliotolewa utakuwa katika muundo `"HH:MM"` ukitumia saa 24. Kwa mfano `"14:00"` ni saa 2 asubuhi.
- Mfuatano wa kwanza utakuwa wakati uliweka gari lako maegesho, na wa pili utakuwa wakati uliuchukua.
- Ikiwa wakati wa kuchukua ni mapema kuliko wakati wa kuingia, ina maana maegesho yalidumu hadi baada ya saa sita usiku hadi siku inayofuata.

Sheria za ada:

- Kila saa ya maegesho inagharimu $3.
- Saa zisizokamilika zinarundikwa hadi saa kamili inayofuata.
- Ikiwa maegesho yalidumu usiku kucha (baada ya saa sita usiku), ada ya ziada ya $10 ya usiku inatozwa.
- Kuna ada ya kiwango cha chini ya $5 (inatumika tu ikiwa jumla ingekuwa chini ya $5).

Rudisha gharama jumla katika muundo `"$cost"`, `"$5"` kwa mfano.

# --hints--

`calculateParkingFee("09:00", "11:00")` inapaswa kurudisha `"$6"`.

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` inapaswa kurudisha `"$5"`.

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` inapaswa kurudisha `"$9"`.

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` inapaswa kurudisha `"$27"`.

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` inapaswa kurudisha `"$34"`.

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` inapaswa kurudisha `"$82"`.

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
