---
id: 699c8e045ee7cb94ed2322d8
title: "Zoezi la 217: Vipande vya chess vilivyokamatwa"
challengeType: 28
dashedName: challenge-217
---

# --description--

Kutolewa safu ya data ya mfuatano wa herufi zinazowakilisha vipande vya chess ulivyo navyo bado kwenye ubao, hesabu thamani ya vipande ambavyo mpinzani wako amevikamata.

Katika chess, unaanza na vipande 16:

| Kipande | Kifupisho      | Idadi    | Thamani |
| ------- | -------------- | -------- | ------- |
| Pawn    | `"P"`      | 8        | 1       |
| Rook    | `"R"`      | 2        | 5       |
| Knight  | `"N"`      | 2        | 3       |
| Bishop  | `"B"`      | 2        | 3       |
| Queen   | `"Q"`      | 1        | 9       |
| King    | `"K"`      | 1        | 0       |

- Safu ya data iliyotolewa itakuwa na vifupisho vilivyotajwa hapo juu tu.
- Vipande vyovyote kati ya 16 ambavyo havipo kwenye safu ya data iliyotolewa vimekamatwa.
- Rudisha jumla ya thamani ya vipande vyote vilivyokamatwa, isipokuwa...
- Ikiwa Mfalme amekamatwa, rudisha `"Checkmate"`.

# --hints--

`getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` inapaswa kurudisha `8`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8);
```

`getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"])` inapaswa kurudisha `26`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"]), 26);
```

`getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` inapaswa kurudisha `16`.

```js
assert.equal(getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16);
```

`getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` inapaswa kurudisha `4`.

```js
assert.equal(getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4);
```

`getCapturedValue(["P", "K"])` inapaswa kurudisha `38`.

```js
assert.equal(getCapturedValue(["P", "K"]), 38);
```

`getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` inapaswa kurudisha `0`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0);
```

`getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` inapaswa kurudisha `"Checkmate"`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate");
```

# --seed--

## --seed-contents--

```js
function getCapturedValue(pieces) {

  return pieces;
}
```

# --solutions--

```js
function getCapturedValue(pieces) {
  if (!pieces.includes("K")) return "Checkmate";

  const values = { P:1, R:5, N:3, B:3, Q:9, K:0 };
  let remaining = "PPPPPPPPRRNNBBQK";

  for (const p of pieces) {
    remaining = remaining.replace(p, "");
  }

  return [...remaining].reduce((sum, p) => sum + values[p], 0);
}
```
