---
id: 69a890af247de743333bd4d1
title: "Zoezi la 228: Usiku wa Filamu"
challengeType: 28
dashedName: challenge-228
---

# --description--

Ukipata mfuatano wa herufi unaoonyesha siku ya wiki, mfuatano mwingine wa herufi unaoonyesha saa ya kuonyesha filamu, na nambari kamili ya tiketi, rudisha jumla ya gharama ya tiketi za filamu kwa onyesho hilo.

Siku iliyotolewa itakuwa moja kati ya:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

Saa ya kuonyesha itatolewa kwa muundo `"H:MMam"` au `"H:MMpm"`. Kwa mfano `"10:00am"` au `"10:00pm"`.

Rudisha jumla ya gharama kwa muundo `"$D.CC"` ukitumia sheria hizi:

- Mwisho wa wiki (Ijumaa - Jumapili): $12.00 kwa tiketi moja.
- Siku za wiki (Jumatatu - Alhamisi): $10.00 kwa tiketi moja.
- Onyesho la mchana (kabla ya saa 11 jioni): pungua $2.00 kwa tiketi moja (isipokuwa siku ya Jumanne).
- Jumatano: tiketi zote ni $5.00 kila moja.

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` inapaswa kurudisha `"$12.00"`.

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` inapaswa kurudisha `"$10.00"`.

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` inapaswa kurudisha `"$10.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` inapaswa kurudisha `"$30.00"`.

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` inapaswa kurudisha `"$32.00"`.

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` inapaswa kurudisha `"$50.00"`.

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` inapaswa kurudisha `"$5.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
