---
id: 69a890af247de743333bd4cc
title: "Zoezi la 223: Mfasiri wa QR"
challengeType: 28
dashedName: challenge-223
---

# --description--

Ikiwa umepewa matriki ya 6x6 (safu ya data ya mfuatano wa herufi), inayowakilisha msimbo wa QR, rudisha mfuatano wa data ya binary ndani ya msimbo huo.

- Msimbo wa QR unaweza kupewa kwa mzunguko wowote wa nyongeza za digrii 90.
- Msimbo ulio katika mwelekeo sahihi una kikundi cha `1` cha 2x2 (alama za mwelekeo) katika kona za chini-kushoto, juu-kushoto, na juu-kulia.
- Alama tatu za mwelekeo za 2x2 si sehemu ya data ya binary.
- Data ya binary husomwa kutoka kushoto kwenda kulia, juu hadi chini (kama kitabu) wakati msimbo wa QR uko katika mwelekeo sahihi.
- Msimbo utakuwa na mwelekeo halali mmoja tu.

Kwa mfano, ukipata:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

au ukipata msimbo uleule katika mwelekeo tofauti:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

Rudisha `"000000000000000000000001"`, data yote ya binary isipokuwa alama tatu za mwelekeo za 2x2.

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])` inapaswa kurudisha `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"])` inapaswa kurudisha `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])` inapaswa kurudisha `"001101000011000000110100"`.

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])` inapaswa kurudisha `"010001000100010101010110"`.

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])` inapaswa kurudisha `"010000100100100101001110"`.

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```
