---
id: 69a890af247de743333bd4d0
title: "Zoezi la 227: Muda wa Kupumzika"
challengeType: 28
dashedName: challenge-227
---

# --description--

Ukizingatia alama mbili za muda, ya kwanza ikiwakilisha wakati mtumizi alimaliza mtihani, na ya pili ikiwakilisha wakati wa sasa, amua kama mtumizi anaweza kuchukua mtihani tena.

- Alama zote mbili za muda zitatolewa kwa muundo: `"YYYY-MM-DDTHH:MM:SS"`, kwa mfano `"2026-03-25T14:00:00"`. Kumbuka kuwa wakati ni saa 24.
- Mtumizi lazima asubiri angalau masaa 48 kabla ya kuchukua mtihani tena.

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` inapaswa kurudisha `true`.

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` inapaswa kurudisha `false`.

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` inapaswa kurudisha `true`.

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` inapaswa kurudisha `false`.

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
