---
id: 69b1028d6e265413d0198a2a
title: "Zoezi la 231: Kagua ISBN-10"
challengeType: 28
dashedName: challenge-231
---

# --description--

Kutolewa mfuatano wa herufi, amua kama ni ISBN-10 halali.

ISBN-10 ina viunganishi vya mkato (hyphens) (`"-"`) na herufi 10 nyingine. Baada ya kuondoa viunganishi vya mkato (`"-"`):

- Herufi 9 za kwanza lazima ziwe nambari, na
- Herufi ya mwisho inaweza kuwa nambari au herufi `"X"`, inayowakilisha nambari 10.

Ili kuikagua:

- Zidisha kila nambari (au thamani) kwa nafasi yake (zidisha nambari ya kwanza kwa 1, ya pili kwa 2, na kadhalika).
- Jumlisha matokeo yote.
- Ikiwa jumla inaweza kugawika kwa 11, basi ni halali.

# --hints--

`isValidIsbn10("0-306-40615-2")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
