---
id: 69b1028d6e265413d0198a2f
title: "Zoezi la 236: Historia ya kivinjari"
challengeType: 28
dashedName: challenge-236
---

# --description--

Ukizingatia safu ya data ya amri za kivinjari, rudisha safu ya data yenye maadili mawili: historia kama safu ya URL, na kielezo cha ukurasa wa sasa.

Amri halali ni:

- `"URL"` - Ambapo URL ni anwani ya mtandao (kwa mfano `"freecodecamp.org"`). Inasogeza kwenye URL iliyotolewa, inaongeza kwenye historia katika nafasi inayofuata, na inatupa historia yoyote ya mbele.
- `"Back"` - inasogeza kwenye ukurasa uliopita katika historia, au inabaki kwenye ukurasa wa sasa ikiwa hakuna mwingine.
- `"Forward"` - inasogeza kwenye ukurasa unaofuata katika historia, au inabaki kwenye ukurasa wa sasa ikiwa hakuna mwingine.

Kwa mfano, ukipokea `["freecodecamp.org", "freecodecamp.org/learn", "Back"]`, rudisha `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

# --hints--

`getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` inapaswa kurudisha `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

```js
assert.deepEqual(getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` inapaswa kurudisha `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact",  "example.com/blog", "Back", "Back", "Forward"])` inapaswa kurudisha `[["example.com", "example.com/contact", "example.com/blog"], 1]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` inapaswa kurudisha `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "Back"])` inapaswa kurudisha `[["example.com", "example.com/about"], 0]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "Forward"])` inapaswa kurudisha `[["example.com", "example.com/about"], 1]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1]);
```

# --seed--

## --seed-contents--

```js
function getBrowserHistory(commands) {

  return commands;
}
```

# --solutions--

```js
function getBrowserHistory(commands) {
  const stack = [];
  let index = -1;

  for (const command of commands) {
    if (command === "Back") {
      if (index > 0) index--;
    } else if (command === "Forward") {
      if (index < stack.length - 1) index++;
    } else {
      stack.splice(index + 1);
      stack.push(command);
      index++;
    }
  }

  return [stack, index];
}
```
