---
id: 69b1028d6e265413d0198a2e
title: "Zoezi la 235: Fibonacci yenye herufi kubwa"
challengeType: 28
dashedName: challenge-235
---

# --description--

Ukitoa mfuatano wa herufi, rudi na mfuatano mpya ambapo kila herufi imeandikwa kwa herufi kubwa ikiwa kielezo chake ni nambari ya Fibonacci, na kwa herufi ndogo vinginevyo.

Mfululizo wa Fibonacci ni mfululizo wa nambari ambapo kila nambari ni jumla ya nambari mbili zilizotangulia. Nambari kumi za kwanza katika mfululizo ni `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

- Herufi ya kwanza iko kwenye kielezo `0`.
- Ikiwa kielezo cha herufi zisizo za herufi ni nambari ya Fibonacci, ziache bila kubadilika.

# --hints--

`capitalizeFibonacci("hello world")` inapaswa kurudisha `"HELLo woRld"`.

```js
assert.equal(capitalizeFibonacci("hello world"), "HELLo woRld");
```

`capitalizeFibonacci("HELLO WORLD")` inapaswa kurudisha `"HELLo woRld"`.

```js
assert.equal(capitalizeFibonacci("HELLO WORLD"), "HELLo woRld");
```

`capitalizeFibonacci("hello, world!")` inapaswa kurudisha `"HELLo, wOrld!"`.

```js
assert.equal(capitalizeFibonacci("hello, world!"), "HELLo, wOrld!");
```

`capitalizeFibonacci("The quick brown fox jumped over the lazy dog.")` inapaswa kurudisha `"THE qUicK broWn fox jUmped over thE lazy dog."`.

```js
assert.equal(capitalizeFibonacci("The quick brown fox jumped over the lazy dog."), "THE qUicK broWn fox jUmped over thE lazy dog.");
```

`capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl.")` inapaswa kurudisha `"LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl."`.

```js
assert.equal(capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl."), "LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl.");
```

# --seed--

## --seed-contents--

```js
function capitalizeFibonacci(str) {

  return str;
}
```

# --solutions--

```js
function capitalizeFibonacci(str) {
  const fibs = new Set([0, 1]);
  let a = 0, b = 1;
  while (b < str.length) {
    [a, b] = [b, a + b];
    fibs.add(b);
  }

  return str
    .split("")
    .map((char, i) => {
      if (!/[a-zA-Z]/.test(char)) return char;
      return fibs.has(i) ? char.toUpperCase() : char.toLowerCase();
    })
    .join("");
}
```
