---
id: 69b5b2be76ec8135a7fbe974
title: "Zoezi la 245: Safu ya matrix ya mviringo"
challengeType: 28
dashedName: challenge-245
---

# --description--

Ikiwa umepewa matrix ya 2D, rudisha safu ya data tambarare yenye thamani zake zote kwa mpangilio wa kuzunguka saa.

Safu ya data itakayorejeshwa inapaswa kuwa na thamani ya juu-kushoto kwanza, kisha kusogea kulia kwenye safu ya juu, kisha chini kwenye safu ya kulia, kisha kushoto kwenye safu ya chini, kisha juu kwenye safu ya kushoto. Rudia mchakato huu kuelekea ndani kwa tabaka zozote zilizobaki.

Kwa mfano, ukipata:

```json
[
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9]
]
```

Rudisha `[1, 2, 3, 6, 9, 8, 7, 4, 5]`.

# --hints--

`spiralMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` inapaswa kurudisha `[1, 2, 3, 6, 9, 8, 7, 4, 5]`.

```js
assert.deepEqual(spiralMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [1, 2, 3, 6, 9, 8, 7, 4, 5]);
```

`spiralMatrix([["a", "b", "c", "d"], ["l", "m", "n", "e"], ["k", "p", "o", "f"], ["j", "i", "h", "g"]])` inapaswa kurudisha `["a", "b", "c", "d", "e", "f", "g", "h", "i", "j", "k", "l", "m", "n", "o", "p"]`.

```js
assert.deepEqual(spiralMatrix([["a", "b", "c", "d"], ["l", "m", "n", "e"], ["k", "p", "o", "f"], ["j", "i", "h", "g"]]), ["a", "b", "c", "d", "e", "f", "g", "h", "i", "j", "k", "l", "m", "n", "o", "p"]);
```

`spiralMatrix([[true, false, false], [false, true, true], [false, true, false], [true, true, false]])` inapaswa kurudisha `[true, false, false, true, false, false, true, true, false, false, true, true]`.

```js
assert.deepEqual(spiralMatrix([[true, false, false], [false, true, true], [false, true, false], [true, true, false]]), [true, false, false, true, false, false, true, true, false, false, true, true]);
```

`spiralMatrix([[25, 24, 23, 22, 21], [10, 9, 8, 7, 20], [11, 2, 1, 6, 19], [12, 3, 4, 5, 18], [13, 14, 15, 16, 17]])` inapaswa kurudisha `[25, 24, 23, 22, 21, 20, 19, 18, 17, 16, 15, 14, 13, 12, 11, 10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`.

```js
assert.deepEqual(spiralMatrix([[25, 24, 23, 22, 21], [10, 9, 8, 7, 20], [11, 2, 1, 6, 19], [12, 3, 4, 5, 18], [13, 14, 15, 16, 17]]), [25, 24, 23, 22, 21, 20, 19, 18, 17, 16, 15, 14, 13, 12, 11, 10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js
function spiralMatrix(matrix) {

    return matrix
}
```

# --solutions--

```js
function spiralMatrix(matrix) {
  const result = [];
  let top = 0, bottom = matrix.length - 1;
  let left = 0, right = matrix[0].length - 1;

  while (top <= bottom && left <= right) {
    for (let i = left; i <= right; i++) result.push(matrix[top][i]);
    top++;

    for (let i = top; i <= bottom; i++) result.push(matrix[i][right]);
    right--;

    if (top <= bottom) {
      for (let i = right; i >= left; i--) result.push(matrix[bottom][i]);
      bottom--;
    }

    if (left <= right) {
      for (let i = bottom; i >= top; i--) result.push(matrix[i][left]);
      left++;
    }
  }

  return result;
}
```
