---
id: 69b5b2be76ec8135a7fbe973
title: "Zoezi la 244: Shambulio la Rook na Bishop"
challengeType: 28
dashedName: challenge-244
---

# --description--

Ukipatiwa mfuatano wa herufi unaoonyesha mahali pa rook kwenye ubao wa chess, na mwingine unaoonyesha mahali pa bishop, amua kama kipande kimoja kinaweza kushambulia kingine.

Ubao wa chess wa kawaida ni wa 8x8, na safu zimepewa lebo `A` hadi `H` (kutoka kushoto kwenda kulia) na mistari imepewa lebo `1` hadi `8` (kutoka chini kwenda juu). Inaonekana hivi:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- Rooks wanaweza kusogea mstatili wowote kwa mwelekeo wa usawa au wima.
- Bishops wanaweza kusogea mstatili wowote kwa mwelekeo wa mwinuko.
- Kipande kimoja kinaweza kushambulia kingine ikiwa kinaweza kusogea hadi mahali pa kipande hicho.

Rudisha:

- `"rook"` ikiwa rook anaweza kushambulia bishop.
- `"bishop"` ikiwa bishop anaweza kushambulia rook.
- `"neither"` ikiwa hakuna kipande kinachoweza kushambulia kingine.

# --hints--

`rookBishopAttack("A1", "A5")` inapaswa kurudisha `"rook"`.

```js
assert.equal(rookBishopAttack("A1", "A5"), "rook");
```

`rookBishopAttack("C3", "F6")` inapaswa kurudisha `"bishop"`.

```js
assert.equal(rookBishopAttack("C3", "F6"), "bishop");
```

`rookBishopAttack("D4", "D7")` inapaswa kurudisha `"rook"`.

```js
assert.equal(rookBishopAttack("D4", "D7"), "rook");
```

`rookBishopAttack("B7", "H1")` inapaswa kurudisha `"bishop"`.

```js
assert.equal(rookBishopAttack("B7", "H1"), "bishop");
```

`rookBishopAttack("B3", "C5")` inapaswa kurudisha `"neither"`.

```js
assert.equal(rookBishopAttack("B3", "C5"), "neither");
```

`rookBishopAttack("G3", "E8")` inapaswa kurudisha `"neither"`.

```js
assert.equal(rookBishopAttack("G3", "E8"), "neither");
```

# --seed--

## --seed-contents--

```js
function rookBishopAttack(rook, bishop) {

  return rook;
}
```

# --solutions--

```js
function rookBishopAttack(rook, bishop) {
  const colDiff = Math.abs(rook.charCodeAt(0) - bishop.charCodeAt(0));
  const rowDiff = Math.abs(rook[1] - bishop[1]);

  if (rook[0] === bishop[0] || rook[1] === bishop[1]) return "rook";
  if (colDiff === rowDiff) return "bishop";
  return "neither";
}
```
