---
id: 69b559d2903b9e4afe9075f7
title: "Zoezi la 237: Uthibitisho wa Usemi wa Kihesabu"
challengeType: 28
dashedName: challenge-237
---

# --description--

Kutolewa mfuatano wa herufi unaowakilisha usemi wa kihesabu, amua kama ni sahihi.

- Upande wa kushoto unaweza kuwa na hadi nambari kamili chanya tatu na viongezaji `+`, `-`, `*`, na `/`.
- Usemi wa kihesabu utatolewa kwa muundo: `"number operator number = number"` (ikiwa na nambari mbili au tatu upande wa kushoto). Kwa mfano: `"2 + 2 = 4"` au `"2 + 3 - 1 = 4"`.
- Upande wa kulia daima utakuwa nambari kamili moja.

Fuata mpangilio wa kawaida wa shughuli: kuzidisha na kugawanya hufanyiwa tathmini kabla ya kuongeza na kutoa, kutoka kushoto kwenda kulia.

# --hints--

`isValidEquation("2 + 2 = 4")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidEquation("2 + 2 = 4"));
```

`isValidEquation("2 + 3 - 1 = 4")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidEquation("2 + 3 - 1 = 4"));
```

`isValidEquation("8 / 2 = 4")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidEquation("8 / 2 = 4"));
```

`isValidEquation("10 * 5 = 50")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidEquation("10 * 5 = 50"));
```

`isValidEquation("2 - 2 = 0")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidEquation("2 - 2 = 0"));
```

`isValidEquation("2 + 9 / 3 = 5")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidEquation("2 + 9 / 3 = 5"));
```

`isValidEquation("20 - 2 * 3 = 14")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidEquation("20 - 2 * 3 = 14"));
```

`isValidEquation("2 + 5 = 6")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidEquation("2 + 5 = 6"));
```

`isValidEquation("10 - 2 * 3 = 24")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidEquation("10 - 2 * 3 = 24"));
```

`isValidEquation("3 + 9 / 3 = 4")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidEquation("3 + 9 / 3 = 4"));
```

# --seed--

## --seed-contents--

```js
function isValidEquation(equation) {

  return equation;
}
```

# --solutions--

```js
function isValidEquation(equation) {
  const [left, right] = equation.split(" = ");
  const tokens = left.split(" ");

  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    if (operator === "*" || operator === "/") {
      const leftOperand = Number(tokens[i - 1]);
      const rightOperand = Number(tokens[i + 1]);
      const computed = operator === "*" ? leftOperand * rightOperand : leftOperand / rightOperand;
      tokens.splice(i - 1, 3, String(computed));
      i -= 2;
    }
  }

  let result = Number(tokens[0]);
  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    const operand = Number(tokens[i + 1]);
    result = operator === "+" ? result + operand : result - operand;
  }

  return result === Number(right);
}
```
