---
id: 69bc6cb30c1d112a2e110a09
title: "Zoezi la 252: Mpenzi wa Ngazi wa Kipekee"
challengeType: 28
dashedName: challenge-252
---

# --description--

Ukizingatia idadi ya ngazi, rudisha ni kwa njia ngapi tofauti mtu anaweza kupanda kwa kuchukua hatua 1 au 2 kwa wakati mmoja.

# --hints--

`getUniqueClimbs(4)` inapaswa kurudisha `5`.

```js
assert.equal(getUniqueClimbs(4), 5);
```

`getUniqueClimbs(5)` inapaswa kurudisha `8`.

```js
assert.equal(getUniqueClimbs(5), 8);
```

`getUniqueClimbs(10)` inapaswa kurudisha `89`.

```js
assert.equal(getUniqueClimbs(10), 89);
```

`getUniqueClimbs(18)` inapaswa kurudisha `4181`.

```js
assert.equal(getUniqueClimbs(18), 4181);
```

`getUniqueClimbs(29)` inapaswa kurudisha `832040`.

```js
assert.equal(getUniqueClimbs(29), 832040);
```

`getUniqueClimbs(50)` inapaswa kurudisha `20365011074`.

```js
assert.equal(getUniqueClimbs(50), 20365011074);
```

# --seed--

## --seed-contents--

```js
function getUniqueClimbs(steps) {

  return steps;
}
```

# --solutions--

```js
function getUniqueClimbs(steps) {
  if (steps <= 0) return 0;
  if (steps === 1) return 1;
  if (steps === 2) return 2;

  let prev2 = 1, prev1 = 2;
  for (let i = 3; i <= steps; i++) {
    [prev2, prev1] = [prev1, prev2 + prev1];
  }
  return prev1;
}
```
