---
id: 69bc6cb30c1d112a2e110a08
title: "Zoezi la 251: Tafuta Jumla ya Safu ya Data"
challengeType: 28
dashedName: challenge-251
---

# --description--

Ukipatiwa safu ya data ya nambari na nambari lengwa, rudisha seti ndogo ya kwanza yenye nambari mbili au zaidi ambazo jumla yake ni sawa na nambari lengwa.

- Seti ndogo ya "kwanza" ni ile ambayo vipengele vyake vina kielezo cha chini kabisa kinachowezekana, ukizingatia kielezo cha mapema zaidi kwanza.
- Kila nambari katika safu ya data inaweza kutumika mara moja tu.
- Ikiwa hakuna seti ndogo halali, rudisha `"Sum not found"`.

Rudisha nambari zinazolingana kama safu ya data kwa mpangilio wa kuonekana kwao katika safu ya data ya awali.

# --hints--

`findSum([1, 3, 5, 7], 6)` inapaswa kurudisha `[1, 5]`.

```js
assert.deepEqual(findSum([1, 3, 5, 7], 6), [1, 5]);
```

`findSum([1, 2, 3, 4, 5], 5)` inapaswa kurudisha `[1, 4]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 5), [1, 4]);
```

`findSum([1, 2, 3, 4, 5], 6)` inapaswa kurudisha `[1, 2, 3]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 6), [1, 2, 3]);
```

`findSum([-1, -2, 3, 4], 1)` inapaswa kurudisha `[-1, -2, 4]`.

```js
assert.deepEqual(findSum([-1, -2, 3, 4], 1), [-1, -2, 4]);
```

`findSum([3, 1, 4, 1, 5, 9, 2, 6], 10)` inapaswa kurudisha `[3, 1, 4, 2]`.

```js
assert.deepEqual(findSum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2]);
```

`findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` inapaswa kurudisha `[1, 2, 3, 5, 9]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9]);
```

`findSum([7, 9, 4, 2, 5], 10)` inapaswa kurudisha `"Sum not found"`.

```js
assert.equal(findSum([7, 9, 4, 2, 5], 10), "Sum not found");
```

# --seed--

## --seed-contents--

```js
function findSum(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findSum(arr, target) {
  function search(start, remaining, current) {
    if (remaining === 0 && current.length >= 2) return current;
    if (start >= arr.length) return null;

    for (let i = start; i < arr.length; i++) {
      const result = search(i + 1, remaining - arr[i], [...current, arr[i]]);
      if (result) return result;
    }

    return null;
  }

  return search(0, target, []) ?? "Sum not found";
}
```
