---
id: 69bc6cb30c1d112a2e110a04
title: "Zoezi la 247: Herufi ya Mwisho"
challengeType: 28
dashedName: challenge-247
---

# --description--

Ukitolewa mfuatano wa herufi, rudisha herufi inayojitokeza mwisho zaidi katika alfabeti.

- Ikiwa herufi mbili au zaidi zinafanana kuwa za mwisho katika alfabeti, rudisha ile ya kwanza.
- Puuza herufi zisizo za alfabeti zote.

# --hints--

`getLastLetter("world")` inapaswa kurudisha `"w"`.

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` inapaswa kurudisha `"W"`.

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` inapaswa kurudisha `"z"`.

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` inapaswa kurudisha `"L"`.

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` inapaswa kurudisha `"T"`.

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
