---
id: 69bc6cb30c1d112a2e110a06
title: "Zoezi la 249: Hisabati ya mfuatano wa herufi"
challengeType: 28
dashedName: challenge-249
---

# --description--

Ikiwa umepewa mfuatano wa herufi wenye nambari na herufi nyingine, fanya hesabu kwa nambari kulingana na idadi ya herufi zisizo za nambari kati ya nambari hizo.

- Ikiwa idadi ya herufi zinazotenganisha nambari mbili ni nambari hata, tumia kuongeza.
- Ikiwa ni nambari isiyo ya hata, tumia kutoa.
- Nambari zinazofuata moja baada ya nyingine huunda nambari moja.
- Operesheni zinafanywa kutoka kushoto kwenda kulia.
- Puuza herufi za mwanzo na mwisho ambazo si nambari.

Kwa mfano, ukipata `"3ab10c8"`, rudisha `5`. Ongeza 3 na 10 kupata 13 kwa sababu kuna idadi ya herufi hata kati yao. Kisha toa 8 kutoka 13 kwa sababu kuna idadi ya herufi isiyo ya hata kati ya matokeo na 8.

# --hints--

`doMath("3ab10c8")` inapaswa kurudisha `5`.

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` inapaswa kurudisha `2`.

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` inapaswa kurudisha `12`.

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` inapaswa kurudisha `15`.

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` inapaswa kurudisha `67`.

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
