---
id: 69c5f3d787b1725d5f00c8b7
title: "Zoezi la 253: Tafuta kifupi"
challengeType: 28
dashedName: challenge-253
---

# --description--

Kutolewa mfuatano wa herufi unaowakilisha kifupi, rudisha jina kamili la shirika ambalo kifupi hicho kinamiliki kutoka kwenye orodha ifuatayo:

- `"National Avocado Storage Authority"`
- `"Cats Infiltration Agency"`
- `"Fluffy Beanbag Inspectors"`
- `"Department Of Jelly"`
- `"Wild Honey Organization"`
- `"Eating Pancakes Administration"`

Kila herufi katika kifupi kilichotolewa inapaswa kuendana na herufi ya kwanza ya kila neno katika shirika kinachomilikiwa, kwa mpangilio ule ule.

# --hints--

`findOrg("NASA")` inapaswa kurudisha `"National Avocado Storage Authority"`.

```js
assert.equal(findOrg("NASA"), "National Avocado Storage Authority");
```

`findOrg("CIA")` inapaswa kurudisha `"Cats Infiltration Agency"`.

```js
assert.equal(findOrg("CIA"), "Cats Infiltration Agency");
```

`findOrg("FBI")` inapaswa kurudisha `"Fluffy Beanbag Inspectors"`.

```js
assert.equal(findOrg("FBI"), "Fluffy Beanbag Inspectors");
```

`findOrg("DOJ")` inapaswa kurudisha `"Department Of Jelly"`.

```js
assert.equal(findOrg("DOJ"), "Department Of Jelly");
```

`findOrg("WHO")` inapaswa kurudisha `"Wild Honey Organization"`.

```js
assert.equal(findOrg("WHO"), "Wild Honey Organization");
```

`findOrg("EPA")` inapaswa kurudisha `"Eating Pancakes Administration"`.

```js
assert.equal(findOrg("EPA"), "Eating Pancakes Administration");
```

# --seed--

## --seed-contents--

```js
function findOrg(acronym) {

  return acronym;
}
```

# --solutions--

```js
function findOrg(acronym) {
  const orgs = [
    "National Avocado Storage Authority",
    "Cats Infiltration Agency",
    "Fluffy Beanbag Inspectors",
    "Department Of Jelly",
    "Wild Honey Organization",
    "Eating Pancakes Administration"
  ];

  for (let org of orgs) {
    const initials = org
      .split(" ")
      .map(word => word[0])
      .join("");

    if (initials.toUpperCase() === acronym.toUpperCase()) {
      return org;
    }
  }

  return "not found";
}
```
