---
id: 69c6ff713a52713463aa7927
title: "Zoezi la 261: Maneno ya nambari"
challengeType: 28
dashedName: challenge-261
---

# --description--

Ukipata nambari kamili kutoka 0 hadi 99, rudisha usemi wake wa neno kwa Kiingereza.

- `0` inarudisha `"zero"`.
- Nambari 1-19 zina majina yao ya kipekee (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- Marudio ya 10 kuanzia 20-90 yana majina yao (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- Nambari 21-99 ambazo si marudio ya 10 zinaandikwa kama maneno mawili yaliyounganishwa na alama ya msalaba. Kwa mfano `"forty-two"` na `"fifty-three"`.

# --hints--

`getNumberWords(0)` inapaswa kurudisha `"zero"`.

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` inapaswa kurudisha `"ten"`.

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` inapaswa kurudisha `"nineteen"`.

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` inapaswa kurudisha `"thirty"`.

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` inapaswa kurudisha `"fifty-three"`.

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` inapaswa kurudisha `"seven"`.

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` inapaswa kurudisha `"twelve"`.

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` inapaswa kurudisha `"sixty"`.

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` inapaswa kurudisha `"sixty-seven"`.

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` inapaswa kurudisha `"ninety-eight"`.

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
