---
id: 69c6ff713a52713463aa7929
title: "Changamoto 263: Msamiati wa binary"
challengeType: 28
dashedName: challenge-263
---

# --description--

Ikiwa utapewa herufi, amua kama uwakilishi wake wa binary wa biti 8 unaweza kupatikana katika gridi ifuatayo, kwa usawa au wima katika mwelekeo wowote:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

Kwa mfano, `"A"` ina uwakilishi wa binary `01000001`, unaoonekana katika safu ya kwanza kutoka kushoto kwenda kulia.

# --hints--

`isInCrossword("I")` inapaswa kurudisha `true`.

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` inapaswa kurudisha `true`.

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` inapaswa kurudisha `true`.

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` inapaswa kurudisha `true`.

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` inapaswa kurudisha `false`.

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` inapaswa kurudisha `false`.

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` inapaswa kurudisha `false`.

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` inapaswa kurudisha `false`.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
