---
id: 69c6ff713a52713463aa7926
title: "Zoezi la 260: Alama za neno"
challengeType: 28
dashedName: challenge-260
---

# --description--

Ukitoa neno, rudisha alama zake kwa kutumia jedwali la thamani za herufi la kawaida:

| Herufi | Thamani |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- Herufi kubwa na herufi ndogo zina thamani sawa.

# --hints--

`getWordScore("hi")` inapaswa kurudisha `17`.

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` inapaswa kurudisha `52`.

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` inapaswa kurudisha `169`.

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` inapaswa kurudisha `94`.

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
