---
id: 69cfca90e8a0a6d4d6871c53
title: "Zoezi la 269: Milo rafiki kwa wale wenye mzio"
challengeType: 28
dashedName: challenge-269
---

# --description--

Kutolewa safu ya data ya milo na safu ya data ya vitu vinavyosababisha mzio vya kuepuka, rudisha majina ya milo yote ambayo haina vitu vyovyote vya mzio vilivyotajwa.

- Kila mlo uko katika muundo `[meal, allergens]`, ambapo `meal` ni jina la mlo, na `allergens` ni safu ya data ya vitu vinavyosababisha mzio vilivyomo kwenye mlo huo. Kwa mfano, `["pasta", ["wheat", "milk"]]`.
- Vitu vinavyosababisha mzio vya kuepuka vitakuwa safu ya data ya mfuatano wa herufi.

Rudisha majina ya milo salama kwa mpangilio ule ule uliopewa. Ikiwa hakuna mlo salama, rudisha safu ya data tupu.

# --hints--

`getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"])` inapaswa kurudisha `["salad"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"]), ["salad"]);
```

`getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"])` inapaswa kurudisha `["fried rice", "chicken parmesan"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"]), ["fried rice", "chicken parmesan"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"])` inapaswa kurudisha `["oatmeal", "granola", "toast"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"]), ["oatmeal", "granola", "toast"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"])` inapaswa kurudisha `["granola", "yogurt", "eggs"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"]), ["granola", "yogurt", "eggs"]);
```

# --seed--

## --seed-contents--

```js
function getAllergenFriendlyMeals(meals, allergens) {

  return meals;
}
```

# --solutions--

```js
function getAllergenFriendlyMeals(meals, allergens) {
  return meals
    .filter(([name, mealAllergens]) =>
      !allergens.some(a => mealAllergens.includes(a)))
    .map(([name]) => name);
}
```
