---
id: 69cfca90e8a0a6d4d6871c52
title: "Zoezi la 268: Nambari ya Narcissistic"
challengeType: 28
dashedName: challenge-268
---

# --description--

Ikiwa nambari kamili chanya imetolewa, amua kama ni nambari ya narcissistic.

- Nambari ni ya narcissistic ikiwa jumla ya kila tarakimu yake iliyoinuliwa kwa nguvu ya jumla ya tarakimu ni sawa na nambari yenyewe.

Kwa mfano, 153 ina tarakimu 3, na 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153, kwa hivyo ni narcissistic.

# --hints--

`isNarcissistic(153)` inapaswa kurudisha `true`.

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` inapaswa kurudisha `false`.

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` inapaswa kurudisha `true`.

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` inapaswa kurudisha `false`.

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` inapaswa kurudisha `true`.

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` inapaswa kurudisha `false`.

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` inapaswa kurudisha `true`.

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` inapaswa kurudisha `false`.

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
