---
id: 69cfca90e8a0a6d4d6871c4f
title: "Zoezi la 265: Mabano yaliyopangwa ndani zaidi"
challengeType: 28
dashedName: challenge-265
---

# --description--

Ikiwa umepewa mfuatano wa herufi unaoonyesha mabano yaliyo sawa, rudisha yaliyomo katika mabano yaliyopangwa ndani zaidi.

- Mabano yanaweza kuwa ya aina yoyote kati ya tatu: `()`, `[]`, na `{}`.
- Ingizo litakuwa na kundi moja tu lililopangwa ndani zaidi.

Kwa mfano, ukipata `"(hello (world))"`, rudisha `"world"`.

# --hints--

`getDeepestBrackets("(hello (world))")` inapaswa kurudisha `"world"`.

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")` inapaswa kurudisha `"inner"`.

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")` inapaswa kurudisha `"e"`.

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` inapaswa kurudisha `"fox"`.

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` inapaswa kurudisha `"C"`.

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```
