---
id: 69e2383af7832c8032603b90
title: "Zoezi la 274: Mtu Mkongwe Zaidi"
challengeType: 28
dashedName: challenge-274
---

# --description--

Ukitolewa safu ya vitu, kila kimoja kikiwa na sifa za `"name"` na `"age"`, rudisha safu ya data inayojumuisha majina ya watu wakongwe zaidi.

Ikiwa watu wengi wana umri sawa wa kuwa wakongwe zaidi, rudisha majina yao yote kwa mpangilio walivyoonekana katika ingizo.

# --hints--

`getOldest([{ name: "Brenda", age: 40 }])` inapaswa kurudisha `["Brenda"]`.

```js
assert.deepEqual(getOldest([{ name: "Brenda", age: 40 }]), ["Brenda"]);
```

`getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }])` inapaswa kurudisha `["Alice"]`.

```js
assert.deepEqual(getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }]), ["Alice"]);
```

`getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }])` inapaswa kurudisha `["Bill", "Carol"]`.

```js
assert.deepEqual(getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }]), ["Bill", "Carol"]);
```

`getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }])` inapaswa kurudisha `["George", "Holly", "Zach"]`.

```js
assert.deepEqual(getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }]), ["George", "Holly", "Zach"]);
```

# --seed--

## --seed-contents--

```js
function getOldest(people) {

  return people;
}
```

# --solutions--

```js
function getOldest(people) {
  const maxAge = Math.max(...people.map(p => p.age));
  return people.filter(p => p.age === maxAge).map(p => p.name);
}
```
