---
id: 69e2383af7832c8032603b93
title: "Zoezi la 277: Picha ya Kioo"
challengeType: 28
dashedName: challenge-277
---

# --description--

Ukizingatia mfuatano wa herufi miwili, amua kama mfuatano wa pili ni picha ya kioo ya mfuatano wa kwanza.

Picha ya kioo huundwa kwa kubadilisha mfuatano wa herufi na kuchukua kila herufi na kuibadilisha na sawa yake katika kioo.

- Herufi zinazolingana zinaonekana kama zilivyo katika kioo:

`W`, `T`, `Y`, `U`, `I`, `O`, `H`, `A`, `X`, `V`, `M`, `w`, `o`, `x`, `v`, `0`, `8`, `=`, `+`, `:`, `|`, `-`, `_`, `*`, `^`, `!`, `.`, na nafasi (` `).

- Pairs zinazobadilishana katika kioo hubadilishana kati yao:

| Herufi | Hubadilishana na |
| - | - |
| `[` | `]` |
| `{` | `}` |
| `<` | `>` |
| `b` | `d` |
| `p` | `q` |
| `(` | `)` |

Kama mfuatano wowote una herufi isiyo kwenye orodha hapo juu, haina picha ya kioo inayoweza kuundwa kutoka kwa herufi hizo.

Kwa mfano, picha ya kioo ya `"[HOW]"` ni `"[WOH]"`.

# --hints--

`isMirrorImage("[HOW]", "[WOH]")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirrorImage("[HOW]", "[WOH]"));
```

`isMirrorImage("MOM", "MOM")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirrorImage("MOM", "MOM"));
```

`isMirrorImage("vow", "wov")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirrorImage("vow", "wov"));
```

`isMirrorImage("TIM", "TIM")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMirrorImage("TIM", "TIM"));
```

`isMirrorImage("{WOW}", "}WOW{")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMirrorImage("{WOW}", "}WOW{"));
```

`isMirrorImage("XXVII", "IIV%X")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMirrorImage("XXVII", "IIV%X"));
```

`isMirrorImage("><(((*>", "<*)))><")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirrorImage("><(((*>", "<*)))><"));
```

`isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW"));
```

# --seed--

## --seed-contents--

```js
function isMirrorImage(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirrorImage(str1, str2) {
  const symmetric = new Set('WTYUIOHA XVMwoxv08=+:|-_*^!.');
  const pairs = { '[': ']', ']': '[', '{': '}', '}': '{', '<': '>', '>': '<', 'b': 'd', 'd': 'b', 'p': 'q', 'q': 'p', '(': ')', ')': '(' };
  
  const mirrored = str1.split('').reverse().map(c => {
    if (symmetric.has(c)) return c;
    if (pairs[c]) return pairs[c];
    return null;
  });
  
  if (mirrored.includes(null)) return false;
  return mirrored.join('') === str2;
}
```
