---
id: 69e2383af7832c8032603b94
title: "Zoezi la 278: Mchambuzi wa Oda ya Kahawa"
challengeType: 28
dashedName: challenge-278
---

# --description--

Kutolewa mfuatano wa herufi wa oda ya kahawa, tambua vitu vyovyote vya menyu na rudisha oda iliyopangwa vizuri.

Tumia vitu vya menyu na bei zifuatazo:

| Kitu | Bei | 
| - | - |
| `"cold brew"` | $4.50 |
| `"oat latte"` | $5.00 |
| `"cappuccino"` | $4.75 |
| `"espresso"` | $3.00 |
| `"vanilla syrup"` | $0.75 |
| `"caramel drizzle"` | $0.60 |
| `"extra shot"` | $0.50 |
| `"oat milk"` | $0.75 |
| `"cream"` | $0.75 |

Rudisha mfuatano wa herufi wenye vitu vilivyolingana vikiwa vimeunganishwa na `" + "`, ikifuatiwa na nukta mbili na nafasi (`": "`), na jumla ya bei.

Kwa mfano, ukipokea `"I'd like an oat latte with vanilla syrup and an extra shot please."`, rudisha `"oat latte + vanilla syrup + extra shot: $6.25"`

Vitu vinapaswa kuonekana kwa mpangilio walivyo kwenye menyu na jumla ya bei inapaswa kuwa na sehemu mbili za desimali kila wakati.

# --hints--

`formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please.")` inapaswa kurudisha `"oat latte + vanilla syrup + extra shot: $6.25"`.

```js
assert.equal(formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25");
```

`formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")` inapaswa kurudisha `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"`.

```js
assert.equal(formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85");
```

`formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot.")` inapaswa kurudisha `"cold brew + extra shot + cream: $5.75"`.

```js
assert.equal(formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75");
```

`formatCoffeeOrder("Just an espresso please.")` inapaswa kurudisha `"espresso: $3.00"`.

```js
assert.equal(formatCoffeeOrder("Just an espresso please."), "espresso: $3.00");
```

`formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")` inapaswa kurudisha `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"`.

```js
assert.equal(formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60");
```

# --seed--

## --seed-contents--

```js
function formatCoffeeOrder(order) {

  return order;
}
```

# --solutions--

```js
function formatCoffeeOrder(order) {
  const menu = [
    { name: "cold brew", price: 4.50 },
    { name: "oat latte", price: 5.00 },
    { name: "cappuccino", price: 4.75 },
    { name: "espresso", price: 3.00 },
    { name: "vanilla syrup", price: 0.75 },
    { name: "caramel drizzle", price: 0.60 },
    { name: "extra shot", price: 0.50 },
    { name: "oat milk", price: 0.75 },
    { name: "cream", price: 0.75 },
  ];

  const found = menu.filter(item => order.toLowerCase().includes(item.name));
  const total = found.reduce((sum, item) => sum + item.price, 0);
  return `${found.map(item => item.name).join(" + ")}: $${total.toFixed(2)}`;
}
```
