---
id: 69e2383af7832c8032603b8f
title: "Zoezi la 273: Kagua nambari halali ya ISBN-13"
challengeType: 28
dashedName: challenge-273
---

# --description--

Ikiwa umepewa mfuatano wa herufi, amua kama ni nambari halali ya ISBN-13.

ISBN-13 halali:

- Ina herufi za nambari tu na alama za msalaba (-)
- Ina nambari 13 kamili baada ya kuondoa alama za msalaba
- Inapita ukaguzi ufuatao:
  1. Zidisha kila nambari kwa 1 au 3, kwa mzunguko (zidisha nambari ya kwanza kwa 1, ya pili kwa 3, ya tatu kwa 1, na kadhalika).
  2. Jumla ya matokeo lazima iweze kugawika kwa 10 bila kubaki.

# --hints--

`isValidIsbn13("9780306406157")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
