---
id: 69f35a5bb823ed620fcb7cbb
title: "Zoezi la 282: Deni la usingizi"
challengeType: 28
dashedName: challenge-282
---

# --description--

Ukizingatia safu ya data ya saa ulizolala kila usiku hadi leo, na nambari lengwa ya saa kwa usiku, rudisha ni saa ngapi za usingizi unazohitaji usiku wa leo kufuta deni lako la usingizi.

- Jumuisha saa za usiku wa leo katika jumla ya muda unaohitajika kufikia malengo.
- Ikiwa umelala vya kutosha kufikia au kuzidi lengo la usiku wa leo, rudisha `0`.

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)` inapaswa kurudisha `20`.

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)` inapaswa kurudisha `10`.

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)` inapaswa kurudisha `4`.

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)` inapaswa kurudisha `0`.

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)` inapaswa kurudisha `0`.

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
