---
id: 69f35a5bb823ed620fcb7cbc
title: "Zoezi la 283: Mfuatano wa herufi wa zipu"
challengeType: 28
dashedName: challenge-283
---

# --description--

Ukipata mfuatano wa herufi miwili, rudisha mfuatano mpya unaochanganya herufi zao moja kwa moja kwa mpangilio. Ikiwa mfuatano mmoja ni mrefu zaidi, ongeza herufi zilizobaki mwishoni.

Anza na herufi ya kwanza ya mfuatano wa kwanza.

# --hints--

`zipStrings("abc", "123")` inapaswa kurudisha `"a1b2c3"`.

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` inapaswa kurudisha `"abcdefghijklmnopqrstuvwxyz"`.

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` inapaswa kurudisha `"dnaiyght"`.

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` inapaswa kurudisha `"pjyatvhaosncript"`.

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` inapaswa kurudisha `"freeCodeCamp"`.

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
