---
id: 69f35a5bb823ed620fcb7cbd
title: "Zoezi la 284: I kabla ya E"
challengeType: 28
dashedName: challenge-284
---

# --description--

Ukitoa neno au sentensi, rudi toleo lililosahihishwa ambapo kila neno linafuata kanuni ya "I kabla ya E isipokuwa baada ya C".

- Ikiwa neno lina `"ei"` lisilotanguliwa na `"c"`, libadilishe na `"ie"`.
- Ikiwa neno lina `"ie"` linalotanguliwa na `"c"`, libadilishe na `"ei"`.
- Maneno mengine yote yaachwe bila kubadilishwa.

# --hints--

`iBeforeE("beleive")` inapaswa kurudisha `"believe"`.

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")` inapaswa kurudisha `"receive"`.

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")` inapaswa kurudisha `"we received a brief"`.

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")` inapaswa kurudisha `"she believed the friendly niece could perceive the grief"`.

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")` inapaswa kurudisha `"we received relief after the thief gave us a brief piece of fierce deceit"`.

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```
