---
id: 69f35a5bb823ed620fcb7cb9
title: "Zoezi la 280: Tarehe ya Mongo ID"
challengeType: 28
dashedName: challenge-280
---

# --description--

Ukitoa mfuatano wa herufi wa MongoDB ID, rudisha wakati wake wa kuundwa kama mfuatano wa herufi wa ISO 8601.

- MongoDB ID ni mfuatano wa herufi 24 wa hex. Herufi nane za kwanza zinawakilisha alama ya wakati ya Unix (kwa sekunde) iliyowakilishwa kama nambari ya msingi-16.

Kwa mfano, `"6a094b50bcf86cd799439011"` ina alama ya wakati ya `"6a094b50"` kwa hex, ambayo ni `1778994000` kwa desimali, ikiwakilisha wakati wa kuundwa wa `"2026-05-17T05:00:00.000Z"`.

# --hints--

`mongoIdToDate("6a094b50bcf86cd799439011")` inapaswa kurudisha `"2026-05-17T05:00:00.000Z"`.

```js
assert.equal(mongoIdToDate("6a094b50bcf86cd799439011"), "2026-05-17T05:00:00.000Z");
```

`mongoIdToDate("695344eb1f4a4c1123042128")` inapaswa kurudisha `"2025-12-30T03:20:11.000Z"`.

```js
assert.equal(mongoIdToDate("695344eb1f4a4c1123042128"), "2025-12-30T03:20:11.000Z");
```

`mongoIdToDate("386da62df34123ac54617e56")` inapaswa kurudisha `"2000-01-01T07:01:01.000Z"`.

```js
assert.equal(mongoIdToDate("386da62df34123ac54617e56"), "2000-01-01T07:01:01.000Z");
```

`mongoIdToDate("69f571c3d7711807afd3dd55")` inapaswa kurudisha `"2026-05-02T03:38:43.000Z"`.

```js
assert.equal(mongoIdToDate("69f571c3d7711807afd3dd55"), "2026-05-02T03:38:43.000Z");
```

`mongoIdToDate("68adce01c0e1144d0a90295a")` inapaswa kurudisha `"2025-08-26T15:08:49.000Z"`.

```js
assert.equal(mongoIdToDate("68adce01c0e1144d0a90295a"), "2025-08-26T15:08:49.000Z");
```

# --seed--

## --seed-contents--

```js
function mongoIdToDate(id) {

  return id;
}
```

# --solutions--

```js
function mongoIdToDate(id) {
  const timestamp = parseInt(id.slice(0, 8), 16);
  return new Date(timestamp * 1000).toISOString();
}
```
