---
id: 69f8c998d78ad3171a0713bd
title: "Zoezi la 290: Sherehe ya pizza"
challengeType: 28
dashedName: challenge-290
---

# --description--

Kutolewa safu ya data ya saa zilizofanyakaziwa leo kwa kila mtu, rudisha idadi ya pizza za kuagiza kwa ajili ya sherehe ya pizza.

- Gawanya saa zilizofanyakaziwa na kila mtu kwa 3 kupata idadi ya vipande vya pizza anavyopaswa.
- Huwezi kula kipande kidogo, kwa hivyo zungusha idadi ya vipande vya kila mtu hadi nambari nzima inayofuata.
- Kila mtu anapata angalau vipande viwili.
- Kila pizza ina vipande 8. Zungusha jumla ya pizza hadi pizza nzima inayofuata.

# --hints--

`getPizzasToOrder([8, 8, 8])` inapaswa kurudisha `2`.

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` inapaswa kurudisha `3`.

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` inapaswa kurudisha `2`.

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` inapaswa kurudisha `3`.

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` inapaswa kurudisha `1`.

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` inapaswa kurudisha `5`.

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
