---
id: 69f8c998d78ad3171a0713bf
title: "Zoezi la 292: Uwiano mpana wa sura"
challengeType: 28
dashedName: challenge-292
---

# --description--

Ukipatiwa mfuatano wa herufi miwili wa vipimo tofauti vya picha, rudisha uwiano wa sura wa picha yenye uwiano mkubwa zaidi wa upana kwa urefu.

- Mfuatano wa herufi uliotolewa utakuwa katika muundo `"WxH"`, kwa mfano, `"1920x1080"`.
- Uwiano wa sura ni uwiano wa upana kwa urefu, uliopunguzwa hadi nambari kamili ndogo kabisa. Kwa mfano, `"1920x1080"` hupunguzwa kuwa `"16:9"`.
- Rudisha mfuatano wa herufi katika muundo `"W:H"`, kwa mfano, `"16:9"`.

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` inapaswa kurudisha `"16:9"`.

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` inapaswa kurudisha `"4:3"`.

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` inapaswa kurudisha `"2:1"`.

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` inapaswa kurudisha `"9:16"`.

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` inapaswa kurudisha `"43:18"`.

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` inapaswa kurudisha `"2089:8539"`.

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
