---
id: 69f8c998d78ad3171a0713ba
title: "Zoezi la 287: Mrekebishaji wa Nambari za Kirumi"
challengeType: 28
dashedName: challenge-287
---

# --description--

Kutolewa mfuatano wa herufi wa nambari za Kirumi zisizopangwa vizuri, rudisha thamani yake kwa noti ya kawaida ya nambari za Kirumi.

Ingizo litakuwa likitumia tu noti ya kuongeza, kwa hivyo kila alama inaongeza thamani yake kwa jumla. Kama ukumbusho, hizi ni alama na thamani zao:

| Alama | Thamani |
| ------ | ----- |
| `"I"`  | 1     |
| `"V"`  | 5     |
| `"X"`  | 10    |
| `"L"`  | 50    |
| `"C"`  | 100   |
| `"D"`  | 500   |
| `"M"`  | 1000  |

Unapofanya upya usimbaji, tumia alama kubwa zaidi inayowezekana katika kila hatua, ukitumia jozi za kutoa (subtractive pairs) (`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`) inapohitajika.

# --hints--

`fixNumerals("XIIIII")` inapaswa kurudisha `"XV"`.

```js
assert.equal(fixNumerals("XIIIII"), "XV");
```

`fixNumerals("IIIILX")` inapaswa kurudisha `"LXIV"`.

```js
assert.equal(fixNumerals("IIIILX"), "LXIV");
```

`fixNumerals("XXVVVIIIII")` inapaswa kurudisha `"XL"`.

```js
assert.equal(fixNumerals("XXVVVIIIII"), "XL");
```

`fixNumerals("MDCCLXXXXVIIII")` inapaswa kurudisha `"MDCCXCIX"`.

```js
assert.equal(fixNumerals("MDCCLXXXXVIIII"), "MDCCXCIX");
```

`fixNumerals("IIIIVVVVXXXXLLLLCCDD")` inapaswa kurudisha `"MCDLXIV"`.

```js
assert.equal(fixNumerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV");
```

`fixNumerals("ILCDMIVDIIXLCVCXDL")` inapaswa kurudisha `"MMCMLXXXIV"`.

```js
assert.equal(fixNumerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV");
```

# --seed--

## --seed-contents--

```js
function fixNumerals(str) {

  return str;
}
```

# --solutions--

```js
function fixNumerals(str) {
  const values = { I: 1, V: 5, X: 10, L: 50, C: 100, D: 500, M: 1000 };
  const encoding = [
    [1000, "M"], [900, "CM"], [500, "D"], [400, "CD"],
    [100, "C"], [90, "XC"], [50, "L"], [40, "XL"],
    [10, "X"], [9, "IX"], [5, "V"], [4, "IV"], [1, "I"]
  ];

  let total = [...str].reduce((sum, c) => sum + values[c], 0);

  let result = "";
  for (const [value, symbol] of encoding) {
    while (total >= value) {
      result += symbol;
      total -= value;
    }
  }
  return result;
}
```
