---
id: 69f8c998d78ad3171a0713bc
title: "Zoezi la 289: Jumla ya Tofauti"
challengeType: 28
dashedName: challenge-289
---

# --description--

Ukipata safu ya nambari, rudisha jumla ya tofauti kati ya kila nambari na ile inayofuata.

Kwa mfano, ukipata `[1, 3, 4]`, rudisha `3` (2 + 1).

# --hints--

`sumOfDifferences([1, 3, 4])` inapaswa kurudisha `3`.

```js
assert.equal(sumOfDifferences([1, 3, 4]), 3);
```

`sumOfDifferences([5, -3, 3, 9, 10])` inapaswa kurudisha `5`.

```js
assert.equal(sumOfDifferences([5, -3, 3, 9, 10]), 5);
```

`sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7])` inapaswa kurudisha `-16`.

```js
assert.equal(sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16);
```

`sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])` inapaswa kurudisha `25`.

```js
assert.equal(sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25);
```

# --seed--

## --seed-contents--

```js
function sumOfDifferences(arr) {

  return arr;
}
```

# --solutions--

```js
function sumOfDifferences(arr) {
  return arr.slice(0, -1).reduce((sum, n, i) => sum + (arr[i + 1] - n), 0);
}
```
