---
id: 69f8c998d78ad3171a0713b9
title: "Changamoto 286: Masuala Waliyo wazi"
challengeType: 28
dashedName: challenge-286
---

# --description--

Ukizingatia safu ya nambari za masuala na safu nyingine ya nambari za ombi la kuvuta mabadiliko (PR), rudisha safu ya masuala ambayo bado yako wazi baada ya maombi yote ya PR kuunganishwa.

- Ombi la PR linafungua suala ikiwa nambari zao ni mzunguko wa kila mmoja. Kwa mfano, suala 123 lingefungwa na PR 231 au 312.
- Ombi la PR halifungi suala lenye nambari sawa kabisa. Kwa mfano, PR 123 halifungi suala 123. Hivyo suala lenye nambari zote sawa haliwezi kufungwa.
- Nambari yoyote inaweza kuwa na sifuri za mbele zilizokatwa. Kwa mfano, PR 201 lingefunga suala 12 (012, mzunguko wa 201). Vivyo hivyo, suala 201 lingefungwa na PR 12.

Rudisha masuala yaliyo wazi yaliyobaki kwa mpangilio waliotolewa.

# --hints--

`getOpenIssues([123, 234], [231])` inapaswa kurudisha `[234]`.

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` inapaswa kurudisha `[345, 16]`.

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` inapaswa kurudisha `[456, 332, 15]`.

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` inapaswa kurudisha `[115, 296, 24]`.

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` inapaswa kurudisha `[95, 395, 754, 296, 709, 237, 1802]`.

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
