---
id: 69f8c998d78ad3171a0713bb
title: "Zoezi la 288: Nambari ya siri"
challengeType: 28
dashedName: challenge-288
---

# --description--

Ikiwa umepewa nambari ya siri na makisio, amua kama makisio ni sahihi.

Rudisha:

- `"higher"` ikiwa nambari ya siri ni kubwa kuliko makisio.
- `"lower"` ikiwa nambari ya siri ni ndogo kuliko makisio.
- `"you got it!"` ikiwa makisio ni sahihi.

# --hints--

`guessNumber(50, 30)` inapaswa kurudisha `"higher"`.

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` inapaswa kurudisha `"lower"`.

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` inapaswa kurudisha `"you got it!"`.

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` inapaswa kurudisha `"higher"`.

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` inapaswa kurudisha `"lower"`.

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` inapaswa kurudisha `"you got it!"`.

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
