---
id: 69f90c1329a94b37e2a2086d
title: "Zoezi la 294: Mchanganyiko wa mabano ya kawaida"
challengeType: 28
dashedName: challenge-294
---

# --description--

Ukitoa nambari kamili, `n`, rudisha idadi ya mchanganyiko halali wa jozi `n` za mabano ya kawaida.

- Mchanganyiko halali ni mfuatano wa herufi ambapo kila mabano ya kawaida ya kufungua yana mabano ya kawaida ya kufunga yanayolingana, na hakuna mabano ya kufunga yanayotokea kabla ya mabano ya kufungua yanayolingana nayo.

Kwa mfano, ukitoa `2`, kuna mchanganyiko 2 halali:

```json
(())
()()
```

# --hints--

`getCombinations(2)` inapaswa kurudisha `2`.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` inapaswa kurudisha `5`.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` inapaswa kurudisha `42`.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` inapaswa kurudisha `1430`.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` inapaswa kurudisha `742900`.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
