---
id: 69f90c1329a94b37e2a2086c
title: "Zoezi la 293: Mkono bora"
challengeType: 28
dashedName: challenge-293
---

# --description--

Kutolewa safu ya data ya mfuatano wa herufi mitano inayowakilisha kadi za kucheza, rudisha jina la mkono bora.

- Kadi kila moja inawakilishwa kama mfuatano wa herufi mbili: daraja ikifuatiwa na rangi, `"2h"` kwa mfano.
  - Madaraja, kutoka chini hadi juu, ni: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"`, na `"A"`.
  - Rangi ni: `"h"`, `"d"`, `"c"`, na `"s"`.
- Aces (`"A"`) zinaweza kutumika kama za juu au za chini katika mfululizo.

Mikono, kwa mpangilio kutoka mbaya hadi bora, ni:

| Jina | Maelezo |
| - | - |
| `"High Card"` | Hakuna jozi au bora zaidi |
| `"Pair"` | Jozi moja ya daraja moja |
| `"Two Pair"` | Jozi mbili za madaraja tofauti |
| `"Three of a Kind"` | Tatu za daraja moja |
| `"Straight"` | Mfululizo wa madaraja matano |
| `"Flush"` | Matano ya rangi moja |
| `"Full House"` | Tatu za daraja moja, na jozi moja ya daraja nyingine |
| `"Four of a Kind"` | Nne za daraja moja |
| `"Straight Flush"` | Mfululizo wa madaraja matano ya rangi moja |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` za rangi moja |

Rudisha jina la mkono bora.

# --hints--

`getBestHand(["7s", "7h", "7d", "2c", "5h"])` inapaswa kurudisha `"Three of a Kind"`.

```js
assert.equal(getBestHand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind");
```

`getBestHand(["Ks", "Kh", "Kd", "4s", "4h"])` inapaswa kurudisha `"Full House"`.

```js
assert.equal(getBestHand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House");
```

`getBestHand(["2h", "5h", "7h", "9h", "Jh"])` inapaswa kurudisha `"Flush"`.

```js
assert.equal(getBestHand(["2h", "5h", "7h", "9h", "Jh"]), "Flush");
```

`getBestHand(["As", "Ah", "Ad", "Ac", "Kh"])` inapaswa kurudisha `"Four of a Kind"`.

```js
assert.equal(getBestHand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind");
```

`getBestHand(["Ts", "Th", "9d", "9c", "8h"])` inapaswa kurudisha `"Two Pair"`.

```js
assert.equal(getBestHand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair");
```

`getBestHand(["9c", "8c", "7c", "6c", "5c"])` inapaswa kurudisha `"Straight Flush"`.

```js
assert.equal(getBestHand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush");
```

`getBestHand(["As", "Kh", "Jd", "8c", "5h"])` inapaswa kurudisha `"High Card"`.

```js
assert.equal(getBestHand(["As", "Kh", "Jd", "8c", "5h"]), "High Card");
```

`getBestHand(["As", "2h", "3d", "4c", "5h"])` inapaswa kurudisha `"Straight"`.

```js
assert.equal(getBestHand(["As", "2h", "3d", "4c", "5h"]), "Straight");
```

`getBestHand(["Ts", "Th", "7c", "6d", "5h"])` inapaswa kurudisha `"Pair"`.

```js
assert.equal(getBestHand(["Ts", "Th", "7c", "6d", "5h"]), "Pair");
```

`getBestHand(["As", "Ks", "Qs", "Js", "Ts"])` inapaswa kurudisha `"Royal Flush"`.

```js
assert.equal(getBestHand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush");
```

# --seed--

## --seed-contents--

```js
function getBestHand(cards) {

  return cards;
}
```

# --solutions--

```js
function getBestHand(cards) {
  const rankOrder = "23456789TJQKA";

  const ranks = cards.map(c => c[0]);
  const suits = cards.map(c => c[1]);

  const isFlush = suits.every(s => s === suits[0]);

  const rankIndices = ranks.map(r => rankOrder.indexOf(r)).sort((a, b) => a - b);

  function isStraight(indices) {
    const isRegular = indices.every((r, i) => i === 0 || r === indices[i - 1] + 1);
    const isLowAce = JSON.stringify(indices) === JSON.stringify([0, 1, 2, 3, 12]);
    return isRegular || isLowAce;
  }

  const straight = isStraight(rankIndices);

  const counts = {};
  ranks.forEach(r => counts[r] = (counts[r] || 0) + 1);
  const freq = Object.values(counts).sort((a, b) => b - a);

  if (isFlush && JSON.stringify(rankIndices) === JSON.stringify([8, 9, 10, 11, 12])) return "Royal Flush";
  if (isFlush && straight) return "Straight Flush";
  if (freq[0] === 4) return "Four of a Kind";
  if (freq[0] === 3 && freq[1] === 2) return "Full House";
  if (isFlush) return "Flush";
  if (straight) return "Straight";
  if (freq[0] === 3) return "Three of a Kind";
  if (freq[0] === 2 && freq[1] === 2) return "Two Pair";
  if (freq[0] === 2) return "Pair";
  return "High Card";
}
```
