---
id: 6a0dcd03ee4e68698080ef66
title: "Zoezi la 302: Kuathirika na Muda wa Safari ya Ndege"
challengeType: 28
dashedName: challenge-302
---

# --description--

Ukizingatia jiji la kuondoka, jiji la kufikia, muda wa safari ya ndege kwa saa, na mwelekeo wa kusafiri, rudisha idadi ya saa za kuathirika na muda wa safari ya ndege (jet lag) ambazo msafiri anazipata.

Miji iliyotolewa itakuwa kutoka kwenye orodha ifuatayo inayojumuisha mbali zao za UTC:

| Jiji | Mbali |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Ili kuhesabu saa za kuathirika na muda wa safari ya ndege:

1. Tafuta tofauti ya saa za mkoa wa muda kati ya miji miwili.
2. Amua kigezo cha mwelekeo. Ikiwa unasafiri `"east"`, ni 1.5, vinginevyo ni 1.0.
3. Pata saa za kuathirika na formula: tofauti ya mkoa wa muda + (muda wa safari * 0.1) * kigezo cha mwelekeo

Rudisha saa za kuathirika zikiwa zimezungushwa hadi sehemu moja ya desimali.

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` inapaswa kurudisha `6.5`.

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` inapaswa kurudisha `5.8`.

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` inapaswa kurudisha `1.6`.

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` inapaswa kurudisha `4.7`.

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` inapaswa kurudisha `17.5`.

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` inapaswa kurudisha `5.9`.

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` inapaswa kurudisha `9.5`.

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
