---
id: 6a0dcd03ee4e68698080ef6b
title: "Zoezi la 307: Kanuni za Mipangilio ya Miji"
challengeType: 28
dashedName: challenge-307
---

# --description--

Ikiwa umepewa gridi ya pande mbili (safu ya data ya safu za data) inayowakilisha mpangilio wa majengo ya jiji, rudisha kuratibu za majengo yote yanayokiuka kanuni za mipangilio.

Kila seli katika gridi ina moja ya lebo kutoka kwenye jedwali hapa chini. Jengo linakiuka kanuni ikiwa yoyote ya majirani zake (hadi nne), wima au mlalo, ni aina ambayo haliwezi kuwa jirani nayo.

| Lebo | Aina | Haiwezi kuwa jirani na |
| - | - | - |
| `"i"` | viwanda | `"R"`, `"I"` |
| `"A"` | kilimo | `"C"` |
| `"R"` | makazi | `"i"`, `"C"` |
| `"I"` | taasisi | `"i"` |
| `"C"` | biashara | `"R"`, `"A"` |
| `""` (mfuatano wa herufi tupu) | haijajengwa | hakuna vikwazo |

Rudisha kuratibu za seli zote zinazokiuka kama safu ya jozi za `[row, col]`, kwa mpangilio wowote. Ikiwa hakuna ukiukaji wowote, rudisha safu ya data tupu.

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` inapaswa kurudisha `[[0, 0], [0, 1]]`.

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` inapaswa kurudisha `[[0, 1], [1, 1]]`.

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` inapaswa kurudisha `[]`.

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` inapaswa kurudisha `[[0, 1], [0, 2], [0, 3]]`.

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` inapaswa kurudisha `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
