---
id: 6a0dcc730cb92a616f86f0c5
title: "Zoezi la 301: Mzigo wa Mwisho"
challengeType: 28
dashedName: challenge-301
---

# --description--

Ukizingatia idadi ya vijiko vya sabuni ya kuosha nguo ulivyo navyo na safu ya data ya vijiko ulivyotumia katika siku zilizopita, rudisha idadi ya siku kamili za sabuni ulizo nazo.

Hesabu wastani wa matumizi yako ya kila siku kutoka kwa historia ya matumizi na chukulia matumizi hayo kila siku kuanzia sasa.

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` inapaswa kurudisha `5`.

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` inapaswa kurudisha `7`.

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` inapaswa kurudisha `11`.

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` inapaswa kurudisha `12`.

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` inapaswa kurudisha `1`.

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
