---
id: 6a0dcc730cb92a616f86f0c2
title: "Zoezi la 298: Kagua muundo sehemu ya 4"
challengeType: 28
dashedName: challenge-298
---

# --description--

Ikiwa umepewa kitu (JavaScript) au kamusi (Python), amua kama kinakidhi muundo ufuatao:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles,
  supporter?: boolean
}
```

- Alama ya bomba (`|`) ina maana ya "au". `role` lazima iwe mojawapo ya thamani zilizoorodheshwa za `Roles`.
- Alama ya swali (`?`) baada ya `supporter` ina maana kuwa sehemu hiyo ni hiari, lakini ni aina iliyobainishwa ikiwa ipo.
- Funguo za ziada zinaruhusiwa

# --hints--

`isValidSchema({ username: "vivian", posts: 1, verified: false, role: "user", supporter: true })` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidSchema({ username: "vivian", posts: 1, verified: false, role: "user", supporter: true }));
```

`isValidSchema({ username: "rudolph", posts: 15, verified: true, role: "creator" })` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidSchema({ username: "rudolph", posts: 15, verified: true, role: "creator" }));
```

`isValidSchema({ username: "hernandez", posts: 35, verified: true, role: "moderator", supporter: false, followers: 55 })` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidSchema({ username: "hernandez", posts: 35, verified: true, role: "moderator", supporter: false, followers: 55 }));
```

`isValidSchema({ username: "julia", posts: 50, verified: true, role: "admin", supporter: "true" })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: "julia", posts: 50, verified: true, role: "admin", supporter: "true" }));
```

`isValidSchema({ username: "bernard", posts: 0, verified: true, role: "friend", supporter: true })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: "bernard", posts: 0, verified: true, role: "friend", supporter: true }));
```

`isValidSchema({ username: "felix", posts: 40, verified: "yes", role: "staff", supporter: false })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: "felix", posts: 40, verified: "yes", role: "staff", supporter: false }));
```

`isValidSchema({ username: "jimmy", posts: true, verified: false, role: "creator", supporter: true })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: "jimmy", posts: true, verified: false, role: "creator", supporter: true }));
```

`isValidSchema({ username: true, posts: 30, verified: true, role: "moderator", supporter: false })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: true, posts: 30, verified: true, role: "moderator", supporter: false }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role) &&
    (obj.supporter === undefined || typeof obj.supporter === 'boolean')
  );
}
```
