---
id: 6a0dcc730cb92a616f86f0bf
title: "Zoezi la 295: Kagua muundo sehemu ya 1"
challengeType: 28
dashedName: challenge-295
---

# --description--

Ikiwa umepewa kitu (JavaScript) au kamusi (Python), thibitisha kama kinakidhi muundo ufuatao:

```json
{
  username: string
}
```

- Funguo za ziada zinaruhusiwa

# --hints--

`isValidSchema({ username: "bob" })` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
