---
id: 6a0dcd03ee4e68698080ef6c
title: "Zoezi la 308: Kagua kadi ya mkopo"
challengeType: 28
dashedName: challenge-308
---

# --description--

Ikiwa umepewa mfuatano wa nambari za kadi ya mkopo, amua kama ni nambari halali kwa kutumia njia ifuatayo:

- Anza kutoka kwa nambari ya pili kutoka mwisho, rudisha mara mbili kila nambari nyingine ukielekea kushoto.
- Ikiwa rudisha mara mbili nambari itatoa nambari kubwa kuliko 9, toa 9.
- Jumlisha nambari zote (ziliorodheshwa mara mbili na zisizorodheshwa).
- Ikiwa jumla ni sehemu ya 10, nambari ni halali.

# --hints--

`isValidCard("4532015112830366")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
