---
id: 6a0dcc730cb92a616f86f0c0
title: "Zoezi la 296: Kagua Mpangilio Sehemu ya 2"
challengeType: 28
dashedName: challenge-296
---

# --description--

Ikiwa umepewa kitu (JavaScript) au kamusi (Python), amua kama kinakidhi mpangilio ufuatao:

```json
{
  username: string,
  posts: number,
  verified: boolean
}
```

- Funguo za ziada zinaruhusiwa

# --hints--

`isValidSchema({ username: "alice", posts: 10, verified: false })` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidSchema({ username: "alice", posts: 10, verified: false }));
```

`isValidSchema({ username: "carol", posts: 15, verified: true, followers: 25 })` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidSchema({ username: "carol", posts: 15, verified: true, followers: 25 }));
```

`isValidSchema({ username: "frank", posts: "21", verified: true })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: "frank", posts: "21", verified: true }));
```

`isValidSchema({ username: "sam", posts: 17, verified: "false" })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: "sam", posts: 17, verified: "false" }));
```

`isValidSchema({ username: "bill", verified: true })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: "bill", verified: true }));
```

`isValidSchema({ username: "fred", verified: true })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: "fred", verified: true }));
```

`isValidSchema({ username: 5, posts: 10, verified: true })` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidSchema({ username: 5, posts: 10, verified: true }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean'
  );
}
```
