---
id: 6a15cadf5f240d05a2649559
title: "Changamoto 315: Siku Ndefu Za Majira ya Joto"
challengeType: 28
dashedName: challenge-315
---

# --description--

Leo ni siku ndefu za majira ya joto, siku ndefu zaidi ya mwaka katika Nusu ya Kaskazini na fupi zaidi katika Nusu ya Kusini. Ukipata latitudo, rudisha mfuatano wa herufi unaowakilisha saa za mchana na usiku.

- Latitudo itakuwa kati ya 90 (kifungu cha kaskazini) na -90 (kifungu cha kusini), ikijumuisha
- Idadi ya saa za mchana = 12 + (latitudo / 90) * 12
- Punguza saa za mchana hadi nambari jamaa inayokaribia zaidi

Rudisha mfuatano wa herufi 24 ukitumia `"☀️"` kwa saa za mchana na `"🌑"` kwa saa za usiku, ambapo:

- Kila herufi inawakilisha saa moja, kuanzia saa sita usiku (saa 0)
- Machweo na machweo ya jua yanajitokeza kwa usawa karibu na saa sita mchana

Kwa mfano, latitudo ya `0` (mwinuko wa katikati ya dunia) ina saa 12 za mchana, hivyo machweo ya jua ni saa 6:00 asubuhi na machweo ni saa 6:00 jioni. Rudisha: `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`.

# --hints--

`getDaytimeHours(0)` inapaswa kurudisha `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`.

```js
assert.equal(getDaytimeHours(0), "🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(90)` inapaswa kurudisha `"☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️"`.

```js
assert.equal(getDaytimeHours(90), "☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️");
```

`getDaytimeHours(-90)` inapaswa kurudisha `"🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
assert.equal(getDaytimeHours(-90), "🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(-33)` inapaswa kurudisha `"🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
assert.equal(getDaytimeHours(-33), "🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(66.5)` inapaswa kurudisha `"🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑"`.

```js
assert.equal(getDaytimeHours(66.5), "🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑");
```

`getDaytimeHours(40)` inapaswa kurudisha `"🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑"`.

```js
assert.equal(getDaytimeHours(40), "🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑");
```

`getDaytimeHours(-50)` inapaswa kurudisha `"🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
assert.equal(getDaytimeHours(-50), "🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑");
```

# --seed--

## --seed-contents--

```js
function getDaytimeHours(latitude) {

  return latitude;
}
```

# --solutions--

```js
function getDaytimeHours(latitude) {
  const daylightHours = Math.round((12 + (latitude / 90) * 12) / 2) * 2;
  const sunrise = 12 - daylightHours / 2;
  const sunset = 12 + daylightHours / 2;

  let result = "";
  for (let hour = 0; hour < 24; hour++) {
    result += hour >= sunrise && hour < sunset ? "☀️" : "🌑";
  }

  return result;
}
```
