---
id: 6a15cadf5f240d05a264955d
title: "Zoezi la 319: Mchambuzi wa Frontmatter"
challengeType: 28
dashedName: challenge-319
---

# --description--

Ukitoa mfuatano wa herufi unaowakilisha kipande cha frontmatter, chambua na urudishe kitu (JavaScript) au kamusi (Python) yenye funguo na thamani.

Frontmatter imefungwa ndani ya vizingiti vya `---` na ina jozi za `key: value` ndani yake, moja kwa mstari. Kwa mfano:

```md
---
title: My Post
draft: false
views: 100
---
```

Inapaswa kurudisha:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- Nambari, thamani za kweli au siyo kweli, na mfuatano wa herufi zote zinapaswa kurudishwa kama aina zao husika.
- Mfuatano uliotolewa utakuwa na mistari mipya iliyotenganishwa na tabia ya mistari mipya (`"\n"`). Mfano ulio juu utatolewa kama: `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`.

# --hints--

`parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` inapaswa kurudisha `{ title: "My Post", draft: false, views: 100 }`.

```js
assert.deepEqual(parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---"), { title: "My Post", draft: false, views: 100 });
```

`parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` inapaswa kurudisha `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false }`.

```js
assert.deepEqual(parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---"), { id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false });
```

`parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` inapaswa kurudisha `{ version: "1.0.0", url: "https://example.com", private: true }`.

```js
assert.deepEqual(parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---"), { version: "1.0.0", url: "https://example.com", private: true });
```

`parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---")` inapaswa kurudisha `{ rating: 4.5, price: 9.99 }`.

```js
assert.deepEqual(parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---"), { rating: 4.5, price: 9.99 });
```

# --seed--

## --seed-contents--

```js
function parseFrontmatter(str) {

  return str;
}
```

# --solutions--

```js
function parseFrontmatter(str) {
  const lines = str.split("\n").filter(line => line !== "---" && line !== "");

  return lines.reduce((obj, line) => {
    const [key, ...rest] = line.split(":");
    const value = rest.join(":").trim();

    if (value === "true") obj[key.trim()] = true;
    else if (value === "false") obj[key.trim()] = false;
    else if (!isNaN(value)) obj[key.trim()] = Number(value);
    else obj[key.trim()] = value;

    return obj;
  }, {});
}
```
