---
id: 6a15cadf5f240d05a264955b
title: "Zoezi la 317: Kihesabu cha BMI"
challengeType: 28
dashedName: challenge-317
---

# --description--

Ukichukulia uzito kwa pauni na urefu kwa inchi, rudisha BMI (Kielezo cha Uzito wa Mwili) kilichozungushwa hadi sehemu moja ya desimali.

Ili kupata BMI: gawanya uzito kwa urefu uliyozidishwa na mwenyewe, kisha zidisha matokeo na 703.

# --hints--

`calculateBmi(180, 70)` inapaswa kurudisha `25.8`.

```js
assert.equal(calculateBmi(180, 70), 25.8);
```

`calculateBmi(140, 64)` inapaswa kurudisha `24.0`.

```js
assert.equal(calculateBmi(140, 64), 24.0);
```

`calculateBmi(160, 76)` inapaswa kurudisha `19.5`.

```js
assert.equal(calculateBmi(160, 76), 19.5);
```

`calculateBmi(200, 60)` inapaswa kurudisha `39.1`.

```js
assert.equal(calculateBmi(200, 60), 39.1);
```

`calculateBmi(150, 68)` inapaswa kurudisha `22.8`.

```js
assert.equal(calculateBmi(150, 68), 22.8);
```

# --seed--

## --seed-contents--

```js
function calculateBmi(weight, height) {

  return weight;
}
```

# --solutions--

```js
function calculateBmi(weight, height) {
  return Math.round((weight / (height * height)) * 703 * 10) / 10;
}
```
