---
id: 6a15cadf5f240d05a2649556
title: "Changamoto 312: Gharama ya Kuangalia Moja kwa Moja"
challengeType: 28
dashedName: challenge-312
---

# --description--

Kutolewa safu ya data inayowakilisha filamu zilizo kwenye kikapu cha huduma yako ya kuangalia moja kwa moja, na mfuatano wa herufi kwa ngazi yako ya usajili, rudisha jumla ya gharama za filamu hizo.

Kila kipengee katika kikapu ni kitu chenye `"format"` (`"HD"` au `"4K"`) na `"type"` (`"rent"` au `"buy"`). Gharama zao ni:

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

Tumia punguzo zifuatazo za ngazi ya usajili:

- `"none"`: bei kamili
- `"basic"`: punguzo la 10%
- `"premium"`: punguzo la 25%

Rudisha jumla ya gharama iliyozungushwa hadi sehemu mbili za desimali katika muundo wa `"$D.CC"`.

# --hints--

`getStreamingBill([{ format: "HD", type: "rent" }], "none")` inapaswa kurudisha `"$3.99"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }], "none"), "$3.99");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "premium")` inapaswa kurudisha `"$12.73"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "premium"), "$12.73");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic")` inapaswa kurudisha `"$18.87"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic"), "$18.87");
```

`getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium")` inapaswa kurudisha `"$29.98"`.

```js
assert.equal(getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium"), "$29.98");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none")` inapaswa kurudisha `"$42.96"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none"), "$42.96");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }, { format: "HD", type: "buy" }], "basic")` inapaswa kurudisha `"$50.36"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }, { format: "HD", type: "buy" }], "basic"), "$50.36");
```

# --seed--

## --seed-contents--

```js
function getStreamingBill(cart, subscription) {

  return cart;
}
```

# --solutions--

```js
function getStreamingBill(cart, subscription) {
  const pricing = {
    HD: { rent: 3.99, buy: 12.99 },
    "4K": { rent: 5.99, buy: 19.99 },
  };

  const discounts = {
    none: 1,
    basic: 0.9,
    premium: 0.75,
  };

  const multiplier = discounts[subscription];

  const total = cart.reduce((sum, item) => {
    const price = pricing[item.format][item.type];
    return sum + price * multiplier;
  }, 0);

  return `$${total.toFixed(2)}`;
}
```
