---
id: 6a15cadf5f240d05a264955c
title: "Changamoto 318: Mabadiliko ya DNA"
challengeType: 28
dashedName: challenge-318
---

# --description--

Ukizingatia safu mbili za DNA zenye urefu sawa, rudisha safu ya data ya vielezo ambavyo safu hizo zinatofautiana (mabadiliko).

- Safu za DNA ni mfuatano wa herufi zilizoundwa na herufi `"A"`, `"T"`, `"C"`, na `"G"`
- Rudisha vielezo kwa mpangilio wa kuongezeka
- Ikiwa hakuna mabadiliko, rudisha safu ya data tupu

# --hints--

`detectMutations("ATCG", "ATGG")` inapaswa kurudisha `[2]`.

```js
assert.deepEqual(detectMutations("ATCG", "ATGG"), [2]);
```

`detectMutations("ATGCGTACGTTAGC", "ATGCATACGATTGC")` inapaswa kurudisha `[4, 9, 11]`.

```js
assert.deepEqual(detectMutations("ATGCGTACGTTAGC", "ATGCATACGATTGC"), [4, 9, 11]);
```

`detectMutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG")` inapaswa kurudisha `[]`.

```js
assert.deepEqual(detectMutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG"), []);
```

`detectMutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG")` inapaswa kurudisha `[15, 16, 17, 18]`.

```js
assert.deepEqual(detectMutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG"), [15, 16, 17, 18]);
```

`detectMutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT")` inapaswa kurudisha `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]`.

```js
assert.deepEqual(detectMutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]);
```

# --seed--

## --seed-contents--

```js
function detectMutations(strand1, strand2) {

  return strand1;
}
```

# --solutions--

```js
function detectMutations(strand1, strand2) {
  return strand1
    .split("")
    .reduce((mutations, base, i) => {
      if (base !== strand2[i]) mutations.push(i);
      return mutations;
    }, []);
}
```
