---
id: 6a151d32d772271dd2448c2d
title: "Zoezi la 311: Mchawi wa maneno"
challengeType: 28
dashedName: challenge-311
---

# --description--

Kutolewa mfuatano wa misimbo ya maneno unayoitumia, hesabu jumla ya alama.

Kila herufi katika mfuatano inawakilisha neno:

| Msimbo | Neno | Kategoria | Alama ya msingi |
| - | - | - | - |
| `"f"` | Moto | Uharibifu | 3 |
| `"l"` | Umeme | Uharibifu | 3 |
| `"i"` | Barafu | Udhibiti | 2 |
| `"w"` | Upepo | Udhibiti | 2 |
| `"h"` | Tiba | Urejeshaji | 1 |
| `"s"` | Mlinzi | Urejeshaji | 1 |

Kiongezaji cha mchanganyiko kinawekwa kulingana na ni maneno mangapi mfululizo yameitumiwa kutoka katika makategoria tofauti:

- Neno la kwanza daima linapata alama ya msingi.
- Kila neno linalofuata kutoka kategoria tofauti na lile lililotangulia huongeza kiongezaji kwa 1.
- Kutumia neno kutoka kategoria ile ile kama lile lililotangulia kunarejesha kiongezaji kuwa 1.
- Alama kwa kila neno ni alama yake ya msingi ikizidishwa na kiongezaji cha sasa.

Rudisha jumla ya alama kutoka kwa mfululizo wa maneno.

# --hints--

`cast("fihwl")` inapaswa kurudisha `33`.

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` inapaswa kurudisha `45`.

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` inapaswa kurudisha `37`.

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` inapaswa kurudisha `50`.

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` inapaswa kurudisha `101`.

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
