---
id: 6a151d32d772271dd2448c2c
title: "Zoezi la 310: Kiingereza cha Uingereza hadi Kiingereza cha Marekani"
challengeType: 28
dashedName: challenge-310
---

# --description--

Ukitoa sentensi, badilisha tahajia zozote za Kiingereza cha Uingereza kuwa sawa na za Kiingereza cha Marekani kwa kutumia jedwali la kutafuta lililofuata na rudisha sentensi iliyosasishwa:

| Kiingereza cha Uingereza | Kiingereza cha Marekani |
| - | - |
| `"colour"` | `"color"` |
| `"flavour"` | `"flavor"` |
| `"honour"` | `"honor"` |
| `"neighbour"` | `"neighbor"` |
| `"labour"` | `"labor"` |
| `"humour"` | `"humor"` |
| `"centre"` | `"center"` |
| `"fibre"` | `"fiber"` |
| `"defence"` | `"defense"` |
| `"offence"` | `"offense"` |
| `"organise"` | `"organize"` |
| `"recognise"` | `"recognize"` |
| `"analyse"` | `"analyze"` |

- Mabadiliko yanapaswa kuwa yasiyoegemea herufi kubwa au ndogo. Kwa mfano, `"Colour"` inapaswa kubadilika kuwa `"Color"`.
- Ingizo linaweza kuwa na maneno yanayojengwa juu ya tahajia halisi ya mzizi katika jedwali ambayo pia yanahitaji kubadilishwa. Kwa mfano, `"colouring"` inapaswa kubadilika kuwa `"coloring"`, na `"disorganised"` inapaswa kubadilika kuwa `"disorganized"`.

# --hints--

`britishToAmerican("I love the colour blue.")` inapaswa kurudisha `"I love the color blue."`

```js
assert.equal(britishToAmerican("I love the colour blue."), "I love the color blue.");
```

`britishToAmerican("The fibre optic cable is new.")` inapaswa kurudisha `"The fiber optic cable is new."`

```js
assert.equal(britishToAmerican("The fibre optic cable is new."), "The fiber optic cable is new.");
```

`britishToAmerican("It's an honour to meet someone with such humour.")` inapaswa kurudisha `"It's an honor to meet someone with such humor."`

```js
assert.equal(britishToAmerican("It's an honour to meet someone with such humour."), "It's an honor to meet someone with such humor.");
```

`britishToAmerican("The unrecognised artist analysed his colour palette at the centre.")` inapaswa kurudisha `"The unrecognized artist analyzed his color palette at the center."`

```js
assert.equal(britishToAmerican("The unrecognised artist analysed his colour palette at the centre."), "The unrecognized artist analyzed his color palette at the center.");
```

`britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful.")` inapaswa kurudisha `The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.`

```js
assert.equal(britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful."), "The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.");
```

# --seed--

## --seed-contents--

```js
function britishToAmerican(sentence) {

  return sentence;
}
```

# --solutions--

```js
function britishToAmerican(sentence) {
  const lookup = {
    colour: "color",
    flavour: "flavor",
    honour: "honor",
    neighbour: "neighbor",
    labour: "labor",
    humour: "humor",
    centre: "center",
    fibre: "fiber",
    defence: "defense",
    offence: "offense",
    organise: "organize",
    recognise: "recognize",
    analyse: "analyze"
  };

  let result = sentence;

  for (const [british, american] of Object.entries(lookup)) {
    const regex = new RegExp(british, "gi");
    result = result.replace(regex, match => {
      return match[0] === match[0].toUpperCase()
        ? american[0].toUpperCase() + american.slice(1)
        : american;
    });
  }

  return result;
}
```
