---
id: 6a15cadf5f240d05a264955a
title: "Zoezi la 316: Kuzungumza 1337"
challengeType: 28
dashedName: challenge-316
---

# --description--

Ukitoa mfuatano wa herufi ndogo, rudisha ukitafsiriwa kuwa lugha ya leet kwa kubadilisha herufi zilizo hapa chini na mbadala zao za leet:

| Herufi | Leet |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- Herufi ambazo hazina mbadala hazibadilishwi.

# --hints--

`makeLeet("cool")` inapaswa kurudisha `"c001"`.

```js
assert.equal(makeLeet("cool"), "c001");
```

`makeLeet("leet")` inapaswa kurudisha `"1337"`.

```js
assert.equal(makeLeet("leet"), "1337");
```

`makeLeet("hacker")` inapaswa kurudisha `"h4ck3r"`.

```js
assert.equal(makeLeet("hacker"), "h4ck3r");
```

`makeLeet("satellite")` inapaswa kurudisha `"547311173"`.

```js
assert.equal(makeLeet("satellite"), "547311173");
```

`makeLeet("abcdefghijklmnopqrstuvwxyz")` inapaswa kurudisha `"4bcd3f9h1jk1mn0pqr57uvwxyz"`.

```js
assert.equal(makeLeet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz");
```

# --seed--

## --seed-contents--

```js
function makeLeet(str) {

  return str;
}
```

# --solutions--

```js
function makeLeet(str) {
  const map = { a: "4", e: "3", g: "9", i: "1", l: "1", o: "0", t: "7", s: "5" };
  return str.split("").map(char => map[char] || char).join("");
}
```
