---
id: 6a15cadf5f240d05a264955f
title: "Changamoto 321: Ustafishaji wa Kipindi"
challengeType: 28
dashedName: challenge-321
---

# --description--

Ipe neno, amua kama linaweza kuandikwa kwa kutumia alama za vipengele kutoka kwenye jedwali la vipengele vya kipindi.

- Puuza tofauti za herufi kubwa na ndogo wakati wa kuandika neno. `"neon"` linaweza kuandikwa kwa alama `"Ne"`, `"O"`, na `"N"`.

Hapa kuna orodha kamili ya alama za vipengele:

```json
["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];
```

Rudisha safu ya data ya vipengele vilivyotumika kuandika neno, kwa herufi zao za asili na kwa mpangilio wa kuandika neno. Au, safu ya data tupu ikiwa haiwezi kuandikwa.

# --hints--

`getPeriodicSpelling("neon")` inapaswa kurudisha `["Ne", "O", "N"]`.

```js
assert.deepEqual(getPeriodicSpelling("neon"), ["Ne", "O", "N"]);
```

`getPeriodicSpelling("rational")` inapaswa kurudisha `["Ra", "Ti", "O", "N", "Al"]`.

```js
assert.deepEqual(getPeriodicSpelling("rational"), ["Ra", "Ti", "O", "N", "Al"]);
```

`getPeriodicSpelling("yarn")` inapaswa kurudisha `["Y", "Ar", "N"]`.

```js
assert.deepEqual(getPeriodicSpelling("yarn"), ["Y", "Ar", "N"]);
```

`getPeriodicSpelling("carbon")` inapaswa kurudisha `["C", "Ar", "B", "O", "N"]` au `["Ca", "Rb", "O", "N"]`.

```js
const result = JSON.stringify(getPeriodicSpelling("carbon"));
const path1 = JSON.stringify(["C", "Ar", "B", "O", "N"]);
const path2 = JSON.stringify(["Ca", "Rb", "O", "N"]);
assert.isTrue(result === path1 || result === path2);
```


`getPeriodicSpelling("noisy")` inapaswa kurudisha `["N", "O", "I", "S", "Y"]` au `["No", "I", "S", "Y"]`.

```js
const result = JSON.stringify(getPeriodicSpelling("noisy"));
const path1 = JSON.stringify(["N", "O", "I", "S", "Y"]);
const path2 = JSON.stringify(["No", "I", "S", "Y"]);
assert.isTrue(result === path1 || result === path2);
```


`getPeriodicSpelling("bicycles")` inapaswa kurudisha `["B", "I", "C", "Y", "Cl", "Es"]` au `["Bi", "C", "Y", "Cl", "Es"]`.

```js
const result = JSON.stringify(getPeriodicSpelling("bicycles"));
const path1 = JSON.stringify(["B", "I", "C", "Y", "Cl", "Es"]);
const path2 = JSON.stringify(["Bi", "C", "Y", "Cl", "Es"]);
assert.isTrue(result === path1 || result === path2);
```

`getPeriodicSpelling("optics")` inapaswa kurudisha `["O", "P", "Ti", "C", "S"]`, `["O", "P", "Ti", "Cs"]`, `["O", "Pt", "I", "C", "S"]`, au `["O", "Pt", "I", "Cs"]`.

```js
const result = JSON.stringify(getPeriodicSpelling("optics"));
const path1 = JSON.stringify(["O", "P", "Ti", "C", "S"]);
const path2 = JSON.stringify(["O", "P", "Ti", "Cs"]);
const path3 = JSON.stringify(["O", "Pt", "I", "C", "S"]);
const path4 = JSON.stringify(["O", "Pt", "I", "Cs"]);
assert.isTrue(result === path1 || result === path2 || result === path3 || result === path4);
```

`getPeriodicSpelling("value")` inapaswa kurudisha `[]`.

```js
assert.deepEqual(getPeriodicSpelling("value"), []);
```

# --seed--

## --seed-contents--

```js
function getPeriodicSpelling(word) {

  return word;
}
```

# --solutions--

```js
function getPeriodicSpelling(word) {
  const elements = ["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];

  const lower = new Map(elements.map(e => [e.toLowerCase(), e]));

  function spell(word, i) {
    if (i === word.length) return [];
    const one = word.slice(i, i + 1);
    const two = word.slice(i, i + 2);
    if (lower.has(two)) {
      const rest = spell(word, i + 2);
      if (rest !== null) return [lower.get(two), ...rest];
    }
    if (lower.has(one)) {
      const rest = spell(word, i + 1);
      if (rest !== null) return [lower.get(one), ...rest];
    }
    return null;
  }

  return spell(word.toLowerCase(), 0) ?? [];
}
```
