---
id: 6a15cadf5f240d05a264955e
title: "Changamoto 320: Benki ya Damu"
challengeType: 28
dashedName: challenge-320
---

# --description--

Kutolewa safu ya data ya hesabu ya damu katika benki ya damu na safu ya data ya maombi ya aina ya damu ya wagonjwa, rudisha mfuatano wa herufi katika muundo wa `"X of Y patients served"`. Ambapo `X` ni idadi kubwa zaidi ya wagonjwa wanaoweza kupokea damu kutoka kwa hesabu ya benki, na `Y` ni jumla ya wagonjwa wote.

Kila maingizo katika safu zote mbili ni mojawapo ya aina zifuatazo za damu: `"AB"`, `"A"`, `"B"`, au `"O"`.

Sheria za ulinganifu:

- `"AB"` inaweza kupokea kutoka kwa aina yoyote ya damu.
- `"A"` inaweza kupokea kutoka kwa `"A"` na `"O"`.
- `"B"` inaweza kupokea kutoka kwa `"B"` na `"O"`.
- `"O"` inaweza kupokea tu kutoka kwa `"O"`.

Maingizo yanayojirudia katika safu zilizotolewa yanawakilisha wingi.

# --hints--

`triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` inapaswa kurudisha `"4 of 4 patients served"`.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` inapaswa kurudisha `"3 of 5 patients served"`.

```js
assert.equal(triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served");
```

`triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` inapaswa kurudisha `"4 of 5 patients served"`.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served");
```

`triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` inapaswa kurudisha `"4 of 4 patients served"`.

```js
assert.equal(triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` inapaswa kurudisha `"8 of 11 patients served"`.

```js
assert.equal(triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served");
```

`triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` inapaswa kurudisha `"13 of 14 patients served"`.

```js
assert.equal(triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served");
```

# --seed--

## --seed-contents--

```js
function triageBlood(bank, patients) {

  return bank;
}
```

# --solutions--

```js
function triageBlood(bank, patients) {
  const b = {};
  const p = {};
  for (const t of bank) b[t] = (b[t] || 0) + 1;
  for (const t of patients) p[t] = (p[t] || 0) + 1;

  const serve = (patient, donors) => {
    for (const donor of donors) {
      const n = Math.min(p[patient] || 0, b[donor] || 0);
      p[patient] = (p[patient] || 0) - n;
      b[donor] = (b[donor] || 0) - n;
      saved += n;
    }
  };

  let saved = 0;
  serve("O",  ["O"]);
  serve("A",  ["A", "O"]);
  serve("B",  ["B", "O"]);
  serve("AB", ["AB", "A", "B", "O"]);
  return `${saved} of ${patients.length} patients served`;
}
```
