---
id: 6a15cadf5f240d05a2649558
title: "Changamoto 314: Gawanyio msingi"
challengeType: 28
dashedName: challenge-314
---

# --description--

Ikiwa umepewa nambari kamili kubwa kuliko 1, rudisha gawanyio msingi wake kama safu ya nambari kwa mpangilio wa kuongezeka.

Gawanyio msingi ni seti ya namba msingi zinazozidiana pamoja kutoa nambari kamili iliyotolewa. Kila nambari ina seti moja tu. Kwa mfano, gawanyio msingi la 20 ni `[2, 2, 5]` kwa sababu 2 * 2 * 5 = 20.

Ikiwa nambari kamili iliyotolewa ni namba msingi yenyewe, irudishe katika safu yenye kipengele kimoja.

# --hints--

`primeFactorization(20)` inapaswa kurudisha `[2, 2, 5]`.

```js
assert.deepEqual(primeFactorization(20), [2, 2, 5]);
```

`primeFactorization(17)` inapaswa kurudisha `[17]`.

```js
assert.deepEqual(primeFactorization(17), [17]);
```

`primeFactorization(15)` inapaswa kurudisha `[3, 5]`.

```js
assert.deepEqual(primeFactorization(15), [3, 5]);
```

`primeFactorization(35)` inapaswa kurudisha `[5, 7]`.

```js
assert.deepEqual(primeFactorization(35), [5, 7]);
```

`primeFactorization(999)` inapaswa kurudisha `[3, 3, 3, 37]`.

```js
assert.deepEqual(primeFactorization(999), [3, 3, 3, 37]);
```

`primeFactorization(360)` inapaswa kurudisha `[2, 2, 2, 3, 3, 5]`.

```js
assert.deepEqual(primeFactorization(360), [2, 2, 2, 3, 3, 5]);
```

`primeFactorization(510510)` inapaswa kurudisha `[2, 3, 5, 7, 11, 13, 17]`.

```js
assert.deepEqual(primeFactorization(510510), [2, 3, 5, 7, 11, 13, 17]);
```

# --seed--

## --seed-contents--

```js
function primeFactorization(n) {

  return n;
}
```

# --solutions--

```js
function primeFactorization(n) {
  const factors = [];
  let divisor = 2;

  while (divisor * divisor <= n) {
    while (n % divisor === 0) {
      factors.push(divisor);
      n /= divisor;
    }
    divisor++;
  }

  if (n > 1) factors.push(n);

  return factors;
}
```
