---
id: 6a19dd69388e5a6e59e4f2ef
title: "Changamoto 326: Faida Kuu"
challengeType: 28
dashedName: challenge-326
---

# --description--

Ukizingatia safu ya data ya bei za hisa kila siku na bajeti (katika dola), hesabu faida kubwa zaidi unayoweza kupata kwa kununua na kuuza hisa katika kipindi kilichotolewa.

- Unaweza kuuza tu baada ya kununua.
- Unaweza kufanya ununuzi mmoja na muamala mmoja wa kuuza tu. Mara unapozaa, huwezi kununua tena.
- Unaweza kununua hisa kamili tu.
- Rudisha faida kubwa zaidi inayowezekana kama mfuatano wa herufi, ukizungusha chini hadi senti inayokaribu na kuipanga kwa sehemu mbili za desimali.

# --hints--

`getMaxProfit([5, 6], 50)` inapaswa kurudisha `"10.00"`.

```js
assert.equal(getMaxProfit([5, 6], 50), "10.00");
```

`getMaxProfit([8, 2, 5, 10], 20)` inapaswa kurudisha `"80.00"`.

```js
assert.equal(getMaxProfit([8, 2, 5, 10], 20), "80.00");
```

`getMaxProfit([4, 5, 3, 6], 20)` inapaswa kurudisha `"18.00"`.

```js
assert.equal(getMaxProfit([4, 5, 3, 6], 20), "18.00");
```

`getMaxProfit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200)` inapaswa kurudisha `"8.31"`.

```js
assert.equal(getMaxProfit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200), "8.31");
```

`getMaxProfit([15.38, 15.01, 14.99, 14.62, 14.28], 80)` inapaswa kurudisha `"0.00"`.

```js
assert.equal(getMaxProfit([15.38, 15.01, 14.99, 14.62, 14.28], 80), "0.00");
```

`getMaxProfit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25)` inapaswa kurudisha `"73.80"`.

```js
assert.equal(getMaxProfit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25), "73.80");
```

# --seed--

## --seed-contents--

```js
function getMaxProfit(prices, budget) {

  return prices;
}
```

# --solutions--

```js
function getMaxProfit(prices, budget) {
  let maxProfit = 0;

  for (let buy = 0; buy < prices.length - 1; buy++) {
    for (let sell = buy + 1; sell < prices.length; sell++) {
      const shares = Math.floor(budget / prices[buy]);
      const profit = shares * (prices[sell] - prices[buy]);
      if (profit > maxProfit) maxProfit = profit;
    }
  }

  return (Math.floor(maxProfit * 100) / 100).toFixed(2);
}
```
