---
id: 6a19b1062d1b153d8ac76d73
title: "Zoezi la 325: Nambari ya Bahati"
challengeType: 28
dashedName: challenge-325
---

# --description--

Kutolewa mfuatano wa herufi wa jina la kwanza na la mwisho la mtu, hesabu nambari yake ya bahati kwa kutumia sheria zifuatazo:

- Majina ya kwanza na ya mwisho yamegawanyika kwa nafasi
- Tafuta idadi ya vokali na konsonanti kwa kila jina
- Zidisha idadi ndogo ya vokali na konsonanti kwa kila mmoja kisha zidisha na urefu wa jina dogo
- Fanya vivyo hivyo kwa idadi kubwa mbili na jina kubwa
- Toa thamani ndogo kutoka kwa kubwa kupata nambari yao ya bahati

Ikiwa thamani ya mwisho ni sifuri (`0`), rudisha `13`.

# --hints--

`getLuckyNumber("John Doe")` inapaswa kurudisha `21`.

```js
assert.equal(getLuckyNumber("John Doe"), 21);
```

`getLuckyNumber("Olivia Lewis")` inapaswa kurudisha `52`.

```js
assert.equal(getLuckyNumber("Olivia Lewis"), 52);
```

`getLuckyNumber("James Wilson")` inapaswa kurudisha `18`.

```js
assert.equal(getLuckyNumber("James Wilson"), 18);
```

`getLuckyNumber("Elizabeth Hernandez")` inapaswa kurudisha `81`.

```js
assert.equal(getLuckyNumber("Elizabeth Hernandez"), 81);
```

`getLuckyNumber("Mike Walker")` inapaswa kurudisha `32`.

```js
assert.equal(getLuckyNumber("Mike Walker"), 32);
```

`getLuckyNumber("Chloe Perez")` inapaswa kurudisha `13`.

```js
assert.equal(getLuckyNumber("Chloe Perez"), 13);
```

# --seed--

## --seed-contents--

```js
function getLuckyNumber(name) {

  return name;
}
```

# --solutions--

```js
function getLuckyNumber(name) {
  const vowels = "aeiouAEIOU";

  function nameStats(str) {
    const vowelCount = str.split("").filter(c => vowels.includes(c)).length;
    const consonantCount = str.split("").filter(c => c.match(/[a-zA-Z]/) && !vowels.includes(c)).length;
    return { vowelCount, consonantCount, length: str.length };
  }

  const [first, last] = name.split(" ");
  const firstStats = nameStats(first);
  const lastStats = nameStats(last);

  const largerVowels = Math.max(firstStats.vowelCount, lastStats.vowelCount);
  const smallerVowels = Math.min(firstStats.vowelCount, lastStats.vowelCount);
  const largerConsonants = Math.max(firstStats.consonantCount, lastStats.consonantCount);
  const smallerConsonants = Math.min(firstStats.consonantCount, lastStats.consonantCount);
  const largerLength = Math.max(firstStats.length, lastStats.length);
  const smallerLength = Math.min(firstStats.length, lastStats.length);

  const largerValue = largerVowels * largerConsonants * largerLength;
  const smallerValue = smallerVowels * smallerConsonants * smallerLength;
  const result = largerValue - smallerValue;

  return result === 0 ? 13 : result;
}
```
