---
id: 6a19b1062d1b153d8ac76d72
title: "Zoezi la 324: Hesabu ya Herufi Zinazojirudia"
challengeType: 28
dashedName: challenge-324
---

# --description--

Ukitoa mfuatano wa herufi miwili, rudisha hesabu ya herufi kutoka mfuatano wa pili zinazopatikana katika wa kwanza.

- Herufi zinazojirudia katika mfuatano wa pili zinahesabiwa tofauti.

# --hints--

`duplicateCharacterCount("aloha", "hei")` inapaswa kurudisha `1`.

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")` inapaswa kurudisha `4`.

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")` inapaswa kurudisha `3`.

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")` inapaswa kurudisha `0`.

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")` inapaswa kurudisha `5`.

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")` inapaswa kurudisha `2`.

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")` inapaswa kurudisha `26`.

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
