---
id: 6a1d9f98e819ed70a0e994df
title: "Changamoto 333: Uainishaji wa Masuala 2"
challengeType: 28
dashedName: challenge-333
---

# --description--

Kutolewa kichwa cha suala na safu ya lebo za sasa, rudisha safu iliyosasishwa ya lebo kulingana na sheria zifuatazo:

Kama suala halina lebo yoyote, ongeza:

- `"bug"` na `"needs triage"` ikiwa kichwa kina `"error"` au `"bug"`
- `"enhancement"` na `"discussing"` ikiwa kichwa kina `"feature"` au `"add"`

Vinginevyo, kama lebo zilizotolewa zina:

- `"needs triage"` na kichwa kina `"simple"` au `"easy"`, ondoa `"needs triage"` na ongeza `"good first issue"`
- `"discussing"` na kichwa kina `"planned"` au `"next"`, ondoa `"discussing"` na ongeza `"on the roadmap"`
- Vinginevyo, kama `"needs triage"` au `"discussing"` ipo, iondoe na ongeza `"help wanted"`

Kama kichwa kina:

- `"security"`, ongeza lebo ya `"critical"`

# --hints--

`triageIssue("app crashes with error", [])` inapaswa kurudisha `["bug", "needs triage"]`.

```js
assert.deepEqual(triageIssue("app crashes with error", []).sort(), ["bug", "needs triage"].sort());
```

`triageIssue("app crashes with error", ["bug", "needs triage"])` inapaswa kurudisha `["bug", "help wanted"]`.

```js
assert.deepEqual(triageIssue("app crashes with error", ["bug", "needs triage"]).sort(), ["bug", "help wanted"].sort());
```

`triageIssue("add dark mode", [])` inapaswa kurudisha `["enhancement", "discussing"]`.

```js
assert.deepEqual(triageIssue("add dark mode", []).sort(), ["enhancement", "discussing"].sort());
```

`triageIssue("add dark mode", ["enhancement", "discussing"])` inapaswa kurudisha `["enhancement", "help wanted"]`.

```js
assert.deepEqual(triageIssue("add dark mode", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted"].sort());
```

`triageIssue("xss security bug", [])` inapaswa kurudisha `["bug", "needs triage", "critical"]`.

```js
assert.deepEqual(triageIssue("xss security bug", []).sort(), ["bug", "needs triage", "critical"].sort());
```

`triageIssue("security vulnerability in auth", [])` inapaswa kurudisha `["critical"]`.

```js
assert.deepEqual(triageIssue("security vulnerability in auth", []).sort(), ["critical"].sort());
```

`triageIssue("easy a11y fix", ["bug", "needs triage"])` inapaswa kurudisha `["bug", "good first issue"]`.

```js
assert.deepEqual(triageIssue("easy a11y fix", ["bug", "needs triage"]).sort(), ["bug", "good first issue"].sort());
```

`triageIssue("planned api migration", ["enhancement", "discussing"])` inapaswa kurudisha `["enhancement", "on the roadmap"]`.

```js
assert.deepEqual(triageIssue("planned api migration", ["enhancement", "discussing"]).sort(), ["enhancement", "on the roadmap"].sort());
```

`triageIssue("improve security", ["enhancement", "discussing"])` inapaswa kurudisha `["enhancement", "help wanted", "critical"]`.

```js
assert.deepEqual(triageIssue("improve security", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted", "critical"].sort());
```

# --seed--

## --seed-contents--

```js
function triageIssue(title, labels) {

  return title;
}
```

# --solutions--

```js
function triageIssue(title, labels) {
  const result = [...labels];
  const t = title.toLowerCase();

  const has = label => result.includes(label);
  const add = label => { if (!has(label)) result.push(label); };
  const remove = label => { const i = result.indexOf(label); if (i > -1) result.splice(i, 1); };

  if (result.length === 0) {
    if (t.includes("error") || t.includes("bug")) {
      add("bug"); add("needs triage");
    } else if (t.includes("feature") || t.includes("add")) {
      add("enhancement"); add("discussing");
    }
  } else {
    if (has("needs triage") && (t.includes("simple") || t.includes("easy"))) {
      remove("needs triage"); add("good first issue");
    } else if (has("discussing") && (t.includes("planned") || t.includes("next"))) {
      remove("discussing"); add("on the roadmap");
    } else if (has("needs triage") || has("discussing")) {
      remove("needs triage"); remove("discussing"); add("help wanted");
    }
  }

  if (t.includes("security")) {
    add("critical");
  }

  return result;
}
```
