---
id: 6a1d9f98e819ed70a0e994dd
title: "Zoezi la 331: Mara ya Karibu Zaidi"
challengeType: 28
dashedName: challenge-331
---

# --description--

Ukizingatia nambari kamili mbili, punguza ya kwanza hadi mara ya karibu zaidi ya ya pili.

# --hints--

`roundToNearestMultiple(5, 3)` inapaswa kurudisha `6`.

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)` inapaswa kurudisha `16`.

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)` inapaswa kurudisha `45`.

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)` inapaswa kurudisha `33`.

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)` inapaswa kurudisha `96`.

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
