---
id: 6a1d9f98e819ed70a0e994dc
title: "Zoezi la 330: maneno ya herufi ndogo"
challengeType: 28
dashedName: challenge-330
---

# --description--

Ukipata mfuatano wa herufi, rudisha maneno ambayo yote ni ya herufi ndogo tu, kwa mpangilio wake wa awali na ukiwa na nafasi kati ya kila neno.

# --hints--

`getLowercaseWords("hello GOOD world")` inapaswa kurudisha `"hello world"`.

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")` inapaswa kurudisha `"these are all lowercase"`.

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")` inapaswa kurudisha `"less is more"`.

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")` inapaswa kurudisha `"eat pizza every day"`.

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` inapaswa kurudisha `"the quick brown fox jumped over the lazy dog"`.

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
