---
id: 6a1d9f98e819ed70a0e994e2
title: "Zoezi la 336: Ulinganifu wa Horoscope"
challengeType: 28
dashedName: challenge-336
---

# --description--

Ukitoa mfuatano wa herufi wa alama mbili za nyota, rudisha asilimia yao ya ulinganifu.

Alama zimepangwa katika gurudumu lenye nafasi 12 kwa mpangilio huu: `"Aries"`, `"Taurus"`, `"Gemini"`, `"Cancer"`, `"Leo"`, `"Virgo"`, `"Libra"`, `"Scorpio"`, `"Sagittarius"`, `"Capricorn"`, `"Aquarius"`, `"Pisces"`, zikizunguka kurudi `"Aries"` baada ya `"Pisces"`. Tafuta umbali mfupi zaidi kati ya alama hizo mbili na rudisha ulinganifu:

| Umbali | Ulinganifu |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscopeMatch("Libra", "Sagittarius")` inapaswa kurudisha `"80%"`.

```js
assert.equal(horoscopeMatch("Libra", "Sagittarius"), "80%");
```

`horoscopeMatch("Gemini", "Scorpio")` inapaswa kurudisha `"20%"`.

```js
assert.equal(horoscopeMatch("Gemini", "Scorpio"), "20%");
```

`horoscopeMatch("Pisces", "Aries")` inapaswa kurudisha `"40%"`.

```js
assert.equal(horoscopeMatch("Pisces", "Aries"), "40%");
```

`horoscopeMatch("Capricorn", "Cancer")` inapaswa kurudisha `"50%"`.

```js
assert.equal(horoscopeMatch("Capricorn", "Cancer"), "50%");
```

`horoscopeMatch("Aquarius", "Aquarius")` inapaswa kurudisha `"100%"`.

```js
assert.equal(horoscopeMatch("Aquarius", "Aquarius"), "100%");
```

`horoscopeMatch("Virgo", "Taurus")` inapaswa kurudisha `"90%"`.

```js
assert.equal(horoscopeMatch("Virgo", "Taurus"), "90%");
```

`horoscopeMatch("Leo", "Scorpio")` inapaswa kurudisha `"30%"`.

```js
assert.equal(horoscopeMatch("Leo", "Scorpio"), "30%");
```

# --seed--

## --seed-contents--

```js
function horoscopeMatch(sign1, sign2) {

  return sign1;
}
```

# --solutions--

```js
function horoscopeMatch(sign1, sign2) {
  const signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"];
  
  const a = signs.indexOf(sign1) + 1;
  const b = signs.indexOf(sign2) + 1;
  const diff = Math.abs(a - b);
  const distance = Math.min(diff, 12 - diff);

  const compatibility = { 0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%" };
  return compatibility[distance];
}
```
