---
id: 6a1d9f98e819ed70a0e994e0
title: "Changamoto 334: Mabadiliko Sahihi"
challengeType: 28
dashedName: challenge-334
---

# --description--

Ukitoa kiasi kamili cha nambari kamili katika senti, rudisha idadi ya njia tofauti za kutengeneza mabadiliko sahihi kwa kutumia pennies (1 senti), nickels (5 senti), dimes (10 senti), na quarters (25 senti).

# --hints--

`exactChange(3)` inapaswa kurudisha `1`.

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)` inapaswa kurudisha `2`.

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)` inapaswa kurudisha `6`.

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)` inapaswa kurudisha `24`.

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)` inapaswa kurudisha `73`.

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)` inapaswa kurudisha `213`.

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
