---
id: 6a1d9f98e819ed70a0e994de
title: "Zoezi la 332: Usimamizi wa Masuala"
challengeType: 28
dashedName: challenge-332
---

# --description--

Kutolewa idadi ya millisekunde tangu chapisho la mwisho kuhusu suala fulani, na ujumbe wa mwisho uliotumwa kuhusu suala hilo, amua unapaswa kufanya nini na suala hilo kulingana na sheria hizi:

- Ikiwa ujumbe wa mwisho ni chini ya siku 7 zilizopita, rudisha `"leave it"`
- Ikiwa ujumbe wa mwisho ni siku 7 au zaidi zilizopita na maudhui yake yana `"bump"` (bila kujali herufi kubwa au ndogo), rudisha `"close it"`
- Vinginevyo, rudisha `"bump it"`

# --hints--

`triageIssue(86400000, "Lets fix it")` inapaswa kurudisha `"leave it"`.

```js
assert.equal(triageIssue(86400000, "Lets fix it"), "leave it");
```

`triageIssue(1209600000, "still waiting")` inapaswa kurudisha `"bump it"`.

```js
assert.equal(triageIssue(1209600000, "still waiting"), "bump it");
```

`triageIssue(864000000, "bump")` inapaswa kurudisha `"close it"`.

```js
assert.equal(triageIssue(864000000,  "bump"), "close it");
```

`triageIssue(604800000, "Do we still want this?")` inapaswa kurudisha `"bump it"`.

```js
assert.equal(triageIssue(604800000, "Do we still want this?"), "bump it");
```

`triageIssue(604800000, "Bumping this")` inapaswa kurudisha `"close it"`.

```js
assert.equal(triageIssue(604800000, "Bumping this"), "close it");
```

`triageIssue(345600000, "I'll make a PR")` inapaswa kurudisha `"leave it"`.

```js
assert.equal(triageIssue(345600000, "I'll make a PR"), "leave it");
```

# --seed--

## --seed-contents--

```js
function triageIssue(ms, message) {

  return ms;
}
```

# --solutions--

```js
function triageIssue(ms, message) {
  const sevenDays = 7 * 24 * 60 * 60 * 1000;

  if (ms < sevenDays) return "leave it";
  if (message.toLowerCase().includes("bump")) return "close it";
  return "bump it";
}
```
