---
id: 6a1d9f98e819ed70a0e994e1
title: "Changamoto 335: Kadi Tano za Yatima"
challengeType: 28
dashedName: challenge-335
---

# --description--

Kutolewa safu ya data ya kadi tano za yatima zenye thamani 1-6, rudisha mkono bora zaidi unaowezekana.

Hapa kuna mikono iliyopangwa kutoka chini hadi juu:

| Mkono | Maelezo |
|------|-------------|
| `"no pair"` | Hakuna jozi au bora zaidi |
| `"pair"` | Kadi mbili zenye thamani sawa |
| `"two pair"` | Jozi mbili tofauti |
| `"three of a kind"` | Kadi tatu zenye thamani sawa |
| `"small straight"` | Thamani nne mfululizo |
| `"large straight"` | Thamani tano mfululizo |
| `"full house"` | Kadi tatu sawa na jozi moja |
| `"four of a kind"` | Kadi nne zenye thamani sawa |
| `"five of a kind"` | Kadi zote tano zenye thamani sawa |

# --hints--

`fiveDice([1, 1, 1, 1, 1])` inapaswa kurudisha `"five of a kind"`.

```js
assert.equal(fiveDice([1, 1, 1, 1, 1]), "five of a kind");
```

`fiveDice([5, 5, 5, 6, 5])` inapaswa kurudisha `"four of a kind"`.

```js
assert.equal(fiveDice([5, 5, 5, 6, 5]), "four of a kind");
```

`fiveDice([2, 5, 6, 4, 3])` inapaswa kurudisha `"large straight"`.

```js
assert.equal(fiveDice([2, 5, 6, 4, 3]), "large straight");
```

`fiveDice([4, 3, 3, 3, 1])` inapaswa kurudisha `"three of a kind"`.

```js
assert.equal(fiveDice([4, 3, 3, 3, 1]), "three of a kind");
```

`fiveDice([4, 6, 2, 6, 5])` inapaswa kurudisha `"pair"`.

```js
assert.equal(fiveDice([4, 6, 2, 6, 5]), "pair");
```

`fiveDice([1, 4, 5, 6, 2])` inapaswa kurudisha `"no pair"`.

```js
assert.equal(fiveDice([1, 4, 5, 6, 2]), "no pair");
```

`fiveDice([1, 3, 4, 6, 2])` inapaswa kurudisha `"small straight"`.

```js
assert.equal(fiveDice([1, 3, 4, 6, 2]), "small straight");
```

`fiveDice([2, 2, 5, 2, 5])` inapaswa kurudisha `"full house"`.

```js
assert.equal(fiveDice([2, 2, 5, 2, 5]), "full house");
```

`fiveDice([6, 4, 5, 6, 4])` inapaswa kurudisha `"two pair"`.

```js
assert.equal(fiveDice([6, 4, 5, 6, 4]), "two pair");
```

# --seed--

## --seed-contents--

```js
function fiveDice(dice) {

  return dice;
}
```

# --solutions--

```js
function fiveDice(dice) {
  const counts = {};
  for (const d of dice) counts[d] = (counts[d] || 0) + 1;
  const values = Object.values(counts).sort((a, b) => b - a);
  const keys = Object.keys(counts).map(Number).sort((a, b) => a - b);

  const isLargeStraight = keys.length === 5 && keys[4] - keys[0] === 4;
  const isSmallStraight = !isLargeStraight && (() => {
    for (let i = 0; i <= 1; i++) {
      const slice = keys.slice(i, i + 4);
      if (slice.length === 4 && slice[3] - slice[0] === 3 &&
        new Set(slice).size === 4) return true;
    }
    return false;
  })();

  if (values[0] === 5) return "five of a kind";
  if (values[0] === 4) return "four of a kind";
  if (values[0] === 3 && values[1] === 2) return "full house";
  if (isLargeStraight) return "large straight";
  if (isSmallStraight) return "small straight";
  if (values[0] === 3) return "three of a kind";
  if (values[0] === 2 && values[1] === 2) return "two pair";
  if (values[0] === 2) return "pair";
  return "no pair";
}
```
