---
id: 6a1d9f98e819ed70a0e994da
title: "Zoezi la 328: Utaratibu wa Kaprekar"
challengeType: 28
dashedName: challenge-328
---

# --description--

Ukipata nambari yenye tarakimu 4, rudisha idadi ya mara unazohitaji kutumia utaratibu wa Kaprekar hadi kufikia 6174.

Utaratibu wa Kaprekar hufanya kazi kama ifuatavyo:

- Panga tarakimu kwa mpangilio wa kushuka ili kuunda nambari kubwa zaidi
- Panga tarakimu kwa mpangilio wa kuongezeka ili kuunda nambari ndogo zaidi (ongeza sifuri za mbele ikiwa ni lazima)
- Toa nambari ndogo kutoka kwa nambari kubwa
- Rudia kwa nambari mpya

# --hints--

`kaprekar(1234)` inapaswa kurudisha `3`.

```js
assert.equal(kaprekar(1234), 3);
```

`kaprekar(2025)` inapaswa kurudisha `6`.

```js
assert.equal(kaprekar(2025), 6);
```

`kaprekar(7173)` inapaswa kurudisha `4`.

```js
assert.equal(kaprekar(7173), 4);
```

`kaprekar(3164)` inapaswa kurudisha `7`.

```js
assert.equal(kaprekar(3164), 7);
```

`kaprekar(8082)` inapaswa kurudisha `2`.

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
