---
id: 6a1d9f98e819ed70a0e994d9
title: "Zoezi la 327: Hajirisha hifadhidata"
challengeType: 28
dashedName: challenge-327
---

# --description--

Ukitoa vitu viwili vya hifadhidata, rudisha kitu cha pili kikiwa na vigezo vyovyote vilivyokosekana kutoka kwa cha kwanza vimejazwa.

- Sehemu ambazo tayari zipo katika rekodi hazipaswi kuandikwa upya.

# --hints--

`migrateRecord({ username: "", posts: 0 }, { verified: true })` inapaswa kurudisha `{ username: "", posts: 0, verified: true }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { verified: true }), { username: "", posts: 0, verified: true });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 })` inapaswa kurudisha `{ username: "camper", posts: 5 }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 }), { username: "camper", posts: 5 });
```

`migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" })` inapaswa kurudisha `{ username: "camper", posts: 0, verified: false }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" }), { username: "camper", posts: 0, verified: false });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" })` inapaswa kurudisha `{ username: "camper", role: "admin", posts: 0 }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" }), { username: "camper", role: "admin", posts: 0 });
```

`migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" })` inapaswa kurudisha `{ username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false }`.

```js
assert.deepEqual(migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" }), { username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false });
```

# --seed--

## --seed-contents--

```js
function migrateRecord(schema, record) {

  return schema;
}
```

# --solutions--

```js
function migrateRecord(schema, record) {
  return { ...schema, ...record };
}
```
