---
id: 6a2037a68a0bc2aef0006002
title: "Zoezi la 340: Mbadalishaji wa Pig Latin"
challengeType: 28
dashedName: challenge-340
---

# --description--

Kutolewa mfuatano wa herufi, ubadilishe kuwa Pig Latin ukitumia sheria zifuatazo:

- Ikiwa neno linaanza na vokali (`"a"`, `"e"`, `"i"`, `"o"`, au `"u"`), ongeza `"way"` mwishoni. Kwa mfano, `"universe"` hubadilika kuwa `"universeway"`.
- Ikiwa neno linaanza na konsonanti moja au zaidi, hamisha konsonanti hizo mwishoni na ongeza `"ay"`. Kwa mfano, `"hello"` hubadilika kuwa `"ellohay"`.
- Hifadhi herufi ya kwanza kama ilivyo. Kwa mfano, `"Hello"` hubadilika kuwa `"Ellohay"`.

# --hints--

`pigLatin("universe")` inapaswa kurudisha `"universeway"`.

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")` inapaswa kurudisha `"ellohay"`.

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")` inapaswa kurudisha `"ellohay universeway"`.

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")` inapaswa kurudisha `"Ellohay universeway"`.

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")` inapaswa kurudisha `"Igpay Atinlay isway unfay"`.

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")` inapaswa kurudisha `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
