---
id: 6a2037a68a0bc2aef0006004
title: "Zoezi la 342: Matarajio ya Kura za Kadi"
challengeType: 28
dashedName: challenge-342
---

# --description--

Ukizingatia idadi ya kadi za pande sita za kurusha na jumla lengwa, rudisha matarajio ya kurusha jumla hiyo kama mfuatano wa herufi katika muundo wa `"1 in X"`.

- Idadi ya kadi itakuwa kati ya 1 na 6.
- Jumla lengwa daima inaweza kufikiwa kwa idadi iliyotolewa ya kadi.
- Zungusha `"X"` hadi nambari kamili iliyo karibu zaidi.

# --hints--

`getOdds(1, 5)` inapaswa kurudisha `"1 in 6"`.

```js
assert.equal(getOdds(1, 5), "1 in 6");
```

`getOdds(2, 4)` inapaswa kurudisha `"1 in 12"`.

```js
assert.equal(getOdds(2, 4), "1 in 12");
```

`getOdds(3, 10)` inapaswa kurudisha `"1 in 8"`.

```js
assert.equal(getOdds(3, 10), "1 in 8");
```

`getOdds(4, 7)` inapaswa kurudisha `"1 in 65"`.

```js
assert.equal(getOdds(4, 7), "1 in 65");
```

`getOdds(5, 26)` inapaswa kurudisha `"1 in 111"`.

```js
assert.equal(getOdds(5, 26), "1 in 111");
```

`getOdds(6, 35)` inapaswa kurudisha `"1 in 7776"`.

```js
assert.equal(getOdds(6, 35), "1 in 7776");
```

# --seed--

## --seed-contents--

```js
function getOdds(dice, target) {

  return dice;
}
```

# --solutions--

```js
function getOdds(dice, target) {
  let validCombos = 0;
  const totalCombos = Math.pow(6, dice);

  function roll(diceLeft, sum) {
    if (diceLeft === 0) {
      if (sum === target) validCombos++;
      return;
    }
    for (let face = 1; face <= 6; face++) {
      roll(diceLeft - 1, sum + face);
    }
  }

  roll(dice, 0);
  return `1 in ${Math.round(totalCombos / validCombos)}`;
}
```
