---
id: 6a2037a68a0bc2aef0006007
title: "Zoezi la 345: Mchanganyiko wa Maneno"
challengeType: 28
dashedName: challenge-345
---

# --description--

Ukitoa maneno mawili, rudisha neno jipya kwa kuunganisha nusu ya kwanza ya neno la kwanza na nusu ya pili ya neno la pili.

- Kwa maneno yenye urefu usio wa kawaida, nusu ya kwanza ni nusu fupi zaidi.

# --hints--

`blendWords("turtle", "toucan")` inapaswa kurudisha `"turcan"`.

```js
assert.equal(blendWords("turtle", "toucan"), "turcan");
```

`blendWords("chipmunk", "flamingo")` inapaswa kurudisha `"chipingo"`.

```js
assert.equal(blendWords("chipmunk", "flamingo"), "chipingo");
```

`blendWords("falcon", "pelican")` inapaswa kurudisha `"falican"`.

```js
assert.equal(blendWords("falcon", "pelican"), "falican");
```

`blendWords("hyena", "iguana")` inapaswa kurudisha `"hyana"`.

```js
assert.equal(blendWords("hyena", "iguana"), "hyana");
```

`blendWords("scorpion", "gorilla")` inapaswa kurudisha `"scorilla"`.

```js
assert.equal(blendWords("scorpion", "gorilla"), "scorilla");
```

`blendWords("platypus", "wolverine")` inapaswa kurudisha `"platerine"`.

```js
assert.equal(blendWords("platypus", "wolverine"), "platerine");
```

# --seed--

## --seed-contents--

```js
function blendWords(word1, word2) {

  return word1;
}
```

# --solutions--

```js
function blendWords(word1, word2) {
  const firstHalf = word1.slice(0, Math.floor(word1.length / 2));
  const secondHalf = word2.slice(Math.floor(word2.length / 2));
  return firstHalf + secondHalf;
}
```
