---
id: 6a2037a68a0bc2aef0005fff
title: "Zoezi la 337: Kihesabu cha Tally"
challengeType: 28
dashedName: challenge-337
---

# --description--

Kutolewa mfuatano wa herufi wa alama za tally, rudisha jumla ya hesabu inayowakilishwa.

- Kila bomba `"|"` wakilisha hesabu moja.
- Kila alama ya tano inawakilishwa kama mstari wa mbele `"/"`, ukikamilisha kundi la tano (`"||||/"`).
- Makundi yanatenganishwa kwa nafasi.

# --hints--

`getTallyCount("||||")` inapaswa kurudisha `4`.

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` inapaswa kurudisha `5`.

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` inapaswa kurudisha `8`.

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` inapaswa kurudisha `17`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` inapaswa kurudisha `41`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
