---
id: 6a2037a68a0bc2aef0006006
title: "Zoezi la 344: Uwiano wa Dhahabu"
challengeType: 28
dashedName: challenge-344
---

# --description--

Ukizingatia nambari mbili, amua kama uwiano wao unakaribia uwiano wa dhahabu.

- Tumia uwiano wa dhahabu wa `1.618`
- Ruhusu uvumilivu wa `0.01`

# --hints--

`isGoldenRatio(21, 34)` inapaswa kurudisha `true`.

```js
assert.isTrue(isGoldenRatio(21, 34));
```

`isGoldenRatio(15, 20)` inapaswa kurudisha `false`.

```js
assert.isFalse(isGoldenRatio(15, 20));
```

`isGoldenRatio(8, 13)` inapaswa kurudisha `true`.

```js
assert.isTrue(isGoldenRatio(8, 13));
```

`isGoldenRatio(10, 16)` inapaswa kurudisha `false`.

```js
assert.isFalse(isGoldenRatio(10, 16));
```

`isGoldenRatio(1618, 1000)` inapaswa kurudisha `true`.

```js
assert.isTrue(isGoldenRatio(1618, 1000));
```

`isGoldenRatio(88, 55)` inapaswa kurudisha `false`.

```js
assert.isFalse(isGoldenRatio(88, 55));
```

# --seed--

## --seed-contents--

```js
function isGoldenRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function isGoldenRatio(a, b) {
  const goldenRatio = 1.618;
  const ratio = Math.max(a, b) / Math.min(a, b);
  return Math.abs(ratio - goldenRatio) <= 0.01;
}
```
