---
id: 6a2037a68a0bc2aef0006003
title: "Zoezi la 341: Hesabu ya siku za kuzaliwa"
challengeType: 28
dashedName: challenge-341
---

# --description--

Ukizingatia tarehe ya leo na siku ya kuzaliwa, rudisha idadi ya siku hadi siku ya kuzaliwa inayofuata ya mtu huyo.

- Tarehe ya leo hutolewa kama mfuatano wa herufi katika muundo wa `"YYYY-MM-DD"`, ukiwa na sifuri za kuongoza, kwa mfano: `"2026-07-16"`.
- Siku ya kuzaliwa hutolewa kama mfuatano wa herufi katika muundo wa `"M/D"`, bila sifuri za kuongoza, kwa mfano: `"9/7"`.
- Ikiwa leo ni siku yao ya kuzaliwa, rudisha idadi ya siku hadi siku yao inayofuata ya kuzaliwa (si `0`).
- Miaka ya kurukia inapaswa kuzingatiwa.

# --hints--

`daysUntilBirthday("2026-07-16", "9/7")` inapaswa kurudisha `53`.

```js
assert.equal(daysUntilBirthday("2026-07-16", "9/7"), 53);
```

`daysUntilBirthday("2026-07-16", "3/22")` inapaswa kurudisha `249`.

```js
assert.equal(daysUntilBirthday("2026-07-16", "3/22"), 249);
```

`daysUntilBirthday("2026-07-16", "7/16")` inapaswa kurudisha `365`.

```js
assert.equal(daysUntilBirthday("2026-07-16", "7/16"), 365);
```

`daysUntilBirthday("2024-02-28", "3/1")` inapaswa kurudisha `2`.

```js
assert.equal(daysUntilBirthday("2024-02-28", "3/1"), 2);
```

`daysUntilBirthday("2023-04-24", "12/30")` inapaswa kurudisha `250`.

```js
assert.equal(daysUntilBirthday("2023-04-24", "12/30"), 250);
```

`daysUntilBirthday("2024-03-01", "2/29")` inapaswa kurudisha `1460`.

```js
assert.equal(daysUntilBirthday("2024-03-01", "2/29"), 1460);
```

`daysUntilBirthday("2096-03-01", "2/29")` inapaswa kurudisha `2920`.

```js
assert.equal(daysUntilBirthday("2096-03-01", "2/29"), 2920);
```

# --seed--

## --seed-contents--

```js
function daysUntilBirthday(today, birthday) {

  return today;
}
```

# --solutions--

```js
function daysUntilBirthday(today, birthday) {
  const [year, month, day] = today.split("-").map(Number);
  const [bMonth, bDay] = birthday.split("/").map(Number);

  const todayDate = new Date(year, month - 1, day);

  function isLeapYear(y) {
    return (y % 4 === 0 && y % 100 !== 0) || y % 400 === 0;
  }

  function nextValidBirthday(fromYear) {
    for (let y = fromYear; y <= fromYear + 8; y++) {
      if (bMonth === 2 && bDay === 29 && !isLeapYear(y)) continue;
      const date = new Date(y, bMonth - 1, bDay);
      if (date > todayDate) return date;
    }
  }

  const nextBirthday = nextValidBirthday(year);
  const diff = nextBirthday - todayDate;
  return Math.round(diff / (1000 * 60 * 60 * 24));
}
```
