---
id: 6a2037a68a0bc2aef0006000
title: "Zoezi la 338: Kihesabu cha Umri wa Mnyama"
challengeType: 28
dashedName: challenge-338
---

# --description--

Ukitoa aina ya mnyama na umri wake kwa miaka ya binadamu, rudisha umri sawa katika miaka ya mnyama kwa kutumia jedwali la kubadilisha lifuatalo:

| Mnyama | Kizidishi |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`petYears("dog", 5)` inapaswa kurudisha `35`.

```js
assert.equal(petYears("dog", 5), 35);
```

`petYears("cat", 9)` inapaswa kurudisha `54`.

```js
assert.equal(petYears("cat", 9), 54);
```

`petYears("rabbit", 3)` inapaswa kurudisha `24`.

```js
assert.equal(petYears("rabbit", 3), 24);
```

`petYears("hamster", 4)` inapaswa kurudisha `120`.

```js
assert.equal(petYears("hamster", 4), 120);
```

`petYears("guinea pig", 5)` inapaswa kurudisha `60`.

```js
assert.equal(petYears("guinea pig", 5), 60);
```

`petYears("goldfish", 2)` inapaswa kurudisha `12`.

```js
assert.equal(petYears("goldfish", 2), 12);
```

`petYears("bird", 1)` inapaswa kurudisha `5`.

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
