---
id: 6a22d77ddf034bc4e35b1d58
title: "Zoezi la 351: Nambari ya Pronic"
challengeType: 28
dashedName: challenge-351
---

# --description--

Ukipata nambari, amua kama ni nambari ya pronic.

Nambari ya pronic ni matokeo ya kuzidisha nambari kamili mbili mfululizo. Kwa mfano, 6 ni pronic kwa sababu 2 * 3 = 6.

# --hints--

`isPronic(6)` inapaswa kurudisha `true`.

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` inapaswa kurudisha `false`.

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` inapaswa kurudisha `true`.

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` inapaswa kurudisha `true`.

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` inapaswa kurudisha `false`.

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` inapaswa kurudisha `true`.

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
