---
id: 6a22d77ddf034bc4e35b1d5a
title: "Changamoto 353: Kiwango cha Utofauti wa Rangi 2"
challengeType: 28
dashedName: challenge-353
---

# --description--

Ukizingatia thamani mbili za mwangaza wa uhusiano na thamani ya kweli au siyo kweli inayoonyesha kama maandishi ni makubwa, rudisha kiwango cha utofauti wa rangi cha WCAG kwa kutumia njia ifuatayo:

Hesabu uwiano wa utofauti kwa kuongeza 0.05 kwa kila thamani ya mwangaza, kisha gawanya ile yenye mwangaza zaidi kwa ile yenye mwangaza mdogo. Ile yenye mwangaza zaidi itakuwa hoja ya kwanza kila mara.

Rudisha kiwango kulingana na uwiano wa utofauti kwa kutumia jedwali lifuatalo:

| Kiwango | Maandishi ya Kawaida | Maandishi Makubwa |
|--------|--------------------|------------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | chini ya 4.5 | chini ya 3.0 |

# --hints--

`getContrastRating(1.0, 0.0, false)` inapaswa kurudisha `"AAA"`.

```js
assert.equal(getContrastRating(1.0, 0.0, false), "AAA");
```

`getContrastRating(0.9015, 0.1364, false)` inapaswa kurudisha `"AA"`.

```js
assert.equal(getContrastRating(0.9015, 0.1364, false), "AA");
```

`getContrastRating(0.8965, 0.1628, false)` inapaswa kurudisha `"Fail"`.

```js
assert.equal(getContrastRating(0.8965, 0.1628, false), "Fail");
```

`getContrastRating(0.7469, 0.0957, true)` inapaswa kurudisha `"AAA"`.

```js
assert.equal(getContrastRating(0.7469, 0.0957, true), "AAA");
```

`getContrastRating(0.7489, 0.2018, true)` inapaswa kurudisha `"AA"`.

```js
assert.equal(getContrastRating(0.7489, 0.2018, true), "AA");
```

`getContrastRating(0.6571, 0.1974, true)` inapaswa kurudisha `"Fail"`.

```js
assert.equal(getContrastRating(0.6571, 0.1974, true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(l1, l2, isLargeText) {

  return l1;
}
```

# --solutions--

```js
function getContrastRating(l1, l2, isLargeText) {
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
