---
id: 6a22d77ddf034bc4e35b1d54
title: "Changamoto 347: Nadharia ya Michezo"
challengeType: 28
dashedName: challenge-347
---

# --description--

Kutolewa mfuatano wa herufi wenye urefu sawa unaowakilisha mikakati ya wachezaji wawili kwa mchezo, rudisha alama kama safu ya data `[player1, player2]`.

- Mfuatano wa herufi uliotolewa utakuwa na herufi moja kati ya mbili tu: `"C"` (shirikiana) au `"D"` (kataa).
- Kila herufi inawakilisha raundi moja, na alama hutolewa kama ifuatavyo:
  - Ikiwa wachezaji wote wawili watashirikiana, kila mmoja atapata alama 3.
  - Ikiwa wachezaji wote wawili watakata, kila mmoja atapata alama 1.
  - Ikiwa mchezaji mmoja atakata na mwingine atashirikiana, mkatakata atapata alama 5 na mshirikaji atapata alama 0.

# --hints--

`playGame("CCCC", "CCCC")` inapaswa kurudisha `[12, 12]`.

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")` inapaswa kurudisha `[4, 4]`.

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")` inapaswa kurudisha `[5, 10]`.

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")` inapaswa kurudisha `[24, 34]`.

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` inapaswa kurudisha `[66, 21]`.

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
