---
id: 6a22d77ddf034bc4e35b1d5b
title: "Changamoto 354: Kiwango cha Utofauti 3"
challengeType: 28
dashedName: challenge-354
---

# --description--

Kutolewa safu mbili za data zinazowakilisha thamani za RGB na thamani ya kweli au siyo kweli inayoonyesha kama maandishi ni makubwa, rudisha kiwango cha utofauti cha WCAG kwa kutumia njia ifuatayo:

Kwanza, badilisha kila thamani ya RGB kuwa mwangaza wa jamaa:

- Gawanya kila chaneli `[R, G, B]` kwa 255 kupata thamani kati ya `0` na `1`
- Tumia fomula ya kurekebisha gamma kwa kila chaneli:
  - Ikiwa thamani ya chaneli ni ndogo au sawa na `0.04045`: `channel / 12.92`
  - Vinginevyo: `((channel + 0.055) / 1.055) ^ 2.4`
- Hesabu mwangaza: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Kisha, hesabu uwiano wa utofauti kwa kuongeza 0.05 kwa kila thamani ya mwangaza, kisha gawanya ile yenye mwanga zaidi kwa ile yenye giza zaidi. Ile yenye mwanga zaidi itakuwa hoja ya kwanza kila mara.

Rudisha kiwango kulingana na uwiano wa utofauti kwa kutumia jedwali lifuatalo:

| Kiwango | Maandishi ya Kawaida | Maandishi Makubwa |
|--------|--------------------|------------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | chini ya 4.5 | chini ya 3.0 |

# --hints--

`getContrastRating([255, 255, 255], [0, 0, 0], false)` inapaswa kurudisha `"AAA"`.

```js
assert.equal(getContrastRating([255, 255, 255], [0, 0, 0], false), "AAA");
```

`getContrastRating([215, 188, 188], [55, 55, 55], false)` inapaswa kurudisha `"AA"`.

```js
assert.equal(getContrastRating([215, 188, 188], [55, 55, 55], false), "AA");
```

`getContrastRating([143, 144, 210], [46, 47, 61], false)` inapaswa kurudisha `"Fail"`.

```js
assert.equal(getContrastRating([143, 144, 210], [46, 47, 61], false), "Fail");
```

`getContrastRating([167, 167, 210], [53, 10, 53], true)` inapaswa kurudisha `"AAA"`.

```js
assert.equal(getContrastRating([167, 167, 210], [53, 10, 53], true), "AAA");
```

`getContrastRating([135, 147, 155], [60, 70, 90], true)` inapaswa kurudisha `"AA"`.

```js
assert.equal(getContrastRating([135, 147, 155], [60, 70, 90], true), "AA");
```

`getContrastRating([125, 210, 195], [105, 130, 90], true)` inapaswa kurudisha `"Fail"`.

```js
assert.equal(getContrastRating([125, 210, 195], [105, 130, 90], true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {

  return rgb1;
}
```

# --solutions--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {
  function toLuminance([r, g, b]) {
    return [r, g, b].map(channel => {
      channel = channel / 255;
      return channel <= 0.04045
        ? channel / 12.92
        : ((channel + 0.055) / 1.055) ** 2.4;
    }).reduce((sum, c, i) => sum + c * [0.2126, 0.7152, 0.0722][i], 0);
  }

  const l1 = toLuminance(rgb1);
  const l2 = toLuminance(rgb2);
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
