---
id: 6a22d77ddf034bc4e35b1d57
title: "Zoezi la 350: Umbali wa Herufi"
challengeType: 28
dashedName: challenge-350
---

# --description--

Ukitoa mfuatano wa herufi miwili wenye urefu sawa, rudisha jumla ya umbali mfupi zaidi kati ya kila jozi ya herufi.

- Ingizo litakuwa na herufi ndogo tu
- Alfabeti hutendewa kama mduara, kwa hivyo umbali kati ya `a` na `z` ni 1.

# --hints--

`letterDistance("abc", "bcd")` inapaswa kurudisha `3`.

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")` inapaswa kurudisha `9`.

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")` inapaswa kurudisha `10`.

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")` inapaswa kurudisha `43`.

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")` inapaswa kurudisha `47`.

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")` inapaswa kurudisha `55`.

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
