---
id: 6a22d77ddf034bc4e35b1d59
title: "Changamoto 352: Kiwango cha uwiano wa utofauti 1"
challengeType: 28
dashedName: challenge-352
---

# --description--

Ukizingatia uwiano wa utofauti na thamani ya kweli au siyo kweli inayoonyesha kama maandishi ni makubwa, rudisha kiwango cha WCAG ukitumia jedwali lifuatalo:

| Kiwango | Maandishi ya kawaida | Maandishi makubwa |
|--------|---------------------|-------------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | chini ya 4.5 | chini ya 3.0 |

# --hints--

`getContrastRating("7.5", false)` inapaswa kurudisha `"AAA"`.

```js
assert.equal(getContrastRating("7.5", false), "AAA");
```

`getContrastRating("4.8", false)` inapaswa kurudisha `"AA"`.

```js
assert.equal(getContrastRating("4.8", false), "AA");
```

`getContrastRating("4.2", false)` inapaswa kurudisha `"Fail"`.

```js
assert.equal(getContrastRating("4.2", false), "Fail");
```

`getContrastRating("4.5", true)` inapaswa kurudisha `"AAA"`.

```js
assert.equal(getContrastRating("4.5", true), "AAA");
```

`getContrastRating("3.0", true)` inapaswa kurudisha `"AA"`.

```js
assert.equal(getContrastRating("3.0", true), "AA");
```

`getContrastRating("2.7", false)` inapaswa kurudisha `"Fail"`.

```js
assert.equal(getContrastRating("2.7", false), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(ratio, isLargeText) {

  return ratio;
}
```

# --solutions--

```js
function getContrastRating(ratio, isLargeText) {
  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
