---
id: 6a26df3e2988bcdded204893
title: "Changamoto 357: Mnyororo wa Chakula"
challengeType: 28
dashedName: challenge-357
---

# --description--

Kutolewa safu ya data ya jozi za `[predator, prey]`, rudisha mnyororo wa chakula kutoka mnyama mwangizi wa juu hadi chini kabisa.

- Mnyama mwangizi wa juu ni mnyama ambaye hakuwahi kuwa mtego wa mnyama mwingine.
- Rudisha mnyororo huo kama safu ya data ya mfuatano wa herufi.

# --hints--

`getFoodChain([["cat", "mouse"]])` inapaswa kurudisha `["cat", "mouse"]`.

```js
assert.deepEqual(getFoodChain([["cat", "mouse"]]), ["cat", "mouse"]);
```

`getFoodChain([["wolf", "deer"], ["deer", "grass"]])` inapaswa kurudisha `["wolf", "deer", "grass"]`.

```js
assert.deepEqual(getFoodChain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"]);
```

`getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` inapaswa kurudisha `["hawk", "snake", "frog", "fly"]`.

```js
assert.deepEqual(getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"]);
```

`getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` inapaswa kurudisha `["eagle", "fox", "rabbit", "grass"]`.

```js
assert.deepEqual(getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"]);
```

`getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` inapaswa kurudisha `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]`.

```js
assert.deepEqual(getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"]);
```

# --seed--

## --seed-contents--

```js
function getFoodChain(pairs) {

  return pairs;
}
```

# --solutions--

```js
function getFoodChain(pairs) {
  const prey = new Set(pairs.map(([, p]) => p));
  const map = Object.fromEntries(pairs);
  const apex = pairs.map(([p]) => p).find(p => !prey.has(p));
  const chain = [apex];
  while (map[chain.at(-1)]) chain.push(map[chain.at(-1)]);
  return chain;
}
```
