---
id: 6a26df3e2988bcdded204896
title: "Zoezi la 360: Muda Unaosemwa"
challengeType: 28
dashedName: challenge-360
---

# --description--

Ukichukulia idadi ya sekunde, rudisha muda huo kwa Kiingereza kinachosemwa.

- Gawanya muda kuwa saa, dakika, na sekunde.
- Ruka thamani yoyote ya sifuri.
- Tumia umoja au wingi ipasavyo (`"1 hour"`, `"2 hours"`).
- Ikiwa ipo, ungana vitengo viwili vya mwisho kwa `"and"`, na vitengo vya pili na vya tatu kutoka mwisho kwa koma (`"1 hour, 2 minutes and 3 seconds"`).

# --hints--

`getSpokenDuration(3723)` inapaswa kurudisha `"1 hour, 2 minutes and 3 seconds"`.

```js
assert.equal(getSpokenDuration(3723), "1 hour, 2 minutes and 3 seconds");
```

`getSpokenDuration(7295)` inapaswa kurudisha `"2 hours, 1 minute and 35 seconds"`.

```js
assert.equal(getSpokenDuration(7295), "2 hours, 1 minute and 35 seconds");
```

`getSpokenDuration(8521)` inapaswa kurudisha `"2 hours, 22 minutes and 1 second"`.

```js
assert.equal(getSpokenDuration(8521), "2 hours, 22 minutes and 1 second");
```

`getSpokenDuration(435)` inapaswa kurudisha `"7 minutes and 15 seconds"`.

```js
assert.equal(getSpokenDuration(435), "7 minutes and 15 seconds");
```

`getSpokenDuration(14455)` inapaswa kurudisha `"4 hours and 55 seconds"`.

```js
assert.equal(getSpokenDuration(14455), "4 hours and 55 seconds");
```

`getSpokenDuration(72000)` inapaswa kurudisha `"20 hours"`.

```js
assert.equal(getSpokenDuration(72000), "20 hours");
```

`getSpokenDuration(1)` inapaswa kurudisha `"1 second"`.

```js
assert.equal(getSpokenDuration(1), "1 second");
```

# --seed--

## --seed-contents--

```js
function getSpokenDuration(seconds) {

  return seconds;
}
```

# --solutions--

```js
function getSpokenDuration(seconds) {
  const hours = Math.floor(seconds / 3600);
  const minutes = Math.floor((seconds % 3600) / 60);
  const secs = seconds % 60;

  const parts = [];
  if (hours) parts.push(`${hours} ${hours === 1 ? "hour" : "hours"}`);
  if (minutes) parts.push(`${minutes} ${minutes === 1 ? "minute" : "minutes"}`);
  if (secs) parts.push(`${secs} ${secs === 1 ? "second" : "seconds"}`);

  if (parts.length === 1) return parts[0];
  if (parts.length === 2) return `${parts[0]} and ${parts[1]}`;
  return `${parts[0]}, ${parts[1]} and ${parts[2]}`;
}
```
