---
id: 6a26df3e2988bcdded204895
title: "Zoezi la 359: Kihesabu cha Kipimo cha Golfi"
challengeType: 28
dashedName: challenge-359
---

# --description--

Ukizingatia safu ya data ya alama za golfi na safu ya data inayolingana ya thamani za par za kozi, rudisha kielezo cha kipimo cha mpigaji golfi kwa kutumia njia ifuatayo:

- Hesabu tofauti kwa kila raundi kwa kutoa par kutoka kwa alama, kisha rudisha wastani wa tofauti zote uliozungushwa hadi sehemu moja ya desimali.

# --hints--

`calculateHandicap([72, 72, 72], [72, 72, 72])` inapaswa kurudisha `0`.

```js
assert.equal(calculateHandicap([72, 72, 72], [72, 72, 72]), 0);
```

`calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72])` inapaswa kurudisha `6`.

```js
assert.equal(calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72]), 6);
```

`calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36])` inapaswa kurudisha `8.3`.

```js
assert.equal(calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3);
```

`calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])` inapaswa kurudisha `8.8`.

```js
assert.equal(calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8);
```

`calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])` inapaswa kurudisha `11.7`.

```js
assert.equal(calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7);
```

# --seed--

## --seed-contents--

```js
function calculateHandicap(scores, pars) {

  return scores;
}
```

# --solutions--

```js
function calculateHandicap(scores, pars) {
  const differentials = scores.map((score, i) => score - pars[i]);
  const average = differentials.reduce((sum, d) => sum + d, 0) / differentials.length;
  return Math.round(average * 10) / 10;
}
```
