---
id: 6a26df95efa55a2524399743
title: "Changamoto 362: Kagua Nonogramu"
challengeType: 28
dashedName: challenge-362
---

# --description--

Ukizingatia safu ya nambari za vidokezo na safu ya seli, amua kama seli hizo zinakidhi kidokezo cha nonogramu.

- Kidokezo ni safu ya nambari zinazowakilisha urefu wa seli zilizojaa mfululizo, kwa mpangilio. Kwa mfano, kidokezo cha `[3, 2]` kinamaanisha kuwa kunapaswa kuwa na seli 3 zilizojaa mfululizo zikifuatiwa na seli 2 zilizojaa mfululizo, zikitenganishwa na angalau seli moja tupu.
- Safu ni safu ya 1 (zilizojaa) na 0 (tupu).

# --hints--

`isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1])` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1]));
```

`isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1])` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1]));
```

`isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]));
```

`isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]));
```

# --seed--

## --seed-contents--

```js
function isValidNonogram(clue, cells) {

  return clue;
}
```

# --solutions--

```js
function isValidNonogram(clue, cells) {
  const runs = [];
  let count = 0;

  for (let i = 0; i <= cells.length; i++) {
    if (cells[i] === 1) {
      count++;
    } else if (count > 0) {
      runs.push(count);
      count = 0;
    }
  }

  if (runs.length !== clue.length) return false;
  return runs.every((run, i) => run === clue[i]);
}
```
