---
id: 69f0a37c3bf194a7c9c62492
title: Jenga Daftari la Mkopo wa Vifaa
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

Wewe ni msanidi programu aliyepewa jukumu la kuunda programu ya daftari kusimamia vifaa vya vifaa vya IT vinavyotolewa na Huduma ya Seva ya kampuni yako. Daftari ni kitu ambacho funguo zake ni lebo za mali na thamani zake ni vitu:

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuunda mfano wa daftari kama kitu ambapo kila kifaa kinafunguliwa kwa lebo yake ya mali na kina mali za `type`, `status`, `borrower` (ikiwa na `name` na `email`), na `dueDate`.

1. Unapaswa kutekeleza kitendakazi cha `checkoutDevice(ledger, assetTag, borrower)` kinachochukua kitu cha daftari na kukifanya nakala bila kubadilisha asili. Katika nakala hiyo, kinapaswa kuweka `name` na `email` za kitu cha `borrower`, kusasisha `status` kuwa `"CheckedOut"`, na kurudisha kitu cha aina `{ ledger: updatedLedger, message: confirmationString }`, ambapo `message` inajumuisha lebo ya mali na `name` ya kitu cha `borrower`.

   - Ikiwa `assetTag` haipo katika daftari, kitendakazi kinapaswa kurudisha daftari bila kubadilika pamoja na `message` inayojumuisha lebo ya mali na kueleza kuwa haikupatikana.
   - Ikiwa kifaa tayari kiko katika hali ya `"CheckedOut"`, kitendakazi kinapaswa kurudisha daftari bila kubadilika na kuacha mkopaji wa sasa mahali pake, pamoja na `message` inayojumuisha lebo ya mali na kueleza kuwa kifaa tayari kimechukuliwa.

1. Unapaswa kutekeleza kitendakazi cha `checkinDevice(ledger, assetTag)` kinachochukua kitu cha daftari na kukifanya nakala bila kubadilisha asili. Katika nakala hiyo, kinapaswa kufuta `name` na `email` za kitu cha `borrower`, kuweka upya `dueDate` kuwa `""`, kusasisha `status` kuwa `"CheckedIn"`, na kurudisha kitu cha aina `{ ledger: updatedLedger, message: confirmationString }`, ambapo `message` inajumuisha lebo ya mali.

   - Ikiwa `assetTag` haipo katika daftari, kitendakazi kinapaswa kurudisha daftari bila kubadilika pamoja na `message` inayojumuisha lebo ya mali na kueleza kuwa haikupatikana.

1. Unapaswa kutekeleza kitendakazi cha `listOverdueDevices(ledger, today)` kinachochukua kitu cha daftari na kurudisha safu ya vitu vya kifaa ambavyo vimechelewa, vikiwa vimepangwa kwa `dueDate` kwa mpangilio wa kuongezeka. Kifaa kimechelewa wakati `status` yake ni `"CheckedOut"` na `dueDate` yake iko kabla kabisa ya `today`. Zote `today` na `dueDate` ni mfuatano wa herufi katika mpangilio wa mwezi/siku/mwaka, na mwezi na siku vinaweza kuwa na au bila nambari za kuanzia sifuri (kwa mfano, `"9/5/2025"` na `"09/05/2025"` zote ni halali).

1. Hupaswi kutumia kitu cha JavaScript `Date` mahali popote katika msimbo wako.

1. Unapaswa kutekeleza kitendakazi cha `serializeLedger(ledger)` kinachobadilisha kitu cha daftari kuwa mfuatano wa herufi wa JSON, na kitendakazi cha `loadLedger(json)` kinachobadilisha mfuatano wa herufi wa JSON kurudi kuwa kitu cha JavaScript.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `checkoutDevice`.

```js
assert.isFunction(checkoutDevice);
```

Unapaswa kuunda kitendakazi kinachoitwa `checkinDevice`.

```js
assert.isFunction(checkinDevice);
```

Unapaswa kuunda kitendakazi kinachoitwa `listOverdueDevices`.

```js
assert.isFunction(listOverdueDevices);
```

Unapaswa kuunda kitendakazi kinachoitwa `serializeLedger`.

```js
assert.isFunction(serializeLedger);
```

Unapaswa kuunda kitendakazi kinachoitwa `loadLedger`.

```js
assert.isFunction(loadLedger);
```

Kitendakazi chako cha `checkoutDevice` hakipaswi kubadilisha daftari la asili.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

Kitendakazi chako cha `checkoutDevice` kinapaswa kuweka `name` na `email` za mkopaji, na kuweka `status` kuwa `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

Kitendakazi chako cha `checkoutDevice` kinapaswa kurudisha ujumbe wa uthibitisho unaojumuisha lebo ya mali na `name` ya mkopaji.

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

Kitendakazi chako cha `checkoutDevice` kinapaswa kuhifadhi vifaa visivyohusiana na sehemu zisizobadilika katika daftari lililorudishwa.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

Kitendakazi chako cha `checkoutDevice` kinapaswa kurudisha daftari bila kubadilika na ujumbe unaojumuisha lebo ya mali wakati lebo ya mali haipatikani.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

Kitendakazi chako cha `checkoutDevice` kinapaswa kurudisha daftari bila kubadilika, kuacha mkopaji wa sasa mahali pake, na kurudisha ujumbe unaojumuisha lebo ya mali wakati kifaa tayari kimechukuliwa.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

Kitendakazi chako cha `checkinDevice` hakipaswi kubadilisha daftari la asili.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

Kitendakazi chako cha `checkinDevice` kinapaswa kufuta `name` na `email` za mkopaji, kuweka upya `dueDate` kuwa `""`, na kuweka `status` kuwa `"CheckedIn"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

Kitendakazi chako cha `checkinDevice` kinapaswa kurudisha ujumbe wa uthibitisho unaojumuisha lebo ya mali.

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

Kitendakazi chako cha `checkinDevice` kinapaswa kuhifadhi vifaa visivyohusiana na sehemu zisizobadilika katika daftari lililorudishwa.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

Kitendakazi chako cha `checkinDevice` kinapaswa kurudisha daftari bila kubadilika na ujumbe unaojumuisha lebo ya mali wakati lebo ya mali haipatikani.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

Kitendakazi chako cha `listOverdueDevices` kinapaswa kurudisha safu ya data.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

Kitendakazi chako cha `listOverdueDevices` kinapaswa kurudisha vifaa vilivyochelewa tu vyenye `status` ya `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

Kitendakazi chako cha `listOverdueDevices` kinapaswa kurudisha vifaa vilivyochelewa vikiwa vimepangwa kwa `dueDate` kwa mpangilio wa kuongezeka.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

Kitendakazi chako cha `listOverdueDevices` kinapaswa kupanga kwa usahihi hata kama mwezi au siku hazijajazwa kwa sifuri mwanzoni.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

Hupaswi kutumia kitu cha JavaScript `Date` mahali popote katika msimbo wako.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

Kitendakazi chako cha `serializeLedger` kinapaswa kurudisha mfuatano wa herufi wa JSON unaowakilisha daftari lililotumwa kwake.

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

Kitendakazi chako cha `loadLedger` kinapaswa kurudisha kitu cha JavaScript ambacho mfuatano wa herufi wa JSON unawakilisha.

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
