---
id: 5688e62ea601b2482ff8422b
title: Construir uma Consulta de Perfil
challengeType:  26
dashedName: lab-profile-lookup
---

# --description--

Neste laboratório, você construirá uma busca de perfil que consulta informações sobre pessoas em uma lista de contatos.

Para este exemplo, imagine que há um contato chamado Akira Laine, o `lookUpProfile("Akira", "lastName")` deve retornar `Laine`. 

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `lookUpProfile` que recebe um nome e uma propriedade como argumentos. 
2. Você deve recuperar as informações de contato do array `contacts` fornecido. 
3. Se a função receber um nome de contato e a propriedade existir no contato relacionado, então o valor da propriedade deve ser retornado.
4. Se o nome passado para a função não corresponder a nenhum contato no array de contatos, então a função deve retornar `"Nenhum contato encontrado"`.
5. Se a propriedade não existir em um contato encontrado, então a função deve retornar `"Propriedade inexistente"`. 

# --before-each--

```js
contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];
```

# --hints--

O array `contacts` ainda deve estar presente no escopo global.Redefina a lição para recuperá-la caso você a tenha excluído.

```js
assert.isArray(contacts);
```

Você deve ter uma função chamada `lookUpProfile`.

```js
assert.isFunction(lookUpProfile);
```

`lookUpProfile("Kristian", "lastName")` deve retornar uma string.

```js
assert.isString(lookUpProfile("Kristian", "lastName"));
```

`lookUpProfile("Kristian", "lastName")` deve retornar a string `Vos`

```js
assert.strictEqual(lookUpProfile('Kristian', 'lastName'), 'Vos');
```

`lookUpProfile("Sherlock", "likes")` deve retornar `["Intriguing Cases", "Violin"]`

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")` deve retornar um array

```js
assert.isArray(lookUpProfile('Harry', 'likes'));
```

`lookUpProfile("Bob", "number")` deve retornar a string `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'number'),'No such contact');
```

`lookUpProfile("Bob", "potato")` deve retornar a string `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'potato'), 'No such contact');
```

`lookUpProfile("Akira", "address")` deve retornar a string `No such property`

```js
assert.strictEqual(lookUpProfile('Akira', 'address'), 'No such property');
```

# --seed--

## --seed-contents--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


```

# --solutions--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
