---
id: 56533eb9ac21ba0edf2244e2
title: Msimbo wa Caesar
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16003
dashedName: caesars-cipher
---

# --description--

Moja ya <dfn>mismbo</dfn> rahisi na inayojulikana sana ni <dfn>msimbo wa Caesar</dfn>, unaojulikana pia kama <dfn>msimbo wa mabadiliko</dfn>. Katika msimbo wa mabadiliko maana za herufi hubadilishwa kwa kiwango fulani kilichowekwa.

Matumizi ya kawaida ya kisasa ni msimbo wa <a href="https://www.freecodecamp.org/news/how-to-code-the-caesar-cipher-an-introduction-to-basic-encryption-3bf77b4e19f7/" target="_blank" rel="noopener noreferrer nofollow">ROT13</a>, ambapo thamani za herufi hubadilishwa kwa nafasi 13. Hivyo `A ↔ N`, `B ↔ O` na kadhalika.

Andika kitendakazi kinachopokea mfuatano wa herufi uliosimbwa kwa <a href="https://www.freecodecamp.org/news/how-to-code-the-caesar-cipher-an-introduction-to-basic-encryption-3bf77b4e19f7/" target="_blank" rel="noopener noreferrer nofollow">ROT13</a> kama ingizo na kurudisha mfuatano wa herufi uliotafsiriwa.

Herufi zote zitakuwa kubwa. Usibadilishe herufi zisizo za alfabeti (yaani nafasi, alama za uandishi), lakini zipitishe kama zilivyo.

# --hints--

`rot13("SERR PBQR PNZC")` inapaswa kutafsiriwa kuwa mfuatano wa herufi `FREE CODE CAMP`

```js
assert(rot13('SERR PBQR PNZC') === 'FREE CODE CAMP');
```

`rot13("SERR CVMMN!")` inapaswa kutafsiriwa kuwa mfuatano wa herufi `FREE PIZZA!`

```js
assert(rot13('SERR CVMMN!') === 'FREE PIZZA!');
```

`rot13("SERR YBIR?")` inapaswa kutafsiriwa kuwa mfuatano wa herufi `FREE LOVE?`

```js
assert(rot13('SERR YBIR?') === 'FREE LOVE?');
```

`rot13("GUR DHVPX OEBJA SBK WHZCF BIRE GUR YNML QBT.")` inapaswa kutafsiriwa kuwa mfuatano wa herufi `THE QUICK BROWN FOX JUMPS OVER THE LAZY DOG.`

```js
assert(
  rot13('GUR DHVPX OEBJA SBK WHZCF BIRE GUR YNML QBT.') ===
    'THE QUICK BROWN FOX JUMPS OVER THE LAZY DOG.'
);
```

# --seed--

## --seed-contents--

```js
function rot13(str) {
  return str;
}

rot13("SERR PBQR PNZC");
```

# --solutions--

```js
var lookup = {
  'A': 'N','B': 'O','C': 'P','D': 'Q',
  'E': 'R','F': 'S','G': 'T','H': 'U',
  'I': 'V','J': 'W','K': 'X','L': 'Y',
  'M': 'Z','N': 'A','O': 'B','P': 'C',
  'Q': 'D','R': 'E','S': 'F','T': 'G',
  'U': 'H','V': 'I','W': 'J','X': 'K',
  'Y': 'L','Z': 'M'
};

function rot13(encodedStr) {
  var codeArr = encodedStr.split("");  // String to Array
  var decodedArr = []; // Your Result goes here
  // Only change code below this line

  decodedArr = codeArr.map(function(letter) {
    if(lookup.hasOwnProperty(letter)) {
      letter = lookup[letter];
    }
    return letter;
  });

  // Only change code above this line
  return decodedArr.join(""); // Array to String
}
```
