---
id: a7f4d8f2483413a6ce226cac
title: Mbadilishaji wa nambari za Kirumi
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16044
dashedName: roman-numeral-converter
---

# --description--

Badilisha nambari iliyotolewa kuwa nambari ya Kirumi.

| Nambari za Kirumi | Nambari za Kiarabu |
|-------------------|--------------------|
| M                 | 1000               |
| CM                | 900                |
| D                 | 500                |
| CD                | 400                |
| C                 | 100                |
| XC                | 90                 |
| L                 | 50                 |
| XL                | 40                 |
| X                 | 10                 |
| IX                | 9                  |
| V                 | 5                  |
| IV                | 4                  |
| I                 | 1                  |

Majibu yote ya nambari za Kirumi yanapaswa kutolewa kwa herufi kubwa.

# --hints--

`convertToRoman(2)` inapaswa kurudisha mfuatano wa herufi `II`.

```js
assert.deepEqual(convertToRoman(2), 'II');
```

`convertToRoman(3)` inapaswa kurudisha mfuatano wa herufi `III`.

```js
assert.deepEqual(convertToRoman(3), 'III');
```

`convertToRoman(4)` inapaswa kurudisha mfuatano wa herufi `IV`.

```js
assert.deepEqual(convertToRoman(4), 'IV');
```

`convertToRoman(5)` inapaswa kurudisha mfuatano wa herufi `V`.

```js
assert.deepEqual(convertToRoman(5), 'V');
```

`convertToRoman(9)` inapaswa kurudisha mfuatano wa herufi `IX`.

```js
assert.deepEqual(convertToRoman(9), 'IX');
```

`convertToRoman(12)` inapaswa kurudisha mfuatano wa herufi `XII`.

```js
assert.deepEqual(convertToRoman(12), 'XII');
```

`convertToRoman(16)` inapaswa kurudisha mfuatano wa herufi `XVI`.

```js
assert.deepEqual(convertToRoman(16), 'XVI');
```

`convertToRoman(29)` inapaswa kurudisha mfuatano wa herufi `XXIX`.

```js
assert.deepEqual(convertToRoman(29), 'XXIX');
```

`convertToRoman(44)` inapaswa kurudisha mfuatano wa herufi `XLIV`.

```js
assert.deepEqual(convertToRoman(44), 'XLIV');
```

`convertToRoman(45)` inapaswa kurudisha mfuatano wa herufi `XLV`.

```js
assert.deepEqual(convertToRoman(45), 'XLV');
```

`convertToRoman(68)` inapaswa kurudisha mfuatano wa herufi `LXVIII`

```js
assert.deepEqual(convertToRoman(68), 'LXVIII');
```

`convertToRoman(83)` inapaswa kurudisha mfuatano wa herufi `LXXXIII`

```js
assert.deepEqual(convertToRoman(83), 'LXXXIII');
```

`convertToRoman(97)` inapaswa kurudisha mfuatano wa herufi `XCVII`

```js
assert.deepEqual(convertToRoman(97), 'XCVII');
```

`convertToRoman(99)` inapaswa kurudisha mfuatano wa herufi `XCIX`

```js
assert.deepEqual(convertToRoman(99), 'XCIX');
```

`convertToRoman(400)` inapaswa kurudisha mfuatano wa herufi `CD`

```js
assert.deepEqual(convertToRoman(400), 'CD');
```

`convertToRoman(500)` inapaswa kurudisha mfuatano wa herufi `D`

```js
assert.deepEqual(convertToRoman(500), 'D');
```

`convertToRoman(501)` inapaswa kurudisha mfuatano wa herufi `DI`

```js
assert.deepEqual(convertToRoman(501), 'DI');
```

`convertToRoman(649)` inapaswa kurudisha mfuatano wa herufi `DCXLIX`

```js
assert.deepEqual(convertToRoman(649), 'DCXLIX');
```

`convertToRoman(798)` inapaswa kurudisha mfuatano wa herufi `DCCXCVIII`

```js
assert.deepEqual(convertToRoman(798), 'DCCXCVIII');
```

`convertToRoman(891)` inapaswa kurudisha mfuatano wa herufi `DCCCXCI`

```js
assert.deepEqual(convertToRoman(891), 'DCCCXCI');
```

`convertToRoman(1000)` inapaswa kurudisha mfuatano wa herufi `M`

```js
assert.deepEqual(convertToRoman(1000), 'M');
```

`convertToRoman(1004)` inapaswa kurudisha mfuatano wa herufi `MIV`

```js
assert.deepEqual(convertToRoman(1004), 'MIV');
```

`convertToRoman(1006)` inapaswa kurudisha mfuatano wa herufi `MVI`

```js
assert.deepEqual(convertToRoman(1006), 'MVI');
```

`convertToRoman(1023)` inapaswa kurudisha mfuatano wa herufi `MXXIII`

```js
assert.deepEqual(convertToRoman(1023), 'MXXIII');
```

`convertToRoman(2014)` inapaswa kurudisha mfuatano wa herufi `MMXIV`

```js
assert.deepEqual(convertToRoman(2014), 'MMXIV');
```

`convertToRoman(3999)` inapaswa kurudisha mfuatano wa herufi `MMMCMXCIX`

```js
assert.deepEqual(convertToRoman(3999), 'MMMCMXCIX');
```

# --seed--

## --seed-contents--

```js
function convertToRoman(num) {
 return num;
}

convertToRoman(36);
```

# --solutions--

```js
function convertToRoman(num) {
  var ref = [['M', 1000], ['CM', 900], ['D', 500], ['CD', 400], ['C', 100], ['XC', 90], ['L', 50], ['XL', 40], ['X', 10], ['IX', 9], ['V', 5], ['IV', 4], ['I', 1]];
  var res = [];
  ref.forEach(function(p) {
    while (num >= p[1]) {
      res.push(p[0]);
      num -= p[1];
    }
  });
  return res.join('');
}
```
