---
id: aaa48de84e1ecc7c742e1124
title: Kagua palindrome
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16004
dashedName: palindrome-checker
---

# --description--

Rudisha `true` ikiwa mfuatano wa herufi uliotolewa ni palindrome. Vinginevyo, rudisha `false`.

<dfn>Palindrome</dfn> ni neno au sentensi inayosomwa sawa mbele na nyuma, bila kuzingatia alama za uandishi, herufi kubwa au ndogo, na nafasi.

**Kumbuka:** Utahitaji kuondoa **herufi zote zisizo za herufi na nambari** (alama za uandishi, nafasi na alama) na kubadilisha kila kitu kuwa herufi kubwa au ndogo ili kuangalia palindromu.

Tutapitisha mfuatano wa herufi wenye miundo tofauti, kama `racecar`, `RaceCar`, na `race CAR` miongoni mwa mengine.

Pia tutapitisha mfuatano wa herufi wenye alama maalum, kama `2A3*3a2`, `2A3 3a2`, na `2_A3*3#A2`.

# --hints--

`palindrome("eye")` inapaswa kurudisha thamani ya kweli au siyo kweli.

```js
assert(typeof palindrome('eye') === 'boolean');
```

`palindrome("eye")` inapaswa kurudisha `true`.

```js
assert(palindrome('eye') === true);
```

`palindrome("_eye")` inapaswa kurudisha `true`.

```js
assert(palindrome('_eye') === true);
```

`palindrome("race car")` inapaswa kurudisha `true`.

```js
assert(palindrome('race car') === true);
```

`palindrome("not a palindrome")` inapaswa kurudisha `false`.

```js
assert(palindrome('not a palindrome') === false);
```

`palindrome("A man, a plan, a canal. Panama")` inapaswa kurudisha `true`.

```js
assert(palindrome('A man, a plan, a canal. Panama') === true);
```

`palindrome("never odd or even")` inapaswa kurudisha `true`.

```js
assert(palindrome('never odd or even') === true);
```

`palindrome("nope")` inapaswa kurudisha `false`.

```js
assert(palindrome('nope') === false);
```

`palindrome("almostomla")` inapaswa kurudisha `false`.

```js
assert(palindrome('almostomla') === false);
```

`palindrome("My age is 0, 0 si ega ym.")` inapaswa kurudisha `true`.

```js
assert(palindrome('My age is 0, 0 si ega ym.') === true);
```

`palindrome("1 eye for of 1 eye.")` inapaswa kurudisha `false`.

```js
assert(palindrome('1 eye for of 1 eye.') === false);
```

`palindrome("0_0 (: /-\ :) 0-0")` inapaswa kurudisha `true`.

```js
assert(palindrome('0_0 (: /- :) 0-0') === true);
```

`palindrome("five|\_/|four")` inapaswa kurudisha `false`.

```js
assert(palindrome('five|_/|four') === false);
```

# --seed--

## --seed-contents--

```js
function palindrome(str) {
  return true;
}

palindrome("eye");
```

# --solutions--

```js
function palindrome(str) {
  var string = str.toLowerCase().split(/[^A-Za-z0-9]/gi).join('');
  var aux = string.split('');
  if (aux.join('') === aux.reverse().join('')){
    return true;
  }

  return false;
}
```
