---
id: 694470b9a65b278f981ec62e
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Testen Sie abschließend Ihre `countup`-Funktion, indem Sie sie mit dem Argument `5` innerhalb einer `console.log()`-Anweisung aufrufen.

Die Ausgabe sollte `[1, 2, 3, 4, 5]` sein.

# --hints--

Sie sollten die Funktion `countup` mit dem Argument `5` innerhalb einer `console.log()`-Anweisung aufrufen.

```js
assert.deepStrictEqual(countup(5), [1, 2, 3, 4, 5]);
assert.match(code, /console\.log\s*\(\s*countup\s*\(\s*5\s*\)\s*\)/)
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
console.log(countup(5));
```
