---
id: 69447054012a22f38c944ff0
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Fügen Sie eine `if`-Anweisung in Ihre `countup`-Funktion ein.

Prüfen Sie, ob `number` kleiner als `1` ist, und geben Sie ein leeres Array zurück.

Dies dient als Basisfall für Ihre Rekursion.

# --hints--

Wenn `number` kleiner als `1` ist, sollte die Funktion ein leeres Array zurückgeben.

```js
assert.deepStrictEqual(countup(0), []);
```

Wenn `number` **nicht** kleiner als `1` ist, sollte die Funktion **nichts** zurückgeben.

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
