---
id: 69446f18fdcb11aa516f7c5c
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop erstellen Sie eine `countup`-Funktion, die mit Rekursion ein Array mit Zahlen von 1 bis zu einer vorgegebenen Zahl erzeugt.

Beginnen Sie damit, eine leere Funktion namens `countup` zu erstellen, die einen einzelnen Parameter annimmt.

# --hints--

Sie sollten eine Funktion namens `countup` haben.

```js
assert.isFunction(countup);
```

Ihre `countup`-Funktion sollte einen einzelnen Parameter akzeptieren.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
