---
id: 6948705160b66cf93a916e29
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Deklarieren Sie ein leeres Array mit dem Namen `countArray` innerhalb der `countup`-Funktion direkt über Ihrer `if`-Anweisung.

Dieses Array speichert bei jedem Funktionsaufruf den Wert von `number`.

# --hints--

Sie sollten ein leeres Array mit dem Namen `countArray` deklarieren.

```js
assert.match(countup.toString(), /(let|const|var)\s+countArray\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
