---
id: 694870bcf3c6874112b6f951
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Für alle Zahlen größer als und gleich `1` erstellen Sie einen `else`-Block und rufen darin die `countup`-Funktion mit `number - 1` auf, wobei Sie das Ergebnis `countArray` zuweisen.

Verwenden Sie dann `.push()`, um `number` zu `countArray` hinzuzufügen, und geben Sie schließlich `countArray` zurück.

So stellen Sie sicher, dass die Rekursion schließlich den Basisfall erreicht und beendet wird.

# --hints--

Sie sollten das Ergebnis des Aufrufs von `countup` mit `number - 1` `countArray` zuweisen.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

Sie sollten `number` in `countArray` einfügen.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

Sie sollten `countArray` zurückgeben.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
