---
id: 6948705160b66cf93a916e29
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Deklarieren Sie ein leeres Array mit dem Namen `countArray` innerhalb der `countup`-Funktion direkt über Ihrer `if`-Anweisung.

Dieses Array speichert bei jedem Funktionsaufruf den Wert von `number`.

# --hints--

Sie sollten ein leeres Array mit dem Namen `countArray` deklarieren.

```js
const explorer = await __helpers.Explorer(code);
const { countup } = explorer.functions;
assert.exists(countup);
const { countArray } = countup.variables;
assert.exists(countArray);
assert.isTrue(countArray.value.matches('[]'));
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
