---
id: 66c8ba41a77db20a93f9d7a1
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個貸款資格檢查器應用程式來複習 JavaScript 條件判斷。

該應用程式將根據使用者的年收入和信用分數，檢查其是否符合雙拼住宅、汽車和公寓貸款的資格。信用分數是一個數字，用來表述你管理借款能力的好壞。

要開始，請創建以下變數和值。

| Variable Name | Value   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

你應該創建一個 `minIncomeForDuplex` 變數並設定為 `60000`。

```js
assert.equal(minIncomeForDuplex, 60000);
```

你應該創建一個 `minCreditScoreForDuplex` 並設定為 `700`。

```js
assert.equal(minCreditScoreForDuplex, 700);
```

你應該創建一個 `minIncomeForCar` 變數並設定為 `30000`。

```js
assert.equal(minIncomeForCar, 30000);
```

你應該創建一個 `minCreditScoreForCar` 變數並設定為 `650`。

```js
assert.equal(minCreditScoreForCar, 650);
```

你應該創建一個 `minIncomeForCondo` 變數並設定為 `45000`。

```js
assert.equal(minIncomeForCondo, 45000);
```

你應該創建一個 `minCreditScoreForCondo` 變數並設定為 `680`。

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
