---
id: 66c8ba975ee7230e29f6c4b2
title: 步驟7
challengeType: 1
dashedName: step-7
---

# --description--

現在，是測試 `getLoanMessage` 函數的時候了。

使用下表創建 `4` 變量及其值：

| 變量名稱            | Value                        |
| --------------- | ---------------------------- |
| `duplexLoanMsg` | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`  | `getLoanMessage(65000, 690)` |
| `carLoanMsg`    | `getLoanMessage(45000, 660)` |
| `noLoanMsg`     | `getLoanMessage(25000, 550)` |

然後，將每個變量記錄到控制檯，查看信息。

至此，你的貸款資格檢查項目就完成了！

# --hints--

你應該有一個名爲 `duplexLoanMsg` 的變量。

```js
assert.isNotNull(duplexLoanMsg);
```

你的 `duplexLoanMsg` 變量應設置爲 `getLoanMessage(85000, 850)`。

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

你應將 `duplexLoanMsg` 變量記錄到控制檯。

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

你應該有一個名爲 `condoLoanMsg` 的變量。

```js
assert.isNotNull(condoLoanMsg);
```

你的 `condoLoanMsg` 變量應設置爲 `getLoanMessage(65000, 690)`。

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

你應將 `condoLoanMsg` 變量記錄到控制檯。

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

你應該有一個名爲 `carLoanMsg` 的變量。

```js
assert.isNotNull(carLoanMsg);
```

你的 `carLoanMsg` 變量應設置爲 `getLoanMessage(45000, 660)`。

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

你應將 `carLoanMsg` 變量記錄到控制檯。

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

你應該有一個名爲 `noLoanMsg` 的變量。

```js
assert.isNotNull(noLoanMsg);
```

你的 `noLoanMsg` 變量應設置爲 `getLoanMessage(25000, 550)`。

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

你應將 `noLoanMsg` 變量記錄到控制檯。

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
