---
id: 66c8ba975ee7230e29f6c4ac
title: 步驟2
challengeType: 1
dashedName: step-2
---

# --description--

當用戶符合貸款條件時，你需要在控制檯中向他們顯示一條信息。

爲此，你需要建立一個函數，在這個函數中，你需要進行一些檢查，以返回申請人符合哪些貸款條件。

創建一個空的 `getLoanMessage` 函數，參數爲 `annualIncome` 和 `creditScore` 。

# --hints--

你應創建一個 `getLoanMessage` 函數。

```js
assert.isFunction(getLoanMessage)
```

你的 `getLoanMessage` 函數應將 `annualIncome` 作爲 `creditScore` 作爲參數。

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` 應該是一個空函數。

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
