---
id: 66c8ba975ee7230e29f6c4af
title: 步驟4
challengeType: 1
dashedName: step-4
---

# --description--

如果申請人的年收入大於或等於 `minIncomeForCondo`，且如果他們的信用分數大於或等於 `minCreditScoreForCondo`，那麼他們就有資格申請公寓和汽車貸款。

在 `getLoanMessage` 函數中檢查是否爲真。 如果是，則返回字符串 `"You qualify for a condo and car loan."`

# --hints--

如果申請人符合公寓和汽車貸款條件，你的 `getLoanMessage` 函數應返回一個字符串。

```js
assert.isString(getLoanMessage(45000, 680));
```

你的 `getLoanMessage` 函數應返回字符串 `"You qualify for a condo and car loan."`。

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

如果申請人的年收入和信用評分不符合公寓貸款的要求，你的 `getLoanMessage` 函數應返回 `undefined` 。

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }

--fcc-editable-region--
}
```
