---
id: 66c8ba975ee7230e29f6c4ac
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

當使用者符合貸款資格時，你會想在主控台中向他們顯示一則訊息。

為此，你將在其中建置一個函式（程式），該函式（程式）內會有一些檢查，這些檢查會傳回申請人有資格申請的貸款種類。

創建一個空的 `getLoanMessage` 函式（程式），並帶有 `annualIncome` 和 `creditScore` 參數。

# --hints--

你應該創建一個 `getLoanMessage` 函式（程式）

```js
assert.isFunction(getLoanMessage)
```

你的 `getLoanMessage` 函式（程式）應該有 `annualIncome` 和 `creditScore` 作為參數。

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` 應該是一個空的函式（程式）。

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
