---
id: 66c8ba975ee7230e29f6c4ad
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

要根據 `annualIncome` 和 `creditScore` 檢查使用者符合哪種貸款資格，你必須在 `getLoanMessage` 函式（程式）內使用 `if/else if` 述語或三元運算子。然後你會在每個條件的區塊中傳回適當的訊息。

從雙拼貸款開始，檢查 `annualIncome` 是否大於或相等於 `minIncomeForDuplex` 且 `creditScore` 是否大於或相等於 `minCreditScoreForDuplex`。

如果該條件為真，申請人即有資格申請雙拼貸款及其他貸款。因此，在檢查內，傳回字串 `"You qualify for a duplex, condo, and car loan."`

# --hints--

你的 `getLoanMessage` 函式應該傳回一個字串。

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

你的 `getLoanMessage` 函式應該傳回字串 `"You qualify for a duplex, condo, and car loan."`。

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

如果申請人的年收入和信用評分不符合雙拼貸款的要求，你的 `getLoanMessage` 函式應該傳回 `undefined`。

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
