---
id: 66c8ba975ee7230e29f6c4b2
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

現在，是時候測試你的 `getLoanMessage` 函式（程式）了。

使用下表來創建 `4` 個變數及其值：

| Variable Name | Value   |
| ----------- | ------- |
| `duplexLoanMsg`       | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`   | `getLoanMessage(65000, 690)`   |
| `carLoanMsg`   | `getLoanMessage(45000, 660)`   |
| `noLoanMsg`   | `getLoanMessage(25000, 550)`   |

之後，將每個變數記錄到主控台以查看訊息。

這樣，你的貸款資格檢查器專案就完成了！

# --hints--

你應該有一個名為 `duplexLoanMsg` 的變數。

```js
assert.isNotNull(duplexLoanMsg);
```

你的 `duplexLoanMsg` 變數應該設定為 `getLoanMessage(85000, 850)` 的結果。

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

你應該將 `duplexLoanMsg` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

你應該有一個名為 `condoLoanMsg` 的變數。

```js
assert.isNotNull(condoLoanMsg);
```

你的 `condoLoanMsg` 變數應該設定為 `getLoanMessage(65000, 690)` 的結果。

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

你應該將 `condoLoanMsg` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

你應該有一個名為 `carLoanMsg` 的變數。

```js
assert.isNotNull(carLoanMsg);
```

你的 `carLoanMsg` 變數應該設定為 `getLoanMessage(45000, 660)` 的結果。

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

你應該將 `carLoanMsg` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

你應該有一個名為 `noLoanMsg` 的變數。

```js
assert.isNotNull(noLoanMsg);
```

你的 `noLoanMsg` 變數應該設定為 `getLoanMessage(25000, 550)` 的結果。

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

你應該將 `noLoanMsg` 變數記錄到主控台。

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
