---
id: 66c8ba975ee7230e29f6c4b1
title: 步驟6
challengeType: 1
dashedName: step-6
---

# --description--

如果申請人的年收入和信用評分低於 `minIncomeForCar` 和 `minCreditScoreForCar`，那麼他們就沒有資格獲得任何貸款。 因此，返回字符串 `"You don't qualify for any loans."`

# --hints--

如果申請人不符合任何貸款條件，你的 `getLoanMessage` 函數應返回一個字符串。

```js
assert.isString(getLoanMessage(25000, 550));
```

如果申請人不符合任何貸款要求，你的 `getLoanMessage` 函數應返回字符串 `"You don't qualify for any loans."`。

```js
assert.strictEqual(getLoanMessage(15000, 600), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(25000, 550), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(20000, 500), "You don't qualify for any loans.");
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
--fcc-editable-region--
  }

--fcc-editable-region--
}
```
