---
id: 66c8ba975ee7230e29f6c4af
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

如果申請人的年收入大於或相等於 `minIncomeForCondo`，且他們的信用分數大於或相等於 `minCreditScoreForCondo`，那麼他們符合申請公寓和汽車貸款的資格。

在 `getLoanMessage` 函式中檢查是否為真。如果是，傳回字串 `"You qualify for a condo and car loan."`

# --hints--

你的 `getLoanMessage` 函式應該在申請人符合公寓和汽車貸款資格時傳回字串。

```js
assert.isString(getLoanMessage(45000, 680));
```

你的 `getLoanMessage` 函式應該傳回字串 `"You qualify for a condo and car loan."`。

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

如果申請人的年收入和信用評分不符合公寓貸款的要求，你的 `getLoanMessage` 函式應該傳回 `undefined`。

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
