---
id: 66c8ba975ee7230e29f6c4b0
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

現在，你應該檢查申請人是否僅經過資格修飾、合格的汽車貸款。如果是，傳回字串 `"You qualify for a car loan."`。

# --hints--

如果申請人符合汽車貸款資格，你的 `getLoanMessage` 函式（程式）應該傳回一個字串。

```js
assert.isString(getLoanMessage(30000, 650));
```

你的 `getLoanMessage` 函式（程式）應該傳回字串 `"You qualify for a car loan."`。

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

如果申請人的年收入和信用評分不符合汽車貸款的要求，你的 `getLoanMessage` 函式（程式）應該傳回 `undefined`。

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 
  
--fcc-editable-region--
}
```
