---
id: 66fbcdfed6848e48947dbbc9
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите трекер рецептів, використовуючи JavaScript-об’єкти.

Почніть зі створення порожнього масиву з назвою `recipes`. Саме в цей масив ви пізніше додаватимете об’єкти рецептів.

# --hints--

Ви маєте створити масив з назвою `recipes`.

```js
assert.isArray(recipes);
```

Ваш масив `recipes` має бути порожнім.

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
