---
id: 66fbcf750a62784cf11f5632
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Створіть функцію `getTotalIngredients`, яка приймає один аргумент, що представляє масив з інгредієнтами, і повертає число інгредієнтів із масиву, переданого функції.

# --hints--

Ви повинні створити функцію `getTotalIngredients`.

```js
assert.isFunction(getTotalIngredients);
```

Ваша функція `getTotalIngredients` має мати один параметр.

```js
assert.lengthOf(getTotalIngredients, 1);
```

Ваша функція `getTotalIngredients` має повертати число.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

`getTotalIngredients(recipe1.ingredients)` має повертати `4`.

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)` має повертати `5`.

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)` має повертати `3`.

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
