---
id: 66fbcf750a62784cf11f5630
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер ви маєте додати три об’єкти у масив `recipes`. Для цього можна використати метод `push()`.

Використайте метод `push()`, щоб додати всі об’єкти рецептів у масив `recipes`. Переконайтеся, що додали `recipe1`, `recipe2` та `recipe3` у такому порядку.

Також видаліть змінні `recipe1Name`, `recipe2CookingTime` і `recipe3Ingredients`, а також інструкції `console.log`, які виводять ці змінні.

# --hints--

Ви маєте додати всі об’єкти рецептів у масив `recipes`. Переконайтеся, що додали їх у порядку оголошення.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

Ви маєте видалити змінну `recipe1Name`.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

Ви маєте видалити інструкцію `console.log(recipe1Name);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

Ви маєте видалити змінну `recipe2CookingTime`.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

Ви маєте видалити інструкцію `console.log(recipe2CookingTime);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

Ви маєте видалити змінну `recipe3Ingredients`.

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

Ви маєте видалити інструкцію `console.log(recipe3Ingredients);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
