---
id: 66fbcf750a62784cf11f562e
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Створіть об’єкт `recipe2` з такими властивостями та значеннями:

| Ключ | Значення   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

Ви повинні створити об’єкт з ім’ям `recipe2`.

```js
assert.isObject(recipe2);
```

Ваш об’єкт `recipe2` має мати властивість `name`.

```js
assert.property(recipe2, "name");
```

Властивість `name` має бути встановлена в `Chicken Curry`.

```js
assert.equal(recipe2.name, "Chicken Curry");
```

Ваш об’єкт `recipe2` має мати властивість `ingredients`.

```js
assert.property(recipe2, "ingredients");
```

Властивість `ingredients` має містити масив.

```js
assert.isArray(recipe2.ingredients);
```

Масив у властивості `ingredients` має містити `chicken breast`, `coconut milk`, `curry powder`, `onion` і `garlic`.

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

Ваш об’єкт `recipe2` має мати властивість `cookingTime`.

```js
assert.property(recipe2, "cookingTime");
```

Значення властивості `cookingTime` має бути числом.

```js
assert.isNumber(recipe2.cookingTime);
```

Властивість `cookingTime` має бути встановлена в `42`.

```js
assert.equal(recipe2.cookingTime, 42);
```

Ваш об’єкт `recipe2` має мати властивість `totalIngredients`.

```js
assert.property(recipe2, "totalIngredients");
```

Властивість `totalIngredients` має бути встановлена в `null`.

```js
assert.isNull(recipe2.totalIngredients);
```

Ваш об’єкт `recipe2` має мати властивість `difficultyLevel`.

```js
assert.property(recipe2, "difficultyLevel");
```

Властивість `difficultyLevel` має бути встановлена в порожній рядок.

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
