---
id: 66fbcf750a62784cf11f5633
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Створіть `getDifficultyLevel` функцію, яка приймає число, що вказує час приготування, як параметр.

Якщо час приготування менший або дорівнює `30`, функція має повертати `"easy"`. Якщо він менший або дорівнює `60`, функція має повертати `"medium"`. Інакше функція має повертати `"hard"`.

# --hints--

Ви маєте створити `getDifficultyLevel` функцію.

```js
assert.isFunction(getDifficultyLevel);
```

Ваша `getDifficultyLevel` функція має мати один параметр.

```js
assert.lengthOf(getDifficultyLevel, 1);
```

Ваша `getDifficultyLevel` функція має повертати `"easy"`, коли час приготування менший або дорівнює `30`.

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

Ваша `getDifficultyLevel` функція має повертати `"medium"`, коли час приготування більший за `31` і менший або дорівнює `60`.

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

Ваша `getDifficultyLevel` функція має повертати `"hard"`, коли час приготування більший за `60`.

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
