---
id: 66fbcf750a62784cf11f562d
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Створіть об’єкт з іменем `recipe1`. Всередині об’єкта `recipe1` створіть властивість `name` зі значенням `"Spaghetti Carbonara"`.

Також всередині об’єкта `recipe1` створіть властивість `ingredients` зі значенням у вигляді масиву. Масив повинен містити `spaghetti`, `Parmesan cheese`, `pancetta` і `black pepper`.

# --hints--

Ви повинні створити об’єкт з іменем `recipe1`.

```js
assert.isObject(recipe1);
```

Ваш об’єкт `recipe1` має містити властивість `name`.

```js
assert.property(recipe1, 'name');
```

Властивість `name` має бути встановлена в `Spaghetti Carbonara`.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

Ваш об’єкт `recipe1` має містити властивість `ingredients`.

```js
assert.property(recipe1, 'ingredients');
```

Властивість `ingredients` має мати значення у вигляді масиву.

```js
assert.isArray(recipe1.ingredients);
```

Масив у властивості `ingredients` має містити `spaghetti`, `Parmesan cheese`, `pancetta` і `black pepper`.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
