---
id: 66fbcf750a62784cf11f5634
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Настав час протестувати кожну з функцій. Ви можете використати будь-який із рецептів для цього, але цей урок почнеться з `recipe1`.

Створіть дві нові змінні: `recipe1TotalIngredients` і `recipe1DifficultyLevel`. Призначте їм значення, викликавши відповідну функцію для кожної змінної та передавши відповідну властивість `recipe1`.

Нарешті, виведіть кожну змінну в консоль, щоб побачити результати.

# --hints--

Ви повинні створити змінну `recipe1TotalIngredients`.

```js
assert.isNotNull(recipe1TotalIngredients);
```

Ваша змінна `recipe1TotalIngredients` має бути призначена `getTotalIngredients(recipe1.ingredients)`.

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

Ви повинні вивести `recipe1TotalIngredients` у консоль.

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

Ви повинні створити змінну `recipe1DifficultyLevel`.

```js
assert.isNotNull(recipe1DifficultyLevel);
```

Ваша змінна `recipe1DifficultyLevel` має бути призначена `getDifficultyLevel(recipe1.cookingTime)`.

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

Ви повинні вивести `recipe1DifficultyLevel` у консоль.

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
