---
id: 66ff02792b3042a2b811f475
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Додайте наступні властивості до об’єкта `recipe1`:

| Ключ | Значення   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

Властивості з `null` та порожніми рядковими значеннями буде оновлено пізніше після їх обчислення.

# --hints--

Ваш об’єкт `recipe1` має містити властивість `cookingTime`.

```js
assert.property(recipe1, "cookingTime");
```

Властивість `cookingTime` має бути встановлена в `22`.

```js
assert.equal(recipe1.cookingTime, 22);
```

Ваш об’єкт `recipe1` має містити властивість `totalIngredients`.

```js
assert.property(recipe1, "totalIngredients");
```

Властивість `totalIngredients` має бути встановлена в `null`.

```js
assert.isNull(recipe1.totalIngredients);
```

Ваш об’єкт `recipe1` має містити властивість `difficultyLevel`.

```js
assert.property(recipe1, "difficultyLevel");
```

Властивість `difficultyLevel` має бути встановлена в порожній рядок.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
