---
id: 670e4f45f7116c0f216a5177
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Перед тим як рухатися далі, потренуйтеся отримувати доступ до властивостей об’єкта.

Ви можете використовувати або крапкову (`.`), або дужкову (`[]`) нотацію для цього. Ось приклад:

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Отримайте властивість `name` об’єкта `recipe1` і призначте її змінній `recipe1Name`.

Далі отримайте властивість `cookingTime` об’єкта `recipe2` і призначте її змінній `recipe2CookingTime`.

Нарешті, отримайте властивість `ingredients` об’єкта `recipe3` і призначте її змінній `recipe3Ingredients`.

Переконайтеся, що всі створені змінні виведені в консоль.

# --hints--

Ви повинні створити змінну `recipe1Name`.

```js
assert.isDefined(recipe1Name);
```

Ви повинні призначити значення властивості `name` об’єкта `recipe1` вашій змінній `recipe1name`.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

Ви повинні вивести `recipe1Name` у консоль.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

Ви повинні створити змінну `recipe2CookingTime`.

```js
assert.isDefined(recipe2CookingTime);
```

Ви повинні призначити значення властивості `cookingTime` об’єкта `recipe2` вашій змінній `recipe2CookingTime`.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

Ви повинні вивести `recipe2CookingTime` у консоль.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

Ви повинні створити змінну `recipe3Ingredients`.

```js
assert.isDefined(recipe3Ingredients);
```

Ви повинні призначити значення властивості `ingredients` об’єкта `recipe3` вашій змінній `recipe3Ingredients`.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

Ви повинні вивести `recipe3Ingredients` у консоль.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
