---
id: 69161f4c0c9c8c48b3c843c5
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Створіть об’єкт `recipe3` з такими властивостями та значеннями:

|Ключ|Значення|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

Ви маєте створити об’єкт з іменем `recipe3`.

```js
assert.isObject(recipe3);
```

Ваш об’єкт `recipe3` має містити властивість `name`.

```js
assert.property(recipe3, "name");
```

Властивість `name` має бути встановлена в `Vegetable Stir Fry`.

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

Ваш об’єкт `recipe3` має містити властивість `ingredients`.

```js
assert.property(recipe3, "ingredients");
```

Властивість `ingredients` має мати значення у вигляді масиву.

```js
assert.isArray(recipe3.ingredients);
```

Масив у властивості `ingredients` має містити `broccoli`, `carrot` та `bell pepper`.

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

Ваш об’єкт `recipe3` має містити властивість `cookingTime`.

```js
assert.property(recipe3, "cookingTime");
```

Значення властивості `cookingTime` має бути числом.

```js
assert.isNumber(recipe3.cookingTime);
```

Властивість `cookingTime` має бути встановлена в `15`.

```js
assert.equal(recipe3.cookingTime, 15);
```

Ваш об’єкт `recipe3` має містити властивість `totalIngredients`.

```js
assert.property(recipe3, "totalIngredients");
```

Властивість `totalIngredients` має бути встановлена в `null`.

```js
assert.isNull(recipe3.totalIngredients);
```

Ваш об’єкт `recipe3` має містити властивість `difficultyLevel`.

```js
assert.property(recipe3, "difficultyLevel");
```

Властивість `difficultyLevel` має бути встановлена в порожній рядок.

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
