---
id: 64dc8b1faf9c60304ca8a9b6
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Projekt lösen Sie das mathematische Rätsel, das als Turm von Hanoi bekannt ist. Das Rätsel besteht aus drei Stäben und einer Anzahl von Scheiben mit unterschiedlichen Durchmessern.

Das Ziel dieses Rätsels ist es, die Scheiben vom ersten Stab auf den dritten Stab zu bewegen, wobei bestimmte Regeln beachtet werden müssen, die das Platzieren einer größeren Scheibe auf einer kleineren verbieten.

Beginnen Sie damit, ein leeres Wörterbuch mit dem Namen `rods` zu erstellen, um die Stäbe darzustellen.

# --hints--

Sie sollten eine Variable mit dem Namen `rods` haben.

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

Ihre `rods`-Variable sollte ein leeres Wörterbuch sein.

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
