---
id: 64dc8ea01436383a88256d7f
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Das Rätsel beginnt damit, dass die Scheiben auf der ersten Stange gestapelt sind, in abnehmender Größe, mit der kleinsten Scheibe oben und der größten Scheibe unten. Sie müssen Ihre erste Liste mit Zahlen füllen, die die verschiedenen Scheibengrößen darstellen.

Anstatt die Elemente manuell zur ersten Liste hinzuzufügen, erzeugen Sie eine Folge von Zahlen, die von `3` bis `1` herunterzählt, indem Sie die `range()`-Funktion verwenden, und weisen Sie sie `rods['A']` zu. Hierbei steht `3` für die größte Scheibe am unteren Ende des Stapels und `1` für die kleinste Scheibe oben auf dem Stapel.

Die Syntax lautet `range(x, y, h)`, wobei `x` die Startzahl (einschließlich), `y` die letzte Zahl (ohne Einschluss) und `h` die Differenz zwischen einer Zahl und der nächsten in der Folge ist.

# --hints--

Sie sollten die `range()`-Funktion verwenden, um eine Folge von Zahlen `rods['A']` zuzuweisen. Die Syntax zum Aufrufen der `range()`-Funktion ist `range(x, y, h)`.

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

Sie sollten die `range()`-Funktion verwenden, um die Zahlenfolge von `3` bis `1` `rods['A']` zuzuweisen.

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
