---
id: 64dc976bf864693e668d67e8
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Rufen Sie jetzt Ihre Funktion auf und sehen Sie sich die Ausgabe im Terminal an.

# --hints--

Sie sollten Ihre `move()` Funktion aufrufen.

```js
({test: () => assert.match(code, /^move\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    print(rods)
--fcc-editable-region--

--fcc-editable-region--
```
