---
id: 64dc9004f658183afa85f4f1
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Überprüfen Sie nun den Datentyp Ihres `'A'`-Schlüssels, indem Sie ihn an die `type()`-Funktion übergeben und das Ergebnis im Terminal ausgeben.

# --hints--

Sie sollten Ihren `'A'`-Schlüssel an die `type()`-Funktion übergeben.

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

Sie sollten den Typ von `rods['A']` mit der `print()`-Funktion ausgeben.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
