---
id: 64dc922df2919e3c38cead61
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Deklarieren Sie ganz oben in Ihrem Code eine Variable namens `NUMBER_OF_DISKS`, um die Anzahl der Scheiben zu speichern, und weisen Sie ihr den Wert `3` zu. Ersetzen Sie dann das erste Argument, das an die `range()`-Funktion übergeben wird, durch Ihre neue Variable.

# --hints--

Sie sollten eine Variable namens `NUMBER_OF_DISKS` deklarieren.

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

Die `NUMBER_OF_DISKS`-Variable sollte den Wert `3` haben.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

Sie sollten den Wert des ersten Arguments der `range()`-Funktion in die `NUMBER_OF_DISKS`-Variable ändern.

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
