---
id: 64dc90e837670a3ba26a1e61
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Da der Typ jetzt wie gefordert `list` ist, können Sie den `print()`-Aufruf entfernen.

# --hints--

Sie sollten `print(type(rods['A']))` aus Ihrem Code entfernen.

```js
({ test: () => assert.match(code, /\}(?!\s+print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")A\1\s*\]\s*\)\s*\))/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
