---
id: 64dce7ee2494e82bf4205a3f
title: Schritt 24
challengeType: 20
dashedName: step-24
---

# --description--

Der andere Fall, in dem Sie die Scheibe unbedingt von `source` nach `target` verschieben müssen, ist, wenn die Quellliste **nicht** leer ist **und** die letzte Scheibe in `source` kleiner ist als die letzte Scheibe in `target`.

Fügen Sie eine `elif`-Anweisung hinzu, um diese Bedingung zu überprüfen. Setzen Sie dann die Variable `forward` auf `True`, wenn die Bedingung erfüllt ist.

# --hints--

Sie sollten eine `elif`-Anweisung haben, um zu prüfen, ob `rods[source]` nicht leer ist und `rods[source][-1] < rods[target][-1]`.

```js
const allowedRes = [
    "rods\\s*\\[\\s*source\\s*\\]",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>\\s*0",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>=\\s*1",
];
const re = new RegExp(`elif\\s+((${allowedRes.join(")|(")}))\\s+and\\s+(rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*<\\s*rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]|rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*>\\s*rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\])\\s*:`);
assert.match(code, re);
```

Sie sollten `forward` auf `True` innerhalb Ihrer neuen `elif`-Anweisung setzen. Die vorherigen Zuweisungen von `forward` sollten Sie nicht ändern.

```js
const allowedRes = [
    "rods\\s*\\[\\s*source\\s*\\]",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>\\s*0",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>=\\s*1",
];
const re = new RegExp(`elif\\s+((${allowedRes.join(")|(")}))\\s+and\\s+(rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*<\\s*rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]|rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*>\\s*rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\])\\s*:\\s+forward\\s*=\\s*True`);
assert.match(code, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
--fcc-editable-region--            
            if not rods[target]:
                forward = True
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
