---
id: 64dc97ab25730b3e9de86ea5
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Am Ende dieses Projekts werden Sie eine rekursive Lösung für das Turm von Hanoi-Rätsel erstellen, aber jetzt werden Sie einen iterativen Ansatz für dieses Problem erkunden.

Beginnen Sie damit, eine `for`-Schleife in Ihre Funktion einzufügen, die die `number_of_moves` durchläuft und die aktuelle Iterationsnummer ausgibt.

# --hints--

Sie sollten eine `for`-Schleife schreiben, um die Anzahl der Züge durchzulaufen. Verwenden Sie dafür die `range()`-Funktion.

```js
({ test: () => {
    const rgs = [
        /for\s+\w+\s+in\s+range\s*\(\s*number_of_moves\s*\)\s*:/,
        /for\s+\w+\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*\)\s*:/,
        /for\s+\w+\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*,\s*1\s*\)\s*:/
    ]
    const loop = rgs.some(r => code.match(r));
    assert.isTrue(loop);
  }
})
```

Sie sollten bei jeder Iteration die aktuelle Zugnummer ausgeben.

```js
({ test: () => {
    const rgs = [
        /for\s+(\w+)\s+in\s+range\s*\(\s*number_of_moves\s*\)\s*:\s+print\s*\(\s*\1\s*\)/,
        /for\s+(\w+)\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*\)\s*:\s+print\s*\(\s*\1\s*\)/,
        /for\s+(\w+)\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*,\s*1\s*\)\s*:\s+print\s*\(\s*\1\s*\)/
    ]
    const loop = rgs.some(r => code.match(r));
    assert.isTrue(loop);
  }
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
