---
id: 64dc92a9718fb53ca3d1bad7
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Das Turm-von-Hanoi-Rätsel kann in 2<sup>n</sup> - 1 Zügen gelöst werden, wobei n die Anzahl der Scheiben ist. Deklarieren Sie eine Variable namens `number_of_moves` und weisen Sie dieser Variable die Gesamtzahl der Züge zu.

Der Potenzoperator in Python ist `**`.

# --hints--

Sie sollten eine Variable namens `number_of_moves` deklarieren.

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

Der Wert von `number_of_moves` sollte der Ausdruck sein, um die Anzahl der Züge zu berechnen.

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
