---
id: 64dc97005bc7943e2978df0a
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Jetzt können Sie Ihren `print()` Aufruf entfernen. Entfernen Sie dann innerhalb der `move()` Funktion das `pass` Schlüsselwort und geben Sie den Inhalt Ihres `rods` Wörterbuchs aus.

# --hints--

Sie sollten kein `print(number_of_moves)` in Ihrem Code haben.

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

Sie sollten das `pass` Schlüsselwort entfernen und das `rods` Wörterbuch ausgeben.

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
