---
id: 64dc963d866fff3dd0329993
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Geben Sie die Variable aus, die Sie im vorherigen Schritt deklariert haben, und ändern Sie gerne die Anzahl der Scheiben, um zu sehen, wie schnell die erforderliche Mindestanzahl der Züge steigt.

# --hints--

Sie sollten die `number_of_moves`-Variable ausgeben.

```js
({ test: () => assert.match(code, /print\s*\(\s*number_of_moves\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move():
    pass
```
