---
id: 64dc90949ff85e3b37be40d0
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Die `range()`-Funktion gibt eine unveränderliche Folge von Zahlen zurück. Wie Sie sehen, ist der Datentyp von `rods['A']` `range`, aber Sie möchten, dass es eine Liste ist.

Übergeben Sie Ihren `range()`-Aufruf an die `list()`-Funktion, um das zu erreichen.

# --hints--

Sie sollten Ihren `range()`-Aufruf an die `list()`-Funktion übergeben.

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
