---
id: 64dc985e6720a23edac63b51
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Die erlaubten Bewegungen der Scheiben zwischen den Stäben folgen einem sich alle drei Züge wiederholenden Muster. Zum Beispiel sind Bewegungen zwischen Stab `A` und Stab `C` beim ersten, vierten und siebten Zug erlaubt, wenn der Rest der Division der Zugnummer durch 3 gleich 1 ist.

Ersetzen Sie innerhalb der zuvor erstellten `for`-Schleife den bestehenden `print()`-Aufruf durch eine `if`-Anweisung, die ausgelöst wird, wenn `(i + 1) % 3 == 1` zutrifft. Innerhalb dieser `if`-Anweisung geben Sie `f'Move {i + 1} allowed between {source} and {target}'` mit einem f-String aus. Bitte beachten Sie, dass `i + 1` die Zugnummer ist, da `i` während der ersten Iteration null ist.

# --hints--

Sie sollten eine `if`-Anweisung haben, die ausgeführt wird, wenn `(i + 1) % 3 == 1` zutrifft.

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

Sie sollten den String `f'Move {i + 1} allowed between {source} and {target}'` mit einem f-String ausgeben.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
