---
id: 64dc916138dbe03bdfe6cec6
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Das Ziel des Turms von Hanoi ist es, alle Scheiben auf die letzte Stange zu bewegen. Dafür müssen Sie drei einfache Regeln befolgen:

1. Sie dürfen nur die obersten Scheiben bewegen
2. Sie dürfen jeweils nur eine Scheibe bewegen
3. Sie dürfen keine größeren Scheiben auf kleinere legen

Deklarieren Sie unter Ihrem bestehenden Code eine leere Funktion mit dem Namen `move`. Später werden Sie diese Funktion verwenden, um die Scheiben zwischen den Stangen zu bewegen. Um jetzt Fehler zu vermeiden, verwenden Sie das Schlüsselwort `pass` im Funktionskörper.

# --hints--

Sie sollten eine leere Funktion mit dem Namen `move` deklarieren. Denken Sie daran, das Schlüsselwort `pass` mit der richtigen Einrückung im Funktionskörper zu verwenden.

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
