---
id: 64ddce6f6e8066173208235b
title: Schritt 27
challengeType: 20
dashedName: step-27
---

# --description--

Wenn `forward` `False` ist, muss die Scheibe in die entgegengesetzte Richtung bewegt werden. Schreiben Sie eine `else`-Klausel dafür. Geben Sie die Bewegung aus und ändern Sie den Inhalt der Listen entsprechend.

# --hints--

Sie sollten eine `else`-Klausel haben, um die Scheiben in die entgegengesetzte Richtung zu bewegen.

```js
({ test: () => assert.match(code, /rods\s*\[\s*target\s*\]\s*\.append\s*\(\s*rods\s*\[\s*source\s*\]\s*\.pop\s*\(\s*\)\s*\)\s+else\s*:/) })
```

Sie sollten `print(f'Moving disk {rods[target][-1]} from {target} to {source}')` innerhalb Ihrer `else`-Klausel haben.

```js
({ test: () => assert.match(code, /else\s*:\s+print\s*\(\s*f('|")Moving\sdisk\s\{\s*rods\s*\[\s*target\s*\]\s*\[\s*-\s*1\s*\]\s*\}\sfrom\s\{\s*target\s*\}\sto\s\{\s*source\s*\}\1\s*\)/) })
```

Sie sollten das letzte Element aus `rods[target]` entfernen.

```js
({ test: () => assert.match(code, /rods\s*\[\s*target\s*\]\s*\.pop\s*\(\s*\)/) })
```

Sie sollten `rods[source].append(rods[target].pop())` in Ihrem Code haben.

```js
({ test: () => assert.match(code, /rods\s*\[\s*source\s*\]\s*\.append\s*\(\s*rods\s*\[\s*target\s*\]\s*\.pop\s*\(\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if not rods[target]:
                forward = True
            elif rods[source] and rods[source][-1] < rods[target][-1]:
                forward = True
--fcc-editable-region--                
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
                rods[target].append(rods[source].pop())
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
