---
id: 64de1469fccce61940f498c4
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Das `rods`-Dictionary soll die drei Stäbe mit ihren Scheiben darstellen. Geben Sie ihm die Strings `'A'`, `'B'` und `'C'` als Schlüssel und setzen Sie jeden von ihnen auf eine leere Liste.

# --hints--

Ihr `rods`-Dictionary sollte einen `'A'`-Schlüssel enthalten.

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` sollte eine leere Liste sein.

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

Ihr `rods`-Dictionary sollte einen `'B'`-Schlüssel enthalten.

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` sollte eine leere Liste sein.

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

Ihr `rods`-Dictionary sollte einen `'C'`-Schlüssel enthalten.

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` sollte eine leere Liste sein.

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
