---
id: 64df3f1011888113fbd3d81b
title: Schritt 44
challengeType: 20
dashedName: step-44
---

# --description--

Um das Rätsel mit Rekursion zu lösen, ist der erste Schritt, das ursprüngliche Problem in kleinere Teilprobleme zu zerlegen.

Die endgültige Konfiguration mit `n` Scheiben, die in absteigender Reihenfolge auf der dritten Stange gestapelt sind, kann erreicht werden, indem Sie:

- `n - 1` Scheiben von der Ausgangsstange zur Hilfsstange bewegen
- die größte Scheibe von der Ausgangsstange zum Ziel bewegen
- und dann die `n - 1` Scheiben von der Hilfsstange zum Ziel bewegen.

Das Erste, was die `move`-Funktion tun sollte, ist, sich selbst mit `n - 1` als erstem Argument aufzurufen. Wenn Sie dies jedoch ohne Definition eines Basisfalls versuchen, erhalten Sie einen `RecursionError`. Das passiert, weil die Funktion sich unendlich oft selbst aufruft.

Fügen Sie vor Ihrem Kommentar und Ihrem `print()`-Aufruf den rekursiven Funktionsaufruf mit `n - 1` als erstem Argument hinzu und stellen Sie sicher, dass der Funktionskörper nur ausgeführt wird, wenn `n` größer als null ist. Lassen Sie vorerst die anderen Argumente in der gleichen Reihenfolge.

# --hints--

Der Funktionskörper von `move` sollte mit einer `if`-Anweisung beginnen, die ausgelöst wird, wenn `n` größer als null ist.

```js
({ test: () => assert.match(code, /def\s+move\(\s*n\s*,\s*source\s*,\s*auxiliary\s*,\s*target\s*\)\s*:\s+if\s+n\s*>\s*0\s*:/) })
```

Sie sollten `move(n - 1, source, auxiliary, target)` innerhalb der `if`-Anweisung aufrufen. Denken Sie daran, Ihren `print()`-Aufruf einzurücken.

```js
const transformedCode = "\n" + code.replace(/\r/g, '');
const move = __helpers.python.getDef(transformedCode, "move");
const { function_parameters, function_body } = move;

assert.match(function_parameters, /n\s*,\s*source\s*,\s*auxiliary\s*,\s*target/);
const functionBodyIndent = function_body.match(/ +/)[0];
const re = new RegExp(`${functionBodyIndent}if\\s+n\\s*>\\s*0\\s*:\\s+^${functionBodyIndent}( +)move\\(\\s*n\\s*-\\s*1\\s*,\\s*source\\s*,\\s*auxiliary\\s*,\\s*target\\s*\\).+?^${functionBodyIndent}\\1print\\s*\\(\\s*rods\\s*,\\s*("|')\\\\n\\2\\s*\\)`, "ms");
assert.match(function_body, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
