---
id: 64df45a3ad4f8719e5355244
title: Schritt 46
challengeType: 20
dashedName: step-46
---

# --description--

Die Schritte zum Bewegen von `n - 1` Scheiben können weiter unterteilt werden, bis nur noch eine einzelne Scheibe betrachtet wird. Dies wird der erste ausgeführte Schritt sein. Nach dem ersten Schritt werden die folgenden Schritte durch das Zurückverfolgen der rekursiven Aufrufe erzeugt. Beachten Sie, dass in jedem rekursiven Schritt die Rolle der Stäbe zwischen Quelle, Ziel und Hilfsstab wechselt.

Im Moment gibt jeder rekursive Aufruf das `rods` Wörterbuch aus, ohne Änderungen an den Listen vorzunehmen. Vor dem `print()` Aufruf entfernen Sie das letzte Element aus der `rods[source]` Liste und fügen es der `rods[target]` Liste hinzu.

# --hints--

Sie sollten das letzte Element aus der `rods[source]` Liste entfernen und es der `rods[target]` Liste hinzufügen, bevor Sie den `print` Aufruf tätigen.

```js
({ test: () => assert.isTrue(runPython(`
_log = []
def capture_print(*args):
    _log.append(repr(args))

old_print = print
print = capture_print

old_rods = rods
rods2 = {
    'A': list(range(5, 0, -1)),
    'B': [],
    'C': []
}
rods = rods2

move(5, 'A', 'B', 'C')
rods = old_rods
print = old_print

_expected_prints = [
r"({'A': [5, 4, 3, 2], 'B': [], 'C': [1]}, '\\n')",
r"({'A': [5, 4, 3], 'B': [], 'C': [1, 2]}, '\\n')",
r"({'A': [5, 4], 'B': [], 'C': [1, 2, 3]}, '\\n')",
r"({'A': [5], 'B': [], 'C': [1, 2, 3, 4]}, '\\n')",
r"({'A': [], 'B': [], 'C': [1, 2, 3, 4, 5]}, '\\n')",
]

for args in _log:
    if not _expected_prints:
        break
    if args == _expected_prints[0]:
        _expected_prints.pop(0)

not rods2['A'] and rods2['C'] == list(range(1, 6)) and not _expected_prints
`))})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # display starting configuration
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
