---
id: 64df47b32b92301a815d5ef8
title: Schritt 49
challengeType: 20
dashedName: step-49
---

# --description--

Zunächst befasst sich der gerade hinzugefügte rekursive Aufruf mit dem Teilproblem, `n - 1` Scheiben auf die zweite Stange zu bewegen.

Aus diesem Grund entspricht das Argument `target` Ihrer zweiten Stange, während das Argument `auxiliary` die dritte Stange ist. Beachten Sie, dass diese im Verlauf der Rekursion weiterhin vertauscht werden.

Korrigieren Sie die Reihenfolge der Argumente, indem Sie `target` und `auxiliary` in Ihrem rekursiven Aufruf tauschen.

# --hints--

Sie sollten die Reihenfolge der Argumente in Ihrem `move(n - 1, source, auxiliary, target)` Aufruf ändern.

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
