---
id: 64df496c6a8ddf1b38db1ed6
title: Schritt 50
challengeType: 20
dashedName: step-50
---

# --description--

In einem vorherigen Schritt haben Sie den Code geschrieben, um die größte Scheibe des Teilproblems auf die Zielstange zu bewegen.

Jetzt müssen Sie nur noch einen weiteren rekursiven Aufruf hinzufügen, um die `n - 1` Scheiben, die Sie bereits verschoben haben, zu bewegen. Kopieren Sie den ersten rekursiven Aufruf und fügen Sie ihn am Ende des `if` Blocks ein.

Beachten Sie, dass die Argumente der Funktion nicht in der richtigen Reihenfolge sind. Versuchen Sie, die korrekte Reihenfolge herauszufinden.

# --hints--

Sie sollten `move(n - 1,  auxiliary, source, target)` am Ende Ihrer `move` Funktion aufrufen.

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
