---
id: 650c6082e5586f9e3acfcd3b
title: Schritt 52
challengeType: 20
dashedName: step-52
---

# --description--

Löschen Sie nun das `rods`-Dictionary und machen Sie dessen Schlüssel zu den Variablen `A`, `B` und `C`, jeweils mit den ursprünglichen Werten. Refaktorieren Sie Ihren Code, um diese Änderungen widerzuspiegeln. Wenn Sie die Ausgabe im Terminal sehen, haben Sie diesen Schritt korrekt ausgeführt.

# --hints--

Sie sollten `rods['A']` in eine Variable namens `A` umwandeln und den Wert beibehalten.

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

Sie sollten `rods['B']` in eine Variable namens `B` umwandeln und den Wert beibehalten.

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

Sie sollten `rods['C']` in eine Variable namens `C` umwandeln und den Wert beibehalten.

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

Sie sollten Ihren `print()`-Aufruf so ändern, dass `A`, `B` und `C` ausgegeben werden, anstatt des `rods`-Objekts. Behalten Sie das Zeilenumbruchzeichen im `print()`-Aufruf bei.

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

`target.append(source.pop())` sollte in Ihrem Code vorhanden sein.

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```

