---
id: 65782342850feb3b8d62f936
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Da Sie den Ausdruck `(i + 1) % 3` mehrfach verwenden werden, ist es praktisch, ihn in einer Variablen zu speichern.

Deklarieren Sie direkt über Ihrer `if`-Anweisung eine `remainder`-Variable und weisen Sie dieser Variable den Wert `(i + 1) % 3` zu.

# --hints--

Sie sollten direkt über dem `remainder`-Block eine `if`-Variable deklarieren.

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=/m) })
```

Der Wert Ihrer `remainder`-Variable sollte der Ausdruck `(i + 1) % 3` sein.

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=\s*\(\s*i\s*\+\s*1\s*\)\s*%\s*3/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        
        if (i + 1) % 3 == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
