---
id: 60eebd07ea685b0e777b5583
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Registration Form プロジェクトへようこそ！まず、ブラウザーがどの種類のドキュメントを読み込んでいるか認識できるように、ドキュメントの先頭に `!DOCTYPE html` 宣言を追加してください。

# --hints--

コードには `DOCTYPE` 参照が含まれている必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にはスペースを入れてください。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプは `html` に設定してください。

```js
assert(code.match(/html/gi));
```

`DOCTYPE` 宣言はタイプの後に `>` で閉じてください。

```js
assert(code.match(/html\s*>/gi));
```

`DOCTYPE` 宣言は HTML の先頭に置いてください。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
