---
id: 60f027099a15b00485563dd2
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`DOCTYPE`の下に、`html`属性が`lang`に設定された`en`要素を追加して、コードを置き始める場所を作ってください。

# --hints--

`DOCTYPE`宣言はHTMLの先頭にあるべきです。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html`要素は開始タグを持ち、`lang`属性の値は`en`であるべきです。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html`要素は終了タグを持つべきです。

```js
assert(code.match(/<\/html\s*>/));
```

`html`タグは正しい順序であるべきです。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

`html`要素は1つだけにしてください。

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
