---
id: 60f02e7361b68405e27b62a5
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`body` の中に、`h1` というテキストを持つ `Registration Form` を追加して、コンテンツの見出しコンテキストを提供してください。

# --hints--

`h1` の中に `body` を追加してください。

```js
assert.exists(document.querySelector('body > h1'));
```

`h1` に `Registration Form` というテキストを設定してください。

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
