---
id: 60f027c87bc98f050395c139
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

次に、`head` 要素の中に開く `body` タグと閉じる `html` タグを追加してください。

# --hints--

開く `head` タグがあるはずです。

```js
assert(code.match(/<head\s*>/i));
```

閉じる `head` タグがあるはずです。

```js
assert(code.match(/<\/head\s*>/i));
```

開く `body` タグがあるはずです。

```js
assert(code.match(/<body\s*>/i));
```

閉じる `body` タグがあるはずです。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` と `body` 要素は兄弟要素であるべきです。

```js
assert(document.querySelector('head')?.nextElementSibling.localName === 'body');
```

`head` 要素は `html` 要素の中にあるべきです。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は `html` 要素の中にあるべきです。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

</html>
--fcc-editable-region--
```
