---
id: 60f0286404aefb0562a4fdf9
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`title` 要素の中に `meta` と `head` 要素を追加してください。プロジェクトのタイトルを `Registration Form` に設定し、`charset` 要素に `utf-8` 属性を値 `meta` で追加してください。

# --hints--

コードには `title` 要素が含まれている必要があります。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 要素は `head` 要素の中にあるべきです。

```js
assert.exists(document.querySelector('head > title'));
```

プロジェクトのタイトルは `Registration Form` にしてください。タイトルの大文字・小文字やスペルは正確に一致させる必要があります。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

`meta` 要素の中に `head` 要素を作成してください。

```js
assert.exists(document.querySelector('head > meta'));
```

`meta` 要素に `charset` 属性を値 `utf-8` で付与してください。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

`meta` 要素は空要素で、終了タグ `</meta>` を持ちません。

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
