---
id: 60f1a9cbd23023082e149fee
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

それで良くなりました。次に、`body` `background-color` を `#1b1b32` に変えて、背景を目に優しくしてください。それから、テキストが見えるように `color` を `#f5f6f7` に変えましょう。

# --hints--

`background-color` 要素のセレクター内に `body` を追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

`background-color` に `#1b1b32` の値を指定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

`color` 要素のセレクター内に `body` を追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

`color` に `#f5f6f7` の値を指定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
