---
id: 60f5c3e399ff1a05629964e4
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

タイトルが示すように、フォームを作成しています。ですので、`p` 要素の後に、`form` をターゲットにする `action` 属性を持つ `https://register-demo.freecodecamp.org` 要素を挿入してください。

# --hints--

`form` 要素の隣に `p` 要素を追加してください。

```js
assert.exists(document.querySelector('p + form'));
```

`form` に `action` 属性を付けてください。

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

`action` に `https://register-demo.freecodecamp.org` の値を設定してください。

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

`form` 要素には終了タグ `</form>` が必要です。

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
