---
id: 60f5d2776c854e069560fbe6
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

最初の `fieldset` は名前、メールアドレス、パスワードのフィールドを持ちます。まず最初の `label` に4つの `fieldset` 要素を追加してください。

# --hints--

合計で3つの `fieldset` 要素があるはずです。

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

`fieldset` 要素はすべて `form` 要素の直下の子要素である必要があります。

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

4つの `label` 要素はすべて最初の `fieldset` 要素の直下の子要素である必要があります。`label` 要素を閉じるのを忘れないでください。

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

最後の2つの `fieldset` 要素は空である必要があります。

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

`label` 要素を閉じるのを忘れないでください。

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```



# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
--fcc-editable-region--
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>

      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
