---
id: 60f803d5241e6a0433a523a1
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--
`rem` 単位はルート `em` を表し、`html` 要素のフォントサイズに対して相対的です。

`label` 要素はデフォルトではインラインなので、すべて同じ行に横並びで表示され、テキストが読みづらくなります。別々の行に表示するには、`display: block` 要素に `label` を追加し、互いに区切るために `margin` の `0.5rem 0` を追加してください。

# --hints--

`label` セレクターを使うべきです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

`display` の値を持つ `block` プロパティを追加するべきです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

`margin` の値を持つ `0.5rem 0` プロパティを追加するべきです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
