---
id: 60fab8367d35de04e5cb7929
title: ステップ 31
challengeType: 0
dashedName: step-31
---

# --description--

新しく追加した `I accept the terms and conditions` の中で、`input` 要素の直後にテキスト `label` を追加してください。次に、テキスト `terms and conditions` を以下の場所にリンクしてください。

```md
https://www.freecodecamp.org/news/terms-of-service/
```

# --hints--

3番目の `I accept the terms and conditions` に続く `label` に `fieldset` テキストを追加してください。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label')?.innerText.trim(), 'I accept the terms and conditions');
```

利用規約にリンクするには `a` 要素を使うべきです。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label a'));
```

新しいテキストは `input` の中の `label` 要素の直後に置いてください。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input + a'));
```

`a` 要素に `href` として `https://www.freecodecamp.org/news/terms-of-service/` を設定してください。

```js
assert.match(document.querySelector('fieldset:nth-child(3) + label > input + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

テキスト `a` の周りに `terms and conditions` 要素だけをラップしてください。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent.trim(), 'terms and conditions');
```

アンカー要素内のテキストに余分な前後の空白があります。アンカーテキスト内のスペースは `terms`、`and`、`conditions` の単語間だけにしてください。

```js
const nestedAnchor = document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent;
const innerContent = nestedAnchor?.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
