---
id: 62cc5b1779e4d313466f73c5
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`link` 要素の中に `head` 要素をネストしてください。`rel` 属性を `stylesheet` の値で、`href` 属性を `styles.css` の値で与えてください。

# --hints--

コードには `link` 要素が含まれているはずです。

```js
assert.isNotNull(document.querySelector('link'));
```

既存の `head` タグは変更しないでください。終了タグを削除していないことを確認してください。

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

`link` 要素が1つあるはずです。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link` 要素は空要素で、終了タグ `</link>` を持ってはいけません。

```js
assert.notMatch(code, /<\/link>/);
```

`link` 要素は `head` 要素の中にあるはずです。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 要素は `rel` 属性を持ち、その値は `stylesheet` であるべきです。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 要素は `href` 属性を持ち、その値は `styles.css` であるべきです。

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
