---
id: 63541ef4f96cd82e8e6c788a
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`method` 属性は、`action` 属性で指定された URL にフォームデータを送信する方法を指定します。フォームデータは、`GET` リクエストで URL パラメータとして（`method="get"` を使って）送信するか、`POST` リクエストでリクエストボディ内のデータとして（`method="post"` を使って）送信できます。

フォームデータを `method` リクエストで送信するには、`POST` 属性を設定してください。

# --hints--

新しい `form` 要素を追加しないでください。

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

`form` 要素に `method` 属性を付ける必要があります。

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

`method` 属性は `post` に設定してください。

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
