---
id: 66db57ad34c7089b9b41bfd6
title: Crie um visualizador de lightbox
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

Uma lightbox é usada em sites para exibir conteúdo multimídia de forma mais envolvente por meio do uso de uma janela popup ou modal sobre o conteúdo principal da página.

Neste laboratório, você criará uma galeria lightbox que exibe imagens em tamanho real quando uma miniatura é clicada. Para cada imagem, duas versões são fornecidas: uma miniatura e uma imagem em tamanho real. A imagem em tamanho real é a mesma da miniatura, mas sem o sufixo `-thumbnail`.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter um `div` com uma classe `gallery` dentro do seu `body`.
1. Dentro do elemento `.gallery`, você deve ter três miniaturas de imagem, cada uma com uma classe `gallery-item`. Você deve usar os seguintes links para imagens em miniatura:
   
    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. Você deve ter um `div` com uma classe `lightbox` dentro do seu `body`.
1. Você deve ter um `button` com um `id` de `close-btn` dentro do seu elemento `.lightbox`. Você pode usar `&times;` como seu texto, se desejar.
1. Você deve ter um `img` com um `id` de `lightbox-image` dentro do seu elemento `.lightbox`.
1. Seu elemento `.lightbox` deve ter uma posição fixa para que a pré-visualização seja aberta sobre as imagens atuais.
1. Seu elemento `.lightbox` deve cobrir toda a viewport definindo a altura e a largura para 100% do contêiner. Você deve garantir que o elemento `.lightbox` comece no canto superior esquerdo do contêiner.
1. `.lightbox` deve ter uma cor de fundo. Inicialmente, sua propriedade `display` deve ser definida como `none` para ocultá-la.
1. Quando você clicar em um dos seus elementos `.gallery-item`, a propriedade `display` do elemento `.lightbox` deve ser definida como `flex` para tornar o elemento `.lightbox` e os dois elementos dentro dele visíveis.
1. Quando você clicar em um dos seus elementos `.gallery-item`, o atributo `src` do elemento `#lightbox-image` deve ser definido para uma versão em tamanho real da imagem clicada, removendo `-thumbnail` do atributo `src` da imagem. As imagens em tamanho real estão localizadas nos seguintes links:

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`
  
1. Quando seu elemento `.lightbox` estiver visível e você clicar no `#close-btn` ou no elemento `.lightbox`, o `display` do elemento `.lightbox` deve ser definido novamente como `none`.

**Nota:** Certifique-se de vincular sua folha de estilos e o arquivo JavaScript no seu HTML.

# --hints--

Você deve ter um `div` com a classe `gallery` dentro do seu elemento `body`.

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

Dentro do elemento `.gallery`, você deve ter três elementos `img` com a classe `gallery-item`.

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

Dentro do elemento `.gallery`, você deve ter um elemento `img` com o `src` definido como `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Dentro do elemento `.gallery`, você deve ter um elemento `img` com o `src` definido como `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Dentro do elemento `.gallery`, você deve ter um elemento `img` com o `src` definido como `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Você deve ter um elemento `div` com a classe `lightbox` dentro do seu elemento `body`.

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

Dentro do seu elemento `.lightbox`, você deve ter um elemento `button` com o `id` definido como `close-btn`.

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

Dentro do seu elemento `.lightbox`, você deve ter um elemento `img` com o `id` definido como `lightbox-image`.

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

Seu elemento `.lightbox` deve ter posicionamento fixo.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

Seu elemento `.lightbox` deve cobrir toda a viewport.

```js
const style = new __helpers.CSSHelp(document).getStyle(".lightbox");

const width = style?.width;
const height = style?.height;

assert.oneOf(
  width,
  ["100%", "100vw"]
);

assert.oneOf(
  height,
  ["100%", "100vh"]
);
```

Seu elemento `.lightbox` deve estar alinhado com o canto superior esquerdo do contêiner.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

Seu elemento `.lightbox` deve ter uma cor de fundo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

Seu elemento `.lightbox` deve estar oculto inicialmente.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

Quando você clicar em um dos seus elementos `.gallery-item`, a propriedade `display` do elemento `.lightbox` deve ser definida como `flex` para tornar `.lightbox` e os dois elementos dentro dele visíveis.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

Quando você clicar em um dos seus elementos `.gallery-item`, o atributo `src` do elemento `#lightbox-image` deve ser definido para uma versão em tamanho real da imagem clicada, removendo `-thumbnail` do atributo `src` da imagem.

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

Quando seu elemento `.lightbox` estiver visível e você clicar no botão `#close-btn`, o `display` do elemento `.lightbox` deve ser definido novamente como `none`.

```js
const lightbox = document.querySelector(".lightbox");
const closeBtn = document.getElementById("close-btn"); 
const galleryItem = document.querySelector(".gallery-item");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

closeBtn.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

Quando seu elemento `.lightbox` estiver visível e você clicar no elemento `.lightbox`, o `display` do elemento `.lightbox` deve ser definido novamente como `none`.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

lightbox.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```
