---
id: 66faac4139dbbd5dd632c860
title: 建置一個工作申請表單
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。 

**使用者故事：**

1. 你應該有一個帶有類別 `container` 的 `div` 元素。
1. 在 `div` 元素內，你應該有一個 `form` 元素。
1. 表單應包含一個 `input` 元素，其型別為 `text` 且 id 為 `name`，用於輸入使用者的全名。
1. 你應該有另一個 `input` 元素，其 `type` 為 `email` 且 `id` 為 `email`，用於輸入使用者的電子郵件位址。
1. 表單應該包含一個 id 為 `position` 的 `select` 元素，讓使用者可以選擇職位。
1. 你應該有一個 `fieldset` 元素，其類別為 `radio-group`。
1. 在 `.radio-group` 內，你應該有一組型別為 `radio` 的 `input` 元素及相關的標籤，用於選擇可用性選項（例如，Full-Time、Part-Time）。該群組的 `name` 應該是 `availability`。
1. 你應該有一個 id 為 `message` 的 `textarea` 元素用來輸入訊息。
1. 你應該將每個 `input` 元素與 `label` 元素關聯。
1. 你應該有一個 `button` 元素，其型別為 `submit`，用於提交表單。
1. 為你的 `input` 和 `textarea` 元素添加 `:focus` 偽類，以在聚焦時更改它們的邊框顏色。
1. `input`、`select` 和 `textarea` 元素應該有一個 `:invalid` 偽類別，當偵測到無效輸入時會將邊框顏色改為紅色。
1. `input`、`select` 和 `textarea` 元素應該有一個 `:valid` 偽類別，當輸入有效時會將邊框顏色變更為綠色。
1. `button` 元素應該有一個 `:hover` 偽類，當暫留時會改變背景顏色。
1. 在 `.radio-group input[type="radio"]` 上使用 `:checked` 偽類別，當圓鈕被選取時，為其添加邊框顏色、背景顏色和盒狀陰影。
1. 在圓鈕、單選按鈕上使用 `:checked` 偽類別，當選項被選取時改變相關的 `label` 的文字顏色。
1. 為你的 `input` 元素添加 `:first-of-type` 偽類別，以不同方式樣式化第一個輸入欄（例如，圓角）。

**注意：** 請務必在你的 HTML 中聯結你的樣式表並套用你的 CSS。

# --hints--

你應該有一個帶有類別 `container` 的 `div` 元素。

```js
assert.exists(document.querySelector("div.container"));
```

在 `div` 元素內，你應該有一個 `form` 元素。

```js
assert.exists(document.querySelector("div.container > form"));
```

表單應包含一個 `input` 元素，其型別為 `text` 且 id 為 `name`，用於輸入使用者的全名。

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

你應該有另一個 `input` 元素，其 `type` 為 `email` 且 `id` 為 `email`，用於輸入使用者的電子郵件位址。

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

表單應該包含一個 id 為 `position` 的 `select` 元素，並帶有一些 `option` 元素。

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

你應該有一個帶有類別 `.radio-group` 的 `fieldset` 或 `section` 元素。

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

在 `.radio-group` 裡，你應該有一組型別為 `radio` 的 `input` 元素，用於選擇可用性選項。該群組的 `name` 應該是 `availability`。

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

你應該有一個 id 為 `message` 的 `textarea` 元素用來輸入訊息。

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

你應該將每個 `input` 元素與 `label` 元素關聯。 

```js
const els = document.querySelectorAll('input');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

你應該有一個 `button` 元素，其型別為 `submit`，用於提交表單。

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

你應該為 `input` 和 `textarea` 元素添加 `:focus` 偽類，以便在聚焦時更改它們的邊框顏色。請依照給定順序使用列表選擇器。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'))
```

`input`、`select` 和 `textarea` 元素應該有一個 `:invalid` 偽類，當偵測到無效輸入時會將邊框顏色改為紅色。請依照給定順序使用列表選擇器。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

`input`、`select` 和 `textarea` 元素應該有一個 `:valid` 偽類，當輸入有效時會將邊框顏色改為綠色。請依照給定順序使用列表選擇器。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

`button` 元素應該有一個 `:hover` 偽類，當暫留時會改變背景顏色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

你應該在 `.radio-group input[type="radio"]` 上使用 `:checked` 偽類別，當圓鈕被選取時添加邊框顏色。

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

你應該在 `.radio-group input[type="radio"]` 上使用 `:checked` 偽類別，以便在選取圓鈕時添加背景色。

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

你應該在 `.radio-group input[type="radio"]` 上使用 `:checked` 偽類別，當圓鈕被選取時添加盒子陰影。

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

你應該在圓鈕上使用 `:checked` 偽類別，以在選取選項時更改相關 `label` 的文字顏色。

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

為你的 `input` 元素添加 `:first-of-type` 偽類別，以不同方式樣式化第一個輸入欄。

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
