---
id: 66b62d0ad68488dd76228d6c
title: تصحيح مغير لون الخلفية العشوائي
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

Camperbot يتعلم JavaScript وحاول بناء مغير لون خلفية عشوائي خاص به. لكنه ارتكب بعض الأخطاء على طول الطريق.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تجتاز حتى يعمل المختبر بشكل صحيح.

**قصص المستخدم:**

1. يجب عليك إصلاح الأخطاء النحوية في مصفوفة `darkColorsArr`.
1. يجب عليك إصلاح `ReferenceError` الذي يُرمى عند الإشارة إلى كائن `math` داخل الدالة `getRandomIndex`.
1. يجب عليك تقريب `randomIndex` إلى أقرب عدد صحيح للأسفل.
1. يجب تحديث قيمة المتغير `body` لاستخدام اسم الطريقة الصحيح وإصلاح `TypeError` المرتبط.
1. يجب إصلاح `bgHexCodeSpanElement` ليختار العنصر الذي يحمل `id` الخاص بـ `bg-hex-code` بشكل صحيح.
1. يجب تعديل المتغير `color` داخل `changeBackgroundColor` ليُعيّن إلى لون عشوائي من مصفوفة `darkColorsArr`.
1. يجب تحديث المتغير `btn` ليستهدف عنصر الزر. افتح `index.html` لترى اسم الـ `id` الصحيح لذلك الزر.
1. يجب إصلاح السطر الذي يضيف مستمع حدث `click` إلى `btn` بتمرير مرجع الدالة `changeBackgroundColor` إلى الطريقة.
1. يجب أن تتغير خلفيتك عشوائيًا إلى أحد الألوان في `darkColorsArr` في كل مرة يُنقر فيها على `#btn`.

# --hints--

يجب عليك حل الأخطاء النحوية في مصفوفة `darkColorsArr`. تذكر أن عناصر المصفوفة يجب أن تفصل بفواصل.

```js
assert.isArray(darkColorsArr);
```

يجب إصلاح خطأ الكتابة في السطر `math.random()` بحيث تكون الحروف الكبيرة والصغيرة صحيحة.

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

يجب تقريب ناتج `darkColorsArr.length * Math.random()` إلى أقرب عدد صحيح للأسفل.

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

يجب تحديث قيمة متغير `body` بحيث يستخدم `queryselector` صيغة camel case الصحيحة.

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

يجب إصلاح `"bg-hex-code"` ليشير بشكل صحيح إلى `id` الخاص بـ `bg-hex-code`. تذكر أن محددات `id` في CSS تبدأ برمز الهاش (`#`).

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

يجب استدعاء الدالة `getRandomIndex` داخل السطر `darkColorsArr[getRandomIndex]`.

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

يجب إصلاح اسم `id` في السطر `"#click-btn"` ليطابق اسم الـ `id` الصحيح في ملف `index.html`.

```js
assert.match(btn?.id, /btn/);
```

لا يجب استدعاء الدالة `changeBackgroundColor`. بدلاً من ذلك، مرر مرجع الدالة إلى طريقة `addEventListener`.

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];
  
  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())
  
  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)
  
  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];
  
    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)
  
  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
