---
id: 68c362b379059c388d3874f2
title: 建置一款數位寵物遊戲
challengeType: 25
dashedName: lab-digital-pet-game
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成此實驗。

**使用者故事：**

此應用程式應有兩個視圖。每個視圖的需求包含在其餘的使用者故事中：

第一個視圖：

- 應包含一個接受單一字串輸入的表單。
- 應要求使用者為寵物命名。
- 應包含一個提交按鈕以開始遊戲。
- 提交後應消失。

第二個視圖：

- 應在第一個視圖提交後出現。
- 應包含三個 `<button>` 元素。
- 每個按鈕元素應對應以下三個動作之一：`[EAT, PLAY, SLEEP]`。

  - `EAT` 按鈕應降低寵物的 `Hunger` 值，並增加寵物的 `Energy` 值。
  - `PLAY` 按鈕應降低寵物的 `Energy` 值，並增加寵物的 `Happiness` 值。
  - `SLEEP` 按鈕應增加寵物的 `Hunger` 值，並增加寵物的 `Energy` 值。

- 應視覺化追蹤寵物的三個屬性。這些是：`['Hunger', 'Happiness', 'Energy']`。
- 追蹤這三個屬性的每個元素應：

  - 包含與該屬性相關的文字，不分大小寫，例如 `hunger`、`Hunger`、`hUnGeR` 都是有效的。
  - 不包含與其他屬性相關的文字，即 `Hunger` 追蹤元素不得包含文字 `Happiness`。
  - 包含該屬性當前值的明確數字表示。

- 寵物可用的每個 `Action` 應包含在一個列舉中。
- 每個 `Action` `button` 應包含與該按鈕動作相關的以下 `id` 中的一個：`['eat-action', 'play-action', 'sleep-action']`。
- 寵物應有以下心情，包含在 `PetMood` 列舉中：`['Happy', 'Excited', 'Content', 'Sad', 'Tired', 'Sick', 'Hungry']`。

  - 如果 `Hungry`，寵物應為 `Hunger > 70`。
  - 如果 `Tired`，寵物應為 `Energy < 30`。
  - 如果 `Sad`，寵物應為 `Happiness < 30`。
  - 如果 `Excited` 且 `Happiness > 80`，寵物應為 `Energy > 70`。
  - 如果 `Happy`，寵物應為 `Happiness > 60`。
  - 否則，寵物應為 `Content`。

- 寵物應有多個與每個 `Mood` 相關的影像或表情符號。

  - 這些表情符號或影像應與當前 `PetMood` 相關，並以 `Record` 連結。
  - 每個 `Record` 項目的鍵應為相關的 `PetMood` 列舉值。

# --before-all--

```js
globalThis._PET_NAME = 'Fitzgerald';
globalThis._clock = __FakeTimers.install();
```

# --after-all--

```js
globalThis._clock.uninstall();
delete globalThis._clock;
delete globalThis._PET_NAME;
```

# --hints--

當頁面載入時，應有一個可見的 `form` 元素。

```js
assert.isNotNull(document.querySelector('form'), 'The form element does not exist at load.');
```

此表單應包含一個 id 為 `pet-name` 的輸入欄位。

```js
assert.isNotNull(document.querySelector('#pet-name'));
```

此表單應包含一個 `button`。

```js
assert.isNotNull(document.querySelector('button'));
```

當輸入有值且按下按鈕後，表單應不再可見。

```js
document.querySelector('#pet-name').value = _PET_NAME;
document.querySelector('button').click();
await _clock.tickAsync(500);

assert.isNull(document.querySelector('form'));
```

第二個視圖應包含一個類別為 `pet-name` 的元素。

```js
const petNameEl = document.querySelector('.pet-name');
assert.exists(petNameEl);
```

第二個視圖應包含提供的寵物名稱。

```js
const petNameEl = document.querySelector('.pet-name');
assert.strictEqual(petNameEl.textContent, _PET_NAME);
```

第二個視圖應包含三個 `button`。

```js
const buttons = document.querySelectorAll('button');
assert.equal(buttons.length, 3);
```

恰好有一個按鈕包含 id `#eat-action`。

```js
const eatButton = document.querySelectorAll('button#eat-action');
assert.equal(eatButton.length, 1);
```

恰好有一個按鈕包含 id `#play-action`。

```js
const playButton = document.querySelectorAll('button#play-action');
assert.equal(playButton.length, 1);
```

恰好有一個按鈕包含 id `#sleep-action`。

```js
const sleepButton = document.querySelectorAll('button#sleep-action');
assert.equal(sleepButton.length, 1);
```

應有恰好三個類別為 `.stat` 的元素。

```js
const statEls = document.querySelectorAll('.stat');
assert.equal(statEls.length, 3);
```

每個寵物屬性應包含恰好一個類別為 `.stat-value` 的元素。

```js
const statEls = document.querySelectorAll('.stat');
for (const statEl of statEls) {
  assert.equal(statEl.querySelectorAll('.stat-value').length, 1);
}
```

每個 `.stat-value` 元素應包含一個數字。

```js
const statValueEls = document.querySelectorAll('.stat-value');
for (const statEl of statValueEls) {
  assert.match(statEl.textContent, /-?\d+\.*\d*/);
}
```

每個 `.stat-value` 數字應在 `[0,100]` 範圍內（含端點）。

```js
const statValueEls = document.querySelectorAll('.stat-value');
for (const statEl of statValueEls) {
  const statValue = parseFloat(statEl.textContent);
  assert.isAtMost(statValue, 100);
  assert.isAtLeast(statValue, 0);
}
```

恰好有一個 `.stat` 元素包含文字 `Hunger`，不分大小寫。

```js
const matches = Array.from(document.querySelectorAll('.stat')).filter(el => /hunger/i.test(el.textContent));
assert.equal(matches.length, 1);
```

有一個 `.stat` 元素包含文字 `Energy`，不分大小寫。

```js
const matches = Array.from(document.querySelectorAll('.stat')).filter(el => /energy/i.test(el.textContent));
assert.equal(matches.length, 1);
```

有一個 `.stat` 元素包含文字 `Happiness`，不分大小寫。

```js
const matches = Array.from(document.querySelectorAll('.stat')).filter(el => /happiness/i.test(el.textContent));
assert.equal(matches.length, 1);
```

`Hunger` 屬性應從 0 開始。

```js
const hungerStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /hunger/i.test(el.textContent));
const hungerValue = parseFloat(hungerStatEl.querySelector('.stat-value').textContent);
assert.equal(hungerValue, 0);
```

`Energy` 屬性應從 100 開始。

```js
const energyStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /energy/i.test(el.textContent));
const energyValue = parseFloat(energyStatEl.querySelector('.stat-value').textContent);
assert.equal(energyValue, 100);
```

`Happiness` 屬性應從 100 開始。

```js
const happinessStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /happiness/i.test(el.textContent));
const happinessValue = parseFloat(happinessStatEl.querySelector('.stat-value').textContent);
assert.equal(happinessValue, 100);
```

點擊 `Eat` 按鈕應降低 `Hunger` 值並增加 `Energy` 值。若 `Hunger` 為 0 或 `Energy` 為 100，該值不得改變。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#eat-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 0) assert.isBelow(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

點擊 `Play` 按鈕應降低 `Energy` 值並增加 `Happiness` 值。若 `Energy` 為 0 或 `Happiness` 為 100，該值不得改變。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessBefore = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#play-action').click();
_clock.tick(50);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessAfter = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (energyBefore !== 0) assert.isBelow(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
if (happinessBefore !== 100) assert.isAbove(happinessAfter, happinessBefore);
else assert.equal(happinessAfter, happinessBefore);
```

點擊 `Sleep` 按鈕應增加 `Hunger` 值並增加 `Energy` 值。若 `Hunger` 為 100 或 `Energy` 為 100，該值不得改變。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#sleep-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 100) assert.isAbove(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

`Hunger` 屬性在過度閒置時間後應為 100。

```js
await _clock.tickAsync(1_200_000);

const hungerStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /hunger/i.test(el.textContent));
const hungerValue = parseFloat(hungerStatEl.querySelector('.stat-value').textContent);
assert.equal(hungerValue, 100);
```

`Energy` 屬性在過度閒置時間後應為 100。

```js
const energyStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /energy/i.test(el.textContent));
const energyValue = parseFloat(energyStatEl.querySelector('.stat-value').textContent);
assert.equal(energyValue, 100);
```

`Happiness` 屬性在過度閒置時間後應為 0。

```js
const happinessStatEl = Array.from(document.querySelectorAll('.stat')).find(el => /happiness/i.test(el.textContent));
const happinessValue = parseFloat(happinessStatEl.querySelector('.stat-value').textContent);
assert.equal(happinessValue, 0);
```

點擊 `Eat` 按鈕應降低 `Hunger` 值並增加 `Energy` 值。若 `Hunger` 為 0 或 `Energy` 為 100，該值不得改變。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#eat-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 0) assert.isBelow(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

點擊 `Play` 按鈕應降低 `Energy` 值並增加 `Happiness` 值。若 `Energy` 為 0 或 `Happiness` 為 100，該值不得改變。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessBefore = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#play-action').click();
_clock.tick(50);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
const happinessAfter = parseFloat(statEls().find(el => /happiness/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (energyBefore !== 0) assert.isBelow(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
if (happinessBefore !== 100) assert.isAbove(happinessAfter, happinessBefore);
else assert.equal(happinessAfter, happinessBefore);
```

點擊 `Sleep` 按鈕應增加 `Hunger` 值並增加 `Energy` 值。若 `Hunger` 為 100 或 `Energy` 為 100，該值不得改變。

```js
const statEls = () => Array.from(document.querySelectorAll('.stat'));
const hungerBefore = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyBefore = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
document.querySelector('#sleep-action').click();
_clock.tick(50);
const hungerAfter = parseFloat(statEls().find(el => /hunger/i.test(el.textContent)).querySelector('.stat-value').textContent);
const energyAfter = parseFloat(statEls().find(el => /energy/i.test(el.textContent)).querySelector('.stat-value').textContent);
if (hungerBefore !== 100) assert.isAbove(hungerAfter, hungerBefore);
else assert.equal(hungerAfter, hungerBefore);
if (energyBefore !== 100) assert.isAbove(energyAfter, energyBefore);
else assert.equal(energyAfter, energyBefore);
```

應有一個 `PetMood` 列舉。

```js
assert.match(code, /\s*enum PetMood\s*/);
```

`PetMood` 列舉應包含 `HAPPY`。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*HAPPY[\s*\w+,\s*]*}/);
```

`PetMood` 列舉應包含 `EXCITED`。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*EXCITED[\s*\w+,\s*]*}/);
```

`PetMood` 列舉應包含 `CONTENT`。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*CONTENT[\s*\w+,\s*]*}/);
```

`PetMood` 列舉應包含 `SAD`。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*SAD[\s*\w+,\s*]*}/);
```

`PetMood` 列舉應包含 `TIRED`。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*TIRED[\s*\w+,\s*]*}/);
```

`PetMood` 列舉應包含 `SICK`。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*SICK[\s*\w+,\s*]*}/);
```

`PetMood` 列舉應包含 `HUNGRY`。

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*HUNGRY[\s*\w+,\s*]*}/);
```

應有一個 `Record<PetMood, string` 對應每個 `PetMood` 的值。

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>/);
```

應有一個值對應 `PetMood.HAPPY`。

```js
assert.match(code, /\[PetMood\.HAPPY\]\s*:/);
```

應有一個值對應 `PetMood.EXCITED`。

```js
assert.match(code, /\[PetMood\.EXCITED\]\s*:/);
```

應有一個值對應 `PetMood.CONTENT`。

```js
assert.match(code, /\[PetMood\.CONTENT\]\s*:/);
```

應有一個值對應 `PetMood.SAD`。

```js
assert.match(code, /\[PetMood\.SAD\]\s*:/);
```

應有一個值對應 `PetMood.TIRED`。

```js
assert.match(code, /\[PetMood\.TIRED\]\s*:/);
```

應有一個值對應 `PetMood.SICK`。

```js
assert.match(code, /\[PetMood\.SICK\]\s*:/);
```

應有一個值對應 `PetMood.HUNGRY`。

```js
assert.match(code, /\[PetMood\.HUNGRY\]\s*:/);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="utf-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Digital Pet Game</title>

    <link rel="preconnect" href="https://fonts.googleapis.com">
    <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin>
    <link href="https://fonts.googleapis.com/css2?family=Tektur:wght@400..900&display=swap" rel="stylesheet">

        <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.tsx"
    ></script>
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { PetGame } from './index.tsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<PetGame />);
    </script>
  </body>
</html>
```

```css
:root {
  /* Backgrounds */
  --bg-gradient-start: #dbeafe;
  --bg-gradient-mid: #faf5ff;
  --bg-gradient-end: #fce7f3;

  /* Text */
  --text-primary: #1f2937;
  --text-muted: #6b7280;
  --text-muted-2: #9ca3af;

  /* Card / surface */
  --card-bg-translucent: rgba(255, 255, 255, 0.6);
  --card-bg-solid: rgba(255, 255, 255, 0.8);

  /* Accent / brand */
  --accent-purple: #8b5cf6;
  --accent-purple-dark: #7c3aed;

  /* Actions */
  --action-feed: #10b981;
  --action-feed-dark: #059669;
  --action-play: #3b82f6;
  --action-play-dark: #2563eb;
  --action-sleep: #8b5cf6;
  --action-sleep-dark: #7c3aed;
  --action-heal: #ec4899;
  --action-heal-dark: #db2777;

  /* Stats */
  --stat-high: #10b981;
  --stat-medium: #f59e0b;
  --stat-low: #ef4444;

  /* UI greys */
  --ui-bg-light: #e5e7eb;
  --ui-fill: #8b5cf6; /* used for exp fill as accent */

  /* Shadows */
  --shadow-strong: rgba(0, 0, 0, 0.25);
  --shadow-soft: rgba(0, 0, 0, 0.1);
}

html {
  /*color-scheme: light dark;*/
  font-family: system-ui;
  display: flex;
  justify-content: center;
  font-family:
    -apple-system, BlinkMacSystemFont, "Segoe UI", "Roboto", "Oxygen", "Ubuntu",
    "Cantarell", "Fira Sans", "Droid Sans", "Helvetica Neue", sans-serif;
  -webkit-font-smoothing: antialiased;
  -moz-osx-font-smoothing: grayscale;
  background: linear-gradient(135deg, #dbeafe 0%, #faf5ff 50%, #fce7f3 100%);
  min-height: 100vh;
  padding: 1rem;
}

header {
  text-align: center;
  margin-bottom: 2rem;
}

body {
  max-width: 400px;
}

.base-container {
  background: rgba(255, 255, 255, 0.6);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 1.5rem;
  padding: 2rem;
  box-shadow: 0 25px 50px -12px rgba(0, 0, 0, 0.25);
  .pet-name {
    text-align: center;
  }
}

.game-container {
  margin-bottom: 1.5rem;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
}

/* Stats Grid */
.stats-grid {
  display: grid;
  grid-template-columns: repeat(2, 1fr);
  gap: 1rem;
  margin-bottom: 1.5rem;
}

.stat-bar {
  background: rgba(255, 255, 255, 0.8);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 0.75rem;
  padding: 1rem;
  box-shadow: 0 1px 3px 0 rgba(0, 0, 0, 0.1);
  transition: all 0.2s ease;
}

.stat-bar:hover {
  box-shadow: 0 4px 6px -1px rgba(0, 0, 0, 0.1);
}

.stat-header {
  display: flex;
  align-items: center;
  justify-content: space-between;
  margin-bottom: 0.5rem;
}

.stat-label {
  display: flex;
  align-items: center;
  gap: 0.5rem;
}

.stat-icon {
  font-size: 1.25rem;
}

.stat-name {
  font-weight: 500;
  color: #374151;
}

.stat-value {
  font-weight: bold;
  color: #1f2937;
}

.stat-progress {
  width: 100%;
  height: 0.75rem;
  background: #e5e7eb;
  border-radius: 9999px;
  overflow: hidden;
}

.stat-fill {
  height: 100%;
  transition: width 0.5s ease-out;
  border-radius: 9999px;
}

.stat-fill.high {
  background: #10b981;
}

.stat-fill.medium {
  background: #f59e0b;
}

.stat-fill.low {
  background: #ef4444;
}

.pet-buttons {
  width: auto;
  display: flex;
  flex-direction: row;
  gap: 1rem;
}

.pet-buttons .pet-button {
  min-width: 80px;
  background: #c2214a;
  color: white;
  border-radius: 13%;
  border-style: solid;
  border-color: #870b04;
  box-shadow: rgba(120, 25, 23, 0.2) 2px 2px 4px 2px;
  &:active {
    transform: scale(0.95);
    box-shadow: 0 2px 4px rgba(0, 0, 0, 0.15) inset;
  }
  text-align: center;
  cursor: pointer;
  padding: 0.5rem 1rem;
}

.pet-sprite {
  font-size: xxx-large;
  text-align: center;
}

.data-management * {
  margin: 0 0.5rem 0 0;
}

.info-panel {
  text-align: center;
}
.info-panel #hud,
.info-panel .start-questions {
  display: flex;
  align-items: center;
  flex-direction: column;
  gap: 1rem;
}

.info-panel #hud p {
  margin: 0rem;
}

.info-panel .start-questions div * {
  margin: 0 0.5rem 0 0;
}

#pet-fact {
  color: var(--text-muted-2);
  margin: 0 0 0.5rem 0;
}
#set-name-btn {
  display: inline-block;
  padding: 14px 28px;
  font-size: 1.15rem;
  font-weight: 600;
  background: linear-gradient(135deg, #6b8cff, #5ec2ff);
  color: #fff;
  border: none;
  border-radius: 10px;
  box-shadow: 0 6px 18px rgba(92, 132, 255, 0.25);
  cursor: pointer;
  transition:
    transform 0.08s ease,
    box-shadow 0.12s ease,
    opacity 0.12s;
}

#set-name-btn:hover {
  opacity: 0.98;
}

#set-name-btn:active,
#set-name-btn.pressed {
  transform: translateY(2px) scale(0.995);
  box-shadow: 0 3px 8px rgba(92, 132, 255, 0.18);
}

#set-name-btn:focus {
  outline: 3px solid rgba(94, 190, 255, 0.28);
  outline-offset: 2px;
}
```

```tsx
const { useState } = React;

export const PetGame = () => {};
```

# --solutions--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="utf-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Digital Pet Game</title>

    <link rel="preconnect" href="https://fonts.googleapis.com">
    <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin>
    <link href="https://fonts.googleapis.com/css2?family=Tektur:wght@400..900&display=swap" rel="stylesheet">
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.tsx"
    ></script>
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { PetGame } from './index.tsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<PetGame />);
    </script>
  </body>
</html>
```

```css
:root {
  /* Backgrounds */
  --bg-gradient-start: #dbeafe;
  --bg-gradient-mid: #faf5ff;
  --bg-gradient-end: #fce7f3;

  /* Text */
  --text-primary: #1f2937;
  --text-muted: #6b7280;
  --text-muted-2: #9ca3af;

  /* Card / surface */
  --card-bg-translucent: rgba(255, 255, 255, 0.6);
  --card-bg-solid: rgba(255, 255, 255, 0.8);

  /* Accent / brand */
  --accent-purple: #8b5cf6;
  --accent-purple-dark: #7c3aed;

  /* Actions */
  --action-feed: #10b981;
  --action-feed-dark: #059669;
  --action-play: #3b82f6;
  --action-play-dark: #2563eb;
  --action-sleep: #8b5cf6;
  --action-sleep-dark: #7c3aed;
  --action-heal: #ec4899;
  --action-heal-dark: #db2777;

  /* Stats */
  --stat-high: #10b981;
  --stat-medium: #f59e0b;
  --stat-low: #ef4444;

  /* UI greys */
  --ui-bg-light: #e5e7eb;
  --ui-fill: #8b5cf6; /* used for exp fill as accent */

  /* Shadows */
  --shadow-strong: rgba(0, 0, 0, 0.25);
  --shadow-soft: rgba(0, 0, 0, 0.1);
}

html {
  /*color-scheme: light dark;*/
  font-family: system-ui;
  display: flex;
  justify-content: center;
  font-family:
    -apple-system, BlinkMacSystemFont, "Segoe UI", "Roboto", "Oxygen", "Ubuntu",
    "Cantarell", "Fira Sans", "Droid Sans", "Helvetica Neue", sans-serif;
  -webkit-font-smoothing: antialiased;
  -moz-osx-font-smoothing: grayscale;
  background: linear-gradient(135deg, #dbeafe 0%, #faf5ff 50%, #fce7f3 100%);
  min-height: 100vh;
  padding: 1rem;
}

header {
  text-align: center;
  margin-bottom: 2rem;
}

body {
  max-width: 400px;
}

.base-container {
  background: rgba(255, 255, 255, 0.6);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 1.5rem;
  padding: 2rem;
  box-shadow: 0 25px 50px -12px rgba(0, 0, 0, 0.25);
  .pet-name {
    text-align: center;
  }
}

.game-container {
  margin-bottom: 1.5rem;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
}

/* Stats Grid */
.stats-grid {
  display: grid;
  grid-template-columns: repeat(2, 1fr);
  gap: 1rem;
  margin-bottom: 1.5rem;
}

.stat-bar {
  background: rgba(255, 255, 255, 0.8);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 0.75rem;
  padding: 1rem;
  box-shadow: 0 1px 3px 0 rgba(0, 0, 0, 0.1);
  transition: all 0.2s ease;
}

.stat-bar:hover {
  box-shadow: 0 4px 6px -1px rgba(0, 0, 0, 0.1);
}

.stat-header {
  display: flex;
  align-items: center;
  justify-content: space-between;
  margin-bottom: 0.5rem;
}

.stat-label {
  display: flex;
  align-items: center;
  gap: 0.5rem;
}

.stat-icon {
  font-size: 1.25rem;
}

.stat-name {
  font-weight: 500;
  color: #374151;
}

.stat-value {
  font-weight: bold;
  color: #1f2937;
}

.stat-progress {
  width: 100%;
  height: 0.75rem;
  background: #e5e7eb;
  border-radius: 9999px;
  overflow: hidden;
}

.stat-fill {
  height: 100%;
  transition: width 0.5s ease-out;
  border-radius: 9999px;
}

.stat-fill.high {
  background: #10b981;
}

.stat-fill.medium {
  background: #f59e0b;
}

.stat-fill.low {
  background: #ef4444;
}

.pet-buttons {
  width: auto;
  display: flex;
  flex-direction: row;
  gap: 1rem;
}

.pet-buttons .pet-button {
  min-width: 80px;
  background: #c2214a;
  color: white;
  border-radius: 13%;
  border-style: solid;
  border-color: #870b04;
  box-shadow: rgba(120, 25, 23, 0.2) 2px 2px 4px 2px;
  &:active {
    transform: scale(0.95);
    box-shadow: 0 2px 4px rgba(0, 0, 0, 0.15) inset;
  }
  text-align: center;
  cursor: pointer;
  padding: 0.5rem 1rem;
}

.pet-sprite {
  font-size: xxx-large;
  text-align: center;
}

.data-management * {
  margin: 0 0.5rem 0 0;
}

.info-panel {
  text-align: center;
}
.info-panel #hud,
.info-panel .start-questions {
  display: flex;
  align-items: center;
  flex-direction: column;
  gap: 1rem;
}

.info-panel #hud p {
  margin: 0rem;
}

.info-panel .start-questions div * {
  margin: 0 0.5rem 0 0;
}

#pet-fact {
  color: var(--text-muted-2);
  margin: 0 0 0.5rem 0;
}
#set-name-btn {
  display: inline-block;
  padding: 14px 28px;
  font-size: 1.15rem;
  font-weight: 600;
  background: linear-gradient(135deg, #6b8cff, #5ec2ff);
  color: #fff;
  border: none;
  border-radius: 10px;
  box-shadow: 0 6px 18px rgba(92, 132, 255, 0.25);
  cursor: pointer;
  transition:
    transform 0.08s ease,
    box-shadow 0.12s ease,
    opacity 0.12s;
}

#set-name-btn:hover {
  opacity: 0.98;
}

#set-name-btn:active,
#set-name-btn.pressed {
  transform: translateY(2px) scale(0.995);
  box-shadow: 0 3px 8px rgba(92, 132, 255, 0.18);
}

#set-name-btn:focus {
  outline: 3px solid rgba(94, 190, 255, 0.28);
  outline-offset: 2px;
}
```

```tsx
const { useState, useEffect } = React;

interface StatBarProps {
  label: string;
  value: number;
  icon: string;
  reverse?: boolean;
}

function getStatColorClass(value: number): string {
  if (value >= 70) return "high";
  if (value >= 40) return "medium";
  return "low";
}

function StatBar({ label, value, icon, reverse = false }: StatBarProps) {
  const colorClass = getStatColorClass(!reverse ? value : 100 - value);

  return (
    <div className="stat-bar stat">
      <div className="stat-header">
        <div className="stat-label">
          <span className="stat-icon">{icon}</span>
          <span className="stat-name">{label}</span>
        </div>
        <span className="stat-value">{Math.round(value)}%</span>
      </div>

      <div className="stat-progress">
        <div
          className={`stat-fill ${colorClass}`}
          style={{ width: `${value}%` }}
        />
      </div>
    </div>
  );
}

enum PetAction {
  NONE,
  EAT,
  PLAY,
  SLEEP,
}

interface Pet {
  name: string;
  happiness: number;
  hunger: number;
  energy: number;
  species: string;
}

enum PetMood {
  HAPPY,
  EXCITED,
  CONTENT,
  SAD,
  TIRED,
  SICK,
  HUNGRY,
}

export const STAT_DECAY_RATES = {
  hunger: 10,
  happiness: 5,
  energy: 5,
};

const UPDATE_INTERVAL = 30 * 1000; // 30 seconds

export function calculatePetMood(pt: Pet): PetMood {
  const { hunger, happiness, energy } = pt;

  if (hunger > 70) return PetMood.HUNGRY;
  if (energy < 30) return PetMood.TIRED;
  if (happiness < 30) return PetMood.SAD;
  if (happiness > 80 && energy > 70) return PetMood.EXCITED;
  if (happiness > 60) return PetMood.HAPPY;

  return PetMood.CONTENT;
}

const moodEmojiMap: Record<PetMood, string> = {
  [PetMood.HAPPY]: "😺", // grinning cat
  [PetMood.EXCITED]: "😻", // heart-eyes cat
  [PetMood.CONTENT]: "😸", // smiling cat
  [PetMood.SAD]: "😿", // crying cat
  [PetMood.TIRED]: "😽", // kissing cat (closest to sleepy)
  [PetMood.SICK]: "🙀", // weary cat (looks unwell)
  [PetMood.HUNGRY]: "😹", // joy cat (closest to hungry/funny)
};

interface usePetProps {
  isGameStarted: boolean;
  setGameStarted: React.Dispatch<React.SetStateAction<boolean>>;
}
// hook where all pet logic is handled
function usePet({ isGameStarted }: usePetProps) {
  let [pet, setPet] = useState<Pet>({
    name: "",
    happiness: 100,
    hunger: 0,
    energy: 100,
    species: "Cat",
  });
  // Auto-decay stats over time
  useEffect(() => {
    if (!isGameStarted) return;

    const interval = setInterval(() => {
      setPet((pet) => ({
        ...pet,
        happiness: Math.max(pet.happiness - 5, 0),
        hunger: Math.min(pet.hunger + 5, 100),
      }));
    }, UPDATE_INTERVAL);

    return () => clearInterval(interval);
  }, [isGameStarted]);

  function doAction(action: PetAction) {
    switch (action) {
      case PetAction.EAT:
        feedPet();
        break;
      case PetAction.SLEEP:
        restPet();
        break;
      case PetAction.PLAY:
        playWithPet();
        break;
    }
  }

  function feedPet() {
    setPet((pet) => ({
      ...pet,
      hunger: Math.max(pet.hunger - STAT_DECAY_RATES.hunger, 0),
      energy: Math.min(pet.energy + STAT_DECAY_RATES.energy, 100),
    }));
  }

  function playWithPet() {
    setPet((pet) => ({
      ...pet,
      energy: Math.max(pet.energy - STAT_DECAY_RATES.energy, 0),
      happiness: Math.min(pet.happiness + STAT_DECAY_RATES.happiness, 100),
    }));
  }

  function restPet() {
    setPet((pet) => ({
      ...pet,
      hunger: Math.min(pet.hunger + STAT_DECAY_RATES.hunger, 100),
      energy: Math.min(pet.energy + STAT_DECAY_RATES.energy, 100),
    }));
  }

  return {
    pet,
    doAction,
    setName: (name: string) => setPet({ ...pet, name }),
  };
}

export function PetGame() {
  const [isGameStarted, setGameStarted] = useState(false);

  let { pet, doAction, setName } = usePet({
    isGameStarted,
    setGameStarted,
  });

  const [fact, setFact] = useState<string | null>("");

  useEffect(() => {
    fetch("https://cat-facts-api.freecodecamp.rocks/api/catfacts/random").then(
      res => res.json()
    ).then(
      fact => setFact(fact)
    ).catch(err => {
      setFact("Sorry, we're not able to retrieve your cat fact right now!");
    });
  }, []);

  function startGame() {
    const petName = (document.getElementById("pet-name") as HTMLInputElement)
      .value;
    if (!petName) return; //do not let user to start game if no name selected
    setName(petName);
    setGameStarted(true);
  }

  return (
    <main>
      <header>
        <h1>Digital Pet Game</h1>
        <p>Take care of your virtual companion!</p>
      </header>

      {isGameStarted && (
        <section className="base-container game-container">
          <div className="pet-screen">
            <div className="pet-sprite">
              {moodEmojiMap[calculatePetMood(pet)]}
            </div>
            <h2 className="pet-name">{pet.name}</h2>
          </div>
          <div className="pet-buttons">
            <button
              onClick={() => doAction(PetAction.EAT)}
              className="pet-button pet-buttons-left"
              id="eat-action"
            >
              EAT
            </button>
            <button
              onClick={() => doAction(PetAction.PLAY)}
              className="pet-button pet-buttons-center"
              id="play-action"
            >
              PLAY
            </button>
            <button
              onClick={() => doAction(PetAction.SLEEP)}
              className="pet-button pet-buttons-right"
              id="sleep-action"
            >
              SLEEP
            </button>
          </div>
        </section>
      )}

      {isGameStarted && (
        <section className="stats-grid">
          <StatBar label="Hunger" value={pet.hunger} icon="🍽️" reverse />
          <StatBar label="Happiness" value={pet.happiness} icon="😊" />
          <StatBar label="Energy" value={pet.energy} icon="⚡" />
        </section>
      )}

      <section className="base-container info-panel">
        {!isGameStarted ? (
          <form className="start-questions">
            <label htmlFor="pet-name">What is your pet's name?</label>
            <input
              id="pet-name"
              name="pet-name"
              required={true}
              pattern="[A-Za-z0-9]{1,20}"
            />

            <button id="set-name-btn" onClick={startGame}>
              Start Game
            </button>
          </form>
        ) : (
          <div id="hud">
            <p id="pet-species">Species: {pet.species}</p>
            <p id="pet-fact">
              <b>Pet Fact:</b> {fact}
            </p>
          </div>
        )}
      </section>
    </main>
  );
}
```
