---
id: 587d7dbb367417b2b2512bac
title: Ondoa nafasi zisizohitajika mwanzoni na mwishoni
challengeType: 1
forumTopicId: 301362
dashedName: remove-whitespace-from-start-and-end
---

# --description--

Wakati mwingine herufi za nafasi zisizohitajika karibu na mfuatano wa herufi huwa zipo lakini hazihitajiki. Utaratibu wa kawaida wa kushughulikia mfuatano wa herufi ni kuondoa nafasi hizo mwanzoni na mwishoni mwa mfuatano.

# --instructions--

Andika usemi wa kawaida na tumia njia zinazofaa za mfuatano wa herufi kuondoa nafasi mwanzoni na mwishoni mwa mfuatano wa herufi.

**Kumbuka:** Njia ya `String.prototype.trim()` ingeweza kufanya kazi hapa, lakini utahitaji kukamilisha changamoto hii ukitumia usemi wa kawaida.

# --hints--

`result` inapaswa kuwa sawa na mfuatano wa herufi `Hello, World!`

```js
assert(result === 'Hello, World!');
```

Suluhisho lako halipaswi kutumia njia ya `String.prototype.trim()`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?trim/));
```

Kigezo `result` hakipaswi kuwekwa moja kwa moja kuwa mfuatano wa herufi

```js
assert(!__helpers.removeJSComments(code).match(/result\s*=\s*["'`].*?["'`]/));
```

Thamani ya kigezo `hello` haipaswi kubadilishwa.

```js
assert(hello === '   Hello, World!  ');
```

# --seed--

## --seed-contents--

```js
let hello = "   Hello, World!  ";
let wsRegex = /change/; // Change this line
let result = hello; // Change this line
```

# --solutions--

```js
let hello = "   Hello, World!  ";
let wsRegex = /^(\s+)(.+[^\s])(\s+)$/;
let result = hello.replace(wsRegex, '$2');
```
