---
id: 587d7db3367417b2b2512b8f
title: Linganisha mfuatano wa herufi halisi
challengeType: 1
forumTopicId: 301355
dashedName: match-literal-strings
---

# --description--

Katika changamoto iliyopita, ulitafuta neno `Hello` kwa kutumia usemi wa kawaida `/Hello/`. Usemi huo wa kawaida ulitafuta mlingano halisi wa mfuatano wa herufi `Hello`. Hapa kuna mfano mwingine wa kutafuta mlingano halisi wa mfuatano wa herufi `Kevin`:

```js
let testStr = "Hello, my name is Kevin.";
let testRegex = /Kevin/;
testRegex.test(testStr);
```

Wito huu wa `test` utarudisha `true`.

Aina nyingine yoyote ya `Kevin` haitalingana. Kwa mfano, usemi wa kawaida `/Kevin/` hautalingana na `kevin` au `KEVIN`.

```js
let wrongRegex = /kevin/;
wrongRegex.test(testStr);
```

Wito huu wa `test` utarudisha `false`.

Changamoto zijazo zitaonyesha jinsi ya kulinganisha pia aina hizo nyingine.

# --instructions--

Kamilisha usemi wa kawaida `waldoRegex` ili kupata `"Waldo"` katika mfuatano wa herufi `waldoIsHiding` kwa mlingano halisi.

# --hints--

Usemi wako wa kawaida `waldoRegex` unapaswa kupata mfuatano wa herufi `Waldo`

```js
waldoRegex.lastIndex = 0;
assert(waldoRegex.test(waldoIsHiding));
```

Usemi wako wa kawaida `waldoRegex` haupaswi kutafuta kitu kingine chochote.

```js
waldoRegex.lastIndex = 0;
assert(!waldoRegex.test('Somewhere is hiding in this text.'));
```

Unapaswa kufanya mlingano halisi wa mfuatano wa herufi kwa usemi wako wa kawaida.

```js
assert(!/\/.*\/i/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /search/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```

# --solutions--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /Waldo/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```
