---
id: 587d7db8367417b2b2512ba1
title: Linganisha herufi zote zisizo za nambari
challengeType: 1
forumTopicId: 301347
dashedName: match-all-non-numbers
---

# --description--

Changamoto ya mwisho ilionyesha jinsi ya kutafuta nambari kwa kutumia njia fupi `\d` na herufi ndogo `d`. Unaweza pia kutafuta herufi zisizo za nambari kwa kutumia njia fupi inayofanana inayotumia herufi kubwa `D` badala yake.

Njia fupi ya kutafuta herufi zisizo za nambari ni `\D`. Hii ni sawa na darasa la vitu `[^0-9]`, ambalo linatafuta herufi moja ambayo si nambari kati ya sifuri na tisa.

# --instructions--

Tumia darasa la vitu la njia fupi kwa herufi zisizo za nambari `\D` kuhesabu ni herufi ngapi zisizo za nambari zipo katika majina ya filamu.

# --hints--

Usemi wako wa kawaida unapaswa kutumia njia fupi ya herufi ili kulinganisha herufi zisizo za nambari

```js
assert(/\\D/.test(noNumRegex.source));
```

Usemi wako wa kawaida unapaswa kutumia bendera ya global.

```js
assert(noNumRegex.global);
```

Usemi wako wa kawaida unapaswa kutopata herufi zisizo za nambari katika mfuatano `9`.

```js
assert('9'.match(noNumRegex) == null);
```

Usemi wako wa kawaida unapaswa kupata herufi 6 zisizo za nambari katika mfuatano `Catch 22`.

```js
assert('Catch 22'.match(noNumRegex).length == 6);
```

Usemi wako wa kawaida unapaswa kupata herufi 11 zisizo za nambari katika mfuatano `101 Dalmatians`.

```js
assert('101 Dalmatians'.match(noNumRegex).length == 11);
```

Usemi wako wa kawaida unapaswa kupata herufi 15 zisizo za nambari katika mfuatano `One, Two, Three`.

```js
assert('One, Two, Three'.match(noNumRegex).length == 15);
```

Usemi wako wa kawaida unapaswa kupata herufi 12 zisizo za nambari katika mfuatano `21 Jump Street`.

```js
assert('21 Jump Street'.match(noNumRegex).length == 12);
```

Usemi wako wa kawaida unapaswa kupata herufi 17 zisizo za nambari katika mfuatano `2001: A Space Odyssey`.

```js
assert('2001: A Space Odyssey'.match(noNumRegex).length == 17);
```

# --seed--

## --seed-contents--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /change/; // Change this line
let result = movieName.match(noNumRegex).length;
```

# --solutions--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /\D/g; // Change this line
let result = movieName.match(noNumRegex).length;
```
