---
id: 587d7db9367417b2b2512ba7
title: Eleza idadi halisi ya mechi
challengeType: 1
forumTopicId: 301365
dashedName: specify-exact-number-of-matches
---

# --description--

Unaweza kueleza idadi ya chini na ya juu ya mifumo kwa kutumia viashiria vya wingi kwa mabano ya curly. Wakati mwingine unataka tu idadi maalum ya mechi.

Ili kueleza idadi fulani ya mifumo, weka nambari hiyo moja kati ya mabano ya curly.

Kwa mfano, ili kupatana na neno `hah` tu lenye herufi `a` `3` mara, usemi wako wa kawaida utakuwa `/ha{3}h/`.

```js
let A4 = "haaaah";
let A3 = "haaah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleHA = /ha{3}h/;
multipleHA.test(A4);
multipleHA.test(A3);
multipleHA.test(A100);
```

Kwa mpangilio, wito tatu za `test` zitarudisha `false`, `true`, na `false`.

# --instructions--

Badilisha usemi wa kawaida `timRegex` ili upatane na neno `Timber` tu wakati lina herufi nne za `m`.

# --hints--

Usemi wako wa kawaida unapaswa kutumia mabano ya curly.

```js
assert(timRegex.source.match(/{.*?}/).length > 0);
```

Usemi wako wa kawaida haupaswi kupatana na mfuatano wa herufi `Timber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timber'));
```

Usemi wako wa kawaida haupaswi kupatana na mfuatano wa herufi `Timmber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmber'));
```

Usemi wako wa kawaida haupaswi kupatana na mfuatano wa herufi `Timmmber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmmber'));
```

Usemi wako wa kawaida unapaswa kupatana na mfuatano wa herufi `Timmmmber`

```js
timRegex.lastIndex = 0;
assert(timRegex.test('Timmmmber'));
```

Usemi wako wa kawaida haupaswi kupatana na mfuatano wa herufi `Timber` wenye herufi 30 za `m` ndani yake.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Ti' + 'm'.repeat(30) + 'ber'));
```

# --seed--

## --seed-contents--

```js
let timStr = "Timmmmber";
let timRegex = /change/; // Change this line
let result = timRegex.test(timStr);
```

# --solutions--

```js
let timStr = "Timmmmber";
let timRegex = /Tim{4}ber/; // Change this line
let result = timRegex.test(timStr);
```
