---
id: 587d7db7367417b2b2512b9e
title: Linganisha mifumo ya mfuatano wa herufi inayomalizika
challengeType: 1
forumTopicId: 301352
dashedName: match-ending-string-patterns
---

# --description--

Katika changamoto iliyopita, ulijifunza kutumia ishara ya caret kutafuta mifumo mwanzoni mwa mfuatano wa herufi. Pia kuna njia ya kutafuta mifumo mwishoni mwa mfuatano wa herufi.

Unaweza kutafuta mwishoni mwa mfuatano wa herufi kwa kutumia ishara ya alama ya dola `$` mwishoni mwa usemi wa kawaida.

```js
let theEnding = "This is a never ending story";
let storyRegex = /story$/;
storyRegex.test(theEnding);
let noEnding = "Sometimes a story will have to end";
storyRegex.test(noEnding);
```

Wito wa kwanza wa `test` utarudisha `true`, wakati wito wa pili utarudisha `false`.

# --instructions--

Tumia ishara ya tagi ya nanga (`$`) kulinganisha mfuatano wa herufi `caboose` mwishoni mwa mfuatano wa herufi `caboose`.

# --hints--

Unapaswa kutafuta `caboose` kwa kutumia alama ya dola `$` kama tagi ya nanga katika usemi wako wa kawaida.

```js
assert(lastRegex.source == 'caboose$');
```

Usemi wako wa kawaida haupaswi kutumia bendera yoyote.

```js
assert(lastRegex.flags == '');
```

Unapaswa kulinganisha `caboose` mwishoni mwa mfuatano wa herufi `The last car on a train is the caboose`

```js
lastRegex.lastIndex = 0;
assert(lastRegex.test('The last car on a train is the caboose'));
```

# --seed--

## --seed-contents--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /change/; // Change this line
let result = lastRegex.test(caboose);
```

# --solutions--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /caboose$/; // Change this line
let result = lastRegex.test(caboose);
```
