---
id: 587d7db7367417b2b2512b9f
title: Linganisha herufi na nambari zote
challengeType: 1
forumTopicId: 301346
dashedName: match-all-letters-and-numbers
---

# --description--

Kwa kutumia madarasa ya herufi, uliweza kutafuta herufi zote za alfabeti kwa `[a-z]`. Darasa hili la herufi ni la kawaida kiasi kwamba kuna njia fupi ya kulifanya, ingawa pia linajumuisha herufi chache za ziada.

Darasa la herufi linalokaribia zaidi katika JavaScript kulinganisha alfabeti ni `\w`. Njia hii fupi ni sawa na `[A-Za-z0-9_]`. Darasa hili la herufi linalinganisha herufi kubwa na ndogo pamoja na nambari. Kumbuka, darasa hili la herufi pia linajumuisha alama ya mstari chini (`_`).

```js
let longHand = /[A-Za-z0-9_]+/;
let shortHand = /\w+/;
let numbers = "42";
let varNames = "important_var";
longHand.test(numbers);
shortHand.test(numbers);
longHand.test(varNames);
shortHand.test(varNames);
```

Wito wote wanne wa `test` wangerejesha `true`.

Madarasa haya ya herufi ya njia fupi pia hujulikana kama <dfn>madarasa ya herufi ya njia fupi</dfn>.

# --instructions--

Tumia darasa la herufi la njia fupi `\w` kuhesabu idadi ya herufi na nambari katika nukuu na mfuatano wa herufi mbalimbali.

# --hints--

Usemi wako wa kawaida unapaswa kutumia bendera ya global.

```js
assert(alphabetRegexV2.global);
```

Usemi wako wa kawaida unapaswa kutumia darasa la herufi la njia fupi `\w` kulinganisha herufi zote ambazo ni herufi na nambari.

```js
assert(/\\w/.test(alphabetRegexV2.source));
```

Usemi wako wa kawaida unapaswa kupata herufi na nambari 31 katika mfuatano wa herufi `The five boxing wizards jump quickly.`

```js
assert(
  'The five boxing wizards jump quickly.'.match(alphabetRegexV2).length === 31
);
```

Usemi wako wa kawaida unapaswa kupata herufi na nambari 32 katika mfuatano wa herufi `Pack my box with five dozen liquor jugs.`

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(alphabetRegexV2).length ===
    32
);
```

Usemi wako wa kawaida unapaswa kupata herufi na nambari 30 katika mfuatano wa herufi `How vexingly quick daft zebras jump!`

```js
assert(
  'How vexingly quick daft zebras jump!'.match(alphabetRegexV2).length === 30
);
```

Usemi wako wa kawaida unapaswa kupata herufi na nambari 36 katika mfuatano wa herufi `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.`

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(alphabetRegexV2)
    .length === 36
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /change/; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /\w/g; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```
