---
id: 587d7db5367417b2b2512b94
title: Linganisha chochote na alama ya wildcard ya kipindi
challengeType: 1
forumTopicId: 301348
dashedName: match-anything-with-wildcard-period
---

# --description--

Wakati mwingine hutajui (au huna haja ya kujua) herufi halisi katika mifumo yako. Kufikiria maneno yote yanayolingana, sema, makosa ya tahajia kutachukua muda mrefu. Kwa bahati nzuri, unaweza kuokoa muda kwa kutumia herufi ya wildcard: `.`

Herufi ya wildcard `.` italingana na herufi moja yoyote. Wildcard pia huitwa `dot` na `period`. Unaweza kutumia herufi ya wildcard kama herufi nyingine yoyote katika usemi wa kawaida. Kwa mfano, kama ungetaka kulinganisha `hug`, `huh`, `hut`, na `hum`, unaweza kutumia usemi wa kawaida `/hu./` kulinganisha maneno yote manne.

```js
let humStr = "I'll hum a song";
let hugStr = "Bear hug";
let huRegex = /hu./;
huRegex.test(humStr);
huRegex.test(hugStr);
```

Wito zote mbili za `test` zingerudisha `true`.

# --instructions--

Kamilisha usemi wa kawaida `unRegex` ili ulingane na mfuatano wa herufi `run`, `sun`, `fun`, `pun`, `nun`, na `bun`. Usemi wako wa kawaida unapaswa kutumia herufi ya wildcard.

# --hints--

Unapaswa kutumia njia ya `.test()`.

```js
assert(__helpers.removeJSComments(code).match(/\.test\(.*\)/));
```

Unapaswa kutumia herufi ya wildcard katika usemi wako wa kawaida `unRegex`

```js
assert(/\./.test(unRegex.source));
```

Usemi wako wa kawaida `unRegex` unapaswa kulinganisha `run` katika mfuatano wa herufi `Let us go on a run.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Let us go on a run.'));
```

Usemi wako wa kawaida `unRegex` unapaswa kulinganisha `sun` katika mfuatano wa herufi `The sun is out today.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('The sun is out today.'));
```

Usemi wako wa kawaida `unRegex` unapaswa kulinganisha `fun` katika mfuatano wa herufi `Coding is a lot of fun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Coding is a lot of fun.'));
```

Usemi wako wa kawaida `unRegex` unapaswa kulinganisha `pun` katika mfuatano wa herufi `Seven days without a pun makes one weak.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Seven days without a pun makes one weak.'));
```

Usemi wako wa kawaida `unRegex` unapaswa kulinganisha `nun` katika mfuatano wa herufi `One takes a vow to be a nun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('One takes a vow to be a nun.'));
```

Usemi wako wa kawaida `unRegex` unapaswa kulinganisha `bun` katika mfuatano wa herufi `She got fired from the hot dog stand for putting her hair in a bun.`

```js
unRegex.lastIndex = 0;
assert(
  unRegex.test(
    'She got fired from the hot dog stand for putting her hair in a bun.'
  )
);
```

Usemi wako wa kawaida `unRegex` haupaswi kulinganisha mfuatano wa herufi `There is a bug in my code.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('There is a bug in my code.'));
```

Usemi wako wa kawaida `unRegex` haupaswi kulinganisha mfuatano wa herufi `Catch me if you can.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('Catch me if you can.'));
```

# --seed--

## --seed-contents--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /change/; // Change this line
let result = unRegex.test(exampleStr);
```

# --solutions--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /.un/; // Change this line
let result = unRegex.test(exampleStr);
```
