---
id: 587d7db6367417b2b2512b9b
title: Tafuta herufi kwa kutumia ulinganifu wa mvivu
challengeType: 1
forumTopicId: 301341
dashedName: find-characters-with-lazy-matching
---

# --description--

Katika usemi wa kawaida, ulinganifu wa <dfn>mvivu</dfn> hupata sehemu ndefu zaidi inayowezekana ya mfuatano wa herufi inayolingana na muundo wa usemi wa kawaida na kuirudisha kama ulinganifu. Mbali na huo ni ulinganifu wa <dfn>mvivu kidogo</dfn>, ambao hupata sehemu ndogo kabisa inayowezekana ya mfuatano wa herufi inayokidhi muundo wa usemi wa kawaida.

Unaweza kutumia usemi wa kawaida `/t[a-z]*i/` kwa mfuatano wa herufi `"titanic"`. Usemi huu wa kawaida ni mfano unaoanza na `t`, unaisha na `i`, na una herufi fulani katikati.

Usemi wa kawaida kwa chaguo-msingi huwa mvivu, kwa hivyo ulinganifu utarudisha `["titani"]`. Hupata sehemu kubwa zaidi ya mfuatano wa herufi inayoweza kuendana na muundo.

Hata hivyo, unaweza kutumia tabia ya `?` kubadilisha kuwa ulinganifu wa mvivu kidogo. `"titanic"` ukilinganishwa na usemi wa kawaida uliobadilishwa wa `/t[a-z]*?i/` hurejesha `["ti"]`.

**Kumbuka:** Kusahihisha HTML kwa kutumia usemi wa kawaida haipendekezwi, lakini kutafuta muundo wa mfuatano wa herufi wa HTML kwa usemi wa kawaida ni sawa kabisa.

# --instructions--

Rekebisha usemi wa kawaida `/<.*>/` ili urudishe lebo ya HTML `<h1>` na si maandishi ya `"<h1>Winter is coming</h1>"`. Kumbuka kuwa tabia ya jokeri `.` katika usemi wa kawaida inalingana na herufi yoyote.

# --hints--

Kigezo `result` kinapaswa kuwa safu ya data yenye `<h1>` ndani yake

```js
assert(result[0] == '<h1>');
```

`myRegex` inapaswa kutumia ulinganifu wa mvivu kidogo

```js
assert(/[^\\][\*\+\?]\?/.test(myRegex));
```

`myRegex` haipaswi kujumuisha mfuatano wa herufi `h1`

```js
assert(!myRegex.source.match('h1'));
```

# --seed--

## --seed-contents--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*>/; // Change this line
let result = text.match(myRegex);
```

# --solutions--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*?>/; // Change this line
let result = text.match(myRegex);
```
