---
id: 587d7db4367417b2b2512b93
title: Tafuta Zaidi ya Mechi ya Kwanza
challengeType: 1
forumTopicId: 301342
dashedName: find-more-than-the-first-match
---

# --description--

Hadi sasa, umeweza tu kutoa au kutafuta usemi wa kawaida mara moja.

```js
let testStr = "Repeat, Repeat, Repeat";
let ourRegex = /Repeat/;
testStr.match(ourRegex);
```

Hapa `match` itarudisha `["Repeat"]`.

Ili kutafuta au kutoa usemi wa kawaida zaidi ya mara moja, unaweza kutumia bendera ya utafutaji wa kimataifa: `g`.

```js
let repeatRegex = /Repeat/g;
testStr.match(repeatRegex);
```

Na hapa `match` inarudisha thamani `["Repeat", "Repeat", "Repeat"]`

# --instructions--

Ukichukua usemi wa kawaida `starRegex`, tafuta na utoe maneno yote mawili ya `Twinkle` kutoka kwenye mfuatano wa herufi `twinkleStar`.

**Kumbuka**  
Unaweza kuwa na bendera nyingi kwenye usemi wako wa kawaida kama `/search/gi`

# --hints--

Usemi wako wa kawaida `starRegex` unapaswa kutumia bendera ya kimataifa `g`

```js
assert(starRegex.flags.match(/g/).length == 1);
```

Usemi wako wa kawaida `starRegex` unapaswa kutumia bendera isiyozingatia herufi kubwa au ndogo `i`

```js
assert(starRegex.flags.match(/i/).length == 1);
```

Mechi yako inapaswa kufanana na matukio yote mawili ya neno `Twinkle`

```js
assert(
  result.sort().join() ==
    twinkleStar
      .match(/twinkle/gi)
      .sort()
      .join()
);
```

Mechi yako `result` inapaswa kuwa na vipengele viwili ndani yake.

```js
assert(result.length == 2);
```

# --seed--

## --seed-contents--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /change/; // Change this line
let result = twinkleStar; // Change this line
```

# --solutions--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /twinkle/gi;
let result = twinkleStar.match(starRegex);
```
