---
id: 587d7db4367417b2b2512b91
title: Puuza Herufi Ndogo au Kubwa Wakati wa Kufananisha
challengeType: 1
forumTopicId: 301344
dashedName: ignore-case-while-matching
---

# --description--

Hadi sasa, umeangalia usemi wa kawaida kufanya ulinganifu wa mfuatano wa herufi kwa herufi halisi. Lakini wakati mwingine, unaweza pia kutaka kufananisha tofauti za herufi ndogo na kubwa.

Herufi (au wakati mwingine huitwa aina ya herufi) ni tofauti kati ya herufi kubwa na herufi ndogo. Mifano ya herufi kubwa ni `A`, `B`, na `C`. Mifano ya herufi ndogo ni `a`, `b`, na `c`.

Unaweza kufananisha herufi zote mbili kwa kutumia kile kinachoitwa bendera (flag). Kuna bendera nyingine lakini hapa utazingatia bendera inayopuuza herufi ndogo au kubwa - bendera ya `i`. Unaweza kuitumia kwa kuiambatanisha kwenye usemi wa kawaida. Mfano wa kutumia bendera hii ni `/ignorecase/i`. Usemi huu wa kawaida unaweza kufananisha mfuatano wa herufi `ignorecase`, `igNoreCase`, na `IgnoreCase`.

# --instructions--

Andika usemi wa kawaida `fccRegex` ili kufananisha `freeCodeCamp`, bila kujali aina ya herufi. Usemi wako wa kawaida haupaswi kufananisha vifupisho au mabadiliko yenye nafasi.

# --hints--

Usemi wako wa kawaida unapaswa kufananisha mfuatano wa herufi `freeCodeCamp`

```js
fccRegex.lastIndex = 0; 
assert(fccRegex.test('freeCodeCamp'));
```

Usemi wako wa kawaida unapaswa kufananisha mfuatano wa herufi `FreeCodeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodeCamp'));
```

Usemi wako wa kawaida unapaswa kufananisha mfuatano wa herufi `FreecodeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreecodeCamp'));
```

Usemi wako wa kawaida unapaswa kufananisha mfuatano wa herufi `FreeCodecamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodecamp'));
```

Usemi wako wa kawaida haupaswi kufananisha mfuatano wa herufi `Free Code Camp`

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('Free Code Camp'));
```

Usemi wako wa kawaida unapaswa kufananisha mfuatano wa herufi `FreeCOdeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCOdeCamp'));
```

Usemi wako wa kawaida haupaswi kufananisha mfuatano wa herufi `FCC`

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('FCC'));
```

Usemi wako wa kawaida unapaswa kufananisha mfuatano wa herufi `FrEeCoDeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCoDeCamp'));
```

Usemi wako wa kawaida unapaswa kufananisha mfuatano wa herufi `FrEeCodECamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCodECamp'));
```

Usemi wako wa kawaida unapaswa kufananisha mfuatano wa herufi `FReeCodeCAmp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FReeCodeCAmp'));
```

# --seed--

## --seed-contents--

```js
let myString = "freeCodeCamp";
let fccRegex = /change/; // Change this line
let result = fccRegex.test(myString);
```

# --solutions--

```js
let myString = "freeCodeCamp";
let fccRegex = /freecodecamp/i; // Change this line
let result = fccRegex.test(myString);
```
