---
id: 587d7dba367417b2b2512ba8
title: Angalia kwa yote au hakuna
challengeType: 1
forumTopicId: 301338
dashedName: check-for-all-or-none
---

# --description--

Wakati mwingine mifumo unayotaka kutafuta inaweza kuwa na sehemu ambazo zinaweza kuwepo au kutokuwepo. Hata hivyo, inaweza kuwa muhimu kuangalia uwepo wake hata hivyo.

Unaweza kubainisha uwezekano wa kuwepo kwa kipengele kwa alama ya swali, `?`. Hii inakagua uwepo wa sifuri au moja ya kipengele kilichotangulia. Unaweza kufikiria alama hii kama kusema kipengele kilichotangulia ni hiari.

Kwa mfano, kuna tofauti ndogo kati ya Kiingereza cha Marekani na Kiingereza cha Uingereza na unaweza kutumia alama ya swali kuendana na tahajia zote mbili.

```js
let american = "color";
let british = "colour";
let rainbowRegex= /colou?r/;
rainbowRegex.test(american);
rainbowRegex.test(british);
```

Matumizi yote ya `test` njia yangerudisha `true`.

# --instructions--

Badilisha usemi wa kawaida `favRegex` ili uendane na toleo la Kiingereza cha Marekani (`favorite`) na toleo la Kiingereza cha Uingereza (`favourite`).

# --hints--

Usemi wako wa kawaida unapaswa kutumia alama ya hiari, `?`.

```js
favRegex.lastIndex = 0;
assert(favRegex.source.match(/\?/).length > 0);
```

Usemi wako wa kawaida unapaswa kuendana na mfuatano wa herufi `favorite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favorite'));
```

Usemi wako wa kawaida unapaswa kuendana na mfuatano wa herufi `favourite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favourite'));
```

Usemi wako wa kawaida haupaswi kuendana na mfuatano wa herufi `fav`

```js
favRegex.lastIndex = 0;
assert(!favRegex.test('fav'));
```

# --seed--

## --seed-contents--

```js
let favWord = "favorite";
let favRegex = /change/; // Change this line
let result = favRegex.test(favWord);
```

# --solutions--

```js
let favWord = "favorite";
let favRegex = /favou?r/;
let result = favRegex.test(favWord);
```
