---
id: 587d7db8367417b2b2512ba3
title: Linganisha nafasi tupu
challengeType: 1
forumTopicId: 301359
dashedName: match-whitespace
---

# --description--

Changamoto zilizopita zimehusu kulinganisha herufi za alfabeti na nambari. Pia unaweza kulinganisha nafasi tupu au nafasi kati ya herufi.

Unaweza kutafuta nafasi tupu kwa kutumia `\s`, ambayo ni herufi ndogo `s`. Mfano huu haulinganishi tu nafasi tupu, bali pia hujumuisha alama za kurudisha gari, tab, form feed, na mistari mipya. Unaweza kuifikiria kama darasa la vitu `[ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let spaceRegex = /\s/g;
whiteSpace.match(spaceRegex);
```

Wito huu wa `match` utarudisha `[" ", " "]`.
# --instructions--

Badilisha usemi wa kawaida `countWhiteSpace` ili utafute herufi nyingi za nafasi tupu katika mfuatano wa herufi.

# --hints--

Usemi wako wa kawaida unapaswa kutumia bendera ya global.

```js
assert(countWhiteSpace.global);
```

Usemi wako wa kawaida unapaswa kutumia kifupi cha herufi `\s` ili kulinganisha herufi zote za nafasi tupu.

```js
assert(/\\s/.test(countWhiteSpace.source));
```

Usemi wako wa kawaida unapaswa kupata nafasi nane katika mfuatano wa herufi `Men are from Mars and women are from Venus.`

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countWhiteSpace).length ==
    8
);
```

Usemi wako wa kawaida unapaswa kupata nafasi tatu katika mfuatano wa herufi `Space: the final frontier.`

```js
assert('Space: the final frontier.'.match(countWhiteSpace).length == 3);
```

Usemi wako wa kawaida unapaswa kutopata nafasi yoyote katika mfuatano wa herufi `MindYourPersonalSpace`

```js
assert('MindYourPersonalSpace'.match(countWhiteSpace) == null);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /change/; // Change this line
let result = sample.match(countWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /\s/g;
let result = sample.match(countWhiteSpace);
```
