---
id: 587d7db6367417b2b2512b98
title: Linganisha herufi moja isiyoainishwa
challengeType: 1
forumTopicId: 301358
dashedName: match-single-characters-not-specified
---

# --description--

Hadi sasa, umeunda seti ya herufi unayotaka kulinganisha, lakini pia unaweza kuunda seti ya herufi ambazo hutaki kulinganisha. Aina hizi za seti za herufi huitwa <dfn>seti za herufi zilizokataliwa</dfn>.

Ili kuunda seti ya herufi zilizokataliwa, weka ishara ya caret (`^`) baada ya mabano ya kufungua na kabla ya herufi ambazo hutaki kulinganisha.

Kwa mfano, `/[^aeiou]/gi` inalinganisha herufi zote ambazo si vokali. Kumbuka kuwa herufi kama `.`, `!`, `[`, `@`, `/` na nafasi nyeupe zinalinganishwa - seti ya herufi zilizokataliwa ya vokali inatenga tu herufi za vokali.

# --instructions--

Unda usemi wa kawaida mmoja unaolinganishwa na herufi zote ambazo si nambari wala vokali. Kumbuka kujumuisha bendera zinazofaa katika usemi wa kawaida.

# --hints--

Usemi wako wa kawaida `myRegex` unapaswa kulinganisha vitu 9.

```js
assert(result.length == 9);
```

Usemi wako wa kawaida `myRegex` unapaswa kutumia bendera ya global.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

Usemi wako wa kawaida `myRegex` unapaswa kutumia bendera isiyozingatia herufi kubwa na ndogo.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "3 blind mice.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "3 blind mice.";
let myRegex = /[^0-9aeiou]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
