---
id: 587d7db7367417b2b2512b9d
title: Linganisha mifumo ya mfuatano wa herufi mwanzoni
challengeType: 1
forumTopicId: 301349
dashedName: match-beginning-string-patterns
---

# --description--

Changamoto zilizopita zilionyesha kwamba usemi wa kawaida unaweza kutumika kutafuta mechi nyingi. Pia hutumika kutafuta mifumo katika nafasi maalum ndani ya mfuatano wa herufi.

Katika changamoto ya awali, ulitumia ishara ya caret (`^`) ndani ya seti ya herufi kuunda seti ya herufi isiyokubalika kwa namna ya `[^thingsThatWillNotBeMatched]`. Nje ya seti ya herufi, ishara ya caret hutumika kutafuta mifumo mwanzoni mwa mfuatano wa herufi.

```js
let firstString = "Ricky is first and can be found.";
let firstRegex = /^Ricky/;
firstRegex.test(firstString);
let notFirst = "You can't find Ricky now.";
firstRegex.test(notFirst);
```

Wito wa kwanza wa `test` utarudisha `true`, wakati wito wa pili utarudisha `false`.

# --instructions--

Tumia ishara ya caret katika usemi wa kawaida kutafuta `Cal` tu mwanzoni mwa mfuatano wa herufi `rickyAndCal`.

# --hints--

Usemi wako wa kawaida unapaswa kutafuta mfuatano wa herufi `Cal` ukiwa na herufi kubwa.

```js
assert(calRegex.source == '^Cal');
```

Usemi wako wa kawaida haupaswi kutumia bendera yoyote.

```js
assert(calRegex.flags == '');
```

Usemi wako wa kawaida unapaswa kulinganisha mfuatano wa herufi `Cal` mwanzoni mwa mfuatano.

```js
calRegex.lastIndex = 0;
assert(calRegex.test('Cal and Ricky both like racing.'));
```

Usemi wako wa kawaida haupaswi kulinganisha mfuatano wa herufi `Cal` katikati ya mfuatano.

```js
calRegex.lastIndex = 0;
assert(!calRegex.test('Ricky and Cal both like racing.'));
```

# --seed--

## --seed-contents--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /change/; // Change this line
let result = calRegex.test(rickyAndCal);
```

# --solutions--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /^Cal/; // Change this line
let result = calRegex.test(rickyAndCal);
```
