---
id: 587d7db9367417b2b2512ba4
title: Linganisha herufi zisizo za nafasi
challengeType: 1
forumTopicId: 18210
dashedName: match-non-whitespace-characters
---

# --description--

Umejifunza kuhusu utafuta nafasi kwa kutumia `\s`, kwa herufi ndogo `s`. Pia unaweza kutafuta kila kitu isipokuwa nafasi.

Tafuta herufi zisizo za nafasi kwa kutumia `\S`, ambayo ni herufi kubwa `s`. Mfano huu hautalinganisha nafasi, kurudi kwa gari, tab, fomu ya chakavu, na herufi mpya za mstari. Unaweza kuifikiria kama darasa la vitu `[^ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let nonSpaceRegex = /\S/g;
whiteSpace.match(nonSpaceRegex).length;
```

Thamani inayorejeshwa na njia ya `.length` itakuwa `32`.

# --instructions--

Badilisha usemi wa kawaida `countNonWhiteSpace` ili utafute herufi nyingi zisizo za nafasi katika mfuatano wa herufi.

# --hints--

Usemi wako wa kawaida unapaswa kutumia bendera ya ulimwengu.

```js
assert(countNonWhiteSpace.global);
```

Usemi wako wa kawaida unapaswa kutumia herufi fupi `\S` kulinganisha herufi zote zisizo za nafasi.

```js
assert(/\\S/.test(countNonWhiteSpace.source));
```

Usemi wako wa kawaida unapaswa kupata herufi 35 zisizo za nafasi katika mfuatano wa herufi `Men are from Mars and women are from Venus.`

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countNonWhiteSpace)
    .length == 35
);
```

Usemi wako wa kawaida unapaswa kupata herufi 23 zisizo za nafasi katika mfuatano wa herufi `Space: the final frontier.`

```js
assert('Space: the final frontier.'.match(countNonWhiteSpace).length == 23);
```

Usemi wako wa kawaida unapaswa kupata herufi 21 zisizo za nafasi katika mfuatano wa herufi `MindYourPersonalSpace`

```js
assert('MindYourPersonalSpace'.match(countNonWhiteSpace).length == 21);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /change/; // Change this line
let result = sample.match(countNonWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /\S/g; // Change this line
let result = sample.match(countNonWhiteSpace);
```
