---
id: 587d7db4367417b2b2512b90
title: Linganisha mfuatano wa herufi halisi na uwezekano tofauti
challengeType: 1
forumTopicId: 301345
dashedName: match-a-literal-string-with-different-possibilities
---

# --description--

Kwa kutumia usemi wa kawaida kama `/coding/`, unaweza kutafuta muundo `coding` katika mfuatano mwingine wa herufi.

Hii ni yenye nguvu kutafuta mfuatano mmoja, lakini ina kikomo kwa muundo mmoja tu. Unaweza kutafuta miundo mingi kwa kutumia operator `alternation` au `OR`: `|`.

Operator hii inalinganisha miundo kabla au baada yake. Kwa mfano, kama ungetaka kulinganisha mfuatano wa herufi `yes` au `no`, usemi wa kawaida unaotakiwa ni `/yes|no/`.

Unaweza pia kutafuta zaidi ya miundo miwili tu. Hii unaweza kufanya kwa kuongeza miundo zaidi na operator `OR` kati yao, kama `/yes|no|maybe/`.

# --instructions--

Kamilisha usemi wa kawaida `petRegex` ili ulinganishe wanyama wa kipenzi `dog`, `cat`, `bird`, au `fish`.

# --hints--

Usemi wako wa kawaida `petRegex` unapaswa kurudisha `true` kwa mfuatano wa herufi `John has a pet dog.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('John has a pet dog.'));
```

Usemi wako wa kawaida `petRegex` unapaswa kurudisha `false` kwa mfuatano wa herufi `Emma has a pet rock.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Emma has a pet rock.'));
```

Usemi wako wa kawaida `petRegex` unapaswa kurudisha `true` kwa mfuatano wa herufi `Emma has a pet bird.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Emma has a pet bird.'));
```

Usemi wako wa kawaida `petRegex` unapaswa kurudisha `true` kwa mfuatano wa herufi `Liz has a pet cat.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Liz has a pet cat.'));
```

Usemi wako wa kawaida `petRegex` unapaswa kurudisha `false` kwa mfuatano wa herufi `Kara has a pet dolphin.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Kara has a pet dolphin.'));
```

Usemi wako wa kawaida `petRegex` unapaswa kurudisha `true` kwa mfuatano wa herufi `Alice has a pet fish.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Alice has a pet fish.'));
```

Usemi wako wa kawaida `petRegex` unapaswa kurudisha `false` kwa mfuatano wa herufi `Jimmy has a pet computer.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Jimmy has a pet computer.'));
```

# --seed--

## --seed-contents--

```js
let petString = "James has a pet cat.";
let petRegex = /change/; // Change this line
let result = petRegex.test(petString);
```

# --solutions--

```js
let petString = "James has a pet cat.";
let petRegex = /dog|cat|bird|fish/; // Change this line
let result = petRegex.test(petString);
```
