---
id: 587d7db6367417b2b2512b99
title: Linganisha herufi zinazojirudia mara moja au zaidi
challengeType: 1
forumTopicId: 301350
dashedName: match-characters-that-occur-one-or-more-times
---

# --description--

Wakati mwingine, unahitaji kulinganisha herufi (au kundi la herufi) zinazojitokeza mara moja au zaidi mfululizo. Hii ina maana kwamba hutokea angalau mara moja, na inaweza kurudiwa.

Unaweza kutumia alama ya `+` kuangalia kama hiyo ndiyo hali. Kumbuka, herufi au muundo lazima uwepo mfululizo. Yaani, herufi lazima irudi moja baada ya nyingine.

Kwa mfano, `/a+/g` ingepata mlingano mmoja katika `abc` na kurudisha `["a"]`. Kwa sababu ya `+`, pia ingepata mlingano mmoja katika `aabc` na kurudisha `["aa"]`.

Ikiwa badala yake ingekuwa ikikagua mfuatano wa herufi `abab`, ingepata mlingano miwili na kurudisha `["a", "a"]` kwa sababu herufi za `a` haziko mfululizo - kuna `b` kati yao. Mwishowe, kwa kuwa hakuna `a` katika mfuatano wa herufi `bcd`, haingepata mlingano.

# --instructions--

Unataka kupata mlingano wakati herufi `s` inapotokea mara moja au zaidi katika `Mississippi`. Andika usemi wa kawaida unaotumia alama ya `+`.

# --hints--

Usemi wako wa kawaida `myRegex` unapaswa kutumia alama ya `+` kulinganisha herufi `s` mara moja au zaidi.

```js
assert(/\+/.test(myRegex.source));
```

Usemi wako wa kawaida `myRegex` unapaswa kulinganisha vitu 2.

```js
assert(result.length == 2);
```

Kigezo cha `result` kinapaswa kuwa safu ya data yenye mlingano miwili ya `ss`

```js
assert(result[0] == 'ss' && result[1] == 'ss');
```

# --seed--

## --seed-contents--

```js
let difficultSpelling = "Mississippi";
let myRegex = /change/; // Change this line
let result = difficultSpelling.match(myRegex);
```

# --solutions--

```js
let difficultSpelling = "Mississippi";
let myRegex = /s+/g; // Change this line
let result = difficultSpelling.match(myRegex);
```
