---
id: 587d7db3367417b2b2512b8e
title: Kutumia njia ya test
challengeType: 1
forumTopicId: 301369
dashedName: using-the-test-method
---

# --description--

Usemi wa kawaida hutumika katika lugha za programu kulinganisha sehemu za mfuatano wa herufi. Unaunda mifumo ya kusaidia kufanya kulinganisha hiyo.

Ikiwa unataka kupata neno `the` katika mfuatano wa herufi `The dog chased the cat`, unaweza kutumia usemi wa kawaida ufuatao: `/the/`. Angalia kuwa alama za nukuu hazihitajiki ndani ya usemi wa kawaida.

JavaScript ina njia nyingi za kutumia usemi wa kawaida. Njia moja ya kujaribu usemi wa kawaida ni kutumia njia ya `.test()`. Njia ya `.test()` inachukua usemi wa kawaida, inautumia kwenye mfuatano wa herufi (ambao umewekwa ndani ya mabano ya kawaida), na inarudisha `true` au `false` ikiwa mfumo wako unapata kitu au la.

```js
let testStr = "freeCodeCamp";
let testRegex = /Code/;
testRegex.test(testStr);
```

Njia ya `test` hapa inarudisha `true`.

# --instructions--

Tumia usemi wa kawaida `myRegex` kwenye mfuatano wa herufi `myString` kwa kutumia njia ya `.test()`.

# --hints--

Unapaswa kutumia `.test()` kujaribu usemi wa kawaida.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

Matokeo yako yanapaswa kurudisha `true`.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex; // Change this line
```

# --solutions--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex.test(myString); // Change this line
```
