---
id: 587d7db6367417b2b2512b9a
title: Linganisha herufi zinazojitokeza mara sifuri au zaidi
challengeType: 1
forumTopicId: 301351
dashedName: match-characters-that-occur-zero-or-more-times
---

# --description--

Changamoto ya mwisho ilitumia alama ya plus `+` kutafuta herufi zinazojitokeza mara moja au zaidi. Pia kuna chaguo linalolingana na herufi zinazojitokeza mara sifuri au zaidi.

Herufi inayofanya hivi ni nyota au asterisk: `*`.

```js
let soccerWord = "gooooooooal!";
let gPhrase = "gut feeling";
let oPhrase = "over the moon";
let goRegex = /go*/;
soccerWord.match(goRegex);
gPhrase.match(goRegex);
oPhrase.match(goRegex);
```

Kwa mpangilio, wito tatu `match` zingerudisha thamani `["goooooooo"]`, `["g"]`, na `null`.

# --instructions--

Tengeneza usemi wa kawaida `chewieRegex` unaotumia herufi `*` kulinganisha herufi kubwa `A` inayofuatiwa mara moja na herufi ndogo `a` mara sifuri au zaidi katika `chewieQuote`. Usemi wako wa kawaida hauhitaji bendera wala madarasa ya herufi, na haupaswi kulinganisha nukuu nyingine yoyote.

# --hints--

Usemi wako wa kawaida `chewieRegex` unapaswa kutumia herufi `*` kulinganisha herufi `a` mara sifuri au zaidi.

```js
assert(/\*/.test(chewieRegex.source));
```

Usemi wako wa kawaida unapaswa kulinganisha mfuatano wa herufi `A` katika `chewieQuote`.

```js
assert(result[0][0] === 'A');
```

Usemi wako wa kawaida unapaswa kulinganisha mfuatano wa herufi `Aaaaaaaaaaaaaaaa` katika `chewieQuote`.

```js
assert(result[0] === 'Aaaaaaaaaaaaaaaa');
```

Usemi wako wa kawaida `chewieRegex` unapaswa kulinganisha herufi 16 katika `chewieQuote`.

```js
assert(result[0].length === 16);
```

Usemi wako wa kawaida haupaswi kulinganisha herufi yoyote katika mfuatano wa herufi `He made a fair move. Screaming about it can't help you.`

```js
assert(
  !"He made a fair move. Screaming about it can't help you.".match(chewieRegex)
);
```

Usemi wako wa kawaida haupaswi kulinganisha herufi yoyote katika mfuatano wa herufi `Let him have it. It's not wise to upset a Wookiee.`

```js
assert(
  !"Let him have it. It's not wise to upset a Wookiee.".match(chewieRegex)
);
```

# --seed--

## --seed-contents--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
// Only change code below this line
let chewieRegex = /change/; // Change this line
// Only change code above this line

let result = chewieQuote.match(chewieRegex);
```

# --solutions--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
let chewieRegex = /Aa*/;
let result = chewieQuote.match(chewieRegex);
```
