---
id: 587d7db5367417b2b2512b96
title: Linganisha herufi za alfabeti
challengeType: 1
forumTopicId: 301354
dashedName: match-letters-of-the-alphabet
---

# --description--

Umeona jinsi unavyoweza kutumia <dfn>seti za herufi</dfn> kubainisha kundi la herufi za kulinganisha, lakini hiyo ni kuandika sana wakati unahitaji kulinganisha anuwai kubwa ya herufi (kwa mfano, kila herufi katika alfabeti). Kwa bahati nzuri, kuna kipengele kilichojengwa ndani kinachofanya hii kuwa fupi na rahisi.

Ndani ya seti ya herufi, unaweza kufafanua anuwai ya herufi za kulinganisha kwa kutumia herufi ya msalaba: `-`.

Kwa mfano, ili kulinganisha herufi ndogo `a` hadi `e` utatumia `[a-e]`.

```js
let catStr = "cat";
let batStr = "bat";
let matStr = "mat";
let bgRegex = /[a-e]at/;
catStr.match(bgRegex);
batStr.match(bgRegex);
matStr.match(bgRegex);
```

Kwa mpangilio, wito tatu za `match` zitarudisha thamani za `["cat"]`, `["bat"]`, na `null`.

# --instructions--

Linganisha herufi zote katika mfuatano wa herufi `quoteSample`.

**Kumbuka**: Hakikisha unalinganisha herufi kubwa na herufi ndogo zote.

# --hints--

Usemi wako wa kawaida `alphabetRegex` unapaswa kulinganisha vitu 35.

```js
assert(result.length == 35);
```

Usemi wako wa kawaida `alphabetRegex` unapaswa kutumia bendera ya global.

```js
assert(alphabetRegex.flags.match(/g/).length == 1);
```

Usemi wako wa kawaida `alphabetRegex` unapaswa kutumia bendera isiyojali herufi kubwa au ndogo.

```js
assert(alphabetRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /change/; // Change this line
let result = alphabetRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /[a-z]/gi; // Change this line
let result = quoteSample.match(alphabetRegex); // Change this line
```
