---
id: 587d7db4367417b2b2512b92
title: Toa mechi zinazolingana
challengeType: 1
forumTopicId: 301340
dashedName: extract-matches
---

# --description--

Hadi sasa, umekuwa ukikagua tu kama muundo upo au haupo ndani ya mfuatano wa herufi. Pia unaweza kutoa mechi halisi ulizozipata kwa kutumia njia ya `.match()`.

Ili kutumia njia ya `.match()`, tumia njia hiyo kwenye mfuatano wa herufi na uingize usemi wa kawaida ndani ya mabano ya kawaida.

Hapa kuna mfano:

```js
"Hello, World!".match(/Hello/);
let ourStr = "Regular expressions";
let ourRegex = /expressions/;
ourStr.match(ourRegex);
```

Hapa `match` ya kwanza itarudisha `["Hello"]` na ya pili itarudisha `["expressions"]`.

Kumbuka kuwa sintaksia ya `.match` ni "kinyume" cha njia ya `.test` ambayo umekuwa ukitumia hadi sasa:

```js
'string'.match(/regex/);
/regex/.test('string');
```

# --instructions--

Tumia njia ya `.match()` kutoa mfuatano wa herufi `coding`.

# --hints--

`result` inapaswa kuwa na mfuatano wa herufi `coding`

```js
assert(result.join() === 'coding');
```

Usemi wako wa kawaida `codingRegex` unapaswa kutafuta mfuatano wa herufi `coding`

```js
assert(codingRegex.source === 'coding');
```

Unapaswa kutumia njia ya `.match()`.

```js
assert(__helpers.removeJSComments(code).match(/\.match\(.*\)/));
```

# --seed--

## --seed-contents--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /change/; // Change this line
let result = extractStr; // Change this line
```

# --solutions--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /coding/; // Change this line
let result = extractStr.match(codingRegex); // Change this line
```
