---
id: 587d7db5367417b2b2512b97
title: Linganisha nambari na herufi za alfabeti
challengeType: 1
forumTopicId: 301356
dashedName: match-numbers-and-letters-of-the-alphabet
---

# --description--

Kutumia alama ya msalaba (`-`) kuoanisha safu ya herufi hakuzuiliwi kwa herufi tu. Pia hufanya kazi kuoanisha safu ya nambari.

Kwa mfano, `/[0-5]/` huoanisha nambari yoyote kati ya `0` na `5`, ikijumuisha `0` na `5`.

Pia, inawezekana kuunganisha safu ya herufi na nambari katika seti moja ya tabia.

```js
let jennyStr = "Jenny8675309";
let myRegex = /[a-z0-9]/ig;
jennyStr.match(myRegex);
```

# --instructions--

Tengeneza usemi wa kawaida unaoendana na safu ya herufi kati ya `h` na `s`, na safu ya nambari kati ya `2` na `6`. Kumbuka kujumuisha vigezo vinavyofaa katika usemi wa kawaida.

# --hints--

Usemi wako wa kawaida `myRegex` unapaswa kuoanisha vitu 17.

```js
assert(result.length == 17);
```

Usemi wako wa kawaida `myRegex` unapaswa kutumia kigezo cha global.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

Usemi wako wa kawaida `myRegex` unapaswa kutumia kigezo cha kutotambua herufi kubwa au ndogo.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /[h-s2-6]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
