---
id: 66a9c7a4521d2b1b1ec6dcf0
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

이제 표에 학생 두 명을 더 추가할 차례입니다.

이전 단계와 같은 패턴을 따라 두 번째 학생 표 행을 추가하세요. 표 데이터 요소에는 다음 데이터를 사용하세요:

- 성: `Doe`
- 이름: `Samantha`
- 학년: `92`

세 번째 학생 표 행에는 다음 데이터를 표 데이터 요소로 사용하세요:

- 성: `Rodriguez`
- 이름: `Marcus`
- 학년: `88`

# --hints--

두 번째 `tr`가 `tbody` 요소 안에 있어야 합니다.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

두 번째 `td` 요소 안에 텍스트 `Doe`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

두 번째 `td` 요소 안에 텍스트 `Samantha`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

두 번째 `td` 요소 안에 텍스트 `92`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

세 번째 `tr`가 `tbody` 요소 안에 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

세 번째 `td` 요소 안에 텍스트 `Rodriguez`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

세 번째 `td` 요소 안에 텍스트 `Marcus`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

세 번째 `td` 요소 안에 텍스트 `88`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
