---
id: 66a9c14e3b34c719e34bae20
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

학생 데이터를 표에 추가하려면 표 행과 표 데이터 요소를 사용해야 합니다.

<dfn>table data</dfn> 요소인 `td`는 표에서 셀을 만드는 데 사용됩니다.

다음은 스포츠 선수 표에 `td` 요소를 사용하는 예입니다:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

표 본문 요소 안에 표 행 요소인 `tr`를 추가하고, 그 안에 세 개의 표 데이터 요소인 `td`를 넣으세요.

첫 번째 표 데이터 요소에는 `Davis`의 성이 들어가야 합니다.

두 번째 표 데이터 요소에는 `Alex`의 이름이 들어가야 합니다.

세 번째 표 데이터 요소에는 `54`의 학년이 들어가야 합니다.

# --hints--

`tr` 요소 안에 `tbody` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

`td` 요소 안에 텍스트 `Davis`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

`td` 요소 안에 텍스트 `Alex`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

`td` 요소 안에 텍스트 `54`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

`td` 요소 안에 세 개의 `tr` 요소가 모두 있어야 합니다.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
