---
id: 66a9bcf00f13a418368a272e
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

표 헤드 요소는 <dfn>표 행</dfn> 요소인 `tr`로 구성되며, 이 요소는 <dfn>표 헤더 셀</dfn> 요소인 `th`을 포함합니다.

다음은 스포츠 표에 `tr`와 `th` 요소를 사용한 예시입니다:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

`thead` 요소 안에 `tr` 요소를 추가하세요.

`tr` 요소 안에 세 개의 `th` 요소를 추가하세요.

첫 번째 `th` 요소에는 `Last Name` 텍스트가 들어가야 합니다. 두 번째 `th` 요소에는 `First Name` 텍스트가 들어가야 합니다. 세 번째 `th` 요소에는 `Grade` 텍스트가 들어가야 합니다.

# --hints--

열린 `tr` 태그가 있어야 합니다.

```js
assert.match(code, /<tr>/i);
```

닫힌 `tr` 태그가 있어야 합니다.

```js
assert.match(code, /<\/tr>/i);
```

`tr` 요소는 `thead` 요소 안에 있어야 합니다.

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

첫 번째 `th` 요소에는 `Last Name` 텍스트가 있어야 합니다.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

두 번째 `th` 요소에는 `First Name` 텍스트가 있어야 합니다.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

세 번째 `th` 요소에는 `Grade` 텍스트가 있어야 합니다.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

모든 `th` 요소는 `tr` 요소 안에 있어야 합니다.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
