---
id: 66a9ca099e3ecb1be10a2696
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

이제 표에 학생 두 명을 더 추가할 차례입니다.

이전 단계와 같은 패턴을 따라 네 번째 학생 표 행을 추가하세요. 표 데이터 요소에는 다음 데이터를 사용하세요:

- 성: `Thompson`
- 이름: `Jane`
- 학년: `77`

다섯 번째 학생 표 행에는 다음 데이터를 사용하세요:

- 성: `Williams`
- 이름: `Natalie`
- 학년: `83`

# --hints--

네 번째 `tr`가 `tbody` 요소 안에 있어야 합니다.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

네 번째 `td` 요소 안에 텍스트 `Thompson`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

네 번째 `td` 요소 안에 텍스트 `Jane`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

네 번째 `td` 요소 안에 텍스트 `77`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

다섯 번째 `tr`가 `tbody` 요소 안에 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

다섯 번째 `td` 요소 안에 텍스트 `Williams`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

다섯 번째 `td` 요소 안에 텍스트 `Natalie`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

다섯 번째 `td` 요소 안에 텍스트 `83`가 있는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
