---
id: 66a9bb8578459a175432e7d0
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

표의 첫 번째 섹션에서는 학생의 이름, 성, 그리고 기말 시험 점수에 해당하는 열 레이블을 나타내는 헤더 내용을 그룹화해야 합니다.

<dfn>table head</dfn> 요소인 `thead`는 표에서 헤더 내용을 그룹화하는 데 사용됩니다.

다음은 `thead` 요소를 사용한 예시입니다:

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

`caption` 요소 아래에 table head 요소를 추가하세요.

# --hints--

열린 `thead` 태그가 있어야 합니다.

```js
assert.match(code, /<thead>/i);
```

닫힌 `thead` 태그가 있어야 합니다.

```js
assert.match(code, /<\/thead>/i);
```

`thead` 요소는 `table` 요소 안에 있어야 합니다.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
