---
id: 66a9cb26ec6bd41cf6c82bc5
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

테이블에 추가할 마지막 섹션은 <dfn>table foot</dfn> 요소인 `tfoot`입니다. table foot 요소는 테이블에 있는 모든 학생의 평균 점수를 표시하는 데 사용됩니다.

테이블에 `tfoot` 요소를 추가하세요.

`tfoot` 요소 안에 `tr` 요소를 추가하세요.

`tr` 요소 안에 두 개의 `td` 요소를 추가하세요.

첫 번째 `td` 요소에는 텍스트 `Average Grade`가 들어가야 합니다. 두 번째 `td` 요소에는 텍스트 `78.8`가 들어가야 합니다.

# --hints--

`tfoot` 여는 태그가 있어야 합니다.

```js
assert.match(code, /<tfoot>/);
```

`tfoot` 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/tfoot>/);
```

`tr` 요소 안에 `tfoot` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

첫 번째 `td` 요소에는 텍스트 `Average Grade`가 들어가야 합니다. `tr` 요소 안에 있는지 확인하세요.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

두 번째 `td` 요소에는 텍스트 `78.8`가 들어가야 합니다. `tr` 요소 안에 있는지 확인하세요.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

`tr` 요소 안에는 두 개의 `td` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
