---
id: 66a9b8f14b963916a3baa732
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`table`에 캡션을 추가하려면 <dfn>table caption</dfn> 요소를 사용할 수 있습니다.

다음은 `caption` 요소를 사용한 예입니다:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

`table` 요소 안에 `caption` 요소를 중첩하고 텍스트 `Calculus Final Exam Grades`를 넣으세요.

# --hints--

열린 `caption` 태그가 있어야 합니다.

```js
assert.match(code, /<caption>/i);
```

닫힌 `caption` 태그가 있어야 합니다.

```js
assert.match(code, /<\/caption>/i);
```

`caption` 요소가 `table` 요소 안에 있어야 합니다.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

`caption`에 텍스트 `Calculus Final Exam Grades`가 있어야 합니다.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--
      
    --fcc-editable-region--      
    </table>
  </body>
</html>
```
