---
id: 66a9c006e4f58819396605b3
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

head 섹션을 완료했으니 이제 <dfn>table body</dfn>, `tbody`를 추가할 차례입니다. table body는 모든 학생 이름과 성적을 나타냅니다.

테이블에 table body 요소를 추가하세요.

# --hints--

여는 `tbody` 태그가 있어야 합니다.

```js
assert.match(code, /<tbody>/i);
```

닫는 `tbody` 태그가 있어야 합니다.

```js
assert.match(code, /<\/tbody>/i);
```

여는 `tbody` 태그가 닫는 `table` 태그보다 앞에 와야 합니다.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
