---
id: 66a9b48a211a73155621d0df
title: 1단계
challengeType: 0
dashedName: step-1
---

# --description--

이 워크숍에서는 학생 그룹의 기말 시험 표를 만들어 HTML 표를 다루는 연습을 합니다.

이전 수업에서 표 형식 데이터를 나타내기 위해 `table` 요소를 다루는 방법을 배웠습니다.

`body` 요소 안에 `table` 요소를 중첩하세요.

# --hints--

열린 `table` 태그가 있어야 합니다.

```js
assert.match(code, /<table>/i);
```

닫힌 `table` 태그가 있어야 합니다.

```js
assert.match(code, /<\/table>/i);
```

`table` 요소는 `body` 요소 안에 있어야 합니다.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
