---
id: aaa48de84e1ecc7c742e1124
title: Перевірка паліндрома
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16004
dashedName: palindrome-checker
---

# --description--

Поверніть `true`, якщо заданий рядок є паліндромом. Інакше поверніть `false`.

<dfn>Паліндром</dfn> — це слово або речення, яке читається однаково вперед і назад, ігноруючи пунктуацію, регістр і пробіли.

**Примітка:** Для перевірки паліндромів потрібно видалити **всі неалфавітно-цифрові символи** (пунктуацію, пробіли та символи) і привести все до одного регістру (нижнього або верхнього).

Ми передамо рядки з різними форматами, наприклад `racecar`, `RaceCar` та `race CAR` серед інших.

Також ми передамо рядки зі спеціальними символами, такими як `2A3*3a2`, `2A3 3a2` та `2_A3*3#A2`.

# --hints--

`palindrome("eye")` має повертати булевий тип.

```js
assert(typeof palindrome('eye') === 'boolean');
```

`palindrome("eye")` має повертати `true`.

```js
assert(palindrome('eye') === true);
```

`palindrome("_eye")` має повертати `true`.

```js
assert(palindrome('_eye') === true);
```

`palindrome("race car")` має повертати `true`.

```js
assert(palindrome('race car') === true);
```

`palindrome("not a palindrome")` має повертати `false`.

```js
assert(palindrome('not a palindrome') === false);
```

`palindrome("A man, a plan, a canal. Panama")` має повертати `true`.

```js
assert(palindrome('A man, a plan, a canal. Panama') === true);
```

`palindrome("never odd or even")` має повертати `true`.

```js
assert(palindrome('never odd or even') === true);
```

`palindrome("nope")` має повертати `false`.

```js
assert(palindrome('nope') === false);
```

`palindrome("almostomla")` має повертати `false`.

```js
assert(palindrome('almostomla') === false);
```

`palindrome("My age is 0, 0 si ega ym.")` має повертати `true`.

```js
assert(palindrome('My age is 0, 0 si ega ym.') === true);
```

`palindrome("1 eye for of 1 eye.")` має повертати `false`.

```js
assert(palindrome('1 eye for of 1 eye.') === false);
```

`palindrome("0_0 (: /-\ :) 0-0")` має повертати `true`.

```js
assert(palindrome('0_0 (: /- :) 0-0') === true);
```

`palindrome("five|\_/|four")` має повертати `false`.

```js
assert(palindrome('five|_/|four') === false);
```

# --seed--

## --seed-contents--

```js
function palindrome(str) {
  return true;
}

palindrome("eye");
```

# --solutions--

```js
function palindrome(str) {
  var string = str.toLowerCase().split(/[^A-Za-z0-9]/gi).join('');
  var aux = string.split('');
  if (aux.join('') === aux.reverse().join('')){
    return true;
  }

  return false;
}
```
