---
id: 587d78ad367417b2b2512afa
title: 플렉스랩 속성으로 행 또는 열 감싸기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cQv9ZtG'
forumTopicId: 301114
dashedName: use-the-flex-wrap-property-to-wrap-a-row-or-column
---

# --description--

CSS 플렉스박스에는 플렉스 컨테이너를 여러 행(또는 열)으로 나누는 기능이 있습니다. 기본적으로 플렉스 컨테이너는 모든 플렉스 항목을 한데 모읍니다. 예를 들어, 행은 모두 한 줄에 배치됩니다.

하지만 `flex-wrap` 속성을 사용하면 CSS가 항목을 감싸도록 지시합니다. 즉, 여분의 항목이 새 행이나 열로 이동합니다. 감싸기가 발생하는 중단점은 항목 크기와 컨테이너 크기에 따라 달라집니다.

CSS에는 감싸기 방향에 대한 옵션도 있습니다:

<ul><li><code>nowrap</code>: 기본 설정이며, 항목을 감싸지 않습니다.</li><li><code>wrap</code>: 행일 경우 위에서 아래로, 열일 경우 왼쪽에서 오른쪽으로 여러 줄에 걸쳐 항목을 감쌉니다.</li><li><code>wrap-reverse</code>: 행일 경우 아래에서 위로, 열일 경우 오른쪽에서 왼쪽으로 여러 줄에 걸쳐 항목을 감쌉니다.</li></ul>

# --instructions--

현재 레이아웃에는 한 행에 너무 많은 상자가 있습니다. `flex-wrap` 요소에 `#box-container` CSS 속성을 추가하고, 값으로 `wrap`을 지정하세요.

# --hints--

`#box-container` 요소에 `flex-wrap` 속성을 `wrap` 값으로 설정해야 합니다.

```js
const boxContainer = document.querySelector('#box-container');
const flexWrap = window.getComputedStyle(boxContainer)['flex-wrap'];
assert.strictEqual(flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;
    flex-wrap: wrap;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```
