---
id: 587d78ae367417b2b2512afe
title: flex 단축 속성 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cbpW2tE'
forumTopicId: 301112
dashedName: use-the-flex-shorthand-property
---

# --description--

여러 flex 속성을 한 번에 설정할 수 있는 단축키가 있습니다. `flex-grow`, `flex-shrink`, 그리고 `flex-basis` 속성은 모두 `flex` 속성을 사용해 함께 설정할 수 있습니다.

예를 들어, `flex: 1 0 10px;`는 아이템을 `flex-grow: 1;`, `flex-shrink: 0;`, 그리고 `flex-basis: 10px;`로 설정합니다.

기본 속성 설정은 `flex: 0 1 auto;`입니다.

# --instructions--

`flex` CSS 속성을 `#box-1`과 `#box-2` 모두에 추가하세요. `#box-1`에 값을 주어 `flex-grow`이 `2`, `flex-shrink`가 `2`, 그리고 `flex-basis`이 `150px`가 되게 하세요. `#box-2`에 값을 주어 `flex-grow`이 `1`, `flex-shrink`가 `1`, 그리고 `flex-basis`가 `150px`가 되게 하세요.

이 값들은 컨테이너가 300px보다 클 때 `#box-1`이 `#box-2`보다 두 배 빠른 속도로 여분 공간을 채우도록 하고, 컨테이너가 300px보다 작을 때 `#box-2`보다 두 배 빠른 속도로 줄어들게 합니다. 300px은 두 박스의 `flex-basis` 값의 합입니다.

# --hints--

`#box-1` 요소는 `flex` 속성을 `2 2 150px` 값으로 설정해야 합니다.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];

assert.equal(flexGrow, '2');
assert.equal(flexShrink, '2');
assert.equal(flexBasis, '150px');
```

`#box-2` 요소는 `flex` 속성을 `1 1 150px` 값으로 설정해야 합니다.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];

assert.equal(flexGrow, '1');
assert.equal(flexShrink, '1');
assert.equal(flexBasis, '150px');
```

코드는 `flex` 속성을 `#box-1`과 `#box-2`에 사용해야 합니다.

```js
assert.lengthOf(code.match(/flex:\s*?\d\s+?\d\s+?150px;/g), 2);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    flex: 2 2 150px;
    height: 200px;
  }

  #box-2 {
    background-color: orangered;
    flex: 1 1 150px;
    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
