---
id: 587d78ad367417b2b2512af8
title: align-items 속성을 사용해 요소 정렬하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c8aggtk'
forumTopicId: 301101
dashedName: align-elements-using-the-align-items-property
---

# --description--

`align-items` 속성은 `justify-content`와 비슷합니다. `justify-content` 속성이 플렉스 항목을 주축을 따라 정렬했다는 점을 기억하세요. 행의 경우 주축은 수평선이고 열의 경우 수직선입니다.

플렉스 컨테이너에는 주축과 반대 방향인 **교차 축**도 있습니다. 행에서는 교차 축이 수직이고 열에서는 교차 축이 수평입니다.

CSS는 플렉스 항목을 교차 축을 따라 정렬하는 `align-items` 속성을 제공합니다. 행의 경우, 이 속성은 컨테이너 내에서 전체 행의 항목을 위아래로 어떻게 밀어 넣을지 CSS에 알려줍니다. 열의 경우, 컨테이너 내에서 모든 항목을 좌우로 어떻게 밀어 넣을지 알려줍니다.

`align-items`에 사용할 수 있는 다양한 값은 다음과 같습니다:

<ul><li><code>flex-start</code>: 항목을 플렉스 컨테이너의 시작 부분에 정렬합니다. 행에서는 항목을 컨테이너 상단에 정렬하고, 열에서는 컨테이너 왼쪽에 정렬합니다.</li><li><code>flex-end</code>: 항목을 플렉스 컨테이너의 끝 부분에 정렬합니다. 행에서는 항목을 컨테이너 하단에 정렬하고, 열에서는 컨테이너 오른쪽에 정렬합니다.</li><li><code>center</code>: 항목을 중앙에 정렬합니다. 행에서는 항목을 수직으로 중앙에 정렬(항목 위아래에 동일한 공간)하고, 열에서는 항목을 수평으로 중앙에 정렬(항목 좌우에 동일한 공간)합니다.</li><li><code>stretch</code>: 항목을 늘려 플렉스 컨테이너를 채웁니다. 예를 들어, 행 항목은 플렉스 컨테이너를 위에서 아래로 채우도록 늘어납니다. <code>align-items</code> 값이 지정되지 않으면 기본값입니다.</li><li><code>baseline</code>: 항목을 기준선에 맞춰 정렬합니다. 기준선은 텍스트 개념으로, 글자가 놓이는 선이라고 생각하세요.</li></ul>

# --instructions--

예제를 통해 이 속성이 어떻게 작동하는지 보여줍니다. `align-items` 요소에 CSS 속성 `#box-container`를 추가하고, 값으로 `center`을 지정하세요.

**보너스**  
코드 편집기에서 `align-items` 속성의 다른 옵션을 시도해 차이를 확인해 보세요. 단, `center` 값만 이 도전 과제를 통과할 수 있습니다.

# --hints--

`#box-container` 요소에 `align-items` 속성을 `center` 값으로 설정해야 합니다.

```js
const boxContainer = document.querySelector('#box-container');
const alignment = window.getComputedStyle(boxContainer)['align-items'];
assert.strictEqual(alignment, 'center');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;
    align-items: center;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```
