---
id: 587d78ab367417b2b2512af2
title: 플렉스 방향 속성으로 행 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cBEkbfJ'
forumTopicId: 301110
dashedName: use-the-flex-direction-property-to-make-a-row
---

# --description--

요소에 `display: flex`를 추가하면 플렉스 컨테이너가 됩니다. 이렇게 하면 해당 요소의 자식들을 행이나 열로 정렬할 수 있습니다. 부모 항목에 `flex-direction` 속성을 추가하고 값을 row 또는 column으로 설정하면 됩니다. 행을 만들면 자식들이 가로로 정렬되고, 열을 만들면 자식들이 세로로 정렬됩니다.

`flex-direction`에 사용할 수 있는 다른 옵션은 `row-reverse`와 `column-reverse`입니다.

**참고:** `flex-direction` 속성의 기본값은 `row`입니다.

# --instructions--

`flex-direction` 요소에 CSS 속성 `#box-container`를 추가하고 값을 `row-reverse`로 설정하세요.

# --hints--

`#box-container` 요소에 `flex-direction` 속성을 `row-reverse`로 설정해야 합니다.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'row-reverse');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: row-reverse;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
