---
id: 587d78ac367417b2b2512af6
title: justify-content 속성을 사용해 요소 정렬하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c43gnHm'
forumTopicId: 301102
dashedName: align-elements-using-the-justify-content-property
---

# --description--

플렉스 컨테이너 안의 플렉스 항목들이 컨테이너의 모든 공간을 채우지 못하는 경우가 있습니다. 플렉스 항목들을 특정 방식으로 정렬하고 간격을 조절하고 싶을 때가 많습니다. 다행히 `justify-content` 속성에는 이를 위한 여러 옵션이 있습니다. 하지만 그 옵션들을 살펴보기 전에 알아야 할 중요한 용어가 있습니다.

<a href="https://www.freecodecamp.org/news/flexbox-the-ultimate-css-flex-cheatsheet/" target="_blank" rel="noopener noreferrer nofollow">플렉스박스 속성에 대한 자세한 내용은 여기에서 읽어보세요</a>

플렉스 컨테이너를 행(row)으로 설정하면 플렉스 항목들이 왼쪽에서 오른쪽으로 나란히 배치됩니다. 플렉스 컨테이너를 열(column)로 설정하면 플렉스 항목들이 위에서 아래로 쌓입니다. 이때 플렉스 항목들이 배치되는 방향을 **주축**이라고 합니다. 행의 경우 주축은 각 항목을 가로지르는 수평선이고, 열의 경우 주축은 항목들을 세로로 관통하는 선입니다.

주축을 따라 플렉스 항목들을 어떻게 간격을 두고 배치할지에 대한 여러 옵션이 있습니다. 가장 자주 쓰이는 것 중 하나가 `justify-content: center;`로, 모든 플렉스 항목을 플렉스 컨테이너 안에서 가운데로 정렬합니다. 그 외 옵션은 다음과 같습니다:

<ul><li><code>flex-start</code>: 항목들을 플렉스 컨테이너의 시작 부분에 정렬합니다. 행의 경우 항목들을 컨테이너의 왼쪽으로 밀어냅니다. 열의 경우 항목들을 컨테이너의 위쪽으로 밀어냅니다. <code>justify-content</code>가 지정되지 않으면 기본 정렬값입니다.</li><li><code>flex-end</code>: 항목들을 플렉스 컨테이너의 끝 부분에 정렬합니다. 행의 경우 항목들을 컨테이너의 오른쪽으로 밀어냅니다. 열의 경우 항목들을 컨테이너의 아래쪽으로 밀어냅니다.</li><li><code>space-between</code>: 항목들을 주축의 중앙에 정렬하되, 항목들 사이에 여분의 공간을 둡니다. 첫 번째와 마지막 항목은 플렉스 컨테이너의 가장자리에 붙습니다. 예를 들어 행에서는 첫 번째 항목이 컨테이너 왼쪽에, 마지막 항목이 오른쪽에 붙고 나머지 공간이 나머지 항목들 사이에 균등하게 분배됩니다.</li><li><code>space-around</code>: <code>space-between</code>과 비슷하지만 첫 번째와 마지막 항목이 컨테이너 가장자리에 고정되지 않고, 모든 항목 주위에 공간이 분배되며 컨테이너 양 끝에는 반만큼의 공간이 있습니다.</li><li><code>space-evenly</code>: 플렉스 항목들 사이에 공간을 균등하게 분배하며, 컨테이너 양 끝에도 동일한 크기의 공간이 있습니다.</li></ul>

# --instructions--

예제를 통해 이 속성이 어떻게 작동하는지 보여드리겠습니다. `justify-content` 요소에 CSS 속성 `#box-container`를 추가하고, 값으로 `center`를 지정하세요.

**보너스**  
코드 편집기에서 `justify-content` 속성의 다른 옵션들을 시도해 보면서 차이를 확인하세요. 단, `center` 값만 이 도전 과제를 통과할 수 있습니다.

# --hints--

`#box-container` 요소에 `justify-content` 속성을 `center` 값으로 설정해야 합니다.

```js
const boxContainer = document.querySelector('#box-container');
const justifyDirection =
  window.getComputedStyle(boxContainer)['justify-content'];
assert.strictEqual(justifyDirection, 'center');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 100%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 100%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;
    justify-content: center;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 100%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 100%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
