---
id: 587d78ab367417b2b2512af0
title: '두 박스를 배치하기 위해 display: flex 사용하기'
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cgz3QS7'
forumTopicId: 301105
dashedName: use-display-flex-to-position-two-boxes
---

# --description--

이 섹션은 CSS를 사용해 요소를 유연하게 배치하는 방법을 보여주기 위해 번갈아 가며 도전 과제 스타일을 사용합니다. 먼저 도전 과제가 이론을 설명하고, 그다음 간단한 트윗 컴포넌트를 사용해 flexbox 개념을 적용하는 실습 도전 과제가 이어집니다.

요소에 CSS 속성 `display: flex;`를 지정하면 다른 flex 속성을 사용해 반응형 페이지를 만들 수 있습니다.

# --instructions--

`display` CSS 속성을 `#box-container`에 추가하고 값을 `flex`로 설정하세요.

# --hints--

`#box-container`에 `display` 속성을 `flex` 값으로 설정해야 합니다.

```js
const boxContainer = document.querySelector('#box-container');
const displayStyle = window.getComputedStyle(boxContainer)['display'];
assert.strictEqual(displayStyle, 'flex');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    height: 500px;

  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    height: 500px;
    display: flex;
  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
