---
id: 587d78ae367417b2b2512aff
title: order 속성으로 항목 순서 재배열하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvNAG'
forumTopicId: 301116
dashedName: use-the-order-property-to-rearrange-items
---

# --description--

`order` 속성은 플렉스 컨테이너 내에서 플렉스 항목이 나타나는 순서를 CSS에 알려주는 데 사용합니다. 기본적으로 항목은 소스 HTML에 나오는 순서대로 표시됩니다. 이 속성은 숫자를 값으로 가지며, 음수도 사용할 수 있습니다.

# --instructions--

`order` CSS 속성을 `#box-1`와 `#box-2` 모두에 추가하세요. `#box-1`에는 `2` 값을, `#box-2`에는 `1` 값을 주세요.

# --hints--

`#box-1` 요소에 `order` 속성을 `2` 값으로 설정하세요.

```js
const boxOne = document.querySelector('#box-1');
const order = window.getComputedStyle(boxOne)['order'];
assert.strictEqual(order, '2');
```

`#box-2` 요소에 `order` 속성을 `1` 값으로 설정하세요.

```js
const boxTwo = document.querySelector('#box-2');
const order = window.getComputedStyle(boxTwo)['order'];
assert.strictEqual(order, '1');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    order: 2;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    order: 1;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
