---
id: 587d78ae367417b2b2512afc
title: flex-grow 속성으로 아이템 확장하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c2p78cg'
forumTopicId: 301111
dashedName: use-the-flex-grow-property-to-expand-items
---

# --description--

`flex-shrink`의 반대는 `flex-grow` 속성입니다. `flex-shrink`가 컨테이너가 줄어들 때 아이템 크기를 제어한다는 점을 기억하세요. `flex-grow` 속성은 부모 컨테이너가 확장될 때 아이템 크기를 제어합니다.

지난 도전 과제의 비슷한 예를 사용하면, 한 아이템의 `flex-grow` 값이 `1`이고 다른 아이템의 `flex-grow` 값이 `3`이라면, `3` 값을 가진 아이템이 다른 아이템보다 세 배 더 크게 성장합니다.

# --instructions--

`flex-grow` CSS 속성을 `#box-1`과 `#box-2` 모두에 추가하세요. `#box-1`에는 `1` 값을, `#box-2`에는 `2` 값을 주세요.

# --hints--

`#box-1` 요소에 `flex-grow` 속성을 `1` 값으로 설정하세요.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
assert.equal(flexGrow, '1');
```

`#box-2` 요소에 `flex-grow` 속성을 `2` 값으로 설정하세요.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
assert.equal(flexGrow, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-grow: 1;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-grow: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
