---
id: 587d78ae367417b2b2512afd
title: 아이템의 초기 크기를 설정하기 위해 flex-basis 속성 사용
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c3d9nCa'
forumTopicId: 301108
dashedName: use-the-flex-basis-property-to-set-the-initial-size-of-an-item
---

# --description--

`flex-basis` 속성은 CSS가 `flex-shrink` 또는 `flex-grow`로 조정하기 전에 아이템의 초기 크기를 지정합니다.

`flex-basis` 속성에 사용하는 단위는 다른 크기 속성들(`px`, `em`, `%` 등)과 동일합니다. 값 `auto`은 콘텐츠를 기준으로 아이템 크기를 정합니다.

# --instructions--

`flex-basis`를 사용해 박스들의 초기 크기를 설정하세요. `flex-basis` CSS 속성을 `#box-1`과 `#box-2` 모두에 추가하세요. `#box-1`에는 `10em` 값을, `#box-2`에는 `20em` 값을 주세요.

# --hints--

`#box-1` 요소는 `flex-basis` 속성을 가져야 합니다.

```js
const boxOne = document.querySelector('#box-1');
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

`#box-1` 요소는 `flex-basis` 값으로 `10em`을 가져야 합니다.

```js
assert.match(code, /#box-1\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?10em;/g);
```

`#box-2` 요소는 `flex-basis` 속성을 가져야 합니다.

```js
const boxTwo = document.querySelector('#box-2');
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

`#box-2` 요소는 `flex-basis` 값으로 `20em`를 가져야 합니다.

```js
assert.match(code, /#box-2\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?20em;/g);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-basis: 10em;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-basis: 20em;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
