---
id: 587d78af367417b2b2512b00
title: align-self 속성 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvzfv'
forumTopicId: 301107
dashedName: use-the-align-self-property
---

# --description--

플렉스 항목에 대한 마지막 속성은 `align-self`입니다. 이 속성은 모든 항목을 한꺼번에 설정하는 대신 각 항목의 정렬을 개별적으로 조정할 수 있게 해줍니다. 이는 CSS 속성 `float`, `clear`, 그리고 `vertical-align`를 사용하는 다른 일반적인 조정 기법들이 플렉스 항목에는 작동하지 않기 때문에 유용합니다.

`align-self`는 `align-items`와 동일한 값을 허용하며 `align-items` 속성에 의해 설정된 값을 재정의합니다.

# --instructions--

`align-self` CSS 속성을 `#box-1`와 `#box-2` 모두에 추가하세요. `#box-1`에는 `center` 값을, `#box-2`에는 `flex-end` 값을 주세요.

# --hints--

`#box-1` 요소에는 `align-self` 속성을 `center` 값으로 설정해야 합니다.

```js
const boxOne = document.querySelector('#box-1');
const alignment = window.getComputedStyle(boxOne)['align-self'];
assert.strictEqual(alignment, 'center');
```

`#box-2` 요소에는 `align-self` 속성을 `flex-end` 값으로 설정해야 합니다.

```js
const boxTwo = document.querySelector('#box-2');
const alignment = window.getComputedStyle(boxTwo)['align-self'];
assert.strictEqual(alignment, 'flex-end');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    align-self: center;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    align-self: flex-end;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
