---
id: 587d78ac367417b2b2512af4
title: 플렉스 방향 속성으로 컬럼 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cZmWeA4'
forumTopicId: 301109
dashedName: use-the-flex-direction-property-to-make-a-column
---

# --description--

마지막 두 도전 과제에서는 `flex-direction` 속성을 `row`로 설정했습니다. 이 속성은 플렉스 컨테이너의 자식 요소들을 수직으로 쌓아 컬럼을 만들 수도 있습니다.

# --instructions--

`flex-direction` 요소에 CSS 속성 `#box-container`를 추가하고, 그 값을 `column`로 지정하세요.

# --hints--

`#box-container` 요소는 `flex-direction` 속성을 `column`로 설정해야 합니다.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: column;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
