---
id: 587d78ad367417b2b2512afb
title: flex-shrink 속성으로 항목 축소하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PBfr'
forumTopicId: 301113
dashedName: use-the-flex-shrink-property-to-shrink-items
---

# --description--

지금까지 도전 과제에서 나온 모든 속성은 플렉스 항목의 부모인 플렉스 컨테이너에 적용되었습니다. 하지만 플렉스 항목에 유용한 속성도 몇 가지 있습니다.

첫 번째는 `flex-shrink` 속성입니다. 이 속성을 사용하면 플렉스 컨테이너가 너무 작을 때 항목이 축소될 수 있습니다. 항목은 부모 컨테이너의 너비가 그 안에 있는 모든 플렉스 항목의 너비 합보다 작을 때 축소됩니다.

`flex-shrink` 속성은 숫자를 값으로 받습니다. 숫자가 클수록 컨테이너 내 다른 항목에 비해 더 많이 축소됩니다. 예를 들어, 한 항목의 `flex-shrink` 값이 `1`이고 다른 항목의 `flex-shrink` 값이 `3`라면, `3` 값을 가진 항목이 다른 항목보다 세 배 더 많이 축소됩니다.

# --instructions--

`flex-shrink` CSS 속성을 `#box-1`와 `#box-2` 모두에 추가하세요. `#box-1`에는 `1` 값을, `#box-2`에는 `2` 값을 주세요.

# --hints--

`#box-1` 요소에 `flex-shrink` 속성을 `1` 값으로 설정하세요.

```js
const boxOne = document.querySelector('#box-1');
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
assert.equal(flexShrink, '1');
```

`#box-2` 요소에 `flex-shrink` 속성을 `2` 값으로 설정하세요.

```js
const boxTwo = document.querySelector('#box-2');
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
assert.equal(flexShrink, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;
    flex-shrink: 1;
  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;
    flex-shrink: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
