---
id: 6814d8e1516e86b171929de4
title: "Aufgabe 1: Vokalgleichgewicht"
challengeType: 28
dashedName: challenge-1
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob die Anzahl der Vokale in der ersten Hälfte des Strings gleich der Anzahl der Vokale in der zweiten Hälfte ist.

- Der String kann beliebige Zeichen enthalten.
- Die Buchstaben `a`, `e`, `i`, `o` und `u`, sowohl in Groß- als auch in Kleinbuchstaben, gelten als Vokale.
- Wenn die Anzahl der Zeichen im String ungerade ist, ignorieren Sie das mittlere Zeichen.

# --hints--

`isBalanced("racecar")` sollte `true` zurückgeben.

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")` sollte `true` zurückgeben.

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")` sollte `false` zurückgeben.

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")` sollte `false` zurückgeben.

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")` sollte `true` zurückgeben.

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")` sollte `false` zurückgeben.

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")` sollte `true` zurückgeben.

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const half = Math.floor(s.length / 2);

  let firstHalf = s.slice(0, half);
  let secondHalf = s.length % 2 === 0 ? s.slice(half) : s.slice(half + 1);

  const countVowels = str =>
    str
      .toLowerCase()
      .split('')
      .filter(c => vowels.includes(c))
      .length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
