---
id: 681cb1a2dab50c87ddb2e514
title: "Aufgabe 3: Fibonacci-Folge"
challengeType: 28
dashedName: challenge-3
---

# --description--

Die Fibonacci-Folge ist eine Zahlenreihe, bei der jede Zahl die Summe der beiden vorhergehenden ist. Beginnt man mit `0` und `1`, sind die ersten 10 Zahlen der Folge `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

Gegeben ist ein Array, das die ersten beiden Zahlen einer Fibonacci-Folge enthält, sowie eine ganze Zahl, die die Länge der Folge angibt. Geben Sie ein Array zurück, das die Folge der angegebenen Länge enthält.

- Ihre Funktion soll Folgen jeder Länge größer oder gleich null verarbeiten.
- Wenn die Länge null ist, geben Sie ein leeres Array zurück.
- Beachten Sie, dass die Startzahlen Teil der Folge sind.

# --hints--

`fibonacciSequence([0, 1], 20)` sollte `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]` zurückgeben.

```js
assert.deepEqual(fibonacciSequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]);
```

`fibonacciSequence([21, 32], 1)` sollte `[21]` zurückgeben.

```js
assert.deepEqual(fibonacciSequence([21, 32], 1), [21]);
```

`fibonacciSequence([0, 1], 0)` sollte `[]` zurückgeben.

```js
assert.deepEqual(fibonacciSequence([0, 1], 0), []);
```

`fibonacciSequence([10, 20], 2)` sollte `[10, 20]` zurückgeben.

```js
assert.deepEqual(fibonacciSequence([10, 20], 2), [10, 20]);
```

`fibonacciSequence([123456789, 987654321], 5)` sollte `[123456789, 987654321, 1111111110, 2098765431, 3209876541]` zurückgeben.

```js
assert.deepEqual(fibonacciSequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541]);
```

# --seed--

## --seed-contents--

```js
function fibonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function fibonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2];
    sequence.push(nextValue);
  }
  return sequence;
}
```
