---
id: 681cb05adab50c87ddb2e513
title: "Aufgabe 2: Basisprüfung"
challengeType: 28
dashedName: challenge-2
---

# --description--

Gegeben ist ein String, der eine Zahl darstellt, und eine ganze Zahl als Basis von 2 bis 36. Bestimmen Sie, ob die Zahl in dieser Basis gültig ist.

- Der String kann ganze Zahlen sowie Groß- oder Kleinbuchstaben enthalten.
- Die Prüfung soll nicht zwischen Groß- und Kleinbuchstaben unterscheiden.
- Die Basis kann eine beliebige Zahl von 2 bis 36 sein.
- Eine Zahl ist gültig, wenn jedes Zeichen eine gültige Ziffer in der angegebenen Basis ist.
- Beispiele für gültige Ziffern in verschiedenen Basen:
  - Basis 2: 0-1
  - Basis 8: 0-7
  - Basis 10: 0-9
  - Basis 16: 0-9 und A-F
  - Basis 36: 0-9 und A-Z

# --hints--

`isValidNumber("10101", 2)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` sollte `false` zurückgeben.

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` sollte `false` zurückgeben.

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` sollte `false` zurückgeben.

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` sollte `false` zurückgeben.

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` sollte `false` zurückgeben.

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNumber("z", 36))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
