---
id: 681cb1afdab50c87ddb2e516
title: "Aufgabe 5: Jbelmud-Text"
challengeType: 28
dashedName: challenge-5
---

# --description--

Gegeben ist ein String. Geben Sie eine durcheinandergebrachte Version dieses Strings zurück, bei der jedes Wort nach folgenden Vorgaben umgewandelt wird:

- Der erste und der letzte Buchstabe der Wörter bleiben an ihrer Stelle
- Alle Buchstaben zwischen dem ersten und letzten Buchstaben werden alphabetisch sortiert
- Die Eingabe-Strings enthalten keine Satzzeichen und bestehen vollständig aus Kleinbuchstaben

# --hints--

`jbelmu("hello world")` sollte `"hello wlord"` zurückgeben.

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` sollte `"i love jbelmud text"` zurückgeben.

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` sollte `"faccdeeemorp is my faiortve pacle to laern to cdoe"` zurückgeben.

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` sollte `"the qciuk borwn fox jmpus oevr the lazy dog"` zurückgeben.

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
