---
id: 681cb1b0dab50c87ddb2e519
title: "Aufgabe 8: Fakultätsberechnung"
challengeType: 28
dashedName: challenge-8
---

# --description--

Gegeben ist eine ganze Zahl von null bis 20. Geben Sie die Fakultät dieser Zahl zurück. Die Fakultät einer Zahl ist das Produkt aller Zahlen zwischen 1 und der angegebenen Zahl.

- Die Fakultät von null ist 1.

# --hints--

`factorial(0)` sollte `1` zurückgeben.

```js
assert.equal(factorial(0), 1);
```

`factorial(5)` sollte `120` zurückgeben.

```js
assert.equal(factorial(5), 120);
```

`factorial(20)` sollte `2432902008176640000` zurückgeben.

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
