---
id: 681cb1afdab50c87ddb2e517
title: "Aufgabe 6: Anagrammprüfer"
challengeType: 28
dashedName: challenge-6
---

# --description--

Gegeben sind zwei Strings. Bestimmen Sie, ob sie Anagramme voneinander sind (d. h. dieselben Zeichen in beliebiger Reihenfolge enthalten).

- Groß- und Kleinschreibung sowie Leerzeichen ignorieren.

# --hints--

`areAnagrams("listen", "silent")` sollte `true` zurückgeben.

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` sollte `true` zurückgeben.

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` sollte `true` zurückgeben.

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` sollte `false` zurückgeben.

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` sollte `false` zurückgeben.

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` sollte `false` zurückgeben.

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
