---
id: 681cb1b0dab50c87ddb2e51a
title: "Aufgabe 9: Quadratsumme"
challengeType: 28
dashedName: challenge-9
---

# --description--

Gegeben ist eine positive ganze Zahl bis zu 1.000. Geben Sie die Summe aller quadrierten ganzen Zahlen von 1 bis zu dieser Zahl zurück.

# --hints--

`sumOfSquares(5)` sollte `55` zurückgeben.

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` sollte `385` zurückgeben.

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` sollte `5525` zurückgeben.

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` sollte `41791750` zurückgeben.

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` sollte `333833500` zurückgeben.

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
