---
id: 681cb1b0dab50c87ddb2e51b
title: "Aufgabe 10: 3 Strikes"
challengeType: 28
dashedName: challenge-10
---

# --description--

Gegeben ist eine ganze Zahl zwischen 1 und 10.000. Geben Sie die Anzahl der Zahlen von 1 bis zu dieser Zahl zurück, deren Quadrat mindestens eine Ziffer 3 enthält.

# --hints--

`squaresWithThree(1)` sollte `0` zurückgeben.

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` sollte `1` zurückgeben.

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` sollte `19` zurückgeben.

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` sollte `326` zurückgeben.

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` sollte `4531` zurückgeben.

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
