---
id: 681cb1afdab50c87ddb2e515
title: "Aufgabe 4: S  P  A  C  E  J  A  M"
challengeType: 28
dashedName: challenge-4
---

# --description--

Gegeben ist ein String. Entfernen Sie alle Leerzeichen aus dem String, fügen Sie zwischen jedem Zeichen zwei Leerzeichen ein, wandeln Sie alle Buchstaben in Großbuchstaben um und geben Sie das Ergebnis zurück.

- Nicht alphabetische Zeichen bleiben unverändert (außer Leerzeichen).

# --hints--

`spaceJam("freeCodeCamp")` sollte `"F  R  E  E  C  O  D  E  C  A  M  P"` zurückgeben.

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")` sollte `"F  R  E  E  C  O  D  E  C  A  M  P"` zurückgeben.

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Hello World?!")` sollte `"H  E  L  L  O  W  O  R  L  D  ?  !"` zurückgeben.

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")` sollte `"C  @  T  $  &  D  0  G  $"` zurückgeben.

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")` sollte `"A  L  L  Y  O  U  R  B  A  S  E"` zurückgeben.

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
