---
id: 681cb1b0dab50c87ddb2e518
title: "Aufgabe 7: Gezielt summieren"
challengeType: 28
dashedName: challenge-7
---

# --description--

Gegeben ist ein Array von Zahlen und ein ganzzahliges Ziel. Finden Sie zwei eindeutige Zahlen im Array, deren Summe dem Zielwert entspricht. Geben Sie ein Array mit den Indizes dieser beiden Zahlen zurück, oder `"Target not found"`, wenn keine zwei Zahlen die Summe des Ziels ergeben.

- Das zurückgegebene Array sollte die Indizes in aufsteigender Reihenfolge enthalten.

# --hints--

`findTarget([2, 7, 11, 15], 9)` sollte `[0, 1]` zurückgeben.

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)` sollte `[1, 2]` zurückgeben.

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)` sollte `[4, 5]` zurückgeben.

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)` sollte `"Target not found"` zurückgeben.

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```
