---
id: 6821ec02237de8297eaee79a
title: "Aufgabe 24: Pangramm"
challengeType: 28
dashedName: challenge-24
---

# --description--

Gegeben sind ein Wort oder Satz und ein String aus Kleinbuchstaben. Bestimmen Sie, ob das Wort oder der Satz alle Buchstaben aus dem gegebenen Satz mindestens einmal verwendet und keine anderen Buchstaben enthält.

- Ignorieren Sie nicht-alfabetische Zeichen im Wort oder Satz.
- Ignorieren Sie die Groß- und Kleinschreibung im Wort oder Satz.

# --hints--

`isPangram("hello", "helo")` sollte `true` zurückgeben

```js
assert.isTrue(isPangram("hello", "helo"));
```

`isPangram("hello", "hel")` sollte `false` zurückgeben

```js
assert.isFalse(isPangram("hello", "hel"));
```

`isPangram("hello", "helow")` sollte `false` zurückgeben

```js
assert.isFalse(isPangram("hello", "helow"));
```

`isPangram("hello world", "helowrd")` sollte `true` zurückgeben

```js
assert.isTrue(isPangram("hello world", "helowrd"));
```

`isPangram("Hello World!", "helowrd")` sollte `true` zurückgeben

```js
assert.isTrue(isPangram("Hello World!", "helowrd"));
```

`isPangram("Hello World!", "heliowrd")` sollte `false` zurückgeben

```js
assert.isFalse(isPangram("Hello World!", "heliowrd"));
```

`isPangram("freeCodeCamp", "frcdmp")` sollte `false` zurückgeben

```js
assert.isFalse(isPangram("freeCodeCamp", "frcdmp"));
```

`isPangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz")` sollte `true` zurückgeben

```js
assert.isTrue(isPangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz"));
```

# --seed--

## --seed-contents--

```js
function isPangram(sentence, letters) {

  return sentence;
}
```

# --solutions--

```js
function isPangram(sentence, letters) {
  const usedLetters = [];
  for (let i = 0; i < sentence.length; i++) {
    const letter = sentence[i].toLowerCase();
    if (!usedLetters.includes(letter) && /[a-z]/.test(letter)) {
      usedLetters.push(letter);
    }
  }

  const sortedLetters = letters.split('').sort().join('');
  const sortedUsedLetters = usedLetters.sort().join('');

  return sortedLetters === sortedUsedLetters;
}
```
