---
id: 6821ebda237de8297eaee792
title: "Aufgabe 16: Runde Klammern umkehren"
challengeType: 28
dashedName: challenge-16
---

# --description--

Gegeben ist ein String, der korrekt verschachtelte runde Klammern enthält. Geben Sie die dekodierte Version des Strings zurück, indem Sie die folgenden Regeln anwenden:

- Alle Zeichen innerhalb jedes Klammernpaares sollen umgekehrt werden.
- Die runden Klammern sollen im Endergebnis entfernt werden.
- Wenn Klammern verschachtelt sind, soll zuerst das innerste Paar umgekehrt werden, und dessen Ergebnis dann in die Umkehrung des äußeren Paares einfließen.
- Gehen Sie davon aus, dass alle Klammern gleichmäßig ausbalanciert und korrekt verschachtelt sind.

# --hints--

`decode("(f(b(dc)e)a)")` sollte `"abcdef"` zurückgeben.

```js
assert.equal(decode("(f(b(dc)e)a)"), "abcdef");
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` sollte `"Can you read this?"` zurückgeben.

```js
assert.equal(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?");
```

`decode("f(Ce(re))o((e(aC)m)d)p")` sollte `"freeCodeCamp"` zurückgeben.

```js
assert.equal(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(s) {

  return s;
}
```

# --solutions--

```js
function decode(s) {
  while (s.includes(')')) {
    const closeIndex = s.indexOf(')');
    const openIndex = s.lastIndexOf('(', closeIndex);
    const before = s.slice(0, openIndex);
    const group = s.slice(openIndex + 1, closeIndex).split('').reverse().join('');
    const after = s.slice(closeIndex + 1);
    s = before + group + after;
  }

  return s;
}
```
