---
id: 6821ebf8237de8297eaee798
title: "Aufgabe 22: Tribonacci-Folge"
challengeType: 28
dashedName: challenge-22
---

# --description--

Die Tribonacci-Folge ist eine Zahlenreihe, bei der jede Zahl die Summe der drei vorhergehenden ist. Beginnt man mit `0`, `0` und `1`, sind die ersten 10 Zahlen der Folge `0`, `0`, `1`, `1`, `2`, `4`, `7`, `13`, `24`, `44`.

Gegeben ist ein Array mit den ersten drei Zahlen einer Tribonacci-Folge und eine ganze Zahl, die die Länge der Folge angibt. Geben Sie ein Array zurück, das die Folge mit der angegebenen Länge enthält.

- Ihre Funktion soll Folgen jeder Länge größer oder gleich null verarbeiten.
- Wenn die Länge null ist, geben Sie ein leeres Array zurück.
- Beachten Sie, dass die Startzahlen Teil der Folge sind.

# --hints--

`tribonacciSequence([0, 0, 1], 20)` sollte `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]` zurückgeben.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)` sollte `[21]` zurückgeben.


```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)` sollte `[]` zurückgeben.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)` sollte `[10, 20]` zurückgeben.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)` sollte `[10, 20, 30]` zurückgeben.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)` sollte `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]` zurückgeben.

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```
