---
id: 68216eb60f957572e7c340c4
title: "Aufgabe 11: Meilen-Pace"
challengeType: 28
dashedName: challenge-11
---

# --description--

Gegeben ist eine Anzahl gelaufener Meilen und eine Zeit in `"MM:SS"` (Minuten:Sekunden), die benötigt wurde, um diese Meilen zu laufen. Geben Sie einen String zurück, der die durchschnittliche Zeit angibt, die für jede Meile benötigt wurde, im Format `"MM:SS"`.

- Fügen Sie bei Bedarf führende Nullen hinzu.

# --hints--

`milePace(3, "24:00")` sollte `"08:00"` zurückgeben.

```js
assert.equal(milePace(3, "24:00"), "08:00");
```

`milePace(1, "06:45")` sollte `"06:45"` zurückgeben.

```js
assert.equal(milePace(1, "06:45"), "06:45");
```

`milePace(2, "07:00")` sollte `"03:30"` zurückgeben.

```js
assert.equal(milePace(2, "07:00"), "03:30");
```

`milePace(26.2, "120:35")` sollte `"04:36"` zurückgeben.

```js
assert.equal(milePace(26.2, "120:35"), "04:36");
```

# --seed--

## --seed-contents--

```js
function milePace(miles, duration) {

  return miles;
}
```

# --solutions--

```js
function milePace(miles, duration) {
  const [minutes, seconds] = duration.split(':').map(Number);
  const totalSeconds = minutes * 60 + seconds;
  const avgSecondsPerMile = totalSeconds / miles;

  const avgMinutes = Math.floor(avgSecondsPerMile / 60);
  const avgSeconds = Math.round(avgSecondsPerMile % 60);

  const paddedMinutes = avgMinutes.toString().padStart(2, '0');
  const paddedSeconds = avgSeconds.toString().padStart(2, '0');

  return `${paddedMinutes}:${paddedSeconds}`;
}
```
