---
id: 6821ebfd237de8297eaee799
title: "Aufgabe 23: RGB in Hex"
challengeType: 28
dashedName: challenge-23
---

# --description--

Gegeben ist ein CSS `rgb(r, g, b)`-Farbstring. Geben Sie dessen hexadezimale Entsprechung zurück.

Hier sind einige Beispielausgaben für eine gegebene Eingabe:

| Eingabe | Ausgabe  |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- Verwenden Sie Kleinbuchstaben für alle Buchstaben.
- Geben Sie einen `#` gefolgt von sechs Zeichen zurück. Verwenden Sie keine Kurzschreibweisen.

# --hints--

`rgbToHex("rgb(255, 255, 255)")` sollte `"#ffffff"` zurückgeben.

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")` sollte `"#010b6f"` zurückgeben.

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")` sollte `"#add8e6"` zurückgeben.

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")` sollte `"#4f7bc9"` zurückgeben.

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```
