---
id: 68216ef80f957572e7c340c5
title: "Aufgabe 12: Nachrichten-Decoder"
challengeType: 28
dashedName: challenge-12
---

# --description--

Gegeben ist ein geheimer Nachrichten-String und eine ganze Zahl, die angibt, um wie viele Buchstaben die Nachricht verschoben wurde, um sie zu kodieren. Geben Sie den dekodierten String zurück.

- Eine positive Zahl bedeutet, dass die Nachricht im Alphabet nach vorne verschoben wurde.
- Eine negative Zahl bedeutet, dass die Nachricht im Alphabet nach hinten verschoben wurde.
- Die Groß- und Kleinschreibung ist wichtig, dekodierte Zeichen sollten die Groß- oder Kleinschreibung ihrer kodierten Gegenstücke beibehalten.
- Nicht-alfabetische Zeichen sollten nicht dekodiert werden.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` sollte `"This is a secret message."` zurückgeben

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` sollte `"Hello World!"` zurückgeben

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` sollte `"Are you okay?"` zurückgeben

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` sollte `"freeCodeCamp"` zurückgeben

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
