---
id: 6821ebee237de8297eaee796
title: "Aufgabe 20: Array-Duplikate"
challengeType: 28
dashedName: challenge-20
---

# --description--

Gegeben ist ein Array ganzer Zahlen. Geben Sie ein Array ganzer Zahlen zurück, die im ursprünglichen Array mehr als einmal vorkommen, sortiert in aufsteigender Reihenfolge. Wenn keine Werte mehr als einmal vorkommen, geben Sie ein leeres Array zurück.

- Nehmen Sie im zurückgegebenen Array nur eine Instanz jedes Werts auf.

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` sollte `[]` zurückgeben.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` sollte `[1, 2]` zurückgeben.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` sollte `[-6, 0, 2, 4, 5, 23]` zurückgeben.

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
