---
id: 6821ebea237de8297eaee795
title: "Aufgabe 19: Kerzenlicht"
challengeType: 28
dashedName: challenge-19
---

# --description--

Gegeben ist eine ganze Zahl, die die Anzahl der Kerzen darstellt, mit denen Sie starten, und eine ganze Zahl, die angibt, wie viele abgebrannte Kerzen benötigt werden, um eine neue zu erzeugen. Geben Sie die Anzahl der Kerzen zurück, die Sie insgesamt verwendet haben, nachdem Sie so viele wie möglich hergestellt und abgebrannt haben.

Wenn Sie zum Beispiel 7 Kerzen haben und es 2 abgebrannte Kerzen braucht, um eine neue zu machen:

1. Brennen Sie 7 Kerzen ab, um 7 Reste zu erhalten,
2. Recyceln Sie 6 Reste zu 3 neuen Kerzen (1 Rest bleibt übrig),
3. Brennen Sie 3 Kerzen ab, um 3 weitere Reste zu erhalten (insgesamt 4),
4. Recyceln Sie 4 Reste zu 2 neuen Kerzen,
5. Brennen Sie 2 Kerzen ab, um 2 Reste zu erhalten,
6. Recyceln Sie 2 Reste zu 1 neuer Kerze,
7. Brennen Sie 1 Kerze ab.

In diesem Beispiel haben Sie insgesamt 13 Kerzen abgebrannt.

# --hints--

`burnCandles(7, 2)` sollte `13` zurückgeben

```js
assert.equal(burnCandles(7, 2), 13);
```

`burnCandles(10, 5)` sollte `12` zurückgeben

```js
assert.equal(burnCandles(10, 5), 12);
```

`burnCandles(20, 3)` sollte `29` zurückgeben

```js
assert.equal(burnCandles(20, 3), 29);
```

`burnCandles(17, 4)` sollte `22` zurückgeben

```js
assert.equal(burnCandles(17, 4), 22);
```

`burnCandles(2345, 3)` sollte `3517` zurückgeben

```js
assert.equal(burnCandles(2345, 3), 3517);
```

# --seed--

## --seed-contents--

```js
function burnCandles(candles, leftoversNeeded) {

  return candles;
}
```

# --solutions--

```js
function burnCandles(candles, leftoversNeeded) {
  let totalBurned = 0;
  let unusedLeftovers = 0;

  while (candles > 0) {
    totalBurned += candles;
    const leftovers = candles + unusedLeftovers;
    candles = Math.floor(leftovers / leftoversNeeded);
    unusedLeftovers = leftovers % leftoversNeeded;
  }

  return totalBurned;
}
```
