---
id: 6821ebdf237de8297eaee793
title: "Aufgabe 17: Reihenfolge der Operationen"
challengeType: 28
dashedName: challenge-17
---

# --description--

Gegeben sind ein Array von ganzen Zahlen und ein Array von String-Operatoren. Wenden Sie die Operationen nacheinander von links nach rechts auf die Zahlen an. Wiederholen Sie die Operationen bei Bedarf, bis alle Zahlen verwendet wurden. Geben Sie das Endergebnis zurück.

Zum Beispiel, gegeben `[1, 2, 3, 4, 5]` und `['+', '*']`, geben Sie das Ergebnis der Auswertung von `1 + 2 * 3 + 4 * 5` von links nach rechts zurück, wobei die übliche Reihenfolge der Operationen ignoriert wird.

- Gültige Operatoren sind `+`, `-`, `*`, `/` und `%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` sollte `3` zurückgeben

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` sollte `38` zurückgeben

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` sollte `60` zurückgeben

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` sollte `30` zurückgeben

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])` sollte `-2` zurückgeben

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```
