---
id: 6821ebce237de8297eaee790
title: "Aufgabe 14: Charakterkampf"
challengeType: 28
dashedName: challenge-14
---

# --description--

Gegeben sind zwei Strings, die Ihre Armee und eine gegnerische Armee darstellen. Jeder Charakter aus Ihrer Armee kämpft gegen den Charakter an derselben Position der gegnerischen Armee nach folgenden Regeln:

- Die Charaktere `a-z` haben jeweils eine Stärke von `1-26`.
- Die Charaktere `A-Z` haben jeweils eine Stärke von `27-52`.
- Ziffern `0-9` haben eine Stärke entsprechend ihrem Zahlenwert.
- Alle anderen Charaktere haben einen Wert von null.
- Jeder Charakter kann nur einen Kampf führen.

Für jeden Kampf gewinnt der stärkere Charakter. Die Armee mit den meisten Siegen gewinnt den Krieg. Geben Sie folgende Werte zurück:

- `"Opponent retreated"`, wenn Ihre Armee mehr Charaktere hat als die gegnerische Armee.
- `"We retreated"`, wenn die gegnerische Armee mehr Charaktere hat als Ihre.
- `"We won"`, wenn Ihre Armee mehr Kämpfe gewonnen hat.
- `"We lost"`, wenn die gegnerische Armee mehr Kämpfe gewonnen hat.
- `"It was a tie"`, wenn beide Armeen gleich viele Kämpfe gewonnen haben.

# --hints--

`battle("Hello", "World")` sollte `"We lost"` zurückgeben.

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")` sollte `"We won"` zurückgeben.

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")` sollte `"We won"` zurückgeben.

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")` sollte `"Opponent retreated"` zurückgeben.

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")` sollte `"We retreated"` zurückgeben.

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")` sollte `"It was a tie"` zurückgeben.

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")` sollte `"It was a tie"` zurückgeben.

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```
