---
id: 6821ebe4237de8297eaee794
title: "Aufgabe 18: Zweitbeste"
challengeType: 28
dashedName: challenge-18
---

# --description--

Gegeben ist ein Array ganzer Zahlen, die den Preis verschiedener Laptops darstellen, und eine ganze Zahl, die Ihr Budget darstellt. Geben Sie zurück:

1. Den zweitteuersten Laptop, wenn er innerhalb Ihres Budgets liegt, oder
2. Den teuersten Laptop, der innerhalb Ihres Budgets liegt, oder
3. `0`, wenn kein Laptop innerhalb Ihres Budgets liegt.

- Doppelte Preise sollten ignoriert werden.

# --hints--

`getLaptopCost([1500, 2000, 1800, 1400], 1900)` sollte `1800` zurückgeben

```js
assert.equal(getLaptopCost([1500, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900)` sollte `1800` zurückgeben

```js
assert.equal(getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([2099, 1599, 1899, 1499], 2200)` sollte `1899` zurückgeben

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 2200), 1899);
```

`getLaptopCost([2099, 1599, 1899, 1499], 1000)` sollte `0` zurückgeben

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 1000), 0);
```

`getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` sollte `1400` zurückgeben

```js
assert.equal(getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400);
```

# --seed--

## --seed-contents--

```js
function getLaptopCost(laptops, budget) {

  return laptops;
}
```

# --solutions--

```js
function getLaptopCost(laptops, budget) {
  laptops = [...new Set(laptops)].sort((a, b) => b - a);

  if (budget >= laptops[1]) return laptops[1];
  if (budget < laptops[laptops.length - 1]) return 0;

  for (let i = 2; i < laptops.length; i++) {
    if (budget >= laptops[i]) return laptops[i];
  }
}
```
