---
id: 6821ebd4237de8297eaee791
title: "Aufgabe 15: CamelCase"
challengeType: 28
dashedName: challenge-15
---

# --description--

Gegeben ist ein String. Geben Sie seine CamelCase-Version zurück, wobei folgende Regeln gelten:

- Wörter im String-Argument sind durch ein oder mehrere Zeichen aus der folgenden Menge getrennt: Leerzeichen (` `), Bindestrich (`-`) oder Unterstrich (`_`). Behandeln Sie jede Folge dieser Zeichen als Worttrennung.
- Das erste Wort soll komplett in Kleinbuchstaben sein.
- Jedes folgende Wort soll mit einem Großbuchstaben beginnen, der Rest davon in Kleinbuchstaben.
- Alle Leerzeichen und Trennzeichen sollen entfernt werden.

# --hints--

`toCamelCase("hello world")` sollte `"helloWorld"` zurückgeben.

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")` sollte `"helloWorld"` zurückgeben.

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("secret agent-X")` sollte `"secretAgentX"` zurückgeben.

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")` sollte `"freeCodeCamp"` zurückgeben.

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` sollte `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"` zurückgeben.

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```
