---
id: 6821ebc9237de8297eaee78f
title: "Aufgabe 13: Unnatürliche Primzahl"
challengeType: 28
dashedName: challenge-13
---

# --description--

Gegeben ist eine ganze Zahl. Bestimmen Sie, ob diese Zahl eine Primzahl oder eine negative Primzahl ist.

- Eine Primzahl ist eine positive ganze Zahl größer als 1, die nur durch 1 und sich selbst teilbar ist.
- Eine negative Primzahl ist die negative Version einer positiven Primzahl.
- `1` und `0` werden nicht als Primzahlen betrachtet.

# --hints--

`isUnnaturalPrime(1)` sollte `false` zurückgeben.

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)` sollte `false` zurückgeben.

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)` sollte `true` zurückgeben.

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)` sollte `true` zurückgeben.

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)` sollte `false` zurückgeben.

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)` sollte `true` zurückgeben.

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)` sollte `true` zurückgeben.

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)` sollte `false` zurückgeben.

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)` sollte `false` zurückgeben.

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
