---
id: 68adce01c0e1144d0a902958
title: "Aufgabe 26: IPv4-Validator"
challengeType: 28
dashedName: challenge-26
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob es sich um eine gültige IPv4-Adresse handelt. Eine gültige IPv4-Adresse besteht aus vier Ganzzahlen, die durch Punkte getrennt sind (`.`). Jede Zahl muss die folgenden Bedingungen erfüllen:

- Sie liegt inklusive zwischen 0 und 255.
- Sie darf keine führenden Nullen haben (z. B. ist 0 erlaubt, 01 nicht).
- Es sind nur numerische Zeichen erlaubt.

# --hints--

`isValidIPv4("192.168.1.1")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.01.50.111")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
