---
id: 68adce01c0e1144d0a902956
title: "Aufgabe 25: Vokalwiederholung"
challengeType: 28
dashedName: challenge-25
---

# --description--

Gegeben ist ein String. Geben Sie eine neue Version des Strings zurück, bei der jeder Vokal jeweils um eine Wiederholung mehr erscheint als der vorherige Vokal, den Sie gefunden haben. Zum Beispiel bleibt der erste Vokal im Satz unverändert. Der zweite Vokal erscheint zweimal hintereinander. Der dritte Vokal erscheint dreimal hintereinander, und so weiter.

- Die Buchstaben `a`, `e`, `i`, `o` und `u`, sowohl in Groß- als auch in Kleinbuchstaben, gelten als Vokale.
- Der ursprüngliche Vokal behält seine Groß- oder Kleinschreibung.
- Wiederholte Vokale sollten in Kleinbuchstaben sein.
- Alle Nicht-Vokal-Zeichen behalten ihre ursprüngliche Groß- oder Kleinschreibung.

# --hints--

`repeatVowels("hello world")` sollte `"helloo wooorld"` zurückgeben.

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` sollte `"freeeCooodeeeeCaaaaamp"` zurückgeben.

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` sollte `"AEeIiiOoooUuuuu"` zurückgeben.

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` sollte `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"` zurückgeben.

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
