---
id: 68adce01c0e1144d0a90295e
title: "Aufgabe 29: Akronym-Generator"
challengeType: 28
dashedName: challenge-29
---

# --description--

Gegeben ist ein String, der ein oder mehrere Wörter enthält. Geben Sie ein Akronym der Wörter zurück, das folgende Bedingungen erfüllt:

- Das Akronym soll aus dem jeweils ersten Buchstaben jedes Wortes bestehen, der großgeschrieben wird, sofern nicht anders angegeben.
- Das Akronym soll die ersten Buchstaben dieser Wörter ignorieren, es sei denn, sie sind das erste Wort des gegebenen Strings: `a`, `for`, `an`, `and`, `by` und `of`.
- Die Buchstaben des Akronyms sollen in der Reihenfolge zurückgegeben werden, in der sie vorkommen.
- Das Akronym darf keine Leerzeichen enthalten.

# --hints--

`buildAcronym("Search Engine Optimization")` sollte `"SEO"` zurückgeben.

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")` sollte `"FAQ"` zurückgeben.

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")` sollte `"NASA"` zurückgeben.

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` sollte `"FBI"` zurückgeben.

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")` sollte `"FYI"` zurückgeben.

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("By the way")` sollte `"BTW"` zurückgeben.

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` sollte `"AUHWPOTIMSH"` zurückgeben.

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
