---
id: 68adce01c0e1144d0a90295c
title: "Aufgabe 28: Parser für römische Zahlen"
challengeType: 28
dashedName: challenge-28
---

# --description--

Gegeben ist ein String, der eine römische Zahl darstellt. Geben Sie ihren ganzzahligen Wert zurück.

Römische Zahlen bestehen aus den folgenden Symbolen und Werten:

| Symbol | Wert  |
|--------|-------|
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

- Die Zahlen werden von links nach rechts gelesen. Wenn eine kleinere Zahl vor einer größeren steht, wird ihr Wert subtrahiert. Andernfalls werden die Werte addiert.

# --hints--

`parseRomanNumeral("III")` sollte `3` zurückgeben.

```js
assert.equal(parseRomanNumeral("III"), 3);
```

`parseRomanNumeral("IV")` sollte `4` zurückgeben.

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

`parseRomanNumeral("XXVI")` sollte `26` zurückgeben.

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

`parseRomanNumeral("XCIX")` sollte `99` zurückgeben.

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

`parseRomanNumeral("CDLX")` sollte `460` zurückgeben.

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

`parseRomanNumeral("DIV")` sollte `504` zurückgeben.

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

`parseRomanNumeral("MMXXV")` sollte `2025` zurückgeben.

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
