---
id: 68af0687ef34c76c28ffa54b
title: "Aufgabe 32: Satz umkehren"
challengeType: 28
dashedName: challenge-32
---

# --description--

Gegeben ist ein String mit Wörtern. Geben Sie einen neuen String zurück, in dem die Wörter in umgekehrter Reihenfolge stehen. Zum Beispiel sollte das erste Wort am Ende des zurückgegebenen Strings stehen, und das letzte Wort sollte am Anfang des zurückgegebenen Strings stehen.

- Im gegebenen String können die Wörter durch ein oder mehrere Leerzeichen getrennt sein.
- Der zurückgegebene String sollte nur ein Leerzeichen zwischen den Wörtern enthalten.

# --hints--

`reverseSentence("world hello")` sollte `"hello world"` zurückgeben.

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` sollte `"git commit push"` zurückgeben.

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` sollte `"sudo install npm"` zurückgeben.

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` sollte `"export function default import"` zurückgeben.

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
