---
id: 68af0687ef34c76c28ffa54f
title: "Aufgabe 34: Fehlende Zahlen"
challengeType: 28
dashedName: challenge-34
---

# --description--

Gegeben ist ein Array von ganzen Zahlen von 1 bis `n`, einschließlich. Geben Sie ein Array aller fehlenden ganzen Zahlen zwischen 1 und `n` zurück (wobei `n` die größte Zahl im gegebenen Array ist).

- Das gegebene Array kann unsortiert sein und Duplikate enthalten.
- Das zurückgegebene Array sollte in aufsteigender Reihenfolge sein.
- Wenn keine ganzen Zahlen fehlen, geben Sie ein leeres Array zurück.

# --hints--

`findMissingNumbers([1, 3, 5])` sollte `[2, 4]` zurückgeben.

```js
assert.deepEqual(findMissingNumbers([1, 3, 5]), [2, 4]);
```

`findMissingNumbers([1, 2, 3, 4, 5])` sollte `[]` zurückgeben.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 10])` sollte `[2, 3, 4, 5, 6, 7, 8, 9]` zurückgeben.

```js
assert.deepEqual(findMissingNumbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([10, 1, 10, 1, 10, 1])` sollte `[2, 3, 4, 5, 6, 7, 8, 9]` zurückgeben.

```js
assert.deepEqual(findMissingNumbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([3, 1, 4, 1, 5, 9])` sollte `[2, 6, 7, 8]` zurückgeben.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])` sollte `[11]` zurückgeben.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11]);
```

# --seed--

## --seed-contents--

```js
function findMissingNumbers(arr) {

  return arr;
}
```

# --solutions--

```js
function findMissingNumbers(arr) {
  if (arr.length === 0) return [];
  
  const n = Math.max(...arr);
  const set = new Set(arr);
  const result = [];

  for (let i = 1; i <= n; i++) {
    if (!set.has(i)) {
      result.push(i);
    }
  }
  return result;
}
```
