---
id: 68af0687ef34c76c28ffa547
title: "Aufgabe 30: Einzigartige Zeichen"
challengeType: 28
dashedName: challenge-30
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob alle Zeichen im String einzigartig sind.

- Groß- und Kleinbuchstaben sollen als unterschiedliche Zeichen betrachtet werden.

# --hints--

`allUnique("abc")` sollte `true` zurückgeben.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` sollte `true` zurückgeben.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` sollte `true` zurückgeben.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` sollte `true` zurückgeben.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` sollte `false` zurückgeben.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` sollte `false` zurückgeben.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` sollte `false` zurückgeben.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
