---
id: 68af0687ef34c76c28ffa54d
title: "Aufgabe 33: Bildschirmzeit"
challengeType: 28
dashedName: challenge-33
---

# --description--

Gegeben ist ein Eingabe-Array mit sieben ganzen Zahlen, die eine Woche Zeit darstellen, wobei jede ganze Zahl die Anzahl der Stunden angibt, die Sie an diesem Tag am Telefon verbracht haben. Bestimmen Sie, ob die Bildschirmzeit zu hoch ist, basierend auf diesen Vorgaben:

- Wenn an einem einzelnen Tag 10 Stunden oder mehr erreicht werden, ist das zu viel.
- Wenn der Durchschnitt von drei aufeinanderfolgenden Tagen größer oder gleich 8 Stunden ist, ist das zu viel.
- Wenn der Durchschnitt der sieben Tage größer oder gleich 6 Stunden ist, ist das zu viel.

# --hints--

`tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7])` sollte `false` zurückgeben.

```js
assert.isFalse(tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7]));
```

`tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3])` sollte `false` zurückgeben.

```js
assert.isFalse(tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3]));
```

`tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6])` sollte `false` zurückgeben.

```js
assert.isFalse(tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6]));
```

`tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4])` sollte `true` zurückgeben.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4]));
```

`tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0])` sollte `true` zurückgeben.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0]));
```

`tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4])` sollte `true` zurückgeben.

```js
assert.isTrue(tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4]));
```

`tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6])` sollte `true` zurückgeben.

```js
assert.isTrue(tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6]));
```

# --seed--

## --seed-contents--

```js
function tooMuchScreenTime(hours) {

  return hours;
}
```

# --solutions--

```js
function tooMuchScreenTime(hours) {
  for (let h of hours) {
    if (h >= 10) return true;
  }

  for (let i = 0; i <= 4; i++) {
    const avg3 = (hours[i] + hours[i+1] + hours[i+2]) / 3;
    if (avg3 >= 8) return true;
  }

  const weeklyAvg = hours.reduce((sum, h) => sum + h, 0) / 7;
  if (weeklyAvg >= 6) return true;

  return false;
}
```
