---
id: 68b06e589bf2273243814777
title: "Aufgabe 39: Den Tank füllen"
challengeType: 28
dashedName: challenge-39
---

# --description--

Gegeben sind die Größe eines Kraftstofftanks, der aktuelle Kraftstoffstand und der Preis pro Gallone. Geben Sie die Kosten zurück, um den Tank vollständig zu füllen.

- `tankSize` ist die Gesamtkapazität des Tanks in Gallonen.
- `fuelLevel` ist die aktuelle Menge Kraftstoff im Tank in Gallonen.
- `pricePerGallon` ist der Preis für eine Gallone Kraftstoff.
- Der zurückgegebene Wert sollte auf zwei Dezimalstellen gerundet sein im Format: `"$d.dd"`.

# --hints--

`costToFill(20, 0, 4.00)` sollte `"$80.00"` zurückgeben.

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` sollte `"$17.50"` zurückgeben.

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` sollte `"$29.25"` zurückgeben.

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` sollte `"$0.00"` zurückgeben.

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` sollte `"$21.89"` zurückgeben.

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
