---
id: 68b06e589bf2273243814771
title: "Aufgabe 36: Thermostatregler"
challengeType: 28
dashedName: challenge-36
---

# --description--

Gegeben sind die aktuelle Temperatur eines Raums und eine Zieltemperatur. Geben Sie einen String zurück, der angibt, wie die Raumtemperatur basierend auf diesen Vorgaben angepasst werden soll:

- Geben Sie `"heat"` zurück, wenn die aktuelle Temperatur unter der Zieltemperatur liegt.
- Geben Sie `"cool"` zurück, wenn die aktuelle Temperatur über der Zieltemperatur liegt.
- Geben Sie `"hold"` zurück, wenn die aktuelle Temperatur gleich der Zieltemperatur ist.

# --hints--

`adjustThermostat(68, 72)` sollte `"heat"` zurückgeben.

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` sollte `"cool"` zurückgeben.

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` sollte `"hold"` zurückgeben.

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` sollte `"heat"` zurückgeben.

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` sollte `"cool"` zurückgeben.

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` sollte `"hold"` zurückgeben.

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
