---
id: 68b06e589bf227324381476f
title: "Aufgabe 35: Wortfrequenz"
challengeType: 28
dashedName: challenge-35
---

# --description--

Gegeben ist ein Absatz. Geben Sie ein Array mit den drei am häufigsten vorkommenden Wörtern zurück.

- Die Wörter im Absatz sind durch Leerzeichen getrennt.
- Ignorieren Sie die Groß- und Kleinschreibung im gegebenen Absatz. Behandeln Sie zum Beispiel `Hello` und `hello` als dasselbe Wort.
- Ignorieren Sie die Interpunktion im gegebenen Absatz. Interpunktion besteht aus Kommas (`,`), Punkten (`.`) und Ausrufezeichen (`!`).
- Das zurückgegebene Array sollte alle Wörter in Kleinbuchstaben enthalten.
- Das zurückgegebene Array sollte in absteigender Reihenfolge sein, wobei das am häufigsten vorkommende Wort zuerst steht.

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` sollte `["coding", "python", "in"]` zurückgeben.

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` sollte `["i", "like", "coding"]` zurückgeben.

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` sollte `["debug", "test", "deploy"]` zurückgeben.

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
