---
id: 68b06e589bf2273243814773
title: "Aufgabe 37: Satzanfang großschreiben"
challengeType: 28
dashedName: challenge-37
---

# --description--

Gegeben ist ein Absatz. Geben Sie einen neuen Absatz zurück, bei dem der erste Buchstabe jedes Satzes großgeschrieben ist.

- Alle anderen Zeichen sollen unverändert bleiben.
- Sätze können mit einem Punkt (`.`), einem oder mehreren Fragezeichen (`?`) oder einem oder mehreren Ausrufezeichen (`!`) enden.

# --hints--

`capitalize("this is a simple sentence.")` sollte `"This is a simple sentence."` zurückgeben.

```js
assert.equal(capitalize("this is a simple sentence."), "This is a simple sentence.");
```

`capitalize("hello world. how are you?")` sollte `"Hello world. How are you?"` zurückgeben.

```js
assert.equal(capitalize("hello world. how are you?"), "Hello world. How are you?");
```

`capitalize("i did today's coding challenge... it was fun!!")` sollte `"I did today's coding challenge... It was fun!!"` zurückgeben.

```js
assert.equal(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!");
```

`capitalize("crazy!!!strange???unconventional...sentences.")` sollte `"Crazy!!!Strange???Unconventional...Sentences."` zurückgeben.

```js
assert.equal(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.");
```

`capitalize("there's a space before this period . why is there a space before that period ?")` sollte `"There's a space before this period . Why is there a space before that period ?"` zurückgeben.

```js
assert.equal(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?");
```

# --seed--

## --seed-contents--

```js
function capitalize(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function capitalize(paragraph) {
  return paragraph.replace(/([.!?]+)(\s*)([a-z])/g, (match, punc, spaces, char) => {
    return punc + spaces + char.toUpperCase();
  })
  .replace(/^([a-z])/, (match, char) => char.toUpperCase());
}
```
