---
id: 68b06e589bf2273243814775
title: "Aufgabe 38: Slug-Generator"
challengeType: 28
dashedName: challenge-38
---

# --description--

Gegeben ist ein String. Geben Sie eine URL-freundliche Version des Strings zurück, die folgende Bedingungen erfüllt:

- Alle Buchstaben sollen Kleinbuchstaben sein.
- Alle Zeichen, die keine Buchstaben, Zahlen oder Leerzeichen sind, sollen entfernt werden.
- Alle Leerzeichen sollen durch den URL-kodierten Leerzeichen-Code `%20` ersetzt werden.
- Aufeinanderfolgende Leerzeichen sollen durch ein einzelnes `%20` ersetzt werden.
- Der zurückgegebene String darf keine führenden oder abschließenden `%20` enthalten.

# --hints--

`generateSlug("helloWorld")` sollte `"helloworld"` zurückgeben.

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` sollte `"hello%20world"` zurückgeben.

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` sollte `"helloworld"` zurückgeben.

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` sollte `"hello%20world"` zurückgeben.

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` sollte `"h3110%20w0r1d"` zurückgeben.

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
