---
id: 68b1f72371a5ac895ac70a06
title: "Aufgabe 42: Videos speichern"
challengeType: 28
dashedName: challenge-42
---

# --description--

Gegeben sind eine Videogröße, eine Einheit für die Videogröße, eine Festplattenkapazität und eine Einheit für die Festplatte. Geben Sie die Anzahl der Videos zurück, die auf der Festplatte gespeichert werden können, unter Verwendung der folgenden Vorgaben:

- Die Einheit für die Videogröße kann Bytes (`"B"`), Kilobytes (`"KB"`), Megabytes (`"MB"`) oder Gigabytes (`"GB"`) sein.
- Wenn keine der oben genannten Videoeinheiten angegeben ist, geben Sie `"Invalid video unit"` zurück.
- Die Einheit der Festplattenkapazität kann Gigabytes (`"GB"`) oder Terabytes (`"TB"`) sein.
- Wenn keine der oben genannten Festplatteneinheiten angegeben ist, geben Sie `"Invalid drive unit"` zurück.
- Geben Sie die Anzahl der ganzen Videos zurück, die auf die Festplatte passen.
- Verwenden Sie die folgenden Umrechnungen:

| Einheit | Entspricht |
|:-------:|:---------:|
| 1 B     | 1 B       |
| 1 KB    | 1000 B    |
| 1 MB    | 1000 KB   |
| 1 GB    | 1000 MB   |
| 1 TB    | 1000 GB   |

Beispielsweise bestimmen Sie bei den Argumenten `500`, `"MB"`, `100` und `"GB"`, wie viele 500 MB Videos auf eine 100 GB Festplatte passen.

# --hints--

`numberOfVideos(500, "MB", 100, "GB")` sollte `200` zurückgeben.

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")` sollte `"Invalid video unit"` zurückgeben.

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")` sollte `"Invalid drive unit"` zurückgeben.

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")` sollte `4000` zurückgeben.

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")` sollte `1466` zurückgeben.

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```
