---
id: 68b1f72371a5ac895ac70a08
title: "Aufgabe 43: Ziffern vs Buchstaben"
challengeType: 28
dashedName: challenge-43
---

# --description--

Gegeben ist ein String. Geben Sie `"digits"` zurück, wenn der String mehr Ziffern als Buchstaben enthält, `"letters"`, wenn er mehr Buchstaben als Ziffern enthält, und `"tie"`, wenn er gleich viele Ziffern und Buchstaben enthält.

- Ziffern bestehen aus `0-9`.
- Buchstaben bestehen aus `a-z` in Groß- oder Kleinbuchstaben.
- Ignorieren Sie alle anderen Zeichen.

# --hints--

`digitsOrLetters("abc123")` sollte `"tie"` zurückgeben.

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` sollte `"letters"` zurückgeben.

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` sollte `"digits"` zurückgeben.

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` sollte `"letters"` zurückgeben.

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` sollte `"digits"` zurückgeben.

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` sollte `"tie"` zurückgeben.

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
