---
id: 68b1f72371a5ac895ac70a02
title: "Aufgabe 40: Fotos speichern"
challengeType: 28
dashedName: challenge-40
---

# --description--

Gegeben ist die Größe eines Fotos in Megabyte (MB) und die Kapazität der Festplatte in Gigabyte (GB). Geben Sie die Anzahl der Fotos zurück, die die Festplatte unter den folgenden Bedingungen speichern kann:

- 1 Gigabyte entspricht 1000 Megabyte.
- Geben Sie die Anzahl der ganzen Fotos zurück, die die Festplatte speichern kann.

# --hints--

`numberOfPhotos(1, 1)` sollte `1000` zurückgeben.

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

`numberOfPhotos(2, 1)` sollte `500` zurückgeben.

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

`numberOfPhotos(4, 256)` sollte `64000` zurückgeben.

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

`numberOfPhotos(3.5, 750)` sollte `214285` zurückgeben.

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

`numberOfPhotos(3.5, 5.5)` sollte `1571` zurückgeben.

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```
