---
id: 68b1f72371a5ac895ac70a0a
title: "Aufgabe 44: String-Spiegel"
challengeType: 28
dashedName: challenge-44
---

# --description--

Gegeben sind zwei Strings. Bestimmen Sie, ob der zweite String ein Spiegelbild des ersten ist.

- Ein String gilt als Spiegelbild, wenn er dieselben Buchstaben in umgekehrter Reihenfolge enthält.
- Groß- und Kleinbuchstaben werden als unterschiedlich behandelt.
- Alle nicht alphabetischen Zeichen werden ignoriert.

# --hints--

`isMirror("helloworld", "helloworld")` sollte `false` zurückgeben.

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` sollte `true` zurückgeben.

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` sollte `true` zurückgeben.

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` sollte `false` zurückgeben.

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` sollte `false` zurückgeben.

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` sollte `true` zurückgeben.

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` sollte `true` zurückgeben.

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
