---
id: 68b1f72371a5ac895ac70a04
title: "Aufgabe 41: Dateispeicherung"
challengeType: 28
dashedName: challenge-41
---

# --description--

Gegeben sind eine Dateigröße, eine Einheit für die Dateigröße und die Kapazität der Festplatte in Gigabyte (GB). Geben Sie die Anzahl der Dateien zurück, die die Festplatte unter den folgenden Bedingungen speichern kann:

- Die Einheit der Dateigröße kann Bytes (`"B"`), Kilobytes (`"KB"`) oder Megabytes (`"MB"`) sein.
- Geben Sie die Anzahl der ganzen Dateien zurück, die auf die Festplatte passen.
- Verwenden Sie die folgenden Umrechnungen:

| Einheit | Entspricht |
|:-------:|:---------:|
| 1 B     |   1 B     |
| 1 KB    |  1000 B   |
| 1 MB    |  1000 KB  |
| 1 GB    |  1000 MB  |

Beispielsweise bestimmen Sie bei den Argumenten `500`, `"KB"` und `1`, wie viele 500 KB große Dateien auf eine 1 GB Festplatte passen.

# --hints--

`numberOfFiles(500, "KB", 1)` sollte `2000` zurückgeben.

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)` sollte `20000` zurückgeben.

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)` sollte `200` zurückgeben.

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)` sollte `366210` zurückgeben.

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)` sollte `453514` zurückgeben.

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)` sollte `166666` zurückgeben.

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
