---
id: 68b7687dded630607aceccad
title: "Aufgabe 46: Zweitgrößte Zahl"
challengeType: 28
dashedName: challenge-46
---

# --description--

Gegeben ist ein Array. Geben Sie die zweitgrößte unterschiedliche Zahl zurück.

# --hints--

`secondLargest([1, 2, 3, 4])` sollte `3` zurückgeben.

```js
assert.equal(secondLargest([1, 2, 3, 4]), 3);
```

`secondLargest([20, 139, 94, 67, 31])` sollte `94` zurückgeben.

```js
assert.equal(secondLargest([20, 139, 94, 67, 31]), 94);
```

`secondLargest([2, 3, 4, 6, 6])` sollte `4` zurückgeben.

```js
assert.equal(secondLargest([2, 3, 4, 6, 6]), 4);
```

`secondLargest([10, -17, 55.5, 44, 91, 0])` sollte `55.5` zurückgeben.

```js
assert.equal(secondLargest([10, -17, 55.5, 44, 91, 0]), 55.5);
```

`secondLargest([1, 0, -1, 0, 1, 0, -1, 1, 0])` sollte `0` zurückgeben.

```js
assert.equal(secondLargest([1, 0, -1, 0, 1, 0, -1, 1, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function secondLargest(arr) {

  return arr;
}
```

# --solutions--

```js
function secondLargest(arr) {
  const unique = [...new Set(arr)];
  unique.sort((a, b) => b - a);
  return unique[1];
}
```
