---
id: 68b7687dded630607aceccab
title: "Aufgabe 45: Perfektes Quadrat"
challengeType: 28
dashedName: challenge-45
---

# --description--

Gegeben ist eine ganze Zahl. Bestimmen Sie, ob sie ein perfektes Quadrat ist.

- Eine Zahl ist ein perfektes Quadrat, wenn Sie eine ganze Zahl mit sich selbst multiplizieren können, um diese Zahl zu erhalten. Zum Beispiel ist 9 ein perfektes Quadrat, weil Sie 3 mit sich selbst multiplizieren können, um 9 zu erhalten.

# --hints--

`isPerfectSquare(9)` sollte `true` zurückgeben.

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` sollte `true` zurückgeben.

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` sollte `true` zurückgeben.

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` sollte `false` zurückgeben.

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` sollte `false` zurückgeben.

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` sollte `false` zurückgeben.

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` sollte `true` zurückgeben.

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` sollte `true` zurückgeben.

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
