---
id: 68b7687dded630607aceccaf
title: "Aufgabe 47: Beim Rasen erwischt"
challengeType: 28
dashedName: challenge-47
---

# --description--

Gegeben ist ein Array von Zahlen, die die Geschwindigkeit darstellen, mit der Fahrzeuge beobachtet wurden, sowie eine Zahl, die das Tempolimit angibt. Geben Sie ein Array mit zwei Elementen zurück: die Anzahl der Fahrzeuge, die zu schnell gefahren sind, gefolgt vom durchschnittlichen Wert, um den diese Fahrzeuge das Tempolimit überschritten haben.

- Wenn keine Fahrzeuge zu schnell gefahren sind, geben Sie `[0, 0]` zurück.

# --hints--

`speeding([50, 60, 55], 60)` sollte `[0, 0]` zurückgeben.

```js
assert.deepEqual(speeding([50, 60, 55], 60), [0, 0]);
```

`speeding([58, 50, 60, 55], 55)` sollte `[2, 4]` zurückgeben.

```js
assert.deepEqual(speeding([58, 50, 60, 55], 55), [2, 4]);
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)` sollte `[4, 8.5]` zurückgeben.

```js
assert.deepEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5]);
```

`speeding([100, 105, 95, 102], 100)` sollte `[2, 3.5]` zurückgeben.

```js
assert.deepEqual(speeding([100, 105, 95, 102], 100), [2, 3.5]);
```

`speeding([40, 45, 44, 50, 112, 39], 55)` sollte `[1, 57]` zurückgeben.

```js
assert.deepEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57]);
```

# --seed--

## --seed-contents--

```js
function speeding(speeds, limit) {

  return speeds;
}
```

# --solutions--

```js
function speeding(speeds, limit) {
  let speeding = 0;
  let totalOver = 0;

  for (let speed of speeds) {
    if (speed > limit) {
      speeding++;
      totalOver += (speed - limit);
    }
  }

  if (speeding === 0) return [0, 0];

  return [speeding, totalOver / speeding];
}
```
