---
id: 68b7cadffed0e75a517da675
title: "Aufgabe 53: Dezimal zu Binär"
challengeType: 28
dashedName: challenge-53
---

# --description--

Gegeben ist eine nicht-negative ganze Zahl. Geben Sie ihre binäre Darstellung als String zurück.

Eine Binärzahl verwendet nur die Ziffern `0` und `1`, um eine Zahl darzustellen. Um eine Dezimalzahl in Binär umzuwandeln, teilen Sie die Zahl wiederholt durch `2` und notieren Sie den Rest. Wiederholen Sie dies, bis die Zahl null ist. Lesen Sie die notierten Reste von zuletzt zu zuerst. Zum Beispiel, um `12` in Binär umzuwandeln:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` in Binär ist `1100`.

# --hints--

`toBinary(5)` sollte `"101"` zurückgeben.

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)` sollte `"1100"` zurückgeben.

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)` sollte `"110010"` zurückgeben.

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)` sollte `"1100011"` zurückgeben.

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
