---
id: 68b7687dded630607aceccb3
title: "Aufgabe 49: CSV-Kopfzeilenparser"
challengeType: 28
dashedName: challenge-49
---

# --description--

Gegeben ist die erste Zeile einer durch Kommas getrennten Werte-Datei (CSV). Geben Sie ein Array zurück, das die Überschriften enthält.

- Die erste Zeile einer CSV-Datei enthält Überschriften, die durch Kommas getrennt sind.
- Entfernen Sie alle führenden oder nachfolgenden Leerzeichen von jeder Überschrift.

# --hints--

`getHeadings("name,age,city")` sollte `["name", "age", "city"]` zurückgeben.

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` sollte `["first name", "last name", "phone"]` zurückgeben.

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` sollte `["username", "email", "signup date"]` zurückgeben.

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
