---
id: 68b7687dded630607aceccb1
title: "Aufgabe 48: Spam-Erkennung"
challengeType: 28
dashedName: challenge-48
---

# --description--

Gegeben ist eine Telefonnummer im Format `"+A (BBB) CCC-DDDD"`, wobei jeder Buchstabe eine Ziffer wie folgt darstellt:

- `A` steht für die Landesvorwahl und kann eine beliebige Anzahl von Ziffern enthalten.
- `BBB` steht für die Vorwahl und besteht immer aus drei Ziffern.
- `CCC` und `DDDD` stehen für die lokale Nummer und haben jeweils immer drei bzw. vier Ziffern.

Bestimmen Sie, ob es sich um eine Spam-Nummer handelt, basierend auf den folgenden Kriterien:

- Die Landesvorwahl ist länger als zwei Ziffern oder beginnt nicht mit einer Null (`0`).
- Die Vorwahl ist größer als 900 oder kleiner als 200.
- Die Summe der ersten drei Ziffern der lokalen Nummer erscheint innerhalb der letzten vier Ziffern der lokalen Nummer.
- Die Nummer enthält dieselbe Ziffer vier- oder mehrmals hintereinander (ohne Berücksichtigung der Formatierungszeichen).

# --hints--

`isSpam("+0 (200) 234-0182")` sollte `false` zurückgeben.

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` sollte `true` zurückgeben.

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` sollte `true` zurückgeben.

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` sollte `true` zurückgeben.

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` sollte `true` zurückgeben.

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` sollte `true` zurückgeben.

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` sollte `true` zurückgeben.

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` sollte `false` zurückgeben.

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
