---
id: 68b7cadffed0e75a517da677
title: "Aufgabe 54: P@ssw0rd-Str3ngth!"
challengeType: 28
dashedName: challenge-54
---

# --description--

Gegeben ist ein Passwort-String. Geben Sie `"weak"`, `"medium"` oder `"strong"` zurück, basierend auf der Stärke des Passworts.

Ein Passwort wird nach folgenden Regeln bewertet:

- Es ist mindestens 8 Zeichen lang.
- Es enthält sowohl Groß- als auch Kleinbuchstaben.
- Es enthält mindestens eine Zahl.
- Es enthält mindestens ein Sonderzeichen aus dieser Menge: `!`, `@`, `#`, `$`, `%`, `^`, `&` oder `*`.

Geben Sie `"weak"` zurück, wenn das Passwort weniger als zwei der Regeln erfüllt.
Geben Sie `"medium"` zurück, wenn das Passwort 2 oder 3 der Regeln erfüllt.
Geben Sie `"strong"` zurück, wenn das Passwort alle 4 Regeln erfüllt.

# --hints--

`checkStrength("123456")` sollte `"weak"` zurückgeben.

```js
assert.equal(checkStrength("123456"), "weak");
```

`checkStrength("pass!!!")` sollte `"weak"` zurückgeben.

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

`checkStrength("Qwerty")` sollte `"weak"` zurückgeben.

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

`checkStrength("PASSWORD")` sollte `"weak"` zurückgeben.

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

`checkStrength("PASSWORD!")` sollte `"medium"` zurückgeben.

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

`checkStrength("PassWord%^!")` sollte `"medium"` zurückgeben.

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

`checkStrength("qwerty12345")` sollte `"medium"` zurückgeben.

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

`checkStrength("S3cur3P@ssw0rd")` sollte `"strong"` zurückgeben.

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

`checkStrength("C0d3&Fun!")` sollte `"strong"` zurückgeben.

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```
