---
id: 68b7cadffed0e75a517da673
title: "Aufgabe 52: Binär zu Dezimal"
challengeType: 28
dashedName: challenge-52
---

# --description--

Gegeben ist ein String, der eine Binärzahl darstellt. Geben Sie deren dezimale Entsprechung als Zahl zurück.

Eine Binärzahl verwendet nur die Ziffern `0` und `1`, um eine beliebige Zahl darzustellen. Um Binärzahlen in Dezimalzahlen umzuwandeln, multiplizieren Sie jede Ziffer mit einer Potenz von `2` und addieren diese zusammen. Beginnen Sie damit, die ganz rechts stehende Ziffer mit `2^0` zu multiplizieren, die nächste Ziffer links davon mit `2^1` und so weiter. Sobald alle Ziffern mit einer Potenz von `2` multipliziert wurden, addieren Sie das Ergebnis zusammen.

Zum Beispiel entspricht die Binärzahl `101` der Dezimalzahl `5`, weil:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")` sollte `5` zurückgeben.

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")` sollte `10` zurückgeben.

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")` sollte `18` zurückgeben.

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")` sollte `85` zurückgeben.

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
