---
id: 68c1a929005bf54d342aa8d5
title: "Aufgabe 58: Space Week Tag 4: Landeplatz"
challengeType: 28
dashedName: challenge-58
---

# --description--

Am vierten Tag der Space Week erhalten Sie eine Matrix von Zahlen (ein Array von Arrays), die potenzielle Landeplätze für Ihren Rover darstellen. Finden Sie den sichersten Landeplatz basierend auf den folgenden Regeln:

- Jeder Platz in der Matrix enthält eine Zahl ab `0-9`, inklusive.
- Jede `0` stellt einen potenziellen Landeplatz dar.
- Jede Zahl außer `0` ist zu gefährlich zum Landen. Je höher die Zahl, desto gefährlicher.
- Der sicherste Platz ist definiert als die `0`-Zelle, deren umliegende Zellen (bis zu 4 Nachbarn, Diagonalen ignorieren) die geringste Gesamtrisiko haben.
- Ignorieren Sie Nachbarn außerhalb der Grenzen (Ecken und Kanten haben einfach weniger Nachbarn).
- Geben Sie die Indizes des sichersten Landeplatzes zurück. Es wird immer nur einen sichersten Platz geben.

Zum Beispiel, gegeben:

```js
[
  [1, 0],
  [2, 0]
]
```

Geben Sie `[0, 1]` zurück, die Indizes für die `0` im ersten Array.

# --hints--

`findLandingSpot([[1, 0], [2, 0]])` sollte `[0, 1]` zurückgeben.

```js
assert.deepEqual(findLandingSpot([[1, 0], [2, 0]]), [0, 1]);
```

`findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` sollte `[1, 1]` zurückgeben.

```js
assert.deepEqual(findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1]);
```

`findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` sollte `[2, 2]` zurückgeben.

```js
assert.deepEqual(findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2]);
```

`findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` sollte `[2, 1]` zurückgeben.

```js
assert.deepEqual(findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1]);
```

# --seed--

## --seed-contents--

```js
function findLandingSpot(matrix) {

  return matrix;
}
```

# --solutions--

```js
function findLandingSpot(matrix) {
  let bestSpot = null;
  let lowestNeighborSum = Infinity;

  for (let i = 0; i < matrix.length; i++) {
    for (let j = 0; j < matrix[i].length; j++) {

      if (matrix[i][j] === 0) {
        let currentNeighborSum = 0;

        if (i > 0) currentNeighborSum += matrix[i - 1][j];
        if (j < matrix[i].length - 1) currentNeighborSum += matrix[i][j + 1];
        if (i < matrix.length - 1) currentNeighborSum += matrix[i + 1][j];
        if (j > 0) currentNeighborSum += matrix[i][j - 1];

        if (currentNeighborSum < lowestNeighborSum) {
          lowestNeighborSum = currentNeighborSum;
          bestSpot = [i, j];
        }
      }
    }
  }

  return bestSpot;
}
```
