---
id: 68c1a929005bf54d342aa8d4
title: "Aufgabe 57: Space Week Tag 3: Phone Home"
challengeType: 28
dashedName: challenge-57
---

# --description--

Für den dritten Tag der Space Week erhalten Sie ein Array von Zahlen, die Entfernungen (in Kilometern) zwischen Ihnen, Satelliten und Ihrem Heimatplaneten in einer Kommunikationsroute darstellen. Bestimmen Sie, wie lange eine Nachricht, die durch die Route gesendet wird, benötigt, um ihren Zielplaneten zu erreichen, unter den folgenden Bedingungen:

- Der erste Wert im Array ist die Entfernung von Ihrem Standort zum ersten Satelliten.
- Jeder folgende Wert, außer dem letzten, ist die Entfernung zum nächsten Satelliten.
- Der letzte Wert im Array ist die Entfernung vom vorherigen Satelliten zu Ihrem Heimatplaneten.
- Die Nachricht reist mit 300.000 km/s.
- Jeder Satellit, den die Nachricht **durchläuft**, fügt eine Übertragungsverzögerung von 0,5 Sekunden hinzu.
- Geben Sie eine Zahl zurück, die auf 4 Dezimalstellen gerundet ist, wobei nachfolgende Nullen entfernt werden.

# --hints--

`sendMessage([300000, 300000])` sollte `2.5` zurückgeben.

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` sollte `3.0627` zurückgeben.

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` sollte `364.5` zurückgeben.

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` sollte `1674.3333` zurückgeben.

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` sollte `2.4086` zurückgeben.

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` sollte `21.1597` zurückgeben.

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
