---
id: 68c1a929005bf54d342aa8d2
title: "Aufgabe 55: Space Week Tag 1: Sternklassifikation"
challengeType: 28
dashedName: challenge-55
---

# --description--

Der 4. Oktober markiert den Beginn der World Space Week. Die nächsten sieben Tage erwarten Sie astronomisch inspirierte Programmieraufgaben.

Für die heutige Aufgabe erhalten Sie die Oberflächentemperatur eines Sterns in Kelvin (K) und sollen dessen Sternklassifikation anhand der folgenden Bereiche bestimmen:

- `"O"`: 30.000 K oder höher
- `"B"`: 10.000 K – 29.999 K
- `"A"`: 7.500 K – 9.999 K
- `"F"`: 6.000 K – 7.499 K
- `"G"`: 5.200 K – 5.999 K
- `"K"`: 3.700 K – 5.199 K
- `"M"`: 0 K – 3.699 K

- Geben Sie die Klassifikation des gegebenen Sterns zurück.

# --hints--

`classification(5778)` sollte `"G"` zurückgeben.

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` sollte `"M"` zurückgeben.

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` sollte `"A"` zurückgeben.

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` sollte `"K"` zurückgeben.

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` sollte `"M"` zurückgeben.

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` sollte `"O"` zurückgeben.

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` sollte `"F"` zurückgeben.

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` sollte `"B"` zurückgeben.

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
