---
id: 68c1a929005bf54d342aa8d6
title: "Aufgabe 59: Space Week Tag 5: Goldilocks-Zone"
challengeType: 28
dashedName: challenge-59
---

# --description--

Am fünften Tag der Space Week berechnen Sie die „Goldilocks-Zone“ eines Sterns – den Bereich um einen Stern, in dem die Bedingungen „genau richtig“ sind, damit flüssiges Wasser existieren kann.

Gegeben ist die Masse eines Sterns. Geben Sie ein Array mit den Start- und Endabständen seiner Goldilocks-Zone in Astronomischen Einheiten zurück.

Zur Berechnung der Goldilocks-Zone:

1. Bestimmen Sie die Leuchtkraft des Sterns, indem Sie seine Masse hoch 3,5 potenzieren.
2. Der Beginn der Zone ist 0,95 mal die Quadratwurzel seiner Leuchtkraft.
3. Das Ende der Zone ist 1,37 mal die Quadratwurzel seiner Leuchtkraft.

- Geben Sie die Abstände auf zwei Dezimalstellen gerundet zurück.

Zum Beispiel: Bei einer Masse von `1` geben Sie `[0.95, 1.37]` zurück.

# --hints--

`goldilocksZone(1)` sollte `[0.95, 1.37]` zurückgeben.

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)` sollte `[0.28, 0.41]` zurückgeben.

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)` sollte `[21.85, 31.51]` zurückgeben.

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)` sollte `[9.38, 13.52]` zurückgeben.

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)` sollte `[179.69, 259.13]` zurückgeben.

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
