---
id: 68c497f3aaefc9fd9f1b0e24
title: "Aufgabe 60: Space Week Tag 6: Mondphase"
challengeType: 28
dashedName: challenge-60
---

# --description--

Für den sechsten Tag der Space Week erhalten Sie ein Datum im Format `"YYYY-MM-DD"` und müssen die Mondphase für diesen Tag anhand der folgenden Regeln bestimmen:

Verwenden Sie einen vereinfachten Mondzyklus von 28 Tagen, der in vier gleiche Phasen unterteilt ist:

- `"New"`: Tage 1 - 7
- `"Waxing"`: Tage 8 - 14
- `"Full"`: Tage 15 - 21
- `"Waning"`: Tage 22 - 28

Nach Tag 28 wiederholt sich der Zyklus mit Tag 1, Neumond.

- Verwenden Sie `"2000-01-06"` als Referenz-Neumond (Tag 1 des Zyklus), um die Phase des angegebenen Tages zu bestimmen.
- Sie erhalten keine Daten vor dem Referenzdatum.
- Geben Sie die korrekte Phase als String zurück.

**Hinweis:** Tag 1 stellt den Tag des Neumonds dar, das heißt, es sind 0 Tage seit dem letzten Neumond vergangen.

# --hints--

`moonPhase("2000-01-12")` sollte `"New"` zurückgeben.

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` sollte `"Waxing"` zurückgeben.

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` sollte `"Full"` zurückgeben.

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` sollte `"Waning"` zurückgeben.

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` sollte `"New"` zurückgeben.

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
