---
id: 68c497f3aaefc9fd9f1b0e26
title: "Aufgabe 62: Hexadezimal zu Dezimal"
challengeType: 28
dashedName: challenge-62
---

# --description--

Gegeben ist ein String, der eine hexadezimale Zahl (Basis 16) darstellt. Geben Sie ihren dezimalen (Basis 10) Wert als Ganzzahl zurück.

Hexadezimal ist ein Zahlensystem, das 16 Ziffern verwendet:

- `0-9` stellen die Werte `0` bis `9` dar.
- `A-F` stellen die Werte `10` bis `15` dar.

Hier ist eine teilweise Umrechnungstabelle:

| Hexadezimal | Dezimal |
|:-----------:|:-------:|
| 0           | 0       |
| 1           | 1       |
| ...         | ...     |
| 9           | 9       |
| A           | 10      |
| ...         | ...     |
| F           | 15      |
| 10          | 16      |
| ...         | ...     |
| 9F          | 159     |
| A0          | 160     |
| ...         | ...     |
| FF          | 255     |
| 100         | 256     |

- Der String enthält nur die Zeichen `0–9` und `A–F`.

# --hints--

`hexToDecimal("A")` sollte `10` zurückgeben.

```js
assert.equal(hexToDecimal("A"), 10);
```

`hexToDecimal("15")` sollte `21` zurückgeben.

```js
assert.equal(hexToDecimal("15"), 21);
```

`hexToDecimal("2E")` sollte `46` zurückgeben.

```js
assert.equal(hexToDecimal("2E"), 46);
```

`hexToDecimal("FF")` sollte `255` zurückgeben.

```js
assert.equal(hexToDecimal("FF"), 255);
```

`hexToDecimal("A3F")` sollte `2623` zurückgeben.

```js
assert.equal(hexToDecimal("A3F"), 2623);
```

# --seed--

## --seed-contents--

```js
function hexToDecimal(hex) {

  return hex;
}
```

# --solutions--

```js
function hexToDecimal(hex) {

  return parseInt(hex, 16);
}
```
