---
id: 68c497f3aaefc9fd9f1b0e25
title: "Aufgabe 61: Space Week Tag 7: Starttreibstoff"
challengeType: 28
dashedName: challenge-61
---

# --description--

Für den letzten Tag der Space Week erhalten Sie die Masse in Kilogramm (kg) einer Nutzlast, die Sie in den Orbit schicken möchten. Bestimmen Sie die Menge an Treibstoff, die benötigt wird, um Ihre Nutzlast mit den folgenden Regeln in den Orbit zu bringen:

- Raketen benötigen 1 kg Treibstoff pro 5 kg Masse, die sie heben müssen.
- Treibstoff hat selbst Masse. Wenn Sie also Treibstoff hinzufügen, steigt die zu hebende Masse, was mehr Treibstoff erfordert, was die Masse erhöht, und so weiter.
- Um die gesamte benötigte Treibstoffmenge zu berechnen: Beginnen Sie mit der Nutzlastmasse, berechnen Sie den dafür benötigten Treibstoff, addieren Sie diesen Treibstoff zur Gesamtmasse und berechnen Sie erneut. Wiederholen Sie diesen Vorgang, bis der zusätzlich benötigte Treibstoff weniger als 1 kg beträgt, dann stoppen Sie.
- Ignorieren Sie die Masse der Rakete selbst. Berechnen Sie nur den Treibstoff, der benötigt wird, um die Nutzlast und ihren eigenen Treibstoff zu heben.

Zum Beispiel: Bei einer Nutzlastmasse von 50 kg benötigen Sie 10 kg Treibstoff, um sie zu heben (Nutzlast / 5), was die Gesamtmasse auf 60 kg erhöht, die 12 kg zum Heben benötigt (2 zusätzliche kg), was die Gesamtmasse auf 62 kg erhöht, die 12,4 kg zum Heben benötigt – 0,4 zusätzliche kg – was weniger als 1 zusätzliche kg ist, also stoppen wir hier. Die Gesamtmasse zum Heben beträgt 62,4 kg, davon sind 50 kg die ursprüngliche Nutzlast und 12,4 kg Treibstoff.

- Geben Sie die benötigte Treibstoffmenge auf eine Nachkommastelle gerundet zurück.

# --hints--

`launchFuel(50)` sollte `12.4` zurückgeben.

```js
assert.equal(launchFuel(50), 12.4);
```

`launchFuel(500)` sollte `124.8` zurückgeben.

```js
assert.equal(launchFuel(500), 124.8);
```

`launchFuel(243)` sollte `60.7` zurückgeben.

```js
assert.equal(launchFuel(243), 60.7);
```

`launchFuel(11000)` sollte `2749.8` zurückgeben.

```js
assert.equal(launchFuel(11000), 2749.8);
```

`launchFuel(6214)` sollte `1553.4` zurückgeben.

```js
assert.equal(launchFuel(6214), 1553.4);
```

# --seed--

## --seed-contents--

```js
function launchFuel(payload) {

  return payload;
}
```

# --solutions--

```js
function launchFuel(payload) {
  let totalMass = payload;
  let additionalFuel = totalMass / 5;
  let totalFuel = additionalFuel;

  while (additionalFuel >= 1) {
    totalMass += additionalFuel;
    additionalFuel = (totalMass / 5) - totalFuel;
    totalFuel += additionalFuel;
  }

  return Number(totalFuel.toFixed(1));
}
```
