---
id: 68cae5b538ff798bbd4da008
title: "Aufgabe 70: HTML-Attribut-Extraktor"
challengeType: 28
dashedName: challenge-70
---

# --description--

Gegeben ist ein String mit einem gültigen HTML-Element. Geben Sie die Attribute des Elements nach den folgenden Kriterien zurück:

- Es wird nur ein Element übergeben.
- Attribute liegen im Format `attribute="value"` vor.
- Geben Sie ein Array von Strings zurück, wobei jede Eigenschaft und ihr Wert durch ein Komma getrennt sind, im Format `["attribute1, value1", "attribute2, value2"]`.
- Geben Sie die Attribute in der Reihenfolge zurück, in der sie angegeben sind.
- Wenn keine Attribute gefunden werden, geben Sie ein leeres Array zurück.

# --hints--

`extractAttributes('<span class="red"></span>')` sollte `["class, red"]` zurückgeben.

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` sollte `["charset, UTF-8"]` zurückgeben.

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` sollte `[]` zurückgeben.

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` sollte `["name, email", "type, email", "required, true"]` zurückgeben.

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` sollte `["id, submit", "class, btn btn-primary"]` zurückgeben.

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
