---
id: 68cae5b538ff798bbd4da004
title: "Aufgabe 66: HTML-Tag-Entferner"
challengeType: 28
dashedName: challenge-66
---

# --description--

Gegeben ist ein String mit HTML-Code. Entfernen Sie die Tags und geben Sie den reinen Textinhalt zurück.

- Der Eingabe-String enthält nur gültiges HTML.
- HTML-Tags können verschachtelt sein.
- Entfernen Sie die Tags und alle Attribute.

Zum Beispiel sollte `'<a href="#">Click here</a>'` `"Click here"` zurückgeben.

# --hints--

`stripTags('<a href="#">Click here</a>')` sollte `"Click here"` zurückgeben.

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')` sollte `"Hello World!"` zurückgeben.

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')` sollte einen leeren String (`""`) zurückgeben.

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` sollte `sectionsection` zurückgeben.

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
