---
id: 68cae5b538ff798bbd4da001
title: "Aufgabe 63: Wortgefecht"
challengeType: 28
dashedName: challenge-63
---

# --description--

Gegeben sind zwei Sätze, die Ihr Team und ein gegnerisches Team darstellen, wobei jedes Wort Ihres Teams gegen das entsprechende Wort des gegnerischen Teams antritt. Bestimmen Sie, welches Team gewinnt, anhand der folgenden Regeln:

- Die gegebenen Sätze enthalten immer die gleiche Anzahl an Wörtern.
- Wörter sind durch ein einzelnes Leerzeichen getrennt und bestehen nur aus Buchstaben.
- Der Wert jedes Wortes ist die Summe seiner Buchstaben.
- Die Buchstaben `a` bis `z` entsprechen den Werten `1` bis `26`. Zum Beispiel ist `a` `1`, und `z` ist `26`.
- Ein Großbuchstabe verdoppelt den Wert des Buchstabens. Zum Beispiel ist `A` `2`, und `Z` ist `52`.
- Die Wörter treten der Reihe nach gegeneinander an: das erste Wort Ihres Teams gegen das erste Wort des gegnerischen Teams, und so weiter.
- Ein Wort gewinnt, wenn sein Wert größer ist als der Wert des gegnerischen Wortes.
- Das Team mit mehr gewonnenen Wörtern ist der Sieger.

Geben Sie `"We win"` zurück, wenn Ihr Team gewinnt, `"We lose"`, wenn Ihr Team verliert, und `"Draw"`, wenn beide Teams gleich viele Siege haben.

# --hints--

`battle("hello world", "hello word")` sollte `"We win"` zurückgeben.

```js
assert.equal(battle("hello world", "hello word"), "We win");
```

`battle("Hello world", "hello world")` sollte `"We win"` zurückgeben.

```js
assert.equal(battle("Hello world", "hello world"), "We win");
```

`battle("lorem ipsum", "kitty ipsum")` sollte `"We lose"` zurückgeben.

```js
assert.equal(battle("lorem ipsum", "kitty ipsum"), "We lose");
```

`battle("hello world", "world hello")` sollte `"Draw"` zurückgeben.

```js
assert.equal(battle("hello world", "world hello"), "Draw");
```

`battle("git checkout", "git switch")` sollte `"We win"` zurückgeben.

```js
assert.equal(battle("git checkout", "git switch"), "We win");
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` sollte `"We lose"` zurückgeben.

```js
assert.equal(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose");
```

`battle("We must never surrender", "Our team must win")` sollte `"Draw"` zurückgeben.

```js
assert.equal(battle("We must never surrender", "Our team must win"), "Draw");
```

# --seed--

## --seed-contents--

```js
function battle(ourTeam, opponent) {

  return ourTeam;
}
```

# --solutions--

```js
function getCharacterValue(char) {
  const lowerCaseValue = char.toLowerCase().charCodeAt(0) - 96;
  return char >= 'A' && char <= 'Z' ? lowerCaseValue * 2 : lowerCaseValue;
}

function getWordValue(word) {
  let wordValue = 0;
  for (let i=0; i<word.length; i++) {
    wordValue += getCharacterValue(word[i])
  }

  return wordValue;
}

function battle(ourTeam, opponent) {
  const myWords = ourTeam.split(' ');
  const opponentWords = opponent.split(' ');

  let myWins = 0, opponentWins = 0;

  for (let i=0; i<myWords.length; i++) {
    const myWordValue = getWordValue(myWords[i]);
    const opponentWordValue = getWordValue(opponentWords[i])

    if (myWordValue > opponentWordValue) myWins++;
    if (opponentWordValue > myWordValue) opponentWins++;
  }

  if (myWins > opponentWins) return 'We win';
  if (opponentWins > myWins) return 'We lose';
  return 'Draw';
}
```
