---
id: 68cae5b538ff798bbd4da007
title: "Aufgabe 69: Fehlende Socken"
challengeType: 28
dashedName: challenge-69
---

# --description--

Gegeben ist eine ganze Zahl, die die Anzahl der Sockenpaare darstellt, mit denen Sie gestartet sind, und eine weitere ganze Zahl, die angibt, wie viele Waschzyklen Sie bereits durchlaufen haben. Geben Sie die Anzahl der vollständigen Sockenpaare zurück, die Sie aktuell haben, unter Verwendung der folgenden Bedingungen:

- Alle 2 Waschzyklen verlieren Sie eine einzelne Socke.
- Alle 3 Waschzyklen finden Sie eine einzelne fehlende Socke.
- Alle 5 Waschzyklen ist eine einzelne Socke abgenutzt und muss weggeworfen werden.
- Alle 10 Waschzyklen kaufen Sie ein Paar Socken.
- Sie können niemals weniger als null Socken insgesamt haben.
- Die Regeln können sich überschneiden. Zum Beispiel verlieren Sie im 10. Waschzyklus eine einzelne Socke, werfen eine einzelne Socke weg und kaufen ein neues Paar Socken.
- Geben Sie die Anzahl der vollständigen Sockenpaare zurück.

# --hints--

`sockPairs(2, 5)` sollte `1` zurückgeben.

```js
assert.equal(sockPairs(2, 5), 1);
```

`sockPairs(1, 2)` sollte `0` zurückgeben.

```js
assert.equal(sockPairs(1, 2), 0);
```

`sockPairs(5, 11)` sollte `4` zurückgeben.

```js
assert.equal(sockPairs(5, 11), 4);
```

`sockPairs(6, 25)` sollte `3` zurückgeben.

```js
assert.equal(sockPairs(6, 25), 3);
```

`sockPairs(1, 8)` sollte `0` zurückgeben.

```js
assert.equal(sockPairs(1, 8), 0);
```

# --seed--

## --seed-contents--

```js
function sockPairs(pairs, cycles) {

  return pairs;
}
```

# --solutions--

```js
function sockPairs(pairs, cycles) {
  let socks = pairs * 2;

  for (let i = 1; i <= cycles; i++) {
    if (i % 2 === 0) socks -= 1;
    if (i % 3 === 0) socks += 1;
    if (i % 5 === 0) socks -= 1;
    if (i % 10 === 0) socks += 2;

    if (socks < 0) socks = 0;
  }

  return Math.floor(socks / 2);
}
```
