---
id: 68cae5b538ff798bbd4da002
title: "Aufgabe 64: 24-Stunden-Format in 12-Stunden-Format umwandeln"
challengeType: 28
dashedName: challenge-64
---

# --description--

Gegeben ist ein String, der eine Uhrzeit im 24-Stunden-Format von `"HHMM"` darstellt. Geben Sie die Uhrzeit im entsprechenden 12-Stunden-Format von `"H:MM AM"` oder `"H:MM PM"` zurück.

- Die gegebene Eingabe ist immer ein vierstelliger String im 24-Stunden-Format, von `"0000"` bis `"2359"`.

# --hints--

`to12("1124")` sollte `"11:24 AM"` zurückgeben.

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")` sollte `"9:00 AM"` zurückgeben.

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")` sollte `"2:55 PM"` zurückgeben.

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")` sollte `"11:46 PM"` zurückgeben.

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")` sollte `"12:30 AM"` zurückgeben.

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```
