---
id: 68cae5b538ff798bbd4da009
title: "Aufgabe 71: Trinkgeldrechner"
challengeType: 28
dashedName: challenge-71
---

# --description--

Gegeben sind der Preis Ihrer Mahlzeit und ein benutzerdefinierter Trinkgeldprozentsatz. Geben Sie ein Array mit drei Trinkgeldbeträgen zurück: 15 %, 20 % und den benutzerdefinierten Betrag.

- Die Preise werden im Format `"$N.NN"` angegeben.
- Benutzerdefinierte Trinkgeldprozentsätze werden in diesem Format `"25%"` angegeben.
- Geben Sie die Beträge im gleichen `"$N.NN"`-Format zurück, auf zwei Dezimalstellen gerundet.

Zum Beispiel, bei einem `"$10.00"`-Mahlzeitpreis und einem `"25%"`-benutzerdefinierten Trinkgeldwert, geben Sie `["$1.50", "$2.00", "$2.50"]` zurück.

# --hints--

`calculateTips("$10.00", "25%")` sollte `["$1.50", "$2.00", "$2.50"]` zurückgeben.

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` sollte `["$13.45", "$17.93", "$23.31"]` zurückgeben.

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` sollte `["$2.98", "$3.97", "$1.79"]` zurückgeben.

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
