---
id: 68cae5b538ff798bbd4da005
title: "Aufgabe 67: E-Mail-Validator"
challengeType: 28
dashedName: challenge-67
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob es sich um eine gültige E-Mail-Adresse handelt, unter Verwendung der folgenden Bedingungen:

- Sie muss genau ein `@`-Symbol enthalten.
- Der lokale Teil (vor dem `@`):
  - Darf nur Buchstaben (`a-z`, `A-Z`), Ziffern (`0-9`), Punkte (`.`), Unterstriche (`_`) oder Bindestriche (`-`) enthalten.
  - Darf nicht mit einem Punkt beginnen oder enden.
- Der Domain-Teil (nach dem `@`):
  - Muss mindestens einen Punkt enthalten.
  - Muss mit einem Punkt gefolgt von mindestens zwei Buchstaben enden.
- Weder der lokale noch der Domain-Teil dürfen zwei Punkte hintereinander haben.

# --hints--

`validate("a@b.cd")` sollte `true` zurückgeben.

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")` sollte `true` zurückgeben.

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")` sollte `false` zurückgeben.

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")` sollte `false` zurückgeben.

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")` sollte `false` zurückgeben.

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")` sollte `true` zurückgeben.

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")` sollte `false` zurückgeben.

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")` sollte `false` zurückgeben.

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("develop..ment_user@c0D!NG.R.CKS")` sollte `false` zurückgeben.

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

`validate("git@commit@push.io")` sollte `false` zurückgeben.

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
