---
id: 68cae5b538ff798bbd4da006
title: "Aufgabe 68: Kreditkarten-Maskierer"
challengeType: 28
dashedName: challenge-68
---

# --description--

Gegeben ist ein String mit Kreditkartennummern. Geben Sie eine maskierte Version davon zurück, die folgende Bedingungen erfüllt:

- Der String enthält vier Gruppen mit jeweils vier Ziffern (`0-9`), wobei alle Gruppen durch ein einzelnes Leerzeichen oder einen einzelnen Bindestrich (`-`) getrennt sind.
- Ersetzen Sie alle Zahlen außer den letzten vier durch ein Sternchen (`*`).
- Lassen Sie die übrigen Zeichen unverändert.

Beispielsweise soll bei Eingabe von `"4012-8888-8888-1881"` `"****-****-****-1881"` zurückgegeben werden.

# --hints--

`mask("4012-8888-8888-1881")` sollte `"****-****-****-1881"` zurückgeben.

```js
assert.equal(mask("4012-8888-8888-1881"), "****-****-****-1881");
```

`mask("5105 1051 0510 5100")` sollte `"**** **** **** 5100"` zurückgeben.

```js
assert.equal(mask("5105 1051 0510 5100"), "**** **** **** 5100");
```

`mask("6011 1111 1111 1117")` sollte `"**** **** **** 1117"` zurückgeben.

```js
assert.equal(mask("6011 1111 1111 1117"), "**** **** **** 1117");
```

`mask("2223-0000-4845-0010")` sollte `"****-****-****-0010"` zurückgeben.

```js
assert.equal(mask("2223-0000-4845-0010"), "****-****-****-0010");
```

# --seed--

## --seed-contents--

```js
function mask(card) {

  return card;
}
```

# --solutions--

```js
function mask(card) {
  const split = card.split(card.includes('-') ? '-' : ' ');
  
  return card.includes('-') ?
    `****-****-****-${split[3]}` :
    `**** **** **** ${split[3]}`;
}
```
