---
id: 68cae5b538ff798bbd4da003
title: "Aufgabe 65: String zählen"
challengeType: 28
dashedName: challenge-65
---

# --description--

Gegeben sind zwei Strings. Bestimmen Sie, wie oft der zweite String im ersten vorkommt.

- Der Muster-String kann sich im ersten String überlappen. Zum Beispiel enthält `"aaa"` `"aa"` zweimal. Die ersten zwei `a` und die zweiten zwei.

# --hints--

`count('abcdefg', 'def')` sollte `1` zurückgeben.

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` sollte `0` zurückgeben.

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` sollte `2` zurückgeben.

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` sollte `3` zurückgeben.

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` sollte `10` zurückgeben.

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
