---
id: 68cae5b538ff798bbd4da00a
title: "Aufgabe 72: Thermostatregler 2"
challengeType: 28
dashedName: challenge-72
---

# --description--

Gegeben ist die aktuelle Raumtemperatur in Fahrenheit und eine Zieltemperatur in Celsius. Geben Sie einen String zurück, der angibt, wie die Raumtemperatur basierend auf diesen Vorgaben angepasst werden soll:

- Geben Sie `"Heat: X degrees Fahrenheit"` zurück, wenn die aktuelle Temperatur unter dem Zielwert liegt. Dabei ist `X` die Anzahl der Grad in Fahrenheit, um die der Raum geheizt werden muss, um das Ziel zu erreichen, auf eine Nachkommastelle gerundet.
- Geben Sie `"Cool: X degrees Fahrenheit"` zurück, wenn die aktuelle Temperatur über dem Zielwert liegt. Dabei ist `X` die Anzahl der Grad in Fahrenheit, um die der Raum gekühlt werden muss, um das Ziel zu erreichen, auf eine Nachkommastelle gerundet.
- Geben Sie `"Hold"` zurück, wenn die aktuelle Temperatur dem Zielwert entspricht.

Um Celsius in Fahrenheit umzurechnen, multiplizieren Sie die Celsius-Temperatur mit 1,8 und addieren Sie 32 zum Ergebnis (`F = (C * 1.8) + 32`).

# --hints--

`adjustThermostat(32, 0)` sollte `"Hold"` zurückgeben.

```js
assert.equal(adjustThermostat(32, 0), "Hold");
```

`adjustThermostat(70, 25)` sollte `"Heat: 7.0 degrees Fahrenheit"` zurückgeben.

```js
assert.equal(adjustThermostat(70, 25), "Heat: 7.0 degrees Fahrenheit");
```

`adjustThermostat(72, 18)` sollte `"Cool: 7.6 degrees Fahrenheit"` zurückgeben.

```js
assert.equal(adjustThermostat(72, 18), "Cool: 7.6 degrees Fahrenheit");
```

`adjustThermostat(212, 100)` sollte `"Hold"` zurückgeben.

```js
assert.equal(adjustThermostat(212, 100), "Hold");
```

`adjustThermostat(59, 22)` sollte `"Heat: 12.6 degrees Fahrenheit"` zurückgeben.

```js
assert.equal(adjustThermostat(59, 22), "Heat: 12.6 degrees Fahrenheit");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(currentF, targetC) {

  return currentF;
}
```

# --solutions--

```js
function adjustThermostat(currentF, targetC) {
  const targetF = targetC * 1.8 + 32;
  const diff = Math.abs(targetF - currentF).toFixed(1);

  if (currentF < targetF) {
    return `Heat: ${diff} degrees Fahrenheit`;
  } else if (currentF > targetF) {
    return `Cool: ${diff} degrees Fahrenheit`;
  } else {
    return "Hold"
  }
}
```
