---
id: 68d2ba1468508398389487ce
title: "Aufgabe 73: Weise sprechen, müssen Sie"
challengeType: 28
dashedName: challenge-73
---

# --description--

Gegeben ist ein Satz. Geben Sie eine Version davon zurück, die wie ein Rat eines weisen Lehrers klingt, und zwar nach den folgenden Regeln:

- Wörter sind durch ein einzelnes Leerzeichen getrennt.
- Finden Sie das erste Vorkommen eines der folgenden Wörter im Satz: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Verschieben Sie alle Wörter vor und einschließlich dieses Wortes ans Ende des Satzes und:
  - Bewahren Sie die Reihenfolge der Wörter beim Verschieben.
  - Schreiben Sie sie alle in Kleinbuchstaben.
  - Fügen Sie ein Komma und ein Leerzeichen davor ein.
- Großschreiben Sie den ersten Buchstaben des neuen ersten Wortes im Satz.
- Alle gegebenen Sätze enden mit einem einzelnen Satzzeichen. Behalten Sie die ursprüngliche Zeichensetzung bei und setzen Sie sie ans Ende des neuen Satzes.
- Geben Sie den neuen Satz zurück. Achten Sie darauf, dass zwischen den Wörtern jeweils nur ein Leerzeichen steht und keine Leerzeichen am Anfang oder Ende des Satzes sind.

Zum Beispiel, bei Eingabe `"You must speak wisely."` geben Sie `"Speak wisely, you must."` zurück.

# --hints--

`wiseSpeak("You must speak wisely.")` sollte `"Speak wisely, you must."` zurückgeben.

```js
assert.equal(wiseSpeak("You must speak wisely."), "Speak wisely, you must.");
```

`wiseSpeak("You can do it!")` sollte `"Do it, you can!"` zurückgeben.

```js
assert.equal(wiseSpeak("You can do it!"), "Do it, you can!");
```

`wiseSpeak("Do you think you will complete this?")` sollte `"Complete this, do you think you will?"` zurückgeben.

```js
assert.equal(wiseSpeak("Do you think you will complete this?"), "Complete this, do you think you will?");
```

`wiseSpeak("All your base are belong to us.")` sollte `"Belong to us, all your base are."` zurückgeben.

```js
assert.equal(wiseSpeak("All your base are belong to us."), "Belong to us, all your base are.");
```

`wiseSpeak("You have much to learn.")` sollte `"Much to learn, you have."` zurückgeben.

```js
assert.equal(wiseSpeak("You have much to learn."), "Much to learn, you have.");
```

# --seed--

## --seed-contents--

```js
function wiseSpeak(sentence) {

  return sentence;
}
```

# --solutions--

```js
function wiseSpeak(sentence) {
  const triggers = ["have", "must", "are", "will", "can"];
  const punctuation = sentence[sentence.length - 1];
  const words = sentence.split(" ");

  const triggerIndex = words.findIndex(w => triggers.includes(w));

  const toMove = words.slice(0, triggerIndex + 1).map(w => w.toLowerCase());
  const remaining = words.slice(triggerIndex + 1);

  const newStart = remaining.map((w, i) => {
    return i === 0 ? w[0].toUpperCase() + w.slice(1) :
      i === remaining.length - 1 ? w.slice(0, w.length - 1) : w;
  })

  return newStart.join(" ") + ", " + toMove.join(" ") + punctuation;
}
```
