---
id: 68d2ba1468508398389487d0
title: "Aufgabe 75: Versteckter Schatz"
challengeType: 28
dashedName: challenge-75
---

# --description--

Gegeben ist ein 2D-Array, das eine Karte des Meeresbodens mit einem versteckten Schatz darstellt, sowie ein Array mit den Koordinaten (`[row, column]`) für den nächsten Tauchgang Ihrer Schatzsuche. Geben Sie `"Empty"`, `"Found"` oder `"Recovered"` nach den folgenden Regeln zurück:

- Das gegebene 2D-Array enthält genau einen noch nicht geborgenen Schatz, der mehrere Zellen einnimmt.
- Jede Zelle im 2D-Array enthält einen der folgenden Werte:
  - `"-"`: Kein Schatz.
  - `"O"`: Ein Teil des Schatzes, der noch nicht gefunden wurde.
  - `"X"`: Ein Teil des Schatzes, der bereits gefunden wurde.

- Wenn der Tauchplatz keinen Schatz enthält, geben Sie `"Empty"` zurück.
- Wenn der Tauchplatz einen Schatz findet, aber mindestens ein anderer Teil des Schatzes noch nicht gefunden wurde, geben Sie `"Found"` zurück.
- Wenn der Tauchplatz den letzten noch nicht gefundenen Teil des Schatzes findet, geben Sie `"Recovered"` zurück.

Zum Beispiel, gegeben:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

Und `[2, 1]` für die Koordinaten des Tauchplatzes, geben Sie `"Recovered"` zurück, weil der Tauchgang den letzten noch nicht gefundenen Teil des Schatzes gefunden hat.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` sollte `"Recovered"` zurückgeben.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` sollte `"Empty"` zurückgeben.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` sollte `"Found"` zurückgeben.

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` sollte `"Found"` zurückgeben.

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` sollte `"Recovered"` zurückgeben.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` sollte `"Empty"` zurückgeben.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```
