---
id: 68d2ba1468508398389487cf
title: "Aufgabe 74: Lieblingslieder"
challengeType: 28
dashedName: challenge-74
---

# --description--

Erinnern Sie sich an iPods? Das erste Modell kam vor genau 24 Jahren heraus, am 23. Oktober 2001.

Gegeben ist ein Array von Song-Objekten, die Ihre iPod-Playlist darstellen. Geben Sie ein Array mit den Titeln der zwei meistgespielten Lieder zurück, wobei das meistgespielte Lied an erster Stelle steht.

- Jedes Objekt hat eine `"title"`-Eigenschaft (String) und eine `"plays"`-Eigenschaft (Ganzzahl).

# --hints--

`favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` sollte `["Sync or Swim", "Earbud Blues"]` zurückgeben.

```js
assert.deepEqual(favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"]);
```

`favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` sollte `["Clickwheel Love", "99 Downloads"]` zurückgeben.

```js
assert.deepEqual(favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"]);
```

`favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` sollte `["Song B", "Song C"]` zurückgeben.

```js
assert.deepEqual(favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"]);
```

# --seed--

## --seed-contents--

```js
function favoriteSongs(playlist) {

  return playlist;
}
```

# --solutions--

```js
function favoriteSongs(playlist) {
  const sorted = [...playlist].sort((a, b) => b.plays - a.plays);
  return sorted.slice(0, 2).map(song => song.title);
}
```
