---
id: 68d30fc57588d97fd3027b30
title: "Aufgabe 79: Navigator"
challengeType: 28
dashedName: challenge-79
---

# --description--

Am 28. Oktober 1994 wurde Netscape Navigator veröffentlicht, der Millionen half, das frühe Web zu erkunden.

Gegeben ist ein Array von Browserbefehlen, die Sie im Netscape Navigator ausgeführt haben. Geben Sie die aktuelle Seite zurück, auf der Sie sich nach Ausführung aller Befehle befinden, und zwar nach den folgenden Regeln:

- Sie starten immer auf der `"Home"`-Seite, die nicht im Befehls-Array enthalten ist.
- Gültige Befehle sind:
  - `"Visit Page"`: Dabei ist `"Page"` der Name der Seite, die Sie besuchen. Zum Beispiel bringt Sie `"Visit About"` zur `"About"`-Seite. Wenn Sie eine neue Seite besuchen, verwerfen Sie unbedingt jegliche Vorwärts-Historie.
  - `"Back"`: Bringt Sie zur vorherigen Seite in Ihrer Historie oder bleibt auf der aktuellen Seite, falls keine vorhanden ist.
  - `"Forward"`: Bringt Sie vorwärts in der Historie zur Seite, von der Sie gekommen sind, oder bleibt auf der aktuellen Seite, falls keine vorhanden ist.

Zum Beispiel, gegeben `["Visit About Us", "Back", "Forward"]`, geben Sie `"About Us"` zurück.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` sollte `"About Us"` zurückgeben.

```js
assert.equal(navigate(["Visit About Us", "Back", "Forward"]), "About Us");
```

`navigate(["Forward"])` sollte `"Home"` zurückgeben.

```js
assert.equal(navigate(["Forward"]), "Home");
```

`navigate(["Back"])` sollte `"Home"` zurückgeben.

```js
assert.equal(navigate(["Back"]), "Home");
```

`navigate(["Visit About Us", "Visit Gallery"])` sollte `"Gallery"` zurückgeben.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery"]), "Gallery");
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` sollte `"Home"` zurückgeben.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home");
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` sollte `"Contact"` zurückgeben.

```js
assert.equal(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact");
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` sollte `"Visit Us"` zurückgeben.

```js
assert.equal(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us");
```

# --seed--

## --seed-contents--

```js
function navigate(commands) {

  return commands;
}
```

# --solutions--

```js
function navigate(commands) {
  const history = ["Home"];
  let currentPageIndex = 0;

  for (const command of commands) {
    if (command.startsWith("Visit")) {
      history.splice(currentPageIndex + 1);
      history.push(command.slice(6));
      currentPageIndex++;
    } else if(command === "Back" && currentPageIndex > 0) {
      currentPageIndex--;
    } else if (command === "Forward" && currentPageIndex < history.length - 1) {
      currentPageIndex++;
    }
  }

  return history[currentPageIndex];
}
```
