---
id: 68d30845cc08266018fc46bc
title: "Aufgabe 76: Komplementäre DNA"
challengeType: 28
dashedName: challenge-76
---

# --description--

Gegeben ist ein String, der eine DNA-Folge darstellt. Geben Sie den komplementären Strang anhand der folgenden Regeln zurück:

- DNA besteht aus den Buchstaben `"A"`, `"C"`, `"G"` und `"T"`.
- Die Buchstaben `"A"` und `"T"` sind komplementär zueinander.
- Die Buchstaben `"C"` und `"G"` sind komplementär zueinander.

Zum Beispiel, bei `"ACGT"` geben Sie `"TGCA"` zurück.

# --hints--

`complementaryDNA("ACGT")` sollte `"TGCA"` zurückgeben.

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` sollte `"TACGCATGCAATCG"` zurückgeben.

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` sollte `"CCGAATGCTAGCTTC"` zurückgeben.

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` sollte `"CTAGATCGATCCGATCGATC"` zurückgeben.

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
