---
id: 68d30845cc08266018fc46be
title: "Aufgabe 78: Folge ganzer Zahlen"
challengeType: 28
dashedName: challenge-78
---

# --description--

Gegeben ist eine positive ganze Zahl. Geben Sie einen String zurück, der alle ganzen Zahlen von `1` bis einschließlich der angegebenen Zahl in numerischer Reihenfolge enthält.

Zum Beispiel, bei `5`, geben Sie `"12345"` zurück.

# --hints--

`sequence(5)` sollte `"12345"` zurückgeben.

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)` sollte `"12345678910"` zurückgeben.

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)` sollte `"1"` zurückgeben.

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)` sollte `"123456789101112131415161718192021222324252627"` zurückgeben.

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
