---
id: 68d30845cc08266018fc46bd
title: "Aufgabe 77: Dauerformatierer"
challengeType: 28
dashedName: challenge-77
---

# --description--

Gegeben ist eine ganze Zahl von Sekunden. Geben Sie einen String zurück, der dieselbe Dauer im Format `"H:MM:SS"` darstellt, wobei `"H"` die Anzahl der Stunden, `"MM"` die Anzahl der Minuten und `"SS"` die Anzahl der Sekunden ist. Geben Sie die Zeit nach folgenden Regeln zurück:

- Sekunden: Sollten immer zweistellig sein.
- Minuten: Führende Nullen sollten weggelassen werden, wenn sie nicht benötigt werden. Verwenden Sie `"0"`, wenn die Dauer weniger als eine Minute beträgt.
- Stunden: Sollten nur angegeben werden, wenn sie größer als null sind.

# --hints--

`format(500)` sollte `"8:20"` zurückgeben.

```js
assert.equal(format(500), "8:20");
```

`format(4000)` sollte `"1:06:40"` zurückgeben.

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)` sollte `"0:01"` zurückgeben.

```js
assert.equal(format(1), "0:01");
```

`format(5555)` sollte `"1:32:35"` zurückgeben.

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)` sollte `"27:46:39"` zurückgeben.

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
