---
id: 68e39ed6106dac2f0a98fd65
title: "Aufgabe 83: Validierung der Signatur"
challengeType: 28
dashedName: challenge-83
---

# --description--

Gegeben sind ein Nachrichten-String, ein geheimer Schlüssel-String und eine Signaturzahl. Bestimmen Sie, ob die Signatur mit dieser Codierungsmethode gültig ist:

- Buchstaben in der Nachricht und im geheimen Schlüssel haben folgende Werte:
  - `a` bis `z` haben jeweils die Werte `1` bis `26`.
  - `A` bis `Z` haben jeweils die Werte `27` bis `52`.
- Alle anderen Zeichen haben keinen Wert.
- Berechnen Sie die Signatur, indem Sie die Summe der Nachricht plus die Summe des geheimen Schlüssels nehmen.

Zum Beispiel, gegeben die Nachricht `"foo"` und den geheimen Schlüssel `"bar"`, wäre die Signatur `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

Prüfen Sie, ob die berechnete Signatur mit der angegebenen Signatur übereinstimmt.

# --hints--

`verify("foo", "bar", 57)` sollte `true` zurückgeben.

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` sollte `false` zurückgeben.

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` sollte `true` zurückgeben.

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` sollte `false` zurückgeben.

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` sollte `true` zurückgeben.

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` sollte `true` zurückgeben.

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
