---
id: 68e39ed6106dac2f0a98fd64
title: "Aufgabe 82: SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--

Gegeben ist ein String, der einen Variablennamen darstellt. Wandeln Sie ihn in „spooky case“ um, unter Beachtung folgender Vorgaben:

- Ersetzen Sie alle Unterstriche (`_`) und Bindestriche (`-`) durch eine Tilde (`~`).
- Großschreiben Sie den ersten Buchstaben des Strings und jeden zweiten Buchstaben danach. Die Tilde wird beim Zählen ignoriert. Alle anderen Buchstaben werden in Kleinbuchstaben geschrieben.

Zum Beispiel, bei Eingabe `hello_world` geben Sie `HeLlO~wOrLd` zurück.

# --hints--

`spookify("hello_world")` sollte `"HeLlO~wOrLd"` zurückgeben.

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")` sollte `"SpOoKy~CaSe"` zurückgeben.

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")` sollte `"TrIcK~oR~tReAt"` zurückgeben.

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")` sollte `"C~a~N~d~Y~~b~O~w~L"` zurückgeben.

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` sollte `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"` zurückgeben.

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```
