---
id: 68e39ed6106dac2f0a98fd62
title: "Aufgabe 80: E-Mail-Sortierer"
challengeType: 28
dashedName: challenge-80
---

# --description--

Am 29. Oktober 1971 wurde die erste E-Mail überhaupt verschickt, die das `username@domain`-Format einführte, das wir noch heute verwenden. Mittlerweile gibt es Milliarden von E-Mail-Adressen.

In dieser Aufgabe erhalten Sie eine Liste von E-Mail-Adressen, die Sie alphabetisch nach dem Domainnamen sortieren sollen (der Teil nach dem `@`) und anschließend nach dem Benutzernamen (der Teil vor dem `@`).

- Die Sortierung soll nicht zwischen Groß- und Kleinschreibung unterscheiden.
- Wenn mehrere E-Mails dieselbe Domain haben, sortieren Sie sie nach ihrem Benutzernamen.
- Geben Sie ein Array mit den sortierten Adressen zurück.
- Die zurückgegebenen Adressen sollen ihre ursprüngliche Groß- und Kleinschreibung beibehalten.

Zum Beispiel, wenn `["jill@mail.com", "john@example.com", "jane@example.com"]` gegeben ist, geben Sie `["jane@example.com", "john@example.com", "jill@mail.com"]` zurück.

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` sollte `["jane@example.com", "john@example.com", "jill@mail.com"]` zurückgeben.

```js
assert.deepEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"]);
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` sollte `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]` zurückgeben.

```js
assert.deepEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"]);
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` sollte `["user@x.com", "user@y.com", "user@z.com"]` zurückgeben.

```js
assert.deepEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"]);
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` sollte `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]` zurückgeben.

```js
assert.deepEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]);
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` sollte `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]` zurückgeben.

```js
assert.deepEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]);
```

# --seed--

## --seed-contents--

```js
function sort(emails) {

  return emails;
}
```

# --solutions--

```js
function sort(emails) {
  return emails.slice().sort((a, b) => {
    const [userA, domainA] = a.split('@');
    const [userB, domainB] = b.split('@');

    const domainCompare = domainA.toLowerCase().localeCompare(domainB.toLowerCase());
    if (domainCompare !== 0) return domainCompare;

    return userA.toLowerCase().localeCompare(userB.toLowerCase());
  });
}
```
