---
id: 68e39ed6106dac2f0a98fd63
title: "Aufgabe 81: n-te Primzahl"
challengeType: 28
dashedName: challenge-81
---

# --description--

Eine Primzahl ist eine positive ganze Zahl größer als 1, die nur durch 1 und sich selbst teilbar ist. Die ersten fünf Primzahlen sind `2`, `3`, `5`, `7` und `11`.

Gegeben ist eine positive ganze Zahl `n`. Geben Sie die `n`te Primzahl zurück. Zum Beispiel, bei `5` geben Sie die 5. Primzahl zurück: `11`.

# --hints--

`nthPrime(5)` sollte `11` zurückgeben.

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)` sollte `29` zurückgeben.

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)` sollte `53` zurückgeben.

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)` sollte `523` zurückgeben.

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)` sollte `7919` zurückgeben.

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
