---
id: 68ee9e3066cfd4eb2328e8a4
title: "Aufgabe 85: Wortzähler"
challengeType: 28
dashedName: challenge-85
---

# --description--

Gegeben ist ein Satz-String. Geben Sie die Anzahl der Wörter zurück, die im Satz enthalten sind.

- Wörter sind jede Folge von Nicht-Leerzeichen-Zeichen und werden durch ein einzelnes Leerzeichen getrennt.

# --hints--

`countWords("Hello world")` sollte `2` zurückgeben.

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` sollte `9` zurückgeben.

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` sollte `4` zurückgeben.

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` sollte `7` zurückgeben.

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` sollte `7` zurückgeben.

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
