---
id: 68ee9e3066cfd4eb2328e8a8
title: "Aufgabe 89: Karten zählen"
challengeType: 28
dashedName: challenge-89
---

# --description--

Ein Standard-Kartenspiel enthält 13 einzigartige Karten in jeder Farbe. Gegeben ist eine ganze Zahl, die die Anzahl der Karten angibt, die aus dem Deck gezogen werden sollen. Geben Sie die Anzahl der einzigartigen Kombinationen von Karten zurück, die Sie ziehen können.

- Die Reihenfolge spielt keine Rolle. Die Karte A zuerst zu ziehen und dann Karte B ist dasselbe wie Karte B zuerst zu ziehen und dann Karte A.

Zum Beispiel, gegeben `52`, geben Sie `1` zurück. Es gibt nur eine Kombination von 52 Karten, wenn Sie aus einem 52-Karten-Deck ziehen. Und gegeben `2`, geben Sie `1326` zurück. Es gibt 1326 Kartenkombinationen, wenn Sie 2 Karten aus dem Deck ziehen.

# --hints--

`combinations(52)` sollte `1` zurückgeben.

```js
assert.equal(combinations(52), 1);
```

`combinations(1)` sollte `52` zurückgeben.

```js
assert.equal(combinations(1), 52);
```

`combinations(2)` sollte `1326` zurückgeben.

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)` sollte `2598960` zurückgeben.

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)` sollte `15820024220` zurückgeben.

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)` sollte `1326` zurückgeben.

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```
