---
id: 68ee9e3066cfd4eb2328e8a5
title: "Aufgabe 86: Bildersuche"
challengeType: 28
dashedName: challenge-86
---

# --description--

Am 4. November 2001 startete Google seine Bildersuche, mit der Menschen Bilder anhand von Suchbegriffen finden konnten. In dieser Aufgabe werden Sie die Bildersuche nachahmen.

Gegeben ist ein Array von Bildnamen und ein Suchbegriff. Geben Sie ein Array mit Bildnamen zurück, die den Suchbegriff enthalten.

- Ignorieren Sie die Groß- und Kleinschreibung bei der Übereinstimmung der Suchbegriffe.
- Geben Sie die Bilder in der gleichen Reihenfolge zurück, in der sie im Eingabe-Array erscheinen.

# --hints--

`imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` sollte `["dog.png"]` zurückgeben.

```js
assert.deepEqual(imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"]);
```

`imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` sollte `["Sunset.jpg", "sunflower.jpeg"]` zurückgeben.

```js
assert.deepEqual(imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"]);
```

`imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` sollte `["Moon.png", "stars.png"]` zurückgeben.

```js
assert.deepEqual(imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"]);
```

`imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` sollte `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]` zurückgeben.

```js
assert.deepEqual(imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]);
```

# --seed--

## --seed-contents--

```js
function imageSearch(images, term) {

  return images;
}
```

# --solutions--

```js
function imageSearch(images, term) {
  const lowerTerm = term.toLowerCase();
  return images.filter(img => img.toLowerCase().includes(lowerTerm));
}
```
