---
id: 68ee9e3066cfd4eb2328e8a6
title: "Aufgabe 87: Matrix-Ersteller"
challengeType: 28
dashedName: challenge-87
---

# --description--

Gegeben sind zwei ganze Zahlen (eine Anzahl von Zeilen und eine Anzahl von Spalten). Geben Sie eine Matrix (ein Array von Arrays) zurück, die mit Nullen (`0`) in der angegebenen Größe gefüllt ist.

Zum Beispiel, bei `2` und `3`, geben Sie zurück:

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`buildMatrix(2, 3)` sollte `[[0, 0, 0], [0, 0, 0]]` zurückgeben.

```js
assert.deepEqual(buildMatrix(2, 3), [[0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(3, 2)` sollte `[[0, 0], [0, 0], [0, 0]]` zurückgeben.

```js
assert.deepEqual(buildMatrix(3, 2), [[0, 0], [0, 0], [0, 0]]);
```

`buildMatrix(4, 3)` sollte `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]` zurückgeben.

```js
assert.deepEqual(buildMatrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(9, 1)` sollte `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]` zurückgeben.

```js
assert.deepEqual(buildMatrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]]);
```

# --seed--

## --seed-contents--

```js
function buildMatrix(rows, cols) {

  return rows;
}
```

# --solutions--

```js
function buildMatrix(rows, cols) {
  const matrix = [];
  for (let i = 0; i < rows; i++) {
    const row = new Array(cols).fill(0);
    matrix.push(row);
  }
  return matrix;
}
```
