---
id: 68f6587287ad1f4ad39b0c81
title: "Aufgabe 95: Array verschieben"
challengeType: 28
dashedName: challenge-95
---

# --description--

Gegeben ist ein Array und eine ganze Zahl, die angibt, um wie viele Positionen das Array verschoben werden soll. Geben Sie das verschobene Array zurück.

- Eine positive ganze Zahl verschiebt das Array nach links.
- Eine negative ganze Zahl verschiebt das Array nach rechts.
- Die Verschiebung erfolgt zyklisch im Array.

Zum Beispiel, wenn `[1, 2, 3]` und `1` gegeben sind, verschieben Sie das Array 1 nach links und geben `[2, 3, 1]` zurück.

# --hints--

`shiftArray([1, 2, 3], 1)` sollte `[2, 3, 1]` zurückgeben.

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` sollte `[3, 1, 2]` zurückgeben.

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` sollte `["charlie", "alpha", "bravo"]` zurückgeben.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` sollte `["bravo", "charlie", "alpha"]` zurückgeben.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` sollte `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]` zurückgeben.

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
