---
id: 68f6587287ad1f4ad39b0c84
title: "Aufgabe 98: Rechteckanzahl"
challengeType: 28
dashedName: challenge-98
---

# --description--

Gegeben sind zwei positive ganze Zahlen, die die Breite und Höhe eines Rechtecks darstellen. Bestimmen Sie, wie viele Rechtecke in das gegebene Rechteck passen.

- Zählen Sie nur Rechtecke mit ganzzahliger Breite und Höhe.

Zum Beispiel, gegeben `1` und `3`, geben Sie `6` zurück. Drei 1x1-Rechtecke, zwei 1x2-Rechtecke und ein 1x3-Rechteck.

# --hints--

`countRectangles(1, 3)` sollte `6` zurückgeben.

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)` sollte `18` zurückgeben.

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)` sollte `3` zurückgeben.

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)` sollte `150` zurückgeben.

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)` sollte `12540` zurückgeben.

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
