---
id: 68f6587287ad1f4ad39b0c83
title: "Aufgabe 97: GGT"
challengeType: 28
dashedName: challenge-97
---

# --description--

Gegeben sind zwei positive ganze Zahlen. Geben Sie ihren größten gemeinsamen Teiler (GGT) zurück.

- Der GGT von zwei ganzen Zahlen ist die größte Zahl, die beide Zahlen ohne Rest teilt.

Zum Beispiel sind die Teiler von `4` `1`, `2` und `4`. Die Teiler von `6` sind `1`, `2`, `3` und `6`. Geben Sie also für `4` und `6` `2` zurück, die größte Zahl, die in beiden Teiler-Mengen vorkommt.

# --hints--

`gcd(4, 6)` sollte `2` zurückgeben.

```js
assert.equal(gcd(4, 6), 2);
```

`gcd(20, 15)` sollte `5` zurückgeben.

```js
assert.equal(gcd(20, 15), 5);
```

`gcd(13, 17)` sollte `1` zurückgeben.

```js
assert.equal(gcd(13, 17), 1);
```

`gcd(654, 456)` sollte `6` zurückgeben.

```js
assert.equal(gcd(654, 456), 6);
```

`gcd(3456, 4320)` sollte `864` zurückgeben.

```js
assert.equal(gcd(3456, 4320), 864);
```

# --seed--

## --seed-contents--

```js
function gcd(x, y) {

  return x;
}
```

# --solutions--

```js
function gcd(x, y) {
  while (y !== 0) {
    const temp = y;
    y = x % y;
    x = temp;
  }
  return x;
}
```
