---
id: 68f6587287ad1f4ad39b0c85
title: "Aufgabe 99: Fingerabdruck-Test"
challengeType: 28
dashedName: challenge-99
---

# --description--

Gegeben sind zwei Strings, die Fingerabdrücke darstellen. Bestimmen Sie, ob sie übereinstimmen, anhand der folgenden Regeln:

- Jeder Fingerabdruck besteht nur aus Kleinbuchstaben (`a-z`).
- Zwei Fingerabdrücke gelten als Übereinstimmung, wenn:
  - Sie die gleiche Länge haben.
  - Die Anzahl der unterschiedlichen Zeichen 10 % der Länge des Fingerabdrucks nicht überschreitet.

# --hints--

`isMatch("helloworld", "helloworld")` sollte `true` zurückgeben.

```js
assert.isTrue(isMatch("helloworld", "helloworld"));
```

`isMatch("helloworld", "helloworlds")` sollte `false` zurückgeben.

```js
assert.isFalse(isMatch("helloworld", "helloworlds"));
```

`isMatch("helloworld", "jelloworld")` sollte `true` zurückgeben.

```js
assert.isTrue(isMatch("helloworld", "jelloworld"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog")` sollte `true` zurückgeben.

```js
assert.isTrue(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog"));
```

`isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog")` sollte `true` zurückgeben.

```js
assert.isTrue(isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat")` sollte `false` zurückgeben.

```js
assert.isFalse(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat"));
```

# --seed--

## --seed-contents--

```js
function isMatch(fingerprintA, fingerprintB) {

  return fingerprintA;
}
```

# --solutions--

```js
function isMatch(fingerprintA, fingerprintB) {
  if (fingerprintA.length !== fingerprintB.length) return false;

  const length = fingerprintA.length;
  let mismatches = 0;

  for (let i = 0; i < length; i++) {
    if (fingerprintA[i] !== fingerprintB[i]) {
      mismatches++;
      if (mismatches > length * 0.1) return false;
    }
  }

  return true;
}
```
