---
id: 68f6587287ad1f4ad39b0c7c
title: "Aufgabe 90: Zeichenlimit"
challengeType: 28
dashedName: challenge-90
---

# --description--

In dieser Aufgabe erhalten Sie einen String und müssen bestimmen, ob er in einen Social-Media-Beitrag passt. Geben Sie die folgenden Strings basierend auf den angegebenen Regeln zurück:

- `"short post"`, wenn er innerhalb des 40-Zeichen-Limits liegt.
- `"long post"`, wenn er länger als 40 Zeichen ist und innerhalb des 80-Zeichen-Limits liegt.
- `"invalid post"`, wenn er zu lang ist, um in eines der beiden Limits zu passen.

# --hints--

`canPost("Hello world")` sollte `"short post"` zurückgeben.

```js
assert.equal(canPost("Hello world"), "short post");
```

`canPost("This is a longer message but still under eighty characters.")` sollte `"long post"` zurückgeben.

```js
assert.equal(canPost("This is a longer message but still under eighty characters."), "long post");
```

`canPost("This message is too long to fit into either of the character limits for a social media post.")` sollte `"invalid post"` zurückgeben.

```js
assert.equal(canPost("This message is too long to fit into either of the character limits for a social media post."), "invalid post");
```

# --seed--

## --seed-contents--

```js
function canPost(message) {

  return message;
}
```

# --solutions--

```js
function canPost(message) {
  if (message.length <= 40) {
    return "short post"
  } else if (message.length <= 80) {
    return "long post"
  } else {
    return "invalid post"
  }
}
```
