---
id: 68f6587287ad1f4ad39b0c82
title: "Aufgabe 96: Ist es das Wochenende?"
challengeType: 28
dashedName: challenge-96
---

# --description--

Gegeben ist ein Datum im Format `"YYYY-MM-DD"`. Geben Sie die Anzahl der Tage bis zum Wochenende zurück.

- Das Wochenende beginnt am Samstag.
- Wenn das angegebene Datum Samstag oder Sonntag ist, geben Sie `"It's the weekend!"` zurück.
- Andernfalls geben Sie `"X days until the weekend."` zurück, wobei `X` die Anzahl der Tage bis Samstag ist.
- Wenn `X` `1` ist, verwenden Sie `"day"` (Singular) anstelle von `"days"` (Plural).
- Stellen Sie sicher, dass die Berechnung Ihre lokale Zeitzone ignoriert.

# --hints--

`daysUntilWeekend("2025-11-14")` sollte `"1 day until the weekend."` zurückgeben.

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` sollte `"3 days until the weekend."` zurückgeben.

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` sollte `"It's the weekend!"` zurückgeben.

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` sollte `"4 days until the weekend."` zurückgeben.

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` sollte `"5 days until the weekend."` zurückgeben.

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` sollte `"It's the weekend!"` zurückgeben.

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
