---
id: 68f6587287ad1f4ad39b0c7e
title: "Aufgabe 92: Erweiterungsextraktor"
challengeType: 28
dashedName: challenge-92
---

# --description--

Gegeben ist ein String, der einen Dateinamen darstellt. Geben Sie die Erweiterung der Datei zurück.

- Die Erweiterung ist der Teil des Dateinamens, der nach dem letzten Punkt (`.`) kommt.
- Wenn der Dateiname keinen Punkt enthält oder mit einem Punkt endet, geben Sie `"none"` zurück.
- Die Erweiterung soll unverändert zurückgegeben werden, einschließlich Groß- und Kleinschreibung.

# --hints--

`getExtension("document.txt")` sollte `"txt"` zurückgeben.

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` sollte `"none"` zurückgeben.

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` sollte `"PNG"` zurückgeben.

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` sollte `"gitignore"` zurückgeben.

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` sollte `"gz"` zurückgeben.

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` sollte `"none"` zurückgeben.

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
