---
id: 68f6587287ad1f4ad39b0c7d
title: "Aufgabe 91: Wortsuche"
challengeType: 28
dashedName: challenge-91
---

# --description--

Gegeben ist eine Matrix (ein Array von Arrays) aus einzelnen Buchstaben und ein Wort, das gefunden werden soll. Geben Sie die Start- und Endindizes des Wortes in der Matrix zurück.

- Die gegebene Matrix ist mit Kleinbuchstaben (`a-z`) gefüllt.
- Das zu findende Wort kommt genau einmal in der Matrix vor.
- Das zu findende Wort liegt immer in einer geraden Linie in einer der folgenden Richtungen:
  - von links nach rechts
  - von rechts nach links
  - von oben nach unten
  - von unten nach oben

Zum Beispiel, gegeben die Matrix:

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

Und das Wort `"cat"`, geben Sie zurück:

```md
[[0, 1], [2, 1]]
```

Dabei sind `[0, 1]` die Indizes für den `"c"` (Anfang des Wortes) und `[2, 1]` die Indizes für den `"t"` (Ende des Wortes).

# --hints--

`findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")` sollte `[[0, 1], [2, 1]]` zurückgeben.

```js
assert.deepEqual(findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]]);
```

`findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")` sollte `[[0, 0], [0, 2]]` zurückgeben.

```js
assert.deepEqual(findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]]);
```

`findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")` sollte `[[3, 3], [0, 3]]` zurückgeben.

```js
assert.deepEqual(findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]]);
```

`findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")` sollte `[[1, 3], [1, 1]]` zurückgeben.

```js
assert.deepEqual(findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]]);
```

# --seed--

## --seed-contents--

```js
function findWord(matrix, word) {

  return matrix;
}
```

# --solutions--

```js
function findWord(matrix, word) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const len = word.length;

  const directions = [
    [0, 1],
    [0, -1],
    [1, 0],
    [-1, 0]
  ];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      for (let [dr, dc] of directions) {
        let match = true;
        for (let i = 0; i < len; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols || matrix[nr][nc] !== word[i]) {
            match = false;
            break;
          }
        }
        if (match) {
          return [
            [r, c],
            [r + dr * (len - 1), c + dc * (len - 1)]
          ];
        }
      }
    }
  }
}
```
