---
id: 68f6587287ad1f4ad39b0c80
title: "Aufgabe 94: Generator für E-Mail-Signaturen"
challengeType: 28
dashedName: challenge-94
---

# --description--

Gegeben sind Strings für den Namen, den Titel und das Unternehmen einer Person. Geben Sie eine E-Mail-Signatur als einen einzigen String zurück, der die folgenden Regeln erfüllt:

- Der Name soll zuerst erscheinen, vorangestellt von einem Präfix, das von dem ersten Buchstaben des Namens abhängt. Für Namen, die (unabhängig von Groß- oder Kleinschreibung) mit folgendem Buchstaben beginnen:
  - `A-I`: Verwenden Sie `>>` als Präfix.
  - `J-R`: Verwenden Sie `--` als Präfix.
  - `S-Z`: Verwenden Sie `::` als Präfix.
- Nach dem Namen folgt ein Komma und ein Leerzeichen (`, `).
- Der Titel und das Unternehmen folgen nach dem Komma und Leerzeichen, getrennt durch `" at "` (mit Leerzeichen davor und danach).

Zum Beispiel, bei den Eingaben `"Quinn Waverly"`, `"Founder and CEO"` und `"TechCo"` geben Sie `"--Quinn Waverly, Founder and CEO at TechCo"` zurück.

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")` sollte `"--Quinn Waverly, Founder and CEO at TechCo"` zurückgeben.

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")` sollte `">>Alice Reed, Engineer at TechCo"` zurückgeben.

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")` sollte `"::Tina Vaughn, Developer at example.com"` zurückgeben.

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")` sollte `">>B. B., Product Tester at AcmeCorp"` zurückgeben.

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")` sollte `"::windstorm, Cloud Architect at Atmospheronics"` zurückgeben.

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
