---
id: 68f6587287ad1f4ad39b0c7f
title: "Aufgabe 93: Vokale und Konsonanten"
challengeType: 28
dashedName: challenge-93
---

# --description--

Gegeben ist ein String. Geben Sie ein Array zurück, das die Anzahl der Vokale und die Anzahl der Konsonanten im String enthält.

- Vokale bestehen aus `a`, `e`, `i`, `o`, `u` in beliebiger Groß- oder Kleinschreibung.
- Konsonanten bestehen aus allen anderen Buchstaben in beliebiger Groß- oder Kleinschreibung.
- Ignorieren Sie alle Nicht-Buchstaben-Zeichen.

Zum Beispiel, bei `"Hello World"` geben Sie `[3, 7]` zurück.

# --hints--

`count("Hello World")` sollte `[3, 7]` zurückgeben.

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")` sollte `[3, 7]` zurückgeben.

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")` sollte `[1, 5]` zurückgeben.

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")` sollte `[5, 7]` zurückgeben.

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")` sollte `[3, 7]` zurückgeben.

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")` sollte `[11, 24]` zurückgeben.

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```
