---
id: 68ffb91507a5b645769328cc
title: "Aufgabe 109: Wie alt bin ich eigentlich?"
challengeType: 28
dashedName: challenge-109
---

# --description--

Gegeben ist das Geburtsdatum einer Person im Format `YYYY-MM-DD`. Geben Sie das Alter der Person zum 27. November 2025 zurück.

- Gehen Sie davon aus, dass alle Geburtstage gültige Daten vor dem 27. November 2025 sind.
- Geben Sie das Alter als ganze Zahl zurück.
- Berücksichtigen Sie unbedingt, ob die Person im Jahr 2025 bereits Geburtstag hatte.

# --hints--

`calculateAge("2000-11-20")` sollte `25` zurückgeben.

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` sollte `24` zurückgeben.

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` sollte `11` zurückgeben.

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` sollte `31` zurückgeben.

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` sollte `30` zurückgeben.

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
