---
id: 68ffb91507a5b645769328c5
title: "Aufgabe 102: Längstes Wort"
challengeType: 28
dashedName: challenge-102
---

# --description--

Gegeben ist ein Satz-String. Geben Sie das längste Wort im Satz zurück.

- Wörter sind durch ein einzelnes Leerzeichen getrennt.
- Nur Buchstaben (`a-z`, Groß- und Kleinschreibung wird nicht unterschieden) zählen zur Wortlänge.
- Wenn mehrere Wörter die gleiche Länge haben, geben Sie das zuerst vorkommende zurück.
- Geben Sie das Wort so zurück, wie es im gegebenen String erscheint, jedoch ohne Satzzeichen.

# --hints--

`longestWord("The quick red fox")` sollte `"quick"` zurückgeben.

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` sollte `"challenge"` zurückgeben.

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` sollte `"This"` zurückgeben.

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` sollte `"exclamation"` zurückgeben.

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` sollte `"tie"` zurückgeben.

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` sollte `"Wouldnt"` zurückgeben.

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
