---
id: 68ffb91507a5b645769328c3
title: "Aufgabe 100: 100 Zeichen"
challengeType: 28
dashedName: challenge-100
---

# --description--

Willkommen zur 100. täglichen Programmieraufgabe!

Gegeben ist ein String. Wiederholen Sie seine Zeichen so lange, bis das Ergebnis genau 100 Zeichen lang ist. Wenn Ihre Wiederholungen über 100 Zeichen hinausgehen, kürzen Sie den Überschuss, sodass es genau 100 sind.

# --hints--

`oneHundred("One hundred ")` sollte `"One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One "` zurückgeben.

```js
assert.equal(oneHundred("One hundred "), "One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One ");
```

`oneHundred("freeCodeCamp ")` sollte `"freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC"` zurückgeben.

```js
assert.equal(oneHundred("freeCodeCamp "), "freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC");
```

`oneHundred("daily challenges ")` sollte `"daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge"` zurückgeben.

```js
assert.equal(oneHundred("daily challenges "), "daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge");
```

`oneHundred("!")` sollte `"!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!"` zurückgeben.

```js
assert.equal(oneHundred("!"), "!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!");
```

# --seed--

## --seed-contents--

```js
function oneHundred(chars) {

  return chars;
}
```

# --solutions--

```js
function oneHundred(chars) {
  let result = "";
  let i = 0;

  while (result.length < 100) {
    result += chars[i % chars.length];
    i++;
  }

  return result.slice(0, 100);
}
```
