---
id: 68ffb91507a5b645769328ca
title: "Aufgabe 107: FizzBuzz"
challengeType: 28
dashedName: challenge-107
---

# --description--

Gegeben ist eine ganze Zahl (`n`). Geben Sie ein Array von ganzen Zahlen von `1` bis `n` (einschließlich) zurück, wobei Zahlen, die Vielfache von sind:

- 3 durch `"Fizz"` ersetzt werden.
- 5 durch `"Buzz"` ersetzt werden.
- 3 und 5 durch `"FizzBuzz"` ersetzt werden.

# --hints--

`fizzBuzz(2)` sollte `[1, 2]` zurückgeben.

```js
assert.deepEqual(fizzBuzz(2), [1, 2]);
```

`fizzBuzz(4)` sollte `[1, 2, "Fizz", 4]` zurückgeben.

```js
assert.deepEqual(fizzBuzz(4), [1, 2, "Fizz", 4]);
```

`fizzBuzz(8)` sollte `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8]` zurückgeben.

```js
assert.deepEqual(fizzBuzz(8), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8]);
```

`fizzBuzz(20)` sollte `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]` zurückgeben.

```js
assert.deepEqual(fizzBuzz(20), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]);
```

`fizzBuzz(50)` sollte `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]` zurückgeben.

```js
assert.deepEqual(fizzBuzz(50), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzz(n) {
  const result = [];
  for (let i = 1; i <= n; i++) {
    if (i % 15 === 0) {
      result.push("FizzBuzz");
    } else if (i % 3 === 0) {
      result.push("Fizz");
    } else if (i % 5 === 0) {
      result.push("Buzz");
    } else {
      result.push(i);
    }
  }
  return result;
}
```
