---
id: 68ffb91507a5b645769328c6
title: "Aufgabe 103: kgV"
challengeType: 28
dashedName: challenge-103
---

# --description--

Gegeben sind zwei ganze Zahlen. Geben Sie das kleinste gemeinsame Vielfache (kgV) der beiden Zahlen zurück.

Das kgV von zwei Zahlen ist die kleinste positive ganze Zahl, die ein Vielfaches beider Zahlen ist. Zum Beispiel, gegeben `4` und `6`, geben Sie `12` zurück, weil:

- Vielfache von `4` sind `4`, `8`, `12` und so weiter.
- Vielfache von `6` sind `6`, `12`, `18` und so weiter.
- `12` ist die kleinste Zahl, die ein Vielfaches von beiden ist.

# --hints--

`lcm(4, 6)` sollte `12` zurückgeben.

```js
assert.equal(lcm(4, 6), 12);
```

`lcm(9, 6)` sollte `18` zurückgeben.

```js
assert.equal(lcm(9, 6), 18);
```

`lcm(10, 100)` sollte `100` zurückgeben.

```js
assert.equal(lcm(10, 100), 100);
```

`lcm(13, 17)` sollte `221` zurückgeben.

```js
assert.equal(lcm(13, 17), 221);
```

`lcm(45, 70)` sollte `630` zurückgeben.

```js
assert.equal(lcm(45, 70), 630);
```

# --seed--

## --seed-contents--

```js
function lcm(a, b) {

  return a;
}
```

# --solutions--

```js
function lcm(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  return Math.abs(a * b) / gcd(a, b);
}
```
