---
id: 68ffb91507a5b645769328c8
title: "Aufgabe 105: Zeichen zählen"
challengeType: 28
dashedName: challenge-105
---

# --description--

Gegeben ist ein Satz-String. Geben Sie ein Array zurück, das die Anzahl jedes Zeichens in alphabetischer Reihenfolge enthält.

- Behandeln Sie Groß- und Kleinbuchstaben beim Zählen als denselben Buchstaben.
- Ignorieren Sie Zahlen, Leerzeichen, Satzzeichen usw.
- Geben Sie die Anzahl und den Buchstaben im Format `"letter count"` zurück. Zum Beispiel `"a 3"`.
- Alle zurückgegebenen Buchstaben sollten Kleinbuchstaben sein.
- Geben Sie keine Anzahl für Buchstaben zurück, die im gegebenen String nicht vorkommen.

# --hints--

`countCharacters("hello world")` sollte `["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]` zurückgeben.

```js
assert.deepEqual(countCharacters("hello world"), ["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]);
```

`countCharacters("I love coding challenges!")` sollte `["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]` zurückgeben.

```js
assert.deepEqual(countCharacters("I love coding challenges!"), ["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]);
```

`countCharacters("// TODO: Complete this challenge ASAP!")` sollte `["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]` zurückgeben.

```js
assert.deepEqual(countCharacters("// TODO: Complete this challenge ASAP!"), ["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]);
```

# --seed--

## --seed-contents--

```js
function countCharacters(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countCharacters(sentence) {
  const counts = {};

  for (let char of sentence.toLowerCase()) {
    if (char >= "a" && char <= "z") {
      counts[char] = (counts[char] || 0) + 1;
    }
  }

  return Object.keys(counts)
               .sort()
               .map(letter => `${letter} ${counts[letter]}`);
}
```
