---
id: 68ffb91507a5b645769328c9
title: "Aufgabe 106: Nachrichtenvalidierer"
challengeType: 28
dashedName: challenge-106
---

# --description--

Gegeben sind ein Nachrichten-String und ein Validierungs-String. Bestimmen Sie, ob die Nachricht gültig ist.

- Eine Nachricht ist gültig, wenn jedes Wort in der Nachricht mit dem entsprechenden Buchstaben im Validierungs-String in der richtigen Reihenfolge beginnt.
- Die Groß- und Kleinschreibung der Buchstaben spielt keine Rolle.
- Die Wörter in der Nachricht sind durch einzelne Leerzeichen getrennt.

# --hints--

`isValidMessage("hello world", "hw")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
