---
id: 68ffb91507a5b645769328c7
title: "Aufgabe 104: Rezept skalieren"
challengeType: 28
dashedName: challenge-104
---

# --description--

Gegeben ist ein Array mit Zutaten für ein Rezept und eine Zahl, um das Rezept zu skalieren. Geben Sie ein Array mit den entsprechend skalierten Mengen zurück.

- Jeder Eintrag im gegebenen Array ist ein String im Format: `"quantity unit ingredient"`. Zum Beispiel `"2 C Flour"`.
- Skalieren Sie die Menge mit der angegebenen Zahl. Schließen Sie keine nachfolgenden Nullen ein und wandeln Sie keine Einheiten um.
- Geben Sie die skalierten Einträge in der gleichen Reihenfolge zurück, in der sie gegeben sind.

# --hints--

`scaleRecipe(["2 C Flour", "1.5 T Sugar"], 2)` sollte `["4 C Flour", "3 T Sugar"]` zurückgeben.

```js
assert.deepEqual(scaleRecipe(["2 C Flour", "1.5 T Sugar"], 2), ["4 C Flour", "3 T Sugar"]);
```

`scaleRecipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5)` sollte `["6 T Flour", "1.5 C Milk", "3 T Oil"]` zurückgeben.

```js
assert.deepEqual(scaleRecipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5), ["6 T Flour", "1.5 C Milk", "3 T Oil"]);
```

`scaleRecipe(["3 C Milk", "2 C Oats"], 0.5)` sollte `["1.5 C Milk", "1 C Oats"]` zurückgeben.

```js
assert.deepEqual(scaleRecipe(["3 C Milk", "2 C Oats"], 0.5), ["1.5 C Milk", "1 C Oats"]);
```

`scaleRecipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5)` sollte `["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]` zurückgeben.

```js
assert.deepEqual(scaleRecipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5), ["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]);
```

# --seed--

## --seed-contents--

```js
function scaleRecipe(ingredients, scale) {

  return ingredients;
}
```

# --solutions--

```js
function scaleRecipe(ingredients, scale) {
  return ingredients.map(item => {
    const [quantityStr, unit, ...rest] = item.split(" ");
    const ingredientName = rest.join(" ");
    
    const quantity = parseFloat(quantityStr) * scale;
    
    return `${quantity} ${unit} ${ingredientName}`;
  });
}
```
