---
id: 68ffb91507a5b645769328c4
title: "Aufgabe 101: Markdown-Überschrift-Konverter"
challengeType: 28
dashedName: challenge-101
---

# --description--

Gegeben ist ein String, der eine Markdown-Überschrift darstellt. Geben Sie die entsprechende HTML-Überschrift zurück.

Eine gültige Markdown-Überschrift muss:

- Mit null oder mehr Leerzeichen beginnen, gefolgt von
- 1 bis 6 Rautezeichen (`#`) in Folge, dann
- Mindestens einem Leerzeichen. Und schließlich,
- Dem Überschriftentext.

Die Anzahl der Rautezeichen bestimmt die Überschriftenebene. Zum Beispiel entspricht ein Rautezeichen einem `h1`-Tag, und sechs Rautezeichen entsprechen einem `h6`-Tag.

Wenn der gegebene String nicht genau das oben beschriebene Format hat, geben Sie `"Invalid format"` zurück.

Beispielsweise, wenn `"# My level 1 heading"` gegeben ist, geben Sie `"<h1>My level 1 heading</h1>"` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings möglicherweise nicht an, wenn Nachrichten protokolliert werden. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`convert("# My level 1 heading")` sollte `"<h1>My level 1 heading</h1>"` zurückgeben.

```js
assert.equal(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>");
```

`convert("My heading")` sollte `"Invalid format"` zurückgeben.

```js
assert.equal(convert("My heading"), "Invalid format");
```

`convert("##### My level 5 heading")` sollte `"<h5>My level 5 heading</h5>"` zurückgeben.

```js
assert.equal(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>");
```

`convert("#My heading")` sollte `"Invalid format"` zurückgeben.

```js
assert.equal(convert("#My heading"), "Invalid format");
```

`convert("  ###  My level 3 heading")` sollte `"<h3>My level 3 heading</h3>"` zurückgeben.

```js
assert.equal(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>");
```

`convert("####### My level 7 heading")` sollte `"Invalid format"` zurückgeben.

```js
assert.equal(convert("####### My level 7 heading"), "Invalid format");
```

`convert("## My #2 heading")` sollte `"<h2>My #2 heading</h2>"` zurückgeben.

```js
assert.equal(convert("## My #2 heading"), "<h2>My #2 heading</h2>");
```

# --seed--

## --seed-contents--

```js
function convert(heading) {

  return heading;
}
```

# --solutions--

```js
function convert(heading) {
  const trimmed = heading.trim()

  if (!/^#{1,6}\s+\S+/.test(trimmed)) return "Invalid format"

  const hashCount = trimmed.indexOf(' ');
  const text = trimmed.slice(hashCount).trimStart();

  return `<h${hashCount}>${text}</h${hashCount}>`;
}
```
