---
id: 69162d64f96574d9bb629f04
title: "Aufgabe 119: String-Komprimierung"
challengeType: 28
dashedName: challenge-119
---

# --description--

Gegeben ist ein String-Satz. Geben Sie eine komprimierte Version des Satzes zurück, bei der aufeinanderfolgende doppelte Wörter durch das Wort gefolgt von der Anzahl der Wiederholungen in runden Klammern ersetzt werden.

- Nur aufeinanderfolgende Duplikate werden komprimiert.
- Wörter sind durch einzelne Leerzeichen getrennt.

Zum Beispiel, gegeben `"yes yes yes please"`, geben Sie `"yes(3) please"` zurück.

# --hints--

`compressString("yes yes yes please")` sollte `"yes(3) please"` zurückgeben.

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

`compressString("I have have have apples")` sollte `"I have(3) apples"` zurückgeben.

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

`compressString("one one three and to the the the the")` sollte `"one(2) three and to the(4)"` zurückgeben.

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

`compressString("route route route route route route tee tee tee tee tee tee")` sollte `"route(6) tee(6)"` zurückgeben.

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```
