---
id: 69162d64f96574d9bb629f01
title: "Aufgabe 116: Anzahl der Permutationen"
challengeType: 28
dashedName: challenge-116
---

# --description--

Gegeben ist ein String. Geben Sie die Anzahl der unterschiedlichen Permutationen zurück, die aus seinen Zeichen gebildet werden können.

- Eine Permutation ist jede Umordnung der Zeichen im String.
- Zählen Sie keine doppelten Permutationen.
- Wenn der String wiederholte Zeichen enthält, sollen wiederholte Anordnungen nur einmal gezählt werden.
- Der String enthält nur Buchstaben (`A-Z`, `a-z`).

Zum Beispiel, für `"abb"` geben Sie `3` zurück, weil es drei einzigartige Möglichkeiten gibt, die Buchstaben anzuordnen: `"abb"`, `"bab"` und `"bba"`.

# --hints--

`countPermutations("abb")` sollte `3` zurückgeben.

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")` sollte `6` zurückgeben.

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")` sollte `630` zurückgeben.

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")` sollte `39916800` zurückgeben.

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
