---
id: 69162d64f96574d9bb629efd
title: "Aufgabe 112: KI-Detektor"
challengeType: 28
dashedName: challenge-112
---

# --description--

Die heutige Aufgabe ist inspiriert von der Veröffentlichung von ChatGPT am 30. November 2022.

Gegeben ist ein String mit einem oder mehreren Sätzen. Bestimmen Sie, ob dieser wahrscheinlich von einer KI generiert wurde, anhand der folgenden Regeln:

- Er enthält zwei oder mehr Bindestriche (`-`).
- Er enthält zwei oder mehr Klammerpaare (`()`). Text kann innerhalb der Klammern stehen.
- Er enthält drei oder mehr Wörter mit 7 oder mehr Buchstaben.

- Wörter sind durch ein einzelnes Leerzeichen getrennt und bestehen nur aus Buchstaben (`A-Z`). Satzzeichen oder andere Nicht-Buchstaben zählen nicht als Teil eines Wortes.

Wenn der gegebene Satz eine der oben genannten Regeln erfüllt, geben Sie `"AI"` zurück, andernfalls `"Human"`.

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` sollte `"Human"` zurückgeben.

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` sollte `"Human"` zurückgeben.

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` sollte `"AI"` zurückgeben.

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` sollte `"AI"` zurückgeben.

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` sollte `"AI"` zurückgeben.

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
