---
id: 69162d64f96574d9bb629f00
title: "Aufgabe 115: Markdown geordnetes Listenelement konvertieren"
challengeType: 28
dashedName: challenge-115
---

# --description--

Gegeben ist ein String, der ein Element einer geordneten Liste in Markdown darstellt. Geben Sie den entsprechenden HTML-String zurück.

Ein gültiges Element einer geordneten Liste in Markdown muss:

- Mit null oder mehr Leerzeichen beginnen, gefolgt von
- Einer Zahl (1 oder größer) und einem Punkt (`.`), gefolgt von
- Mindestens einem Leerzeichen, und dann
- Dem Text des Listenelements.

Wenn der String nicht genau dieses Format hat, geben Sie `"Invalid format"` zurück. Andernfalls umschließen Sie den Text des Listenelements mit `li`-Tags und geben den String zurück.

Zum Beispiel, bei `"1. My item"` geben Sie `"<li>My item</li>"` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings möglicherweise nicht an, wenn Nachrichten protokolliert werden. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`convertListItem("1. My item")` sollte `"<li>My item</li>"` zurückgeben.

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` sollte `"<li>Another item</li>"` zurückgeben.

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` sollte `"Invalid format"` zurückgeben.

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` sollte `"<li>list item text</li>"` zurückgeben.

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` sollte `"Invalid format"` zurückgeben.

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` sollte `"Invalid format"` zurückgeben.

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
