---
id: 69162d64f96574d9bb629efe
title: "Aufgabe 113: Meilen in Kilometer"
challengeType: 28
dashedName: challenge-113
---

# --description--

Gegeben ist eine Entfernung in Meilen als Zahl. Geben Sie die entsprechende Entfernung in Kilometern zurück.

- Die Eingabe ist immer eine nicht negative Zahl.
- 1 Meile entspricht 1,60934 Kilometern.
- Runden Sie das Ergebnis auf zwei Dezimalstellen.
- Entfernen Sie unnötige nachgestellte Nullen vom gerundeten Ergebnis.

# --hints--

`convertToKm(1)` sollte `1.61` zurückgeben.

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)` sollte `33.8` zurückgeben.

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)` sollte `5.63` zurückgeben.

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)` sollte `0` zurückgeben.

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)` sollte `1` zurückgeben.

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
