---
id: 69162d64f96574d9bb629f03
title: "Aufgabe 118: Datumsformatierer"
challengeType: 28
dashedName: challenge-118
---

# --description--

Gegeben ist ein Datum im Format `"Month day, year"`. Geben Sie das Datum im Format `"YYYY-MM-DD"` zurück.

- Der angegebene Monat wird der vollständige englische Monatsname sein. Zum Beispiel: `"January"`, `"February"` usw.
- Im Rückgabewert fügen Sie bei Monat und Tag führende Nullen hinzu, falls nötig, damit jeweils zwei Ziffern angezeigt werden.

Zum Beispiel, bei Eingabe von `"December 6, 2025"`, geben Sie `"2025-12-06"` zurück.

# --hints--

`formatDate("December 6, 2025")` sollte `"2025-12-06"` zurückgeben.

```js
assert.equal(formatDate("December 6, 2025"), "2025-12-06");
```

`formatDate("January 1, 2000")` sollte `"2000-01-01"` zurückgeben.

```js
assert.equal(formatDate("January 1, 2000"), "2000-01-01");
```

`formatDate("November 11, 1111")` sollte `"1111-11-11"` zurückgeben.

```js
assert.equal(formatDate("November 11, 1111"), "1111-11-11");
```

`formatDate("September 7, 512")` sollte `"512-09-07"` zurückgeben.

```js
assert.equal(formatDate("September 7, 512"), "512-09-07");
```

`formatDate("May 4, 1950")` sollte `"1950-05-04"` zurückgeben.

```js
assert.equal(formatDate("May 4, 1950"), "1950-05-04");
```

`formatDate("February 29, 1992")` sollte `"1992-02-29"` zurückgeben.

```js
assert.equal(formatDate("February 29, 1992"), "1992-02-29");
```

# --seed--

## --seed-contents--

```js
function formatDate(dateString) {

  return dateString;
}
```

# --solutions--

```js
function formatDate(dateString) {
  const months = {
    January: "01",
    February: "02",
    March: "03",
    April: "04",
    May: "05",
    June: "06",
    July: "07",
    August: "08",
    September: "09",
    October: "10",
    November: "11",
    December: "12"
  };

  const [month, day, year] = dateString.replace(",", "").split(" ");

  return `${year}-${months[month]}-${day.padStart(2, "0")}`;
}
```
