---
id: 69162d64f96574d9bb629eff
title: "Aufgabe 114: Camel Case in Snake Case"
challengeType: 28
dashedName: challenge-114
---

# --description--

Gegeben ist ein String im Camel Case. Geben Sie die Snake Case-Version des Strings zurück, wobei Sie die folgenden Regeln beachten:

- Der Eingabe-String enthält nur Buchstaben (`A-Z` und `a-z`) und beginnt immer mit einem Kleinbuchstaben.
- Jeder Großbuchstabe im Camel Case-String markiert den Beginn eines neuen Wortes.
- Wandeln Sie alle Buchstaben in Kleinbuchstaben um.
- Trennen Sie die Wörter mit einem Unterstrich (`_`).

# --hints--

`toSnake("helloWorld")` sollte `"hello_world"` zurückgeben.

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` sollte `"my_variable_name"` zurückgeben.

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` sollte `"freecodecamp_daily_challenges"` zurückgeben.

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
