---
id: 69162d64f96574d9bb629efc
title: "Aufgabe 111: Ballflugbahn"
challengeType: 28
dashedName: challenge-111
---

# --description--

Die heutige Aufgabe ist inspiriert vom Videospiel Pong, das am 29. November 1972 veröffentlicht wurde.

Gegeben ist eine Matrix (Array von Arrays), die die Position des Balls (`2`) und die vorherige Position des Balls (`1`) enthält. Geben Sie die Matrixindizes für die nächste Position des Balls zurück.

- Der Ball bewegt sich immer in einer geraden Linie.
- Die Bewegungsrichtung wird bestimmt durch die Bewegung des Balls von `1` zu `2`.
- Die Ränder der Matrix gelten als Wände. Wenn der Ball eine:
  - obere oder untere Wand trifft, prallt er ab, indem er seine vertikale Richtung umkehrt.
  - linke oder rechte Wand trifft, prallt er ab, indem er seine horizontale Richtung umkehrt.
  - Ecke trifft, prallt er ab, indem er beide Richtungen umkehrt.

# --hints--

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]])` sollte `[2, 3]` zurückgeben.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]]), [2, 3]);
```

`getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]])` sollte `[3, 0]` zurückgeben.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]]), [3, 0]);
```

`getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]])` sollte `[1, 2]` zurückgeben.

```js
assert.deepEqual(getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]]), [1, 2]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]])` sollte `[1, 1]` zurückgeben.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]]), [1, 1]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]])` sollte `[2, 2]` zurückgeben.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]]), [2, 2]);
```

# --seed--

## --seed-contents--

```js
function getNextLocation(matrix) {

  return matrix;
}
```

# --solutions--

```js
function getNextLocation(matrix) {
  let prev = null;
  let curr = null;

  for (let row = 0; row < matrix.length; row++) {
    for (let col = 0; col < matrix[0].length; col++) {
      if (matrix[row][col] === 1) prev = [row, col];
      if (matrix[row][col] === 2) curr = [row, col];
    }
  }

  const [prevRow, prevCol] = prev;
  const [currRow, currCol] = curr;

  let dirX = currCol - prevCol;
  let dirY = currRow - prevRow;

  let nextRow = currRow + dirY;
  let nextCol = currCol + dirX;

  const maxRow = matrix.length - 1;
  const maxCol = matrix[0].length - 1;


  if (nextCol < 0 || nextCol > maxCol) dirX *= -1;
  if (nextRow < 0 || nextRow > maxRow) dirY *= -1;

  nextRow = currRow + dirY;
  nextCol = currCol + dirX;

  return [nextRow, nextCol];
}
```
