---
id: 69162d64f96574d9bb629efb
title: "Aufgabe 110: Wort-Rater"
challengeType: 28
dashedName: challenge-110
---

# --description--

Gegeben sind zwei Strings gleicher Länge, ein geheimes Wort und ein Rateversuch. Vergleichen Sie den Rateversuch mit dem geheimen Wort anhand der folgenden Regeln:

- Das geheime Wort und der Rateversuch bestehen nur aus Großbuchstaben (`"A"` bis `"Z"`);
- Für jeden Buchstaben im Rateversuch ersetzen Sie ihn durch eine Zahl, je nachdem, wie er mit dem geheimen Wort übereinstimmt:
  - `"2"`, wenn der Buchstabe im geheimen Wort an der richtigen Position steht.
  - `"1"`, wenn der Buchstabe im geheimen Wort vorkommt, aber an der falschen Position.
  - `"0"`, wenn der Buchstabe nicht im geheimen Wort vorkommt.
- Jeder Buchstabe im geheimen Wort kann höchstens einmal verwendet werden.
- Exakte Übereinstimmungen (`"2"`) werden zuerst zugewiesen, danach werden partielle Übereinstimmungen (`"1"`) von links nach rechts für die verbleibenden Buchstaben zugewiesen.
- Wenn ein Buchstabe im Rateversuch mehrfach vorkommt, kann er nur so oft übereinstimmen, wie er im geheimen Wort vorkommt.

Zum Beispiel, bei einem geheimen Wort `"APPLE"` und einem Rateversuch `"POPPA"`, geben Sie `"10201"` zurück:

Der erste `"P"` steht nicht an der richtigen Stelle (`"1"`), der `"O"` ist nicht im geheimen Wort (`"0"`), der zweite `"P"` steht an der richtigen Stelle (`"2"`), der dritte `"P"` ist eine Null (`"0"`), weil die zwei `"P"` im geheimen Wort bereits verwendet wurden, und der `"A"` steht nicht an der richtigen Stelle (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` sollte `"10201"` zurückgeben.

```js
assert.equal(compare("APPLE", "POPPA"), "10201");
```

`compare("REACT", "TRACE")` sollte `"11221"` zurückgeben.

```js
assert.equal(compare("REACT", "TRACE"), "11221");
```

`compare("DEBUGS", "PYTHON")` sollte `"000000"` zurückgeben.

```js
assert.equal(compare("DEBUGS", "PYTHON"), "000000");
```

`compare("JAVASCRIPT", "TYPESCRIPT")` sollte `"0000222222"` zurückgeben.

```js
assert.equal(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222");
```

`compare("ORANGE", "ROUNDS")` sollte `"110200"` zurückgeben.

```js
assert.equal(compare("ORANGE", "ROUNDS"), "110200");
```

`compare("WIRELESS", "ETHERNET")` sollte `"10021000"` zurückgeben.

```js
assert.equal(compare("WIRELESS", "ETHERNET"), "10021000");
```

# --seed--

## --seed-contents--

```js
function compare(word, guess) {

  return word;
}
```

# --solutions--

```js
function compare(word, guess) {
  const n = word.length;
  let result = Array(n).fill("0");
  let used = Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    if (guess[i] === word[i]) {
      result[i] = "2";
      used[i] = true;
    }
  }

  for (let i = 0; i < n; i++) {
    if (result[i] === "0") {
      for (let j = 0; j < n; j++) {
        if (!used[j] && guess[i] === word[j]) {
          result[i] = "1";
          used[j] = true;
          break;
        }
      }
    }
  }

  return result.join("");
}
```
