---
id: 69162d64f96574d9bb629f02
title: "Aufgabe 117: Symmetrische Differenz"
challengeType: 28
dashedName: challenge-117
---

# --description--

Gegeben sind zwei Arrays. Geben Sie ein neues Array zurück, das die symmetrische Differenz dieser Arrays enthält.

- Die symmetrische Differenz zwischen zwei Mengen ist die Menge der Werte, die in einer der beiden Mengen, aber nicht in beiden vorkommen.
- Geben Sie die Werte in der Reihenfolge zurück, in der sie zuerst in den Eingabe-Arrays erscheinen.

# --hints--

`difference([1, 2, 3], [3, 4, 5])` sollte `[1, 2, 4, 5]` zurückgeben.

```js
assert.deepEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5]);
```

`difference(["a", "b"], ["c", "b"])` sollte `["a", "c"]` zurückgeben.

```js
assert.deepEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"]);
```

`difference([1, "a", 2], [2, "b", "a"])` sollte `[1, "b"]` zurückgeben.

```js
assert.deepEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"]);
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` sollte `[2, 4, 6, 8]` zurückgeben.

```js
assert.deepEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8]);
```

# --seed--

## --seed-contents--

```js
function difference(arr1, arr2) {

  return arr1;
}
```

# --solutions--

```js
function difference(arr1, arr2) {
  const diff = [];

  for (const v of arr1) {

    if (!diff.includes(v) && !arr2.includes(v)) diff.push(v)
  }
  
  for (const v of arr2) {
    if (!diff.includes(v) && !arr1.includes(v)) diff.push(v)
  }

  return diff;
}
```
