---
id: 691b559495c5cb5a37b9b480
title: "Aufgabe 120: Pfund in Kilogramm"
challengeType: 28
dashedName: challenge-120
---

# --description--

Gegeben ist ein Gewicht in Pfund als Zahl. Geben Sie den String `"(lbs) pounds equals (kgs) kilograms."` zurück.

- Ersetzen Sie `"(lbs)"` durch die Eingabezahl.
- Ersetzen Sie `"(kgs)"` durch die in Kilogramm umgerechnete Eingabe, auf zwei Dezimalstellen gerundet, und geben Sie immer zwei Dezimalstellen im Wert an.
- 1 Pfund entspricht 0,453592 Kilogramm.
- Wenn die Eingabe `1` ist, verwenden Sie `"pound"` anstelle von `"pounds"`.
- Wenn der umgerechnete Wert `1` ist, verwenden Sie `"kilogram"` anstelle von `"kilograms"`.

# --hints--

`convertToKgs(1)` sollte `"1 pound equals 0.45 kilograms."` zurückgeben.

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)` sollte `"0 pounds equals 0.00 kilograms."` zurückgeben.

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)` sollte `"100 pounds equals 45.36 kilograms."` zurückgeben.

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)` sollte `"2.5 pounds equals 1.13 kilograms."` zurückgeben.

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)` sollte `"2.20462 pounds equals 1.00 kilogram."` zurückgeben.

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
