---
id: 691b559495c5cb5a37b9b481
title: "Aufgabe 121: Am häufigsten"
challengeType: 28
dashedName: challenge-121
---

# --description--

Gegeben ist ein Array von Elementen. Geben Sie das Element zurück, das am häufigsten vorkommt.

- Es wird immer genau ein Element geben, das am häufigsten vorkommt.

# --hints--

`mostFrequent(["a", "b", "a", "c"])` sollte `"a"` zurückgeben.

```js
assert.equal(mostFrequent(["a", "b", "a", "c"]), "a");
```

`mostFrequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9])` sollte `2` zurückgeben.

```js
assert.equal(mostFrequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9]), 2);
```

`mostFrequent([true, false, "false", "true", false])` sollte `false` zurückgeben.

```js
assert.isFalse(mostFrequent([true, false, "false", "true", false]));
```

`mostFrequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60])` sollte `40` zurückgeben.

```js
assert.equal(mostFrequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60]), 40);
```

# --seed--

## --seed-contents--

```js
function mostFrequent(arr) {
  return arr;
}
```

# --solutions--

```js
function mostFrequent(arr) {
  const freq = {};
  let maxCount = 0;
  let mostElem = null;

  for (const el of arr) {
    freq[el] = (freq[el] || 0) + 1;
    if (freq[el] > maxCount) {
      maxCount = freq[el];
      mostElem = el;
    }
  }

  return mostElem;
}
```
