---
id: 691b559495c5cb5a37b9b487
title: "Aufgabe 127: Geschwindigkeitsprüfung"
challengeType: 28
dashedName: challenge-127
---

# --description--

Gegeben ist die Geschwindigkeit, mit der Sie in Meilen pro Stunde (MPH) fahren, und eine Geschwindigkeitsbegrenzung in Kilometern pro Stunde (KPH). Bestimmen Sie, ob Sie zu schnell fahren und ob Sie eine Verwarnung oder ein Bußgeld erhalten.

- 1 Meile entspricht 1,60934 Kilometern.
- Wenn Sie mit einer Geschwindigkeit fahren, die kleiner oder gleich der Geschwindigkeitsbegrenzung ist, geben Sie `"Not Speeding"` zurück.
- Wenn Sie bis zu 5 KPH über der Geschwindigkeitsbegrenzung fahren, geben Sie `"Warning"` zurück.
- Wenn Sie mehr als 5 KPH über der Geschwindigkeitsbegrenzung fahren, geben Sie `"Ticket"` zurück.

# --hints--

`speedCheck(30, 70)` sollte `"Not Speeding"` zurückgeben.

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)` sollte `"Warning"` zurückgeben.

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)` sollte `"Not Speeding"` zurückgeben.

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)` sollte `"Ticket"` zurückgeben.

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)` sollte `"Warning"` zurückgeben.

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)` sollte `"Ticket"` zurückgeben.

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```
