---
id: 691b559495c5cb5a37b9b484
title: "Aufgabe 124: Inventaraktualisierung"
challengeType: 28
dashedName: challenge-124
---

# --description--

Gegeben ist ein 2D-Array, das das Inventar Ihres Geschäfts darstellt, und ein weiteres 2D-Array, das eine erhaltene Lieferung darstellt. Geben Sie Ihr aktualisiertes Inventar zurück.

- Jedes Element in den Arrays hat das Format: `[quantity, "item"]`, wobei `quantity` eine ganze Zahl und `"item"` ein String ist.
- Aktualisieren Sie die Artikel im Inventar, indem Sie die Menge aller passenden Artikel aus der Lieferung hinzufügen.
- Wenn ein empfangener Artikel im aktuellen Inventar nicht existiert, fügen Sie ihn als neuen Eintrag am Ende des Inventars hinzu.
- Geben Sie das Inventar in der Reihenfolge zurück, in der es angegeben wurde, wobei neue Artikel am Ende in der Reihenfolge erscheinen, in der sie in der Lieferung vorkommen.

Zum Beispiel, wenn das Inventar `[[2, "apples"], [5, "bananas"]]` und die Lieferung `[[1, "apples"], [3, "bananas"]]` ist, geben Sie `[[3, "apples"], [8, "bananas"]]` zurück.

# --hints--

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])` sollte `[[3, "apples"], [8, "bananas"]]` zurückgeben.

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]]);
```

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])` sollte `[[3, "apples"], [8, "bananas"], [4, "oranges"]]` zurückgeben.

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]]);
```

`updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])` sollte `[[10, "apples"], [30, "bananas"], [20, "oranges"]]` zurückgeben.

```js
assert.deepEqual(updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]]);
```

`updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])` sollte `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]` zurückgeben.

```js
assert.deepEqual(updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]);
```

# --seed--

## --seed-contents--

```js
function updateInventory(inventory, shipment) {
  return inventory;
}
```

# --solutions--

```js
function updateInventory(inventory, shipment) {
  const inventoryMap = new Map();
  inventory.forEach(([qty, item], index) => {
    inventoryMap.set(item, index);
  });

  shipment.forEach(([qty, item]) => {
    if (inventoryMap.has(item)) {
      const index = inventoryMap.get(item);
      inventory[index][0] += qty;
    } else {
      inventory.push([qty, item]);
      inventoryMap.set(item, inventory.length - 1);
    }
  });

  return inventory;
}
```
