---
id: 691b559495c5cb5a37b9b488
title: "Aufgabe 128: Konsonanten zählen"
challengeType: 28
dashedName: challenge-128
---

# --description--

Gegeben ist ein String und eine Zielzahl. Bestimmen Sie, ob der String genau die Zielzahl an Konsonanten enthält.

- Konsonanten sind alle alphabetischen Zeichen außer `"a"`, `"e"`, `"i"`, `"o"` und `"u"` in beliebiger Groß- oder Kleinschreibung.
- Ignorieren Sie beim Zählen Ziffern, Satzzeichen, Leerzeichen und andere Nicht-Buchstaben-Zeichen.

# --hints--

`hasConsonantCount("helloworld", 7)` sollte `true` zurückgeben.

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)` sollte `false` zurückgeben.

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)` sollte `true` zurückgeben.

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` sollte `false` zurückgeben.

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` sollte `true` zurückgeben.

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
