---
id: 691b5597f0f3e85a588a5c94
title: "Aufgabe 130: Schachbrett"
challengeType: 28
dashedName: challenge-130
---

# --description--

Gegeben ist ein Array mit zwei Zahlen, wobei die erste die Anzahl der Zeilen und die zweite die Anzahl der Spalten angibt. Geben Sie eine Matrix (ein Array von Arrays) zurück, die mit `"X"` und `"O"` Zeichen der angegebenen Größe gefüllt ist.

- Die Zeichen sollen sich wie bei einem Schachbrett abwechseln.
- Die Zelle oben links muss immer `"X"` sein.

Zum Beispiel, bei `[3, 3]`, geben Sie zurück:

```json
[
  ["X", "O", "X"],
  ["O", "X", "O"],
  ["X", "O", "X"]
]
```

# --hints--

`createBoard([3, 3])` sollte `[["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]` zurückgeben.

```js
assert.deepEqual(createBoard([3, 3]), [["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]);
```

`createBoard([6, 1])` sollte `[["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]` zurückgeben.

```js
assert.deepEqual(createBoard([6, 1]), [["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]);
```

`createBoard([2, 10])` sollte `[["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]` zurückgeben.

```js
assert.deepEqual(createBoard([2, 10]), [["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]);
```

`createBoard([5, 4])` sollte `[["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]` zurückgeben.

```js
assert.deepEqual(createBoard([5, 4]), [["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]);
```

# --seed--

## --seed-contents--

```js
function createBoard(dimensions) {
  return dimensions;
}
```

# --solutions--

```js
function createBoard(dimensions) {
  const [rows, cols] = dimensions;
  const board = [];

  for (let r = 0; r < rows; r++) {
    const row = [];
    for (let c = 0; c < cols; c++) {
      row.push((r + c) % 2 === 0 ? "X" : "O");
    }
    board.push(row);
  }

  return board;
}
```
