---
id: 691b559495c5cb5a37b9b482
title: "Aufgabe 122: Markdown fett Parser"
challengeType: 28
dashedName: challenge-122
---

# --description--

Gegeben ist ein String, der möglicherweise fettgedruckten Text in Markdown enthält. Geben Sie den entsprechenden HTML-String zurück.

- Fettgedruckter Text in Markdown ist jeder Text, der mit zwei Sternchen (`**`) oder zwei Unterstrichen (`__`) beginnt und endet.
- Zwischen dem Text und den Sternchen oder Unterstrichen dürfen keine Leerzeichen sein, aber im Text selbst können Leerzeichen enthalten sein.
- Wandeln Sie alle fettgedruckten Vorkommen in HTML-`b`-Tags um und geben Sie den String zurück.

Zum Beispiel, bei Eingabe von `"**This is bold**"` geben Sie `"<b>This is bold</b>"` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings beim Protokollieren von Nachrichten möglicherweise nicht an. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parseBold("**This is bold**")` sollte `"<b>This is bold</b>"` zurückgeben.

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")` sollte `"<b>This is also bold</b>"` zurückgeben.

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")` sollte `"**This is not bold **"` zurückgeben.

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")` sollte `"__ This is also not bold__"` zurückgeben.

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")` sollte `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."` zurückgeben.

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```
