---
id: 691b559495c5cb5a37b9b485
title: "Aufgabe 125: Game of Life"
challengeType: 28
dashedName: challenge-125
---

# --description--

Gegeben ist eine Matrix (Array von Arrays), die den aktuellen Zustand im Game of Life von Conway darstellt. Geben Sie den nächsten Zustand der Matrix zurück, indem Sie diese Regeln anwenden:

- Jede Zelle ist entweder `1` (lebendig) oder `0` (tot).
- Die Nachbarn einer Zelle sind die bis zu acht umliegenden Zellen (vertikal, horizontal und diagonal).
- Zellen am Rand haben weniger als acht Nachbarn.

Regeln zur Aktualisierung jeder Zelle:

- Jede lebendige Zelle mit weniger als zwei lebendigen Nachbarn stirbt (Unterbevölkerung).
- Jede lebendige Zelle mit zwei oder drei lebendigen Nachbarn bleibt am Leben.
- Jede lebendige Zelle mit mehr als drei lebendigen Nachbarn stirbt (Überbevölkerung).
- Jede tote Zelle mit genau drei lebendigen Nachbarn wird lebendig (Fortpflanzung).

Zum Beispiel, gegeben:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

geben Sie zurück:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Jede Zelle wird entsprechend der Anzahl lebendiger Nachbarn aktualisiert. Zum Beispiel bleibt `[0][0]` tot (2 lebendige Nachbarn), `[0][1]` bleibt lebendig (2 lebendige Nachbarn), `[0][2]` stirbt (3 lebendige Nachbarn) und so weiter.

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` sollte `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]` zurückgeben.

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` sollte `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]` zurückgeben.

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` sollte `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]` zurückgeben.

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` sollte `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]` zurückgeben.

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
