---
id: 691b559495c5cb5a37b9b486
title: "Aufgabe 126: Großschreibung"
challengeType: 28
dashedName: challenge-126
---

# --description--

Gegeben ist ein String title. Geben Sie einen neuen String zurück, der im Titel-Format formatiert ist, und zwar nach folgenden Regeln:

- Großschreiben Sie den ersten Buchstaben jedes Wortes.
- Schreiben Sie alle anderen Buchstaben jedes Wortes in Kleinbuchstaben.
- Wörter sind immer durch ein einzelnes Leerzeichen getrennt.

# --hints--

`titleCase("hello world")` sollte `"Hello World"` zurückgeben.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` sollte `"The Quick Brown Fox"` zurückgeben.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` sollte `"Javascript And Python"` zurückgeben.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` sollte `"Avocado Toast For Breakfast"` zurückgeben.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
