---
id: 691b559495c5cb5a37b9b489
title: "Aufgabe 129: Markdown-Blockzitat-Parser"
challengeType: 28
dashedName: challenge-129
---

# --description--

Gegeben ist ein String, der ein Blockzitat in Markdown enthält. Geben Sie den entsprechenden HTML-String zurück.

Ein Blockzitat in Markdown ist jede Zeile, die:

- Mit null oder mehr Leerzeichen beginnt
- Gefolgt von einem Größer-als-Zeichen (`>`)
- Dann einem oder mehreren Leerzeichen
- Und schließlich dem Text des Blockzitats.

Geben Sie den Text des Blockzitats zurück, umgeben von öffnenden und schließenden HTML-`blockquote`-Tags.

Zum Beispiel, gegeben `"> This is a quote"`, geben Sie `<blockquote>This is a quote</blockquote>` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings möglicherweise nicht an, wenn Nachrichten protokolliert werden. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parseBlockquote("> This is a quote")` sollte `"<blockquote>This is a quote</blockquote>"` zurückgeben.

```js
assert.equal(parseBlockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>");
```

`parseBlockquote(" > This is also a quote")` sollte `"<blockquote>This is also a quote</blockquote>"` zurückgeben.

```js
assert.equal(parseBlockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>");
```

`parseBlockquote("  >    So  Is  This")` sollte `"<blockquote>So  Is  This</blockquote>"` zurückgeben.

```js
assert.equal(parseBlockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>");
```

# --seed--

## --seed-contents--

```js
function parseBlockquote(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBlockquote(markdown) {
  const blockquoteRegex = /^\s*>\s+(.+)$/;
  const match = markdown.match(blockquoteRegex);

  const text = match[1];
  return `<blockquote>${text}</blockquote>`;
}
```
