---
id: 6925e2068081f40f549ced1b
title: "Aufgabe 137: Snowflake-Generator"
challengeType: 28
dashedName: challenge-137
---

# --description--

Gegeben ist ein mehrzeiliger String, der Zeilenumbrüche mit dem Zeichen `\n` darstellt. Geben Sie einen neuen String zurück, bei dem jede Zeile horizontal gespiegelt und an das Ende der ursprünglichen Zeile angehängt wird.

- Spiegeln Sie eine Zeile, indem Sie alle ihre Zeichen, einschließlich Leerzeichen, umkehren.

Zum Beispiel, gegeben `"* \n *\n* "`, das in der Konsole wie folgt ausgegeben wird:

```sh
* 
 *
* 
```

Geben Sie `"*  *\n ** \n*  *"` zurück, das in der Konsole wie folgt ausgegeben wird:

```sh
*  *
 ** 
*  *
```

Achten Sie genau auf die Leerzeichen in den gegebenen und zurückgegebenen Strings. Schneiden Sie keine davon ab.

# --hints--

`generateSnowflake("* \n *\n* ")` sollte `"*  *\n ** \n*  *"` zurückgeben.

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")` sollte `"X=~~=X"` zurückgeben.

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")` sollte `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "` zurückgeben.

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")` sollte `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"` zurückgeben.

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")` sollte `"*  --  *\n * -- * \n*  --  *"` zurückgeben.

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```
