---
id: 6925e2068081f40f549ced1c
title: "Aufgabe 138: Summe der Teiler"
challengeType: 28
dashedName: challenge-138
---

# --description--

Gegeben ist eine positive ganze Zahl. Geben Sie die Summe aller ihrer Teiler zurück.

- Ein Teiler ist jede ganze Zahl, die die Zahl ohne Rest teilt (der Rest ist `0`).
- Zählen Sie jeden Teiler nur einmal.

Zum Beispiel, gegeben `6`, geben Sie `12` zurück, weil die Teiler von `6` `1`, `2`, `3` und `6` sind, und die Summe davon `12` ergibt.

# --hints--

`sumDivisors(6)` sollte `12` zurückgeben.

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)` sollte `14` zurückgeben.

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)` sollte `56` zurückgeben.

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)` sollte `224` zurückgeben.

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)` sollte `806` zurückgeben.

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)` sollte `23520` zurückgeben.

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
