---
id: 6925e2068081f40f549ced1d
title: "Aufgabe 139: Schere, Stein, Papier"
challengeType: 28
dashedName: challenge-139
---

# --description--

Gegeben sind zwei Strings, der erste stellt Spieler 1 dar und der zweite Spieler 2. Bestimmen Sie den Gewinner eines Spiels Schere, Stein, Papier.

- Die Eingabestrings sind immer `"Rock"`, `"Paper"` oder `"Scissors"`.
- `"Rock"` schlägt `"Scissors"`.
- `"Paper"` schlägt `"Rock"`.
- `"Scissors"` schlägt `"Paper"`.

Geben Sie zurück:

- `"Player 1 wins"`, wenn Spieler 1 gewinnt.
- `"Player 2 wins"`, wenn Spieler 2 gewinnt.
- `"Tie"`, wenn beide Spieler dieselbe Option wählen.

# --hints--

`rockPaperScissors("Rock", "Rock")` sollte `"Tie"` zurückgeben.

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")` sollte `"Player 2 wins"` zurückgeben.

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")` sollte `"Player 1 wins"` zurückgeben.

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")` sollte `"Player 1 wins"` zurückgeben.

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")` sollte `"Tie"` zurückgeben.

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")` sollte `"Player 2 wins"` zurückgeben.

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
