---
id: 6925e2068081f40f549ced1a
title: "Aufgabe 136: Markdown-Bildparser"
challengeType: 28
dashedName: challenge-136
---

# --description--

Gegeben ist ein String mit einem Bild in Markdown. Geben Sie den entsprechenden HTML-String zurück.

Ein Markdown-Bild hat folgendes Format: `"![alt text](image_url)"`. Dabei gilt:

- `alt text` ist die Beschreibung des Bildes (der Wert des `alt`-Attributs).
- `image_url` ist die Quell-URL des Bildes (der Wert des `src`-Attributs).

Geben Sie einen String mit dem HTML-`img`-Tag zurück, bei dem das `src`-Attribut auf die Bild-URL gesetzt ist und das `alt`-Attribut den Alternativtext enthält.

Zum Beispiel, bei Eingabe von `"![Cute cat](cat.png)"` geben Sie `'<img src="cat.png" alt="Cute cat">'` zurück;

- Achten Sie darauf, dass Tag, Reihenfolge der Attribute, Leerzeichen und Anführungszeichen genau wie oben sind.

Hinweis: Die Konsole zeigt HTML-Tags in Strings beim Protokollieren von Nachrichten möglicherweise nicht an – prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parseImage("![Cute cat](cat.png)")` sollte `'<img src="cat.png" alt="Cute cat">'` zurückgeben.

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` sollte `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'` zurückgeben.

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` sollte `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'` zurückgeben.

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
