---
id: 69272dcf1c24b44fd79137c3
title: "Aufgabe 140: SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

Gegeben ist ein String, der einen Variablennamen darstellt. Geben Sie den Variablennamen zurück, konvertiert in SCREAMING_SNAKE_CASE.

Die gegebenen Variablennamen sind in einem der folgenden Formate geschrieben:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

In den oben genannten Formaten sind Wörter durch einen Unterstrich (`_`), einen Bindestrich (`-`) getrennt oder ein neues Wort beginnt mit einem Großbuchstaben.

Um in SCREAMING_SNAKE_CASE zu konvertieren:

- Machen Sie alle Buchstaben groß
- Trennen Sie Wörter mit einem Unterstrich (`_`)

# --hints--

`toScreamingSnakeCase("userEmail")` sollte `"USER_EMAIL"` zurückgeben.

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` sollte `"USER_PASSWORD"` zurückgeben.

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` sollte `"USER_ID"` zurückgeben.

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` sollte `"USER_ADDRESS"` zurückgeben.

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` sollte `"USERNAME"` zurückgeben.

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

`toScreamingSnakeCase("my_variable_name")` sollte `"MY_VARIABLE_NAME"` zurückgeben.

```js
assert.equal(toScreamingSnakeCase("my_variable_name"), "MY_VARIABLE_NAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();
}
```
