---
id: 69272dcf1c24b44fd79137c6
title: "Aufgabe 143: Markdown-Parser für Kursivschrift"
challengeType: 28
dashedName: challenge-143
---

# --description--

Gegeben ist ein String, der möglicherweise kursiven Text in Markdown enthält. Geben Sie den entsprechenden HTML-String zurück.

- Kursiver Text in Markdown ist jeder Text, der mit einem einzelnen Sternchen (`*`) oder einem einzelnen Unterstrich (`_`) beginnt und endet.
- Zwischen dem Text und dem Sternchen oder Unterstrich dürfen keine Leerzeichen stehen, aber im Text selbst können Leerzeichen enthalten sein.
- Wandeln Sie alle kursiven Vorkommen in HTML-`i`-Tags um und geben Sie den String zurück.

Zum Beispiel, bei Eingabe von `"*This is italic*"` geben Sie `"<i>This is italic</i>"` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings beim Protokollieren von Nachrichten möglicherweise nicht an. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parseItalics("*This is italic*")` sollte `"<i>This is italic</i>"` zurückgeben.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` sollte `"<i>This is also italic</i>"` zurückgeben.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` sollte `"*This is not italic *"` zurückgeben.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` sollte `"_ This is also not italic_"` zurückgeben.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` sollte `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."` zurückgeben.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
