---
id: 69272dcf1c24b44fd79137c5
title: "Aufgabe 142: Summe des Strings"
challengeType: 28
dashedName: challenge-142
---

# --description--

Gegeben ist ein String, der Ziffern und andere Zeichen enthält. Geben Sie die Summe aller Zahlen im String zurück.

- Behandeln Sie aufeinanderfolgende Ziffern als eine einzige Zahl. Zum Beispiel zählt `"13"` als 13, nicht als 1 + 3.
- Ignorieren Sie alle Nicht-Ziffern-Zeichen.

# --hints--

`stringSum("3apples2bananas")` sollte `5` zurückgeben.

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` sollte `17` zurückgeben.

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` sollte `125344` zurückgeben.

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` sollte `321` zurückgeben.

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` sollte `1653` zurückgeben.

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
