---
id: 69272dcf1c24b44fd79137c4
title: "Aufgabe 141: Startkraftstoff"
challengeType: 28
dashedName: challenge-141
---

# --description--

Gegeben sind die Anzahl der Gallonen Kraftstoff, die sich derzeit in Ihrem Flugzeug befinden, und die erforderliche Anzahl Liter Kraftstoff, um Ihr Ziel zu erreichen. Bestimmen Sie, wie viele zusätzliche Gallonen Kraftstoff Sie hinzufügen sollten.

- 1 Gallone entspricht 3,78541 Litern.
- Wenn das Flugzeug bereits genug Kraftstoff hat, geben Sie `0` zurück.
- Sie können nur ganze Gallonen hinzufügen.
- Geben Sie keine Dezimalstellen in der Rückgabe an.

# --hints--

`fuelToAdd(0, 1)` sollte `1` zurückgeben.

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)` sollte `6` zurückgeben.

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)` sollte `0` zurückgeben.

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)` sollte `4520` zurückgeben.

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)` sollte `12209` zurückgeben.

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
