---
id: 69306364df283fcaff2e1ada
title: "Aufgabe 149: vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

Gegeben ist ein String. Geben Sie einen neuen String zurück, bei dem alle Vokale in Großbuchstaben und alle anderen alphabetischen Zeichen in Kleinbuchstaben umgewandelt sind.

- Vokale sind `"a"`, `"e"`, `"i"`, `"o"` und `"u"` in beliebiger Schreibweise.
- Nicht-alphabetische Zeichen bleiben unverändert.

# --hints--

`vowelCase("vowelcase")` sollte `"vOwElcAsE"` zurückgeben.

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` sollte `"cOdIng Is fUn"` zurückgeben.

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` sollte `"hEllO, wOrld!"` zurückgeben.

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` sollte `"gIt chErry-pIck"` zurückgeben.

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` sollte `"hEAd~1"` zurückgeben.

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
