---
id: 69306364df283fcaff2e1ad6
title: "Aufgabe 145: N-te Fibonacci-Zahl"
challengeType: 28
dashedName: challenge-145
---

# --description--

Gegeben ist eine ganze Zahl `n`. Geben Sie die `n`te Zahl in der Fibonacci-Folge zurück.

Die Fibonacci-Folge ist eine Zahlenreihe, bei der jede Zahl die Summe der beiden vorhergehenden ist. Die ersten 10 Zahlen der Folge sind `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

# --hints--

`nthFibonacci(4)` sollte `2` zurückgeben.

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)` sollte `34` zurückgeben.

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)` sollte `377` zurückgeben.

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)` sollte `63245986` zurückgeben.

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)` sollte `1304969544928657` zurückgeben.

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
