---
id: 69306364df283fcaff2e1ad8
title: "Aufgabe 147: Schaltjahrrechner"
challengeType: 28
dashedName: challenge-147
---

# --description--

Gegeben ist ein ganzzahliges Jahr. Bestimmen Sie, ob es ein Schaltjahr ist.

Ein Jahr ist ein Schaltjahr, wenn es die folgenden Regeln erfüllt:

- Das Jahr ist durch 4 ohne Rest teilbar, und
- Das Jahr ist nicht durch 100 ohne Rest teilbar, außer
- Das Jahr ist durch 400 ohne Rest teilbar.

# --hints--

`isLeapYear(2024)` sollte `true` zurückgeben.

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)` sollte `false` zurückgeben.

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)` sollte `false` zurückgeben.

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)` sollte `true` zurückgeben.

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)` sollte `false` zurückgeben.

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)` sollte `true` zurückgeben.

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)` sollte `false` zurückgeben.

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
