---
id: 69306364df283fcaff2e1ad7
title: "Aufgabe 146: Linkshändiger Sitz an der Tabelle"
challengeType: 28
dashedName: challenge-146
---

# --description--

Gegeben ist eine 4x2-Matrix (Array von Arrays), die die Sitzordnung für eine Mahlzeit darstellt. Bestimmen Sie, wie viele Plätze ein linkshändiger Mensch einnehmen kann.

- Ein linkshändiger Mensch darf nicht dort sitzen, wo sich unmittelbar links neben ihm ein rechtshändiger Mensch befindet.

In der gegebenen Matrix:

- Ein `"R"` ist ein Sitz, der von einem rechtshändigen Menschen besetzt ist.
- Ein `"L"` ist ein Sitz, der von einem linkshändigen Menschen besetzt ist.
- Ein `"U"` ist ein unbesetzter Sitz.
- Nur unbesetzte Sitze sind verfügbar.
- Die Sitze in der oberen Reihe sind nach „unten“ ausgerichtet, und die Sitze in der unteren Reihe sind nach „oben“ ausgerichtet (wie an einem Tisch), daher sind links und rechts relativ zur Sitzorientierung zu verstehen.
- Ecksitze haben nur einen angrenzenden Sitz.

Zum Beispiel in der gegebenen Matrix:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

Der Sitz oben links kann nicht besetzt werden, weil sich links davon ein rechtshändiger Mensch befindet. Die anderen beiden freien Sitze können besetzt werden, da sich links davon kein rechtshändiger Mensch befindet.

# --hints--

`findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` sollte `2` zurückgeben.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2);
```

`findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` sollte `8` zurückgeben.

```js
assert.equal(findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8);
```

`findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` sollte `0` zurückgeben.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0);
```

`findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` sollte `1` zurückgeben.

```js
assert.equal(findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1);
```

`findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` sollte `5` zurückgeben.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5);
```

# --seed--

## --seed-contents--

```js
function findLeftHandedSeats(table) {

  return table;
}
```

# --solutions--

```js
function findLeftHandedSeats(table) {
  let availableSeats = 0;

  const [topRow, bottomRow] = table;

  for (let i=0; i<topRow.length; i++) {
    if (topRow[i] === 'U') {
      if (i === 3) {
        availableSeats++;
      } else if (topRow[i+1] !== 'R') {
        availableSeats++;
      }
    }
  }

  for (let i=0; i<bottomRow.length; i++) {
    if (bottomRow[i] === 'U') {
      if (i === 0) {
        availableSeats++;
      } else if (bottomRow[i-1] !== 'R') {
        availableSeats++;
      }
    }
  }

  return availableSeats;
}
```
