---
id: 69373793f5a867f769cde136
title: "Aufgabe 151: Sortiertes Array?"
challengeType: 28
dashedName: challenge-151
---

# --description--

Gegeben ist ein Array von Zahlen. Bestimmen Sie, ob die Zahlen in aufsteigender Reihenfolge, absteigender Reihenfolge oder in keiner von beiden sortiert sind.

Wenn das gegebene Array:

- In aufsteigender Reihenfolge (von niedrig nach hoch) ist, geben Sie `"Ascending"` zurück.
- In absteigender Reihenfolge (von hoch nach niedrig) ist, geben Sie `"Descending"` zurück.
- Weder in aufsteigender noch in absteigender Reihenfolge sortiert ist, geben Sie `"Not sorted"` zurück.

# --hints--

`isSorted([1, 2, 3, 4, 5])` sollte `"Ascending"` zurückgeben.

```js
assert.equal(isSorted([1, 2, 3, 4, 5]), "Ascending");
```

`isSorted([10, 8, 6, 4, 2])` sollte `"Descending"` zurückgeben.

```js
assert.equal(isSorted([10, 8, 6, 4, 2]), "Descending");
```

`isSorted([1, 3, 2, 4, 5])` sollte `"Not sorted"` zurückgeben.

```js
assert.equal(isSorted([1, 3, 2, 4, 5]), "Not sorted");
```

`isSorted([3.14, 2.71, 1.61, 0.57])` sollte `"Descending"` zurückgeben.

```js
assert.equal(isSorted([3.14, 2.71, 1.61, 0.57]), "Descending");
```

`isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])` sollte `"Ascending"` zurückgeben.

```js
assert.equal(isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending");
```

`isSorted([0.4, 0.5, 0.3])` sollte `"Not sorted"` zurückgeben.

```js
assert.equal(isSorted([0.4, 0.5, 0.3]), "Not sorted");
```

# --seed--

## --seed-contents--

```js
function isSorted(arr) {

  return arr;
}
```

# --solutions--

```js
function isSorted(arr) {
  let ascending = true;
  let descending = true;

  for (let i = 1; i < arr.length; i++) {
    if (arr[i] < arr[i - 1]) ascending = false;
    if (arr[i] > arr[i - 1]) descending = false;
  }

  if (ascending) return "Ascending";
  if (descending) return "Descending";
  return "Not sorted";
}
```
