---
id: 69373793f5a867f769cde138
title: "Aufgabe 153: Tic-Tac-Toe"
challengeType: 28
dashedName: challenge-153
---

# --description--

Gegeben ist eine 3×3-Matrix (ein Array von Arrays), die ein abgeschlossenes Tic-Tac-Toe-Spiel darstellt. Bestimmen Sie den Gewinner.

- Jedes Element in der gegebenen Matrix ist entweder ein `"X"` oder `"O"`.

Ein Spieler gewinnt, wenn er drei seiner Zeichen in einer Reihe hat – horizontal, vertikal oder diagonal.

Geben Sie zurück:

- `"X wins"`, wenn Spieler X drei in einer Reihe hat.
- `"O wins"`, wenn Spieler O drei in einer Reihe hat.
- `"Draw"`, wenn kein Spieler drei in einer Reihe hat.

# --hints--

`ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` sollte `"X wins"` zurückgeben.

```js
assert.equal(ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins");
```

`ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]])` sollte `"X wins"` zurückgeben.

```js
assert.equal(ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` sollte `"Draw"` zurückgeben.

```js
assert.equal(ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw");
```

`ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]])` sollte `"O wins"` zurückgeben.

```js
assert.equal(ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` sollte `"X wins"` zurückgeben.

```js
assert.equal(ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` sollte `"Draw"` zurückgeben.

```js
assert.equal(ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw");
```

# --seed--

## --seed-contents--

```js
function ticTacToe(board) {

  return board;
}
```

# --solutions--

```js
function ticTacToe(board) {
  const lines = [];

  for (let row of board) {
    lines.push(row);
  }

  for (let c = 0; c < 3; c++) {
    lines.push([board[0][c], board[1][c], board[2][c]]);
  }

  lines.push([board[0][0], board[1][1], board[2][2]]);
  lines.push([board[0][2], board[1][1], board[2][0]]);

  for (let line of lines) {
    if (line.every(cell => cell === "X")) return "X wins";
    if (line.every(cell => cell === "O")) return "O wins";
  }

  return "Draw";
}
```
