---
id: 69373793f5a867f769cde139
title: "Aufgabe 154: Par für das Loch"
challengeType: 28
dashedName: challenge-154
---

# --description--

Gegeben sind zwei ganze Zahlen: das Par für ein Golfloch und die Anzahl der Striche, die ein Golfer für dieses Loch benötigt hat. Geben Sie die Punktzahl des Golfers mit Golfterminen zurück.

Geben Sie zurück:

- `"Hole in one!"`, wenn es einen Strich gebraucht hat.
- `"Eagle"`, wenn es zwei Striche weniger als Par gebraucht hat.
- `"Birdie"`, wenn es einen Strich weniger als Par gebraucht hat.
- `"Par"`, wenn es genauso viele Striche wie Par gebraucht hat.
- `"Bogey"`, wenn es einen Strich mehr als Par gebraucht hat.
- `"Double bogey"`, wenn es zwei Striche mehr als Par gebraucht hat.

# --hints--

`golfScore(3, 3)` sollte `"Par"` zurückgeben.

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)` sollte `"Birdie"` zurückgeben.

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)` sollte `"Hole in one!"` zurückgeben.

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)` sollte `"Double bogey"` zurückgeben.

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)` sollte `"Bogey"` zurückgeben.

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)` sollte `"Eagle"` zurückgeben.

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```
