---
id: 69373793f5a867f769cde135
title: "Aufgabe 150: Markdown-Parser für ungeordnete Listen"
challengeType: 28
dashedName: challenge-150
---

# --description--

Gegeben ist der String einer gültigen ungeordneten Liste in Markdown. Geben Sie den entsprechenden HTML-String zurück.

Eine ungeordnete Liste besteht aus einem oder mehreren Listenelementen. Ein gültiges Listenelement steht in einer eigenen Zeile und:

- Beginnt mit einem Bindestrich (`"-"`), gefolgt von
- Mindestens einem Leerzeichen, und dann
- Dem Text des Listenelements.

Die Liste wird als ein einzelner String mit Zeilenumbrüchen dargestellt, die durch das Newline-Zeichen (`"\n"`) getrennt sind. Die Newline-Zeichen dürfen nicht im Text des Listenelements enthalten sein.

Fassen Sie jedes Listenelement in HTML-`li`-Tags ein und die gesamte Liste der Elemente in `ul`-Tags.

Beispielsweise soll bei Eingabe von `"- Item A\n- Item B"` `"<ul><li>Item A</li><li>Item B</li></ul>"` zurückgegeben werden.

Hinweis: Die Konsole zeigt HTML-Tags in Strings möglicherweise nicht an, wenn Nachrichten protokolliert werden. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parseUnorderedList("- Item A\n- Item B")` sollte `"<ul><li>Item A</li><li>Item B</li></ul>"` zurückgeben.

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("-  JavaScript\n-  Python")` sollte `"<ul><li>JavaScript</li><li>Python</li></ul>"` zurückgeben.

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` sollte `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"` zurückgeben.

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")` sollte `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"` zurückgeben.

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```
