---
id: 69373793f5a867f769cde13a
title: "Aufgabe 155: Die Pflanze setzen"
challengeType: 28
dashedName: challenge-155
---

# --description--

Gegeben ist eine ganze Zahl, die die Größe Ihres Bauernfelds darstellt, sowie `"acres"` oder `"hectares"`, die die Einheit für die Größe Ihres Bauernfelds angeben, und eine Pflanzenart. Bestimmen Sie, wie viele Pflanzen dieser Art auf Ihr Feld passen.

- 1 Acre entspricht 4046,86 Quadratmetern.
- 1 Hektar entspricht 10.000 Quadratmetern.

Hier ist eine Liste von Pflanzen, die als Eingabe gegeben werden, und wie viel Platz eine einzelne Pflanze benötigt:

|Pflanze|Platz pro Pflanze|
|-|-|
|`"corn"`|1 Quadratmeter|
|`"wheat"`|0,1 Quadratmeter|
|`"soybeans"`|0,5 Quadratmeter|
|`"tomatoes"`|0,25 Quadratmeter|
|`"lettuce"`|0,2 Quadratmeter|

Geben Sie die Anzahl der Pflanzen zurück, die auf das Feld passen, abgerundet auf die nächstkleinere ganze Pflanze.

# --hints--

`getNumberOfPlants(1, "acres", "corn")` sollte `4046` zurückgeben.

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")` sollte `100000` zurückgeben.

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")` sollte `161874` zurückgeben.

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")` sollte `150000` zurückgeben.

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")` sollte `271139` zurückgeben.

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
