---
id: 69373793f5a867f769cde137
title: "Aufgabe 152: Zirkuläre Primzahl"
challengeType: 28
dashedName: challenge-152
---

# --description--

Gegeben ist eine ganze Zahl. Bestimmen Sie, ob sie eine zirkuläre Primzahl ist.

Eine zirkuläre Primzahl ist eine ganze Zahl, bei der alle Drehungen ihrer Ziffern selbst Primzahlen sind.

Zum Beispiel ist `197` eine zirkuläre Primzahl, weil alle Drehungen ihrer Ziffern: `197`, `971` und `719` Primzahlen sind.

# --hints--

`isCircularPrime(197)` sollte `true` zurückgeben.

```js
assert.isTrue(isCircularPrime(197));
```

`isCircularPrime(23)` sollte `false` zurückgeben.

```js
assert.isFalse(isCircularPrime(23));
```

`isCircularPrime(13)` sollte `true` zurückgeben.

```js
assert.isTrue(isCircularPrime(13));
```

`isCircularPrime(89)` sollte `false` zurückgeben.

```js
assert.isFalse(isCircularPrime(89));
```

`isCircularPrime(1193)` sollte `true` zurückgeben.

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
