---
id: 6939b873185d8e00d453563f
title: "Aufgabe 160: Springerzüge"
challengeType: 28
dashedName: challenge-160
---

# --description--

Gegeben ist die Position eines Springers auf einem Schachbrett. Geben Sie die Anzahl der gültigen Felder zurück, auf die der Springer ziehen kann.

Ein Standard-Schachbrett ist 8x8 groß, mit Spalten, die von `A` bis `H` beschriftet sind (von links nach rechts), und Reihen, die von `1` bis `8` beschriftet sind (von unten nach oben). Es sieht so aus:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Ein Springer zieht in einer „L“-Form: zwei Felder in eine Richtung (horizontal oder vertikal) und ein Feld in die senkrechte Richtung.

Das bedeutet, ein Springer kann bis zu acht mögliche Positionen erreichen, aber weniger, wenn er sich nahe am Rand des Brettes befindet. Zum Beispiel könnte ein Springer auf `A1` nur auf `B3` oder `C2` ziehen.

# --hints--

`knightMoves("A1")` sollte `2` zurückgeben.

```js
assert.equal(knightMoves("A1"), 2);
```

`knightMoves("D4")` sollte `8` zurückgeben.

```js
assert.equal(knightMoves("D4"), 8);
```

`knightMoves("G6")` sollte `6` zurückgeben.

```js
assert.equal(knightMoves("G6"), 6);
```

`knightMoves("B8")` sollte `3` zurückgeben.

```js
assert.equal(knightMoves("B8"), 3);
```

`knightMoves("H3")` sollte `4` zurückgeben.

```js
assert.equal(knightMoves("H3"), 4);
```

# --seed--

## --seed-contents--

```js
function knightMoves(position) {

  return position;
}
```

# --solutions--

```js
function knightMoves(position) {
  const col = position[0].charCodeAt(0) - 65;
  const row = parseInt(position[1]) - 1;

  const moves = [
    [2, 1], [2, -1],
    [-2, 1], [-2, -1],
    [1, 2], [1, -2],
    [-1, 2], [-1, -2]
  ];

  let valid = 0;

  for (const [dx, dy] of moves) {
    const newCol = col + dx;
    const newRow = row + dy;

    if (newCol >= 0 && newCol < 8 && newRow >= 0 && newRow < 8) {
      valid++;
    }
  }

  return valid;
}
```
