---
id: 6939b873185d8e00d453563b
title: "Aufgabe 156: Ungerade oder gerade?"
challengeType: 28
dashedName: challenge-156
---

# --description--

Gegeben ist eine positive ganze Zahl. Geben Sie `"Odd"` zurück, wenn sie ungerade ist, und `"Even"`, wenn sie gerade ist.

# --hints--

`oddOrEven(1)` sollte `"Odd"` zurückgeben.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` sollte `"Even"` zurückgeben.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` sollte `"Odd"` zurückgeben.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` sollte `"Even"` zurückgeben.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` sollte `"Odd"` zurückgeben.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
