---
id: 6939b873185d8e00d453563d
title: "Aufgabe 158: Array tauschen"
challengeType: 28
dashedName: challenge-158
---

# --description--

Gegeben ist ein Array mit zwei Werten. Geben Sie ein Array zurück, in dem die Werte vertauscht sind.

Zum Beispiel, wenn `["A", "B"]` gegeben ist, geben Sie `["B", "A"]` zurück.

# --hints--

`arraySwap(["A", "B"])` sollte `["B", "A"]` zurückgeben.

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` sollte `[20, 25]` zurückgeben.

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` sollte `[false, true]` zurückgeben.

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` sollte `[1, "1"]` zurückgeben.

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
