---
id: 6939b873185d8e00d453563e
title: "Aufgabe 159: Ganze Zahl als Hypotenuse"
challengeType: 28
dashedName: challenge-159
---

# --description--

Gegeben sind zwei positive ganze Zahlen, die die Längen der beiden Katheten (die beiden kurzen Seiten) eines rechtwinkligen Dreiecks darstellen. Bestimmen Sie, ob die Hypotenuse eine ganze Zahl ist.

Die Länge der Hypotenuse wird berechnet, indem die Quadrate der beiden Kathetenlängen addiert und anschließend die Quadratwurzel dieser Summe gezogen wird (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

`isIntegerHypotenuse(3, 4)` sollte `true` zurückgeben.

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)` sollte `false` zurückgeben.

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)` sollte `true` zurückgeben.

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)` sollte `false` zurückgeben.

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)` sollte `true` zurückgeben.

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)` sollte `false` zurückgeben.

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
