---
id: 6939b873185d8e00d453563c
title: "Aufgabe 157: Markdown-Link-Parser"
challengeType: 28
dashedName: challenge-157
---

# --description--

Gegeben ist der String eines Links in Markdown. Geben Sie den entsprechenden HTML-String zurück.

Ein Markdown-Bild hat folgendes Format: `"[link_text](link_url)"`. Geben Sie den String des HTML-`a`-Tags zurück, wobei `href` auf `link_url` gesetzt ist und `link_text` als Inhalt des Tags dient.

Zum Beispiel, gegeben `"[freeCodeCamp](https://freecodecamp.org/)"`, geben Sie `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'` zurück;

Hinweis: Die Konsole zeigt HTML-Tags in Strings beim Protokollieren von Nachrichten möglicherweise nicht an – prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")` sollte `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'` zurückgeben.

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` sollte `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'` zurückgeben.

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` sollte `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'` zurückgeben.

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```
