---
id: 694596b0585c11170ac7c7fa
title: "Aufgabe 162: Energieverbrauch"
challengeType: 28
dashedName: challenge-162
---

# --description--

Gegeben ist die Anzahl der während eines Trainings verbrannten Kalorien und die Anzahl der Wattstunden, die Ihre elektronischen Geräte während dieses Trainings verbraucht haben. Bestimmen Sie, welches von beiden mehr Energie verbraucht hat.

Um sie zu vergleichen, wandeln Sie beide Werte mit den folgenden Umrechnungen in Joule um:

- 1 Kalorie entspricht 4184 Joule.
- 1 Wattstunde entspricht 3600 Joule.

Geben Sie zurück:

- `"Workout"`, wenn das Training mehr Energie verbraucht hat.
- `"Devices"`, wenn das Gerät mehr Energie verbraucht hat.
- `"Equal"`, wenn beide die gleiche Energiemenge verbraucht haben.

# --hints--

`compareEnergy(250, 50)` sollte `"Workout"` zurückgeben.

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` sollte `"Devices"` zurückgeben.

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` sollte `"Equal"` zurückgeben.

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` sollte `"Workout"` zurückgeben.

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` sollte `"Devices"` zurückgeben.

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` sollte `"Equal"` zurückgeben.

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
