---
id: 694596b0585c11170ac7c7fb
title: "Aufgabe 163: Konsonanten-Großschreibung"
challengeType: 28
dashedName: challenge-163
---

# --description--

Gegeben ist ein String, der einen Variablennamen darstellt. Wandeln Sie ihn in Konsonanten-Großschreibung um, indem Sie folgende Regeln anwenden:

- Alle Konsonanten sollen in Großbuchstaben umgewandelt werden.
- Alle Vokale (`a`, `e`, `i`, `o`, `u` in beliebiger Schreibweise) sollen in Kleinbuchstaben umgewandelt werden.
- Alle Bindestriche (`-`) sollen in Unterstriche (`_`) umgewandelt werden.

# --hints--

`toConsonantCase("helloworld")` sollte `"HeLLoWoRLD"` zurückgeben.

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")` sollte `"HeLLoWoRLD"` zurückgeben.

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")` sollte `"_HeLLo_WoRLD_"` zurückgeben.

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")` sollte `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"` zurückgeben.

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
