---
id: 694596b0585c11170ac7c7fc
title: "Aufgabe 164: Markdown Inline-Code-Parser"
challengeType: 28
dashedName: challenge-164
---

# --description--

Gegeben ist ein Markdown-String, der einen oder mehrere inline Code-Blöcke enthält. Geben Sie den entsprechenden HTML-String zurück.

Inline-Code-Blöcke in Markdown verwenden einen einzelnen Backtick (``` ` ```) am Anfang und Ende des Code-Block-Texts.

Geben Sie den gegebenen String zurück, wobei alle Code-Blöcke in HTML-`code`-Tags umgewandelt sind.

Wenn zum Beispiel der String ```"Use `let` to declare the variable."``` gegeben ist, geben Sie `"Use <code>let</code> to declare the variable."` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings möglicherweise nicht an, wenn Nachrichten protokolliert werden. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

```parseInlineCode("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
assert.equal(parseInlineCode("Use `let` to declare the variable."), "Use <code>let</code> to declare the variable.");
```

```parseInlineCode("Use `let` or `const` to declare a variable.")``` should return `"Use <code>let</code> or <code>const</code> to declare a variable."`.

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

```parseInlineCode("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
assert.equal(parseInlineCode("Run `npm install` then `npm start`."), "Run <code>npm install</code> then <code>npm start</code>.");
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```
