---
id: 696655d24b614176d4c9b78a
title: "Aufgabe 167: Bingo! Buchstabe"
challengeType: 28
dashedName: challenge-167
---

# --description--

Gibt eine Zahl vor, geben Sie den dazugehörigen Bingo-Buchstaben zurück (großgeschrieben). Bingo-Zahlen sind wie folgt gruppiert:

| Buchstabe | Zahlenbereich |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

# --hints--

`getBingoLetter(75)` sollte `"O"` zurückgeben.

```js
assert.equal(getBingoLetter(75), "O");
```

`getBingoLetter(54)` sollte `"G"` zurückgeben.

```js
assert.equal(getBingoLetter(54), "G");
```

`getBingoLetter(25)` sollte `"I"` zurückgeben.

```js
assert.equal(getBingoLetter(25), "I");
```

`getBingoLetter(38)` sollte `"N"` zurückgeben.

```js
assert.equal(getBingoLetter(38), "N");
```

`getBingoLetter(11)` sollte `"B"` zurückgeben.

```js
assert.equal(getBingoLetter(11), "B");
```

# --seed--

## --seed-contents--

```js
function getBingoLetter(n) {

  return n;
}
```

# --solutions--

```js
function getBingoLetter(n) {
  if (n >= 1 && n <= 15) return "B";
  if (n >= 16 && n <= 30) return "I";
  if (n >= 31 && n <= 45) return "N";
  if (n >= 46 && n <= 60) return "G";
  if (n >= 61 && n <= 75) return "O";
  return null;
}
```
