---
id: 696655d24b614176d4c9b78d
title: "Aufgabe 170: Ungerader oder gerader Tag"
challengeType: 28
dashedName: challenge-170
---

# --description--

Gegeben ist ein Zeitstempel (Anzahl der Millisekunden seit dem Unix-Epoch). Geben Sie zurück:

- `"odd"`, wenn der Tag des Monats für diesen Zeitstempel ungerade ist.
- `"even"`, wenn der Tag des Monats für diesen Zeitstempel gerade ist.

Zum Beispiel, bei `1769472000000`, einem Zeitstempel für den 27. Januar 2026, geben Sie `"odd"` zurück, weil die Tageszahl (`27`) eine ungerade Zahl ist.

Hinweis: Der Zeitstempel ist in Millisekunden angegeben, und Sie sollten das Datum in der UTC-Zeitzone verwenden, nicht in Ihrer lokalen Zeit.

# --hints--

`oddOrEvenDay(1769472000000)` sollte `"odd"` zurückgeben.

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` sollte `"even"` zurückgeben.

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` sollte `"odd"` zurückgeben.

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` sollte `"odd"` zurückgeben.

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` sollte `"even"` zurückgeben.

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
