---
id: 696655d24b614176d4c9b78b
title: "Aufgabe 168: Skaliertes Bild"
challengeType: 28
dashedName: challenge-168
---

# --description--

Gegeben ist ein String, der die Breite und Höhe eines Bildes darstellt, sowie eine Zahl, um das Bild zu skalieren. Geben Sie die skalierten Breite und Höhe zurück.

- Der Eingabe-String hat das Format `"WxH"`. Zum Beispiel `"800x600"`.
- Der Skalierungsfaktor ist eine Zahl, mit der Breite und Höhe multipliziert werden.

Geben Sie die skalierten Abmessungen im gleichen `"WxH"`-Format zurück.

# --hints--

`scaleImage("800x600", 2)` sollte `"1600x1200"` zurückgeben.

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

`scaleImage("100x100", 10)` sollte `"1000x1000"` zurückgeben.

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

`scaleImage("1024x768", 0.5)` sollte `"512x384"` zurückgeben.

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

`scaleImage("300x200", 1.5)` sollte `"450x300"` zurückgeben.

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```
