---
id: 696655d24b614176d4c9b789
title: "Aufgabe 166: Hex-Validator"
challengeType: 28
dashedName: challenge-166
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob es sich um eine gültige CSS-Hex-Farbe handelt. Eine gültige CSS-Hex-Farbe muss:

- mit einem `#` beginnen und
- entweder von 3 oder 6 hexadezimalen Zeichen gefolgt sein.

Hexadezimale Zeichen sind die Zahlen `0` bis `9` und die Buchstaben `a` bis `f` (Groß- und Kleinschreibung ist egal).

# --hints--

`isValidHex("#123")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
