---
id: 696655d24b614176d4c9b78c
title: "Aufgabe 169: FizzBuzz Mini"
challengeType: 28
dashedName: challenge-169
---

# --description--

Gegeben ist eine ganze Zahl. Geben Sie einen String zurück, der auf den folgenden Regeln basiert:

- Wenn die Zahl durch 3 teilbar ist, geben Sie `"Fizz"` zurück.
- Wenn die Zahl durch 5 teilbar ist, geben Sie `"Buzz"` zurück.
- Wenn die Zahl sowohl durch 3 als auch durch 5 teilbar ist, geben Sie `"FizzBuzz"` zurück.
- Andernfalls geben Sie die gegebene Zahl als String zurück.

# --hints--

`fizzBuzzMini(3)` sollte `"Fizz"` zurückgeben.

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` sollte `"4"` zurückgeben.

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` sollte `"Buzz"` zurückgeben.

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` sollte `"FizzBuzz"` zurückgeben.

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` sollte `"98"` zurückgeben.

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
