---
id: 69738771fb5a7b8b24cca2a0
title: "Aufgabe 174: Sternzeichenfinder"
challengeType: 28
dashedName: challenge-174
---

# --description--

Gegeben ist ein Datumsstring im Format `"YYYY-MM-DD"`. Geben Sie das Sternzeichen für dieses Datum anhand des folgenden Diagramms zurück:

| Datumsbereich | Sternzeichen |
| - | - |
| 21. März – 19. April | `"Aries"` |
| 20. April – 20. Mai | `"Taurus"` |
| 21. Mai – 20. Juni | `"Gemini"` |
| 21. Juni – 22. Juli | `"Cancer"` |
| 23. Juli – 22. August | `"Leo"` |
| 23. August – 22. September | `"Virgo"` |
| 23. September – 22. Oktober | `"Libra"` |
| 23. Oktober – 21. November | `"Scorpio"` |
| 22. November – 21. Dezember | `"Sagittarius"` |
| 22. Dezember – 19. Januar | `"Capricorn"` |
| 20. Januar – 18. Februar | `"Aquarius"` |
| 19. Februar – 20. März | `"Pisces"` |

- Die Sternzeichen basieren nur auf Monat und Tag, das Jahr können Sie ignorieren.

# --hints--

`getSign("2026-01-31")` sollte `"Aquarius"` zurückgeben.

```js
assert.equal(getSign("2026-01-31"), "Aquarius");
```

`getSign("2001-06-10")` sollte `"Gemini"` zurückgeben.

```js
assert.equal(getSign("2001-06-10"), "Gemini");
```

`getSign("1985-09-07")` sollte `"Virgo"` zurückgeben.

```js
assert.equal(getSign("1985-09-07"), "Virgo");
```

`getSign("2023-03-19")` sollte `"Pisces"` zurückgeben.

```js
assert.equal(getSign("2023-03-19"), "Pisces");
```

`getSign("2045-11-05")` sollte `"Scorpio"` zurückgeben.

```js
assert.equal(getSign("2045-11-05"), "Scorpio");
```

`getSign("1985-12-06")` sollte `"Sagittarius"` zurückgeben.

```js
assert.equal(getSign("1985-12-06"), "Sagittarius");
```

`getSign("2025-12-30")` sollte `"Capricorn"` zurückgeben.

```js
assert.equal(getSign("2025-12-30"), "Capricorn");
```

`getSign("2018-10-08")` sollte `"Libra"` zurückgeben.

```js
assert.equal(getSign("2018-10-08"), "Libra");
```

`getSign("1958-05-04")` sollte `"Taurus"` zurückgeben.

```js
assert.equal(getSign("1958-05-04"), "Taurus");
```

# --seed--

## --seed-contents--

```js
function getSign(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getSign(dateStr) {
  const [, monthStr, dayStr] = dateStr.split("-");
  const month = Number(monthStr);
  const day = Number(dayStr);

  if ((month === 3 && day >= 21) || (month === 4 && day <= 19)) return "Aries";
  if ((month === 4 && day >= 20) || (month === 5 && day <= 20)) return "Taurus";
  if ((month === 5 && day >= 21) || (month === 6 && day <= 20)) return "Gemini";
  if ((month === 6 && day >= 21) || (month === 7 && day <= 22)) return "Cancer";
  if ((month === 7 && day >= 23) || (month === 8 && day <= 22)) return "Leo";
  if ((month === 8 && day >= 23) || (month === 9 && day <= 22)) return "Virgo";
  if ((month === 9 && day >= 23) || (month === 10 && day <= 22)) return "Libra";
  if ((month === 10 && day >= 23) || (month === 11 && day <= 21)) return "Scorpio";
  if ((month === 11 && day >= 22) || (month === 12 && day <= 21)) return "Sagittarius";
  if ((month === 12 && day >= 22) || (month === 1 && day <= 19)) return "Capricorn";
  if ((month === 1 && day >= 20) || (month === 2 && day <= 18)) return "Aquarius";
  if ((month === 2 && day >= 19) || (month === 3 && day <= 20)) return "Pisces";
}
```
