---
id: 69738771fb5a7b8b24cca2a3
title: "Aufgabe 177: String-Spiegel"
challengeType: 28
dashedName: challenge-177
---

# --description--

Gegeben ist ein String. Geben Sie einen neuen String zurück, der aus dem gegebenen String besteht, an dessen Ende eine umgekehrte Kopie von ihm angehängt wurde.

# --hints--

`mirror("freeCodeCamp")` sollte `"freeCodeCamppmaCedoCeerf"` zurückgeben.

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` sollte `"RaceCarraCecaR"` zurückgeben.

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` sollte `"helloworlddlrowolleh"` zurückgeben.

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` sollte `"The quick brown fox......xof nworb kciuq ehT"` zurückgeben.

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
