---
id: 69738771fb5a7b8b24cca2a1
title: "Aufgabe 175: Digitaler Detox"
challengeType: 28
dashedName: challenge-175
---

# --description--

Gegeben ist ein Array Ihrer Login-Protokolle. Bestimmen Sie, ob Sie Ihr Ziel für den digitalen Detox erreicht haben.

Jedes Protokoll ist ein String im Format `"YYYY-MM-DD HH:mm:ss"`.

Sie haben Ihr Ziel für den digitalen Detox erreicht, wenn beide der folgenden Aussagen zutreffen:

- Sie haben sich innerhalb eines beliebigen Vier-Stunden-Zeitraums höchstens einmal eingeloggt.
- Sie haben sich an keinem einzelnen Tag mehr als zweimal eingeloggt.

# --hints--

`digitalDetox(["2026-02-01 08:00:00", "2026-02-01 12:30:00"])` sollte `true` zurückgeben.

```js
assert.isTrue(digitalDetox(["2026-02-01 08:00:00", "2026-02-01 12:30:00"]));
```

`digitalDetox(["2026-02-01 04:00:00", "2026-02-01 07:30:00"])` sollte `false` zurückgeben.

```js
assert.isFalse(digitalDetox(["2026-02-01 04:00:00", "2026-02-01 07:30:00"]));
```

`digitalDetox(["2026-01-31 08:21:30", "2026-01-31 14:30:00", "2026-02-01 08:00:00", "2026-02-01 12:30:00"])` sollte `true` zurückgeben.

```js
assert.isTrue(digitalDetox(["2026-01-31 08:21:30", "2026-01-31 14:30:00", "2026-02-01 08:00:00", "2026-02-01 12:30:00"]));
```

`digitalDetox(["2026-01-31 10:40:21", "2026-01-31 15:19:41", "2026-01-31 21:49:50", "2026-02-01 09:30:00"])` sollte `false` zurückgeben.

```js
assert.isFalse(digitalDetox(["2026-01-31 10:40:21", "2026-01-31 15:19:41", "2026-01-31 21:49:50", "2026-02-01 09:30:00"]));
```

`digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 09:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"])` sollte `true` zurückgeben.

```js
assert.isTrue(digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 20:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 09:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"]));
```

`digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 01:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"])` sollte `false` zurückgeben.

```js
assert.isFalse(digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 01:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"]));
```

# --seed--

## --seed-contents--

```js
function digitalDetox(logs) {

  return logs;
}
```

# --solutions--

```js
function digitalDetox(logs) {
  if (logs.length === 0) return true;

  const times = logs
    .map(log => new Date(log.replace(" ", "T") + "Z").getTime())
    .sort((a, b) => a - b);

  const ONE_DAY = 24 * 60 * 60 * 1000;
  const FOUR_HOURS = 4 * 60 * 60 * 1000;

  for (let i = 1; i < times.length; i++) {
    if (times[i] - times[i - 1] < FOUR_HOURS) {
      return false;
    }
  }

  const dailyCounts = {};
  for (const time of times) {
    const day = new Date(time).toISOString().slice(0, 10);
    dailyCounts[day] = (dailyCounts[day] || 0) + 1;
    if (dailyCounts[day] > 2) {
      return false;
    }
  }

  return true;
}
```
