---
id: 69738771fb5a7b8b24cca29e
title: "Aufgabe 172: Buchstaben und Zahlen"
challengeType: 28
dashedName: challenge-172
---

# --description--

Gegeben ist ein String, der nur aus Buchstaben und Zahlen besteht. Geben Sie einen neuen String zurück, in den jedes Mal ein Bindestrich (`-`) eingefügt wird, wenn der String von einem Buchstaben zu einer Zahl oder von einer Zahl zu einem Buchstaben wechselt.

# --hints--

`separateLettersAndNumbers("ABC123")` sollte `"ABC-123"` zurückgeben.

```js
assert.equal(separateLettersAndNumbers("ABC123"), "ABC-123");
```

`separateLettersAndNumbers("Route66")` sollte `"Route-66"` zurückgeben.

```js
assert.equal(separateLettersAndNumbers("Route66"), "Route-66");
```

`separateLettersAndNumbers("H3LL0W0RLD")` sollte `"H-3-LL-0-W-0-RLD"` zurückgeben.

```js
assert.equal(separateLettersAndNumbers("H3LL0W0RLD"), "H-3-LL-0-W-0-RLD");
```

`separateLettersAndNumbers("a1b2c3d4")` sollte `"a-1-b-2-c-3-d-4"` zurückgeben.

```js
assert.equal(separateLettersAndNumbers("a1b2c3d4"), "a-1-b-2-c-3-d-4");
```

# --seed--

## --seed-contents--

```js
function separateLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function separateLettersAndNumbers(str) {
  if (str.length === 0) return str;

  let result = str[0];

  for (let i = 1; i < str.length; i++) {
    const prev = str[i - 1];
    const curr = str[i];

    const prevIsLetter = /[a-zA-Z]/.test(prev);
    const currIsLetter = /[a-zA-Z]/.test(curr);

    if (prevIsLetter !== currIsLetter) {
      result += "-";
    }

    result += curr;
  }

  return result;
}
```
