---
id: 69738771fb5a7b8b24cca29d
title: "Aufgabe 171: Array abflachen"
challengeType: 28
dashedName: challenge-171
---

# --description--

Gegeben ist ein Array, das verschachtelte Arrays enthält. Geben Sie ein neues Array zurück, in dem alle Werte zu einem einzigen eindimensionalen Array abgeflacht sind. Behalten Sie die ursprüngliche Reihenfolge der Elemente in den Arrays bei.

# --hints--

`flatten([1, [2, 3], 4])` sollte `[1, 2, 3, 4]` zurückgeben.

```js
assert.deepEqual(flatten([1, [2, 3], 4]), [1, 2, 3, 4]);
```

`flatten([5, [4, [3, 2]], 1])` sollte `[5, 4, 3, 2, 1]` zurückgeben.

```js
assert.deepEqual(flatten([5, [4, [3, 2]], 1]), [5, 4, 3, 2, 1]);
```

`flatten(["A", [[[["B"]]]], "C"])` sollte `["A", "B", "C"]` zurückgeben.

```js
assert.deepEqual(flatten(["A", [[[["B"]]]], "C"]), ["A", "B", "C"]);
```

`flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]])` sollte `["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]` zurückgeben.

```js
assert.deepEqual(flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]]), ["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]);
```

`flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]])` sollte `["red","blue","green","yellow","purple","orange","pink","brown"]` zurückgeben.

```js
assert.deepEqual(flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]]), ["red","blue","green","yellow","purple","orange","pink","brown"]);
```

# --seed--

## --seed-contents--

```js
function flatten(arr) {

  return arr;
}
```

# --solutions--

```js
function flatten(arr) {
  let result = [];

  for (const item of arr) {
    if (Array.isArray(item)) {
      result = result.concat(flatten(item));
    } else {
      result.push(item);
    }
  }

  return result;
}
```
