---
id: 69738771fb5a7b8b24cca2a2
title: "Aufgabe 176: Murmeltiertag"
challengeType: 28
dashedName: challenge-176
---

# --description--

Heute ist Murmeltiertag, an dem ein Murmeltier das Wetter vorhersagt, indem es beobachtet, ob es seinen Schatten sieht oder nicht.

Geben Sie basierend auf einem Wert, der das Erscheinen des Murmeltiers darstellt, die korrekte Vorhersage zurück:

- Wenn der gegebene Wert der Boolean `true` ist (das Murmeltier hat seinen Schatten gesehen), geben Sie `"Looks like we'll have six more weeks of winter."` zurück.
- Wenn der Wert der Boolean `false` ist (das Murmeltier hat seinen Schatten nicht gesehen), geben Sie `"It's going to be an early spring."` zurück.
- Wenn der Wert etwas anderes ist (das Murmeltier ist nicht aufgetaucht), geben Sie `"No prediction this year."` zurück.

# --hints--

`groundhogDayPrediction(true)` sollte `"Looks like we'll have six more weeks of winter."` zurückgeben.

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

`groundhogDayPrediction(false)` sollte `"It's going to be an early spring."` zurückgeben.

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

`groundhogDayPrediction(null)` sollte `"No prediction this year."` zurückgeben.

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")` sollte `"No prediction this year."` zurückgeben.

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")` sollte `"No prediction this year."` zurückgeben.

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
