---
id: 69738771fb5a7b8b24cca29f
title: "Aufgabe 173: Gültige Bauernzüge"
challengeType: 28
dashedName: challenge-173
---

# --description--

Gegeben ist die Position eines Ihrer Bauern auf einem Schachbrett. Geben Sie ein Array aller gültigen Felder zurück, auf die er ziehen kann, sortiert in aufsteigender Reihenfolge.

Ein Standard-Schachbrett ist 8x8 groß, mit Spalten, die von `A` bis `H` beschriftet sind (von links nach rechts), und Reihen, die von `1` bis `8` beschriftet sind (von unten nach oben). Es sieht so aus:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Für diese Aufgabe gilt:

- Sie sind der Spieler unten am Brett.
- Bauern können nur ein Feld „nach oben“ ziehen.
- Außer der Bauer steht in der Startreihe (Reihe 2), dann kann er ein oder zwei Felder nach oben ziehen.

Zum Beispiel: Für `"D4"` geben Sie `["D5"]` zurück, das einzige Feld, auf das Ihr Bauer ziehen kann.
Für `"B2"` geben Sie `["B3", "B4"]` zurück, weil er in der Startreihe steht und die Ergebnisse in aufsteigender Reihenfolge sein müssen.

# --hints--

`findPawnMoves("D4")` sollte `["D5"]` zurückgeben.

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")` sollte `["B3", "B4"]` zurückgeben.

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")` sollte `["A8"]` zurückgeben.

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")` sollte `["G3", "G4"]` zurückgeben.

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")` sollte `["E4"]` zurückgeben.

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```
