---
id: 69738771fb5a7b8b24cca2a5
title: "Aufgabe 179: Kleingeld"
challengeType: 28
dashedName: challenge-179
---

# --description--

Gegeben ist ein Array ganzer Zahlen, die die Münzen in Ihrer Tasche darstellen, wobei jede ganze Zahl den Wert einer Münze in Cent angibt. Geben Sie den Gesamtbetrag im Format `"$D.CC"` zurück.

- 100 Cent entsprechen 1 Dollar.
- Im Rückgabewert soll für Beträge unter einem Dollar eine führende Null enthalten sein und immer genau zwei Stellen für die Cent angegeben werden.

# --hints--

`countChange([25, 10, 5, 1])` sollte `"$0.41"` zurückgeben.

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` sollte `"$1.43"` zurückgeben.

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])` sollte `"$37.55"` zurückgeben.

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` sollte `"$0.70"` zurückgeben.

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])` sollte `"$0.01"` zurückgeben.

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])` sollte `"$1.00"` zurückgeben.

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```
