---
id: 69738771fb5a7b8b24cca2a4
title: "Aufgabe 178: Text kürzen"
challengeType: 28
dashedName: challenge-178
---

# --description--

Gegeben ist ein String. Geben Sie ihn unverändert zurück, wenn er 20 Zeichen oder kürzer ist. Ist er länger als 20 Zeichen, kürzen Sie ihn auf die ersten 17 Zeichen und fügen `"..."` am Ende hinzu (so dass er insgesamt 20 Zeichen lang ist) und geben Sie das Ergebnis zurück.

# --hints--

`truncateText("Hello, world!")` sollte `"Hello, world!"` zurückgeben.

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` sollte `"This string shoul..."` zurückgeben.

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` sollte `"Exactly twenty chars"` zurückgeben.

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` sollte `"...................."` zurückgeben.

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
