---
id: 697a49e6ff50d756c9b69362
title: "Aufgabe 185: Winterspiele 2026 Tag 6: Eiskunstlauf"
challengeType: 28
dashedName: challenge-185
---

# --description--

Gegeben ist ein Array mit Wertungen von Juroren und optionalen Strafpunkten. Berechnen Sie die Endpunktzahl für eine Eiskunstlauf-Kür.

Das erste Argument ist ein Array mit 10 Wertungen von Juroren, jeweils eine Zahl von 0 bis 10. Entfernen Sie die höchste und die niedrigste Wertung und summieren Sie die verbleibenden 8 Wertungen, um die Grundpunktzahl zu erhalten.

Alle weiteren an die Funktion übergebenen Argumente sind Strafpunkte. Ziehen Sie alle Strafpunkte von der Grundpunktzahl ab, um die Endpunktzahl zu erhalten.

# --hints--

`computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1)` sollte `64` zurückgeben.

```js
assert.equal(computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1), 64);
```

`computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10])` sollte `80` zurückgeben.

```js
assert.equal(computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10]), 80);
```

`computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1)` sollte `67` zurückgeben.

```js
assert.equal(computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1), 67);
```

`computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0)` sollte `67.5` zurückgeben.

```js
assert.equal(computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0), 67.5);
```

`computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5)` sollte `59` zurückgeben.

```js
assert.equal(computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5), 59);
```

# --seed--

## --seed-contents--

```js
function computeScore(judgeScores, ...penalties) {

  return judgeScores;
}
```

# --solutions--

```js
function computeScore(judgeScores, ...penalties) {
  const sortedScores = [...judgeScores].sort((a, b) => a - b);
  const baseScore = sortedScores.slice(1, 9).reduce((sum, s) => sum + s, 0);
  const totalPenalty = penalties.reduce((sum, p) => sum + p, 0);

  return baseScore - totalPenalty;
}
```
