---
id: 697a49e6ff50d756c9b69360
title: "Aufgabe 183: Winterspiele 2026 Tag 4: Skispringen"
challengeType: 28
dashedName: challenge-183
---

# --description--

Gegeben sind Distanzpunkte, Stilpunkte, ein Windkompensationswert und ein K-Punkt-Bonuswert. Berechnen Sie Ihre Punktzahl für den Skisprung und bestimmen Sie, ob Sie eine Medaille gewonnen haben oder nicht.

- Ihre Punktzahl wird berechnet, indem Sie die vier oben genannten Werte summieren.

Die aktuellen Gesamtpunktzahlen der anderen Springer sind:

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- Wenn Ihre Punktzahl die beste ist, geben Sie `"Gold"` zurück
- Wenn sie die zweitbeste ist, geben Sie `"Silver"` zurück
- Wenn sie die drittbeste ist, geben Sie `"Bronze"` zurück
- Andernfalls geben Sie `"No Medal"` zurück

# --hints--

`skiJumpMedal(125.0, 58.0, 0.0, 6.0)` sollte `"Gold"` zurückgeben.

```js
assert.equal(skiJumpMedal(125.0, 58.0, 0.0, 6.0), "Gold");
```

`skiJumpMedal(119.0, 50.0, 1.0, 4.0)` sollte `"Bronze"` zurückgeben.

```js
assert.equal(skiJumpMedal(119.0, 50.0, 1.0, 4.0), "Bronze");
```

`skiJumpMedal(122.0, 52.0, -1.0, 4.0)` sollte `"Silver"` zurückgeben.

```js
assert.equal(skiJumpMedal(122.0, 52.0, -1.0, 4.0), "Silver");
```

`skiJumpMedal(118.0, 50.5, -1.5, 4.0)` sollte `"No Medal"` zurückgeben.

```js
assert.equal(skiJumpMedal(118.0, 50.5, -1.5, 4.0), "No Medal");
```

`skiJumpMedal(124.0, 50.5, 2.0, 5.0)` sollte `"Gold"` zurückgeben.

```js
assert.equal(skiJumpMedal(124.0, 50.5, 2.0, 5.0), "Gold");
```

`skiJumpMedal(119.0, 49.5, 0.0, 3.0)` sollte `"No Medal"` zurückgeben.

```js
assert.equal(skiJumpMedal(119.0, 49.5, 0.0, 3.0), "No Medal");
```

# --seed--

## --seed-contents--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {

  return distancePoints;
}
```

# --solutions--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {
  const myScore = distancePoints + stylePoints + windComp + kPointBonus;
  const otherScores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0];
  const allScores = [...otherScores, myScore];
  allScores.sort((a, b) => b - a);
  const rank = allScores.indexOf(myScore) + 1;

  if (rank === 1) return "Gold";
  if (rank === 2) return "Silver";
  if (rank === 3) return "Bronze";
  return "No Medal";
}
```
