---
id: 697a49e6ff50d756c9b6935f
title: "Aufgabe 182: Winterspiele 2026 Tag 3: Biathlon"
challengeType: 28
dashedName: challenge-182
---

# --description--

Gegeben ist ein Array aus ganzen Zahlen, wobei jeder Wert die Anzahl der in einer Runde eines Biathlons getroffenen Ziele darstellt. Geben Sie die gesamte Strafdistanz zurück, die der Athlet laufen muss.

- Jede Runde besteht aus 5 Zielen.
- Jedes verfehlte Ziel führt zu einer Strafschleife von 150 Metern.

# --hints--

`calculatePenaltyDistance([4, 4])` sollte `300` zurückgeben.

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])` sollte `0` zurückgeben.

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])` sollte `450` zurückgeben.

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])` sollte `150` zurückgeben.

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])` sollte `1500` zurückgeben.

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```
