---
id: 697a49e9860d24853adef67d
title: "Aufgabe 192: Winterspiele 2026 Tag 13: Nordische Kombination"
challengeType: 28
dashedName: challenge-192
---

# --description--

Gegeben ist ein Array mit Sprungwerten von Athletinnen und Athleten. Berechnen Sie daraus die Startverzögerungen für den Langlaufteil der Nordischen Kombination.

Die Athletin oder der Athlet mit dem höchsten Sprungwert startet zuerst (0 Sekunden Verzögerung). Alle anderen starten später, abhängig davon, wie weit ihr Sprungwert hinter dem besten Sprung zurückliegt.

Um die Verzögerung für jede Athletin und jeden Athleten zu berechnen, subtrahieren Sie den Sprungwert der Athletin oder des Athleten vom besten Sprungwert insgesamt und multiplizieren das Ergebnis mit 1,5. Runden Sie die Verzögerung auf die nächste ganze Zahl auf.

# --hints--

`calculateStartDelays([120, 110, 125])` sollte `[8, 23, 0]` zurückgeben.

```js
assert.deepEqual(calculateStartDelays([120, 110, 125]), [8, 23, 0]);
```

`calculateStartDelays([118, 125, 122, 120])` sollte `[11, 0, 5, 8]` zurückgeben.

```js
assert.deepEqual(calculateStartDelays([118, 125, 122, 120]), [11, 0, 5, 8]);
```

`calculateStartDelays([100, 105, 95, 110, 120, 115, 108])` sollte `[30, 23, 38, 15, 0, 8, 18]` zurückgeben.

```js
assert.deepEqual(calculateStartDelays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18]);
```

`calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` sollte `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]` zurückgeben.

```js
assert.deepEqual(calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12]);
```

# --seed--

## --seed-contents--

```js
function calculateStartDelays(jumpScores) {

  return jumpScores
}
```

# --solutions--

```js
function calculateStartDelays(jumpScores) {
  const bestJump = Math.max(...jumpScores);
  return jumpScores.map(score => Math.ceil((bestJump - score) * 1.5));
}
```
