---
id: 697a49e9860d24853adef67f
title: "Aufgabe 194: Winterspiele 2026 Tag 15: Freestyle-Skifahren"
challengeType: 28
dashedName: challenge-194
---

# --description--

Gegeben ist ein Trickname, der aus zwei Wörtern besteht. Bestimmen Sie, ob es sich um einen gültigen Freestyle-Ski-Tricknamen handelt.

Ein Trick ist gültig, wenn das erste Wort in der Liste der gültigen ersten Wörter steht und das zweite Wort in der Liste der gültigen zweiten Wörter.

- Die beiden Wörter sind durch ein einzelnes Leerzeichen getrennt.

Gültige erste Wörter:

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

Gültige zweite Wörter:

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`isValidTrick("Polar Vortex")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidTrick("Polar Vortex"));
```

`isValidTrick("Solar Icequake")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidTrick("Solar Icequake"));
```

`isValidTrick("Thunder Blizzard")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidTrick("Thunder Blizzard"));
```

`isValidTrick("Phantom Frostbite")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidTrick("Phantom Frostbite"));
```

`isValidTrick("Ghost Avalanche")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidTrick("Ghost Avalanche"));
```

`isValidTrick("Snowstorm Shadow")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidTrick("Snowstorm Shadow"));
```

`isValidTrick("Solar Sky")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidTrick("Solar Sky"));
```

# --seed--

## --seed-contents--

```js
function isValidTrick(trickName) {

  return trickName;
}
```

# --solutions--

```js
function isValidTrick(trickName) {
  const validFirst = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"];
  const validSecond = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"];

  const words = trickName.split(" ");
  const [first, second] = words;
  return validFirst.includes(first) && validSecond.includes(second);
}
```
