---
id: 697a49e9860d24853adef680
title: "Aufgabe 195: Winterspiele 2026 Tag 16: Curling"
challengeType: 28
dashedName: challenge-195
---

# --description--

Gegeben ist eine 5x5-Matrix, die das „Haus“ am Ende einer Curling-Runde darstellt. Bestimmen Sie, welches Team Punkte erzielt und wie viele Punkte es sind.

Das Layout:

- Die mittlere Zelle (Index `[2, 2]`) ist der „Button“.
- Die 8 Zellen, die den Button direkt umgeben, stellen Ring 1 dar.
- Die 16 Zellen am äußeren Rand des Hauses stellen Ring 2 dar.

In der gegebenen Matrix:

- `"."` steht für einen leeren Platz.
- `"R"` steht für einen Platz mit einem roten Stein.
- `"Y"` steht für einen Platz mit einem gelben Stein.

Punkteregeln:

- Pro Runde kann nur ein Team punkten.
- Das Team mit dem Stein, der dem Button am nächsten ist, erzielt Punkte.
- Das punktende Team erhält 1 Punkt für jeden seiner Steine, der näher am Button liegt als der dem Gegner am nächsten liegende Stein.
- Je niedriger die Ringnummer, desto näher ist der Stein am Zentrum.
- Wenn der dem Zentrum am nächsten liegende Stein beider Teams gleich weit entfernt ist, erzielt kein Team Punkte.

Geben Sie zurück:

- Einen String im Format `"team: number_of_points"`, z. B. `"R: 2"`.
- Oder `"No points awarded"`, falls kein Team Punkte erzielt hat.

Zum Beispiel, gegeben:

```js
[
  [".", ".", "R", ".", "."],
  [".", "R", ".", ".", "."],
  ["Y", ".", ".", ".", "."],
  [".", "R", ".", ".", "."],
  [".", ".", ".", ".", "."]
]
```

Geben Sie `"R: 2"` zurück. Die zwei roten Steine in Ring 1 sind gleich nah am Button und die einzigen zwei Steine, die näher sind als der dem Gelb am nächsten liegende Stein.

# --hints--

`scoreCurling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]])` sollte `"R: 2"` zurückgeben.

```js
assert.equal(scoreCurling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]]), "R: 2");
```

`scoreCurling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]])` sollte `"Y: 3"` zurückgeben.

```js
assert.equal(scoreCurling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]]), "Y: 3");
```

`scoreCurling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]])` sollte `"No points awarded"` zurückgeben.

```js
assert.equal(scoreCurling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]]), "No points awarded");
```

`scoreCurling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]])` sollte `"R: 4"` zurückgeben.

```js
assert.equal(scoreCurling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]]), "R: 4");
```

`scoreCurling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]])` sollte `"Y: 1"` zurückgeben.

```js
assert.equal(scoreCurling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]]), "Y: 1");
```

`scoreCurling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]])` sollte `"No points awarded"` zurückgeben.

```js
assert.equal(scoreCurling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]]), "No points awarded");
```

# --seed--

## --seed-contents--

```js
function scoreCurling(house) {

  return house;
}
```

# --solutions--

```js
function scoreCurling(house) {
  const stones = { R: [], Y: [] };

  for (let i = 0; i < 5; i++) {
    for (let j = 0; j < 5; j++) {
      const cell = house[i][j];
      if (cell === 'R' || cell === 'Y') {
        const distance = Math.max(Math.abs(i - 2), Math.abs(j - 2));
        stones[cell].push(distance);
      }
    }
  }

  const redClosest = stones.R.length > 0 ? Math.min(...stones.R) : Infinity;
  const yellowClosest = stones.Y.length > 0 ? Math.min(...stones.Y) : Infinity;

  if (redClosest === Infinity && yellowClosest === Infinity) {
    return "No points awarded";
  }

  if (redClosest === yellowClosest) {
    return "No points awarded";
  }

let winner, opponentClosest, winningStones;
  if (redClosest < yellowClosest) {
    winner = 'R';
    opponentClosest = yellowClosest;
    winningStones = stones.R;
  } else {
    winner = 'Y';
    opponentClosest = redClosest;
    winningStones = stones.Y;
  }

  const points = winningStones.filter(d => d < opponentClosest).length;

  return points === 0 ? "No points awarded" : `${winner}: ${points}`;
}
```
