---
id: 697a49e9860d24853adef67c
title: "Aufgabe 191: Winterspiele 2026 Tag 12: Bobfahren"
challengeType: 28
dashedName: challenge-191
---

# --description--

Gegeben ist ein Array, das die Gewichte der Athletinnen und Athleten eines Bobteams darstellt, sowie eine Zahl, die das Gewicht des Bobs angibt. Bestimmen Sie, ob das Team startberechtigt ist.

- Die Länge des Arrays bestimmt die Teamgröße: 1-, 2- oder 4-Personen-Teams.
- Alle angegebenen Gewichtsangaben sind in Kilogramm (kg).

Der Bob (Schlitten allein) muss ein Mindestgewicht von haben:

- 162 kg für ein 1-Personen-Team
- 170 kg für ein 2-Personen-Team
- 210 kg für ein 4-Personen-Team

Das Gesamtgewicht des Bobs (Athleten plus Schlitten) darf nicht überschreiten:

- 247 kg für ein 1-Personen-Team
- 390 kg für ein 2-Personen-Team
- 630 kg für ein 4-Personen-Team

Geben Sie `"Eligible"` zurück, wenn das Team alle Anforderungen erfüllt, oder `"Not Eligible"`, wenn das Team eine oder mehrere Anforderungen nicht erfüllt.

# --hints--

`checkEligibility([78], 165)` sollte `"Eligible"` zurückgeben.

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)` sollte `"Not Eligible"` zurückgeben.

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)` sollte `"Not Eligible"` zurückgeben.

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)` sollte `"Eligible"` zurückgeben.

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)` sollte `"Not Eligible"` zurückgeben.

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)` sollte `"Not Eligible"` zurückgeben.

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)` sollte `"Eligible"` zurückgeben.

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)` sollte `"Not Eligible"` zurückgeben.

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)` sollte `"Not Eligible"` zurückgeben.

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
