---
id: 697a49e6ff50d756c9b69366
title: "Aufgabe 189: Winterspiele 2026 Tag 10: Alpiner Skisport"
challengeType: 28
dashedName: challenge-189
---

# --description--

Gegeben sind der Höhenunterschied, die horizontale Distanz und der Typ eines Skihangs. Bestimmen Sie die Schwierigkeitsbewertung des Hangs.

Zur Bestimmung der Bewertung:

- Berechnen Sie die Steilheit des Hangs, indem Sie den Höhenunterschied durch die Distanz teilen.
- Berechnen Sie dann die angepasste Steilheit basierend auf dem Hangtyp:
  - `"Downhill"`: multiplizieren Sie die Steilheit mit 1,2
  - `"Slalom"`: multiplizieren Sie die Steilheit mit 0,9
  - `"Giant Slalom"`: multiplizieren Sie die Steilheit mit 1,0

Geben Sie zurück:

- `"Green"`, wenn die angepasste Steilheit kleiner oder gleich 0,1 ist
- `"Blue"`, wenn die angepasste Steilheit größer als 0,1 und kleiner oder gleich 0,25 ist
- `"Black"`, wenn die angepasste Steilheit größer als 0,25 ist

# --hints--

`getHillRating(95, 900, "Slalom")` sollte `"Green"` zurückgeben.

```js
assert.equal(getHillRating(95, 900, "Slalom"), "Green");
```

`getHillRating(620, 2800, "Downhill")` sollte `"Black"` zurückgeben.

```js
assert.equal(getHillRating(620, 2800, "Downhill"), "Black");
```

`getHillRating(420, 1680, "Giant Slalom")` sollte `"Blue"` zurückgeben.

```js
assert.equal(getHillRating(420, 1680, "Giant Slalom"), "Blue");
```

`getHillRating(250, 3000, "Downhill")` sollte `"Green"` zurückgeben.

```js
assert.equal(getHillRating(250, 3000, "Downhill"), "Green");
```

`getHillRating(110, 900, "Slalom")` sollte `"Blue"` zurückgeben.

```js
assert.equal(getHillRating(110, 900, "Slalom"), "Blue");
```

`getHillRating(380, 1500, "Giant Slalom")` sollte `"Black"` zurückgeben.

```js
assert.equal(getHillRating(380, 1500, "Giant Slalom"), "Black");
```

# --seed--

## --seed-contents--

```js
function getHillRating(drop, distance, type) {

  return drop;
}
```

# --solutions--

```js
function getHillRating(drop, distance, type) {
  let steepness = drop / distance;

  if (type === "Downhill") {
    steepness *= 1.2;
  } else if (type === "Slalom") {
    steepness *= 0.9;
  } else if (type === "Giant Slalom") {
    steepness *= 1.0;
  }

  if (steepness <= 0.1) {
    return "Green";
  } else if (steepness <= 0.25) {
    return "Blue";
  } else {
    return "Black";
  }
}
```
