---
id: 697a49e6ff50d756c9b69365
title: "Aufgabe 188: Winterspiele 2026 Tag 9: Skeleton"
challengeType: 28
dashedName: challenge-188
---

# --description--

Gegeben ist ein String, der die Kurven auf einer Skeleton-Bahn darstellt. Bestimmen Sie die Schwierigkeit der Bahn.

- Der gegebene String besteht nur aus den Buchstaben:
  - `"L"` für eine Linkskurve
  - `"R"` für eine Rechtskurve
  - `"S"` für ein gerades Teilstück

- Jede Richtungsänderung fügt 15 Punkte hinzu (ein `"L"` gefolgt von einem `"R"` oder umgekehrt).
- Alle anderen Kurven fügen jeweils 5 Punkte hinzu (alle anderen `"L"` oder `"R"` Zeichen).
- Gerade Teilstücke fügen 0 Punkte hinzu.

Die Schwierigkeit der Bahn basiert auf der Gesamtpunktzahl. Geben Sie zurück:

- `"Easy"`, wenn die Gesamtpunktzahl 0 bis 100 beträgt
- `"Medium"`, wenn die Gesamtpunktzahl 101 bis 200 beträgt
- `"Hard"`, wenn die Gesamtpunktzahl über 200 liegt

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")` sollte `"Easy"` zurückgeben.

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` sollte `"Hard"` zurückgeben.

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")` sollte `"Medium"` zurückgeben.

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` sollte `"Hard"` zurückgeben.

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")` sollte `"Medium"` zurückgeben.

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` sollte `"Easy"` zurückgeben.

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
