---
id: 697a49e6ff50d756c9b6935e
title: "Aufgabe 181: Winterspiele 2026 Tag 2: Snowboarden"
challengeType: 28
dashedName: challenge-181
---

# --description--

Gegeben sind die Ausgangsposition eines Snowboarders und eine Drehung in Grad. Bestimmen Sie die Landeposition.

- Die Position eines Snowboarders ist entweder `"Regular"` oder `"Goofy"`.
- Drehungen bei Tricks sind Vielfache von 90 Grad. Positive Werte bedeuten eine Drehung im Uhrzeigersinn, negative eine gegen den Uhrzeigersinn.
- Die Landeposition wechselt alle 180 Grad Drehung.

Zum Beispiel: Bei `"Regular"` und `90` geben Sie `"Regular"` zurück. Bei `"Regular"` und `180` Grad geben Sie `"Goofy"` zurück.

# --hints--

`getLandingStance("Regular", 90)` sollte `"Regular"` zurückgeben.

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)` sollte `"Goofy"` zurückgeben.

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)` sollte `"Regular"` zurückgeben.

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)` sollte `"Goofy"` zurückgeben.

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)` sollte `"Goofy"` zurückgeben.

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)` sollte `"Regular"` zurückgeben.

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```
