---
id: 697a49e6ff50d756c9b6935d
title: "Aufgabe 180: Winterspiele Tag 1: Eröffnungstag"
challengeType: 28
dashedName: challenge-180
---

# --description--

Die nächsten 17 Tage erwarten Sie Programmieraufgaben, die von den Winterspielen inspiriert sind.

Für die erste Aufgabe erhalten Sie einen zweibuchstabigen Ländercode und sollen das Flaggen-Emoji für dieses Land zurückgeben.

Verwenden Sie diese Liste:

| Land                       | Code       | Flagge   |
| -                         | -:         | :--:     |
| Albanien                  |            `"AL"` | `"🇦🇱"` |
| Andorra                   |            `"AD"` | `"🇦🇩"` |
| Argentinien               |            `"AR"` | `"🇦🇷"` |
| Armenien                  |            `"AM"` | `"🇦🇲"` |
| Australien                |            `"AU"` | `"🇦🇺"` |
| Österreich                |            `"AT"` | `"🇦🇹"` |
| Aserbaidschan             |            `"AZ"` | `"🇦🇿"` |
| Belgien                   |            `"BE"` | `"🇧🇪"` |
| Benin                     |            `"BJ"` | `"🇧🇯"` |
| Bolivien                  |            `"BO"` | `"🇧🇴"` |
| Bosnien und Herzegowina   |            `"BA"` | `"🇧🇦"` |
| Brasilien                 |            `"BR"` | `"🇧🇷"` |
| Bulgarien                 |            `"BG"` | `"🇧🇬"` |
| Kanada                    |            `"CA"` | `"🇨🇦"` |
| Chile                     |            `"CL"` | `"🇨🇱"` |
| China                     |            `"CN"` | `"🇨🇳"` |
| Kolumbien                 |            `"CO"` | `"🇨🇴"` |
| Kroatien                  |            `"HR"` | `"🇭🇷"` |
| Zypern                    |            `"CY"` | `"🇨🇾"` |
| Tschechien                |            `"CZ"` | `"🇨🇿"` |
| Dänemark                  |            `"DK"` | `"🇩🇰"` |
| Ecuador                   |            `"EC"` | `"🇪🇨"` |
| Eritrea                   |            `"ER"` | `"🇪🇷"` |
| Estland                   |            `"EE"` | `"🇪🇪"` |
| Finnland                  |            `"FI"` | `"🇫🇮"` |
| Frankreich                |            `"FR"` | `"🇫🇷"` |
| Georgien                  |            `"GE"` | `"🇬🇪"` |
| Deutschland               |            `"DE"` | `"🇩🇪"` |
| Großbritannien            |            `"GB"` | `"🇬🇧"` |
| Griechenland              |            `"GR"` | `"🇬🇷"` |
| Guinea-Bissau             |            `"GW"` | `"🇬🇼"` |
| Haiti                     |            `"HT"` | `"🇭🇹"` |
| Hongkong                  |            `"HK"` | `"🇭🇰"` |
| Ungarn                    |            `"HU"` | `"🇭🇺"` |
| Island                    |            `"IS"` | `"🇮🇸"` |
| Indien                    |            `"IN"` | `"🇮🇳"` |
| Iran                      |            `"IR"` | `"🇮🇷"` |
| Irland                    |            `"IE"` | `"🇮🇪"` |
| Israel                    |            `"IL"` | `"🇮🇱"` |
| Italien                   |            `"IT"` | `"🇮🇹"` |
| Jamaika                   |            `"JM"` | `"🇯🇲"` |
| Japan                     |            `"JP"` | `"🇯🇵"` |
| Kasachstan                |            `"KZ"` | `"🇰🇿"` |
| Kenia                     |            `"KE"` | `"🇰🇪"` |
| Kosovo                    |            `"XK"` | `"🇽🇰"` |
| Kirgisistan               |            `"KG"` | `"🇰🇬"` |
| Lettland                  |            `"LV"` | `"🇱🇻"` |
| Libanon                   |            `"LB"` | `"🇱🇧"` |
| Liechtenstein             |            `"LI"` | `"🇱🇮"` |
| Litauen                   |            `"LT"` | `"🇱🇹"` |
| Luxemburg                 |            `"LU"` | `"🇱🇺"` |
| Madagaskar                |            `"MG"` | `"🇲🇬"` |
| Malaysia                  |            `"MY"` | `"🇲🇾"` |
| Malta                     |            `"MT"` | `"🇲🇹"` |
| Mexiko                    |            `"MX"` | `"🇲🇽"` |
| Moldawien                 |            `"MD"` | `"🇲🇩"` |
| Monaco                    |            `"MC"` | `"🇲🇨"` |
| Mongolei                  |            `"MN"` | `"🇲🇳"` |
| Montenegro                |            `"ME"` | `"🇲🇪"` |
| Marokko                   |            `"MA"` | `"🇲🇦"` |
| Niederlande               |            `"NL"` | `"🇳🇱"` |
| Neuseeland                |            `"NZ"` | `"🇳🇿"` |
| Nigeria                   |            `"NG"` | `"🇳🇬"` |
| Nordmazedonien            |            `"MK"` | `"🇲🇰"` |
| Norwegen                  |            `"NO"` | `"🇳🇴"` |
| Pakistan                  |            `"PK"` | `"🇵🇰"` |
| Philippinen               |            `"PH"` | `"🇵🇭"` |
| Polen                     |            `"PL"` | `"🇵🇱"` |
| Portugal                  |            `"PT"` | `"🇵🇹"` |
| Puerto Rico               |            `"PR"` | `"🇵🇷"` |
| Rumänien                  |            `"RO"` | `"🇷🇴"` |
| San Marino                |            `"SM"` | `"🇸🇲"` |
| Saudi-Arabien             |            `"SA"` | `"🇸🇦"` |
| Serbien                   |            `"RS"` | `"🇷🇸"` |
| Singapur                  |            `"SG"` | `"🇸🇬"` |
| Slowakei                  |            `"SK"` | `"🇸🇰"` |
| Slowenien                 |            `"SI"` | `"🇸🇮"` |
| Südafrika                 |            `"ZA"` | `"🇿🇦"` |
| Südkorea                  |            `"KR"` | `"🇰🇷"` |
| Spanien                   |            `"ES"` | `"🇪🇸"` |
| Schweden                  |            `"SE"` | `"🇸🇪"` |
| Schweiz                   |            `"CH"` | `"🇨🇭"` |
| Thailand                  |            `"TH"` | `"🇹🇭"` |
| Trinidad und Tobago       |            `"TT"` | `"🇹🇹"` |
| Türkei                    |            `"TR"` | `"🇹🇷"` |
| Ukraine                   |            `"UA"` | `"🇺🇦"` |
| Vereinigte Arabische Emirate |        `"AE"` | `"🇦🇪"` |
| Vereinigte Staaten        |            `"US"` | `"🇺🇸"` |
| Uruguay                   |            `"UY"` | `"🇺🇾"` |
| Usbekistan                |            `"UZ"` | `"🇺🇿"` |
| Venezuela                 |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")` sollte `"🇦🇱"` zurückgeben.

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")` sollte `"🇦🇩"` zurückgeben.

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")` sollte `"🇦🇷"` zurückgeben.

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")` sollte `"🇦🇲"` zurückgeben.

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")` sollte `"🇦🇺"` zurückgeben.

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")` sollte `"🇦🇹"` zurückgeben.

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")` sollte `"🇦🇿"` zurückgeben.

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")` sollte `"🇧🇪"` zurückgeben.

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")` sollte `"🇧🇯"` zurückgeben.

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")` sollte `"🇧🇴"` zurückgeben.

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")` sollte `"🇧🇦"` zurückgeben.

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")` sollte `"🇧🇷"` zurückgeben.

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")` sollte `"🇧🇬"` zurückgeben.

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")` sollte `"🇨🇦"` zurückgeben.

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")` sollte `"🇨🇱"` zurückgeben.

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")` sollte `"🇨🇳"` zurückgeben.

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")` sollte `"🇨🇴"` zurückgeben.

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")` sollte `"🇭🇷"` zurückgeben.

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")` sollte `"🇨🇾"` zurückgeben.

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")` sollte `"🇨🇿"` zurückgeben.

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")` sollte `"🇩🇰"` zurückgeben.

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")` sollte `"🇪🇨"` zurückgeben.

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")` sollte `"🇪🇷"` zurückgeben.

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")` sollte `"🇪🇪"` zurückgeben.

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")` sollte `"🇫🇮"` zurückgeben.

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")` sollte `"🇫🇷"` zurückgeben.

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")` sollte `"🇬🇪"` zurückgeben.

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")` sollte `"🇩🇪"` zurückgeben.

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")` sollte `"🇬🇧"` zurückgeben.

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")` sollte `"🇬🇷"` zurückgeben.

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")` sollte `"🇬🇼"` zurückgeben.

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")` sollte `"🇭🇹"` zurückgeben.

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")` sollte `"🇭🇰"` zurückgeben.

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")` sollte `"🇭🇺"` zurückgeben.

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")` sollte `"🇮🇸"` zurückgeben.

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")` sollte `"🇮🇳"` zurückgeben.

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")` sollte `"🇮🇷"` zurückgeben.

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")` sollte `"🇮🇪"` zurückgeben.

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")` sollte `"🇮🇱"` zurückgeben.

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")` sollte `"🇮🇹"` zurückgeben.

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")` sollte `"🇯🇲"` zurückgeben.

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")` sollte `"🇯🇵"` zurückgeben.

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")` sollte `"🇰🇿"` zurückgeben.

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")` sollte `"🇰🇪"` zurückgeben.

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")` sollte `"🇽🇰"` zurückgeben.

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")` sollte `"🇰🇬"` zurückgeben.

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")` sollte `"🇱🇻"` zurückgeben.

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")` sollte `"🇱🇧"` zurückgeben.

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")` sollte `"🇱🇮"` zurückgeben.

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")` sollte `"🇱🇹"` zurückgeben.

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")` sollte `"🇱🇺"` zurückgeben.

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")` sollte `"🇲🇬"` zurückgeben.

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")` sollte `"🇲🇾"` zurückgeben.

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")` sollte `"🇲🇹"` zurückgeben.

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")` sollte `"🇲🇽"` zurückgeben.

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")` sollte `"🇲🇩"` zurückgeben.

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")` sollte `"🇲🇨"` zurückgeben.

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")` sollte `"🇲🇳"` zurückgeben.

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")` sollte `"🇲🇪"` zurückgeben.

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")` sollte `"🇲🇦"` zurückgeben.

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")` sollte `"🇳🇱"` zurückgeben.

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")` sollte `"🇳🇿"` zurückgeben.

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")` sollte `"🇳🇬"` zurückgeben.

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")` sollte `"🇲🇰"` zurückgeben.

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")` sollte `"🇳🇴"` zurückgeben.

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")` sollte `"🇵🇰"` zurückgeben.

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")` sollte `"🇵🇭"` zurückgeben.

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")` sollte `"🇵🇱"` zurückgeben.

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")` sollte `"🇵🇹"` zurückgeben.

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")` sollte `"🇵🇷"` zurückgeben.

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")` sollte `"🇷🇴"` zurückgeben.

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")` sollte `"🇸🇲"` zurückgeben.

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")` sollte `"🇸🇦"` zurückgeben.

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")` sollte `"🇷🇸"` zurückgeben.

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")` sollte `"🇸🇬"` zurückgeben.

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")` sollte `"🇸🇰"` zurückgeben.

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")` sollte `"🇸🇮"` zurückgeben.

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")` sollte `"🇿🇦"` zurückgeben.

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")` sollte `"🇰🇷"` zurückgeben.

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")` sollte `"🇪🇸"` zurückgeben.

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")` sollte `"🇸🇪"` zurückgeben.

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")` sollte `"🇨🇭"` zurückgeben.

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")` sollte `"🇹🇭"` zurückgeben.

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")` sollte `"🇹🇹"` zurückgeben.

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")` sollte `"🇹🇷"` zurückgeben.

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")` sollte `"🇺🇦"` zurückgeben.

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")` sollte `"🇦🇪"` zurückgeben.

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")` sollte `"🇺🇸"` zurückgeben.

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")` sollte `"🇺🇾"` zurückgeben.

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")` sollte `"🇺🇿"` zurückgeben.

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")` sollte `"🇻🇪"` zurückgeben.

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
