---
id: 697a49e9860d24853adef67b
title: "Aufgabe 190: Winterspiele 2026 Tag 11: Eishockey"
challengeType: 28
dashedName: challenge-190
---

# --description--

Gegeben ist ein Array mit 6 Eishockeyteams und deren Ergebnissen nach den Rundenspielen. Bestimmen Sie die Paarungen für die Halbfinalrunde.

- Jedes Array-Element enthält ein Team und dessen Ergebnis im Format `"TEAM: W-OTW-OTL-L"`. Dabei gilt:
  - `"W"` ist die Anzahl der regulären Siege, die jeweils 3 Punkte wert sind
  - `"OTW"` ist die Anzahl der Siege in der Verlängerung, die jeweils 2 Punkte wert sind
  - `"OTL"` ist die Anzahl der Niederlagen in der Verlängerung, die jeweils 1 Punkt wert sind
  - `"L"` ist die Anzahl der Niederlagen, die jeweils 0 Punkte wert sind
  
Zum Beispiel hätte `"FIN: 2-2-1-0"` nach Addition der Punkte 11 Punkte.
  
Berechnen Sie die Gesamtpunktzahl für jedes Team und geben Sie `"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."` zurück. Zum Beispiel `"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])` sollte `"The semi-final games will be FIN vs SWE and CAN vs USA."` zurückgeben.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.");
```

`getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])` sollte `"The semi-final games will be USA vs CZE and CAN vs FIN."` zurückgeben.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.");
```

`getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])` sollte `"The semi-final games will be CAN vs ITA and USA vs CZE."` zurückgeben.

```js
assert.equal(getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.");
```

`getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])` sollte `"The semi-final games will be JPN vs ITA and AUT vs KOR."` zurückgeben.

```js
assert.equal(getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.");
```

# --seed--

## --seed-contents--

```js
function getSemifinalMatchups(teams) {

  return teams;
}
```

# --solutions--

```js
function getSemifinalMatchups(teams) {
    const teamPoints = teams.map(teamStr => {
        const [name, record] = teamStr.split(": ");
        const [W, OTW, OTL, L] = record.split("-").map(Number);
        const points = W * 3 + OTW * 2 + OTL * 1 + L * 0;
        return { name, points };
    });

    teamPoints.sort((a, b) => b.points - a.points);
    const [first, second, third, fourth] = teamPoints;

    return `The semi-final games will be ${first.name} vs ${fourth.name} and ${second.name} vs ${third.name}.`;
}
```
