---
id: 697a49e6ff50d756c9b6935d
title: "Aufgabe 180: Winterspiele 2026 Tag 1: Erﾃｶffnungstag"
challengeType: 28
dashedName: challenge-180
---

# --description--

Heute beginnt die Winterspiele 2026. Die nﾃ､chsten 17 Tage erwarten Sie Programmieraufgaben, die von ihnen inspiriert sind.

Fﾃｼr die erste Aufgabe erhalten Sie einen zweibuchstabigen Lﾃ､ndercode und sollen das Flaggen-Emoji fﾃｼr dieses Land zurﾃｼckgeben.

Verwenden Sie diese Liste:

| Land                       | Code       | Flagge   |
| -                         | -:         | :--:     |
| Albanien                  |            `"AL"` | `"宖仡宖佬"` |
| Andorra                   |            `"AD"` | `"宖仡宖伃"` |
| Argentinien               |            `"AR"` | `"宖仡宖俉"` |
| Armenien                  |            `"AM"` | `"宖仡宖佾"` |
| Australien                |            `"AU"` | `"宖仡宖倂"` |
| ﾃ穆terreich                |            `"AT"` | `"宖仡宖倁"` |
| Aserbaidschan             |            `"AZ"` | `"宖仡宖偀"` |
| Belgien                   |            `"BE"` | `"宖份宖伋"` |
| Benin                     |            `"BJ"` | `"宖份宖佟"` |
| Bolivien                  |            `"BO"` | `"宖份宖侔"` |
| Bosnien und Herzegowina   |            `"BA"` | `"宖份宖仡"` |
| Brasilien                 |            `"BR"` | `"宖份宖俉"` |
| Bulgarien                 |            `"BG"` | `"宖份宖佈"` |
| Kanada                    |            `"CA"` | `"宖仿宖仡"` |
| Chile                     |            `"CL"` | `"宖仿宖佬"` |
| China                     |            `"CN"` | `"宖仿宖侊"` |
| Kolumbien                 |            `"CO"` | `"宖仿宖侔"` |
| Kroatien                  |            `"HR"` | `"宖佉宖俉"` |
| Zypern                    |            `"CY"` | `"宖仿宖倮"` |
| Tschechien                |            `"CZ"` | `"宖仿宖偀"` |
| Dﾃ､nemark                  |            `"DK"` | `"宖伃宖佪"` |
| Ecuador                   |            `"EC"` | `"宖伋宖仿"` |
| Eritrea                   |            `"ER"` | `"宖伋宖俉"` |
| Estland                   |            `"EE"` | `"宖伋宖伋"` |
| Finnland                  |            `"FI"` | `"宖你宖佖"` |
| Frankreich                |            `"FR"` | `"宖你宖俉"` |
| Georgien                  |            `"GE"` | `"宖佈宖伋"` |
| Deutschland               |            `"DE"` | `"宖伃宖伋"` |
| Groﾃ歟ritannien            |            `"GB"` | `"宖佈宖份"` |
| Griechenland              |            `"GR"` | `"宖佈宖俉"` |
| Guinea-Bissau             |            `"GW"` | `"宖佈宖倘"` |
| Haiti                     |            `"HT"` | `"宖佉宖倁"` |
| Hongkong                  |            `"HK"` | `"宖佉宖佪"` |
| Ungarn                    |            `"HU"` | `"宖佉宖倂"` |
| Island                    |            `"IS"` | `"宖佖宖俠"` |
| Indien                    |            `"IN"` | `"宖佖宖侊"` |
| Iran                      |            `"IR"` | `"宖佖宖俉"` |
| Irland                    |            `"IE"` | `"宖佖宖伋"` |
| Israel                    |            `"IL"` | `"宖佖宖佬"` |
| Italien                   |            `"IT"` | `"宖佖宖倁"` |
| Jamaika                   |            `"JM"` | `"宖佟宖佾"` |
| Japan                     |            `"JP"` | `"宖佟宖侗"` |
| Kasachstan                |            `"KZ"` | `"宖佪宖偀"` |
| Kenia                     |            `"KE"` | `"宖佪宖伋"` |
| Kosovo                    |            `"XK"` | `"宖倧宖佪"` |
| Kirgisistan               |            `"KG"` | `"宖佪宖佈"` |
| Lettland                  |            `"LV"` | `"宖佬宖倎"` |
| Libanon                   |            `"LB"` | `"宖佬宖份"` |
| Liechtenstein             |            `"LI"` | `"宖佬宖佖"` |
| Litauen                   |            `"LT"` | `"宖佬宖倁"` |
| Luxemburg                 |            `"LU"` | `"宖佬宖倂"` |
| Madagaskar                |            `"MG"` | `"宖佾宖佈"` |
| Malaysia                  |            `"MY"` | `"宖佾宖倮"` |
| Malta                     |            `"MT"` | `"宖佾宖倁"` |
| Mexiko                    |            `"MX"` | `"宖佾宖倧"` |
| Moldawien                 |            `"MD"` | `"宖佾宖伃"` |
| Monaco                    |            `"MC"` | `"宖佾宖仿"` |
| Mongolei                  |            `"MN"` | `"宖佾宖侊"` |
| Montenegro                |            `"ME"` | `"宖佾宖伋"` |
| Marokko                   |            `"MA"` | `"宖佾宖仡"` |
| Niederlande               |            `"NL"` | `"宖侊宖佬"` |
| Neuseeland                |            `"NZ"` | `"宖侊宖偀"` |
| Nigeria                   |            `"NG"` | `"宖侊宖佈"` |
| Nordmazedonien            |            `"MK"` | `"宖佾宖佪"` |
| Norwegen                  |            `"NO"` | `"宖侊宖侔"` |
| Pakistan                  |            `"PK"` | `"宖侗宖佪"` |
| Philippinen               |            `"PH"` | `"宖侗宖佉"` |
| Polen                     |            `"PL"` | `"宖侗宖佬"` |
| Portugal                  |            `"PT"` | `"宖侗宖倁"` |
| Puerto Rico               |            `"PR"` | `"宖侗宖俉"` |
| Rumﾃ､nien                  |            `"RO"` | `"宖俉宖侔"` |
| San Marino                |            `"SM"` | `"宖俠宖佾"` |
| Saudi-Arabien             |            `"SA"` | `"宖俠宖仡"` |
| Serbien                   |            `"RS"` | `"宖俉宖俠"` |
| Singapur                  |            `"SG"` | `"宖俠宖佈"` |
| Slowakei                  |            `"SK"` | `"宖俠宖佪"` |
| Slowenien                 |            `"SI"` | `"宖俠宖佖"` |
| Sﾃｼdafrika                 |            `"ZA"` | `"宖偀宖仡"` |
| Sﾃｼdkorea                  |            `"KR"` | `"宖佪宖俉"` |
| Spanien                   |            `"ES"` | `"宖伋宖俠"` |
| Schweden                  |            `"SE"` | `"宖俠宖伋"` |
| Schweiz                   |            `"CH"` | `"宖仿宖佉"` |
| Thailand                  |            `"TH"` | `"宖倁宖佉"` |
| Trinidad und Tobago       |            `"TT"` | `"宖倁宖倁"` |
| Tﾃｼrkei                    |            `"TR"` | `"宖倁宖俉"` |
| Ukraine                   |            `"UA"` | `"宖倂宖仡"` |
| Vereinigte Arabische Emirate |        `"AE"` | `"宖仡宖伋"` |
| Vereinigte Staaten        |            `"US"` | `"宖倂宖俠"` |
| Uruguay                   |            `"UY"` | `"宖倂宖倮"` |
| Usbekistan                |            `"UZ"` | `"宖倂宖偀"` |
| Venezuela                 |            `"VE"` | `"宖倎宖伋"` |

# --hints--

`getFlag("AL")` sollte `"宖仡宖佬"` zurﾃｼckgeben.

```js
assert.equal(getFlag("AL"), "宖仡宖佬");
```

`getFlag("AD")` sollte `"宖仡宖伃"` zurﾃｼckgeben.

```js
assert.equal(getFlag("AD"), "宖仡宖伃");
```

`getFlag("AR")` sollte `"宖仡宖俉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("AR"), "宖仡宖俉")
```

`getFlag("AM")` sollte `"宖仡宖佾"` zurﾃｼckgeben.

```js
assert.equal(getFlag("AM"), "宖仡宖佾")
```

`getFlag("AU")` sollte `"宖仡宖倂"` zurﾃｼckgeben.

```js
assert.equal(getFlag("AU"), "宖仡宖倂")
```

`getFlag("AT")` sollte `"宖仡宖倁"` zurﾃｼckgeben.

```js
assert.equal(getFlag("AT"), "宖仡宖倁")
```

`getFlag("AZ")` sollte `"宖仡宖偀"` zurﾃｼckgeben.

```js
assert.equal(getFlag("AZ"), "宖仡宖偀")
```

`getFlag("BE")` sollte `"宖份宖伋"` zurﾃｼckgeben.

```js
assert.equal(getFlag("BE"), "宖份宖伋")
```

`getFlag("BJ")` sollte `"宖份宖佟"` zurﾃｼckgeben.

```js
assert.equal(getFlag("BJ"), "宖份宖佟")
```

`getFlag("BO")` sollte `"宖份宖侔"` zurﾃｼckgeben.

```js
assert.equal(getFlag("BO"), "宖份宖侔")
```

`getFlag("BA")` sollte `"宖份宖仡"` zurﾃｼckgeben.

```js
assert.equal(getFlag("BA"), "宖份宖仡")
```

`getFlag("BR")` sollte `"宖份宖俉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("BR"), "宖份宖俉")
```

`getFlag("BG")` sollte `"宖份宖佈"` zurﾃｼckgeben.

```js
assert.equal(getFlag("BG"), "宖份宖佈")
```

`getFlag("CA")` sollte `"宖仿宖仡"` zurﾃｼckgeben.

```js
assert.equal(getFlag("CA"), "宖仿宖仡")
```

`getFlag("CL")` sollte `"宖仿宖佬"` zurﾃｼckgeben.

```js
assert.equal(getFlag("CL"), "宖仿宖佬")
```

`getFlag("CN")` sollte `"宖仿宖侊"` zurﾃｼckgeben.

```js
assert.equal(getFlag("CN"), "宖仿宖侊")
```

`getFlag("CO")` sollte `"宖仿宖侔"` zurﾃｼckgeben.

```js
assert.equal(getFlag("CO"), "宖仿宖侔")
```

`getFlag("HR")` sollte `"宖佉宖俉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("HR"), "宖佉宖俉")
```

`getFlag("CY")` sollte `"宖仿宖倮"` zurﾃｼckgeben.

```js
assert.equal(getFlag("CY"), "宖仿宖倮")
```

`getFlag("CZ")` sollte `"宖仿宖偀"` zurﾃｼckgeben.

```js
assert.equal(getFlag("CZ"), "宖仿宖偀")
```

`getFlag("DK")` sollte `"宖伃宖佪"` zurﾃｼckgeben.

```js
assert.equal(getFlag("DK"), "宖伃宖佪")
```

`getFlag("EC")` sollte `"宖伋宖仿"` zurﾃｼckgeben.

```js
assert.equal(getFlag("EC"), "宖伋宖仿")
```

`getFlag("ER")` sollte `"宖伋宖俉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("ER"), "宖伋宖俉")
```

`getFlag("EE")` sollte `"宖伋宖伋"` zurﾃｼckgeben.

```js
assert.equal(getFlag("EE"), "宖伋宖伋")
```

`getFlag("FI")` sollte `"宖你宖佖"` zurﾃｼckgeben.

```js
assert.equal(getFlag("FI"), "宖你宖佖")
```

`getFlag("FR")` sollte `"宖你宖俉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("FR"), "宖你宖俉")
```

`getFlag("GE")` sollte `"宖佈宖伋"` zurﾃｼckgeben.

```js
assert.equal(getFlag("GE"), "宖佈宖伋")
```

`getFlag("DE")` sollte `"宖伃宖伋"` zurﾃｼckgeben.

```js
assert.equal(getFlag("DE"), "宖伃宖伋")
```

`getFlag("GB")` sollte `"宖佈宖份"` zurﾃｼckgeben.

```js
assert.equal(getFlag("GB"), "宖佈宖份")
```

`getFlag("GR")` sollte `"宖佈宖俉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("GR"), "宖佈宖俉")
```

`getFlag("GW")` sollte `"宖佈宖倘"` zurﾃｼckgeben.

```js
assert.equal(getFlag("GW"), "宖佈宖倘")
```

`getFlag("HT")` sollte `"宖佉宖倁"` zurﾃｼckgeben.

```js
assert.equal(getFlag("HT"), "宖佉宖倁")
```

`getFlag("HK")` sollte `"宖佉宖佪"` zurﾃｼckgeben.

```js
assert.equal(getFlag("HK"), "宖佉宖佪")
```

`getFlag("HU")` sollte `"宖佉宖倂"` zurﾃｼckgeben.

```js
assert.equal(getFlag("HU"), "宖佉宖倂")
```

`getFlag("IS")` sollte `"宖佖宖俠"` zurﾃｼckgeben.

```js
assert.equal(getFlag("IS"), "宖佖宖俠")
```

`getFlag("IN")` sollte `"宖佖宖侊"` zurﾃｼckgeben.

```js
assert.equal(getFlag("IN"), "宖佖宖侊")
```

`getFlag("IR")` sollte `"宖佖宖俉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("IR"), "宖佖宖俉")
```

`getFlag("IE")` sollte `"宖佖宖伋"` zurﾃｼckgeben.

```js
assert.equal(getFlag("IE"), "宖佖宖伋")
```

`getFlag("IL")` sollte `"宖佖宖佬"` zurﾃｼckgeben.

```js
assert.equal(getFlag("IL"), "宖佖宖佬")
```

`getFlag("IT")` sollte `"宖佖宖倁"` zurﾃｼckgeben.

```js
assert.equal(getFlag("IT"), "宖佖宖倁")
```

`getFlag("JM")` sollte `"宖佟宖佾"` zurﾃｼckgeben.

```js
assert.equal(getFlag("JM"), "宖佟宖佾")
```

`getFlag("JP")` sollte `"宖佟宖侗"` zurﾃｼckgeben.

```js
assert.equal(getFlag("JP"), "宖佟宖侗")
```

`getFlag("KZ")` sollte `"宖佪宖偀"` zurﾃｼckgeben.

```js
assert.equal(getFlag("KZ"), "宖佪宖偀")
```

`getFlag("KE")` sollte `"宖佪宖伋"` zurﾃｼckgeben.

```js
assert.equal(getFlag("KE"), "宖佪宖伋")
```

`getFlag("XK")` sollte `"宖倧宖佪"` zurﾃｼckgeben.

```js
assert.equal(getFlag("XK"), "宖倧宖佪")
```

`getFlag("KG")` sollte `"宖佪宖佈"` zurﾃｼckgeben.

```js
assert.equal(getFlag("KG"), "宖佪宖佈")
```

`getFlag("LV")` sollte `"宖佬宖倎"` zurﾃｼckgeben.

```js
assert.equal(getFlag("LV"), "宖佬宖倎")
```

`getFlag("LB")` sollte `"宖佬宖份"` zurﾃｼckgeben.

```js
assert.equal(getFlag("LB"), "宖佬宖份")
```

`getFlag("LI")` sollte `"宖佬宖佖"` zurﾃｼckgeben.

```js
assert.equal(getFlag("LI"), "宖佬宖佖")
```

`getFlag("LT")` sollte `"宖佬宖倁"` zurﾃｼckgeben.

```js
assert.equal(getFlag("LT"), "宖佬宖倁")
```

`getFlag("LU")` sollte `"宖佬宖倂"` zurﾃｼckgeben.

```js
assert.equal(getFlag("LU"), "宖佬宖倂")
```

`getFlag("MG")` sollte `"宖佾宖佈"` zurﾃｼckgeben.

```js
assert.equal(getFlag("MG"), "宖佾宖佈")
```

`getFlag("MY")` sollte `"宖佾宖倮"` zurﾃｼckgeben.

```js
assert.equal(getFlag("MY"), "宖佾宖倮")
```

`getFlag("MT")` sollte `"宖佾宖倁"` zurﾃｼckgeben.

```js
assert.equal(getFlag("MT"), "宖佾宖倁")
```

`getFlag("MX")` sollte `"宖佾宖倧"` zurﾃｼckgeben.

```js
assert.equal(getFlag("MX"), "宖佾宖倧")
```

`getFlag("MD")` sollte `"宖佾宖伃"` zurﾃｼckgeben.

```js
assert.equal(getFlag("MD"), "宖佾宖伃")
```

`getFlag("MC")` sollte `"宖佾宖仿"` zurﾃｼckgeben.

```js
assert.equal(getFlag("MC"), "宖佾宖仿")
```

`getFlag("MN")` sollte `"宖佾宖侊"` zurﾃｼckgeben.

```js
assert.equal(getFlag("MN"), "宖佾宖侊")
```

`getFlag("ME")` sollte `"宖佾宖伋"` zurﾃｼckgeben.

```js
assert.equal(getFlag("ME"), "宖佾宖伋")
```

`getFlag("MA")` sollte `"宖佾宖仡"` zurﾃｼckgeben.

```js
assert.equal(getFlag("MA"), "宖佾宖仡")
```

`getFlag("NL")` sollte `"宖侊宖佬"` zurﾃｼckgeben.

```js
assert.equal(getFlag("NL"), "宖侊宖佬")
```

`getFlag("NZ")` sollte `"宖侊宖偀"` zurﾃｼckgeben.

```js
assert.equal(getFlag("NZ"), "宖侊宖偀")
```

`getFlag("NG")` sollte `"宖侊宖佈"` zurﾃｼckgeben.

```js
assert.equal(getFlag("NG"), "宖侊宖佈")
```

`getFlag("MK")` sollte `"宖佾宖佪"` zurﾃｼckgeben.

```js
assert.equal(getFlag("MK"), "宖佾宖佪")
```

`getFlag("NO")` sollte `"宖侊宖侔"` zurﾃｼckgeben.

```js
assert.equal(getFlag("NO"), "宖侊宖侔")
```

`getFlag("PK")` sollte `"宖侗宖佪"` zurﾃｼckgeben.

```js
assert.equal(getFlag("PK"), "宖侗宖佪")
```

`getFlag("PH")` sollte `"宖侗宖佉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("PH"), "宖侗宖佉")
```

`getFlag("PL")` sollte `"宖侗宖佬"` zurﾃｼckgeben.

```js
assert.equal(getFlag("PL"), "宖侗宖佬")
```

`getFlag("PT")` sollte `"宖侗宖倁"` zurﾃｼckgeben.

```js
assert.equal(getFlag("PT"), "宖侗宖倁")
```

`getFlag("PR")` sollte `"宖侗宖俉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("PR"), "宖侗宖俉")
```

`getFlag("RO")` sollte `"宖俉宖侔"` zurﾃｼckgeben.

```js
assert.equal(getFlag("RO"), "宖俉宖侔")
```

`getFlag("SM")` sollte `"宖俠宖佾"` zurﾃｼckgeben.

```js
assert.equal(getFlag("SM"), "宖俠宖佾")
```

`getFlag("SA")` sollte `"宖俠宖仡"` zurﾃｼckgeben.

```js
assert.equal(getFlag("SA"), "宖俠宖仡")
```

`getFlag("RS")` sollte `"宖俉宖俠"` zurﾃｼckgeben.

```js
assert.equal(getFlag("RS"), "宖俉宖俠")
```

`getFlag("SG")` sollte `"宖俠宖佈"` zurﾃｼckgeben.

```js
assert.equal(getFlag("SG"), "宖俠宖佈")
```

`getFlag("SK")` sollte `"宖俠宖佪"` zurﾃｼckgeben.

```js
assert.equal(getFlag("SK"), "宖俠宖佪")
```

`getFlag("SI")` sollte `"宖俠宖佖"` zurﾃｼckgeben.

```js
assert.equal(getFlag("SI"), "宖俠宖佖")
```

`getFlag("ZA")` sollte `"宖偀宖仡"` zurﾃｼckgeben.

```js
assert.equal(getFlag("ZA"), "宖偀宖仡")
```

`getFlag("KR")` sollte `"宖佪宖俉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("KR"), "宖佪宖俉")
```

`getFlag("ES")` sollte `"宖伋宖俠"` zurﾃｼckgeben.

```js
assert.equal(getFlag("ES"), "宖伋宖俠")
```

`getFlag("SE")` sollte `"宖俠宖伋"` zurﾃｼckgeben.

```js
assert.equal(getFlag("SE"), "宖俠宖伋")
```

`getFlag("CH")` sollte `"宖仿宖佉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("CH"), "宖仿宖佉")
```

`getFlag("TH")` sollte `"宖倁宖佉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("TH"), "宖倁宖佉")
```

`getFlag("TT")` sollte `"宖倁宖倁"` zurﾃｼckgeben.

```js
assert.equal(getFlag("TT"), "宖倁宖倁")
```

`getFlag("TR")` sollte `"宖倁宖俉"` zurﾃｼckgeben.

```js
assert.equal(getFlag("TR"), "宖倁宖俉")
```

`getFlag("UA")` sollte `"宖倂宖仡"` zurﾃｼckgeben.

```js
assert.equal(getFlag("UA"), "宖倂宖仡")
```

`getFlag("AE")` sollte `"宖仡宖伋"` zurﾃｼckgeben.

```js
assert.equal(getFlag("AE"), "宖仡宖伋")
```

`getFlag("US")` sollte `"宖倂宖俠"` zurﾃｼckgeben.

```js
assert.equal(getFlag("US"), "宖倂宖俠")
```

`getFlag("UY")` sollte `"宖倂宖倮"` zurﾃｼckgeben.

```js
assert.equal(getFlag("UY"), "宖倂宖倮")
```

`getFlag("UZ")` sollte `"宖倂宖偀"` zurﾃｼckgeben.

```js
assert.equal(getFlag("UZ"), "宖倂宖偀")
```

`getFlag("VE")` sollte `"宖倎宖伋"` zurﾃｼckgeben.

```js
assert.equal(getFlag("VE"), "宖倎宖伋")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
