---
id: 697a49e9860d24853adef67e
title: "Aufgabe 193: Winterspiele 2026 Tag 14: Skibergsteigen"
challengeType: 28
dashedName: challenge-193
---

# --description--

Anhand der Schneetiefe und des Gefälles eines Bergs bestimmen Sie, ob eine Lawinengefahr besteht.

- Die Werte für die Schneetiefe sind `"Shallow"`, `"Moderate"` oder `"Deep"`.
- Die Werte für das Gefälle sind `"Gentle"`, `"Steep"` oder `"Very Steep"`.

Geben Sie `"Safe"` oder `"Risky"` basierend auf dieser Tabelle zurück:

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalancheRisk("Shallow", "Gentle")` sollte `"Safe"` zurückgeben.

```js
assert.equal(avalancheRisk("Shallow", "Gentle"), "Safe");
```

`avalancheRisk("Shallow", "Steep")` sollte `"Safe"` zurückgeben.

```js
assert.equal(avalancheRisk("Shallow", "Steep"), "Safe");
```

`avalancheRisk("Shallow", "Very Steep")` sollte `"Safe"` zurückgeben.

```js
assert.equal(avalancheRisk("Shallow", "Very Steep"), "Safe");
```

`avalancheRisk("Moderate", "Gentle")` sollte `"Safe"` zurückgeben.

```js
assert.equal(avalancheRisk("Moderate", "Gentle"), "Safe");
```

`avalancheRisk("Moderate", "Steep")` sollte `"Risky"` zurückgeben.

```js
assert.equal(avalancheRisk("Moderate", "Steep"), "Risky");
```

`avalancheRisk("Moderate", "Very Steep")` sollte `"Risky"` zurückgeben.

```js
assert.equal(avalancheRisk("Moderate", "Very Steep"), "Risky");
```

`avalancheRisk("Deep", "Gentle")` sollte `"Safe"` zurückgeben.

```js
assert.equal(avalancheRisk("Deep", "Gentle"), "Safe");
```

`avalancheRisk("Deep", "Steep")` sollte `"Risky"` zurückgeben.

```js
assert.equal(avalancheRisk("Deep", "Steep"), "Risky");
```

`avalancheRisk("Deep", "Very Steep")` sollte `"Risky"` zurückgeben.

```js
assert.equal(avalancheRisk("Deep", "Very Steep"), "Risky");
```

# --seed--

## --seed-contents--

```js
function avalancheRisk(snowDepth, slope) {

  return snowDepth;
}
```

# --solutions--

```js
function avalancheRisk(snowDepth, slope) {
  const riskTable = {
    Gentle: {
      Shallow: "Safe",
      Moderate: "Safe",
      Deep: "Safe",
    },
    Steep: {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
    "Very Steep": {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
  };

  return riskTable[slope][snowDepth];
}
```
