---
id: 698a1a863194f1f4e63f645f
title: "Aufgabe 208: Pfaddurchquerung"
challengeType: 28
dashedName: challenge-208
---

# --description--

Gegeben ist ein Array von Strings, die Ihre Pfadkarte darstellen. Geben Sie einen String mit den Bewegungen zurück, die nötig sind, um Ihr Ziel zu erreichen.

Die gegebenen Strings enthalten die Werte:

- `"C"`: Ihren aktuellen Standort
- `"G"`: Ihr Ziel
- `"T"`: Begehbare Teile des Pfads
- `"-"`: Unbegehbare Teile der Karte

Geben Sie einen String mit den Bewegungen zurück, die nötig sind, um dem Pfad von Ihrem Standort bis zu Ihrem Ziel zu folgen, wobei:

- `"R"` eine Bewegung nach rechts ist
- `"D"` eine Bewegung nach unten ist
- `"L"` eine Bewegung nach links ist
- `"U"` eine Bewegung nach oben ist

- Es gibt immer genau einen durchgehenden Pfad ohne Verzweigungen von Ihrem aktuellen Standort bis zu Ihrem Ziel.
- Jeder Pfadabschnitt hat maximal zwei begehbare Nachbarfelder.

Zum Beispiel, gegeben:

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

Geben Sie `"RDDRDD"` zurück.

# --hints--

`navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` sollte `"RDDRDD"` zurückgeben.

```js
assert.equal(navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD");
```

`navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` sollte `"RRUUURR"` zurückgeben.

```js
assert.equal(navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR");
```

`navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` sollte `"DLDDRRRRD"` zurückgeben.

```js
assert.equal(navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD");
```

`navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` sollte `"RRRDDL"` zurückgeben.

```js
assert.equal(navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL");
```

`navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` sollte `"ULLLUUURRRRRRDDDR"` zurückgeben.

```js
assert.equal(navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR");
```

# --seed--

## --seed-contents--

```js
function navigateTrail(map) {

  return map;
}
```

# --solutions--

```js
function navigateTrail(map) {
  const rows = map.length;
  const cols = map[0].length;

  const directions = [
    [0, 1, "R"],
    [1, 0, "D"],
    [0, -1, "L"],
    [-1, 0, "U"],
  ];

  let row, col;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (map[r][c] === "C") {
        row = r;
        col = c;
      }
    }
  }

  let result = "";
  let prev = null;

  while (map[row][col] !== "G") {
    for (let [dr, dc, move] of directions) {
      const newRow = row + dr;
      const newCol = col + dc;

      if (
        newRow >= 0 &&
        newRow < rows &&
        newCol >= 0 &&
        newCol < cols &&
        (map[newRow][newCol] === "T" ||
          map[newRow][newCol] === "G") &&
        (!prev || newRow !== prev[0] || newCol !== prev[1])
      ) {
        result += move;
        prev = [row, col];
        row = newRow;
        col = newCol;
        break;
      }
    }
  }

  return result;
}
```
