---
id: 698a1a73ade5ac0e19180fa9
title: "Aufgabe 206: Werte von Spielkarten"
challengeType: 28
dashedName: challenge-206
---

# --description--

Gegeben ist ein Array von Spielkarten. Geben Sie ein neues Array mit dem Zahlenwert jeder Karte zurück.

Kartenwerte:

- Ein Ass (`"A"`) hat den Wert 1.
- Zahlenkarten (`"2"` - `"10"`) haben ihren aufgedruckten Wert: 2 bis 10.
- Bildkarten: Bube (`"J"`), Dame (`"Q"`) und König (`"K"`) sind jeweils 10 wert.

Farben:

- Jede Karte hat eine Farbe: Pik (`"S"`), Kreuz (`"C"`), Karo (`"D"`) oder Herz (`"H"`).

Kartenformat:

- Jede Karte wird als String dargestellt: `"valueSuit"`. Zum Beispiel: `"AS"` ist das Pik-Ass, `"10H"` ist die Zehn Herz und `"QC"` ist die Kreuz-Dame.

# --hints--

`cardValues(["3H", "4D", "5S"])` sollte `[3, 4, 5]` zurückgeben.

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` sollte `[1, 10, 10, 6, 7]` zurückgeben.

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` sollte `[8, 10, 2, 10, 9]` zurückgeben.

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` sollte `[1, 10]` zurückgeben.

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` sollte `[10, 10, 10, 10, 1]` zurückgeben.

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
