---
id: 698a1a73ade5ac0e19180fa8
title: "Aufgabe 205: Anzahl perfekter Kuben"
challengeType: 28
dashedName: challenge-205
---

# --description--

Gegeben sind zwei ganze Zahlen. Bestimmen Sie, wie viele perfekte Kuben im Bereich zwischen und einschließlich dieser beiden Zahlen liegen.

- Die kleinere Zahl ist nicht zwangsläufig das erste Argument.
- Eine Zahl ist ein perfekter Kubus, wenn es eine ganze Zahl (`n`) gibt, für die `n * n * n = number` gilt. Zum Beispiel ist 27 ein perfekter Kubus, weil `3 * 3 * 3 = 27`.

# --hints--

`countPerfectCubes(3, 30)` sollte `2` zurückgeben.

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)` sollte `3` zurückgeben.

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)` sollte `4` zurückgeben.

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)` sollte `9` zurückgeben.

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)` sollte `41` zurückgeben.

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
