---
id: 698a1a73ade5ac0e19180fa7
title: "Aufgabe 204: Buchstaben summieren"
challengeType: 28
dashedName: challenge-204
---

# --description--

Gegeben ist ein String. Geben Sie die Summe seiner Buchstaben zurück.

- Buchstaben sind A-Z in Groß- oder Kleinbuchstaben
- Buchstabenwerte sind: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26
- Groß- und Kleinbuchstaben haben denselben Wert.
- Ignorieren Sie alle Nicht-Buchstaben-Zeichen.

# --hints--

`sumLetters("Hello")` sollte `52` zurückgeben.

```js
assert.equal(sumLetters("Hello"), 52);
```

`sumLetters("freeCodeCamp")` sollte `94` zurückgeben.

```js
assert.equal(sumLetters("freeCodeCamp"), 94);
```

`sumLetters("The quick brown fox jumps over the lazy dog.")` sollte `473` zurückgeben.

```js
assert.equal(sumLetters("The quick brown fox jumps over the lazy dog."), 473);
```

`sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` sollte `1681` zurückgeben.

```js
assert.equal(sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681);
```

`sumLetters("</404>")` sollte `0` zurückgeben.

```js
assert.equal(sumLetters("</404>"), 0);
```

# --seed--

## --seed-contents--

```js
function sumLetters(str) {

  return str;
}
```

# --solutions--

```js
function sumLetters(str) {
  let total = 0;
  for (let char of str) {
    let upper = char.toUpperCase();
    if (upper >= "A" && upper <= "Z") {
      total += upper.charCodeAt(0) - "A".charCodeAt(0) + 1;
    }
  }
  return total;
}
```
