---
id: 698a1a863194f1f4e63f6460
title: "Aufgabe 209: Elementgröße"
challengeType: 28
dashedName: challenge-209
---

# --description--

Gegeben sind eine Fenstergröße, die Breite eines Elements in Viewport-Breite `"vw"`-Einheiten und die Höhe eines Elements in Viewport-Höhe `"vh"`-Einheiten. Bestimmen Sie die Größe des Elements in Pixeln.

- Die angegebene Fenstergröße und die zurückgegebene Elementgröße sind Strings im Format `"width x height"`, `"1200 x 800"`, zum Beispiel.

- `"vw"`-Einheiten sind der Prozentsatz der Fensterbreite. `"50vw"` ist zum Beispiel 50 % der Fensterbreite.
- `"vh"`-Einheiten sind der Prozentsatz der Fensterhöhe. `"50vh"` ist zum Beispiel 50 % der Fensterhöhe.

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` sollte `"600 x 400"` zurückgeben.

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` sollte `"80 x 240"` zurückgeben.

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` sollte `"70 x 15"` zurückgeben.

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` sollte `"1824 x 1080"` zurückgeben.

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` sollte `"0 x 0"` zurückgeben.

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` sollte `"1440 x 1026"` zurückgeben.

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```
