---
id: 698a1a863194f1f4e63f645e
title: "Aufgabe 207: Kleinste Lücke"
challengeType: 28
dashedName: challenge-207
---

# --description--

Gegeben ist ein String. Geben Sie den Teilstring zwischen den zwei identischen Zeichen zurück, die die kleinste Anzahl von Zeichen zwischen sich haben (kleinste Lücke).

- Es gibt immer mindestens ein Paar übereinstimmender Zeichen.
- Der zurückgegebene Teilstring soll die übereinstimmenden Zeichen ausschließen.
- Wenn zwei oder mehr Lücken gleich lang sind, geben Sie die Zeichen der ersten zurück.

Zum Beispiel, bei `"ABCDAC"`, geben Sie `"DA"` zurück, weil:

- Nur `"A"` und `"C"` wiederholen sich im String.
- Die Anzahl der Zeichen zwischen den beiden `"A"`-Zeichen ist 3, und zwischen den `"C"`-Zeichen ist 2.
- Also geben Sie den String zwischen den beiden `"C"`-Zeichen zurück.

# --hints--

`smallestGap("ABCDAC")` sollte `"DA"` zurückgeben.

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` sollte `"e"` zurückgeben.

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` sollte `"#q6e&rkf(p"` zurückgeben.

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` sollte `""` zurückgeben.

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` sollte `"fox"` zurückgeben.

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
