---
id: 698a1a73ade5ac0e19180fa1
title: "Aufgabe 198: Anzahl der Arbeitstage"
challengeType: 28
dashedName: challenge-198
---

# --description--

Gegeben sind ein Startdatum und ein Enddatum. Geben Sie die Anzahl der Arbeitstage zwischen den beiden Daten zurück.

- Die angegebenen Daten haben das Format `"YYYY-MM-DD"`.
- Wochentage sind Arbeitstage (Montag bis Freitag).
- Wochenenden sind keine Arbeitstage (Samstag und Sonntag).
- Zählen Sie sowohl das Start- als auch das Enddatum mit.

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")` sollte `3` zurückgeben.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")` sollte `4` zurückgeben.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")` sollte `5` zurückgeben.

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")` sollte `7` zurückgeben.

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")` sollte `262` zurückgeben.

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
