---
id: 698a1a73ade5ac0e19180fa0
title: "Aufgabe 197: Blutgruppenverträglichkeit"
challengeType: 28
dashedName: challenge-197
---

# --description--

Gegeben sind eine Blutgruppe des Spenders und eine Blutgruppe des Empfängers. Bestimmen Sie, ob der Spender dem Empfänger Blut geben kann.

Jede Blutgruppe besteht aus:

- Einem Buchstaben: `"A"`, `"B"`, `"AB"` oder `"O"`
- Und einem Rhesusfaktor: `"+"` oder `"-"`

Blutgruppen bestehen aus einem der gültigen Buchstaben, gefolgt von einem Rhesusfaktor. Zum Beispiel sind `"AB+"` und `"O-"` gültige Blutgruppen.

Regeln für den Buchstaben:

- `"O"` kann an andere Buchstabentypen spenden.
- `"A"` kann an `"A"` und `"AB"` spenden.
- `"B"` kann an `"B"` und `"AB"` spenden.
- `"AB"` kann nur an `"AB"` spenden.

Regeln für den Rhesusfaktor:

- Negativ (`"-"`) kann an sowohl `"-"` als auch `"+"` spenden.
- Positiv (`"+"`) kann nur an `"+"` spenden.

Sowohl die Buchstabenregel als auch die Rhesusregel müssen erfüllt sein, damit ein Spender dem Empfänger Blut geben kann.

# --hints--

`canDonate("B+", "B+")` sollte `true` zurückgeben.

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` sollte `true` zurückgeben.

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` sollte `false` zurückgeben.

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` sollte `true` zurückgeben.

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` sollte `false` zurückgeben.

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` sollte `true` zurückgeben.

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` sollte `false` zurückgeben.

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` sollte `true` zurückgeben.

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` sollte `false` zurückgeben.

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` sollte `false` zurückgeben.

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
