---
id: 698a1a73ade5ac0e19180fa5
title: "Aufgabe 202: Satzzeichen hinzufügen"
challengeType: 28
dashedName: challenge-202
---

# --description--

Gegeben ist ein String mit Sätzen, bei denen Punkte fehlen. Fügen Sie an den folgenden Stellen einen Punkt (`"."`) ein:

- Vor jedem Leerzeichen, das unmittelbar vor einem Großbuchstaben steht
- Und am Ende des Strings

Geben Sie den resultierenden String zurück.

# --hints--

`addPunctuation("Hello world")` sollte `"Hello world."` zurückgeben

```js
assert.equal(addPunctuation("Hello world"), "Hello world.");
```

`addPunctuation("Hello world It's nice today")` sollte `"Hello world. It's nice today."` zurückgeben

```js
assert.equal(addPunctuation("Hello world It's nice today"), "Hello world. It's nice today.");
```

`addPunctuation("JavaScript is great Sometimes")` sollte `"JavaScript is great. Sometimes."` zurückgeben

```js
assert.equal(addPunctuation("JavaScript is great Sometimes"), "JavaScript is great. Sometimes.");
```

`addPunctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z")` sollte `"A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z."` zurückgeben

```js
assert.equal(addPunctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z"), "A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z.");
```

`addPunctuation("Wait.. For it")` sollte `"Wait... For it."` zurückgeben

```js
assert.equal(addPunctuation("Wait.. For it"), "Wait... For it.");
```

# --seed--

## --seed-contents--

```js
function addPunctuation(sentences) {

  return sentences;
}
```

# --solutions--

```js
function addPunctuation(sentences) {
  let result = "";

  for (let i = 0; i < sentences.length; i++) {
    if (sentences[i] === " " && i + 1 < sentences.length && /[A-Z]/.test(sentences[i + 1])) {
      result += ".";
    }
    result += sentences[i];
  }
  result += ".";

  return result;
}
```
