---
id: 698a1a73ade5ac0e19180fa6
title: "Aufgabe 203: Geflattet"
challengeType: 28
dashedName: challenge-203
---

# --description--

Gegeben ist ein Array. Bestimmen Sie, ob es flach ist.

- Ein Array ist flach, wenn keines seiner Elemente ein Array ist.

# --hints--

`isFlat([1, 2, 3, 4])` sollte `true` zurückgeben.

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])` sollte `false` zurückgeben.

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])` sollte `true` zurückgeben.

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])` sollte `false` zurückgeben.

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])` sollte `false` zurückgeben.

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```
