---
id: 698a1a73ade5ac0e19180fa3
title: "Aufgabe 200: Buchstaben- und Zahlenzählung"
challengeType: 28
dashedName: challenge-200
---

# --description--

Gegeben ist ein String. Geben Sie eine Nachricht zurück, die zählt, wie viele Buchstaben und Zahlen er enthält.

- Buchstaben sind `A-Z` und `a-z`.
- Zahlen sind `0-9`.
- Ignorieren Sie alle anderen Zeichen.

Geben Sie `"The string has X letters and Y numbers."` zurück, wobei `"X"` die Anzahl der Buchstaben und `"Y"` die Anzahl der Zahlen ist. Wenn eine der beiden Zahlen 1 ist, verwenden Sie die Einzahl für diesen Begriff. Zum Beispiel: `"1 letter"` statt `"1 letters"` und `"1 number"` statt `"1 numbers"`.

# --hints--

`countLettersAndNumbers("helloworld123")` sollte `"The string has 10 letters and 3 numbers."` zurückgeben.

```js
assert.equal(countLettersAndNumbers("helloworld123"), "The string has 10 letters and 3 numbers.");
```

`countLettersAndNumbers("Catch 22")` sollte `"The string has 5 letters and 2 numbers."` zurückgeben.

```js
assert.equal(countLettersAndNumbers("Catch 22"), "The string has 5 letters and 2 numbers.");
```

`countLettersAndNumbers("A1!")` sollte `"The string has 1 letter and 1 number."` zurückgeben.

```js
assert.equal(countLettersAndNumbers("A1!"), "The string has 1 letter and 1 number.");
```

`countLettersAndNumbers("12345")` sollte `"The string has 0 letters and 5 numbers."` zurückgeben.

```js
assert.equal(countLettersAndNumbers("12345"), "The string has 0 letters and 5 numbers.");
```

`countLettersAndNumbers("password")` sollte `"The string has 8 letters and 0 numbers."` zurückgeben.

```js
assert.equal(countLettersAndNumbers("password"), "The string has 8 letters and 0 numbers.");
```

# --seed--

## --seed-contents--

```js
function countLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function countLettersAndNumbers(str) {
  let letterCount = 0;
  let numberCount = 0;

  for (let char of str) {
    if ((char >= 'A' && char <= 'Z') || (char >= 'a' && char <= 'z')) {
      letterCount++;
    } else if (char >= '0' && char <= '9') {
      numberCount++;
    }
  }

  const letterWord = letterCount === 1 ? "letter" : "letters";
  const numberWord = numberCount === 1 ? "number" : "numbers";

  return `The string has ${letterCount} ${letterWord} and ${numberCount} ${numberWord}.`;
}
```
