---
id: 698a1a73ade5ac0e19180fa4
title: "Aufgabe 201: Matrixverschiebung"
challengeType: 28
dashedName: challenge-201
---

# --description--

Gegeben ist eine Matrix (Array von Arrays) mit Zahlen und eine ganze Zahl. Verschieben Sie alle Werte in der Matrix um den angegebenen Betrag.

- Eine positive Verschiebung bewegt die Werte nach rechts.
- Eine negative Verschiebung bewegt die Werte nach links.

Die Werte sollen umschlagen:

- Behandeln Sie die Matrix als eine durchgehende Folge von Werten.
- Wenn ein Wert über das Ende einer Zeile hinaus verschoben wird, setzt er am Anfang der nächsten Zeile fort.
- Wenn ein Wert über die letzte Position der Matrix hinaus verschoben wird, springt er zur ersten Position zurück.
- Das Gleiche gilt umgekehrt bei einer Verschiebung nach links.

Zum Beispiel, gegeben:

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

mit einer Verschiebung von `1`, verschieben Sie alle Zahlen um eins nach rechts:

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shiftMatrix([[1, 2, 3], [4, 5, 6]], 1)` sollte `[[6, 1, 2], [3, 4, 5]]` zurückgeben.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6]], -1)` sollte `[[2, 3, 4], [5, 6, 1]]` zurückgeben.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)` sollte `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]` zurückgeben.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)` sollte `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]` zurückgeben.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)` sollte `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]` zurückgeben.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)` sollte `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]` zurückgeben.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]);
```

# --seed--

## --seed-contents--

```js
function shiftMatrix(matrix, shift) {

  return matrix;
}
```

# --solutions--

```js
function shiftMatrix(matrix, shift) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const flat = matrix.flat();
  const n = flat.length;
  const s = ((shift % n) + n) % n;
  const shifted = flat.slice(-s).concat(flat.slice(0, n - s));
  const result = [];
  for (let i = 0; i < rows; i++) {
    result.push(shifted.slice(i * cols, (i + 1) * cols));
  }

  return result;
}
```
