---
id: 6994cff2290543b3aec9f50f
title: "Aufgabe 210: HSL-Validator"
challengeType: 28
dashedName: challenge-210
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob es sich um einen gültigen CSS `hsl()`-Farbwert handelt.

- Ein gültiger HSL-Wert muss im Format `"hsl(h, s%, l%)"` vorliegen, wobei:
  - `h` (Farbton) eine Zahl zwischen 0 und 360 (einschließlich) sein muss.
  - `s` (Sättigung) ein Prozentsatz zwischen 0 % und 100 % sein muss.
  - `l` (Helligkeit) ein Prozentsatz zwischen 0 % und 100 % sein muss.

- Leerzeichen sind nur erlaubt:
  - Nach der öffnenden Klammer
  - Vor und/oder nach den Kommas
  - Vor und/oder nach der schließenden Klammer

- Optional kann der Wert mit einem Semikolon (`";"`) enden.

Zum Beispiel ist `"hsl(240, 50%, 50%)"` ein gültiger HSL-Wert.

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
