---
id: 6994cff2290543b3aec9f511
title: "Aufgabe 212: Array-Einfügung"
challengeType: 28
dashedName: challenge-212
---

# --description--

Gegeben ist ein Array, ein Wert, der in das Array eingefügt werden soll, und ein Index, an dem der Wert eingefügt werden soll. Geben Sie ein neues Array zurück, in das der Wert an dem angegebenen Index eingefügt wurde.

# --hints--

`insertIntoArray([2, 4, 8, 10], 6, 2)` sollte `[2, 4, 6, 8, 10]` zurückgeben.

```js
assert.deepEqual(insertIntoArray([2, 4, 8, 10], 6, 2), [2, 4, 6, 8, 10]);
```

`insertIntoArray(["the", "quick", "fox"], "brown", 2)` sollte `["the", "quick", "brown", "fox"]` zurückgeben.

```js
assert.deepEqual(insertIntoArray(["the", "quick", "fox"], "brown", 2), ["the", "quick", "brown", "fox"]);
```

`insertIntoArray([], 0, 0)` sollte `[0]` zurückgeben.

```js
assert.deepEqual(insertIntoArray([], 0, 0), [0]);
```

`insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5)` sollte `[0, 1, 1, 2, 3, 5, 8, 13]` zurückgeben.

```js
assert.deepEqual(insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5), [0, 1, 1, 2, 3, 5, 8, 13]);
```

# --seed--

## --seed-contents--

```js
function insertIntoArray(arr, value, index) {

  return arr;
}
```

# --solutions--

```js
function insertIntoArray(arr, value, index) {
  return [
    ...arr.slice(0, index),
    value,
    ...arr.slice(index)
  ];
}
```
