---
id: 6994cff2290543b3aec9f510
title: "Aufgabe 211: Summe eines Arrays"
challengeType: 28
dashedName: challenge-211
---

# --description--

Gegeben ist ein Array von Zahlen. Geben Sie die Summe aller Zahlen zurück.

# --hints--

`sumArray([1, 2, 3, 4, 5])` sollte `15` zurückgeben.

```js
assert.equal(sumArray([1, 2, 3, 4, 5]), 15);
```

`sumArray([42])` sollte `42` zurückgeben.

```js
assert.equal(sumArray([42]), 42);
```

`sumArray([5, -2, 7, -3])` sollte `7` zurückgeben.

```js
assert.equal(sumArray([5, -2, 7, -3]), 7);
```

`sumArray([203, 145, -129, 6293, 523, -919, 845, 2434])` sollte `9395` zurückgeben.

```js
assert.equal(sumArray([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395);
```

`sumArray([0, 0])` sollte `0` zurückgeben.

```js
assert.equal(sumArray([0, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function sumArray(numbers) {

  return numbers;
}
```

# --solutions--

```js
function sumArray(numbers) {

  return numbers.reduce((a, c) => c + a, 0);
}
```
