---
id: 699c8e045ee7cb94ed2322dd
title: "Aufgabe 222: Schatten zur Tagundnachtgleiche"
challengeType: 28
dashedName: challenge-222
---

# --description--

Heute ist die Tagundnachtgleiche, wenn die Sonne direkt über dem Äquator steht und mittags genau im Zenit ist. Gegeben eine Uhrzeit, bestimmen Sie den Schatten, den ein 4 Fuß hoher senkrechter Pfosten wirft.

- Die Uhrzeit wird als String im `"HH:MM"` 24-Stunden-Format angegeben (zum Beispiel ist `"15:00"` 15 Uhr).
- Sie erhalten nur Uhrzeiten in 30-Minuten-Schritten.

Regeln:

- Die Sonne geht um 6 Uhr morgens direkt `"east"` auf und um 18 Uhr direkt `"west"` unter.
- Ein Schatten zeigt immer entgegengesetzt zur Sonne.
- Die Länge des Schattens (in Fuß) ist die Anzahl der Stunden vom Mittag entfernt, hoch drei.
- Vor Sonnenaufgang (vor 6 Uhr), nach Sonnenuntergang (18 Uhr oder später) und mittags gibt es keinen Schatten.

Rückgabe:

- Wenn ein Schatten existiert, geben Sie `"(length)ft (direction)"` zurück. Zum Beispiel `"8ft west"`.
- Andernfalls geben Sie `"No shadow"` zurück.

Beispielsweise wird bei `"10:00"` `"8ft west"` zurückgegeben, weil 10 Uhr 2 Stunden vom Mittag entfernt ist, also 2<sup>3</sup> = 8 Fuß, und der Schatten nach Westen zeigt, weil die Sonne um 10 Uhr im Osten steht.

# --hints--

`getShadow("10:00")` sollte `"8ft west"` zurückgeben.

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")` sollte `"27ft east"` zurückgeben.

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")` sollte `"No shadow"` zurückgeben.

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")` sollte `"166.375ft east"` zurückgeben.

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")` sollte `"No shadow"` zurückgeben.

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")` sollte `"216ft west"` zurückgeben.

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")` sollte `"No shadow"` zurückgeben.

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")` sollte `"91.125ft west"` zurückgeben.

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")` sollte `"No shadow"` zurückgeben.

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```
