---
id: 699c8e045ee7cb94ed2322dc
title: "Aufgabe 221: Invertierte Matrix"
challengeType: 28
dashedName: challenge-221
---

# --description--

Gegeben ist eine Matrix (ein Array von Arrays), die mit zwei unterschiedlichen Werten gefüllt ist. Geben Sie eine neue Matrix zurück, in der alle Vorkommen des einen Werts mit dem anderen vertauscht sind.

Zum Beispiel, gegeben:

```js
[
  ["a", "b"],
  ["a", "a"]
]
```

Geben Sie zurück:

```js
[
  ["b", "a"],
  ["b", "b"]
]
```

# --hints--

`invertMatrix([["a", "b"], ["a", "a"]])` sollte `[["b", "a"], ["b", "b"]]` zurückgeben.

```js
assert.deepEqual(invertMatrix([["a", "b"], ["a", "a"]]), [["b", "a"], ["b", "b"]]);
```

`invertMatrix([[1, 0, 1], [1, 1, 1], [0, 1, 0]])` sollte `[[0, 1, 0], [0, 0, 0], [1, 0, 1]]` zurückgeben.

```js
assert.deepEqual(invertMatrix([[1, 0, 1], [1, 1, 1], [0, 1, 0]]), [[0, 1, 0], [0, 0, 0], [1, 0, 1]]);
```

`invertMatrix([["apple", "banana", "banana", "apple"], ["banana", "apple", "apple", "banana"], ["banana", "banana", "banana", "apple"]])` sollte `[["banana", "apple", "apple", "banana"], ["apple", "banana", "banana", "apple"], ["apple", "apple", "apple", "banana"]]` zurückgeben.

```js
assert.deepEqual(invertMatrix([["apple", "banana", "banana", "apple"], ["banana", "apple", "apple", "banana"], ["banana", "banana", "banana", "apple"]]), [["banana", "apple", "apple", "banana"], ["apple", "banana", "banana", "apple"], ["apple", "apple", "apple", "banana"]]);
```

`invertMatrix([[6, 7, 7, 7, 6], [7, 6, 7, 6, 7], [7, 7, 6, 7, 7], [7, 6, 7, 6, 7], [6, 7, 7, 7, 6]])` sollte `[[7, 6, 6, 6, 7], [6, 7, 6, 7, 6], [6, 6, 7, 6, 6], [6, 7, 6, 7, 6], [7, 6, 6, 6, 7]]` zurückgeben.

```js
assert.deepEqual(invertMatrix([[6, 7, 7, 7, 6], [7, 6, 7, 6, 7], [7, 7, 6, 7, 7], [7, 6, 7, 6, 7], [6, 7, 7, 7, 6]]), [[7, 6, 6, 6, 7], [6, 7, 6, 7, 6], [6, 6, 7, 6, 6], [6, 7, 6, 7, 6], [7, 6, 6, 6, 7]]);
```

`invertMatrix([[1.2, 2.1, 2.1, 2.1], [2.1, 1.2, 2.1, 1.2], [1.2, 1.2, 2.1, 2.1]])` sollte `[[2.1, 1.2, 1.2, 1.2], [1.2, 2.1, 1.2, 2.1], [2.1, 2.1, 1.2, 1.2]]` zurückgeben.

```js
assert.deepEqual(invertMatrix([[1.2, 2.1, 2.1, 2.1], [2.1, 1.2, 2.1, 1.2], [1.2, 1.2, 2.1, 2.1]]), [[2.1, 1.2, 1.2, 1.2], [1.2, 2.1, 1.2, 2.1], [2.1, 2.1, 1.2, 1.2]]);
```

# --seed--

## --seed-contents--

```js
function invertMatrix(matrix) {

  return matrix;
}
```

# --solutions--

```js
function invertMatrix(matrix) {
  if (matrix.length === 0) return [];

  const values = [...new Set(matrix.flat())];
  const [val1, val2] = values;

  return matrix.map(row =>
    row.map(cell => (cell === val1 ? val2 : val1))
  );
}
```
