---
id: 699c8e045ee7cb94ed2322db
title: "Aufgabe 220: Größte Zahl"
challengeType: 28
dashedName: challenge-220
---

# --description--

Gegeben ist ein String aus Zahlen, die durch verschiedene Satzzeichen getrennt sind. Geben Sie die größte Zahl zurück.

- Der gegebene String enthält nur Zahlen und Trennzeichen.
- Trennzeichen können Kommas (`","`), Ausrufezeichen (`"!"`), Fragezeichen (`"?"`), Doppelpunkte (`":"`) oder Semikolons (`";"`) sein.

# --hints--

`largestNumber("1,2")` sollte `2` zurückgeben.

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` sollte `60` zurückgeben.

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` sollte `-402` zurückgeben.

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` sollte `99.9` zurückgeben.

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
