---
id: 699c8e045ee7cb94ed2322d4
title: "Aufgabe 213: Wortlängen-Konverter"
challengeType: 28
dashedName: challenge-213
---

# --description--

Gegeben ist ein String aus Wörtern. Geben Sie einen neuen String zurück, in dem jedes Wort durch seine Länge ersetzt wird.

- Die Wörter im gegebenen String sind durch ein einzelnes Leerzeichen getrennt
- Behalten Sie die Leerzeichen im zurückgegebenen String bei.

Zum Beispiel, bei `"hello world"` geben Sie `"5 5"` zurück.

# --hints--

`convertWords("hello world")` sollte `"5 5"` zurückgeben.

```js
assert.equal(convertWords("hello world"), "5 5");
```

`convertWords("Thanks and happy coding")` sollte `"6 3 5 6"` zurückgeben.

```js
assert.equal(convertWords("Thanks and happy coding"), "6 3 5 6");
```

`convertWords("The quick brown fox jumps over the lazy dog")` sollte `"3 5 5 3 5 4 3 4 3"` zurückgeben.

```js
assert.equal(convertWords("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3");
```

`convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` sollte `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"` zurückgeben.

```js
assert.equal(convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4");
```

# --seed--

## --seed-contents--

```js
function convertWords(str) {

  return str;
}
```

# --solutions--

```js
function convertWords(str) {
  return str
    .split(" ")
    .map(word => word.length)
    .join(" ");
}
```
