---
id: 699c8e045ee7cb94ed2322d6
title: "Aufgabe 215: Parkgebührenrechner"
challengeType: 28
dashedName: challenge-215
---

# --description--

Gegeben sind zwei Strings, die die Zeit darstellen, zu der Sie Ihr Auto geparkt haben, und die Zeit, zu der Sie es abgeholt haben. Berechnen Sie die Parkgebühr.

- Die gegebenen Strings haben das Format `"HH:MM"` und verwenden eine 24-Stunden-Uhr. `"14:00"` entspricht zum Beispiel 14 Uhr.
- Der erste String ist die Zeit, zu der Sie Ihr Auto geparkt haben, der zweite die Zeit, zu der Sie es abgeholt haben.
- Wenn die Abholzeit früher als die Parkzeit ist, bedeutet das, dass das Parken über Mitternacht in den nächsten Tag hinausging.

Gebührenregeln:

- Jede geparkte Stunde kostet 3 $.
- Teilstunden werden auf die nächste volle Stunde aufgerundet.
- Wenn das Parken über Nacht (über Mitternacht) geht, wird eine zusätzliche Übernachtgebühr von 10 $ berechnet.
- Es gibt eine Mindestgebühr von 5 $ (wird nur angewendet, wenn die Gesamtsumme darunter liegen würde).

Geben Sie die Gesamtkosten im Format `"$cost"`, `"$5"` zurück.

# --hints--

`calculateParkingFee("09:00", "11:00")` sollte `"$6"` zurückgeben.

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` sollte `"$5"` zurückgeben.

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` sollte `"$9"` zurückgeben.

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` sollte `"$27"` zurückgeben.

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` sollte `"$34"` zurückgeben.

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` sollte `"$82"` zurückgeben.

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
