---
id: 699c8e045ee7cb94ed2322da
title: "Aufgabe 219: Jubiläumsmeilensteine"
challengeType: 28
dashedName: challenge-219
---

# --description--

Gegeben ist eine ganze Zahl, die angibt, wie viele Jahre ein Paar verheiratet ist. Geben Sie den zuletzt erreichten Jubiläumsmeilenstein gemäß diesem Diagramm zurück:

| Jahre verheiratet | Meilenstein    |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- Wenn sie den ersten Meilenstein noch nicht erreicht haben, geben Sie `"Newlyweds"` zurück.

# --hints--

`getMilestone(0)` sollte `"Newlyweds"` zurückgeben.

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` sollte `"Paper"` zurückgeben.

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` sollte `"Wood"` zurückgeben.

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` sollte `"Tin"` zurückgeben.

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` sollte `"Silver"` zurückgeben.

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` sollte `"Ruby"` zurückgeben.

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` sollte `"Gold"` zurückgeben.

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` sollte `"Diamond"` zurückgeben.

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` sollte `"Platinum"` zurückgeben.

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
