---
id: 699c8e045ee7cb94ed2322d9
title: "Aufgabe 218: Gleichmäßig teilbar"
challengeType: 28
dashedName: challenge-218
---

# --description--

Gegeben sind zwei ganze Zahlen. Bestimmen Sie, ob die erste durch die zweite ohne Rest teilbar ist.

# --hints--

`isEvenlyDivisible(4, 2)` sollte `true` zurückgeben.

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

`isEvenlyDivisible(7, 3)` sollte `false` zurückgeben.

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

`isEvenlyDivisible(5, 10)` sollte `false` zurückgeben.

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

`isEvenlyDivisible(48, 6)` sollte `true` zurückgeben.

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

`isEvenlyDivisible(3186, 9)` sollte `true` zurückgeben.

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

`isEvenlyDivisible(4192, 11)` sollte `false` zurückgeben.

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```
