---
id: 699c8e045ee7cb94ed2322d8
title: "Aufgabe 217: Erfasste Schachfiguren"
challengeType: 28
dashedName: challenge-217
---

# --description--

Gegeben ist ein Array von Strings, die die Schachfiguren darstellen, die Sie noch auf dem Brett haben. Berechnen Sie den Wert der Figuren, die Ihr Gegner erobert hat.

Im Schach starten Sie mit 16 Figuren:

| Figur  | Abkürzung     | Anzahl   | Wert  |
| ------ | --------------| -------- | ----- |
| Bauer  | `"P"`     | 8        | 1     |
| Turm   | `"R"`     | 2        | 5     |
| Springer | `"N"`   | 2        | 3     |
| Läufer | `"B"`     | 2        | 3     |
| Dame   | `"Q"`     | 1        | 9     |
| König  | `"K"`     | 1        | 0     |

- Das gegebene Array enthält nur die oben genannten Abkürzungen.
- Alle der 16 Figuren, die nicht im gegebenen Array enthalten sind, wurden erobert.
- Geben Sie den Gesamtwert aller eroberten Figuren zurück, außer...
- Wenn der König erobert wurde, geben Sie `"Checkmate"` zurück.

# --hints--

`getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` sollte `8` zurückgeben.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8);
```

`getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"])` sollte `26` zurückgeben.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"]), 26);
```

`getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` sollte `16` zurückgeben.

```js
assert.equal(getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16);
```

`getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` sollte `4` zurückgeben.

```js
assert.equal(getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4);
```

`getCapturedValue(["P", "K"])` sollte `38` zurückgeben.

```js
assert.equal(getCapturedValue(["P", "K"]), 38);
```

`getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` sollte `0` zurückgeben.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0);
```

`getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` sollte `"Checkmate"` zurückgeben.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate");
```

# --seed--

## --seed-contents--

```js
function getCapturedValue(pieces) {

  return pieces;
}
```

# --solutions--

```js
function getCapturedValue(pieces) {
  if (!pieces.includes("K")) return "Checkmate";

  const values = { P:1, R:5, N:3, B:3, Q:9, K:0 };
  let remaining = "PPPPPPPPRRNNBBQK";

  for (const p of pieces) {
    remaining = remaining.replace(p, "");
  }

  return [...remaining].reduce((sum, p) => sum + values[p], 0);
}
```
