---
id: 699c8e045ee7cb94ed2322d5
title: "Aufgabe 214: Domino-Ketten-Validator"
challengeType: 28
dashedName: challenge-214
---

# --description--

Gegeben ist ein 2D-Array, das eine Folge von Dominosteinen darstellt. Bestimmen Sie, ob es eine gültige Kette bildet.

- Jedes Element im Array stellt einen Dominostein dar und ist ein Array aus zwei Zahlen von 1 bis 6 (einschließlich).
- Damit die Kette gültig ist, muss die zweite Zahl jedes Dominosteins mit der ersten Zahl des nächsten Dominosteins übereinstimmen.
- Die erste Zahl des ersten Dominosteins und die letzte Zahl des letzten Dominosteins müssen mit nichts übereinstimmen.

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])` sollte `true` zurückgeben.

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])` sollte `false` zurückgeben.

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])` sollte `false` zurückgeben.

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` sollte `true` zurückgeben.

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` sollte `false` zurückgeben.

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```
