---
id: 69a890af247de743333bd4d2
title: "Aufgabe 229: Text kürzen 2"
challengeType: 28
dashedName: challenge-229
---

# --description--

Gegeben ist ein String. Geben Sie einen neuen String zurück, der so gekürzt ist, dass die Gesamtbreite der Zeichen 50 Einheiten nicht überschreitet.

Jedes Zeichen hat eine bestimmte Breite:

| Buchstaben | Breite |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

Die obige Tabelle enthält alle Groß- und Kleinbuchstaben. Zusätzlich:

- Leerzeichen (`" "`) haben eine Breite von 2
- Punkte (`"."`) haben eine Breite von 1

- Wenn der gegebene String 50 Einheiten oder weniger hat, geben Sie den String unverändert zurück, andernfalls
- Kürzen Sie den String und fügen Sie am Ende drei Punkte (`"..."`) hinzu, sodass die Gesamtbreite inklusive der drei Punkte so nah wie möglich an 50 Einheiten liegt, ohne diese zu überschreiten.

# --hints--

`truncateText("The quick brown fox")` sollte `"The quick brown f..."` zurückgeben.

```js
assert.equal(truncateText("The quick brown fox"), "The quick brown f...");
```

`truncateText("The silky smooth sloth")` sollte `"The silky smooth s..."` zurückgeben.

```js
assert.equal(truncateText("The silky smooth sloth"), "The silky smooth s...");
```

`truncateText("THE LOUD BRIGHT BIRD")` sollte `"THE LOUD BRIG..."` zurückgeben.

```js
assert.equal(truncateText("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...");
```

`truncateText("The fast striped zebra")` sollte `"The fast striped z..."` zurückgeben.

```js
assert.equal(truncateText("The fast striped zebra"), "The fast striped z...");
```

`truncateText("The big black bear")` sollte `"The big black bear"` zurückgeben.

```js
assert.equal(truncateText("The big black bear"), "The big black bear");
```

# --seed--

## --seed-contents--

```js
function truncateText(str) {

  return str;
}
```

# --solutions--

```js
function truncateText(str) {
  const MAX_WIDTH = 50;
  const ELLIPSIS = "...";
  const ELLIPSIS_WIDTH = 3;
  const TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH;

  const charGroups = {
    'ilI.': 1,
    'fjrt ': 2,
    'abcdeghkmnopqrstuvwxyzJL': 3,
    'ABCDEFGHKMNOPQRSTUVWXYZ': 4
  };

  function getCharWidth(char) {
    for (const key in charGroups) {
      if (key.includes(char)) return charGroups[key];
    }
    return 3;
  }

  function stringWidth(str) {
    let totalWidth = 0;
    for (const char of str) {
      totalWidth += getCharWidth(char);
    }
    return totalWidth;
  }

  if (stringWidth(str) <= MAX_WIDTH) return str;

  let result = "";
  let totalWidth = 0;

  for (const char of str) {
    const charWidth = getCharWidth(char);

    if (totalWidth + charWidth > TRUNCATE_LIMIT) break;

    result += char;
    totalWidth += charWidth;
  }

  return result + ELLIPSIS;
}
```
