---
id: 69a890af247de743333bd4cd
title: "Aufgabe 224: Kaffee-Röstungsdetektor"
challengeType: 28
dashedName: challenge-224
---

# --description--

Gegeben ist ein String, der die Bohnen darstellt, die zur Herstellung einer Tasse Kaffee verwendet werden. Bestimmen Sie die Röstung der Tasse.

- Der gegebene String enthält die folgenden Zeichen, die jeweils eine Bohnenart darstellen:

  - Ein Apostroph (`'`) ist eine hell geröstete Bohne, die jeweils 1 Punkt wert ist.
  - Ein Bindestrich (`-`) ist eine mittel geröstete Bohne, die jeweils 2 Punkte wert ist.
  - Ein Punkt (`.`) ist eine dunkel geröstete Bohne, die jeweils 3 Punkte wert ist.

- Der Röstgrad wird durch den Durchschnitt aller Bohnen bestimmt.

Geben Sie zurück:

- `"Light"`, wenn der Durchschnitt kleiner als 1,75 ist.
- `"Medium"`, wenn der Durchschnitt zwischen 1,75 und 2,5 liegt.
- `"Dark"`, wenn der Durchschnitt größer als 2,5 ist.

# --hints--

`detectRoast("''-''''''-'-''--''''")` sollte `"Light"` zurückgeben.

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")` sollte `"Medium"` zurückgeben.

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")` sollte `"Medium"` zurückgeben.

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")` sollte `"Dark"` zurückgeben.

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")` sollte `"Medium"` zurückgeben.

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")` sollte `"Dark"` zurückgeben.

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")` sollte `"Light"` zurückgeben.

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```
