---
id: 69a890af247de743333bd4d0
title: "Aufgabe 227: Abkühlzeit"
challengeType: 28
dashedName: challenge-227
---

# --description--

Gegeben sind zwei Zeitstempel, der erste stellt dar, wann ein Benutzer eine Prüfung beendet hat, und der zweite stellt die aktuelle Zeit dar. Bestimmen Sie, ob der Benutzer die Prüfung erneut ablegen kann.

- Beide Zeitstempel werden im Format `"YYYY-MM-DDTHH:MM:SS"` angegeben, zum Beispiel `"2026-03-25T14:00:00"`. Beachten Sie, dass die Zeit im 24-Stunden-Format angegeben ist.
- Ein Benutzer muss mindestens 48 Stunden warten, bevor er eine Prüfung erneut ablegen darf.

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` sollte `true` zurückgeben.

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` sollte `false` zurückgeben.

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` sollte `true` zurückgeben.

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` sollte `false` zurückgeben.

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
