---
id: 69a890af247de743333bd4cc
title: "Aufgabe 223: QR-Decoder"
challengeType: 28
dashedName: challenge-223
---

# --description--

Gegeben ist eine 6x6-Matrix (Array von Strings), die einen QR-Code darstellt. Geben Sie den String der Binärdaten im Code zurück.

- Der QR-Code kann in jeder 90-Grad-Drehung vorliegen.
- Ein korrekt ausgerichteter Code hat eine 2x2-Gruppe von `1` (Orientierungsmarkierungen) in den Ecken unten links, oben links und oben rechts.
- Die drei 2x2-Orientierungsmarkierungen sind nicht Teil der Binärdaten.
- Die Binärdaten werden von links nach rechts, von oben nach unten gelesen (wie in einem Buch), wenn der QR-Code korrekt ausgerichtet ist.
- Ein Code hat immer genau eine gültige Ausrichtung.

Zum Beispiel, gegeben:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

oder derselbe Code mit einer anderen Ausrichtung:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

Geben Sie `"000000000000000000000001"` zurück, alle Binärdaten ohne die drei 2x2-Orientierungsmarkierungen.

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])` sollte `"000000000000000000000001"` zurückgeben.

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"])` sollte `"000000000000000000000001"` zurückgeben.

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])` sollte `"001101000011000000110100"` zurückgeben.

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])` sollte `"010001000100010101010110"` zurückgeben.

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])` sollte `"010000100100100101001110"` zurückgeben.

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```
