---
id: 69a890af247de743333bd4ce
title: "Aufgabe 225: Keine aufeinanderfolgenden Wiederholungen"
challengeType: 28
dashedName: challenge-225
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob er keine sich wiederholenden Zeichen enthält.

- Ein String hat keine Wiederholungen, wenn kein Zeichen zwei- oder mehrmals hintereinander vorkommt.

# --hints--

`hasNoRepeats("hi world")` sollte `true` zurückgeben.

```js
assert.isTrue(hasNoRepeats("hi world"));
```

`hasNoRepeats("hello world")` sollte `false` zurückgeben.

```js
assert.isFalse(hasNoRepeats("hello world"));
```

`hasNoRepeats("abcdefghijklmnopqrstuvwxyz")` sollte `true` zurückgeben.

```js
assert.isTrue(hasNoRepeats("abcdefghijklmnopqrstuvwxyz"));
```

`hasNoRepeats("freeCodeCamp")` sollte `false` zurückgeben.

```js
assert.isFalse(hasNoRepeats("freeCodeCamp"));
```

`hasNoRepeats("The quick brown fox jumped over the lazy dog.")` sollte `true` zurückgeben.

```js
assert.isTrue(hasNoRepeats("The quick brown fox jumped over the lazy dog."));
```

`hasNoRepeats("Mississippi")` sollte `false` zurückgeben.

```js
assert.isFalse(hasNoRepeats("Mississippi"));
```

# --seed--

## --seed-contents--

```js
function hasNoRepeats(str) {

  return str;
}
```

# --solutions--

```js
function hasNoRepeats(str) {
  for (let i = 1; i < str.length; i++) {
    if (str[i] === str[i - 1]) {
      return false;
    }
  }
  return true;
}
```
