---
id: 69a890af247de743333bd4cf
title: "Aufgabe 226: Anzahl der bestandenen Prüfungen"
challengeType: 28
dashedName: challenge-226
---

# --description--

Gegeben ist ein Array mit Prüfungsergebnissen von Studierenden und die Punktzahl, die zum Bestehen erforderlich ist. Geben Sie die Anzahl der Studierenden zurück, die die Prüfung bestanden haben.

# --hints--

`passingCount([90, 85, 75, 60, 50], 70)` sollte `3` zurückgeben.

```js
assert.equal(passingCount([90, 85, 75, 60, 50], 70), 3);
```

`passingCount([100, 80, 75, 88, 72, 74, 79, 71, 60, 92], 75)` sollte `6` zurückgeben.

```js
assert.equal(passingCount([100, 80, 75, 88, 72, 74, 79, 71, 60, 92], 75), 6);
```

`passingCount([79, 60, 88, 72, 74, 59, 75, 71, 80, 92], 60)` sollte `9` zurückgeben.

```js
assert.equal(passingCount([79, 60, 88, 72, 74, 59, 75, 71, 80, 92], 60), 9);
```

`passingCount([76, 79, 80, 70, 71, 65, 79, 78, 59, 72], 85)` sollte `0` zurückgeben.

```js
assert.equal(passingCount([76, 79, 80, 70, 71, 65, 79, 78, 59, 72], 85), 0);
```

`passingCount([84, 65, 98, 53, 58, 71, 91, 80, 92, 70, 73, 83, 86, 69, 84, 77, 72, 58, 69, 75, 66, 68, 72, 96, 90, 63, 88, 63, 80, 67], 60)` sollte `27` zurückgeben.

```js
assert.equal(passingCount([84, 65, 98, 53, 58, 71, 91, 80, 92, 70, 73, 83, 86, 69, 84, 77, 72, 58, 69, 75, 66, 68, 72, 96, 90, 63, 88, 63, 80, 67], 60), 27);
```

# --seed--

## --seed-contents--

```js
function passingCount(scores, passingScore) {

  return scores;
}
```

# --solutions--

```js
function passingCount(scores, passingScore) {
  let count = 0;

  for (const score of scores) {
    if (score >= passingScore) {
      count++;
    }
  }

  return count;
}
```
