---
id: 69a890af247de743333bd4d1
title: "Aufgabe 228: Filmabend"
challengeType: 28
dashedName: challenge-228
---

# --description--

Gegeben sind ein String für den Wochentag, ein weiterer String für die Vorstellungszeit und eine ganze Zahl für die Anzahl der Tickets. Geben Sie die Gesamtkosten der Kinokarten für diese Vorstellung zurück.

Der angegebene Tag wird einer der folgenden sein:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

Die Vorstellungszeit wird im Format `"H:MMam"` oder `"H:MMpm"` angegeben. Zum Beispiel `"10:00am"` oder `"10:00pm"`.

Geben Sie die Gesamtkosten im Format `"$D.CC"` zurück, und zwar nach diesen Regeln:

- Wochenende (Freitag bis Sonntag): 12,00 $ pro Ticket.
- Wochentag (Montag bis Donnerstag): 10,00 $ pro Ticket.
- Matinee (vor 17:00 Uhr): 2,00 $ Rabatt pro Ticket (außer dienstags).
- Dienstags: alle Tickets kosten jeweils 5,00 $.

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` sollte `"$12.00"` zurückgeben.

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` sollte `"$10.00"` zurückgeben.

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` sollte `"$10.00"` zurückgeben.

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` sollte `"$30.00"` zurückgeben.

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` sollte `"$32.00"` zurückgeben.

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` sollte `"$50.00"` zurückgeben.

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` sollte `"$5.00"` zurückgeben.

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
