---
id: 69b1028d6e265413d0198a2c
title: "Aufgabe 233: Weckalarm"
challengeType: 28
dashedName: challenge-233
---

# --description--

Gegeben ist ein String, der die Zeit darstellt, zu der Sie Ihren Wecker gestellt haben, und ein String, der die Zeit darstellt, zu der Sie tatsächlich aufgewacht sind. Bestimmen Sie, ob Sie früh, pünktlich oder zu spät aufgewacht sind.

- Beide Zeiten werden im `"HH:MM"` 24-Stunden-Format angegeben.

Geben Sie zurück:

- `"early"`, wenn Sie vor der Weckzeit aufgewacht sind.
- `"on time"`, wenn Sie zur Weckzeit oder innerhalb des 10-minütigen Schlummerfensters nach der Weckzeit aufgewacht sind.
- `"late"`, wenn Sie mehr als 10 Minuten nach der Weckzeit aufgewacht sind.

Beide Zeiten liegen am selben Tag.

# --hints--

`alarmCheck("07:00", "06:45")` sollte `"early"` zurückgeben.

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")` sollte `"on time"` zurückgeben.

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")` sollte `"on time"` zurückgeben.

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")` sollte `"late"` zurückgeben.

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")` sollte `"on time"` zurückgeben.

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")` sollte `"late"` zurückgeben.

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")` sollte `"early"` zurückgeben.

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```
