---
id: 69b1028d6e265413d0198a2f
title: "Aufgabe 236: Browser-Verlauf"
challengeType: 28
dashedName: challenge-236
---

# --description--

Gegeben ist ein Array von Browser-Befehlen. Geben Sie ein Array mit zwei Werten zurück: den Verlauf als Array von URLs und den Index der aktuellen Seite.

Gültige Befehle sind:

- `"URL"` – wobei URL eine Webadresse ist (zum Beispiel `"freecodecamp.org"`). Navigiert zur angegebenen URL, fügt sie an der nächsten Position zum Verlauf hinzu und verwirft alle Vorwärtsverläufe.
- `"Back"` – bewegt sich zur vorherigen Seite im Verlauf oder bleibt auf der aktuellen Seite, falls keine vorhanden ist.
- `"Forward"` – bewegt sich zur nächsten Seite im Verlauf oder bleibt auf der aktuellen Seite, falls keine vorhanden ist.

Zum Beispiel, bei Eingabe von `["freecodecamp.org", "freecodecamp.org/learn", "Back"]` geben Sie `[["freecodecamp.org", "freecodecamp.org/learn"], 0]` zurück.

# --hints--

`getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` sollte `[["freecodecamp.org", "freecodecamp.org/learn"], 0]` zurückgeben.

```js
assert.deepEqual(getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` sollte `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]` zurückgeben.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact",  "example.com/blog", "Back", "Back", "Forward"])` sollte `[["example.com", "example.com/contact", "example.com/blog"], 1]` zurückgeben.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` sollte `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]` zurückgeben.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "Back"])` sollte `[["example.com", "example.com/about"], 0]` zurückgeben.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "Forward"])` sollte `[["example.com", "example.com/about"], 1]` zurückgeben.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1]);
```

# --seed--

## --seed-contents--

```js
function getBrowserHistory(commands) {

  return commands;
}
```

# --solutions--

```js
function getBrowserHistory(commands) {
  const stack = [];
  let index = -1;

  for (const command of commands) {
    if (command === "Back") {
      if (index > 0) index--;
    } else if (command === "Forward") {
      if (index < stack.length - 1) index++;
    } else {
      stack.splice(index + 1);
      stack.push(command);
      index++;
    }
  }

  return [stack, index];
}
```
