---
id: 69b1028d6e265413d0198a2d
title: "Aufgabe 234: Streichzahl"
challengeType: 28
dashedName: challenge-234
---

# --description--

Gegeben ist ein Array von Zahlen, bei dem alle bis auf eine Zahl einem Muster folgen. Geben Sie ein neues Array zurück, in dem die eine Zahl, die nicht dem Muster folgt, korrigiert wurde.

Das Muster ist eines der folgenden:

- Die Zahlen steigen von einer zur nächsten um einen festen Betrag (Addition).
- Die Zahlen fallen von einer zur nächsten um einen festen Betrag (Subtraktion).

Zum Beispiel, gegeben `[2, 4, 7, 8, 10]`, geben Sie `[2, 4, 6, 8, 10]` zurück.

# --hints--

`fixPrankNumber([2, 4, 7, 8, 10])` sollte `[2, 4, 6, 8, 10]` zurückgeben.

```js
assert.deepEqual(fixPrankNumber([2, 4, 7, 8, 10]), [2, 4, 6, 8, 10]);
```

`fixPrankNumber([10, 10, 8, 7, 6])` sollte `[10, 9, 8, 7, 6]` zurückgeben.

```js
assert.deepEqual(fixPrankNumber([10, 10, 8, 7, 6]), [10, 9, 8, 7, 6]);
```

`fixPrankNumber([12, 24, 36, 48, 61, 72, 84, 96])` sollte `[12, 24, 36, 48, 60, 72, 84, 96]` zurückgeben.

```js
assert.deepEqual(fixPrankNumber([12, 24, 36, 48, 61, 72, 84, 96]), [12, 24, 36, 48, 60, 72, 84, 96]);
```

`fixPrankNumber([4, 1, -2, -5, -8, -5])` sollte `[4, 1, -2, -5, -8, -11]` zurückgeben.

```js
assert.deepEqual(fixPrankNumber([4, 1, -2, -5, -8, -5]), [4, 1, -2, -5, -8, -11]);
```

`fixPrankNumber([0, 100, 200, 300, 150, 500])` sollte `[0, 100, 200, 300, 400, 500]` zurückgeben.

```js
assert.deepEqual(fixPrankNumber([0, 100, 200, 300, 150, 500]), [0, 100, 200, 300, 400, 500]);
```

`fixPrankNumber([400, 425, 400, 375, 350, 325, 300])` sollte `[450, 425, 400, 375, 350, 325, 300]` zurückgeben.

```js
assert.deepEqual(fixPrankNumber([400, 425, 400, 375, 350, 325, 300]), [450, 425, 400, 375, 350, 325, 300]);
```

`fixPrankNumber([-5, 5, 10, 15, 20])` sollte `[0, 5, 10, 15, 20]` zurückgeben.

```js
assert.deepEqual(fixPrankNumber([-5, 5, 10, 15, 20]), [0, 5, 10, 15, 20]);
```

# --seed--

## --seed-contents--

```js
function fixPrankNumber(arr) {

  return arr;
}
```

# --solutions--

```js
function fixPrankNumber(arr) {
  const diffs = arr.slice(1).map((v, i) => v - arr[i]);

  let step = diffs[0];
  for (let i = 0; i < diffs.length - 1; i++) {
    if (diffs[i] === diffs[i + 1]) {
      step = diffs[i];
      break;
    }
  }

  const result = [...arr];

  for (let i = 0; i < result.length - 1; i++) {
    if (result[i + 1] - result[i] !== step) {
      if (i === 0 && result[2] - result[1] === step) {
        result[0] = result[1] - step;
      } else {
        result[i + 1] = result[i] + step;
      }
      break;
    }
  }

  return result;
}
```
