---
id: 69b1028d6e265413d0198a2a
title: "Aufgabe 231: ISBN-10-Validator"
challengeType: 28
dashedName: challenge-231
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob es sich um eine gültige ISBN-10 handelt.

Eine ISBN-10 besteht aus Bindestrichen (`"-"`) und 10 weiteren Zeichen. Nach dem Entfernen der Bindestriche (`"-"`):

- Die ersten 9 Zeichen müssen Ziffern sein, und
- Das letzte Zeichen kann eine Ziffer oder der Buchstabe `"X"` sein, der die Zahl 10 darstellt.

Zur Validierung:

- Multiplizieren Sie jede Ziffer (oder den Wert) mit ihrer Position (multiplizieren Sie die erste Ziffer mit 1, die zweite mit 2 usw.).
- Addieren Sie alle Ergebnisse zusammen.
- Wenn die Summe durch 11 teilbar ist, ist die ISBN gültig.

# --hints--

`isValidIsbn10("0-306-40615-2")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
