---
id: 69b1028d6e265413d0198a29
title: "Aufgabe 230: Pascal-Dreieck-Zeile"
challengeType: 28
dashedName: challenge-230
---

# --description--

Gegeben ist eine ganze Zahl `n`. Geben Sie die `n`te Zeile des Pascal-Dreiecks als Array zurück.

Im Pascal-Dreieck beginnt und endet jede Zeile mit 1, und jeder innere Wert ist die Summe der beiden Werte direkt darüber.

Hier sind die ersten 5 Zeilen des Dreiecks:

```js
    1
   1 1
  1 2 1
 1 3 3 1
1 4 6 4 1
```

# --hints--

`pascalRow(5)` sollte `[1, 4, 6, 4, 1]` zurückgeben.

```js
assert.deepEqual(pascalRow(5), [1, 4, 6, 4, 1]);
```

`pascalRow(3)` sollte `[1, 2, 1]` zurückgeben.

```js
assert.deepEqual(pascalRow(3), [1, 2, 1]);
```

`pascalRow(1)` sollte `[1]` zurückgeben.

```js
assert.deepEqual(pascalRow(1), [1]);
```

`pascalRow(10)` sollte `[1, 9, 36, 84, 126, 126, 84, 36, 9, 1]` zurückgeben.

```js
assert.deepEqual(pascalRow(10), [1, 9, 36, 84, 126, 126, 84, 36, 9, 1]);
```

`pascalRow(15)` sollte `[1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]` zurückgeben.

```js
assert.deepEqual(pascalRow(15), [1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]);
```

# --seed--

## --seed-contents--

```js
function pascalRow(n) {

  return n;
}
```

# --solutions--

```js
function pascalRow(n) {
  if (n === 1) return [1];

  let row = [1];

  for (let i = 2; i <= n; i++) {
    const next = [1];

    for (let j = 1; j < row.length; j++) {
      next.push(row[j - 1] + row[j]);
    }

    next.push(1);
    row = next;
  }

  return row;
}
```
