---
id: 69b1028d6e265413d0198a2e
title: "Aufgabe 235: Großgeschriebene Fibonacci"
challengeType: 28
dashedName: challenge-235
---

# --description--

Gegeben ist ein String. Geben Sie einen neuen String zurück, bei dem jeder Buchstabe großgeschrieben wird, wenn sein Index eine Fibonacci-Zahl ist, und ansonsten in Kleinbuchstaben.

Die Fibonacci-Folge ist eine Zahlenreihe, bei der jede Zahl die Summe der beiden vorhergehenden ist. Die ersten 10 Zahlen der Folge sind `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

- Das erste Zeichen hat den Index `0`.
- Wenn der Index von Nicht-Buchstaben-Zeichen eine Fibonacci-Zahl ist, lassen Sie dieses Zeichen unverändert.

# --hints--

`capitalizeFibonacci("hello world")` sollte `"HELLo woRld"` zurückgeben.

```js
assert.equal(capitalizeFibonacci("hello world"), "HELLo woRld");
```

`capitalizeFibonacci("HELLO WORLD")` sollte `"HELLo woRld"` zurückgeben.

```js
assert.equal(capitalizeFibonacci("HELLO WORLD"), "HELLo woRld");
```

`capitalizeFibonacci("hello, world!")` sollte `"HELLo, wOrld!"` zurückgeben.

```js
assert.equal(capitalizeFibonacci("hello, world!"), "HELLo, wOrld!");
```

`capitalizeFibonacci("The quick brown fox jumped over the lazy dog.")` sollte `"THE qUicK broWn fox jUmped over thE lazy dog."` zurückgeben.

```js
assert.equal(capitalizeFibonacci("The quick brown fox jumped over the lazy dog."), "THE qUicK broWn fox jUmped over thE lazy dog.");
```

`capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl.")` sollte `"LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl."` zurückgeben.

```js
assert.equal(capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl."), "LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl.");
```

# --seed--

## --seed-contents--

```js
function capitalizeFibonacci(str) {

  return str;
}
```

# --solutions--

```js
function capitalizeFibonacci(str) {
  const fibs = new Set([0, 1]);
  let a = 0, b = 1;
  while (b < str.length) {
    [a, b] = [b, a + b];
    fibs.add(b);
  }

  return str
    .split("")
    .map((char, i) => {
      if (!/[a-zA-Z]/.test(char)) return char;
      return fibs.has(i) ? char.toUpperCase() : char.toLowerCase();
    })
    .join("");
}
```
