---
id: 69b5b2be76ec8135a7fbe973
title: "Aufgabe 244: Angriff von Turm und Läufer"
challengeType: 28
dashedName: challenge-244
---

# --description--

Gegeben ist ein String für die Position eines Turms auf einem Schachbrett und ein weiterer für die Position eines Läufers. Bestimmen Sie, ob eine Figur die andere angreifen kann.

Ein Standard-Schachbrett ist 8x8, mit Spalten, die von `A` bis `H` beschriftet sind (von links nach rechts), und Reihen, die von `1` bis `8` beschriftet sind (von unten nach oben). Es sieht so aus:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- Türme können sich beliebig viele Felder horizontal oder vertikal bewegen.
- Läufer können sich beliebig viele Felder diagonal bewegen.
- Eine Figur kann eine andere angreifen, wenn sie sich auf die Position dieser Figur bewegen kann.

Geben Sie zurück:

- `"rook"`, wenn der Turm den Läufer angreifen kann.
- `"bishop"`, wenn der Läufer den Turm angreifen kann.
- `"neither"`, wenn keine der beiden Figuren die andere angreifen kann.

# --hints--

`rookBishopAttack("A1", "A5")` sollte `"rook"` zurückgeben.

```js
assert.equal(rookBishopAttack("A1", "A5"), "rook");
```

`rookBishopAttack("C3", "F6")` sollte `"bishop"` zurückgeben.

```js
assert.equal(rookBishopAttack("C3", "F6"), "bishop");
```

`rookBishopAttack("D4", "D7")` sollte `"rook"` zurückgeben.

```js
assert.equal(rookBishopAttack("D4", "D7"), "rook");
```

`rookBishopAttack("B7", "H1")` sollte `"bishop"` zurückgeben.

```js
assert.equal(rookBishopAttack("B7", "H1"), "bishop");
```

`rookBishopAttack("B3", "C5")` sollte `"neither"` zurückgeben.

```js
assert.equal(rookBishopAttack("B3", "C5"), "neither");
```

`rookBishopAttack("G3", "E8")` sollte `"neither"` zurückgeben.

```js
assert.equal(rookBishopAttack("G3", "E8"), "neither");
```

# --seed--

## --seed-contents--

```js
function rookBishopAttack(rook, bishop) {

  return rook;
}
```

# --solutions--

```js
function rookBishopAttack(rook, bishop) {
  const colDiff = Math.abs(rook.charCodeAt(0) - bishop.charCodeAt(0));
  const rowDiff = Math.abs(rook[1] - bishop[1]);

  if (rook[0] === bishop[0] || rook[1] === bishop[1]) return "rook";
  if (colDiff === rowDiff) return "bishop";
  return "neither";
}
```
