---
id: 69b559d2903b9e4afe9075f8
title: "Aufgabe 238: Ziffernrotation entkommen"
challengeType: 28
dashedName: challenge-238
---

# --description--

Gegeben ist eine positive ganze Zahl. Bestimmen Sie, ob sie selbst oder eine ihrer Rotationen durch die Anzahl ihrer Ziffern ohne Rest teilbar ist.

Eine Rotation bedeutet, die erste Ziffer an das Ende zu verschieben. Zum Beispiel wird aus 123 nach einer Rotation 231.

- Überprüfen Sie zuerst die Rotation `0` (die gegebene Zahl).
- Gegebene Zahlen enthalten keine Nullen.
- Geben Sie die erste Rotationszahl zurück, die diese Bedingung erfüllt, oder `"none"`, wenn keine gefunden wird.

# --hints--

`getRotation(123)` sollte `0` zurückgeben.

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)` sollte `3` zurückgeben.

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)` sollte `"none"` zurückgeben.

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)` sollte `6` zurückgeben.

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
