---
id: 69b58ce40693f140c84c8559
title: "Aufgabe 240: Palindrom-Zeichen"
challengeType: 28
dashedName: challenge-240
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob er ein Palindrom ist, und geben Sie das mittlere Zeichen zurück (wenn die Länge ungerade ist) oder die mittleren zwei Zeichen (wenn die Länge gerade ist).

- Ein Palindrom ist ein String, der vorwärts und rückwärts gleich ist.
- Wenn es kein Palindrom ist, geben Sie `"none"` zurück.

# --hints--

`palindromeLocator("racecar")` sollte `"e"` zurückgeben.

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` sollte `"v"` zurückgeben.

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` sollte `"none"` zurückgeben.

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` sollte `"oo"` zurückgeben.

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` sollte `"00"` zurückgeben.

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
