---
id: 69b58ce40693f140c84c855a
title: "Aufgabe 241: FizzBuzz-Validator"
challengeType: 28
dashedName: challenge-241
---

# --description--

Gegeben ist ein Array mit aufeinanderfolgenden ganzen Zahlen, bei dem Vielfache von 3 und 5 ersetzt wurden. Bestimmen Sie, ob es sich um eine gültige FizzBuzz-Folge handelt.

In einer gültigen FizzBuzz-Folge:

- Vielfache von 3 werden durch „Fizz“ ersetzt.
- Vielfache von 5 werden durch „Buzz“ ersetzt.
- Vielfache von sowohl 3 als auch 5 werden durch „FizzBuzz“ ersetzt.
- Alle anderen Zahlen bleiben ganze Zahlen.

# --hints--

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz"])` sollte `true` zurückgeben.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4, "Buzz"]));
```

`isFizzBuzz([13, 14, "FizzBuzz", 16, 17])` sollte `true` zurückgeben.

```js
assert.isTrue(isFizzBuzz([13, 14, "FizzBuzz", 16, 17]));
```

`isFizzBuzz([1, 2, "Fizz", 4, 5])` sollte `false` zurückgeben.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, 5]));
```

`isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` sollte `true` zurückgeben.

```js
assert.isTrue(isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]));
```

`isFizzBuzz([1, 2, "Fizz", "Buzz", 5])` sollte `false` zurückgeben.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", "Buzz", 5]));
```

`isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` sollte `false` zurückgeben.

```js
assert.isFalse(isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]));
```

`isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` sollte `true` zurückgeben.

```js
assert.isTrue(isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]));
```

# --seed--

## --seed-contents--

```js
function isFizzBuzz(arr) {

  return arr;
}
```

# --solutions--

```js
function isFizzBuzz(arr) {
  let start = null;
  for (let i = 0; i < arr.length; i++) {
    if (typeof arr[i] === "number") {
      start = arr[i] - i;
      break;
    }
  }

  if (start === null) return false;

  for (let i = 0; i < arr.length; i++) {
    const n = start + i;
    const expected = n % 15 === 0 ? "FizzBuzz" : n % 3 === 0 ? "Fizz" : n % 5 === 0 ? "Buzz" : n;
    if (arr[i] !== expected) return false;
  }

  return true;
}
```
