---
id: 69b58ce40693f140c84c855b
title: "Aufgabe 242: Nächste Bingo-Nummer"
challengeType: 28
dashedName: challenge-242
---

# --description--

Gegeben ist eine Bingo-Nummer. Geben Sie die nächste Bingo-Nummer in der Reihenfolge zurück.

Eine Bingo-Nummer besteht aus einem einzelnen Buchstaben, gefolgt von einer Zahl in ihrem Bereich gemäß diesem Diagramm:

| Buchstabe | Zahlenbereich |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

Wenn zum Beispiel `"B10"` gegeben ist, geben Sie `"B11"` zurück, die nächste Bingo-Nummer. Wenn die letzte Bingo-Nummer gegeben ist, geben Sie `"B1"` zurück.

# --hints--

`getNextBingoNumber("B10")` sollte `"B11"` zurückgeben.

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")` sollte `"N34"` zurückgeben.

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")` sollte `"N31"` zurückgeben.

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")` sollte `"O61"` zurückgeben.

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")` sollte `"B1"` zurückgeben.

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```
