---
id: 69b559d2903b9e4afe9075f9
title: "Aufgabe 239: Welcher Tag ist heute?"
challengeType: 28
dashedName: challenge-239
---

# --description--

Gegeben ist ein Unix-Zeitstempel in Millisekunden. Geben Sie den Wochentag zurück.

Gültige Rückgabewerte sind:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Beachten Sie, dass Zeitzonen ignoriert werden müssen.

# --hints--

`getDayOfWeek(1775492249000)` sollte `"Monday"` zurückgeben.

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)` sollte `"Saturday"` zurückgeben.

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)` sollte `"Tuesday"` zurückgeben.

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)` sollte `"Sunday"` zurückgeben.

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)` sollte `"Thursday"` zurückgeben.

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```
