---
id: 69bc6cb30c1d112a2e110a04
title: "Aufgabe 247: Letzter Buchstabe"
challengeType: 28
dashedName: challenge-247
---

# --description--

Gegeben ist ein String. Geben Sie den Buchstaben aus dem String zurück, der im Alphabet zuletzt erscheint.

- Wenn zwei oder mehr Buchstaben im Alphabet gleich weit hinten stehen, geben Sie den ersten davon zurück.
- Ignorieren Sie alle Nicht-Buchstaben-Zeichen.

# --hints--

`getLastLetter("world")` sollte `"w"` zurückgeben.

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` sollte `"W"` zurückgeben.

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` sollte `"z"` zurückgeben.

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` sollte `"L"` zurückgeben.

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` sollte `"T"` zurückgeben.

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
