---
id: 69bc6cb30c1d112a2e110a08
title: "Aufgabe 251: Array-Summenfinder"
challengeType: 28
dashedName: challenge-251
---

# --description--

Gegeben ist ein Array von Zahlen und eine Zielzahl. Geben Sie die erste Teilmenge von zwei oder mehr Zahlen zurück, deren Summe der Zielzahl entspricht.

- Die „erste“ Teilmenge ist diejenige, deren Elemente die niedrigstmöglichen Indizes haben, wobei der früheste Index Vorrang hat.
- Jede Zahl im Array darf nur einmal verwendet werden.
- Wenn keine gültige Teilmenge existiert, geben Sie `"Sum not found"` zurück.

Geben Sie die passenden Zahlen als Array in der Reihenfolge zurück, in der sie im ursprünglichen Array erscheinen.

# --hints--

`findSum([1, 3, 5, 7], 6)` sollte `[1, 5]` zurückgeben.

```js
assert.deepEqual(findSum([1, 3, 5, 7], 6), [1, 5]);
```

`findSum([1, 2, 3, 4, 5], 5)` sollte `[1, 4]` zurückgeben.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 5), [1, 4]);
```

`findSum([1, 2, 3, 4, 5], 6)` sollte `[1, 2, 3]` zurückgeben.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 6), [1, 2, 3]);
```

`findSum([-1, -2, 3, 4], 1)` sollte `[-1, -2, 4]` zurückgeben.

```js
assert.deepEqual(findSum([-1, -2, 3, 4], 1), [-1, -2, 4]);
```

`findSum([3, 1, 4, 1, 5, 9, 2, 6], 10)` sollte `[3, 1, 4, 2]` zurückgeben.

```js
assert.deepEqual(findSum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2]);
```

`findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` sollte `[1, 2, 3, 5, 9]` zurückgeben.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9]);
```

`findSum([7, 9, 4, 2, 5], 10)` sollte `"Sum not found"` zurückgeben.

```js
assert.equal(findSum([7, 9, 4, 2, 5], 10), "Sum not found");
```

# --seed--

## --seed-contents--

```js
function findSum(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findSum(arr, target) {
  function search(start, remaining, current) {
    if (remaining === 0 && current.length >= 2) return current;
    if (start >= arr.length) return null;

    for (let i = start; i < arr.length; i++) {
      const result = search(i + 1, remaining - arr[i], [...current, arr[i]]);
      if (result) return result;
    }

    return null;
  }

  return search(0, target, []) ?? "Sum not found";
}
```
