---
id: 69bc6cb30c1d112a2e110a06
title: "Aufgabe 249: String-Mathematik"
challengeType: 28
dashedName: challenge-249
---

# --description--

Gegeben ist ein String mit Zahlen und anderen Zeichen. Führen Sie mathematische Operationen mit den Zahlen durch, basierend auf der Anzahl der Nicht-Ziffern-Zeichen zwischen den Zahlen.

- Wenn die Anzahl der Zeichen zwischen zwei Zahlen gerade ist, verwenden Sie Addition.
- Wenn sie ungerade ist, verwenden Sie Subtraktion.
- Aufeinanderfolgende Ziffern bilden eine einzelne Zahl.
- Die Operationen werden von links nach rechts angewendet.
- Ignorieren Sie führende und abschließende Zeichen, die keine Ziffern sind.

Zum Beispiel, gegeben `"3ab10c8"`, geben Sie `5` zurück. Addieren Sie 3 und 10, um 13 zu erhalten, weil eine gerade Anzahl von Zeichen zwischen ihnen liegt. Ziehen Sie dann 8 von 13 ab, weil eine ungerade Anzahl von Zeichen zwischen dem Ergebnis und 8 liegt.

# --hints--

`doMath("3ab10c8")` sollte `5` zurückgeben.

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` sollte `2` zurückgeben.

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` sollte `12` zurückgeben.

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` sollte `15` zurückgeben.

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` sollte `67` zurückgeben.

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
