---
id: 69bc6cb30c1d112a2e110a09
title: "Aufgabe 252: Einzigartiger Treppensteiger"
challengeType: 28
dashedName: challenge-252
---

# --description--

Gegeben ist eine Anzahl von Treppenstufen. Geben Sie zurück, auf wie viele verschiedene Arten jemand diese erklimmen kann, wenn er jeweils entweder 1 oder 2 Stufen auf einmal nimmt.

# --hints--

`getUniqueClimbs(4)` sollte `5` zurückgeben.

```js
assert.equal(getUniqueClimbs(4), 5);
```

`getUniqueClimbs(5)` sollte `8` zurückgeben.

```js
assert.equal(getUniqueClimbs(5), 8);
```

`getUniqueClimbs(10)` sollte `89` zurückgeben.

```js
assert.equal(getUniqueClimbs(10), 89);
```

`getUniqueClimbs(18)` sollte `4181` zurückgeben.

```js
assert.equal(getUniqueClimbs(18), 4181);
```

`getUniqueClimbs(29)` sollte `832040` zurückgeben.

```js
assert.equal(getUniqueClimbs(29), 832040);
```

`getUniqueClimbs(50)` sollte `20365011074` zurückgeben.

```js
assert.equal(getUniqueClimbs(50), 20365011074);
```

# --seed--

## --seed-contents--

```js
function getUniqueClimbs(steps) {

  return steps;
}
```

# --solutions--

```js
function getUniqueClimbs(steps) {
  if (steps <= 0) return 0;
  if (steps === 1) return 1;
  if (steps === 2) return 2;

  let prev2 = 1, prev1 = 2;
  for (let i = 3; i <= steps; i++) {
    [prev2, prev1] = [prev1, prev2 + prev1];
  }
  return prev1;
}
```
