---
id: 69bc6cb30c1d112a2e110a05
title: "Aufgabe 248: Sortiertes Array mit Tausch"
challengeType: 28
dashedName: challenge-248
---

# --description--

Gegeben ist ein Array aus ganzen Zahlen. Geben Sie ein neues Array zurück, das folgende Regeln erfüllt:

1. Sortieren Sie die ganzen Zahlen in aufsteigender Reihenfolge.
2. Tauschen Sie dann alle Werte, deren Index ein Vielfaches von 3 ist, mit dem Wert davor.

# --hints--

`sortAndSwap([3, 1, 2, 4, 6, 5])` sollte `[1, 2, 4, 3, 5, 6]` zurückgeben.

```js
assert.deepEqual(sortAndSwap([3, 1, 2, 4, 6, 5]), [1, 2, 4, 3, 5, 6]);
```

`sortAndSwap([9, 7, 5, 3, 1, 2, 4, 6, 8])` sollte `[1, 2, 4, 3, 5, 7, 6, 8, 9]` zurückgeben.

```js
assert.deepEqual(sortAndSwap([9, 7, 5, 3, 1, 2, 4, 6, 8]), [1, 2, 4, 3, 5, 7, 6, 8, 9]);
```

`sortAndSwap([1, 2, 3, 4, 5, 6, 7, 8, 9])` sollte `[1, 2, 4, 3, 5, 7, 6, 8, 9]` zurückgeben.

```js
assert.deepEqual(sortAndSwap([1, 2, 3, 4, 5, 6, 7, 8, 9]), [1, 2, 4, 3, 5, 7, 6, 8, 9]);
```

`sortAndSwap([12, 5, 8, 1, 3, 10, 2, 7, 6, 4, 9, 11])` sollte `[1, 2, 4, 3, 5, 7, 6, 8, 10, 9, 11, 12]` zurückgeben.

```js
assert.deepEqual(sortAndSwap([12, 5, 8, 1, 3, 10, 2, 7, 6, 4, 9, 11]), [1, 2, 4, 3, 5, 7, 6, 8, 10, 9, 11, 12]);
```

`sortAndSwap([100, -50, 0, 75, -25, 50, -75, 25])` sollte `[-75, -50, 0, -25, 25, 75, 50, 100]` zurückgeben.

```js
assert.deepEqual(sortAndSwap([100, -50, 0, 75, -25, 50, -75, 25]), [-75, -50, 0, -25, 25, 75, 50, 100]);
```

`sortAndSwap([5, 9, 13, 77, 88, 313, -10, -65, 0, 8, 99, 101, -4, 2])` sollte `[-65, -10, 0, -4, 2, 8, 5, 9, 77, 13, 88, 101, 99, 313]` zurückgeben.

```js
assert.deepEqual(sortAndSwap([5, 9, 13, 77, 88, 313, -10, -65, 0, 8, 99, 101, -4, 2]), [-65, -10, 0, -4, 2, 8, 5, 9, 77, 13, 88, 101, 99, 313]);
```

# --seed--

## --seed-contents--

```js
function sortAndSwap(arr) {

  return arr;
}
```

# --solutions--

```js
function sortAndSwap(arr) {
  const result = [...arr].sort((a, b) => a - b);

  for (let i = 3; i < result.length; i += 3) {
    [result[i], result[i - 1]] = [result[i - 1], result[i]];
  }

  return result;
}
```
