---
id: 69bc6cb30c1d112a2e110a03
title: "Aufgabe 246: Initialen von Namen"
challengeType: 28
dashedName: challenge-246
---

# --description--

Gegeben ist ein vollständiger Name als String. Geben Sie die Initialen zurück.

- Die Namen, deren Initialen gebildet werden sollen, sind durch ein Leerzeichen getrennt.
- Die Initialen sollen in Großbuchstaben sein.
- Die Initialen sollen durch Punkte getrennt sein.

Zum Beispiel gibt `"Tommy Millwood"` `"T.M."` zurück.

# --hints--

`getInitials("Tommy Millwood")` sollte `"T.M."` zurückgeben.

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")` sollte `"S.P."` zurückgeben.

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")` sollte `"F.C.C."` zurückgeben.

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")` sollte `"D."` zurückgeben.

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")` sollte `"D.V.C.H.N."` zurückgeben.

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
