---
id: 69bc6cb30c1d112a2e110a07
title: "Aufgabe 250: Versteckter Schlüssel"
challengeType: 28
dashedName: challenge-250
---

# --description--

Willkommen zur 250. täglichen Aufgabe!

Gegeben ist ein verschlüsselter String. Entschlüsseln Sie ihn mit einem Verschlüsselungsschlüssel und geben Sie das Ergebnis zurück.

So finden Sie den Schlüssel:

- Schauen Sie sich alle täglichen Aufgaben bis heute an, deren Aufgabennummer ein Vielfaches von 25 ist (einschließlich dieser).
- Nehmen Sie den ersten Buchstaben jedes dieser Aufgabentitel und fügen Sie sie zu einem String zusammen. Beginnt der Titel mit einem Nicht-Buchstaben, finden Sie den ersten Buchstaben darin.

Um die Nachricht zu entschlüsseln, gehen Sie jeden Buchstaben der verschlüsselten Nachricht durch und:

  1. Schauen Sie sich den entsprechenden Buchstaben im Schlüssel an (wiederholen Sie den Schlüssel, falls die Nachricht länger ist als der Schlüssel).
  2. Wandeln Sie den Schlüsselbuchstaben in seine entsprechende Zahl um: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
  3. Verschieben Sie den verschlüsselten Buchstaben im Alphabet um diese Zahl rückwärts.
  4. Wenn die Verschiebung vor `"A"` liegt, springen Sie zum Ende des Alphabets zurück zu `"Z"`.

Zum Beispiel: Wenn die verschlüsselte Nachricht mit `"Y"` beginnt und der erste Schlüsselbuchstabe `"V"` (22) ist, verschieben Sie `"Y"` um 22 Stellen zurück, um `"C"` zu erhalten. Wiederholen Sie diesen Vorgang für jeden Buchstaben, um die gesamte Nachricht zu entschlüsseln.

- Nur Buchstaben werden verschoben, Leerzeichen bleiben unverändert.
- Alle gegebenen und zurückgegebenen Buchstaben sind Großbuchstaben.

# --hints--

`decode("YAVJYNXE")` sollte `"CONGRATS"` zurückgeben.

```js
assert.equal(decode("YAVJYNXE"), "CONGRATS");
```

`decode("YALLUT PQUMJP")` sollte `"CODING LEGEND"` zurückgeben.

```js
assert.equal(decode("YALLUT PQUMJP"), "CODING LEGEND");
```

`decode("UAC DYR EISAKYM")` sollte `"YOU ARE AWESOME"` zurückgeben.

```js
assert.equal(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME");
```

`decode("GQMS NBMZU")` sollte `"KEEP GOING"` zurückgeben.

```js
assert.equal(decode("GQMS NBMZU"), "KEEP GOING");
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` sollte `"A WINNER IS A DREAMER WHO NEVER GIVES UP"` zurückgeben.

```js
assert.equal(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP");
```

# --seed--

## --seed-contents--

```js
function decode(message) {

  return message;
}
```

# --solutions--

```js
function decode(message) {
  const key = "VLHCGMDLNH";
  let keyIndex = 0;
  return message.split('').map(char => {
    if (char === ' ') return char;
    const shift = key[keyIndex % key.length].charCodeAt(0) - 64;
    keyIndex++;
    return String.fromCharCode(((char.charCodeAt(0) - 65 - shift + 26) % 26) + 65);
  }).join('');
}
```
