---
id: 69c5f3d787b1725d5f00c8b9
title: "Aufgabe 255: Earth-Day-Aufräumteam"
challengeType: 28
dashedName: challenge-255
---

# --description--

Heute ist Earth Day. Gegeben ist ein Array von Gegenständen, die Sie aufgeräumt haben. Geben Sie Ihre Gesamtpunktzahl für das Aufräumen basierend auf den folgenden Regeln zurück.

Die gegebenen Gegenstände sind einer der folgenden:

| Gegenstand | Grundwert |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

Ein seltener Gegenstand wird als `["rare", value]` dargestellt. Zum Beispiel `["rare", 80]`. Seltene Gegenstände erhalten keinen Bonus für aufeinanderfolgende Vorkommen.

- Bonus für aufeinanderfolgende Vorkommen: Wenn derselbe Gegenstand hintereinander erscheint, erhält er steigende Bonuspunkte.
  - Erstes aufeinanderfolgendes Vorkommen: Grundwert
  - Zweites: Grundwert + 1
  - Drittes: Grundwert + 2
  - usw.
- Multiplikator für jeden fünften Gegenstand: Jeder fünfte gesammelte Gegenstand erhält einen Multiplikator.
  - Fünfter Gegenstand: *2
  - Zehnter Gegenstand: *3
  - usw.

- Wenden Sie den Multiplikator nach der Berechnung aller Boni an.

# --hints--

`getCleanupScore(["bottle", "straw", "shoe", "battery"])` sollte `44` zurückgeben.

```js
assert.equal(getCleanupScore(["bottle", "straw", "shoe", "battery"]), 44);
```

`getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"])` sollte `58` zurückgeben.

```js
assert.equal(getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"]), 58);
```

`getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` sollte `79` zurückgeben.

```js
assert.equal(getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79);
```

`getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` sollte `358` zurückgeben.

```js
assert.equal(getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358);
```

`getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` sollte `383` zurückgeben.

```js
assert.equal(getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383);
```

# --seed--

## --seed-contents--

```js
function getCleanupScore(items) {

  return items;
}
```

# --solutions--

```js
function getCleanupScore(items) {
  const values = {
    bottle: 10,
    can: 6,
    bag: 8,
    tire: 35,
    straw: 4,
    cardboard: 3,
    newspaper: 3,
    shoe: 12,
    electronics: 25,
    battery: 18,
    mattress: 38
  };

  let total = 0;
  let prevItem = null;
  let streak = 0;

  for (let i = 0; i < items.length; i++) {
    let item = items[i];
    let base;
    let isRare = Array.isArray(item) && item[0] === "rare";

    if (isRare) {
      base = item[1];
      prevItem = null;
      streak = 0;
    } else {
      base = values[item];
      streak = item === prevItem ? streak + 1 : 0;
      prevItem = item;
    }

    let points = base;
    if (!isRare) points += streak;

    if ((i + 1) % 5 === 0) {
      const multiplier = (i + 1) / 5 + 1;
      points *= multiplier;
    }

    total += points;
  }

  return total;
}
```
