---
id: 69c5f3d787b1725d5f00c8bd
title: "Aufgabe 259: FizzBuzz-Explosion"
challengeType: 28
dashedName: challenge-259
---

# --description--

Gegeben ist eine ganze Zahl. Geben Sie die Anzahl der Schritte zurück, die benötigt werden, um das Wort `"fizzbuzz"` in einen String mit mindestens der angegebenen Anzahl von `"z"`-Zeichen zu verwandeln, wobei die folgenden Regeln gelten:

- Beginnen Sie mit dem String `"fizzbuzz"`.
- Wenden Sie in jedem Schritt die Standard-FizzBuzz-Regeln auf die Buchstabenposition im String an (der erste `"f"` hat die Position 1).
  - Wenn die Buchstabenposition durch 3 teilbar ist, ersetzen Sie den Buchstaben durch `"fizz"`
  - Wenn sie durch 5 teilbar ist, ersetzen Sie den Buchstaben durch `"buzz"`
  - Wenn sie durch 3 und 5 teilbar ist, ersetzen Sie den Buchstaben durch `"fizzbuzz"`

Nach 1 Schritt verwandelt sich `"fizzbuzz"` in `"fifizzzbuzzfizzzz"`, das 9 `"z"`-Zeichen enthält.

# --hints--

`explodeFizzbuzz(9)` sollte `1` zurückgeben.

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` sollte `2` zurückgeben.

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` sollte `3` zurückgeben.

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` sollte `4` zurückgeben.

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` sollte `5` zurückgeben.

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` sollte `6` zurückgeben.

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` sollte `11` zurückgeben.

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` sollte `14` zurückgeben.

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
