---
id: 69c5f3d787b1725d5f00c8b7
title: "Aufgabe 253: Akronymfinder"
challengeType: 28
dashedName: challenge-253
---

# --description--

Gegeben ist ein String, der ein Akronym darstellt. Geben Sie den vollständigen Namen der Organisation zurück, zu der es gehört, aus der folgenden Liste:

- `"National Avocado Storage Authority"`
- `"Cats Infiltration Agency"`
- `"Fluffy Beanbag Inspectors"`
- `"Department Of Jelly"`
- `"Wild Honey Organization"`
- `"Eating Pancakes Administration"`

Jeder Buchstabe im gegebenen Akronym soll mit dem ersten Buchstaben jedes Wortes der zugehörigen Organisation übereinstimmen, und zwar in der gleichen Reihenfolge.

# --hints--

`findOrg("NASA")` sollte `"National Avocado Storage Authority"` zurückgeben.

```js
assert.equal(findOrg("NASA"), "National Avocado Storage Authority");
```

`findOrg("CIA")` sollte `"Cats Infiltration Agency"` zurückgeben.

```js
assert.equal(findOrg("CIA"), "Cats Infiltration Agency");
```

`findOrg("FBI")` sollte `"Fluffy Beanbag Inspectors"` zurückgeben.

```js
assert.equal(findOrg("FBI"), "Fluffy Beanbag Inspectors");
```

`findOrg("DOJ")` sollte `"Department Of Jelly"` zurückgeben.

```js
assert.equal(findOrg("DOJ"), "Department Of Jelly");
```

`findOrg("WHO")` sollte `"Wild Honey Organization"` zurückgeben.

```js
assert.equal(findOrg("WHO"), "Wild Honey Organization");
```

`findOrg("EPA")` sollte `"Eating Pancakes Administration"` zurückgeben.

```js
assert.equal(findOrg("EPA"), "Eating Pancakes Administration");
```

# --seed--

## --seed-contents--

```js
function findOrg(acronym) {

  return acronym;
}
```

# --solutions--

```js
function findOrg(acronym) {
  const orgs = [
    "National Avocado Storage Authority",
    "Cats Infiltration Agency",
    "Fluffy Beanbag Inspectors",
    "Department Of Jelly",
    "Wild Honey Organization",
    "Eating Pancakes Administration"
  ];

  for (let org of orgs) {
    const initials = org
      .split(" ")
      .map(word => word[0])
      .join("");

    if (initials.toUpperCase() === acronym.toUpperCase()) {
      return org;
    }
  }

  return "not found";
}
```
