---
id: 69c5f3d787b1725d5f00c8b8
title: "Aufgabe 254: Wörter mit ungerader Länge"
challengeType: 28
dashedName: challenge-254
---

# --description--

Gegeben ist ein String mit Wörtern. Geben Sie nur die Wörter zurück, die eine ungerade Anzahl von Buchstaben haben.

- Die Wörter im gegebenen String sind durch ein einzelnes Leerzeichen getrennt.
- Geben Sie die Wörter durch ein einzelnes Leerzeichen getrennt zurück.

# --hints--

`getOddWords("This is a super good test")` sollte `"a super"` zurückgeben.

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` sollte `"one two three"` zurückgeben.

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` sollte `"split rainbow sprinkles top"` zurückgeben.

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` sollte `"The quick brown fox the river"` zurückgeben.

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
