---
id: 69c5f3d787b1725d5f00c8bc
title: "Aufgabe 258: Wortdekompressor"
challengeType: 28
dashedName: challenge-258
---

# --description--

Gegeben ist ein komprimierter String. Geben Sie die dekomprimierte Version zurück, indem Sie die folgenden Regeln anwenden:

- Der gegebene String besteht aus Wörtern und Zahlen, die durch Leerzeichen getrennt sind.
- Lassen Sie die Wörter unverändert.
- Ersetzen Sie Zahlen durch das Wort an dieser Position, wobei das erste Wort die Position 1 hat.

Zum Beispiel, gegeben `"practice makes perfect and 3 1 2 3"`, geben Sie `"practice makes perfect and perfect practice makes perfect"` zurück.

# --hints--

`decompress("practice makes perfect and 3 1 2 3")` sollte `"practice makes perfect and perfect practice makes perfect"` zurückgeben.

```js
assert.equal(decompress("practice makes perfect and 3 1 2 3"), "practice makes perfect and perfect practice makes perfect");
```

`decompress("hello 1 1")` sollte `"hello hello hello"` zurückgeben.

```js
assert.equal(decompress("hello 1 1"), "hello hello hello");
```

`decompress("the cat sat on 1 mat 4 which 1 2 3")` sollte `"the cat sat on the mat on which the cat sat"` zurückgeben.

```js
assert.equal(decompress("the cat sat on 1 mat 4 which 1 2 3"), "the cat sat on the mat on which the cat sat");
```

`decompress("the more you know 1 2 3 realize 3 don't 4")` sollte `"the more you know the more you realize you don't know"` zurückgeben.

```js
assert.equal(decompress("the more you know 1 2 3 realize 3 don't 4"), "the more you know the more you realize you don't know");
```

`decompress("lorem ipsum dolor sit per elit donec 4 nostra libero 5 7 ligula 4 gravida at 6 vitae a 6 sodales 7 en 7 16 3 nam 13 dignissim risus 16 13 5 27 2 2 15 23 6 5 2 13 23 15 5 21 4 16 27 1 4 5 10 23 2 6 4 21 4 30 6 30 2 6 16 15 18 23 29 27 4 18 sollicitudin 5 9 5 4 10")` sollte `"lorem ipsum dolor sit per elit donec sit nostra libero per donec ligula sit gravida at elit vitae a elit sodales donec en donec at dolor nam ligula dignissim risus at ligula per nam ipsum ipsum gravida en elit per ipsum ligula en gravida per sodales sit at nam lorem sit per libero en ipsum elit sit sodales sit risus elit risus ipsum elit at gravida vitae en dignissim nam sit vitae sollicitudin per nostra per sit libero"` zurückgeben.

```js
assert.equal(decompress("lorem ipsum dolor sit per elit donec 4 nostra libero 5 7 ligula 4 gravida at 6 vitae a 6 sodales 7 en 7 16 3 nam 13 dignissim risus 16 13 5 27 2 2 15 23 6 5 2 13 23 15 5 21 4 16 27 1 4 5 10 23 2 6 4 21 4 30 6 30 2 6 16 15 18 23 29 27 4 18 sollicitudin 5 9 5 4 10"), "lorem ipsum dolor sit per elit donec sit nostra libero per donec ligula sit gravida at elit vitae a elit sodales donec en donec at dolor nam ligula dignissim risus at ligula per nam ipsum ipsum gravida en elit per ipsum ligula en gravida per sodales sit at nam lorem sit per libero en ipsum elit sit sodales sit risus elit risus ipsum elit at gravida vitae en dignissim nam sit vitae sollicitudin per nostra per sit libero");
```

# --seed--

## --seed-contents--

```js
function decompress(str) {

  return str;
}
```

# --solutions--

```js
function decompress(str) {
  const tokens = str.split(" ");
  const words = [];

  for (const token of tokens) {
    if (isNaN(parseInt(token))) {
      words.push(token)
    } else {
      words.push(tokens[token - 1])
    }
  }

  return words.join(" ")
}
```
