---
id: 69c5f3d787b1725d5f00c8ba
title: "Aufgabe 256: Nächste Zeitrichtung"
challengeType: 28
dashedName: challenge-256
---

# --description--

Gegeben sind zwei Zeiten. Bestimmen Sie, ob Sie vom ersten Zeitpunkt zum zweiten schneller vorwärts oder rückwärts gelangen.

- Die Zeiten sind im 24-Stunden-Format angegeben (`"HH:MM"`)
- Die Uhr läuft zyklisch (23:59 geht bei Vorwärtsbewegung zu 00:00 über, und 00:00 geht bei Rückwärtsbewegung zu 23:59)

Geben Sie zurück:

- `"forward"`, wenn Vorwärtsbewegung kürzer ist
- `"backward"`, wenn Rückwärtsbewegung kürzer ist
- `"equal"`, wenn beide Richtungen gleich lange dauern

# --hints--

`getDirection("10:00", "12:00")` sollte `"forward"` zurückgeben.

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` sollte `"backward"` zurückgeben.

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` sollte `"equal"` zurückgeben.

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` sollte `"forward"` zurückgeben.

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` sollte `"backward"` zurückgeben.

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` sollte `"equal"` zurückgeben.

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
