---
id: 69c6ff713a52713463aa7929
title: "Aufgabe 263: Binäres Kreuzworträtsel"
challengeType: 28
dashedName: challenge-263
---

# --description--

Gegeben ist ein Zeichen. Bestimmen Sie, ob seine 8-Bit-Binärdarstellung im folgenden Raster horizontal oder vertikal in beliebiger Richtung zu finden ist:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

Zum Beispiel hat `"A"` die Binärdarstellung `01000001`, die in der ersten Zeile von links nach rechts erscheint.

# --hints--

`isInCrossword("I")` sollte `true` zurückgeben.

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` sollte `true` zurückgeben.

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` sollte `true` zurückgeben.

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` sollte `true` zurückgeben.

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` sollte `false` zurückgeben.

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` sollte `false` zurückgeben.

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` sollte `false` zurückgeben.

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` sollte `false` zurückgeben.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
