---
id: 69c6ff713a52713463aa7927
title: "Aufgabe 261: Zahlwörter"
challengeType: 28
dashedName: challenge-261
---

# --description--

Gegeben ist eine ganze Zahl von 0 bis 99. Geben Sie ihre englische Wortdarstellung zurück.

- `0` gibt `"zero"` zurück.
- Die Zahlen 1–19 haben eigene Namen (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- Vielfache von 10 von 20 bis 90 haben eigene Namen (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- Zahlen von 21 bis 99, die keine Vielfachen von 10 sind, werden als zwei durch einen Bindestrich verbundene Wörter geschrieben. Zum Beispiel `"forty-two"` und `"fifty-three"`.

# --hints--

`getNumberWords(0)` sollte `"zero"` zurückgeben.

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` sollte `"ten"` zurückgeben.

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` sollte `"nineteen"` zurückgeben.

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` sollte `"thirty"` zurückgeben.

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` sollte `"fifty-three"` zurückgeben.

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` sollte `"seven"` zurückgeben.

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` sollte `"twelve"` zurückgeben.

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` sollte `"sixty"` zurückgeben.

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` sollte `"sixty-seven"` zurückgeben.

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` sollte `"ninety-eight"` zurückgeben.

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
