---
id: 69c6ff713a52713463aa7928
title: "Aufgabe 262: URL-Abfrageparser"
challengeType: 28
dashedName: challenge-262
---

# --description--

Gegeben ist eine URL, die eine Abfragezeichenfolge enthält. Zerlegen Sie die Abfragezeichenfolge in ein Objekt (oder Wörterbuch) aus Schlüssel-Wert-Paaren.

- Die Abfragezeichenfolge beginnt nach dem `"?"`,
- jeder Parameter ist durch `"&"` getrennt,
- jedes Schlüssel/Wert-Paar ist durch `"="` getrennt

Zum Beispiel, gegeben `"https://example.com/search?name=Alice&age=30"`, geben Sie zurück:

```json
{
  "name": "Alice",
  "age": "30"
}
```

Alle Werte sollen als Strings zurückgegeben werden.

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")` sollte `{"name": "Alice", "age": "30"}` zurückgeben

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")` sollte `{"skill": "programming", "language": "python"}` zurückgeben

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")` sollte `{"category": "books", "sort": "asc", "page": "2"}` zurückgeben

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")` sollte `{"redirect": "freecodecamp.org/learn", "when": "now"}` zurückgeben

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```
