---
id: 69c6ff713a52713463aa7926
title: "Aufgabe 260: Wortpunktzahl"
challengeType: 28
dashedName: challenge-260
---

# --description--

Gegeben ist ein Wort. Geben Sie dessen Punktzahl zurück, indem Sie eine standardmäßige Buchstaben-Wert-Tabelle verwenden:

| Buchstabe | Wert |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- Groß- und Kleinbuchstaben haben denselben Wert.

# --hints--

`getWordScore("hi")` sollte `17` zurückgeben.

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` sollte `52` zurückgeben.

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` sollte `169` zurückgeben.

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` sollte `94` zurückgeben.

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
