---
id: 69cfca90e8a0a6d4d6871c51
title: "Aufgabe 267: Parsec-Konverter"
challengeType: 28
dashedName: challenge-267
---

# --description--

In einer fernen Galaxie werden Parsec sowohl zur Messung von Zeit als auch von Entfernung verwendet. Gegeben ist eine ganze Zahl von Parsec, geben Sie deren Äquivalent in Zeit oder Entfernung zurück.

- Wenn die gegebene ganze Zahl ungerade ist, stellt sie Zeit dar. Wenn sie gerade ist, stellt sie Entfernung dar.

Verwenden Sie diese Umrechnungswerte:

| Parsec | Zeit/Entfernung |
| - | - |
| 1 | 2 Stunden |
| 2 | 6 Lichtjahre |

Geben Sie den umgerechneten Wert als ganze Zahl zurück.

# --hints--

`convertParsecs(1)` sollte `2` zurückgeben.

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)` sollte `6` zurückgeben.

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)` sollte `62` zurückgeben.

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)` sollte `264` zurückgeben.

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)` sollte `34` zurückgeben.

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)` sollte `42` zurückgeben.

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
