---
id: 69cfca90e8a0a6d4d6871c54
title: "Aufgabe 270: Längster gemeinsamer Teilstring"
challengeType: 28
dashedName: challenge-270
---

# --description--

Gegeben ist ein String. Geben Sie den längsten Teilstring zurück, der mehr als einmal vorkommt.

- Die Teilstrings können sich überlappen.

# --hints--

`getLongestSubstring("abracadabra")` sollte `"abra"` zurückgeben.

```js
assert.equal(getLongestSubstring("abracadabra"), "abra");
```

`getLongestSubstring("hello world hello")` sollte `"hello"` zurückgeben.

```js
assert.equal(getLongestSubstring("hello world hello"), "hello");
```

`getLongestSubstring("mississippi")` sollte `"issi"` zurückgeben.

```js
assert.equal(getLongestSubstring("mississippi"), "issi");
```

`getLongestSubstring("ha ha ha ha ha ha ha")` sollte `"ha ha ha ha ha ha"` zurückgeben.

```js
assert.equal(getLongestSubstring("ha ha ha ha ha ha ha"), "ha ha ha ha ha ha");
```

`getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over")` sollte `"the quick brown fox jumped over"` zurückgeben.

```js
assert.equal(getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over"), "the quick brown fox jumped over");
```

# --seed--

## --seed-contents--

```js
function getLongestSubstring(str) {

  return str;
}
```

# --solutions--

```js
function getLongestSubstring(str) {
  let longest = '';

  for (let len = str.length - 1; len >= 1; len--) {
    for (let i = 0; i <= str.length - len; i++) {
      const sub = str.slice(i, i + len);
      if (str.indexOf(sub) !== str.lastIndexOf(sub)) {
        return sub;
      }
    }
  }

  return longest;
}
```
