---
id: 69cfca90e8a0a6d4d6871c53
title: "Aufgabe 269: Allergikerfreundliche Mahlzeiten"
challengeType: 28
dashedName: challenge-269
---

# --description--

Gegeben sind ein Array von Mahlzeiten und ein Array von Allergenen, die vermieden werden sollen. Geben Sie die Namen aller Mahlzeiten zurück, die keines der angegebenen Allergene enthalten.

- Jede Mahlzeit hat das Format `[meal, allergens]`, wobei `meal` der Name der Mahlzeit ist und `allergens` ein Array der Allergene, die die Mahlzeit enthält. Zum Beispiel `["pasta", ["wheat", "milk"]]`.
- Die zu vermeidenden Allergene sind ein Array von Strings.

Geben Sie die sicheren Mahlzeitenamen in der gleichen Reihenfolge zurück, wie sie angegeben sind. Wenn keine Mahlzeit sicher ist, geben Sie ein leeres Array zurück.

# --hints--

`getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"])` sollte `["salad"]` zurückgeben.

```js
assert.deepEqual(getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"]), ["salad"]);
```

`getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"])` sollte `["fried rice", "chicken parmesan"]` zurückgeben.

```js
assert.deepEqual(getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"]), ["fried rice", "chicken parmesan"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"])` sollte `["oatmeal", "granola", "toast"]` zurückgeben.

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"]), ["oatmeal", "granola", "toast"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"])` sollte `["granola", "yogurt", "eggs"]` zurückgeben.

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"]), ["granola", "yogurt", "eggs"]);
```

# --seed--

## --seed-contents--

```js
function getAllergenFriendlyMeals(meals, allergens) {

  return meals;
}
```

# --solutions--

```js
function getAllergenFriendlyMeals(meals, allergens) {
  return meals
    .filter(([name, mealAllergens]) =>
      !allergens.some(a => mealAllergens.includes(a)))
    .map(([name]) => name);
}
```
