---
id: 69cfca90e8a0a6d4d6871c50
title: "Aufgabe 266: Guten Tag"
challengeType: 28
dashedName: challenge-266
---

# --description--

Gegeben ist ein Zeit-String im `"HH:MM"`-Format (24-Stunden-Uhr). Geben Sie zurück:

- `"Good morning"` für Zeiten von `05:00` bis `11:59`
- `"Good afternoon"` für Zeiten von `12:00` bis `17:59`
- `"Good evening"` für Zeiten von `18:00` bis `21:59`
- `"Good night"` für Zeiten von `22:00` bis `04:59`

# --hints--

`getGreeting("06:30")` sollte `"Good morning"` zurückgeben.

```js
assert.equal(getGreeting("06:30"), "Good morning");
```

`getGreeting("12:00")` sollte `"Good afternoon"` zurückgeben.

```js
assert.equal(getGreeting("12:00"), "Good afternoon");
```

`getGreeting("21:59")` sollte `"Good evening"` zurückgeben.

```js
assert.equal(getGreeting("21:59"), "Good evening");
```

`getGreeting("00:01")` sollte `"Good night"` zurückgeben.

```js
assert.equal(getGreeting("00:01"), "Good night");
```

`getGreeting("11:30")` sollte `"Good morning"` zurückgeben.

```js
assert.equal(getGreeting("11:30"), "Good morning");
```

# --seed--

## --seed-contents--

```js
function getGreeting(time) {

  return time;
}
```

# --solutions--

```js
function getGreeting(time) {
  const [hours, minutes] = time.split(':').map(Number);
  const total = hours * 60 + minutes;

  if (total >= 300 && total < 720) return 'Good morning';
  if (total >= 720 && total < 1080) return 'Good afternoon';
  if (total >= 1080 && total < 1320) return 'Good evening';
  return 'Good night';
}
```
