---
id: 69cfca90e8a0a6d4d6871c4f
title: "Aufgabe 265: Tiefste Klammern"
challengeType: 28
dashedName: challenge-265
---

# --description--

Gegeben ist ein String mit ausgeglichenen Klammern. Geben Sie den Inhalt der am tiefsten verschachtelten Klammern zurück.

- Klammern können eine der drei Typen sein: `()`, `[]` und `{}`.
- Die Eingabe enthält immer genau eine tiefste Gruppe.

Zum Beispiel, bei `"(hello (world))"` geben Sie `"world"` zurück.

# --hints--

`getDeepestBrackets("(hello (world))")` sollte `"world"` zurückgeben.

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")` sollte `"inner"` zurückgeben.

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")` sollte `"e"` zurückgeben.

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` sollte `"fox"` zurückgeben.

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` sollte `"C"` zurückgeben.

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```
