---
id: 69cfca90e8a0a6d4d6871c52
title: "Aufgabe 268: Narzisstische Zahl"
challengeType: 28
dashedName: challenge-268
---

# --description--

Gegeben ist eine positive ganze Zahl. Bestimmen Sie, ob es sich um eine narzisstische Zahl handelt.

- Eine Zahl ist narzisstisch, wenn die Summe jeder ihrer Ziffern, potenziert mit der Gesamtanzahl der Ziffern, der Zahl selbst entspricht.

Zum Beispiel hat 153 drei Ziffern, und 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153, daher ist sie narzisstisch.

# --hints--

`isNarcissistic(153)` sollte `true` zurückgeben.

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` sollte `false` zurückgeben.

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` sollte `true` zurückgeben.

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` sollte `false` zurückgeben.

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` sollte `true` zurückgeben.

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` sollte `false` zurückgeben.

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` sollte `true` zurückgeben.

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` sollte `false` zurückgeben.

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
