---
id: 69cfca90e8a0a6d4d6871c4e
title: "Aufgabe 264: Anagramm-Gruppen"
challengeType: 28
dashedName: challenge-264
---

# --description--

Gegeben ist ein Array von Wörtern. Geben Sie ein zweidimensionales Array zurück, in dem die Wörter in Anagramm-Gruppen zusammengefasst sind.

- Wörter sind Anagramme, wenn sie dieselben Buchstaben in beliebiger Reihenfolge enthalten.
- Jedes Wort gehört genau zu einer Gruppe.
- Die Reihenfolge der Rückgabe spielt keine Rolle.

Zum Beispiel, gegeben `["listen", "silent", "hello", "enlist", "world"]`, geben Sie `[["listen", "silent", "enlist"], ["hello"], ["world"]]` zurück.

# --hints--

`groupAnagrams(["listen", "silent", "hello", "enlist", "world"])` sollte `[["listen", "silent", "enlist"], ["hello"], ["world"]]` zurückgeben.

```js
const groups = groupAnagrams(["listen", "silent", "hello", "enlist", "world"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["enlist", "listen", "silent"], ["hello"], ["world"]]);
```

`groupAnagrams(["eat", "tea", "tan", "ate", "nat", "bat"])` sollte `[["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]` zurückgeben.

```js
const groups = groupAnagrams(["eat", "tea", "tan", "ate", "nat", "bat"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]);
```

`groupAnagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])` sollte `[["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]` zurückgeben.

```js
const groups = groupAnagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]);

```

`groupAnagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])` sollte `[["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]` zurückgeben.

```js
const groups = groupAnagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]);
```

# --seed--

## --seed-contents--

```js
function groupAnagrams(words) {

  return words;
}
```

# --solutions--

```js
function groupAnagrams(words) {
  const map = {};
  for (const word of words) {
    const key = word.split('').sort().join('');
    if (!map[key]) map[key] = [];
    map[key].push(word);
  }
  return Object.values(map);
}
```
