---
id: 69d03e549613fbdbe21d11e0
title: "Aufgabe 271: Medikamentenerinnerung"
challengeType: 28
dashedName: challenge-271
---

# --description--

Gegeben ist ein Array von Medikamenten und ein String, der die aktuelle Uhrzeit darstellt. Geben Sie das nächste Medikament zurück, das Sie einnehmen müssen, und wie lange es noch dauert, bis Sie es einnehmen müssen.

- Jedes Medikament hat das Format `[name, lastTaken]`, wobei `name` der Name des Medikaments und `lastTaken` die Zeit ist, zu der es zuletzt eingenommen wurde.
- Alle angegebenen Zeiten sind im `"HH:MM"`-Format (24-Stunden-Format).

Verwenden Sie den folgenden Medikamentenplan:

| Name | Zeitplan |
| - | - |
| Deployxitrin | 08:00, 16:00 |
| Debuggamanizole | 07:00, 13:00, 21:00 |
| Mergeflictamine | alle 4 Stunden |

Geben Sie einen String im Format `"{name} in Hh Mm"` zurück. Zum Beispiel `"Debuggamanizole in 2h 0m"` oder `"Deployxitrin in 1h 5m"`.

# --hints--

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00")` sollte `"Debuggamanizole in 2h 0m"` zurückgeben.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00"), "Debuggamanizole in 2h 0m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55")` sollte `"Deployxitrin in 1h 5m"` zurückgeben.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55"), "Deployxitrin in 1h 5m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15")` sollte `"Mergeflictamine in 0h 45m"` zurückgeben.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15"), "Mergeflictamine in 0h 45m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59")` sollte `"Debuggamanizole in 0h 1m"` zurückgeben.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59"), "Debuggamanizole in 0h 1m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55")` sollte `"Debuggamanizole in 0h 5m"` zurückgeben.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55"), "Debuggamanizole in 0h 5m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00")` sollte `"Mergeflictamine in 3h 30m"` zurückgeben.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00"), "Mergeflictamine in 3h 30m");
```

# --seed--

## --seed-contents--

```js
function medicationReminder(medications, currentTime) {

  return medications;
}
```

# --solutions--

```js
function medicationReminder(medications, currentTime) {
  const schedule = {
    Deployxitrin: { type: 'fixed', times: ['08:00', '16:00'] },
    Debuggamanizole: { type: 'fixed', times: ['07:00', '13:00', '21:00'] },
    Mergeflictamine: { type: 'interval', hours: 4 }
  };

  const toMinutes = t => {
    const [h, m] = t.split(':').map(Number);
    return h * 60 + m;
  };

  const current = toMinutes(currentTime);
  let earliest = Infinity;
  let nextMed = '';

  for (const [name, lastTaken] of medications) {
    const med = schedule[name];
    let nextDue;

    if (med.type === 'fixed') {
      const futureTimes = med.times
        .map(toMinutes)
        .filter(t => t > current);
      if (futureTimes.length === 0) continue;
      nextDue = Math.min(...futureTimes);
    } else {
      nextDue = toMinutes(lastTaken) + med.hours * 60;
    }

    const diff = nextDue - current;
    if (diff > 0 && diff < earliest) {
      earliest = diff;
      nextMed = name;
    }
  }

  const h = Math.floor(earliest / 60);
  const m = earliest % 60;
  return `${nextMed} in ${h}h ${m}m`;
}
```
