---
id: 69e2383af7832c8032603b93
title: "Aufgabe 277: Spiegelbild"
challengeType: 28
dashedName: challenge-277
---

# --description--

Gegeben sind zwei Strings. Bestimmen Sie, ob der zweite String ein Spiegelbild des ersten ist.

Ein Spiegelbild entsteht, indem man den String umkehrt und jedes Zeichen durch sein Spiegeläquivalent ersetzt.

- Symmetrische Zeichen sehen im Spiegel aus wie sie selbst:

`W`, `T`, `Y`, `U`, `I`, `O`, `H`, `A`, `X`, `V`, `M`, `w`, `o`, `x`, `v`, `0`, `8`, `=`, `+`, `:`, `|`, `-`, `_`, `*`, `^`, `!`, `.`, und das Leerzeichen (` `).

- Spiegelpaare tauschen im Spiegel miteinander:

| Zeichen | Tauscht mit |
| - | - |
| `[` | `]` |
| `{` | `}` |
| `<` | `>` |
| `b` | `d` |
| `p` | `q` |
| `(` | `)` |

Wenn einer der Strings ein Zeichen enthält, das nicht in den obigen Listen steht, kann daraus kein Spiegelbild gebildet werden.

Zum Beispiel ist das Spiegelbild von `"[HOW]"` `"[WOH]"`.

# --hints--

`isMirrorImage("[HOW]", "[WOH]")` sollte `true` zurückgeben.

```js
assert.isTrue(isMirrorImage("[HOW]", "[WOH]"));
```

`isMirrorImage("MOM", "MOM")` sollte `true` zurückgeben.

```js
assert.isTrue(isMirrorImage("MOM", "MOM"));
```

`isMirrorImage("vow", "wov")` sollte `true` zurückgeben.

```js
assert.isTrue(isMirrorImage("vow", "wov"));
```

`isMirrorImage("TIM", "TIM")` sollte `false` zurückgeben.

```js
assert.isFalse(isMirrorImage("TIM", "TIM"));
```

`isMirrorImage("{WOW}", "}WOW{")` sollte `false` zurückgeben.

```js
assert.isFalse(isMirrorImage("{WOW}", "}WOW{"));
```

`isMirrorImage("XXVII", "IIV%X")` sollte `false` zurückgeben.

```js
assert.isFalse(isMirrorImage("XXVII", "IIV%X"));
```

`isMirrorImage("><(((*>", "<*)))><")` sollte `true` zurückgeben.

```js
assert.isTrue(isMirrorImage("><(((*>", "<*)))><"));
```

`isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW")` sollte `true` zurückgeben.

```js
assert.isTrue(isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW"));
```

# --seed--

## --seed-contents--

```js
function isMirrorImage(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirrorImage(str1, str2) {
  const symmetric = new Set('WTYUIOHA XVMwoxv08=+:|-_*^!.');
  const pairs = { '[': ']', ']': '[', '{': '}', '}': '{', '<': '>', '>': '<', 'b': 'd', 'd': 'b', 'p': 'q', 'q': 'p', '(': ')', ')': '(' };
  
  const mirrored = str1.split('').reverse().map(c => {
    if (symmetric.has(c)) return c;
    if (pairs[c]) return pairs[c];
    return null;
  });
  
  if (mirrored.includes(null)) return false;
  return mirrored.join('') === str2;
}
```
