---
id: 69e2383af7832c8032603b8e
title: "Aufgabe 272: Transponierte Matrix"
challengeType: 28
dashedName: challenge-272
---

# --description--

Gegeben ist eine Matrix (ein Array von Arrays). Geben Sie die transponierte Version davon zurück.

Um die Matrix zu transponieren, tauschen Sie die Zeilen und Spalten. Zum Beispiel: Ein Wert an Index `[0, 1]` sollte an Index `[1, 0]` verschoben werden.

Zum Beispiel, gegeben:

```json
[
  [1, 2, 3],
  [4, 5, 6]
]
```

Geben Sie zurück:

```json
[
  [1, 4],
  [2, 5],
  [3, 6]
]
```

# --hints--

`transpose([[1, 2, 3], [4, 5, 6]])` sollte `[[1, 4], [2, 5], [3, 6]]` zurückgeben.

```js
assert.deepEqual(transpose([[1, 2, 3], [4, 5, 6]]), [[1, 4], [2, 5], [3, 6]]);
```

`transpose([[1, 2], [3, 4], [5, 6]])` sollte `[[1, 3, 5], [2, 4, 6]]` zurückgeben.

```js
assert.deepEqual(transpose([[1, 2], [3, 4], [5, 6]]), [[1, 3, 5], [2, 4, 6]]);
```

`transpose([[1, 2], [3, 4], [5, 6], [7, 8]])` sollte `[[1, 3, 5, 7], [2, 4, 6, 8]]` zurückgeben.

```js
assert.deepEqual(transpose([[1, 2], [3, 4], [5, 6], [7, 8]]), [[1, 3, 5, 7], [2, 4, 6, 8]]);
```

`transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]])` sollte `[["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]` zurückgeben.

```js
assert.deepEqual(transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]]), [["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]);
```

`transpose([[true, false, true, false], [false, true, false, true], [true, true, false, false], [false, false, true, true], [true, false, false, true]])` sollte `[[true, false, true, false, true], [false, true, true, false, false], [true, false, false, true, false], [false, true, false, true, true]]` zurückgeben.

```js
assert.deepEqual(transpose([[true, false, true, false], [false, true, false, true], [true, true, false, false], [false, false, true, true], [true, false, false, true]]), [[true, false, true, false, true], [false, true, true, false, false], [true, false, false, true, false], [false, true, false, true, true]]);
```

# --seed--

## --seed-contents--

```js
function transpose(matrix) {

  return matrix;
}
```

# --solutions--

```js
function transpose(matrix) {

  return matrix[0].map((_, i) => matrix.map(row => row[i]));
}
```
