---
id: 69e2383af7832c8032603b94
title: "Aufgabe 278: Kaffee-Bestellparser"
challengeType: 28
dashedName: challenge-278
---

# --description--

Gegeben ist ein String für eine Kaffee-Bestellung. Identifizieren Sie alle Menüartikel und geben Sie eine formatierte Bestellung zurück.

Verwenden Sie die folgenden Menüartikel und Preise:

| Artikel | Preis | 
| - | - |
| `"cold brew"` | 4,50 $ |
| `"oat latte"` | 5,00 $ |
| `"cappuccino"` | 4,75 $ |
| `"espresso"` | 3,00 $ |
| `"vanilla syrup"` | 0,75 $ |
| `"caramel drizzle"` | 0,60 $ |
| `"extra shot"` | 0,50 $ |
| `"oat milk"` | 0,75 $ |
| `"cream"` | 0,75 $ |

Geben Sie einen String zurück, in dem die gefundenen Artikel durch `" + "` verbunden sind, gefolgt von einem Doppelpunkt und einem Leerzeichen (`": "`), sowie dem Gesamtpreis.

Zum Beispiel, bei Eingabe `"I'd like an oat latte with vanilla syrup and an extra shot please."`, geben Sie `"oat latte + vanilla syrup + extra shot: $6.25"` zurück.

Die Artikel sollten in der Reihenfolge erscheinen, wie sie im Menü stehen, und der Gesamtpreis muss immer zwei Dezimalstellen haben.

# --hints--

`formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please.")` sollte `"oat latte + vanilla syrup + extra shot: $6.25"` zurückgeben.

```js
assert.equal(formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25");
```

`formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")` sollte `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"` zurückgeben.

```js
assert.equal(formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85");
```

`formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot.")` sollte `"cold brew + extra shot + cream: $5.75"` zurückgeben.

```js
assert.equal(formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75");
```

`formatCoffeeOrder("Just an espresso please.")` sollte `"espresso: $3.00"` zurückgeben.

```js
assert.equal(formatCoffeeOrder("Just an espresso please."), "espresso: $3.00");
```

`formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")` sollte `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"` zurückgeben.

```js
assert.equal(formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60");
```

# --seed--

## --seed-contents--

```js
function formatCoffeeOrder(order) {

  return order;
}
```

# --solutions--

```js
function formatCoffeeOrder(order) {
  const menu = [
    { name: "cold brew", price: 4.50 },
    { name: "oat latte", price: 5.00 },
    { name: "cappuccino", price: 4.75 },
    { name: "espresso", price: 3.00 },
    { name: "vanilla syrup", price: 0.75 },
    { name: "caramel drizzle", price: 0.60 },
    { name: "extra shot", price: 0.50 },
    { name: "oat milk", price: 0.75 },
    { name: "cream", price: 0.75 },
  ];

  const found = menu.filter(item => order.toLowerCase().includes(item.name));
  const total = found.reduce((sum, item) => sum + item.price, 0);
  return `${found.map(item => item.name).join(" + ")}: $${total.toFixed(2)}`;
}
```
