---
id: 69e2383af7832c8032603b8f
title: "Aufgabe 273: ISBN-13-Validator"
challengeType: 28
dashedName: challenge-273
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob es sich um eine gültige ISBN-13-Nummer handelt.

Eine gültige ISBN-13:

- Enthält nur Ziffern und Bindestriche
- Hat genau 13 Ziffern, nachdem die Bindestriche entfernt wurden
- Besteht die folgende Prüfung:
  1. Multiplizieren Sie jede Ziffer abwechselnd mit 1 oder 3 (die erste Ziffer mit 1, die zweite mit 3, die dritte mit 1 usw.).
  2. Die Summe der Ergebnisse muss durch 10 teilbar sein.

# --hints--

`isValidIsbn13("9780306406157")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
