---
id: 69e2383af7832c8032603b90
title: "Aufgabe 274: Älteste Person"
challengeType: 28
dashedName: challenge-274
---

# --description--

Gegeben ist ein Array von Objekten, von denen jedes eine `"name"`- und eine `"age"`-Eigenschaft besitzt. Geben Sie ein Array zurück, das den Namen der ältesten Person enthält.

Wenn mehrere Personen das höchste Alter teilen, geben Sie alle ihre Namen in der Reihenfolge zurück, in der sie in der Eingabe erscheinen.

# --hints--

`getOldest([{ name: "Brenda", age: 40 }])` sollte `["Brenda"]` zurückgeben.

```js
assert.deepEqual(getOldest([{ name: "Brenda", age: 40 }]), ["Brenda"]);
```

`getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }])` sollte `["Alice"]` zurückgeben.

```js
assert.deepEqual(getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }]), ["Alice"]);
```

`getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }])` sollte `["Bill", "Carol"]` zurückgeben.

```js
assert.deepEqual(getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }]), ["Bill", "Carol"]);
```

`getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }])` sollte `["George", "Holly", "Zach"]` zurückgeben.

```js
assert.deepEqual(getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }]), ["George", "Holly", "Zach"]);
```

# --seed--

## --seed-contents--

```js
function getOldest(people) {

  return people;
}
```

# --solutions--

```js
function getOldest(people) {
  const maxAge = Math.max(...people.map(p => p.age));
  return people.filter(p => p.age === maxAge).map(p => p.name);
}
```
