---
id: 69e2383af7832c8032603b92
title: "Aufgabe 276: Fehlendes Element"
challengeType: 28
dashedName: challenge-276
---

# --description--

Gegeben ist ein Array ganzer Zahlen, das aufsteigend sortiert ist, bis auf ein Element, das an der falschen Stelle steht. Geben Sie den Index dieses Elements zurück.

- Wenn mehr als ein Element als fehl am Platz betrachtet werden könnte, geben Sie den Index des ersten solchen Elements zurück.

# --hints--

`findOffender([1, 6, 2, 3, 4, 5])` sollte `1` zurückgeben.

```js
assert.equal(findOffender([1, 6, 2, 3, 4, 5]), 1);
```

`findOffender([1, 2, 3, 5, 4, 5])` sollte `3` zurückgeben.

```js
assert.equal(findOffender([1, 2, 3, 5, 4, 5]), 3);
```

`findOffender([2, 1])` sollte `0` zurückgeben.

```js
assert.equal(findOffender([2, 1]), 0);
```

`findOffender([2, 4, 1, 6, 8])` sollte `2` zurückgeben.

```js
assert.equal(findOffender([2, 4, 1, 6, 8]), 2);
```

`findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])` sollte `6` zurückgeben.

```js
assert.equal(findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6);
```

# --seed--

## --seed-contents--

```js
function findOffender(arr) {

  return arr;
}
```

# --solutions--

```js
function findOffender(arr) {
  for (let i = 0; i < arr.length; i++) {
    const without = [...arr.slice(0, i), ...arr.slice(i + 1)];
    if (without.every((n, j) => j === 0 || without[j - 1] <= n)) return i;
  }
}
```
