---
id: 69f35a5bb823ed620fcb7cbc
title: "Aufgabe 283: String-Zipper"
challengeType: 28
dashedName: challenge-283
---

# --description--

Gegeben sind zwei Strings. Geben Sie einen neuen String zurück, der die Zeichen der beiden Strings abwechselnd enthält. Wenn ein String länger ist, hängen Sie die verbleibenden Zeichen am Ende an.

Beginnen Sie mit dem ersten Zeichen des ersten Strings.

# --hints--

`zipStrings("abc", "123")` sollte `"a1b2c3"` zurückgeben.

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` sollte `"abcdefghijklmnopqrstuvwxyz"` zurückgeben.

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` sollte `"dnaiyght"` zurückgeben.

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` sollte `"pjyatvhaosncript"` zurückgeben.

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` sollte `"freeCodeCamp"` zurückgeben.

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
