---
id: 69f35a5bb823ed620fcb7cbd
title: "Aufgabe 284: I vor E"
challengeType: 28
dashedName: challenge-284
---

# --description--

Gegeben ist ein Wort oder Satz. Geben Sie eine korrigierte Version zurück, bei der jedes Wort der Regel „I vor E, außer nach C“ folgt.

- Wenn ein Wort `"ei"` enthält, das nicht von `"c"` vorangegangen wird, ersetzen Sie es durch `"ie"`.
- Wenn ein Wort `"ie"` enthält, das von `"c"` vorangegangen wird, ersetzen Sie es durch `"ei"`.
- Alle anderen Wörter bleiben unverändert.

# --hints--

`iBeforeE("beleive")` sollte `"believe"` zurückgeben.

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")` sollte `"receive"` zurückgeben.

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")` sollte `"we received a brief"` zurückgeben.

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")` sollte `"she believed the friendly niece could perceive the grief"` zurückgeben.

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")` sollte `"we received relief after the thief gave us a brief piece of fierce deceit"` zurückgeben.

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```
