---
id: 69f35a5bb823ed620fcb7cba
title: "Aufgabe 281: Bingo-Bereich"
challengeType: 28
dashedName: challenge-281
---

# --description--

Gegeben ist ein Bingo-Buchstabe. Geben Sie den Zahlenbereich zurück, der mit diesem Buchstaben verknüpft ist.

| Buchstabe | Zahlenbereich |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

Geben Sie ein Array mit allen Zahlen im Bereich vom kleinsten bis zum größten Wert zurück.

# --hints--

`getBingoRange("B")` sollte `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15]` zurückgeben.

```js
assert.deepEqual(getBingoRange("B"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15]);
```

`getBingoRange("I")` sollte `[16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30]` zurückgeben.

```js
assert.deepEqual(getBingoRange("I"), [16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30]);
```

`getBingoRange("N")` sollte `[31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45]` zurückgeben.

```js
assert.deepEqual(getBingoRange("N"), [31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45]);
```

`getBingoRange("G")` sollte `[46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60]` zurückgeben.

```js
assert.deepEqual(getBingoRange("G"), [46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60]);
```

`getBingoRange("O")` sollte `[61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75]` zurückgeben.

```js
assert.deepEqual(getBingoRange("O"), [61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75]);
```

# --seed--

## --seed-contents--

```js
function getBingoRange(letter) {

  return letter;
}
```

# --solutions--

```js
function getBingoRange(letter) {
  const ranges = { B: [1, 15], I: [16, 30], N: [31, 45], G: [46, 60], O: [61, 75] };
  const [start, end] = ranges[letter];
  return Array.from({ length: end - start + 1 }, (_, i) => start + i);
}
```
