---
id: 69f35a5bb823ed620fcb7cbe
title: "Aufgabe 285: Treffzeit"
challengeType: 28
dashedName: challenge-285
---

# --description--

Gegeben ist ein 3D-Array, das Verfügbarkeitsfenster für mehrere Personen darstellt. Geben Sie die früheste Zeit zurück, zu der alle eine Stunde frei haben. Falls keine solche Zeit existiert, geben Sie `"None"` zurück.

- Die Verfügbarkeit jeder Person ist ein Array von `[start, end]` Ganzzahlpaaren in 24-Stunden-Zeit. Zum Beispiel würde `[10, 12]` bedeuten, dass die Person von 10 bis 12 verfügbar ist. Startzeiten liegen zwischen 0 und 23, Endzeiten zwischen 1 und 24.

Zum Beispiel, gegeben:

```json
[
  [[10, 12], [15, 16]], // person 1
  [[11, 14], [15, 16]]  // person 2
]
```

Geben Sie `11` zurück, den Beginn ihrer ersten gemeinsamen freien Stunde.

# --hints--

`getMeetingTime([[[10, 12], [15, 16]], [[11, 14], [15, 16]]])` sollte `11` zurückgeben.

```js
assert.equal(getMeetingTime([[[10, 12], [15, 16]], [[11, 14], [15, 16]]]), 11);
```

`getMeetingTime([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]])` sollte `13` zurückgeben.

```js
assert.equal(getMeetingTime([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]]), 13);
```

`getMeetingTime([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]])` sollte `9` zurückgeben.

```js
assert.equal(getMeetingTime([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]]), 9);
```

`getMeetingTime([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]])` sollte `None` zurückgeben.

```js
assert.equal(getMeetingTime([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]]), "None");
```

`getMeetingTime([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]])` sollte `21` zurückgeben.

```js
assert.equal(getMeetingTime([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]]), 21);
```

# --seed--

## --seed-contents--

```js
function getMeetingTime(availability) {

  return availability;
}
```

# --solutions--

```js
function getMeetingTime(availability) {
  for (let hour = 0; hour < 24; hour++) {
    const allFree = availability.every(person =>
      person.some(([start, end]) => start <= hour && hour + 1 <= end)
    );
    if (allFree) return hour;
  }
  return "None";
}
```
