---
id: 69f35a5bb823ed620fcb7cbb
title: "Aufgabe 282: Schlafschuld"
challengeType: 28
dashedName: challenge-282
---

# --description--

Gegeben ist ein Array mit den Stunden, die Sie jede Nacht bis heute geschlafen haben, und eine Zielanzahl von Stunden pro Nacht. Geben Sie zurück, wie viele Stunden Schlaf Sie heute Nacht benötigen, um Ihre Schlafschuld auszugleichen.

- Beziehen Sie die Stunden von heute Nacht in die insgesamt benötigte Zeit zum Aufholen mit ein.
- Wenn Sie bereits genug geschlafen haben, um das heutige Ziel oder mehr zu erreichen, geben Sie `0` zurück.

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)` sollte `20` zurückgeben.

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)` sollte `10` zurückgeben.

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)` sollte `4` zurückgeben.

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)` sollte `0` zurückgeben.

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)` sollte `0` zurückgeben.

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
