---
id: 69f35a5bb823ed620fcb7cb9
title: "Aufgabe 280: Mongo-ID-Datum"
challengeType: 28
dashedName: challenge-280
---

# --description--

Gegeben ist ein MongoDB-ID-String. Geben Sie die Erstellungszeit als ISO-8601-String zurück.

- Eine MongoDB-ID ist ein 24-stelliger Hex-String. Die ersten 8 Zeichen stellen einen Unix-Zeitstempel (in Sekunden) dar, der als Basis-16-Ganzzahl kodiert ist.

Zum Beispiel hat `"6a094b50bcf86cd799439011"` einen Zeitstempel von `"6a094b50"` in Hex, was `1778994000` in Dezimal entspricht und eine Erstellungszeit von `"2026-05-17T05:00:00.000Z"` darstellt.

# --hints--

`mongoIdToDate("6a094b50bcf86cd799439011")` sollte `"2026-05-17T05:00:00.000Z"` zurückgeben.

```js
assert.equal(mongoIdToDate("6a094b50bcf86cd799439011"), "2026-05-17T05:00:00.000Z");
```

`mongoIdToDate("695344eb1f4a4c1123042128")` sollte `"2025-12-30T03:20:11.000Z"` zurückgeben.

```js
assert.equal(mongoIdToDate("695344eb1f4a4c1123042128"), "2025-12-30T03:20:11.000Z");
```

`mongoIdToDate("386da62df34123ac54617e56")` sollte `"2000-01-01T07:01:01.000Z"` zurückgeben.

```js
assert.equal(mongoIdToDate("386da62df34123ac54617e56"), "2000-01-01T07:01:01.000Z");
```

`mongoIdToDate("69f571c3d7711807afd3dd55")` sollte `"2026-05-02T03:38:43.000Z"` zurückgeben.

```js
assert.equal(mongoIdToDate("69f571c3d7711807afd3dd55"), "2026-05-02T03:38:43.000Z");
```

`mongoIdToDate("68adce01c0e1144d0a90295a")` sollte `"2025-08-26T15:08:49.000Z"` zurückgeben.

```js
assert.equal(mongoIdToDate("68adce01c0e1144d0a90295a"), "2025-08-26T15:08:49.000Z");
```

# --seed--

## --seed-contents--

```js
function mongoIdToDate(id) {

  return id;
}
```

# --solutions--

```js
function mongoIdToDate(id) {
  const timestamp = parseInt(id.slice(0, 8), 16);
  return new Date(timestamp * 1000).toISOString();
}
```
