---
id: 69f8c998d78ad3171a0713be
title: "Aufgabe 291: FizzBuzz zählen"
challengeType: 28
dashedName: challenge-291
---

# --description--

Gegeben sind eine Start- und eine Endzahl. Zählen Sie, wie oft „fizz“ und „buzz“ im Bereich (einschließlich) vorkommen.

- Zahlen, die durch 3 teilbar sind, zählen als fizz.
- Zahlen, die durch 5 teilbar sind, zählen als buzz.
- Zahlen, die sowohl durch 3 als auch durch 5 teilbar sind, zählen sowohl als fizz als auch als buzz.

Geben Sie ein Objekt oder Wörterbuch mit den Zählungen im Format `{ fizz, buzz }` zurück.

# --hints--

`fizzBuzzCount(1, 11)` sollte `{fizz: 3, buzz: 2}` zurückgeben.

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)` sollte `{fizz: 9, buzz: 6}` zurückgeben.

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)` sollte `{fizz: 26, buzz: 16}` zurückgeben.

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)` sollte `{fizz: 207, buzz: 125}` zurückgeben.

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)` sollte `{fizz: 4074, buzz: 2444}` zurückgeben.

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```
