---
id: 69f8c998d78ad3171a0713bc
title: "Aufgabe 289: Summe der Differenzen"
challengeType: 28
dashedName: challenge-289
---

# --description--

Gegeben ist ein Array von Zahlen. Geben Sie die Summe der Differenzen zwischen jeder Zahl und der unmittelbar folgenden Zahl zurück.

Zum Beispiel, gegeben `[1, 3, 4]`, geben Sie `3` zurück (2 + 1).

# --hints--

`sumOfDifferences([1, 3, 4])` sollte `3` zurückgeben.

```js
assert.equal(sumOfDifferences([1, 3, 4]), 3);
```

`sumOfDifferences([5, -3, 3, 9, 10])` sollte `5` zurückgeben.

```js
assert.equal(sumOfDifferences([5, -3, 3, 9, 10]), 5);
```

`sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7])` sollte `-16` zurückgeben.

```js
assert.equal(sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16);
```

`sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])` sollte `25` zurückgeben.

```js
assert.equal(sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25);
```

# --seed--

## --seed-contents--

```js
function sumOfDifferences(arr) {

  return arr;
}
```

# --solutions--

```js
function sumOfDifferences(arr) {
  return arr.slice(0, -1).reduce((sum, n, i) => sum + (arr[i + 1] - n), 0);
}
```
