---
id: 69f8c998d78ad3171a0713b9
title: "Aufgabe 286: Offene Issues"
challengeType: 28
dashedName: challenge-286
---

# --description--

Gegeben sind ein Array von Issue-Nummern und ein weiteres Array von Pull-Request-(PR-)Nummern. Geben Sie ein Array der Issues zurück, die nach dem Zusammenführen aller PRs noch offen sind.

- Ein PR schließt ein Issue, wenn deren Ziffern eine Rotation voneinander sind. Zum Beispiel würde Issue 123 durch PR 231 oder 312 geschlossen.
- Ein PR schließt kein Issue mit exakt derselben Nummer. Zum Beispiel schließt PR 123 nicht Issue 123. Ein Issue mit nur derselben Zahl kann also nicht geschlossen werden.
- Beide Nummern können führende Nullen verlieren. Zum Beispiel würde PR 201 Issue 12 (012, eine Rotation von 201) schließen. Ebenso würde Issue 201 durch PR 12 geschlossen.

Geben Sie die verbleibenden offenen Issues in der Reihenfolge zurück, in der sie angegeben wurden.

# --hints--

`getOpenIssues([123, 234], [231])` sollte `[234]` zurückgeben.

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` sollte `[345, 16]` zurückgeben.

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` sollte `[456, 332, 15]` zurückgeben.

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` sollte `[115, 296, 24]` zurückgeben.

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` sollte `[95, 395, 754, 296, 709, 237, 1802]` zurückgeben.

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
