---
id: 69f8c998d78ad3171a0713bd
title: "Aufgabe 290: Pizza-Party"
challengeType: 28
dashedName: challenge-290
---

# --description--

Gegeben ist ein Array mit den heute pro Person gearbeiteten Stunden. Geben Sie die Anzahl der Pizzen zurück, die für eine Pizza-Party bestellt werden müssen.

- Teilen Sie die gearbeiteten Stunden jeder Person durch 3, um die Anzahl der Pizzastücke zu erhalten.
- Sie können keine halben Stücke essen, also runden Sie die Anzahl der Stücke jeder Person auf die nächste ganze Zahl auf.
- Jede Person bekommt mindestens zwei Stücke.
- Jede Pizza hat 8 Stücke. Runden Sie die Gesamtzahl der Pizzen auf die nächste ganze Pizza auf.

# --hints--

`getPizzasToOrder([8, 8, 8])` sollte `2` zurückgeben.

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` sollte `3` zurückgeben.

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` sollte `2` zurückgeben.

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` sollte `3` zurückgeben.

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` sollte `1` zurückgeben.

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` sollte `5` zurückgeben.

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
