---
id: 69f8c998d78ad3171a0713bb
title: "Aufgabe 288: Geheime Zahl"
challengeType: 28
dashedName: challenge-288
---

# --description--

Gegeben sind eine geheime Zahl und eine Vermutung. Bestimmen Sie, ob die Vermutung korrekt ist.

Geben Sie zurück:

- `"higher"`, wenn die geheime Zahl größer als die Vermutung ist.
- `"lower"`, wenn die geheime Zahl kleiner als die Vermutung ist.
- `"you got it!"`, wenn die Vermutung korrekt ist.

# --hints--

`guessNumber(50, 30)` sollte `"higher"` zurückgeben.

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` sollte `"lower"` zurückgeben.

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` sollte `"you got it!"` zurückgeben.

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` sollte `"higher"` zurückgeben.

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` sollte `"lower"` zurückgeben.

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` sollte `"you got it!"` zurückgeben.

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
