---
id: 69f8c998d78ad3171a0713bf
title: "Aufgabe 292: Breiteres Seitenverhältnis"
challengeType: 28
dashedName: challenge-292
---

# --description--

Gegeben sind zwei Strings für unterschiedliche Bildmaße. Geben Sie das Seitenverhältnis des Bildes mit dem größeren Breite-zu-Höhe-Verhältnis zurück.

- Die gegebenen Strings haben das Format `"WxH"`, zum Beispiel `"1920x1080"`.
- Das Seitenverhältnis ist das Verhältnis von Breite zu Höhe, gekürzt auf die kleinsten ganzen Zahlen. Zum Beispiel wird `"1920x1080"` zu `"16:9"` gekürzt.
- Geben Sie einen String im Format `"W:H"` zurück, zum Beispiel `"16:9"`.

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` sollte `"16:9"` zurückgeben.

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` sollte `"4:3"` zurückgeben.

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` sollte `"2:1"` zurückgeben.

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` sollte `"9:16"` zurückgeben.

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` sollte `"43:18"` zurückgeben.

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` sollte `"2089:8539"` zurückgeben.

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
