---
id: 69f8c998d78ad3171a0713ba
title: "Aufgabe 287: Römische Zahl Korrektur"
challengeType: 28
dashedName: challenge-287
---

# --description--

Gegeben ist ein String mit fehlerhaften römischen Zahlen. Geben Sie den Wert in der standardmäßigen römischen Zahlenschreibweise zurück.

Die Eingabe verwendet nur additive Notation, das heißt, jedes Symbol addiert seinen Wert zur Gesamtsumme. Zur Erinnerung, hier sind die Symbole und ihre Werte:

| Symbol | Wert  |
| ------ | ----- |
| `"I"`  | 1     |
| `"V"`  | 5     |
| `"X"`  | 10    |
| `"L"`  | 50    |
| `"C"`  | 100   |
| `"D"`  | 500   |
| `"M"`  | 1000  |

Beim erneuten Kodieren verwenden Sie an jedem Schritt das größtmögliche Symbol und bei Bedarf die subtraktiven Paare (`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`).

# --hints--

`fixNumerals("XIIIII")` sollte `"XV"` zurückgeben.

```js
assert.equal(fixNumerals("XIIIII"), "XV");
```

`fixNumerals("IIIILX")` sollte `"LXIV"` zurückgeben.

```js
assert.equal(fixNumerals("IIIILX"), "LXIV");
```

`fixNumerals("XXVVVIIIII")` sollte `"XL"` zurückgeben.

```js
assert.equal(fixNumerals("XXVVVIIIII"), "XL");
```

`fixNumerals("MDCCLXXXXVIIII")` sollte `"MDCCXCIX"` zurückgeben.

```js
assert.equal(fixNumerals("MDCCLXXXXVIIII"), "MDCCXCIX");
```

`fixNumerals("IIIIVVVVXXXXLLLLCCDD")` sollte `"MCDLXIV"` zurückgeben.

```js
assert.equal(fixNumerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV");
```

`fixNumerals("ILCDMIVDIIXLCVCXDL")` sollte `"MMCMLXXXIV"` zurückgeben.

```js
assert.equal(fixNumerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV");
```

# --seed--

## --seed-contents--

```js
function fixNumerals(str) {

  return str;
}
```

# --solutions--

```js
function fixNumerals(str) {
  const values = { I: 1, V: 5, X: 10, L: 50, C: 100, D: 500, M: 1000 };
  const encoding = [
    [1000, "M"], [900, "CM"], [500, "D"], [400, "CD"],
    [100, "C"], [90, "XC"], [50, "L"], [40, "XL"],
    [10, "X"], [9, "IX"], [5, "V"], [4, "IV"], [1, "I"]
  ];

  let total = [...str].reduce((sum, c) => sum + values[c], 0);

  let result = "";
  for (const [value, symbol] of encoding) {
    while (total >= value) {
      result += symbol;
      total -= value;
    }
  }
  return result;
}
```
