---
id: 69f90c1329a94b37e2a2086c
title: "Aufgabe 293: Beste Hand"
challengeType: 28
dashedName: challenge-293
---

# --description--

Gegeben ist ein Array aus fünf Strings, die Spielkarten darstellen. Geben Sie den Namen der besten Hand zurück.

- Jede Karte wird als zweistelliger String dargestellt: der Rang gefolgt vom Symbol, `"2h"` zum Beispiel.
  - Die Ränge, von niedrig bis hoch, sind: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"` und `"A"`.
  - Die Symbole sind: `"h"`, `"d"`, `"c"` und `"s"`.
- Asse (`"A"`) können in einer Straße sowohl als hoch als auch als niedrig verwendet werden.

Die Hände, sortiert von der schlechtesten bis zur besten, sind:

| Name | Beschreibung |
| - | - |
| `"High Card"` | Kein Paar oder besser |
| `"Pair"` | Zwei Karten gleichen Rangs |
| `"Two Pair"` | Zwei Karten gleichen Rangs und zwei Karten eines anderen Rangs |
| `"Three of a Kind"` | Drei Karten gleichen Rangs |
| `"Straight"` | Fünf aufeinanderfolgende Ränge |
| `"Flush"` | Fünf Karten desselben Symbols |
| `"Full House"` | Drei Karten gleichen Rangs und zwei Karten eines anderen Rangs |
| `"Four of a Kind"` | Vier Karten gleichen Rangs |
| `"Straight Flush"` | Fünf aufeinanderfolgende Ränge desselben Symbols |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` desselben Symbols |

Geben Sie den Namen der besten Hand zurück.

# --hints--

`getBestHand(["7s", "7h", "7d", "2c", "5h"])` sollte `"Three of a Kind"` zurückgeben.

```js
assert.equal(getBestHand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind");
```

`getBestHand(["Ks", "Kh", "Kd", "4s", "4h"])` sollte `"Full House"` zurückgeben.

```js
assert.equal(getBestHand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House");
```

`getBestHand(["2h", "5h", "7h", "9h", "Jh"])` sollte `"Flush"` zurückgeben.

```js
assert.equal(getBestHand(["2h", "5h", "7h", "9h", "Jh"]), "Flush");
```

`getBestHand(["As", "Ah", "Ad", "Ac", "Kh"])` sollte `"Four of a Kind"` zurückgeben.

```js
assert.equal(getBestHand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind");
```

`getBestHand(["Ts", "Th", "9d", "9c", "8h"])` sollte `"Two Pair"` zurückgeben.

```js
assert.equal(getBestHand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair");
```

`getBestHand(["9c", "8c", "7c", "6c", "5c"])` sollte `"Straight Flush"` zurückgeben.

```js
assert.equal(getBestHand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush");
```

`getBestHand(["As", "Kh", "Jd", "8c", "5h"])` sollte `"High Card"` zurückgeben.

```js
assert.equal(getBestHand(["As", "Kh", "Jd", "8c", "5h"]), "High Card");
```

`getBestHand(["As", "2h", "3d", "4c", "5h"])` sollte `"Straight"` zurückgeben.

```js
assert.equal(getBestHand(["As", "2h", "3d", "4c", "5h"]), "Straight");
```

`getBestHand(["Ts", "Th", "7c", "6d", "5h"])` sollte `"Pair"` zurückgeben.

```js
assert.equal(getBestHand(["Ts", "Th", "7c", "6d", "5h"]), "Pair");
```

`getBestHand(["As", "Ks", "Qs", "Js", "Ts"])` sollte `"Royal Flush"` zurückgeben.

```js
assert.equal(getBestHand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush");
```

# --seed--

## --seed-contents--

```js
function getBestHand(cards) {

  return cards;
}
```

# --solutions--

```js
function getBestHand(cards) {
  const rankOrder = "23456789TJQKA";

  const ranks = cards.map(c => c[0]);
  const suits = cards.map(c => c[1]);

  const isFlush = suits.every(s => s === suits[0]);

  const rankIndices = ranks.map(r => rankOrder.indexOf(r)).sort((a, b) => a - b);

  function isStraight(indices) {
    const isRegular = indices.every((r, i) => i === 0 || r === indices[i - 1] + 1);
    const isLowAce = JSON.stringify(indices) === JSON.stringify([0, 1, 2, 3, 12]);
    return isRegular || isLowAce;
  }

  const straight = isStraight(rankIndices);

  const counts = {};
  ranks.forEach(r => counts[r] = (counts[r] || 0) + 1);
  const freq = Object.values(counts).sort((a, b) => b - a);

  if (isFlush && JSON.stringify(rankIndices) === JSON.stringify([8, 9, 10, 11, 12])) return "Royal Flush";
  if (isFlush && straight) return "Straight Flush";
  if (freq[0] === 4) return "Four of a Kind";
  if (freq[0] === 3 && freq[1] === 2) return "Full House";
  if (isFlush) return "Flush";
  if (straight) return "Straight";
  if (freq[0] === 3) return "Three of a Kind";
  if (freq[0] === 2 && freq[1] === 2) return "Two Pair";
  if (freq[0] === 2) return "Pair";
  return "High Card";
}
```
