---
id: 69f90c1329a94b37e2a2086d
title: "Aufgabe 294: Kombinationen von runden Klammern"
challengeType: 28
dashedName: challenge-294
---

# --description--

Gegeben ist eine ganze Zahl, `n`. Geben Sie die Anzahl der gültigen Kombinationen von `n` Paaren runder Klammern zurück.

- Eine gültige Kombination ist ein String, bei dem jede öffnende runde Klammer eine entsprechende schließende runde Klammer hat und keine schließende runde Klammer vor ihrer passenden öffnenden runden Klammer erscheint.

Zum Beispiel gibt es bei `2` 2 gültige Kombinationen:

```json
(())
()()
```

# --hints--

`getCombinations(2)` sollte `2` zurückgeben.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` sollte `5` zurückgeben.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` sollte `42` zurückgeben.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` sollte `1430` zurückgeben.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` sollte `742900` zurückgeben.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
