---
id: 6a0dcc730cb92a616f86f0c5
title: "Aufgabe 301: Letzte Ladung"
challengeType: 28
dashedName: challenge-301
---

# --description--

Gegeben ist die Anzahl der verbleibenden Waschpulverportionen und ein Array, das angibt, wie viele Portionen Sie an den vorherigen Tagen verwendet haben. Geben Sie die Anzahl der vollen Tage zurück, die Sie mit dem verbleibenden Waschpulver noch auskommen.

Berechnen Sie Ihren durchschnittlichen Tagesverbrauch anhand der Verbrauchshistorie und gehen Sie davon aus, dass Sie diese Menge an Verbrauch für jeden weiteren Tag verwenden.

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` sollte `5` zurückgeben.

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` sollte `7` zurückgeben.

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` sollte `11` zurückgeben.

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` sollte `12` zurückgeben.

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` sollte `1` zurückgeben.

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
