---
id: 6a0dcd03ee4e68698080ef6c
title: "Aufgabe 308: Kreditkarten-Validator"
challengeType: 28
dashedName: challenge-308
---

# --description--

Gegeben ist ein String aus Ziffern für eine Kreditkartennummer. Bestimmen Sie, ob es sich um eine gültige Kartennummer handelt, indem Sie die folgende Methode anwenden:

- Beginnen Sie mit der zweitletzten Ziffer und verdoppeln Sie jede zweite Ziffer, die nach links folgt.
- Wenn das Verdoppeln einer Ziffer zu einer Zahl größer als 9 führt, ziehen Sie 9 ab.
- Addieren Sie alle Ziffern (verdoppelte und nicht verdoppelte).
- Wenn die Summe durch 10 teilbar ist, ist die Nummer gültig.

# --hints--

`isValidCard("4532015112830366")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` sollte `true` zurückgeben.

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` sollte `false` zurückgeben.

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
