---
id: 6a0dcd03ee4e68698080ef68
title: "Aufgabe 304: Reiseplananordnungen"
challengeType: 28
dashedName: challenge-304
---

# --description--

Gegeben ist ein Array mit mindestens zwei optionalen Stopps für einen Tagesausflug. Geben Sie die Anzahl der gültigen Reiseplananordnungen zurück.

Der Reiseplan enthält immer `"breakfast"`, `"lunch"` und `"dinner"`; diese werden nicht als Argumente übergeben. Die optionalen Stopps können an beliebiger Stelle im Reiseplan eingefügt werden, unter Beachtung der folgenden Regeln:

- `"breakfast"` steht immer an erster Stelle, mit mindestens einem Stopp vor `"lunch"`.
- `"lunch"` muss vor `"dinner"` erscheinen, mit mindestens einem Stopp dazwischen.
- Maximal ein optionaler Stopp darf nach `"dinner"` erscheinen.

Geben Sie die Anzahl der gültigen Anordnungen zurück.

# --hints--

`getItineraryCount(["library", "park"])` sollte `2` zurückgeben.

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])` sollte `18` zurückgeben.

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])` sollte `120` zurückgeben.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])` sollte `840` zurückgeben.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])` sollte `55440` zurückgeben.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
