---
id: 6a0dcc730cb92a616f86f0c1
title: "Aufgabe 297: Schema-Validator Teil 3"
challengeType: 28
dashedName: challenge-297
---

# --description--

Gegeben ist ein Objekt (JavaScript) oder ein Dictionary (Python). Bestimmen Sie, ob es dem folgenden Schema entspricht:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles
}
```

- Das Pipe-Symbol (`|`) bedeutet „oder“. `role` muss einer der aufgelisteten `Roles` Werte sein.
- Zusätzliche Schlüssel sind erlaubt

# --hints--

`isValidSchema({ username: "henry", posts: 0, verified: true, role: "staff" })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "henry", posts: 0, verified: true, role: "staff" }));
```

`isValidSchema({ username: "sara", posts: 45, verified: false, role: "creator", followers: 70 })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "sara", posts: 45, verified: false, role: "creator", followers: 70 }));
```

`isValidSchema({ username: "penelope", posts: 20, verified: true, role: "admin" })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "penelope", posts: 20, verified: true, role: "admin" }));
```

`isValidSchema({ username: "kevin", posts: 0, verified: false, role: "user" })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "kevin", posts: 0, verified: false, role: "user" }));
```

`isValidSchema({ username: "george", posts: 15, verified: true, role: "moderator" })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "george", posts: 15, verified: true, role: "moderator" }));
```

`isValidSchema({ username: "david", posts: 0, verified: false, role: "guest" })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "david", posts: 0, verified: false, role: "guest" }));
```

`isValidSchema({ username: "wendy", posts: 10, verified: true })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "wendy", posts: 10, verified: true }));
```

`isValidSchema({ username: "fabian", posts: 1, verified: true, role: true })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "fabian", posts: 1, verified: true, role: true }));
```

`isValidSchema({ username: 8, posts: 1, verified: true, role: "user" })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: 8, posts: 1, verified: true, role: "user" }));
```

`isValidSchema({ username: "penny", posts: "10", verified: true, role: "staff" })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "penny", posts: "10", verified: true, role: "staff" }));
```

`isValidSchema({ username: "john", posts: "1", verified: "true", role: "admin" })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "john", posts: "1", verified: "true", role: "admin" }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role)
  );
}
```
