---
id: 6a0dcd03ee4e68698080ef6b
title: "Aufgabe 307: Bebauungsvorschriften"
challengeType: 28
dashedName: challenge-307
---

# --description--

Gegeben ist ein 2D-Raster (Array von Arrays), das die Gebäudestruktur einer Stadt darstellt. Geben Sie die Koordinaten aller Gebäude zurück, die gegen Bebauungsvorschriften verstoßen.

Jede Zelle im Raster enthält eine der Bezeichnungen aus der untenstehenden Tabelle. Ein Gebäude verstößt gegen die Vorschriften, wenn eines seiner (bis zu) 4 Nachbarn, horizontal oder vertikal, ein Typ ist, an den es nicht angrenzen darf.

| Bezeichnung | Typ | Darf nicht angrenzen an |
| - | - | - |
| `"i"` | Industrie | `"R"`, `"I"` |
| `"A"` | Landwirtschaft | `"C"` |
| `"R"` | Wohngebiet | `"i"`, `"C"` |
| `"I"` | Institutionell | `"i"` |
| `"C"` | Gewerbe | `"R"`, `"A"` |
| `""` (leerer String) | unbebaut | keine Einschränkungen |

Geben Sie die Koordinaten aller Zellen mit Verstößen als Array von `[row, col]`-Paaren in beliebiger Reihenfolge zurück. Wenn keine Verstöße vorliegen, geben Sie ein leeres Array zurück.

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` sollte `[[0, 0], [0, 1]]` zurückgeben.

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` sollte `[[0, 1], [1, 1]]` zurückgeben.

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` sollte `[]` zurückgeben.

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` sollte `[[0, 1], [0, 2], [0, 3]]` zurückgeben.

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` sollte `[[2, 3], [2, 4], [3, 1], [3, 2]]` zurückgeben.

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
