---
id: 6a0dcd03ee4e68698080ef66
title: "Aufgabe 302: Jetlag"
challengeType: 28
dashedName: challenge-302
---

# --description--

Gegeben sind eine Abfahrtsstadt, eine Ankunftsstadt, eine Flugdauer in Stunden und eine Reiserichtung. Geben Sie die Anzahl der Jetlag-Stunden zurück, die der Reisende erlebt.

Die angegebenen Städte stammen aus der folgenden Liste mit ihrem UTC-Versatz:

| Stadt | Versatz |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Um die Jetlag-Stunden zu berechnen:

1. Ermitteln Sie die Zeitzonendifferenz in Stunden zwischen den beiden Städten.
2. Bestimmen Sie den Richtungsfaktor. Wenn die Reise `"east"` ist, beträgt er 1,5, andernfalls 1,0.
3. Berechnen Sie die Jetlag-Stunden mit der Formel: Zeitzonendifferenz + (Flugdauer * 0,1) * Richtungsfaktor

Geben Sie die Jetlag-Stunden auf eine Nachkommastelle gerundet zurück.

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` sollte `6.5` zurückgeben.

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` sollte `5.8` zurückgeben.

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` sollte `1.6` zurückgeben.

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` sollte `4.7` zurückgeben.

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` sollte `17.5` zurückgeben.

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` sollte `5.9` zurückgeben.

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` sollte `9.5` zurückgeben.

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
