---
id: 6a0dcd03ee4e68698080ef69
title: "Aufgabe 305: Ideen-Rankings"
challengeType: 28
dashedName: challenge-305
---

# --description--

Gegeben ist ein 2D-Array, bei dem jedes innere Array (in dieser Reihenfolge) einen Ideennamen, eine optimistische Schätzung, eine realistische Schätzung und eine pessimistische Schätzung (in Tagen) enthält. Geben Sie ein Array mit den Ideennamen zurück, sortiert nach der erwarteten Fertigstellungszeit, beginnend mit der kürzesten.

Berechnen Sie die erwartete Fertigstellungszeit für jede Idee mit der folgenden Formel:

- `expected = ((optimistic + 4 * realistic + pessimistic) / 6) * length of idea name`

# --hints--

`analyzeIdeas([["Add logging", 2, 5, 15], ["SEO optimization", 4, 8, 20], ["Fix bug", 1, 3, 5]])` sollte `["Fix bug", "Add logging", "SEO optimization"]` zurückgeben.

```js
assert.deepEqual(analyzeIdeas([["Add logging", 2, 5, 15], ["SEO optimization", 4, 8, 20], ["Fix bug", 1, 3, 5]]), ["Fix bug", "Add logging", "SEO optimization"]);
```

`analyzeIdeas([["Dark mode", 1, 3, 8], ["Real-time collaboration feature", 6, 12, 20], ["Add tooltip", 1, 2, 4]])` sollte `["Add tooltip", "Dark mode", "Real-time collaboration feature"]` zurückgeben.

```js
assert.deepEqual(analyzeIdeas([["Dark mode", 1, 3, 8], ["Real-time collaboration feature", 6, 12, 20], ["Add tooltip", 1, 2, 4]]), ["Add tooltip", "Dark mode", "Real-time collaboration feature"]);
```

`analyzeIdeas([["Update user profile page", 3, 7, 14], ["Add pagination", 2, 5, 10], ["Add tags", 2, 3, 6], ["Fix login bug", 1, 4, 8]])` sollte `["Add tags", "Fix login bug", "Add pagination", "Update user profile page"]` zurückgeben.

```js
assert.deepEqual(analyzeIdeas([["Update user profile page", 3, 7, 14], ["Add pagination", 2, 5, 10], ["Add tags", 2, 3, 6], ["Fix login bug", 1, 4, 8]]), ["Add tags", "Fix login bug", "Add pagination", "Update user profile page"]);
```

`analyzeIdeas([["Migrate database", 14, 25, 40], ["Add chat assistant", 8, 15, 24], ["Redesign onboarding flow", 3, 7, 13], ["Add language support", 6, 11, 18]])` sollte `["Redesign onboarding flow", "Add language support", "Add chat assistant", "Migrate database"]` zurückgeben.

```js
assert.deepEqual(analyzeIdeas([["Migrate database", 14, 25, 40], ["Add chat assistant", 8, 15, 24], ["Redesign onboarding flow", 3, 7, 13], ["Add language support", 6, 11, 18]]), ["Redesign onboarding flow", "Add language support", "Add chat assistant", "Migrate database"]);
```

`analyzeIdeas([["Add email notifications", 3, 7, 10], ["Migrate deployment flow", 6, 10, 16], ["Add push notifications", 2, 6, 10], ["Optimize continuous integration", 5, 8, 15], ["Analyze user patterns", 5, 10, 18], ["Create onboarding curriculum", 6, 15, 25]])` sollte `["Add push notifications", "Add email notifications", "Analyze user patterns", "Migrate deployment flow", "Optimize continuous integration", "Create onboarding curriculum"]` zurückgeben.

```js
assert.deepEqual(analyzeIdeas([["Add email notifications", 3, 7, 10], ["Migrate deployment flow", 6, 10, 16], ["Add push notifications", 2, 6, 10], ["Optimize continuous integration", 5, 8, 15], ["Analyze user patterns", 5, 10, 18], ["Create onboarding curriculum", 6, 15, 25]]), ["Add push notifications", "Add email notifications", "Analyze user patterns", "Migrate deployment flow", "Optimize continuous integration", "Create onboarding curriculum"]);
```

# --seed--

## --seed-contents--

```js
function analyzeIdeas(ideas) {

  return ideas;
}
```

# --solutions--

```js
function analyzeIdeas(ideas) {
  return ideas
    .map(([name, optimistic, realistic, pessimistic]) => ({
      name,
      expected: ((optimistic + 4 * realistic + pessimistic) / 6) * name.length
    }))
    .sort((a, b) => a.expected - b.expected)
    .map(idea => idea.name);
}
```
