---
id: 6a0dcc730cb92a616f86f0bf
title: "Aufgabe 295: Schema-Validator Teil 1"
challengeType: 28
dashedName: challenge-295
---

# --description--

Gegeben ist ein Objekt (JavaScript) oder ein Dictionary (Python). Bestimmen Sie, ob es dem folgenden Schema entspricht:

```json
{
  username: string
}
```

- Zusätzliche Schlüssel sind erlaubt

# --hints--

`isValidSchema({ username: "bob" })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
