---
id: 6a0dcc730cb92a616f86f0c3
title: "Aufgabe 299: Schema-Validator Teil 5"
challengeType: 28
dashedName: challenge-299
---

# --description--

Gegeben ist ein Objekt (JavaScript) oder ein Dictionary (Python). Bestimmen Sie, ob es dem folgenden Schema entspricht:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles,
  supporter?: boolean,
  badges: string[]
}
```

- Das Pipe-Symbol (`|`) bedeutet „oder“. `role` muss einer der aufgelisteten `Roles`-Werte sein.
- Das Fragezeichen (`?`) nach `supporter` bedeutet, dass das Feld optional ist, aber vom angegebenen Typ sein muss, falls es existiert.
- Die Klammern `[]` nach `string` bedeuten, dass `badges` ein Array von Strings sein sollte (oder leer).
- Zusätzliche Schlüssel sind erlaubt

# --hints--

`isValidSchema({ username: "gill", posts: 12, verified: false, role: "creator", supporter: false, badges: [ "early-adopter", "popular" ] })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "gill", posts: 12, verified: false, role: "creator", supporter: false, badges: [ "early-adopter", "popular" ] }));
```

`isValidSchema({ username: "tonya", posts: 299, verified: true, role: "moderator", supporter: true, badges: [ "streak-master", "veteran" ], followers: 1233 })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "tonya", posts: 299, verified: true, role: "moderator", supporter: true, badges: [ "streak-master", "veteran" ], followers: 1233 }));
```

`isValidSchema({ username: "zara", posts: 0, verified: false, role: "user", supporter: false, badges: [] })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "zara", posts: 0, verified: false, role: "user", supporter: false, badges: [] }));
```

`isValidSchema({ username: "nicole", posts: 65, verified: true, role: "admin", supporter: false, badges: [ "first-post", 18 ] })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "nicole", posts: 65, verified: true, role: "admin", supporter: false, badges: [ "first-post", 18 ] }));
```

`isValidSchema({ username: "tim", posts: 25, verified: true, role: "staff", supporter: false })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "tim", posts: 25, verified: true, role: "staff", supporter: false }));
```

`isValidSchema({ username: "charlie", posts: 0, verified: false, role: "user", supporter: "no", badges: [ "first-post", "anniversary" ] })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "charlie", posts: 0, verified: false, role: "user", supporter: "no", badges: [ "first-post", "anniversary" ] }));
```

`isValidSchema({ username: "wanda", posts: 15, verified: true, role: "friend", supporter: true, badges: [ "popular" ] })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "wanda", posts: 15, verified: true, role: "friend", supporter: true, badges: [ "popular" ] }));
```

`isValidSchema({ username: "guy", posts: 5, verified: "false", role: "staff", supporter: true, badges: [ "helper" ] })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "guy", posts: 5, verified: "false", role: "staff", supporter: true, badges: [ "helper" ] }));
```

`isValidSchema({ username: "carrie", verified: true, role: "moderator", supporter: true, badges: [ "helper", "sharer" ] })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "carrie", verified: true, role: "moderator", supporter: true, badges: [ "helper", "sharer" ] }));
```

`isValidSchema({ username: true, posts: 75, verified: true, role: "creator", supporter: true, badges: [ "veteran" ] })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: true, posts: 75, verified: true, role: "creator", supporter: true, badges: [ "veteran" ] }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role) &&
    (obj.supporter === undefined || typeof obj.supporter === 'boolean') &&
    Array.isArray(obj.badges) && obj.badges.every(b => typeof b === 'string')
  );
}
```
