---
id: 6a0dcc730cb92a616f86f0c0
title: "Aufgabe 296: Schema-Validator Teil 2"
challengeType: 28
dashedName: challenge-296
---

# --description--

Gegeben ist ein Objekt (JavaScript) oder ein Dictionary (Python). Bestimmen Sie, ob es dem folgenden Schema entspricht:

```json
{
  username: string,
  posts: number,
  verified: boolean
}
```

- Zusätzliche Schlüssel sind erlaubt

# --hints--

`isValidSchema({ username: "alice", posts: 10, verified: false })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "alice", posts: 10, verified: false }));
```

`isValidSchema({ username: "carol", posts: 15, verified: true, followers: 25 })` sollte `true` zurückgeben.

```js
assert.isTrue(isValidSchema({ username: "carol", posts: 15, verified: true, followers: 25 }));
```

`isValidSchema({ username: "frank", posts: "21", verified: true })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "frank", posts: "21", verified: true }));
```

`isValidSchema({ username: "sam", posts: 17, verified: "false" })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "sam", posts: 17, verified: "false" }));
```

`isValidSchema({ username: "bill", verified: true })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "bill", verified: true }));
```

`isValidSchema({ username: "fred", verified: true })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: "fred", verified: true }));
```

`isValidSchema({ username: 5, posts: 10, verified: true })` sollte `false` zurückgeben.

```js
assert.isFalse(isValidSchema({ username: 5, posts: 10, verified: true }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean'
  );
}
```
