---
id: 6a0dcd03ee4e68698080ef67
title: "Aufgabe 303: Mitbewohner"
challengeType: 28
dashedName: challenge-303
---

# --description--

Gegeben ist ein Array von Personen und deren Mitbewohnergruppe. Geben Sie die Zimmerbelegungen für einen Hotelaufenthalt zurück, wobei folgende Regeln gelten:

- Jede Person hat eine `name` und eine `group` Eigenschaft:

```json
[
  { "name": "Alice", "group": "A" },
  { "name": "Bob", "group": "B" },
  { "name": "Carol", "group": "A" }
]
```

- Personen können nur mit jemandem aus derselben Gruppe ein Zimmer teilen und werden in der Reihenfolge gepaart, in der sie angegeben sind.
- Geben Sie ein Array von Strings zurück, bei dem die Namen für ein geteiltes Zimmer durch `" and "` getrennt sind, und für ein Einzelzimmer nur der Name. Die Namen müssen in der Reihenfolge erscheinen, in der sie gepaart wurden. Für das obige Beispiel geben Sie `["Alice and Carol", "Bob"]` zurück.

# --hints--

`getRoommates([{ "name": "Alice", "group": "A" }, { "name": "Bob", "group": "B" }, { "name": "Carol", "group": "A" }])` sollte `["Alice and Carol", "Bob"]` zurückgeben.

```js
assert.deepEqual(getRoommates([{ "name": "Alice", "group": "A" }, { "name": "Bob", "group": "B" }, { "name": "Carol", "group": "A" }]).sort(), ["Alice and Carol", "Bob"].sort());
```

`getRoommates([{ "name": "John", "group": "C" }, { "name": "Julia", "group": "C" }, { "name": "Jim", "group": "C" }])` sollte `["John and Julia", "Jim"]` zurückgeben.

```js
assert.deepEqual(getRoommates([{ "name": "John", "group": "C" }, { "name": "Julia", "group": "C" }, { "name": "Jim", "group": "C" }]).sort(), ["John and Julia", "Jim"].sort());
```

`getRoommates([{ "name": "Adam", "group": "D" }, { "name": "Abraham", "group": "E" }, { "name": "Austin", "group": "E" }, { "name": "Augustus", "group": "D" }, { "name": "Angelica", "group": "D" }, { "name": "Aaron", "group": "E" }])` sollte `["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"]` zurückgeben.

```js
assert.deepEqual(getRoommates([{ "name": "Adam", "group": "D" }, { "name": "Abraham", "group": "E" }, { "name": "Austin", "group": "E" }, { "name": "Augustus", "group": "D" }, { "name": "Angelica", "group": "D" }, { "name": "Aaron", "group": "E" }]).sort(), ["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"].sort());
```

`getRoommates([{ "name": "Frank", "group": "A" }, { "name": "Emitt", "group": "B" }, { "name": "Daria", "group": "F" }, { "name": "Charles", "group": "D" }, { "name": "Bailey", "group": "A" }, { "name": "Albert", "group": "F" }])` sollte `["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"]` zurückgeben.

```js
assert.deepEqual(getRoommates([{ "name": "Frank", "group": "A" }, { "name": "Emitt", "group": "B" }, { "name": "Daria", "group": "F" }, { "name": "Charles", "group": "D" }, { "name": "Bailey", "group": "A" }, { "name": "Albert", "group": "F" }]).sort(), ["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"].sort());
```

`getRoommates([{ "name": "Kevin", "group": "A" }, { "name": "Yuri", "group": "A" }, { "name": "Hugo", "group": "B" }, { "name": "Violet", "group": "A" }, { "name": "Brett", "group": "A" }, { "name": "Wayne", "group": "B" }])` sollte `["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"]` zurückgeben.

```js
assert.deepEqual(getRoommates([{ "name": "Kevin", "group": "A" }, { "name": "Yuri", "group": "A" }, { "name": "Hugo", "group": "B" }, { "name": "Violet", "group": "A" }, { "name": "Brett", "group": "A" }, { "name": "Wayne", "group": "B" }]).sort(), ["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"].sort());
```

# --seed--

## --seed-contents--

```js
function getRoommates(people) {

  return people;
}
```

# --solutions--

```js
function getRoommates(people) {
  const groups = {};

  for (const { name, group } of people) {
    (groups[group] = groups[group] || []).push(name);
  }

  return Object.values(groups).flatMap(group => {
    const rooms = [];
    for (let i = 0; i < group.length; i += 2) {
      rooms.push(group.slice(i, i + 2).join(" and "));
    }
    return rooms;
  });
}
```
