---
id: 6a15cadf5f240d05a2649556
title: "Aufgabe 312: Streaming-Kosten"
challengeType: 28
dashedName: challenge-312
---

# --description--

Gegeben ist ein Array, das Filme im Warenkorb Ihres Streaming-Dienstes darstellt, und ein String für Ihre Abonnementstufe. Geben Sie die Gesamtkosten der Filme zurück.

Jeder Artikel im Warenkorb ist ein Objekt mit einem `"format"` (`"HD"` oder `"4K"`) und einem `"type"` (`"rent"` oder `"buy"`). Deren Kosten sind:

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

Wenden Sie die folgenden Rabatte für die Abonnementstufen an:

- `"none"`: voller Preis
- `"basic"`: 10 % Rabatt
- `"premium"`: 25 % Rabatt

Geben Sie die Gesamtkosten gerundet auf zwei Dezimalstellen im Format `"$D.CC"` zurück.

# --hints--

`getStreamingBill([{ format: "HD", type: "rent" }], "none")` sollte `"$3.99"` zurückgeben.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }], "none"), "$3.99");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "premium")` sollte `"$12.73"` zurückgeben.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "premium"), "$12.73");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic")` sollte `"$18.87"` zurückgeben.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic"), "$18.87");
```

`getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium")` sollte `"$29.98"` zurückgeben.

```js
assert.equal(getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium"), "$29.98");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none")` sollte `"$42.96"` zurückgeben.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none"), "$42.96");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }, { format: "HD", type: "buy" }], "basic")` sollte `"$50.36"` zurückgeben.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }, { format: "HD", type: "buy" }], "basic"), "$50.36");
```

# --seed--

## --seed-contents--

```js
function getStreamingBill(cart, subscription) {

  return cart;
}
```

# --solutions--

```js
function getStreamingBill(cart, subscription) {
  const pricing = {
    HD: { rent: 3.99, buy: 12.99 },
    "4K": { rent: 5.99, buy: 19.99 },
  };

  const discounts = {
    none: 1,
    basic: 0.9,
    premium: 0.75,
  };

  const multiplier = discounts[subscription];

  const total = cart.reduce((sum, item) => {
    const price = pricing[item.format][item.type];
    return sum + price * multiplier;
  }, 0);

  return `$${total.toFixed(2)}`;
}
```
