---
id: 6a15cadf5f240d05a264955b
title: "Aufgabe 317: BMI-Rechner"
challengeType: 28
dashedName: challenge-317
---

# --description--

Gegeben sind ein Gewicht in Pfund und eine Körpergröße in Zoll. Geben Sie den BMI (Body Mass Index) zurück, gerundet auf eine Nachkommastelle.

Um den BMI zu berechnen: Teilen Sie das Gewicht durch die Körpergröße zum Quadrat und multiplizieren Sie das Ergebnis dann mit 703.

# --hints--

`calculateBmi(180, 70)` sollte `25.8` zurückgeben.

```js
assert.equal(calculateBmi(180, 70), 25.8);
```

`calculateBmi(140, 64)` sollte `24.0` zurückgeben.

```js
assert.equal(calculateBmi(140, 64), 24.0);
```

`calculateBmi(160, 76)` sollte `19.5` zurückgeben.

```js
assert.equal(calculateBmi(160, 76), 19.5);
```

`calculateBmi(200, 60)` sollte `39.1` zurückgeben.

```js
assert.equal(calculateBmi(200, 60), 39.1);
```

`calculateBmi(150, 68)` sollte `22.8` zurückgeben.

```js
assert.equal(calculateBmi(150, 68), 22.8);
```

# --seed--

## --seed-contents--

```js
function calculateBmi(weight, height) {

  return weight;
}
```

# --solutions--

```js
function calculateBmi(weight, height) {
  return Math.round((weight / (height * height)) * 703 * 10) / 10;
}
```
