---
id: 6a151d32d772271dd2448c2d
title: "Aufgabe 311: Zauberwirker"
challengeType: 28
dashedName: challenge-311
---

# --description--

Gegeben ist ein String mit Zaubercodes, die Sie wirken. Berechnen Sie die Gesamtpunktzahl.

Jedes Zeichen im String steht für einen Zauber:

| Code | Zauber | Kategorie | Grundpunktzahl |
| - | - | - | - |
| `"f"` | Feuer | Zerstörung | 3 |
| `"l"` | Blitz | Zerstörung | 3 |
| `"i"` | Eis | Kontrolle | 2 |
| `"w"` | Wind | Kontrolle | 2 |
| `"h"` | Heilung | Wiederherstellung | 1 |
| `"s"` | Schild | Wiederherstellung | 1 |

Ein Kombinationsmultiplikator wird basierend darauf angewendet, wie viele Zauber hintereinander aus unterschiedlichen Kategorien gewirkt wurden:

- Der erste Zauber erzielt immer die Grundpunktzahl.
- Jeder aufeinanderfolgende Zauber aus einer anderen Kategorie als der vorherige erhöht den Multiplikator um 1.
- Wird ein Zauber aus derselben Kategorie wie der vorherige gewirkt, wird der Multiplikator wieder auf 1 zurückgesetzt.
- Die Punktzahl für jeden Zauber ist seine Grundpunktzahl multipliziert mit dem aktuellen Multiplikator.

Geben Sie die Gesamtpunktzahl aus der Folge der Zauber zurück.

# --hints--

`cast("fihwl")` sollte `33` zurückgeben.

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` sollte `45` zurückgeben.

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` sollte `37` zurückgeben.

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` sollte `50` zurückgeben.

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` sollte `101` zurückgeben.

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
