---
id: 6a151d32d772271dd2448c2c
title: "Aufgabe 310: Britisches Englisch zu amerikanischem Englisch"
challengeType: 28
dashedName: challenge-310
---

# --description--

Gegeben ist ein Satz. Wandeln Sie alle britischen Englischschreibweisen mithilfe der folgenden Tabelle in ihre amerikanischen Entsprechungen um und geben Sie den aktualisierten Satz zurück:

| British | American |
| - | - |
| `"colour"` | `"color"` |
| `"flavour"` | `"flavor"` |
| `"honour"` | `"honor"` |
| `"neighbour"` | `"neighbor"` |
| `"labour"` | `"labor"` |
| `"humour"` | `"humor"` |
| `"centre"` | `"center"` |
| `"fibre"` | `"fiber"` |
| `"defence"` | `"defense"` |
| `"offence"` | `"offense"` |
| `"organise"` | `"organize"` |
| `"recognise"` | `"recognize"` |
| `"analyse"` | `"analyze"` |

- Die Ersetzungen sollen nicht auf Groß- und Kleinschreibung achten. Zum Beispiel sollte `"Colour"` zu `"Color"` werden.
- Die Eingabe kann Wörter enthalten, die auf der genauen Schreibweise einer Wurzel aus der Tabelle aufbauen und ebenfalls geändert werden müssen. Zum Beispiel sollte `"colouring"` zu `"coloring"` werden, und `"disorganised"` zu `"disorganized"`.

# --hints--

`britishToAmerican("I love the colour blue.")` sollte `"I love the color blue."` zurückgeben

```js
assert.equal(britishToAmerican("I love the colour blue."), "I love the color blue.");
```

`britishToAmerican("The fibre optic cable is new.")` sollte `"The fiber optic cable is new."` zurückgeben

```js
assert.equal(britishToAmerican("The fibre optic cable is new."), "The fiber optic cable is new.");
```

`britishToAmerican("It's an honour to meet someone with such humour.")` sollte `"It's an honor to meet someone with such humor."` zurückgeben

```js
assert.equal(britishToAmerican("It's an honour to meet someone with such humour."), "It's an honor to meet someone with such humor.");
```

`britishToAmerican("The unrecognised artist analysed his colour palette at the centre.")` sollte `"The unrecognized artist analyzed his color palette at the center."` zurückgeben

```js
assert.equal(britishToAmerican("The unrecognised artist analysed his colour palette at the centre."), "The unrecognized artist analyzed his color palette at the center.");
```

`britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful.")` sollte `"The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful."` zurückgeben

```js
assert.equal(britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful."), "The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.");
```

# --seed--

## --seed-contents--

```js
function britishToAmerican(sentence) {

  return sentence;
}
```

# --solutions--

```js
function britishToAmerican(sentence) {
  const lookup = {
    colour: "color",
    flavour: "flavor",
    honour: "honor",
    neighbour: "neighbor",
    labour: "labor",
    humour: "humor",
    centre: "center",
    fibre: "fiber",
    defence: "defense",
    offence: "offense",
    organise: "organize",
    recognise: "recognize",
    analyse: "analyze"
  };

  let result = sentence;

  for (const [british, american] of Object.entries(lookup)) {
    const regex = new RegExp(british, "gi");
    result = result.replace(regex, match => {
      return match[0] === match[0].toUpperCase()
        ? american[0].toUpperCase() + american.slice(1)
        : american;
    });
  }

  return result;
}
```
