---
id: 6a15cadf5f240d05a2649558
title: "Aufgabe 314: Primfaktorzerlegung"
challengeType: 28
dashedName: challenge-314
---

# --description--

Gegeben ist eine ganze Zahl größer als 1. Geben Sie ihre Primfaktorzerlegung als Array von Zahlen in aufsteigender Reihenfolge zurück.

Eine Primfaktorzerlegung ist die Menge von Primzahlen, die miteinander multipliziert die gegebene Zahl ergeben. Jede Zahl hat genau eine solche Menge. Zum Beispiel ist die Primfaktorzerlegung von 20 `[2, 2, 5]`, weil 2 * 2 * 5 = 20.

Wenn die gegebene Zahl selbst eine Primzahl ist, geben Sie sie in einem Array mit einem Element zurück.

# --hints--

`primeFactorization(20)` sollte `[2, 2, 5]` zurückgeben.

```js
assert.deepEqual(primeFactorization(20), [2, 2, 5]);
```

`primeFactorization(17)` sollte `[17]` zurückgeben.

```js
assert.deepEqual(primeFactorization(17), [17]);
```

`primeFactorization(15)` sollte `[3, 5]` zurückgeben.

```js
assert.deepEqual(primeFactorization(15), [3, 5]);
```

`primeFactorization(35)` sollte `[5, 7]` zurückgeben.

```js
assert.deepEqual(primeFactorization(35), [5, 7]);
```

`primeFactorization(999)` sollte `[3, 3, 3, 37]` zurückgeben.

```js
assert.deepEqual(primeFactorization(999), [3, 3, 3, 37]);
```

`primeFactorization(360)` sollte `[2, 2, 2, 3, 3, 5]` zurückgeben.

```js
assert.deepEqual(primeFactorization(360), [2, 2, 2, 3, 3, 5]);
```

`primeFactorization(510510)` sollte `[2, 3, 5, 7, 11, 13, 17]` zurückgeben.

```js
assert.deepEqual(primeFactorization(510510), [2, 3, 5, 7, 11, 13, 17]);
```

# --seed--

## --seed-contents--

```js
function primeFactorization(n) {

  return n;
}
```

# --solutions--

```js
function primeFactorization(n) {
  const factors = [];
  let divisor = 2;

  while (divisor * divisor <= n) {
    while (n % divisor === 0) {
      factors.push(divisor);
      n /= divisor;
    }
    divisor++;
  }

  if (n > 1) factors.push(n);

  return factors;
}
```
