---
id: 6a15cadf5f240d05a2649557
title: "Aufgabe 313: Mietkosten"
challengeType: 28
dashedName: challenge-313
---

# --description--

Gegeben sind ein Mietzeitstempel, ein Rückgabezeitstempel und eine Mietstufe. Geben Sie die Gesamtkosten der Miete einschließlich etwaiger Verspätungsgebühren zurück.

- Die angegebenen Zeitstempel sind UTC-ISO-Strings, zum Beispiel: `"2026-06-18T18:30:00Z"`.
- Die Mietstufe ist die Anzahl der Tage, bevor die Miete zurückgegeben werden muss: `1`, `3` oder `7`.
- Die Miete muss bis spätestens 12:00 Uhr UTC am letzten Tag der Mietdauer zurückgegeben werden. Zum Beispiel ist eine 1-Tages-Miete, die irgendwann am 15. März ausgeliehen wurde, bis 12:00 Uhr UTC am 16. März zurückzugeben.
- Für jeden Tag nach dem Fälligkeitsdatum und der Fälligkeitszeit fällt eine Verspätungsgebühr an.

Die Preisgestaltung ist wie folgt:

| Stufe | Grundkosten | Verspätungsgebühr pro Tag |
| - | - | - |
| 1 Tag | 4,99 $ | 3,99 $ |
| 3 Tage | 3,99 $ | 2,99 $ |
| 7 Tage | 2,99 $ | 0,99 $ |

Geben Sie die Gesamtkosten auf zwei Dezimalstellen gerundet im Format `"$D.CC"` zurück.

# --hints--

`getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` sollte `"$4.99"` zurückgeben.

```js
assert.equal(getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99");
```

`getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` sollte `"$12.97"` zurückgeben.

```js
assert.equal(getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97");
```

`getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` sollte `"$3.99"` zurückgeben.

```js
assert.equal(getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99");
```

`getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` sollte `"$9.97"` zurückgeben.

```js
assert.equal(getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97");
```

`getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` sollte `"$2.99"` zurückgeben.

```js
assert.equal(getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99");
```

`getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` sollte `"$358.40"` zurückgeben.

```js
assert.equal(getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40");
```

# --seed--

## --seed-contents--

```js
function getRentalCost(rented, returned, tier) {

  return rented;
}
```

# --solutions--

```js
function getRentalCost(rented, returned, tier) {
  const pricing = {
    1: { base: 4.99, late: 3.99 },
    3: { base: 3.99, late: 2.99 },
    7: { base: 2.99, late: 0.99 },
  };

  const rentedDate = new Date(rented);
  const returnDate = new Date(returned);
  const due = new Date(rentedDate);
  due.setUTCDate(due.getUTCDate() + tier);
  due.setUTCHours(12, 0, 0, 0);
  const { base, late } = pricing[tier];

  if (returnDate <= due) {
    return `$${base.toFixed(2)}`;
  }

  const msPerDay = 1000 * 60 * 60 * 24;
  const daysLate = Math.ceil((returnDate - due) / msPerDay);
  const total = base + daysLate * late;

  return `$${total.toFixed(2)}`;
}
```
