---
id: 6a15cadf5f240d05a264955d
title: "Aufgabe 319: Frontmatter-Parser"
challengeType: 28
dashedName: challenge-319
---

# --description--

Gegeben ist ein String, der einen Frontmatter-Block darstellt. Parsen Sie ihn und geben Sie ein Objekt (JavaScript) oder ein Dictionary (Python) mit den Schlüsseln und Werten zurück.

Frontmatter ist von `---`-Begrenzer umschlossen und enthält darin `key: value`-Paare, jeweils eines pro Zeile. Zum Beispiel:

```md
---
title: My Post
draft: false
views: 100
---
```

Soll zurückgeben:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- Zahlen, Booleans und Strings sollten jeweils als ihren entsprechenden Typ zurückgegeben werden.
- Der gegebene String enthält Zeilenumbrüche, die durch das Newline-Zeichen (`"\n"`) getrennt sind. Das obige Beispiel würde so übergeben werden: `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`.

# --hints--

`parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` sollte `{ title: "My Post", draft: false, views: 100 }` zurückgeben.

```js
assert.deepEqual(parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---"), { title: "My Post", draft: false, views: 100 });
```

`parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` sollte `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false }` zurückgeben.

```js
assert.deepEqual(parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---"), { id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false });
```

`parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` sollte `{ version: "1.0.0", url: "https://example.com", private: true }` zurückgeben.

```js
assert.deepEqual(parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---"), { version: "1.0.0", url: "https://example.com", private: true });
```

`parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---")` sollte `{ rating: 4.5, price: 9.99 }` zurückgeben.

```js
assert.deepEqual(parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---"), { rating: 4.5, price: 9.99 });
```

# --seed--

## --seed-contents--

```js
function parseFrontmatter(str) {

  return str;
}
```

# --solutions--

```js
function parseFrontmatter(str) {
  const lines = str.split("\n").filter(line => line !== "---" && line !== "");

  return lines.reduce((obj, line) => {
    const [key, ...rest] = line.split(":");
    const value = rest.join(":").trim();

    if (value === "true") obj[key.trim()] = true;
    else if (value === "false") obj[key.trim()] = false;
    else if (!isNaN(value)) obj[key.trim()] = Number(value);
    else obj[key.trim()] = value;

    return obj;
  }, {});
}
```
