---
id: 6a15cadf5f240d05a264955e
title: "Aufgabe 320: Blutbank"
challengeType: 28
dashedName: challenge-320
---

# --description--

Gegeben ist ein Array mit dem Inventar einer Blutbank und ein Array mit Blutgruppenanfragen von Patientinnen und Patienten. Geben Sie einen String im Format `"X of Y patients served"` zurück. Dabei ist `X` die maximale Anzahl an Patientinnen und Patienten, die aus dem Inventar der Blutbank Blut erhalten können, und `Y` die Gesamtzahl der Patientinnen und Patienten.

Jeder Eintrag in beiden Arrays ist eine der folgenden Blutgruppen: `"AB"`, `"A"`, `"B"` oder `"O"`.

Kompatibilitätsregeln:

- `"AB"` kann von jeder Blutgruppe empfangen.
- `"A"` kann von `"A"` und `"O"` empfangen.
- `"B"` kann von `"B"` und `"O"` empfangen.
- `"O"` kann nur von `"O"` empfangen.

Doppelte Einträge in den gegebenen Arrays stellen die Menge dar.

# --hints--

`triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` sollte `"4 of 4 patients served"` zurückgeben.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` sollte `"3 of 5 patients served"` zurückgeben.

```js
assert.equal(triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served");
```

`triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` sollte `"4 of 5 patients served"` zurückgeben.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served");
```

`triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` sollte `"4 of 4 patients served"` zurückgeben.

```js
assert.equal(triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` sollte `"8 of 11 patients served"` zurückgeben.

```js
assert.equal(triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served");
```

`triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` sollte `"13 of 14 patients served"` zurückgeben.

```js
assert.equal(triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served");
```

# --seed--

## --seed-contents--

```js
function triageBlood(bank, patients) {

  return bank;
}
```

# --solutions--

```js
function triageBlood(bank, patients) {
  const b = {};
  const p = {};
  for (const t of bank) b[t] = (b[t] || 0) + 1;
  for (const t of patients) p[t] = (p[t] || 0) + 1;

  const serve = (patient, donors) => {
    for (const donor of donors) {
      const n = Math.min(p[patient] || 0, b[donor] || 0);
      p[patient] = (p[patient] || 0) - n;
      b[donor] = (b[donor] || 0) - n;
      saved += n;
    }
  };

  let saved = 0;
  serve("O",  ["O"]);
  serve("A",  ["A", "O"]);
  serve("B",  ["B", "O"]);
  serve("AB", ["AB", "A", "B", "O"]);
  return `${saved} of ${patients.length} patients served`;
}
```
