---
id: 6a15cadf5f240d05a264955a
title: "Aufgabe 316: 1337-Sprache"
challengeType: 28
dashedName: challenge-316
---

# --description--

Gegeben ist ein String in Kleinbuchstaben. Geben Sie ihn übersetzt in Leetspeak zurück, indem Sie die unten stehenden Buchstaben durch ihre Leet-Entsprechungen ersetzen:

| Buchstabe | Leet |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- Zeichen ohne Ersatz bleiben unverändert.

# --hints--

`makeLeet("cool")` sollte `"c001"` zurückgeben.

```js
assert.equal(makeLeet("cool"), "c001");
```

`makeLeet("leet")` sollte `"1337"` zurückgeben.

```js
assert.equal(makeLeet("leet"), "1337");
```

`makeLeet("hacker")` sollte `"h4ck3r"` zurückgeben.

```js
assert.equal(makeLeet("hacker"), "h4ck3r");
```

`makeLeet("satellite")` sollte `"547311173"` zurückgeben.

```js
assert.equal(makeLeet("satellite"), "547311173");
```

`makeLeet("abcdefghijklmnopqrstuvwxyz")` sollte `"4bcd3f9h1jk1mn0pqr57uvwxyz"` zurückgeben.

```js
assert.equal(makeLeet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz");
```

# --seed--

## --seed-contents--

```js
function makeLeet(str) {

  return str;
}
```

# --solutions--

```js
function makeLeet(str) {
  const map = { a: "4", e: "3", g: "9", i: "1", l: "1", o: "0", t: "7", s: "5" };
  return str.split("").map(char => map[char] || char).join("");
}
```
