---
id: 6a19b1062d1b153d8ac76d73
title: "Aufgabe 325: Glückszahl"
challengeType: 28
dashedName: challenge-325
---

# --description--

Gegeben ist ein String mit dem Vor- und Nachnamen einer Person. Berechnen Sie deren Glückszahl anhand der folgenden Regeln:

- Vor- und Nachnamen sind durch ein Leerzeichen getrennt
- Ermitteln Sie die Anzahl der Vokale und Konsonanten für jeden Namen
- Multiplizieren Sie die kleinere Anzahl der Vokale und Konsonanten miteinander und anschließend mit der Länge des kleineren Namens
- Machen Sie dasselbe für die beiden größeren Werte und den größeren Namen
- Ziehen Sie den kleineren Wert vom größeren ab, um die Glückszahl zu erhalten

Wenn der Endwert null (`0`) ist, geben Sie `13` zurück.

# --hints--

`getLuckyNumber("John Doe")` sollte `21` zurückgeben.

```js
assert.equal(getLuckyNumber("John Doe"), 21);
```

`getLuckyNumber("Olivia Lewis")` sollte `52` zurückgeben.

```js
assert.equal(getLuckyNumber("Olivia Lewis"), 52);
```

`getLuckyNumber("James Wilson")` sollte `18` zurückgeben.

```js
assert.equal(getLuckyNumber("James Wilson"), 18);
```

`getLuckyNumber("Elizabeth Hernandez")` sollte `81` zurückgeben.

```js
assert.equal(getLuckyNumber("Elizabeth Hernandez"), 81);
```

`getLuckyNumber("Mike Walker")` sollte `32` zurückgeben.

```js
assert.equal(getLuckyNumber("Mike Walker"), 32);
```

`getLuckyNumber("Chloe Perez")` sollte `13` zurückgeben.

```js
assert.equal(getLuckyNumber("Chloe Perez"), 13);
```

# --seed--

## --seed-contents--

```js
function getLuckyNumber(name) {

  return name;
}
```

# --solutions--

```js
function getLuckyNumber(name) {
  const vowels = "aeiouAEIOU";

  function nameStats(str) {
    const vowelCount = str.split("").filter(c => vowels.includes(c)).length;
    const consonantCount = str.split("").filter(c => c.match(/[a-zA-Z]/) && !vowels.includes(c)).length;
    return { vowelCount, consonantCount, length: str.length };
  }

  const [first, last] = name.split(" ");
  const firstStats = nameStats(first);
  const lastStats = nameStats(last);

  const largerVowels = Math.max(firstStats.vowelCount, lastStats.vowelCount);
  const smallerVowels = Math.min(firstStats.vowelCount, lastStats.vowelCount);
  const largerConsonants = Math.max(firstStats.consonantCount, lastStats.consonantCount);
  const smallerConsonants = Math.min(firstStats.consonantCount, lastStats.consonantCount);
  const largerLength = Math.max(firstStats.length, lastStats.length);
  const smallerLength = Math.min(firstStats.length, lastStats.length);

  const largerValue = largerVowels * largerConsonants * largerLength;
  const smallerValue = smallerVowels * smallerConsonants * smallerLength;
  const result = largerValue - smallerValue;

  return result === 0 ? 13 : result;
}
```
