---
id: 6a19b1062d1b153d8ac76d72
title: "Aufgabe 324: Zählen doppelter Zeichen"
challengeType: 28
dashedName: challenge-324
---

# --description--

Gegeben sind zwei Strings. Geben Sie die Anzahl der Zeichen aus dem zweiten String zurück, die im ersten gefunden werden können.

- Doppelte Zeichen im zweiten String werden einzeln gezählt.

# --hints--

`duplicateCharacterCount("aloha", "hei")` sollte `1` zurückgeben.

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")` sollte `4` zurückgeben.

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")` sollte `3` zurückgeben.

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")` sollte `0` zurückgeben.

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")` sollte `5` zurückgeben.

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")` sollte `2` zurückgeben.

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")` sollte `26` zurückgeben.

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
