---
id: 6a19dd69388e5a6e59e4f2ef
title: "Aufgabe 326: Maximaler Gewinn"
challengeType: 28
dashedName: challenge-326
---

# --description--

Gegeben ist ein Array mit täglichen Aktienkursen und ein Budget (in Dollar). Berechnen Sie den maximalen Gewinn, den Sie durch den Kauf und Verkauf der Aktie im angegebenen Zeitraum erzielen können.

- Sie dürfen nur verkaufen, nachdem Sie gekauft haben.
- Sie dürfen höchstens eine Kauf- und eine Verkaufsaktion durchführen. Sobald Sie verkaufen, können Sie nicht erneut kaufen.
- Sie können nur ganze Aktien kaufen.
- Geben Sie den maximal möglichen Gewinn als String zurück, abgerundet auf den nächsten Cent und mit zwei Dezimalstellen formatiert.

# --hints--

`getMaxProfit([5, 6], 50)` sollte `"10.00"` zurückgeben.

```js
assert.equal(getMaxProfit([5, 6], 50), "10.00");
```

`getMaxProfit([8, 2, 5, 10], 20)` sollte `"80.00"` zurückgeben.

```js
assert.equal(getMaxProfit([8, 2, 5, 10], 20), "80.00");
```

`getMaxProfit([4, 5, 3, 6], 20)` sollte `"18.00"` zurückgeben.

```js
assert.equal(getMaxProfit([4, 5, 3, 6], 20), "18.00");
```

`getMaxProfit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200)` sollte `"8.31"` zurückgeben.

```js
assert.equal(getMaxProfit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200), "8.31");
```

`getMaxProfit([15.38, 15.01, 14.99, 14.62, 14.28], 80)` sollte `"0.00"` zurückgeben.

```js
assert.equal(getMaxProfit([15.38, 15.01, 14.99, 14.62, 14.28], 80), "0.00");
```

`getMaxProfit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25)` sollte `"73.80"` zurückgeben.

```js
assert.equal(getMaxProfit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25), "73.80");
```

# --seed--

## --seed-contents--

```js
function getMaxProfit(prices, budget) {

  return prices;
}
```

# --solutions--

```js
function getMaxProfit(prices, budget) {
  let maxProfit = 0;

  for (let buy = 0; buy < prices.length - 1; buy++) {
    for (let sell = buy + 1; sell < prices.length; sell++) {
      const shares = Math.floor(budget / prices[buy]);
      const profit = shares * (prices[sell] - prices[buy]);
      if (profit > maxProfit) maxProfit = profit;
    }
  }

  return (Math.floor(maxProfit * 100) / 100).toFixed(2);
}
```
