---
id: 6a1d9f98e819ed70a0e994dd
title: "Aufgabe 331: Nächstes Vielfaches"
challengeType: 28
dashedName: challenge-331
---

# --description--

Gegeben sind zwei ganze Zahlen. Runden Sie die erste auf das nächste Vielfache der zweiten.

# --hints--

`roundToNearestMultiple(5, 3)` sollte `6` zurückgeben.

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)` sollte `16` zurückgeben.

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)` sollte `45` zurückgeben.

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)` sollte `33` zurückgeben.

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)` sollte `96` zurückgeben.

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
