---
id: 6a1d9f98e819ed70a0e994df
title: "Aufgabe 333: Issue-Triage 2"
challengeType: 28
dashedName: challenge-333
---

# --description--

Angenommen, Sie haben einen Issue-Titel und ein Array mit aktuellen Bezeichnungen. Geben Sie ein aktualisiertes Array von Bezeichnungen zurück, das auf den folgenden Regeln basiert:

Wenn der Issue keine Bezeichnungen hat, fügen Sie hinzu:

- `"bug"` und `"needs triage"`, wenn der Titel `"error"` oder `"bug"` enthält
- `"enhancement"` und `"discussing"`, wenn der Titel `"feature"` oder `"add"` enthält

Andernfalls, wenn die gegebenen Bezeichnungen enthalten:

- `"needs triage"` und der Titel `"simple"` oder `"easy"` enthält, entfernen Sie `"needs triage"` und fügen Sie `"good first issue"` hinzu
- `"discussing"` und der Titel `"planned"` oder `"next"` enthält, entfernen Sie `"discussing"` und fügen Sie `"on the roadmap"` hinzu
- Andernfalls, wenn `"needs triage"` oder `"discussing"` vorhanden ist, entfernen Sie es und fügen Sie `"help wanted"` hinzu

Wenn der Titel enthält:

- `"security"`, fügen Sie eine Bezeichnung `"critical"` hinzu

# --hints--

`triageIssue("app crashes with error", [])` sollte `["bug", "needs triage"]` zurückgeben.

```js
assert.deepEqual(triageIssue("app crashes with error", []).sort(), ["bug", "needs triage"].sort());
```

`triageIssue("app crashes with error", ["bug", "needs triage"])` sollte `["bug", "help wanted"]` zurückgeben.

```js
assert.deepEqual(triageIssue("app crashes with error", ["bug", "needs triage"]).sort(), ["bug", "help wanted"].sort());
```

`triageIssue("add dark mode", [])` sollte `["enhancement", "discussing"]` zurückgeben.

```js
assert.deepEqual(triageIssue("add dark mode", []).sort(), ["enhancement", "discussing"].sort());
```

`triageIssue("add dark mode", ["enhancement", "discussing"])` sollte `["enhancement", "help wanted"]` zurückgeben.

```js
assert.deepEqual(triageIssue("add dark mode", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted"].sort());
```

`triageIssue("xss security bug", [])` sollte `["bug", "needs triage", "critical"]` zurückgeben.

```js
assert.deepEqual(triageIssue("xss security bug", []).sort(), ["bug", "needs triage", "critical"].sort());
```

`triageIssue("security vulnerability in auth", [])` sollte `["critical"]` zurückgeben.

```js
assert.deepEqual(triageIssue("security vulnerability in auth", []).sort(), ["critical"].sort());
```

`triageIssue("easy a11y fix", ["bug", "needs triage"])` sollte `["bug", "good first issue"]` zurückgeben.

```js
assert.deepEqual(triageIssue("easy a11y fix", ["bug", "needs triage"]).sort(), ["bug", "good first issue"].sort());
```

`triageIssue("planned api migration", ["enhancement", "discussing"])` sollte `["enhancement", "on the roadmap"]` zurückgeben.

```js
assert.deepEqual(triageIssue("planned api migration", ["enhancement", "discussing"]).sort(), ["enhancement", "on the roadmap"].sort());
```

`triageIssue("improve security", ["enhancement", "discussing"])` sollte `["enhancement", "help wanted", "critical"]` zurückgeben.

```js
assert.deepEqual(triageIssue("improve security", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted", "critical"].sort());
```

# --seed--

## --seed-contents--

```js
function triageIssue(title, labels) {

  return title;
}
```

# --solutions--

```js
function triageIssue(title, labels) {
  const result = [...labels];
  const t = title.toLowerCase();

  const has = label => result.includes(label);
  const add = label => { if (!has(label)) result.push(label); };
  const remove = label => { const i = result.indexOf(label); if (i > -1) result.splice(i, 1); };

  if (result.length === 0) {
    if (t.includes("error") || t.includes("bug")) {
      add("bug"); add("needs triage");
    } else if (t.includes("feature") || t.includes("add")) {
      add("enhancement"); add("discussing");
    }
  } else {
    if (has("needs triage") && (t.includes("simple") || t.includes("easy"))) {
      remove("needs triage"); add("good first issue");
    } else if (has("discussing") && (t.includes("planned") || t.includes("next"))) {
      remove("discussing"); add("on the roadmap");
    } else if (has("needs triage") || has("discussing")) {
      remove("needs triage"); remove("discussing"); add("help wanted");
    }
  }

  if (t.includes("security")) {
    add("critical");
  }

  return result;
}
```
