---
id: 6a1d9f98e819ed70a0e994e1
title: "Aufgabe 335: Fünf Würfel"
challengeType: 28
dashedName: challenge-335
---

# --description--

Gegeben ist ein Array von fünf Würfeln mit den Werten 1–6. Geben Sie die bestmögliche Hand zurück.

Hier sind die Hände vom niedrigsten zum höchsten Rang:

| Hand | Beschreibung |
|------|-------------|
| `"no pair"` | Kein Paar oder besser |
| `"pair"` | Zwei Würfel mit demselben Wert |
| `"two pair"` | Zwei verschiedene Paare |
| `"three of a kind"` | Drei Würfel mit demselben Wert |
| `"small straight"` | Vier aufeinanderfolgende Werte |
| `"large straight"` | Fünf aufeinanderfolgende Werte |
| `"full house"` | Drilling und ein Paar |
| `"four of a kind"` | Vier Würfel mit demselben Wert |
| `"five of a kind"` | Alle fünf Würfel mit demselben Wert |

# --hints--

`fiveDice([1, 1, 1, 1, 1])` sollte `"five of a kind"` zurückgeben.

```js
assert.equal(fiveDice([1, 1, 1, 1, 1]), "five of a kind");
```

`fiveDice([5, 5, 5, 6, 5])` sollte `"four of a kind"` zurückgeben.

```js
assert.equal(fiveDice([5, 5, 5, 6, 5]), "four of a kind");
```

`fiveDice([2, 5, 6, 4, 3])` sollte `"large straight"` zurückgeben.

```js
assert.equal(fiveDice([2, 5, 6, 4, 3]), "large straight");
```

`fiveDice([4, 3, 3, 3, 1])` sollte `"three of a kind"` zurückgeben.

```js
assert.equal(fiveDice([4, 3, 3, 3, 1]), "three of a kind");
```

`fiveDice([4, 6, 2, 6, 5])` sollte `"pair"` zurückgeben.

```js
assert.equal(fiveDice([4, 6, 2, 6, 5]), "pair");
```

`fiveDice([1, 4, 5, 6, 2])` sollte `"no pair"` zurückgeben.

```js
assert.equal(fiveDice([1, 4, 5, 6, 2]), "no pair");
```

`fiveDice([1, 3, 4, 6, 2])` sollte `"small straight"` zurückgeben.

```js
assert.equal(fiveDice([1, 3, 4, 6, 2]), "small straight");
```

`fiveDice([2, 2, 5, 2, 5])` sollte `"full house"` zurückgeben.

```js
assert.equal(fiveDice([2, 2, 5, 2, 5]), "full house");
```

`fiveDice([6, 4, 5, 6, 4])` sollte `"two pair"` zurückgeben.

```js
assert.equal(fiveDice([6, 4, 5, 6, 4]), "two pair");
```

# --seed--

## --seed-contents--

```js
function fiveDice(dice) {

  return dice;
}
```

# --solutions--

```js
function fiveDice(dice) {
  const counts = {};
  for (const d of dice) counts[d] = (counts[d] || 0) + 1;
  const values = Object.values(counts).sort((a, b) => b - a);
  const keys = Object.keys(counts).map(Number).sort((a, b) => a - b);

  const isLargeStraight = keys.length === 5 && keys[4] - keys[0] === 4;
  const isSmallStraight = !isLargeStraight && (() => {
    for (let i = 0; i <= 1; i++) {
      const slice = keys.slice(i, i + 4);
      if (slice.length === 4 && slice[3] - slice[0] === 3 &&
        new Set(slice).size === 4) return true;
    }
    return false;
  })();

  if (values[0] === 5) return "five of a kind";
  if (values[0] === 4) return "four of a kind";
  if (values[0] === 3 && values[1] === 2) return "full house";
  if (isLargeStraight) return "large straight";
  if (isSmallStraight) return "small straight";
  if (values[0] === 3) return "three of a kind";
  if (values[0] === 2 && values[1] === 2) return "two pair";
  if (values[0] === 2) return "pair";
  return "no pair";
}
```
