---
id: 6a1d9f98e819ed70a0e994db
title: "Aufgabe 329: Farbeimer-Füllung"
challengeType: 28
dashedName: challenge-329
---

# --description--

Gegeben ist ein 2D-Raster, eine Startposition (`[row, col]`) und ein neuer Wert. Ersetzen Sie den Wert an der Startposition und alle verbundenen Zellen mit demselben Wert durch den neuen Wert.

- Zellen sind verbunden, wenn sie horizontal oder vertikal angrenzend sind (nicht diagonal).

Geben Sie das aktualisierte Raster zurück.

# --hints--

`bucketFill([["R", "G"], ["R", "G"]], [0, 1], "B")` sollte `[["R", "B"], ["R", "B"]]` zurückgeben.

```js
assert.deepEqual(bucketFill([["R", "G"], ["R", "G"]], [0, 1], "B"), [["R", "B"], ["R", "B"]]);
```

`bucketFill([["Y", "G", "G"], ["Y", "Y", "Y"], ["B", "Y", "R"]], [1, 2], "B")` sollte `[["B", "G", "G"], ["B", "B", "B"], ["B", "B", "R"]]` zurückgeben.

```js
assert.deepEqual(bucketFill([["Y", "G", "G"], ["Y", "Y", "Y"], ["B", "Y", "R"]], [1, 2], "B"), [["B", "G", "G"], ["B", "B", "B"], ["B", "B", "R"]]);
```

`bucketFill([["O", "O", "P"], ["P", "O", "O"], ["P", "P", "O"]], [2, 0], "R")` sollte `[["O", "O", "P"], ["R", "O", "O"], ["R", "R", "O"]]` zurückgeben.

```js
assert.deepEqual(bucketFill([["O", "O", "P"], ["P", "O", "O"], ["P", "P", "O"]], [2, 0], "R"), [["O", "O", "P"], ["R", "O", "O"], ["R", "R", "O"]]);
```

`bucketFill([["T", "T", "R", "T"], ["R", "T", "R", "T"], ["R", "T", "R", "T"], ["T", "T", "T", "T"]], [0, 3], "Y")` sollte `[["Y", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["Y", "Y", "Y", "Y"]]` zurückgeben.

```js
assert.deepEqual(bucketFill([["T", "T", "R", "T"], ["R", "T", "R", "T"], ["R", "T", "R", "T"], ["T", "T", "T", "T"]], [0, 3], "Y"), [["Y", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["Y", "Y", "Y", "Y"]]);
```

`bucketFill([["G", "B", "G", "B"], ["R", "B", "B", "G"], ["B", "G", "B", "R"], ["B", "G", "G", "B"]], [2, 2], "G")` sollte `[["G", "G", "G", "B"], ["R", "G", "G", "G"], ["B", "G", "G", "R"], ["B", "G", "G", "B"]]` zurückgeben.

```js
assert.deepEqual(bucketFill([["G", "B", "G", "B"], ["R", "B", "B", "G"], ["B", "G", "B", "R"], ["B", "G", "G", "B"]], [2, 2], "G"), [["G", "G", "G", "B"], ["R", "G", "G", "G"], ["B", "G", "G", "R"], ["B", "G", "G", "B"]]);
```

# --seed--

## --seed-contents--

```js
function bucketFill(grid, [row, col], newValue) {

  return grid;
}
```

# --solutions--

```js
function bucketFill(grid, [row, col], newValue) {
  const target = grid[row][col];
  if (target === newValue) return grid;

  function fill(r, c) {
    if (r < 0 || r >= grid.length) return;
    if (c < 0 || c >= grid[0].length) return;
    if (grid[r][c] !== target) return;

    grid[r][c] = newValue;
    fill(r - 1, c);
    fill(r + 1, c);
    fill(r, c - 1);
    fill(r, c + 1);
  }

  fill(row, col);
  return grid;
}
```
