---
id: 6a1d9f98e819ed70a0e994da
title: "Aufgabe 328: Kaprekars Routine"
challengeType: 28
dashedName: challenge-328
---

# --description--

Gegeben ist eine vierstellige Zahl. Geben Sie die Anzahl der Male zurück, die Sie Kaprekars Routine anwenden müssen, bis Sie 6174 erreichen.

Kaprekars Routine funktioniert wie folgt:

- Ordnen Sie die Ziffern in absteigender Reihenfolge, um die größte Zahl zu bilden
- Ordnen Sie die Ziffern in aufsteigender Reihenfolge, um die kleinste Zahl zu bilden (füllen Sie bei Bedarf mit führenden Nullen auf)
- Subtrahieren Sie die kleinere von der größeren Zahl
- Wiederholen Sie den Vorgang mit der neuen Zahl

# --hints--

`kaprekar(1234)` sollte `3` zurückgeben.

```js
assert.equal(kaprekar(1234), 3);
```

`kaprekar(2025)` sollte `6` zurückgeben.

```js
assert.equal(kaprekar(2025), 6);
```

`kaprekar(7173)` sollte `4` zurückgeben.

```js
assert.equal(kaprekar(7173), 4);
```

`kaprekar(3164)` sollte `7` zurückgeben.

```js
assert.equal(kaprekar(3164), 7);
```

`kaprekar(8082)` sollte `2` zurückgeben.

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
