---
id: 6a1d9f98e819ed70a0e994d9
title: "Aufgabe 327: Datenbankmigration"
challengeType: 28
dashedName: challenge-327
---

# --description--

Gegeben sind zwei Datenbankobjekte. Geben Sie das zweite Objekt zurück, wobei alle fehlenden Eigenschaften aus dem ersten Objekt ergänzt werden.

- Felder, die im Datensatz bereits existieren, dürfen nicht überschrieben werden.

# --hints--

`migrateRecord({ username: "", posts: 0 }, { verified: true })` sollte `{ username: "", posts: 0, verified: true }` zurückgeben.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { verified: true }), { username: "", posts: 0, verified: true });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 })` sollte `{ username: "camper", posts: 5 }` zurückgeben.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 }), { username: "camper", posts: 5 });
```

`migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" })` sollte `{ username: "camper", posts: 0, verified: false }` zurückgeben.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" }), { username: "camper", posts: 0, verified: false });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" })` sollte `{ username: "camper", role: "admin", posts: 0 }` zurückgeben.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" }), { username: "camper", role: "admin", posts: 0 });
```

`migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" })` sollte `{ username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false }` zurückgeben.

```js
assert.deepEqual(migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" }), { username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false });
```

# --seed--

## --seed-contents--

```js
function migrateRecord(schema, record) {

  return schema;
}
```

# --solutions--

```js
function migrateRecord(schema, record) {
  return { ...schema, ...record };
}
```
