---
id: 6a1d9f98e819ed70a0e994dc
title: "Aufgabe 330: Kleinbuchstaben-Wörter"
challengeType: 28
dashedName: challenge-330
---

# --description--

Gegeben ist ein String. Geben Sie nur die Wörter zurück, die vollständig aus Kleinbuchstaben bestehen, in ihrer ursprünglichen Reihenfolge und mit einem Leerzeichen zwischen den Wörtern.

# --hints--

`getLowercaseWords("hello GOOD world")` sollte `"hello world"` zurückgeben.

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")` sollte `"these are all lowercase"` zurückgeben.

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")` sollte `"less is more"` zurückgeben.

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")` sollte `"eat pizza every day"` zurückgeben.

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` sollte `"the quick brown fox jumped over the lazy dog"` zurückgeben.

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
