---
id: 6a1d9f98e819ed70a0e994e0
title: "Aufgabe 334: Exakte Wechselgeldbeträge"
challengeType: 28
dashedName: challenge-334
---

# --description--

Gegeben ist ein ganzzahliger Betrag in Cent. Geben Sie die Anzahl der verschiedenen Möglichkeiten zurück, exaktes Wechselgeld mit Pennies (1 Cent), Nickels (5 Cent), Dimes (10 Cent) und Quarters (25 Cent) zu bilden.

# --hints--

`exactChange(3)` sollte `1` zurückgeben.

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)` sollte `2` zurückgeben.

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)` sollte `6` zurückgeben.

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)` sollte `24` zurückgeben.

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)` sollte `73` zurückgeben.

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)` sollte `213` zurückgeben.

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
