---
id: 6a1d9f98e819ed70a0e994e2
title: "Aufgabe 336: Horoskop-Übereinstimmung"
challengeType: 28
dashedName: challenge-336
---

# --description--

Gegeben sind zwei Strings mit Sternzeichen. Geben Sie deren Kompatibilitätsprozentsatz zurück.

Die Zeichen sind in einem Rad mit 12 Positionen in folgender Reihenfolge angeordnet: `"Aries"`, `"Taurus"`, `"Gemini"`, `"Cancer"`, `"Leo"`, `"Virgo"`, `"Libra"`, `"Scorpio"`, `"Sagittarius"`, `"Capricorn"`, `"Aquarius"`, `"Pisces"`, wobei nach `"Pisces"` wieder zu `"Aries"` zurückgekehrt wird. Finden Sie die kürzeste Entfernung zwischen den beiden Zeichen und geben Sie die Kompatibilität zurück:

| Entfernung | Kompatibilität |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscopeMatch("Libra", "Sagittarius")` sollte `"80%"` zurückgeben.

```js
assert.equal(horoscopeMatch("Libra", "Sagittarius"), "80%");
```

`horoscopeMatch("Gemini", "Scorpio")` sollte `"20%"` zurückgeben.

```js
assert.equal(horoscopeMatch("Gemini", "Scorpio"), "20%");
```

`horoscopeMatch("Pisces", "Aries")` sollte `"40%"` zurückgeben.

```js
assert.equal(horoscopeMatch("Pisces", "Aries"), "40%");
```

`horoscopeMatch("Capricorn", "Cancer")` sollte `"50%"` zurückgeben.

```js
assert.equal(horoscopeMatch("Capricorn", "Cancer"), "50%");
```

`horoscopeMatch("Aquarius", "Aquarius")` sollte `"100%"` zurückgeben.

```js
assert.equal(horoscopeMatch("Aquarius", "Aquarius"), "100%");
```

`horoscopeMatch("Virgo", "Taurus")` sollte `"90%"` zurückgeben.

```js
assert.equal(horoscopeMatch("Virgo", "Taurus"), "90%");
```

`horoscopeMatch("Leo", "Scorpio")` sollte `"30%"` zurückgeben.

```js
assert.equal(horoscopeMatch("Leo", "Scorpio"), "30%");
```

# --seed--

## --seed-contents--

```js
function horoscopeMatch(sign1, sign2) {

  return sign1;
}
```

# --solutions--

```js
function horoscopeMatch(sign1, sign2) {
  const signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"];
  
  const a = signs.indexOf(sign1) + 1;
  const b = signs.indexOf(sign2) + 1;
  const diff = Math.abs(a - b);
  const distance = Math.min(diff, 12 - diff);

  const compatibility = { 0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%" };
  return compatibility[distance];
}
```
