---
id: 6a2037a68a0bc2aef0006003
title: "Aufgabe 341: Countdown zum Geburtstag"
challengeType: 28
dashedName: challenge-341
---

# --description--

Geben Sie das heutige Datum und einen Geburtstag vor, und geben Sie die Anzahl der Tage bis zum nächsten Geburtstag der Person zurück.

- Das heutige Datum wird als String im Format `"YYYY-MM-DD"` mit führenden Nullen angegeben, zum Beispiel: `"2026-07-16"`.
- Der Geburtstag wird als String im Format `"M/D"` ohne führende Nullen angegeben, zum Beispiel: `"9/7"`.
- Wenn heute ihr Geburtstag ist, geben Sie die Anzahl der Tage bis zum nächsten Geburtstag zurück (nicht `0`).
- Schaltjahre müssen berücksichtigt werden.

# --hints--

`daysUntilBirthday("2026-07-16", "9/7")` sollte `53` zurückgeben.

```js
assert.equal(daysUntilBirthday("2026-07-16", "9/7"), 53);
```

`daysUntilBirthday("2026-07-16", "3/22")` sollte `249` zurückgeben.

```js
assert.equal(daysUntilBirthday("2026-07-16", "3/22"), 249);
```

`daysUntilBirthday("2026-07-16", "7/16")` sollte `365` zurückgeben.

```js
assert.equal(daysUntilBirthday("2026-07-16", "7/16"), 365);
```

`daysUntilBirthday("2024-02-28", "3/1")` sollte `2` zurückgeben.

```js
assert.equal(daysUntilBirthday("2024-02-28", "3/1"), 2);
```

`daysUntilBirthday("2023-04-24", "12/30")` sollte `250` zurückgeben.

```js
assert.equal(daysUntilBirthday("2023-04-24", "12/30"), 250);
```

`daysUntilBirthday("2024-03-01", "2/29")` sollte `1460` zurückgeben.

```js
assert.equal(daysUntilBirthday("2024-03-01", "2/29"), 1460);
```

`daysUntilBirthday("2096-03-01", "2/29")` sollte `2920` zurückgeben.

```js
assert.equal(daysUntilBirthday("2096-03-01", "2/29"), 2920);
```

# --seed--

## --seed-contents--

```js
function daysUntilBirthday(today, birthday) {

  return today;
}
```

# --solutions--

```js
function daysUntilBirthday(today, birthday) {
  const [year, month, day] = today.split("-").map(Number);
  const [bMonth, bDay] = birthday.split("/").map(Number);

  const todayDate = new Date(year, month - 1, day);

  function isLeapYear(y) {
    return (y % 4 === 0 && y % 100 !== 0) || y % 400 === 0;
  }

  function nextValidBirthday(fromYear) {
    for (let y = fromYear; y <= fromYear + 8; y++) {
      if (bMonth === 2 && bDay === 29 && !isLeapYear(y)) continue;
      const date = new Date(y, bMonth - 1, bDay);
      if (date > todayDate) return date;
    }
  }

  const nextBirthday = nextValidBirthday(year);
  const diff = nextBirthday - todayDate;
  return Math.round(diff / (1000 * 60 * 60 * 24));
}
```
