---
id: 6a2037a68a0bc2aef0006004
title: "Aufgabe 342: Würfelwahrscheinlichkeiten"
challengeType: 28
dashedName: challenge-342
---

# --description--

Gegeben ist eine Anzahl von sechsseitigen Würfeln, die geworfen werden, und eine Zielsumme. Geben Sie die Wahrscheinlichkeit zurück, diese Summe zu würfeln, als String im Format `"1 in X"`.

- Die Anzahl der Würfel liegt zwischen 1 und 6.
- Die Zielsumme ist mit der gegebenen Anzahl von Würfeln immer erreichbar.
- Runden Sie `"X"` auf die nächste ganze Zahl.

# --hints--

`getOdds(1, 5)` sollte `"1 in 6"` zurückgeben.

```js
assert.equal(getOdds(1, 5), "1 in 6");
```

`getOdds(2, 4)` sollte `"1 in 12"` zurückgeben.

```js
assert.equal(getOdds(2, 4), "1 in 12");
```

`getOdds(3, 10)` sollte `"1 in 8"` zurückgeben.

```js
assert.equal(getOdds(3, 10), "1 in 8");
```

`getOdds(4, 7)` sollte `"1 in 65"` zurückgeben.

```js
assert.equal(getOdds(4, 7), "1 in 65");
```

`getOdds(5, 26)` sollte `"1 in 111"` zurückgeben.

```js
assert.equal(getOdds(5, 26), "1 in 111");
```

`getOdds(6, 35)` sollte `"1 in 7776"` zurückgeben.

```js
assert.equal(getOdds(6, 35), "1 in 7776");
```

# --seed--

## --seed-contents--

```js
function getOdds(dice, target) {

  return dice;
}
```

# --solutions--

```js
function getOdds(dice, target) {
  let validCombos = 0;
  const totalCombos = Math.pow(6, dice);

  function roll(diceLeft, sum) {
    if (diceLeft === 0) {
      if (sum === target) validCombos++;
      return;
    }
    for (let face = 1; face <= 6; face++) {
      roll(diceLeft - 1, sum + face);
    }
  }

  roll(dice, 0);
  return `1 in ${Math.round(totalCombos / validCombos)}`;
}
```
