---
id: 6a2037a68a0bc2aef0006002
title: "Aufgabe 340: Pig Latin-Konverter"
challengeType: 28
dashedName: challenge-340
---

# --description--

Gegeben ist ein String, den Sie nach den folgenden Regeln in Pig Latin umwandeln sollen:

- Beginnt ein Wort mit einem Vokal (`"a"`, `"e"`, `"i"`, `"o"` oder `"u"`), fügen Sie am Ende `"way"` hinzu. Zum Beispiel wird aus `"universe"` `"universeway"`.
- Beginnt ein Wort mit einem oder mehreren Konsonanten, verschieben Sie diese an das Ende und fügen `"ay"` hinzu. Zum Beispiel wird aus `"hello"` `"ellohay"`.
- Bewahren Sie die Groß- und Kleinschreibung des ersten Buchstabens. Zum Beispiel wird aus `"Hello"` `"Ellohay"`.

# --hints--

`pigLatin("universe")` sollte `"universeway"` zurückgeben.

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")` sollte `"ellohay"` zurückgeben.

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")` sollte `"ellohay universeway"` zurückgeben.

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")` sollte `"Ellohay universeway"` zurückgeben.

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")` sollte `"Igpay Atinlay isway unfay"` zurückgeben.

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")` sollte `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"` zurückgeben.

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
