---
id: 6a2037a68a0bc2aef0005fff
title: "Aufgabe 337: Zählwerk"
challengeType: 28
dashedName: challenge-337
---

# --description--

Gegeben ist ein String aus Strichlisten. Geben Sie die insgesamt dargestellte Anzahl zurück.

- Jeder senkrechte Strich `"|"` stellt eine Zählung dar.
- Jeder fünfte Strich wird als Schrägstrich `"/"` dargestellt und vervollständigt eine Gruppe von fünf (`"||||/"`).
- Die Gruppen sind durch ein Leerzeichen getrennt.

# --hints--

`getTallyCount("||||")` sollte `4` zurückgeben.

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` sollte `5` zurückgeben.

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` sollte `8` zurückgeben.

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` sollte `17` zurückgeben.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` sollte `41` zurückgeben.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
