---
id: 6a2037a68a0bc2aef0006005
title: "Aufgabe 343: Aufzug-Stopps"
challengeType: 28
dashedName: challenge-343
---

# --description--

Gegeben ist eine Zahl für die aktuelle Etage eines Aufzugs und ein Array mit angefragten Etagen. Geben Sie ein Array mit der Reihenfolge zurück, in der der Aufzug diese Etagen besuchen soll, um die Anzahl der zurückgelegten Etagen zu minimieren.

- Bei Gleichstand zuerst nach oben fahren
- Etagen mit einer Anfrage müssen besucht werden, wenn der Aufzug sie zum ersten Mal passiert

# --hints--

`elevatorStops(5, [2, 8, 3, 9])` sollte `[3, 2, 8, 9]` zurückgeben.

```js
assert.deepEqual(elevatorStops(5, [2, 8, 3, 9]), [3, 2, 8, 9]);
```

`elevatorStops(6, [2, 10, 8, 3, 1, 9])` sollte `[8, 9, 10, 3, 2, 1]` zurückgeben.

```js
assert.deepEqual(elevatorStops(6, [2, 10, 8, 3, 1, 9]), [8, 9, 10, 3, 2, 1]);
```

`elevatorStops(1, [4, 8, 3, 6, 9])` sollte `[3, 4, 6, 8, 9]` zurückgeben.

```js
assert.deepEqual(elevatorStops(1, [4, 8, 3, 6, 9]), [3, 4, 6, 8, 9]);
```

`elevatorStops(12, [6, 10, 7, 3, 1, 4])` sollte `[10, 7, 6, 4, 3, 1]` zurückgeben.

```js
assert.deepEqual(elevatorStops(12, [6, 10, 7, 3, 1, 4]), [10, 7, 6, 4, 3, 1]);
```

`elevatorStops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19])` sollte `[10, 9, 8, 6, 5, 2, 12, 19, 23]` zurückgeben.

```js
assert.deepEqual(elevatorStops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19]), [10, 9, 8, 6, 5, 2, 12, 19, 23]);
```

# --seed--

## --seed-contents--

```js
function elevatorStops(currentFloor, stops) {

  return currentFloor;
}
```

# --solutions--

```js
function elevatorStops(currentFloor, stops) {
  stops = stops.sort((a, b) => a - b);
  const above = stops.filter(s => s > currentFloor);
  const below = stops.filter(s => s < currentFloor).reverse();
  const upDistance = above.length ? above[above.length - 1] - currentFloor : Infinity;
  const downDistance = below.length ? currentFloor - below[below.length - 1] : Infinity;
  return downDistance < upDistance ? [...below, ...above] : [...above, ...below];
}
```
