---
id: 6a2037a68a0bc2aef0006007
title: "Aufgabe 345: Wortmischer"
challengeType: 28
dashedName: challenge-345
---

# --description--

Gegeben sind zwei Wörter. Geben Sie ein neues Wort zurück, indem Sie die erste Hälfte des ersten Wortes mit der zweiten Hälfte des zweiten Wortes kombinieren.

- Bei Wörtern mit ungerader Länge ist die erste Hälfte die kürzere Hälfte.

# --hints--

`blendWords("turtle", "toucan")` sollte `"turcan"` zurückgeben.

```js
assert.equal(blendWords("turtle", "toucan"), "turcan");
```

`blendWords("chipmunk", "flamingo")` sollte `"chipingo"` zurückgeben.

```js
assert.equal(blendWords("chipmunk", "flamingo"), "chipingo");
```

`blendWords("falcon", "pelican")` sollte `"falican"` zurückgeben.

```js
assert.equal(blendWords("falcon", "pelican"), "falican");
```

`blendWords("hyena", "iguana")` sollte `"hyana"` zurückgeben.

```js
assert.equal(blendWords("hyena", "iguana"), "hyana");
```

`blendWords("scorpion", "gorilla")` sollte `"scorilla"` zurückgeben.

```js
assert.equal(blendWords("scorpion", "gorilla"), "scorilla");
```

`blendWords("platypus", "wolverine")` sollte `"platerine"` zurückgeben.

```js
assert.equal(blendWords("platypus", "wolverine"), "platerine");
```

# --seed--

## --seed-contents--

```js
function blendWords(word1, word2) {

  return word1;
}
```

# --solutions--

```js
function blendWords(word1, word2) {
  const firstHalf = word1.slice(0, Math.floor(word1.length / 2));
  const secondHalf = word2.slice(Math.floor(word2.length / 2));
  return firstHalf + secondHalf;
}
```
