---
id: 6a2037a68a0bc2aef0006000
title: "Aufgabe 338: Haustier-Alter-Rechner"
challengeType: 28
dashedName: challenge-338
---

# --description--

Geben Sie einen Haustiertyp und das Alter in Menschenjahren an, und geben Sie das entsprechende Alter in Tierjahren anhand der folgenden Umrechnungstabelle zurück:

| Haustier | Multiplikator |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`petYears("dog", 5)` sollte `35` zurückgeben.

```js
assert.equal(petYears("dog", 5), 35);
```

`petYears("cat", 9)` sollte `54` zurückgeben.

```js
assert.equal(petYears("cat", 9), 54);
```

`petYears("rabbit", 3)` sollte `24` zurückgeben.

```js
assert.equal(petYears("rabbit", 3), 24);
```

`petYears("hamster", 4)` sollte `120` zurückgeben.

```js
assert.equal(petYears("hamster", 4), 120);
```

`petYears("guinea pig", 5)` sollte `60` zurückgeben.

```js
assert.equal(petYears("guinea pig", 5), 60);
```

`petYears("goldfish", 2)` sollte `12` zurückgeben.

```js
assert.equal(petYears("goldfish", 2), 12);
```

`petYears("bird", 1)` sollte `5` zurückgeben.

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
