---
id: 6a2037a68a0bc2aef0006006
title: "Aufgabe 344: Goldener Schnitt"
challengeType: 28
dashedName: challenge-344
---

# --description--

Gegeben sind zwei Zahlen. Bestimmen Sie, ob ihr Verhältnis ungefähr dem goldenen Schnitt entspricht.

- Verwenden Sie einen goldenen Schnitt von `1.618`
- Erlauben Sie eine Toleranz von `0.01`

# --hints--

`isGoldenRatio(21, 34)` sollte `true` zurückgeben.

```js
assert.isTrue(isGoldenRatio(21, 34));
```

`isGoldenRatio(15, 20)` sollte `false` zurückgeben.

```js
assert.isFalse(isGoldenRatio(15, 20));
```

`isGoldenRatio(8, 13)` sollte `true` zurückgeben.

```js
assert.isTrue(isGoldenRatio(8, 13));
```

`isGoldenRatio(10, 16)` sollte `false` zurückgeben.

```js
assert.isFalse(isGoldenRatio(10, 16));
```

`isGoldenRatio(1618, 1000)` sollte `true` zurückgeben.

```js
assert.isTrue(isGoldenRatio(1618, 1000));
```

`isGoldenRatio(88, 55)` sollte `false` zurückgeben.

```js
assert.isFalse(isGoldenRatio(88, 55));
```

# --seed--

## --seed-contents--

```js
function isGoldenRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function isGoldenRatio(a, b) {
  const goldenRatio = 1.618;
  const ratio = Math.max(a, b) / Math.min(a, b);
  return Math.abs(ratio - goldenRatio) <= 0.01;
}
```
