---
id: 6a2037a68a0bc2aef0006001
title: "Aufgabe 339: Array-Teilarrays"
challengeType: 28
dashedName: challenge-339
---

# --description--

Gegeben ist ein Array und eine Chunk-Größe. Geben Sie das Array aufgeteilt in Teilarrays dieser Größe zurück.

- Der letzte Teilarray kann kleiner sein, wenn das Array nicht gleichmäßig teilbar ist.

# --hints--

`chunkArray([1, 2, 3, 4, 5, 6], 3)` sollte `[[1, 2, 3], [4, 5, 6]]` zurückgeben.

```js
assert.deepEqual(chunkArray([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]]);
```

`chunkArray([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` sollte `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]` zurückgeben.

```js
assert.deepEqual(chunkArray([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]);
```

`chunkArray([1, 2, 3, 4, 5], 3)` sollte `[[1, 2, 3], [4, 5]]` zurückgeben.

```js
assert.deepEqual(chunkArray([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]]);
```

`chunkArray(["a", "b", "c", "d", "e"], 1)` sollte `[["a"], ["b"], ["c"], ["d"], ["e"]]` zurückgeben.

```js
assert.deepEqual(chunkArray(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]]);
```

`chunkArray([1, 2, 3], 5)` sollte `[[1, 2, 3]]` zurückgeben.

```js
assert.deepEqual(chunkArray([1, 2, 3], 5), [[1, 2, 3]]);
```

# --seed--

## --seed-contents--

```js
function chunkArray(arr, size) {

  return arr;
}
```

# --solutions--

```js
function chunkArray(arr, size) {
  const result = [];
  for (let i = 0; i < arr.length; i += size) {
    result.push(arr.slice(i, i + size));
  }
  return result;
}
```
