---
id: 6a22d77ddf034bc4e35b1d5a
title: "Aufgabe 353: Kontrastbewertung 2"
challengeType: 28
dashedName: challenge-353
---

# --description--

Gegeben sind zwei relative Leuchtdichtewerte und ein Boolean, der angibt, ob der Text groß ist. Geben Sie die WCAG-Kontrastbewertung mit der folgenden Methode zurück:

Berechnen Sie das Kontrastverhältnis, indem Sie zu jedem Leuchtdichtewert 0,05 addieren und dann den helleren Wert durch den dunkleren teilen. Der hellere Wert ist immer das erste Argument.

Geben Sie die Bewertung basierend auf dem Kontrastverhältnis mit der folgenden Tabelle zurück:

| Bewertung | Normaler Text | Großer Text |
|--------|-------------|------------|
| `"AAA"` | 7,0+ | 4,5+ |
| `"AA"` | 4,5+ | 3,0+ |
| `"Fail"` | unter 4,5 | unter 3,0 |

# --hints--

`getContrastRating(1.0, 0.0, false)` sollte `"AAA"` zurückgeben.

```js
assert.equal(getContrastRating(1.0, 0.0, false), "AAA");
```

`getContrastRating(0.9015, 0.1364, false)` sollte `"AA"` zurückgeben.

```js
assert.equal(getContrastRating(0.9015, 0.1364, false), "AA");
```

`getContrastRating(0.8965, 0.1628, false)` sollte `"Fail"` zurückgeben.

```js
assert.equal(getContrastRating(0.8965, 0.1628, false), "Fail");
```

`getContrastRating(0.7469, 0.0957, true)` sollte `"AAA"` zurückgeben.

```js
assert.equal(getContrastRating(0.7469, 0.0957, true), "AAA");
```

`getContrastRating(0.7489, 0.2018, true)` sollte `"AA"` zurückgeben.

```js
assert.equal(getContrastRating(0.7489, 0.2018, true), "AA");
```

`getContrastRating(0.6571, 0.1974, true)` sollte `"Fail"` zurückgeben.

```js
assert.equal(getContrastRating(0.6571, 0.1974, true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(l1, l2, isLargeText) {

  return l1;
}
```

# --solutions--

```js
function getContrastRating(l1, l2, isLargeText) {
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
