---
id: 6a22d77ddf034bc4e35b1d58
title: "Aufgabe 351: Pronic-Zahl"
challengeType: 28
dashedName: challenge-351
---

# --description--

Gegeben ist eine Zahl. Bestimmen Sie, ob es sich um eine Pronic-Zahl handelt.

Eine Pronic-Zahl ist das Produkt von zwei aufeinanderfolgenden ganzen Zahlen. Zum Beispiel ist 6 eine Pronic-Zahl, weil 2 * 3 = 6 gilt.

# --hints--

`isPronic(6)` sollte `true` zurückgeben.

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` sollte `false` zurückgeben.

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` sollte `true` zurückgeben.

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` sollte `true` zurückgeben.

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` sollte `false` zurückgeben.

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` sollte `true` zurückgeben.

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
