---
id: 6a22d77ddf034bc4e35b1d57
title: "Aufgabe 350: Buchstabenabstand"
challengeType: 28
dashedName: challenge-350
---

# --description--

Gegeben sind zwei Strings gleicher Länge. Geben Sie die Summe der kürzesten Abstände zwischen jedem Buchstabenpaar zurück.

- Die Eingabe enthält nur Kleinbuchstaben
- Das Alphabet wird als Kreis behandelt, daher beträgt der Abstand zwischen `a` und `z` 1.

# --hints--

`letterDistance("abc", "bcd")` sollte `3` zurückgeben.

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")` sollte `9` zurückgeben.

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")` sollte `10` zurückgeben.

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")` sollte `43` zurückgeben.

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")` sollte `47` zurückgeben.

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")` sollte `55` zurückgeben.

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
