---
id: 6a22d77ddf034bc4e35b1d59
title: "Aufgabe 352: Kontrastbewertung 1"
challengeType: 28
dashedName: challenge-352
---

# --description--

Geben Sie ein Kontrastverhältnis und einen Boolean an, der angibt, ob der Text groß ist, und geben Sie die WCAG-Bewertung anhand der folgenden Tabelle zurück:

| Bewertung | Normaler Text | Großer Text |
|--------|-------------|------------|
| `"AAA"` | 7,0+ | 4,5+ |
| `"AA"` | 4,5+ | 3,0+ |
| `"Fail"` | unter 4,5 | unter 3,0 |

# --hints--

`getContrastRating("7.5", false)` sollte `"AAA"` zurückgeben.

```js
assert.equal(getContrastRating("7.5", false), "AAA");
```

`getContrastRating("4.8", false)` sollte `"AA"` zurückgeben.

```js
assert.equal(getContrastRating("4.8", false), "AA");
```

`getContrastRating("4.2", false)` sollte `"Fail"` zurückgeben.

```js
assert.equal(getContrastRating("4.2", false), "Fail");
```

`getContrastRating("4.5", true)` sollte `"AAA"` zurückgeben.

```js
assert.equal(getContrastRating("4.5", true), "AAA");
```

`getContrastRating("3.0", true)` sollte `"AA"` zurückgeben.

```js
assert.equal(getContrastRating("3.0", true), "AA");
```

`getContrastRating("2.7", false)` sollte `"Fail"` zurückgeben.

```js
assert.equal(getContrastRating("2.7", false), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(ratio, isLargeText) {

  return ratio;
}
```

# --solutions--

```js
function getContrastRating(ratio, isLargeText) {
  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
