---
id: 6a22d77ddf034bc4e35b1d54
title: "Aufgabe 347: Spieltheorie"
challengeType: 28
dashedName: challenge-347
---

# --description--

Gegeben sind zwei Strings gleicher Länge, die die Strategien von zwei Spielern für ein Spiel darstellen. Geben Sie die Punktzahlen als Array `[player1, player2]` zurück.

- Die gegebenen Strings enthalten nur einen von zwei Buchstaben: `"C"` (kooperieren) oder `"D"` (defekt).
- Jeder Buchstabe steht für eine Runde, die wie folgt bewertet wird:
  - Wenn beide Spieler kooperieren, erhält jeder 3 Punkte.
  - Wenn beide Spieler defekt spielen, erhält jeder 1 Punkt.
  - Wenn ein Spieler defekt spielt und der andere kooperiert, erhält der Defektor 5 Punkte und der Kooperierende 0 Punkte.

# --hints--

`playGame("CCCC", "CCCC")` sollte `[12, 12]` zurückgeben.

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")` sollte `[4, 4]` zurückgeben.

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")` sollte `[5, 10]` zurückgeben.

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")` sollte `[24, 34]` zurückgeben.

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` sollte `[66, 21]` zurückgeben.

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
