---
id: 6a22d77ddf034bc4e35b1d56
title: "Aufgabe 349: Handysignal"
challengeType: 28
dashedName: challenge-349
---

# --description--

Gegeben ist ein Raster mit drei Messwerten von Mobilfunkmasten. Bestimmen Sie den Standort des Telefons.

- Jede Zelle im Raster ist entweder `0` (kein Turm) oder eine positive ganze Zahl, die die Anzahl der Zellen bis zum Telefon darstellt, gemessen in einer geraden Linie: horizontal, vertikal oder diagonal.
- Geben Sie die `[row, col]` der Zelle zurück, die von allen drei Türmen die korrekte Anzahl an Zellen entfernt ist.
- Es gibt immer genau eine Lösung.

# --hints--

`findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` sollte `[1, 2]` zurückgeben.

```js
assert.deepEqual(findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2]);
```

`findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` sollte `[2, 1]` zurückgeben.

```js
assert.deepEqual(findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1]);
```

`findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` sollte `[2, 2]` zurückgeben.

```js
assert.deepEqual(findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2]);
```

`findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` sollte `[3, 4]` zurückgeben.

```js
assert.deepEqual(findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4]);
```

`findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` sollte `[3, 3]` zurückgeben.

```js
assert.deepEqual(findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3]);
```

# --seed--

## --seed-contents--

```js
function findSignal(grid) {

  return grid;
}
```

# --solutions--

```js
function findSignal(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const towers = [];
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) towers.push([r, c, grid[r][c]]);
    }
  }

  function isValid(r, c, tr, tc, dist) {
    const dr = Math.abs(r - tr);
    const dc = Math.abs(c - tc);
    return (dr === dist || dc === dist) && (dr === 0 || dc === 0 || dr === dc);
  }

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) continue;
      if (towers.every(([tr, tc, dist]) => isValid(r, c, tr, tc, dist))) {
        return [r, c];
      }
    }
  }
}
```
