---
id: 6a26df3e2988bcdded204893
title: "Aufgabe 357: Nahrungskette"
challengeType: 28
dashedName: challenge-357
---

# --description--

Gegeben ist ein Array von `[predator, prey]`-Paaren. Geben Sie die Nahrungskette vom Spitzenprädator bis zum unteren Ende zurück.

- Der Spitzenprädator ist das Tier, das niemals Beute eines anderen Tieres ist.
- Geben Sie die Kette als Array von Strings zurück.

# --hints--

`getFoodChain([["cat", "mouse"]])` sollte `["cat", "mouse"]` zurückgeben.

```js
assert.deepEqual(getFoodChain([["cat", "mouse"]]), ["cat", "mouse"]);
```

`getFoodChain([["wolf", "deer"], ["deer", "grass"]])` sollte `["wolf", "deer", "grass"]` zurückgeben.

```js
assert.deepEqual(getFoodChain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"]);
```

`getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` sollte `["hawk", "snake", "frog", "fly"]` zurückgeben.

```js
assert.deepEqual(getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"]);
```

`getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` sollte `["eagle", "fox", "rabbit", "grass"]` zurückgeben.

```js
assert.deepEqual(getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"]);
```

`getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` sollte `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]` zurückgeben.

```js
assert.deepEqual(getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"]);
```

# --seed--

## --seed-contents--

```js
function getFoodChain(pairs) {

  return pairs;
}
```

# --solutions--

```js
function getFoodChain(pairs) {
  const prey = new Set(pairs.map(([, p]) => p));
  const map = Object.fromEntries(pairs);
  const apex = pairs.map(([p]) => p).find(p => !prey.has(p));
  const chain = [apex];
  while (map[chain.at(-1)]) chain.push(map[chain.at(-1)]);
  return chain;
}
```
