---
id: 6a26df3e2988bcdded204896
title: "Aufgabe 360: Gesprochene Dauer"
challengeType: 28
dashedName: challenge-360
---

# --description--

Geben Sie eine Anzahl von Sekunden vor, und geben Sie die Dauer in gesprochenem Englisch zurück.

- Zerlegen Sie die Dauer in Stunden, Minuten und Sekunden.
- Überspringen Sie alle Nullwerte.
- Verwenden Sie Singular oder Plural, je nachdem, was passend ist (`"1 hour"`, `"2 hours"`).
- Falls vorhanden, verbinden Sie die letzten beiden Einheiten mit `"and"` und die zweit- und drittletzten Einheiten mit einem Komma (`"1 hour, 2 minutes and 3 seconds"`).

# --hints--

`getSpokenDuration(3723)` sollte `"1 hour, 2 minutes and 3 seconds"` zurückgeben.

```js
assert.equal(getSpokenDuration(3723), "1 hour, 2 minutes and 3 seconds");
```

`getSpokenDuration(7295)` sollte `"2 hours, 1 minute and 35 seconds"` zurückgeben.

```js
assert.equal(getSpokenDuration(7295), "2 hours, 1 minute and 35 seconds");
```

`getSpokenDuration(8521)` sollte `"2 hours, 22 minutes and 1 second"` zurückgeben.

```js
assert.equal(getSpokenDuration(8521), "2 hours, 22 minutes and 1 second");
```

`getSpokenDuration(435)` sollte `"7 minutes and 15 seconds"` zurückgeben.

```js
assert.equal(getSpokenDuration(435), "7 minutes and 15 seconds");
```

`getSpokenDuration(14455)` sollte `"4 hours and 55 seconds"` zurückgeben.

```js
assert.equal(getSpokenDuration(14455), "4 hours and 55 seconds");
```

`getSpokenDuration(72000)` sollte `"20 hours"` zurückgeben.

```js
assert.equal(getSpokenDuration(72000), "20 hours");
```

`getSpokenDuration(1)` sollte `"1 second"` zurückgeben.

```js
assert.equal(getSpokenDuration(1), "1 second");
```

# --seed--

## --seed-contents--

```js
function getSpokenDuration(seconds) {

  return seconds;
}
```

# --solutions--

```js
function getSpokenDuration(seconds) {
  const hours = Math.floor(seconds / 3600);
  const minutes = Math.floor((seconds % 3600) / 60);
  const secs = seconds % 60;

  const parts = [];
  if (hours) parts.push(`${hours} ${hours === 1 ? "hour" : "hours"}`);
  if (minutes) parts.push(`${minutes} ${minutes === 1 ? "minute" : "minutes"}`);
  if (secs) parts.push(`${secs} ${secs === 1 ? "second" : "seconds"}`);

  if (parts.length === 1) return parts[0];
  if (parts.length === 2) return `${parts[0]} and ${parts[1]}`;
  return `${parts[0]}, ${parts[1]} and ${parts[2]}`;
}
```
