---
id: 6a26df95efa55a2524399743
title: "Aufgabe 362: Nonogramm-Validator"
challengeType: 28
dashedName: challenge-362
---

# --description--

Gegeben sind ein Array von Hinweiszahlen und ein Array von Zellen. Bestimmen Sie, ob die Zellen den Nonogramm-Hinweis erfüllen.

- Der Hinweis ist ein Array von Zahlen, die die Längen aufeinanderfolgender gefüllter Zellen in der Reihenfolge darstellen. Zum Beispiel bedeutet ein Hinweis von `[3, 2]`, dass es 3 aufeinanderfolgende gefüllte Zellen geben soll, gefolgt von 2 aufeinanderfolgenden gefüllten Zellen, getrennt durch mindestens eine leere Zelle.
- Die Zeile ist ein Array aus 1en (gefüllt) und 0en (leer).

# --hints--

`isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1])` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1]));
```

`isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1])` sollte `false` zurückgeben.

```js
assert.isFalse(isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1]));
```

`isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])` sollte `false` zurückgeben.

```js
assert.isFalse(isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]));
```

`isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])` sollte `true` zurückgeben.

```js
assert.isTrue(isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])` sollte `false` zurückgeben.

```js
assert.isFalse(isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]));
```

# --seed--

## --seed-contents--

```js
function isValidNonogram(clue, cells) {

  return clue;
}
```

# --solutions--

```js
function isValidNonogram(clue, cells) {
  const runs = [];
  let count = 0;

  for (let i = 0; i <= cells.length; i++) {
    if (cells[i] === 1) {
      count++;
    } else if (count > 0) {
      runs.push(count);
      count = 0;
    }
  }

  if (runs.length !== clue.length) return false;
  return runs.every((run, i) => run === clue[i]);
}
```
