---
id: 6a26df3e2988bcdded204897
title: "Aufgabe 361: Gesprochene Zeit"
challengeType: 28
dashedName: challenge-361
---

# --description--

Gegeben sind die Winkel für den Stunden- und Minutenzeiger einer analogen Uhr in Grad (im Uhrzeigersinn von 12 aus gemessen). Geben Sie die Zeit in gesprochenem Englisch zurück.

Wandeln Sie den Winkel des Minutenzeigers in Minuten um (360° = 60 Minuten) und verwenden Sie dann die folgenden Regeln:

| Minuten | Gesprochen |
|---------|------------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (außer 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (außer 45) | `"X minutes to Z"` (wobei `X` 60 minus Minuten ist) |

Dabei ist `Y` die aktuelle Stunde und `Z` die nächste Stunde, beide abgeleitet vom Winkel des Stundenzeigers (360° = 12 Stunden).

Hinweis: Die Zeigerwinkel landen möglicherweise nicht genau auf einer Zahl, berücksichtigen Sie daher eine geeignete Rundung.

# --hints--

`getSpokenTime(90, 0)` sollte `"3 o'clock"` zurückgeben.

```js
assert.equal(getSpokenTime(90, 0), "3 o'clock");
```

`getSpokenTime(160, 120)` sollte `"20 minutes past 5"` zurückgeben.

```js
assert.equal(getSpokenTime(160, 120), "20 minutes past 5");
```

`getSpokenTime(255, 180)` sollte `"half past 8"` zurückgeben.

```js
assert.equal(getSpokenTime(255, 180), "half past 8");
```

`getSpokenTime(67.5, 92)` sollte `"quarter past 2"` zurückgeben.

```js
assert.equal(getSpokenTime(67.5, 92), "quarter past 2");
```

`getSpokenTime(200, 240)` sollte `"20 minutes to 7"` zurückgeben.

```js
assert.equal(getSpokenTime(200, 240), "20 minutes to 7");
```

`getSpokenTime(322.5, 273)` sollte `"quarter to 11"` zurückgeben.

```js
assert.equal(getSpokenTime(322.5, 273), "quarter to 11");
```

`getSpokenTime(117.5, 335)` sollte `"5 minutes to 4"` zurückgeben.

```js
assert.equal(getSpokenTime(117.5, 335), "5 minutes to 4");
```

# --seed--

## --seed-contents--

```js
function getSpokenTime(hourAngle, minuteAngle) {

  return hourAngle;
}
```

# --solutions--

```js
function getSpokenTime(hourAngle, minuteAngle) {
  const minutes = Math.floor(minuteAngle / 6);
  const hour = Math.floor(hourAngle / 30) || 12;
  const nextHour = (hour % 12) + 1;

  if (minutes === 0) return `${hour} o'clock`;
  if (minutes === 15) return `quarter past ${hour}`;
  if (minutes === 30) return `half past ${hour}`;
  if (minutes === 45) return `quarter to ${nextHour}`;
  if (minutes < 30) return `${minutes} minutes past ${hour}`;
  return `${60 - minutes} minutes to ${nextHour}`;
}
```
