---
id: 6a26df3e2988bcdded204894
title: "Aufgabe 358: Emoji-Übersetzer"
challengeType: 28
dashedName: challenge-358
---

# --description--

Gegeben ist ein String aus Emojis. Geben Sie den Satz zurück, indem Sie die folgende Tabelle verwenden:

| Emoji | Wort |
|-------|------|
| 👶 | `"baby"` |
| 🐱 | `"cat"` |
| 🐕 | `"dog"` |
| 🐟 | `"fish"` |
| 🥵 | `"hot"` |
| 🧊 | `"ice"` |
| 🪨 | `"rock"` |
| 🦈 | `"shark"` |
| 🍲 | `"soup"` |
| ⭐ | `"star"` |

Geben Sie die Wörter durch Leerzeichen getrennt zurück.

# --hints--

`getEmojiPhrase("🪨⭐")` sollte `"rock star"` zurückgeben.

```js
assert.equal(getEmojiPhrase("🪨⭐"), "rock star");
```

`getEmojiPhrase("🥵🐕")` sollte `"hot dog"` zurückgeben.

```js
assert.equal(getEmojiPhrase("🥵🐕"), "hot dog");
```

`getEmojiPhrase("👶🦈")` sollte `"baby shark"` zurückgeben.

```js
assert.equal(getEmojiPhrase("👶🦈"), "baby shark");
```

`getEmojiPhrase("⭐🐟")` sollte `"star fish"` zurückgeben.

```js
assert.equal(getEmojiPhrase("⭐🐟"), "star fish");
```

`getEmojiPhrase("🧊🧊👶")` sollte `"ice ice baby"` zurückgeben.

```js
assert.equal(getEmojiPhrase("🧊🧊👶"), "ice ice baby");
```

`getEmojiPhrase("🐱🐟🍲")` sollte `"cat fish soup"` zurückgeben.

```js
assert.equal(getEmojiPhrase("🐱🐟🍲"), "cat fish soup");
```

# --seed--

## --seed-contents--

```js
function getEmojiPhrase(str) {

  return str;
}
```

# --solutions--

```js
function getEmojiPhrase(str) {
  const table = {
    '👶': 'baby',
    '🐱': 'cat',
    '🐕': 'dog',
    '🐟': 'fish',
    '🥵': 'hot',
    '🧊': 'ice',
    '🪨': 'rock',
    '🦈': 'shark',
    '🍲': 'soup',
    '⭐': 'star',
  };

  return [...str].map(emoji => table[emoji]).join(' ');
}
```
