---
id: 6a26df95efa55a2524399745
title: "Aufgabe 364: Zwischen zwei Eimern"
challengeType: 28
dashedName: challenge-364
---

# --description--

Gegeben sind zwei Farbeimer, jeweils mit einer RGB-Farbe und einem Füllstand. Geben Sie die gemischte RGB-Farbe als Array von drei ganzen Zahlen zurück.

- Jeder Eimer ist ein Objekt (JavaScript) oder ein Dictionary (Python) mit einer `color`-Eigenschaft (ein Array von drei ganzen Zahlen `[r, g, b]`) und einer `fullness`-Eigenschaft (0–100).
- Die gemischte Farbe ist ein gewichteter Durchschnitt jedes Kanals der beiden Farben basierend auf dem Füllstand, wobei jeder Kanal auf die nächste ganze Zahl gerundet wird.

# --hints--

`mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 })` sollte `[125, 125, 125]` zurückgeben.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 }), [125, 125, 125]);
```

`mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 })` sollte `[200, 200, 200]` zurückgeben.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 }), [200, 200, 200]);
```

`mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 })` sollte `[100, 150, 200]` zurückgeben.

```js
assert.deepEqual(mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 }), [100, 150, 200]);
```

`mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 })` sollte `[114, 184, 170]` zurückgeben.

```js
assert.deepEqual(mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 }), [114, 184, 170]);
```

`mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 })` sollte `[68, 163, 123]` zurückgeben.

```js
assert.deepEqual(mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 }), [68, 163, 123]);
```

# --seed--

## --seed-contents--

```js
function mixPaint(bucket1, bucket2) {

  return bucket1;
}
```

# --solutions--

```js
function mixPaint(bucket1, bucket2) {
  const total = bucket1.fullness + bucket2.fullness;
  return bucket1.color.map((channel, i) =>
    Math.round((channel * bucket1.fullness + bucket2.color[i] * bucket2.fullness) / total)
  );
}
```
