---
id: 6a26df3e2988bcdded204895
title: "Aufgabe 359: Golf-Handicap-Rechner"
challengeType: 28
dashedName: challenge-359
---

# --description--

Gegeben ist ein Array von Golf-Ergebnissen und ein entsprechendes Array von Kurs-Par-Werten. Geben Sie den Handicap-Index des Golfspielers mit der folgenden Methode zurück:

- Berechnen Sie für jede Runde die Differenz, indem Sie das Par vom Ergebnis subtrahieren, und geben Sie dann den Durchschnitt aller Differenzen gerundet auf eine Nachkommastelle zurück.

# --hints--

`calculateHandicap([72, 72, 72], [72, 72, 72])` sollte `0` zurückgeben.

```js
assert.equal(calculateHandicap([72, 72, 72], [72, 72, 72]), 0);
```

`calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72])` sollte `6` zurückgeben.

```js
assert.equal(calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72]), 6);
```

`calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36])` sollte `8.3` zurückgeben.

```js
assert.equal(calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3);
```

`calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])` sollte `8.8` zurückgeben.

```js
assert.equal(calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8);
```

`calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])` sollte `11.7` zurückgeben.

```js
assert.equal(calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7);
```

# --seed--

## --seed-contents--

```js
function calculateHandicap(scores, pars) {

  return scores;
}
```

# --solutions--

```js
function calculateHandicap(scores, pars) {
  const differentials = scores.map((score, i) => score - pars[i]);
  const average = differentials.reduce((sum, d) => sum + d, 0) / differentials.length;
  return Math.round(average * 10) / 10;
}
```
