---
id: 6a26df95efa55a2524399744
title: "Aufgabe 363: Fülleimer 2"
challengeType: 28
dashedName: challenge-363
---

# --description--

Gegeben ist ein 2D-Raster aus einbuchstabigen Farb-Strings und eine Ziel-Farbe. Geben Sie die minimale Anzahl von Flood-Fill-„Klicks“ zurück, die nötig sind, um das gesamte Raster in die Ziel-Farbe zu ändern.

- Jeder Klick ändert die Farbe der angeklickten Zelle und der gesamten Region verbundener Zellen derselben Farbe in die Ziel-Farbe.
- Zellen sind horizontal und vertikal verbunden (nicht diagonal).

# --hints--

`bucketFill([["R", "R"], ["R", "R"]], "G")` sollte `1` zurückgeben.

```js
assert.equal(bucketFill([["R", "R"], ["R", "R"]], "G"), 1);
```

`bucketFill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B")` sollte `0` zurückgeben.

```js
assert.equal(bucketFill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B"), 0);
```

`bucketFill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R")` sollte `3` zurückgeben.

```js
assert.equal(bucketFill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R"), 3);
```

`bucketFill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P")` sollte `5` zurückgeben.

```js
assert.equal(bucketFill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P"), 5);
```

`bucketFill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y")` sollte `12` zurückgeben.

```js
assert.equal(bucketFill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y"), 12);
```

# --seed--

## --seed-contents--

```js
function bucketFill(grid, targetColor) {

  return grid;
}
```

# --solutions--

```js
function bucketFill(grid, targetColor) {
  const rows = grid.length;
  const cols = grid[0].length;
  const visited = new Set();

  function dfs(r, c, color) {
    const key = `${r},${c}`;
    if (r < 0 || r >= rows || c < 0 || c >= cols) return;
    if (visited.has(key)) return;
    if (grid[r][c] !== color) return;
    visited.add(key);
    dfs(r + 1, c, color);
    dfs(r - 1, c, color);
    dfs(r, c + 1, color);
    dfs(r, c - 1, color);
  }

  let clicks = 0;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const key = `${r},${c}`;
      if (!visited.has(key) && grid[r][c] !== targetColor) {
        dfs(r, c, grid[r][c]);
        clicks++;
      }
    }
  }

  return clicks;
}
```
