---
id: 66ed41f912d0bb1dc62da5dd
title: Створіть JavaScript-бота для вікторини
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні вивести `"Hello! I'm your coding fun fact guide!"` у консоль як привітальне повідомлення для користувача.
1. Ви повинні оголосити три змінні: `botName`, `botLocation` і `favoriteLanguage`, які зберігатимуть ім’я бота, звідки він, а також його улюблену мову програмування відповідно.
1. Ви повинні вивести `"My name is (botName) and I live on (botLocation)."` у консоль.
1. Ви повинні вивести `"My favorite programming language is (favoriteLanguage)."` у консоль.
1. Ви повинні використати `let` для створення змінної `codingFact` і призначити їй рядок із цікавою інформацією про улюблену мову програмування вашого бота, включаючи використання змінної `favoriteLanguage`.
1. Ви повинні вивести `codingFact` у консоль.
1. Ви повинні повторно призначити змінну `codingFact` новим фактом про улюблену мову бота, знову використовуючи змінну `favoriteLanguage`.
1. Ви повинні знову вивести `codingFact` у консоль.
1. Ви повинні знову призначити змінну `codingFact` ще одним новим фактом про улюблену мову бота, використовуючи змінну `favoriteLanguage`.
1. Ви повинні втретє вивести `codingFact` у консоль.
1. Ви повинні вивести `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` у консоль як прощальне повідомлення від бота.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні вивести `"Hello! I'm your coding fun fact guide!"` у консоль.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

Ви повинні оголосити змінну `botName`. Перевірте орфографію та регістр.

```js
assert.exists(botName);
```

Ваша змінна `botName` має бути рядком.

```js
assert.isString(botName);
```

Ви повинні оголосити змінну `botLocation`. Перевірте орфографію та регістр.

```js
assert.exists(botLocation);
```

Ваша змінна `botLocation` має бути рядком.

```js
assert.isString(botLocation);
```

Ви повинні оголосити змінну `favoriteLanguage`. Перевірте орфографію та регістр.

```js
assert.exists(favoriteLanguage);
```

Ваша змінна `favoriteLanguage` має бути рядком.

```js
assert.isString(favoriteLanguage);
```

Ви повинні вивести в консоль `"My name is (botName) and I live on (botLocation)."` і додати змінні до рядка.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

Ви повинні вивести в консоль `"My favorite programming language is (favoriteLanguage)."` і додати змінну до рядка.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

Ви повинні використати `let` для оголошення нової змінної `codingFact`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

Ви повинні призначити `codingFact` значення, яке містить `favoriteLanguage`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

Ви повинні вивести `codingFact` у консоль.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

Ви повинні призначити нове значення `codingFact`, яке також містить `favoriteLanguage`, і вивести його в консоль.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

Ви повинні втретє призначити значення `codingFact`, яке також містить `favoriteLanguage`, і вивести його в консоль.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact); 
```

Ви повинні вивести в консоль `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` і додати значення змінних.

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
