---
id: 6853e58d52ec64349a5760ca
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop implementieren Sie den Algorithmus für den kürzesten Pfad. Sie schreiben eine Python-Funktion, die den kürzesten Pfad zwischen den Knoten in einem Graphen berechnet und außerdem den zurückgelegten Weg zurückgibt.

Zum Beispiel findet der Algorithmus in einem Graphen, in dem Städte durch Straßen mit unterschiedlichen Entfernungen verbunden sind, die kürzeste Route von einer Stadt zur anderen. Wenn Sie von Stadt A nach Stadt D reisen möchten, könnte der Algorithmus feststellen, dass der Weg A ⇨ B ⇨ C ⇨ D (insgesamt 15 km) kürzer ist als der direkte Weg A ⇨ D (20 km).

Um zu beginnen, definieren Sie eine Variable namens `INF` und weisen ihr den Wert `float('inf')` zu, der positive Unendlichkeit darstellt. Später verwenden Sie diese, um eine unendliche Entfernung zwischen zwei Knoten anzugeben.

# --hints--

Sie sollten die Variable `INF` am Anfang Ihrer Datei definieren.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

Sie sollten `float('inf')` der Variable `INF` zuweisen.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
