---
id: 686251fb49ebc489ba60ccb5
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Sie müssen eine 2D-Liste erstellen, um die Adjazenzmatrix des Graphen darzustellen. Diese Matrix wird verwendet, um die Gewichte der Kanten zwischen den Knoten im Graphen darzustellen.

Erstellen Sie eine Variable namens `adj_matrix` und weisen Sie ihr eine 2D-Liste zu, die den Graphen mit den folgenden Gewichten darstellt:

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Sie sollten eine Variable namens `adj_matrix` erstellen.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` sollte eine 2D-Liste sein, die die angegebenen Werte enthält.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
