---
id: 68625ad34aa3f58d7f3b9b07
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Sie müssen die Anzahl der Knoten im Graphen speichern. Dafür benötigen Sie eine Variable, die der Länge der Adjazenzmatrix entspricht.

Erstellen Sie innerhalb der `shortest_path`-Funktion eine Variable `n` und setzen Sie sie auf die Länge der `matrix`.

# --hints--

Sie sollten `n` mit der Länge der `matrix` initialisieren.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("n").is_equivalent("n = len(matrix)")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    pass

--fcc-editable-region--
```
