---
id: 686253f1b051998ad4904e3e
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Sie erstellen nun die Hauptfunktion, die drei Parameter entgegennimmt: die Adjazenzmatrix, den Startknoten und einen optionalen Zielknoten.

Erstellen Sie eine Funktion namens `shortest_path`, die drei Parameter entgegennimmt: `matrix`, `start_node` und `target_node`. Weisen Sie `None` als Standardwert für `target_node` zu.

Der Parameter `target_node` ist standardmäßig auf `None` gesetzt, was bedeutet, dass die Funktion, falls kein Zielknoten angegeben ist, die kürzesten Pfade vom Startknoten zu allen anderen Knoten im Graphen berechnen soll.

Fügen Sie vorerst eine `pass`-Anweisung im Funktionskörper hinzu.

# --hints--

Sie sollten eine Funktion namens `shortest_path` erstellen.

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

Die Funktion sollte drei Parameter entgegennehmen: `matrix`, `start_node` und `target_node`, wobei `target_node` den Standardwert `None` hat.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
