---
id: 686b82af323098e7a5203bf4
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Während der Algorithmus läuft, müssen Sie festhalten, welche Knoten Sie bereits besucht haben, damit Sie diese nicht mehr als einmal verarbeiten.

Erstellen Sie dazu eine Liste namens `visited` und initialisieren Sie sie mit `False` für jeden Knoten.

# --hints--

Sie sollten eine Liste namens `visited` innerhalb der Funktion `shortest_path` haben.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("visited")`
        )
      );
    }
});
```

Sie sollten die Liste `visited` mit `False` für jeden Knoten mithilfe von `[False] * n` initialisieren.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("visited").is_equivalent("visited = [False] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
--fcc-editable-region--
    
--fcc-editable-region--
```
