---
id: 686b7848fbaa41e0c2a7efef
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Jetzt, da Sie Ihre `distances`-Liste initialisiert haben, müssen Sie die Distanz für den Startknoten aktualisieren.

Da die Distanz vom Startknoten zu sich selbst immer `0` ist, setzen Sie den Wert am `start_node`-Index in der `distances`-Liste auf `0`.

# --hints--

Sie sollten `distances[start_node]` auf `0` setzen.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_stmt("distances[start_node] = 0")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n

--fcc-editable-region--
```
