---
id: 686b78bc13d518e1267448e4
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Zusätzlich zur Verfolgung der Distanzen müssen Sie auch die tatsächlich zurückgelegten Pfade speichern, um jeden Knoten zu erreichen.

Sie erstellen eine Liste, in der jeder Eintrag den Pfad speichert, der zu diesem Knoten führt. Anfangs enthält der Pfad jedes Knotens nur sich selbst.

List Comprehensions bieten eine prägnante Möglichkeit, Listen zu erstellen. Zum Beispiel:

```py
[x * 2 for x in range(3)]
```

Erstellen Sie eine Variable namens `paths` und initialisieren Sie sie mit einer List Comprehension, die eine Liste mit `[node_no]` für jedes `node_no` in `range(n)` erstellt.

# --hints--

Sie sollten eine Variable namens `paths` haben.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

Sie sollten `paths` mit einer List Comprehension initialisieren.

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

Ihre List Comprehension sollte `node_no` verwenden, um `range(n)` zu durchlaufen.

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

Ihre List Comprehension sollte `[node_no]` für jedes `node_no` in `range(n)` auswerten.

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
