---
id: 68760464eb9c2e79cc913fac
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

In diesem Schritt fügen Sie eine Schleife hinzu, die einmal für jeden Knoten im Graphen ausgeführt wird. Diese Schleife ermöglicht es dem Algorithmus, Entfernungen und Pfade über mehrere Durchläufe zu aktualisieren.

Erstellen Sie eine `for`-Schleife, die `n`-mal ausgeführt wird. Verwenden Sie `_` als Schleifenvariable, da Sie den Iterationswert nicht benötigen.

Innerhalb der Schleife müssen Sie sich darauf vorbereiten, den nächsten zu verarbeitenden Knoten auszuwählen, indem Sie zwei Variablen erstellen:

- eine, um die bisher kleinste in der aktuellen Iteration gefundene Entfernung zu speichern  
- und eine andere, um den Index des Knotens zu speichern, der diese kleinste Entfernung hat.

Erstellen Sie die Variablen `min_distance` und `current` und setzen Sie sie jeweils auf `INF` und `-1`.

# --hints--

Sie sollten eine `for`-Schleife innerhalb der `shortest_path`-Funktion haben.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

Ihre `for`-Schleife sollte `range(n)` durchlaufen.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

Ihre for-Schleife sollte `_` als Iterationsvariable verwenden.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

Sie sollten innerhalb der for-Schleife eine Variable namens `min_distance` erstellen und sie auf `INF` setzen.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

Sie sollten innerhalb der for-Schleife eine Variable namens `current` erstellen und sie auf `-1` setzen.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
