---
id: 68760671d950767b055bf2f9
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Jetzt müssen Sie jeden Knoten überprüfen, um denjenigen mit der kleinsten bekannten Entfernung zu finden, der noch nicht besucht wurde.

Fügen Sie dazu eine `for`-Schleife innerhalb der Hauptschleife hinzu.

Die Schleife soll jeden `node_no` in `range(n)` durchlaufen, wobei `n` die Anzahl der Knoten im Graphen ist. Fügen Sie `pass` vorerst als Schleifenkörper hinzu.

# --hints--

Sie sollten eine zweite `for`-Schleife innerhalb der `shortest_path`-Funktion haben.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
len(outer_loop.find_bodies()[0].find_for_loops()) >= 1
      `)
      );
    }
});
```

Ihre `for`-Schleife sollte `range(n)` mit `node_no` als Iterationsvariable durchlaufen.

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
inner_loop = outer_loop.find_bodies()[0].find_for_loops()[0]
assert inner_loop.find_for_vars().is_equivalent("node_no")
assert inner_loop.find_for_iter().is_equivalent("range(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
    
    for _ in range(n):  
        min_distance = INF  
        current = -1
--fcc-editable-region--
        
--fcc-editable-region--
```
