---
id: 6876089f7a500b7c782e5b83
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Sie müssen entscheiden, ob der aktuelle Knoten eine bessere Wahl ist als der, den Sie bereits gefunden haben (falls vorhanden). Dazu fügen Sie eine bedingte Anweisung innerhalb der Schleife hinzu.

Die Bedingung soll zwei Dinge tun:

- Prüfen, ob der Knoten noch nicht besucht wurde.
- Die bekannte Entfernung zum aktuellen `min_distance` vergleichen.

Fügen Sie innerhalb der inneren `for`-Schleife eine `if`-Anweisung hinzu, die überprüft, ob `node_no` noch nicht besucht wurde **und** ob `distances[node_no]` kleiner ist als `min_distance`.

Fügen Sie `pass` als Platzhalter innerhalb des bedingten Blocks hinzu.

# --hints--

Sie sollten eine `if`-Anweisung innerhalb Ihrer inneren `for`-Schleife hinzufügen.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

Ihre `if`-Anweisung sollte prüfen, ob der Knoten unbesucht ist (`visited[node_no]` ist falsch) und `distances[node_no]` kleiner ist als `min_distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
    
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
--fcc-editable-region--
            pass
--fcc-editable-region--
```
