---
id: 68777816e95d4936f129819f
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Bevor Sie versuchen, die Entfernung zu einem Nachbarn zu aktualisieren, müssen Sie überprüfen, ob der Nachbar sowohl erreichbar als auch unbesucht ist. Dann berechnen Sie, wie die Gesamtdistanz wäre, um diesen Nachbarn über den aktuellen Knoten zu erreichen.

Fügen Sie innerhalb der `for`-Schleife eine `if`-Anweisung hinzu, die überprüft:

- Dass die `distance` nicht gleich `INF` ist (was bedeutet, dass eine Kante zwischen den Knoten besteht)
- Dass der Nachbar `node_no` noch nicht besucht wurde

Erstellen Sie innerhalb des `if`-Blocks eine Variable namens `new_distance` und weisen Sie ihr die Summe von `distances[current]` (die kürzeste Entfernung zum aktuellen Knoten) und `distance` (die Entfernung vom aktuellen Knoten zum Nachbarn) zu.

# --hints--

Sie sollten eine `if`-Anweisung haben, die überprüft, ob `distance` nicht gleich `INF` ist und ob der Nachbarknoten noch nicht besucht wurde.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

Sie sollten innerhalb des `new_distance`-Blocks eine Variable namens `if` erstellen.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

Sie sollten `new_distance` die Summe von `distances[current]` und `distance` zuweisen.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
