---
id: 687b45e16f38ad5f4e4bd799
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Nachdem Sie nun bestimmt haben, dass ein Knoten angezeigt werden soll, müssen Sie seinen Pfad so formatieren, dass er klar ausgegeben werden kann. Dafür verwenden Sie einen Generatorausdruck.

Ein Generatorausdruck ähnelt einer Listenverständnis, aber anstatt eine Liste zu erstellen, erzeugt er jeden Wert einzeln. Er verwendet runde Klammern `()` anstelle von eckigen Klammern `[]`. Zum Beispiel:

```py
numbers = [1, 2, 3]
squared = (x**2 for x in numbers)  # Generator expression
```

Innerhalb der Schleife, nach der `if`-Anweisung, erstellen Sie eine Variable namens `string_path`. Weisen Sie ihr einen Generatorausdruck zu, der jede Knotennummer in `paths[node_no]` mit `str()` in einen String umwandelt.

Der Generatorausdruck soll jede Knotennummer `n` in `paths[node_no]` durchlaufen.

# --hints--

Sie sollten eine Variable namens `string_path` erstellen.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
for_loop.find_bodies()[0].has_variable("string_path")
      `)
      );
    }
});
```

Sie sollten einen Generatorausdruck `string_path` zuweisen.

```js
({
  test: () => {
      runPython(`
      import ast
      assert isinstance(_Node(_code).find_function("shortest_path").find_for_loops()[1].find_bodies()[0].find_variable("string_path").tree.value, ast.GeneratorExp)
      `)
 }
})
```

Ihr Generatorausdruck sollte `paths[node_no]` durchlaufen.

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert path_var.find_comp_iters()[0].is_equivalent("paths[node_no]")
      `);
    }
});
```

Ihr Generatorausdruck sollte `n` als Iterationsvariable verwenden, um `paths[node_no]` zu durchlaufen.

```js
({
  test: () => {
      runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
string_path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert string_path_var.find_comp_targets()[0].is_equivalent("n")
      `);
    }
});
```

Ihr Generatorausdruck sollte `str(n)` für jedes `n` in `paths[node_no]` auswerten.

```js
({
  test: () => {
      runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
string_path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert string_path_var.find_comp_expr().is_equivalent("str(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no

        if current == -1:
            break
        visited[current] = True

        for node_no in range(n):
            distance = matrix[current][node_no]
            if distance != INF and not visited[node_no]:
                new_distance = distances[current] + distance
                if new_distance < distances[node_no]:
                    distances[node_no] = new_distance
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
    targets = [target_node] if target_node is not None else range(n)
    for node_no in targets:
        if node_no == start_node or distances[node_no] == INF:
            continue

--fcc-editable-region--
```
