---
id: 68e4c0ff9edcc92771d7549a
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Im Moment enthält die `distances`-Liste nur ein Element. Sie benötigen jedoch einen Distanzwert für jeden Knoten im Graphen.

In Python können Sie eine Liste mit einer ganzen Zahl multiplizieren, um sie zu wiederholen. Zum Beispiel erzeugt `[0] * 3` `[0, 0, 0]`.

Aktualisieren Sie die `distances`-Liste so, dass sie `n` Kopien von `INF` enthält, indem Sie `[INF]` mit `n` multiplizieren.

# --hints--

Sie sollten `distances` aktualisieren, damit sie `n` Kopien von `INF` enthält.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
