---
id: 68e4c81577243e4db1676491
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop, eserciterai l'uso dell'HTML semantico costruendo una pagina web che include alcuni consigli di Quincy Larson per ottenere un lavoro da sviluppatore. Il codice boilerplate HTML di base è già pronto per te.

Inizia creando un elemento `h1` con il testo `Quincy's Tips for Getting a Developer Job`.

# --hints--

Dovresti avere un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Il tuo elemento `h1` dovrebbe contenere il testo `Quincy's Tips for Getting a Developer Job`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
