---
id: 68e65ec5348e40a739903e6d
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Ora, crea un elemento paragrafo sotto l'elemento `h1`. All'interno di questo paragrafo aggiungi il testo `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

# --hints--

Dovresti avere un elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Il tuo paragrafo dovrebbe contenere il testo `Learning to code is hard, but as Quincy Larson says, You can become a developer.`. Controlla bene l'ortografia.

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

Il tuo paragrafo dovrebbe essere sotto l'elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
