---
id: 68e964b84f954c322376e592
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Se la fonte di una citazione è un sito web, puoi citarla con l'attributo `cite`. Il valore di questo attributo dovrebbe essere un URL valido. Anche se questo attributo non modifica la presentazione della citazione a blocchi, è molto utile per fornire ai lettori di schermo e ai motori di ricerca più informazioni sulla citazione.

Ecco un esempio di un elemento di citazione in linea con un attributo `cite`:

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

Aggiungi l'attributo `cite` all'elemento di citazione in linea con questo URL:

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

Il tuo elemento `q` dovrebbe avere un attributo `cite`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

L'attributo `cite` del tuo elemento `q` dovrebbe avere il valore `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
