---
id: 68e90335bd53a901c41dfc13
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

In una lezione precedente hai imparato che l'elemento `q` serve a distinguere il testo citato dal contenuto circostante.

Ecco un esempio:

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

La maggior parte dei browser moderni aggiunge automaticamente le virgolette intorno a una citazione in linea quando usi l'elemento `q`.

All'interno del paragrafo, avvolgi `You can become a developer.` in un elemento di citazione in linea, mantenendo il resto del paragrafo invariato.

# --hints--

Dovresti avere un elemento `q`.

```js
assert.exists(document.querySelector('q'));
```

Il tuo elemento `q` dovrebbe essere annidato all'interno del paragrafo che si trova sotto l'elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

Il tuo elemento `q` dovrebbe contenere il testo `You can become a developer.`

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

All'interno del tuo elemento `q`, il testo non deve avere spazi bianchi iniziali o finali (niente spazi, tabulazioni o interruzioni di riga prima del primo carattere o dopo l'ultimo).

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

Dovresti aggiungere i tag `q` solo intorno alla parte appropriata della frase senza modificare il testo originale. L'elemento `p` dovrebbe ancora leggere: `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
