---
id: 68ea7e7b916e8270483689b6
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Sotto l'elemento `h2`, aggiungerai un'altra citazione di Quincy. Questa volta, la citazione non farà parte di un paragrafo più ampio. Invece, l'intero paragrafo sarà una citazione. Per distinguere un testo citato di questo tipo, dovresti usare l'elemento di citazione a blocchi: `blockquote`. Nel browser, vedrai che il testo è leggermente rientrato.

Ecco un esempio di elemento di citazione a blocchi con testo citato:

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Ora, all'interno della prima sezione, aggiungi un elemento di citazione a blocchi sotto l'elemento `h2` con il testo `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

# --hints--

All'interno del primo elemento `section` dovresti avere un elemento `blockquote` sotto l'elemento `h2`.

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

Il tuo elemento `blockquote` dovrebbe contenere il testo `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
