---
id: 68ea9d23318b6481ac3de777
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Proprio come per l'elemento di citazione in linea, puoi anche aggiungere un attributo `cite` a un elemento di citazione a blocco.

Ecco un esempio di elemento di citazione a blocco con un attributo `cite`:

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Ora, aggiungi un attributo `cite` all'elemento di citazione a blocco con l'URL `https://www.freecodecamp.org/news/learn-to-code-book/`.

# --hints--

Il tuo elemento `blockquote` dovrebbe avere un attributo `cite`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.exists(firstBlockquoteEl?.getAttribute('cite'));
```

L'attributo `cite` dovrebbe avere il valore `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
      --fcc-editable-region--
        <blockquote>
      --fcc-editable-region--
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
