---
id: 68eaa9ef8a0867071f195f59
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Finora hai usato l'attributo `cite` per attribuire la fonte della citazione, ma l'attributo non mostra realmente la fonte all'utente.

Se vuoi attribuire visivamente la fonte, puoi aggiungere un elemento di citazione, `cite`, esterno all'elemento di citazione a blocchi. L'elemento di citazione è un elemento HTML che puoi usare per contrassegnare il titolo di un'opera creativa di riferimento, come un libro, un articolo, una canzone, un film, un sito web o un articolo di ricerca.

Ecco un esempio:

```html
<div>
  <blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
    The first thing you should consider about education is this is an economic decision.
  </blockquote>
  <p>&mdash;Quincy Larson, <cite>Is College Still Worth it? Tips from my 20 Years in Adult Education</cite></p>
</div>
```

All'interno dell'elemento `p` sotto l'elemento di citazione a blocchi, racchiudi `How to Learn to Code and Get a Developer Job [Full Book]` in un elemento `cite`.

# --hints--

Dovresti avere un elemento `cite` all'interno dell'elemento `p` sotto l'elemento `blockquote`.

```js
assert.exists(document.querySelector('blockquote + p > cite'));
```

Il tuo elemento `cite` dovrebbe contenere il testo `How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const citeEl = document.querySelector('blockquote + p > cite');
assert.equal(citeEl?.innerText.trim(), 'How to Learn to Code and Get a Developer Job [Full Book]');
```

Ignorando l'elemento `cite` che hai appena aggiunto, il testo dell'elemento `p` dovrebbe essere ancora `—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const pElCleanedText = document.querySelector('blockquote + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
      --fcc-editable-region--
          &mdash;Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]
      --fcc-editable-region--
        </p>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
