---
id: 68ea5ccc1229b85eca38a3e8
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

All'interno del primo elemento `section`, aggiungi un elemento `h2` con il testo `Envisioning Success`.

# --hints--

All'interno del primo elemento `section`, dovresti avere un elemento `h2`.

```js
assert.exists(document.querySelector('section:first-of-type > h2'));
```

Il tuo `h2` dovrebbe avere il testo `Envisioning Success`.

```js
const firstH2El = document.querySelector('section:first-of-type > h2');
assert.equal(firstH2El?.innerText.trim(), `Envisioning Success`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
