---
id: 587d7fb0367417b2b2512bed
title: 认识 Node 的控制台
challengeType: 2
forumTopicId: 301515
dashedName: meet-the-node-console
---

# --description--

Working on these challenges will involve you writing your code using one of the following methods:

- 克隆 <a href="https://github.com/freeCodeCamp/boilerplate-express/" target="_blank" rel="noopener noreferrer nofollow">GitHub 仓库</a> 并完成您的本地项目。
- Use a site builder of your choice to complete the project. Be sure to incorporate all the files from our GitHub repo.

During the development process, it is important to be able to check what’s going on in your code.

Node is just a JavaScript environment. Like client side JavaScript, you can use the console to display useful debug information. On your local machine, you would see console output in a terminal.

We recommend to keep the terminal open while working at these challenges. By reading the output in the terminal, you can see any errors that may occur.

更改文件后必须重新启动服务器。

你可以使用 `Ctrl + C` 从终端停止服务器，然后直接使用 Node（`node mainEntryFile.js`）或使用 `package.json` 文件中的运行脚本 `npm run` 启动服务器。

例如，可以使用 `npm run start` 在终端运行 `"start": "node server.js"` 脚本。

要实现保存文件时服务器自动重启，你可以将 Node 的 `--watch` 标志添加到你的启动脚本 `"start": "node --watch server.js"`，或者你可以安装一个 npm 包，例如 `nodemon`。 我们将把这作为一项练习留给你。

# --instructions--

Modify the `myApp.js` file to log "Hello World" to the console.

# --hints--

`"Hello World"` should be in the console

```js
  const response = await fetch(code + '/_api/hello-console');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.passed, '"Hello World" is not in the server console');
```

