---
id: 587d7fb1367417b2b2512bf1
title: 在指定路由上提供 JSON 服务
challengeType: 2
forumTopicId: 301517
dashedName: serve-json-on-a-specific-route
---

# --description--

HTML 服务器提供 HTML 服务，而 API 提供数据服务。 <dfn>REST</dfn>（REpresentational State Transfer）API 允许以简单的方式进行数据交换，对于客户端不必要知道服务器的细节。 客户只需要知道资源在哪里（URL），以及想执行的动作（动词）。 GET 动词常被用来获取无需修改的信息。 如今，网络上的移动数据首选格式是 JSON， 简而言之，JSON 是一种可以方便地用字符串表示 JavaScript 对象的方式，因此它很容易传输。

让我们通过创建一个在路径 `/json` 上响应 JSON 的路由来创建一个简单的 API。你可以像平常一样使用 `app.get()` 方法。在路由处理器内，使用 `res.json()` 方法，并传入一个对象作为参数。这个方法会结束请求-响应循环，返回数据。在幕后，它会将一个有效的 JavaScript 对象转换成字串，然后设置适当的头部，告诉你的浏览器你正在提供 JSON，并将数据发送回去。一个有效的对象通常结构为 `{key: data}`。`data` 可以是数字、字串、嵌套对象或数组。`data` 也可以是一个变量或函数调用的结果，在这种情况下，它会先被求值，然后再转换成字串。

# --instructions--

当向路由 `/json` 发送 GET 请求，将对象 `{"message": "Hello json"}` 以 JSON 格式返回给客户端， 浏览器访问 `your-app-url/json` 时，应该在屏幕上看到这个消息。

# --hints--

端口 `/json` 应该返回一个 JSON 对象 `{"message": "Hello json"}`

```js
  const response = await fetch(code + '/json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.message,
    'Hello json',
    "The '/json' endpoint does not serve the right data"
  );
```

