---
id: 66c06d618d075c7f7f1b890a
title: 創建一個占卜機
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

在本實驗室中，你將定義五個運勢，並隨機選擇其中一個向用戶顯示運勢。

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶故事：**

1. 你應初始化五個變量 `fortune1`、`fortune2`、`fortune3`、`fortune4` 和 `fortune5`，它們的值是你自定義的一段字符串。 你也可以使用以下選項：

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. 你應在 1 到 5（包括 5）之間選擇一個隨機數，並將其賦值給變量 `randomNumber` 。

3. 你應創建一個 `selectedFortune` 變量，並根據以下規則分配適當的運勢：

   - 如果 `randomNumber` 是 1，就把 `fortune1` 的值賦給 `selectedFortune`。
   - 如果 `randomNumber` 是 2，就把 `fortune2` 的值賦給 `selectedFortune`。
   - 如果 `randomNumber` 是 3，就把 `fortune3` 的值賦給 `selectedFortune`。
   - 如果 `randomNumber` 是 4，就把 `fortune4` 的值賦給 `selectedFortune`。
   - 如果 `randomNumber` 是 5，就把 `fortune5` 的值賦給 `selectedFortune`。

4. 你應將 `selectedFortune` 記錄到控制檯。

# --hints--

你應該用一個字符串值初始化 `fortune1` 。

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

你應該用一個字符串值初始化 `fortune2` 。

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

你應該用一個字符串值初始化 `fortune3` 。

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

你應該用一個字符串值初始化 `fortune4` 。

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

你應該用一個字符串值初始化 `fortune5` 。

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

你應使用 `Math.random()` 方法生成隨機數。

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

你應該生成一個介於 1 和 5 之間的隨機數，並將其賦值給變量 `randomNumber` 。

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

你應該有一個 `selectedFortune` 變量，並根據 `randomNumber` 的值賦值。

```js
assert.isNotNull(selectedFortune);
```

`randomNumber` 應與其運勢相對應。 例如，如果 `randomNumber` 爲 1，`selectedFortune` 應等於 `fortune1`，以此類推。

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

你應將 `selectedFortune` 輸出到控制檯。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
