---
id: 66c06d618d075c7f7f1b890a
title: 建造一個算命師
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

在這個實驗中，你將定義五個 fortune，並隨機選擇其中一個來顯示給使用者。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該用你選擇的字串值初始化五個變數 `fortune1`、`fortune2`、`fortune3`、`fortune4` 和 `fortune5`。如果你願意，可以使用以下選項：

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. 你應該從 1 到 5（含）之間選擇一個隨機數，並將它指定給變數 `randomNumber`。

3. 你應該創建一個 `selectedFortune` 變數，並根據這些規則指定適當的運勢：

   - 如果 `randomNumber` 是 1，將 `fortune1` 的值指定給 `selectedFortune`。
   - 如果 `randomNumber` 是 2，將 `fortune2` 的值指定給 `selectedFortune`。
   - 如果 `randomNumber` 是 3，將 `fortune3` 的值指定給 `selectedFortune`。
   - 如果 `randomNumber` 是 4，將 `fortune4` 的值指定給 `selectedFortune`。
   - 如果 `randomNumber` 是 5，將 `fortune5` 的值指定給 `selectedFortune`。

4. 你應該將 `selectedFortune` 記錄到主控台。

# --hints--

你應該用字串值初始化 `fortune1`。

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

你應該用字串值初始化 `fortune2`。

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

你應該用字串值初始化 `fortune3`。

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

你應該用字串值初始化 `fortune4`。

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

你應該用字串值初始化 `fortune5`。

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

你應該使用 `Math.random()` 方法來產生隨機數。

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

你應該產生一個介於 1 和 5（含）之間的隨機數，並將它指定給變數 `randomNumber`。

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

你應該有一個 `selectedFortune` 變數，其值是根據 `randomNumber` 的值指定的。

```js
assert.isNotNull(selectedFortune);
```

`randomNumber` 應該對應到它的運勢。例如，如果 `randomNumber` 是 1，則 `selectedFortune` 應該與 `fortune1` 相等，依此類推。

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

你應該將 `selectedFortune` 輸出到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
