---
id: 684ccbdc1ef89e042b744e09
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Наступна частина практичного заняття — додати кнопку закриття.

Під вашим елементом `p` додайте елемент `button` з множиною `aria-label`, встановленою в `"Dismiss"`. Всередині цього елемента `button` додайте елемент `span` з атрибутом `aria-hidden`. Текст елемента `span` має бути `&times;`.

*ПРИМІТКА*: Оскільки це практичне заняття буде зосереджене лише на роботі зі стилями Tailwind CSS, до `button` не буде додано жодної інтерактивності.

# --hints--

Ви повинні мати елемент `button`.

```js
assert.exists(document.querySelector("button"));
```
 
Ваш елемент `button` має мати атрибут `aria-label`.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

Ваш елемент `button` має мати атрибут `aria-label`, встановлений у `"Dismiss"`.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

Ваш елемент `button` має містити вкладений елемент `span`.

```js
assert.exists(document.querySelector("button span"));
```

Ваш елемент `span` має мати атрибут `aria-hidden`.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

Ваш елемент `span` має містити текст `&times;`.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```
