---
id: 684cd8bf5445de089af9c837
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Тепер настав час застосувати колір кордону. Щоб встановити властивість `border-color` за допомогою Tailwind CSS, ви можете використати утилітарний клас із таким шаблоном: `border-color-number`, де `number` представляє відтінок кольору кордону.

Ось приклад застосування фіолетового кольору кордону до елемента:

```html
<p class="border-purple-400">This is border-color.</p>
```

У вашому відкриваючому тегу `div` застосуйте червоний колір кордону з відтінком `300`.

# --hints--

Ваш відкриваючий тег `div` має містити утилітарний клас `border-red-300`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
