---
id: 684cd027744c2a0520ba4d58
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS використовує утилітарні класи, які застосовуються безпосередньо до розмітки. Ці класи включають стилі для шрифтів, кольорів, макетів на основі гнучкого контейнера, сітки та інше.

Щоб додати колір фону в Tailwind CSS, ви можете використати префікс `bg`, за яким іде назва кольору та число. Число представляє відтінок кольору.

Ось приклад додавання темно-синього кольору фону до елемента абзацу.

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

У вашому відкриваючому тегу `div` додайте клас з утилітарним класом `bg-red-100`, який представляє світло-червоний колір фону.

# --hints--

Ваш відкриваючий тег `div` повинен мати атрибут класу.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

Ваш атрибут класу повинен містити утилітарний клас `bg-red-100`.

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
