---
id: 684cca09aac0fc035db4099f
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Всередині вашого елемента `p` додайте текст `Error! Something went wrong. Please try again.`, де `Error!` має бути обгорнутий всередині елемента `strong`.

# --hints--

У вашому елементі `strong` має бути елемент `p`.

```js
assert.exists(document.querySelector("p strong"));
```

Ваш елемент `strong` має містити текст `Error!`.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

Ваш елемент `p` має містити правильний текст, який закінчується на `Something went wrong. Please try again.`.

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
