---
id: 684cc59eed2706fa50e8933c
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви навчитеся починати працювати з основами Tailwind CSS, створюючи компонент повідомлення про помилку.

На перших кількох кроках ви зосередитеся на створенні ініціальної HTML розмітки.

Всередині елемента `body` додайте елемент `div`. Потім всередині цього елемента `div` додайте елемент `p`.

# --hints--

У вас має бути елемент `div`.

```js
assert.exists(document.querySelector("div"));
```

У вас має бути елемент `p`.

```js
assert.exists(document.querySelector("p"));
```

Ваш елемент `p` має бути всередині елемента `div`.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
