---
id: 684cd59ec94b5e063bf776f0
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Було б добре додати кордон навколо того світло-червоного фону.

Щоб встановити `border-width` у Tailwind CSS, ви можете використати утилітарний клас `border-number`, де `number` представлятиме світлість або товщину вашого кордону.

Ось приклад застосування `border-width` товщиною `4px` з усіх боків:

```html
<p class="border-4">This is how to set the border-width.</p>
```

Всередині вашого відкриваючого тегу `div` додайте правильний утилітарний клас для множини `border-width`, встановленої на `2px` з усіх боків.

# --hints--

Ваш відкриваючий тег `div` має містити клас `border-2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
