---
id: 684cff0bf0a4ca0ad7e42990
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Наступним кроком буде заокруглення кутів кордону. Щоб встановити властивість `border-radius` за допомогою Tailwind CSS, ви можете використати утилітарний клас із таким шаблоном: `rounded-size`, де `size` представляє радіус.

Ось приклад застосування дуже великого радіуса кордону до елемента:

```html
<p class="rounded-xl">This is border radius.</p>
```

У вашому відкриваючому тегу `div` застосуйте радіус кордону розміром `md`.

# --hints--

Ваш відкриваючий тег `div` має містити утилітарний клас `rounded-md`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
