---
id: 684d09cad4ec900e41db2e74
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Зараз елемент `div` розташований у верхньому лівому куті екрану попереднього перегляду. Але було б добре додати трохи простору зверху елемента. Щоб застосувати `margin` до елемента за допомогою Tailwind CSS, ви можете використати такий шаблон: `m-number`, де `number` представляє фіксовану кількість пікселів, які застосовуються до всіх чотирьох сторін.

Якщо ви хочете додати поле лише з одного боку, можна використати будь-який із наступних шаблонів: `mt-number` (поле зверху), `mb-number` (поле знизу), `mr-number` (поле справа) та `ml-number` (поле зліва).

Ось приклад застосування `margin-bottom` до елемента:

```html
<p class="mb-8">This is margin.</p>
```

Всередині відкриваючого тегу `div` використайте правильний клас утиліти, щоб застосувати `margin-top` `4px`.

# --hints--

Ваш відкриваючий тег `div` повинен містити утилітний клас `mt-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
