---
id: 684d160a2370b01310ae0fda
title: Крок 18
challengeType: 0
dashedName: step-18
---

# --description--

Новий червоний колір виглядає гарно, але було б краще, якби текст був більшим. Щоб встановити `font-size` для елемента за допомогою Tailwind CSS, ви можете використати такий шаблон: `text-value`.

Ось приклад встановлення `font-size` для елемента на `0.875rem` або (`14px`):

```html
<p class="text-sm">This is for font-size.</p>
```

Всередині вашого відкриваючого `p` елемента використайте правильний утилітарний клас, щоб встановити `font-size` на `xl`, що представляє `1.25rem` або (`20px`).

# --hints--

Ваш відкриваючий `p` тег має містити утилітарний клас `text-xl`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
