---
id: 684d11ffd15a531225d081e9
title: Крок 17
challengeType: 0
dashedName: step-17
---

# --description--

Тепер, коли стилі застосовані до елемента `div`, настав час застосувати деякі стилі до елемента `p`.

Наразі колір тексту для елемента абзацу чорний. Але було б добре, якби колір був темно-червоним.

У попередньому кроці ви навчились працювати з кольорами фону так: `bg-red-100`. Ви можете зробити щось подібне з кольором тексту, дотримуючись цього шаблону: `text-color-shade`.

Ось приклад, як застосувати синій колір тексту до елемента:

```html
<p class="text-blue-500">This has blue text.</p>
```

Почніть із додавання атрибута класу до вашого елемента `p`. Потім застосуйте правильний утилітний клас для встановлення червоного кольору тексту з відтінком `700`.

# --hints--

Ваш елемент `p` повинен мати атрибут класу.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

Ваш відкриваючий тег `p` повинен мати утилітний клас `text-red-700`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
