---
id: 684d0eeb9c81e5109e52fe8b
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

Елемент `div` починає виглядати краще. Але тепер кнопка dismiss `button` розташована трохи занадто близько до абзацу.

Щоб створити простір між гнучкими елементами, ви можете використати такий шаблон: `gap-number`. `number` представляє фіксований пробіл між рядками та стовпцями у гнучкому макеті.

Ось приклад застосування `gap` до елемента:

```html
<p class="flex gap-8">This is gap.</p>
```

Всередині вашого відкриваючого тегу `div` використайте правильний утилітний клас для встановлення `gap` на `4`.

# --hints--

Ваш відкриваючий тег `div` повинен містити утилітний клас `gap-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
