---
id: 684d009994a3840b98c9aa3a
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Зараз ширина елемента `div` займає всю ширину сторінки. Щоб застосувати ширину в Tailwind CSS, у вас є кілька варіантів, наприклад, використання шаблону `w-number` для фіксованих ширин або шаблону `w-fraction` для відсотків.

Ось приклад застосування ширини, встановленої на `66%` від батьківського контейнера:

```html
<p class="w-2/3">This is a percentage width.</p>
```

У вашому відкриваючому тегу `div` застосуйте ширину, встановлену на `50%` від батьківського контейнера.

# --hints--

Ваш відкриваючий тег `div` повинен містити утиліту класу `w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
