---
id: 684d17a3b350bb13d3c19c0f
title: Крок 19
challengeType: 0
dashedName: step-19
---

# --description--

Тепер настав час застосувати стилі до елемента `button`.

Почніть із додавання атрибута class до вашого елемента `button`. Потім застосуйте правильний утилітарний клас для встановлення червоного кольору тексту з відтінком `700`.

Далі, всередині відкриваючого тегу `button`, використайте правильний утилітарний клас, щоб встановити `font-size` у `2xl`, що представляє `1.5rem` або (`24px`).

# --hints--

Ваш елемент `button` повинен мати атрибут class.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

Ваш відкриваючий тег `button` повинен мати утилітарний клас `text-red-700`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

Ваш відкриваючий тег `button` повинен мати утилітарний клас `text-2xl`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
