---
id: 684d07cbfa43a20d6a002ed7
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Стилі для елемента `div` починають складатися. Але було б добре, якби навколо тексту був пробіл.

Щоб додати `padding` до елемента за допомогою Tailwind CSS, ви можете використати такий шаблон: `p-number`, де кожна одиниця `number` відповідає 0.25rem (або 4px), які ви хочете застосувати з усіх чотирьох сторін.

Ось приклад додавання `padding` у розмірі `32px` до елемента:

```html
<p class="p-8">This is padding.</p>
```

У вашому відкриваючому тегу `div` використайте правильний клас утиліти, щоб додати `padding` у розмірі `16px`.

# --hints--

Ваш відкриваючий тег `div` має містити клас утиліти `p-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
