---
id: 684d105e488dbd1166a717ae
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Якщо збільшити розмір вікна попереднього перегляду, ви помітите, що всі елементи все ще вирівняні по лівому краю в контейнері. Було б добре, якби гнучкі елементи були по центру.

Щоб розташувати гнучкі елементи вздовж головної осі контейнера за допомогою Tailwind CSS, ви можете використати такий шаблон: `justify-value`. `value` може бути `start`, `end`, `center`, `between` і більше. Будь-яке дійсне значення, яке використовується у ванільному CSS з `justify-content`, можна застосувати тут з Tailwind CSS.

Всередині вашого відкриваючого тегу `div` використайте правильний клас утиліти для встановлення властивості `justify-content` у значення `center`.

# --hints--

Ваш відкриваючий тег `div` має містити утиліту класу `justify-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
