---
id: 684d03df2a0ddf0c735865f0
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Тепер ви бачите, що ширина елемента становить половину розміру `body`. Це добре працює для більших пристроїв, але не так добре для менших.

Tailwind CSS використовує підхід mobile-first у дизайні. Тож якщо ви хочете створити адаптивний дизайн для інших пристроїв, ви можете використовувати адаптивні варіанти Tailwind, такі як `md` і `lg`, ось так:

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

У наведеному прикладі ширина буде встановлена на `66%` батьківського елемента для розмірів пристроїв `1024px` і більших.

У вашому відкриваючому тегу `div` змініть `w-1/2`, щоб він застосовувався лише для розмірів пристроїв `768px` і більших. Ви можете додати префікс `md` до цього класу.

# --hints--

Ваш відкриваючий тег `div` має містити утиліту класу `md:w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

У вашому відкриваючому тегу `w-1/2` більше не має бути класу `div`.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
