---
id: 684d1942a1904a14a3709c6c
title: Крок 20
challengeType: 0
dashedName: step-20
---

# --description--

Стилізація вашого компонента повідомлення про помилку майже завершена, але залишилось виправити одну річ. Оскільки ви встановили різні розміри шрифтів для абзацу та тексту кнопки, вертикальне вирівнювання між цими гнучкими елементами порушене.

Щоб розташувати гнучкі елементи вздовж поперечної осі контейнера за допомогою Tailwind CSS, ви можете використати такий шаблон: `items-value`. `value` може бути будь-яким із `start`, `end`, `center`, `baseline` та іншими.

Всередині вашого відкриваючого тегу `div` використайте правильний клас утиліти для встановлення властивості `align-items` у значення `center`.

І з цією останньою зміною ваш компонент повідомлення про помилку буде завершено!

# --hints--

Ваш відкриваючий тег `div` повинен містити клас утиліти `items-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
