---
id: 68501d32b69f9d7756334ffc
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Якщо збільшити вікно попереднього перегляду, ви повинні побачити елемент `div`, розташований по центру горизонтально на сторінці. Але для менших пристроїв ширина елемента `div` охоплює всю ширину сторінки.

Оновіть клас `mx-auto`, щоб він застосовувався лише для пристроїв `md` і більших.

# --hints--

Ваш відкриваючий тег `div` повинен містити утилітарний клас `md:mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

У вашому відкриваючому тегу `mx-auto` більше не має бути класу `div`.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
