---
id: 61fd6cc9475a784b7776233e
title: 步驟7
challengeType: 0
dashedName: step-7
---

# --description--

在你現有的 `div` 元素下方，添加一個新的 `div` 元素，並將其 `class` 屬性設置爲 `table-wrap`。 這將是你的表格的容器。

在其中添加三個 `table` 元素。 你將使用 CSS 將這些樣式設置爲單個表格，但使用單獨的表格將有助於屏幕閱讀器理解文檔流。

# --hints--

你應創建一個新的 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

你的新 `div` 元素應該有一個設置爲 `table-wrap` 的 `class` 。

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

你的 `.table-wrap` 元素應該在現有的 `div` 之後。

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

你的 `.table-wrap` 元素應該包含三個 `table` 元素。

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
