---
id: 61fd71d596e8f253b9408b39
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

`tr` 元素用於顯示表格的列。為你的 `thead` 元素添加一個 `tr` 元素。在你的新 `tr` 元素中，添加一個 `td` 元素，接著添加三個 `th` 元素。

`td` 元素顯示資料儲存格，而 `th` 元素顯示標題儲存格。

# --hints--

你的 `thead` 元素應該有一個 `tr` 元素。

```js
assert.strictEqual(
  document.querySelector('thead')?.children?.[0]?.localName,
  'tr'
);
```

你的 `tr` 元素應該有一個 `td` 元素作為第一個子元素。

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[0]?.localName,
  'td'
);
```

你的 `tr` 元素應該在 `td` 元素之後有三個 `th` 元素。

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[1]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[2]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[3]?.localName,
  'th'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
