---
id: 61fd719788899952e67692b9
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

`thead` 和 `tbody` 元素用來顯示你的表格中哪一部分是標題，以及哪一部分包含主要的資料或內容。

在你的第一個 `table` 中，於 `caption` 元素下方添加 `thead` 和 `tbody`。

# --hints--

你的第一個 `table` 元素應該有一個 `thead` 元素。

```js
const table = document.querySelectorAll('table')?.[0];
assert.isNotNull(table?.querySelector('thead'));
```

你的第一個 `table` 元素應該有一個 `tbody` 元素。

```js
const table = document.querySelectorAll('table')?.[0];
const tbody = table?.querySelector('tbody');
assert.isNotNull(tbody);
```

你的 `thead` 元素應該緊接在你的 `caption` 元素下方。

```js
assert.strictEqual(
  document.querySelector('caption')?.nextElementSibling?.localName,
  'thead'
);
```

你的 `tbody` 元素應該緊接在你的 `thead` 元素下方。

```js
assert.strictEqual(
  document.querySelector('thead')?.nextElementSibling?.localName,
  'tbody'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
