---
id: 61fd778081276b59d59abad6
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

在你的第一個 `tr` 中，添加一個帶有文字 `Cash This is the cash we currently have on hand.` 的 `th` 元素。將除了 `Cash ` 以外的所有文字包裹在一個 `class` 設為 `description` 的 `span` 元素中。

接著，添加三個 `td` 元素，依序包含以下文字：`$25`、`$30`、`$28`。為第三個 `td` 元素添加一個 `class` 屬性，並設為 `current`。

# --hints--

你的第一個 `tr` 應該有一個 `th` 元素。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.isNotNull(tableRow?.querySelector('th'));
```

你的 `th` 元素應該有文字 `Cash This is the cash we currently have on hand.`

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Cash This is the cash we currently have on hand.'
);
```

你應該將文字 `This is the cash we currently have on hand.` 包裹在 `span` 元素中。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'This is the cash we currently have on hand.'
);
```

你的 `span` 元素應該將 `class` 屬性設為 `description`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

你應該有三個 `td` 元素。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

你的第一個 `td` 元素應該有文字 `$25`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$25');
```

你的第二個 `td` 元素應該有文字 `$30`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$30');
```

你的第三個 `td` 元素應該有文字 `$28`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$28');
```

你的第三個 `td` 元素的 `class` 應該設定為 `current`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[0];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data"></tr>
              <tr class="data"></tr>
              <tr class="total"></tr>
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
