---
id: 61fd67a656743144844941cb
title: 步驟4
challengeType: 0
dashedName: step-4
---

# --description--

屏幕閱讀器根據文檔流閱讀 HTML 元素。 我們最終希望資產負債表有一個 “Balance Sheet” 的標題和一個 “AcmeWidgetCorp” 的子標題。 但是，如果屏幕閱讀器按此順序閱讀則沒有意義。

將現有的 `span` 元素的 `class` 屬性設置爲 `flex`，並在其內部添加兩個 `span` 元素。 給第一個添加文本 `AcmeWidgetCorp`。 給第二個添加文本 `Balance Sheet`。 你將使用 CSS 來顛倒頁面上文字的順序，但 HTML 的順序對於屏幕閱讀器來說更合理。

# --hints--

你現有的 `span` 元素應將 `class` 屬性設置爲 `flex`。

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

你現有的 `span` 元素中應有兩個新的 `span` 元素。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

你的新 `span` 元素不應有 `class` 屬性。

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

你的第一個新 `span` 元素應包含文本 `AcmeWidgetCorp`。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent,
  'AcmeWidgetCorp'
);
```

你的第二個新 `span` 元素應包含文本 `Balance Sheet`。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent,
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
