---
id: 61fd6b7c83dbf54a08cf0498
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在你的 `div` 元素中，為你的三個 `span` 元素添加 `class` 屬性並設為 `year`，並依序添加以下文字：`2019`、`2020` 和 `2021`。

# --hints--

你的 `div` 元素應該有三個 `span` 元素。

```js
assert.lengthOf(document.querySelector('div')?.children, 3);
```

每個 `span` 元素都應該有一個 `class` 屬性設為 `year`。

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert.isTrue(span?.classList?.contains('year')));
```

你的第一個 `span` 應該有文字 `2019`。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[0]?.textContent.trim(),
  '2019'
);
```

你的第二個 `span` 應該有文字 `2020`。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[1]?.textContent.trim(),
  '2020'
);
```

你的第三個 `span` 應該有文字 `2021`。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[2]?.textContent.trim(),
  '2021'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
