---
id: 61fd6ab779390f49148773bb
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在你的 `h1` 元素下方，創建一個 `div` 元素。為它的 `id` 屬性設定為 `years`。你想讓這個特定元素對螢幕閱讀器隱藏，所以為它的 `aria-hidden` 屬性設定為 `true`。

# --hints--

你應該在你的 `h1` 元素之後創建一個新的 `div` 元素。

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling?.localName,
  'div'
);
```

你的新 `div` 元素應該有一個 `id` 屬性設為 `years`。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('id'), 'years');
```

你的新 `div` 元素應該將 `aria-hidden` 屬性設為 `true`。

```js
assert.strictEqual(
  document.querySelector('div')?.getAttribute('aria-hidden'),
  'true'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
