---
id: 61fd70336ebb3e4f62ee81ba
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

HTML 表格使用 `caption` 元素來描述表格的內容。`caption` 元素應該永遠是 `table` 的第一個子元素，但可以使用 `caption-side` CSS 屬性來調整其位置。

為你的第一個 `table` 添加一個 `caption` 元素，並給它文字 `Assets`。

# --hints--

你的第一個 `table` 元素應該有一個 `caption` 元素。

```js
assert.strictEqual(
  document.querySelector('table')?.children?.[0]?.localName,
  'caption'
);
```

你的 `caption` 元素應該有文字 `Assets`。

```js
assert.strictEqual(document.querySelector('caption')?.textContent.trim(), 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
