---
id: 61fd67a656743144844941cb
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

螢幕閱讀器會根據文件流程宣告 HTML 元素。我們最終會希望資產負債表有一個標題為「Balance Sheet」和一個副標題為「AcmeWidgetCorp」。然而，如果由螢幕閱讀器宣告，這個順序並不合理。

為你現有的 `span` 設定 `class` 屬性為 `flex`，並在其中添加兩個 `span` 元素。為第一個設定文字為 `AcmeWidgetCorp`。為第二個設定文字為 `Balance Sheet`。你將使用 CSS 反轉頁面上文字的順序，但 HTML 順序對螢幕閱讀器來說會更有意義。

# --hints--

你現有的 `span` 元素應該將 `class` 屬性設為 `flex`。

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

你現有的 `span` 元素中應該包含兩個新的 `span` 元素。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

你的新 `span` 元素不應該有 `class` 屬性。

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

你的第一個新的 `span` 元素應該有文字 `AcmeWidgetCorp`。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

你的第二個新的 `span` 元素應該有 `Balance Sheet` 的文字。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
