---
id: 61fd75ea7f663457612dba02
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

在你的每個新的 `th` 元素中，巢狀一個 `span` 元素，並將 `class` 設為 `sr-only year`。依序為它們提供以下文字：`2019`、`2020` 和 `2021`。

為你的第三個 `th` 元素添加 `class` 屬性並設為 `current`。

將 `td` 元素留空。此元素僅用於確保你的表格具有四行佈局，並將標題與正確的行關聯。

# --hints--

你的每個 `th` 元素都應該有一個 `span` 元素。

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.lengthOf(th?.children, 1);
  assert.strictEqual(th?.children?.[0]?.localName, 'span');
});
```

你的每個新的 `span` 元素都應該將 `class` 屬性設為 `sr-only year`。

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.isTrue(th?.children?.[0]?.classList?.contains('sr-only'));
  assert.isTrue(th?.children?.[0]?.classList?.contains('year'));
});
```

你的第一個 `span` 元素應該有文字 `2019`。

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent.trim(),
  '2019'
);
```

你的第二個 `span` 元素應該有文字 `2020`。

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent.trim(),
  '2020'
);
```

你的第三個 `span` 元素應該有文字 `2021`。

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent.trim(),
  '2021'
);
```

你的第三個 `th` 元素的 `class` 應設為 `current`。

```js
assert.isTrue(
  document
    .querySelector('table')
    ?.querySelectorAll('th')?.[2]
    ?.classList?.contains('current')
);
```

你的 `td` 元素應該是空的。

```js
const table = document.querySelector('table');
const tableData = table.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '');
assert.isEmpty(tableData?.children);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
