---
id: 61fd6cc9475a784b7776233e
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

在你現有的 `div` 元素下方，為你的表格添加一個 `class` 設為 `table-wrap` 的新 `div` 元素。這將是你的表格的容器。

在其中，為你的文件添加三個 `table` 元素。你將使用 CSS 將這些樣式化為單一的表格，但使用獨立的表格將有助於螢幕閱讀器理解文件流程。

# --hints--

你應該創建一個新的 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

你的新 `div` 元素應該將 `class` 設定為 `table-wrap`。

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

你的 `.table-wrap` 元素應該放在你現有的 `div` 之後。

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

你的 `.table-wrap` 元素應該有三個 `table` 元素。

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
