---
id: 61fd5a93fd62bb35968adeab
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

通過在 `head` 元素中爲樣式表添加 `link` 元素，開始你的項目。

# --hints--

你的代碼中應該有一個 `link` 元素。

```js
assert.match(code, /<link/);
```

你的 `link` 元素應該在你的 `head` 元素中。

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

你的 `link` 元素應該有一個 `rel` 屬性，值爲 `stylesheet`。

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

你的 `link` 元素應該有一個 `href` 屬性，值爲 `styles.css`。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
