---
id: a3bfc1673c0526e06d3ac698
title: 素数の合計を計算する電卓を作ろう
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

このラボでは、数字を受け取り、その数字以下のすべての素数の合計を返す電卓を作成します。

**目的**：以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. 数字を引数に取る`sumPrimes`関数が必要です。
1. `sumPrimes`関数は、指定された数字以下のすべての素数の合計を返す必要があります。
1. 入力された数字が`2`より小さい場合、関数は`0`を返す必要があります。

# --hints--

`sumPrimes`関数が必要です。

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)`は`17`を返すべきです。

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)`は`10`を返すべきです。

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)`は`2`を返すべきです。

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)`は`0`を返すべきです。

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)`は`73156`を返すべきです。

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
