---
id: 694cabb9f1afa9446f36ffc0
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Agora você pode começar a construir sua equipe. Use `addCrewMember()` para adicionar o `firstAstronaut` à sua `squad`.

# --hints--

Você deve fazer a chamada de `addCrewMember()` com `squad` e `firstAstronaut`.

```js
assert.match(__helpers.removeJSComments(code), /addCrewMember\s*\(\s*squad\s*,\s*firstAstronaut\s*\)/);
assert.strictEqual(squad.length, 1);
assert.strictEqual(squad[0], firstAstronaut);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
  console.log(`Added ${astronaut.name} as ${astronaut.role}`);
}

--fcc-editable-region--

--fcc-editable-region--
```
