---
id: 694c84330cbe604821b9deef
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

O controle da missão requer um método para adicionar novos membros à lista. Crie uma função vazia chamada `addCrewMember` que aceite dois parâmetros, `crew` e `astronaut`.

# --hints--

Você deve criar uma function chamada `addCrewMember`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addCrewMember);
```

Sua `addCrewMember` função deve ter os parâmetros `crew` e `astronaut`.

```js
const explorer = await __helpers.Explorer(code);
const { addCrewMember } = explorer.allFunctions;
assert.lengthOf(addCrewMember?.parameters, 2);
assert.equal(addCrewMember?.parameters?.[0]?.toString(), "crew");
assert.equal(addCrewMember?.parameters?.[1]?.toString(), "astronaut");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
