---
id: 694ce8609fb805ba21cebc02
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

O controle da missão precisa de um método para trocar as posições de dois membros da tripulação em uma lista. Crie uma função vazia chamada `swapCrewMembers` com três parâmetros:

- `crew`: um array de objetos astronauta

- `fromIndex`: o índice do primeiro astronauta a ser trocado

- `toIndex`: o índice do segundo astronauta a ser trocado

# --hints--

Você deve criar uma função chamada `swapCrewMembers`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.swapCrewMembers);
```

Sua `swapCrewMembers` função deve ter os parâmetros `crew`, `fromIndex` e `toIndex`.

```js
const explorer = await __helpers.Explorer(code);
const { swapCrewMembers } = explorer.allFunctions;
assert.lengthOf(swapCrewMembers?.parameters, 3);
assert.equal(swapCrewMembers?.parameters?.[0]?.toString(), "crew");
assert.equal(swapCrewMembers?.parameters?.[1]?.toString(), "fromIndex");
assert.equal(swapCrewMembers?.parameters?.[2]?.toString(), "toIndex");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

--fcc-editable-region--

--fcc-editable-region--
```
