---
id: 694c9c186c089f7bdff0f3c6
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Geralmente, é uma boa prática validar sua(s) entrada(s) dentro de uma função. Use um `for` loop para iterar através do array `crew` e verificar se algum membro já possui o mesmo `id` que o `astronaut` de entrada. Se um duplicado for encontrado, registre a mensagem de erro `console.log("Duplicate ID: " + astronaut.id)` e então faça a chamada `return` para sair da função antecipadamente. 

Aqui está um exemplo desta técnica comumente usada:

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

Sua função deve usar um `for` loop para escanear os membros existentes do `crew` (qualquer forma válida de `for (...)` é aceitável).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

Dentro do `for` loop, você deve verificar se há um `id` duplicado. Se um for encontrado, registre `"Duplicate ID: " + astronaut.id"` no console e então faça a chamada `return` para sair da função.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Quando fornecidas entradas válidas, sua função `addCrewMember` não deve invocar uma mensagem de erro `console.log()` nem um `return` antecipado.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
