---
id: 694c84330cbe604821b9deef
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

O controle da missão requer um método para adicionar novos membros à lista. Crie uma função vazia chamada `addCrewMember` que aceite dois parâmetros, `crew` e `astronaut`.

# --hints--

Você deve criar uma `function` chamada `addCrewMember`.

```js
assert.isFunction(addCrewMember);
```

Sua `addCrewMember` função deve ter os parâmetros `crew` e `astronaut`.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
