---
id: 694c40d86fe6a78d20b6559c
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, você vai aproveitar `JavaScript` para construir uma lista de astronautas para uma missão espacial. 

Para começar, crie um array vazio para sua lista chamado `squad`.

# --hints--

Você deve criar uma variável `squad`.

```js
assert.isDefined(squad);
```

Sua variável `squad` deve ser um array.

```js
assert.isArray(squad);
```

Seu `squad` array deve estar vazio.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
