---
id: 694c84106e9023f289c03b84
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Sua equipe começa com um único astronauta comandante. Crie um objeto chamado `firstAstronaut` com as seguintes propriedades:

| Key | Value |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

Você deve criar uma variável `firstAstronaut`.

```js
assert.isDefined(firstAstronaut);
```

Seu `firstAstronaut` deve ser um objeto.

```js
assert.isObject(firstAstronaut);
```

Seu `firstAstronaut` objeto deve incluir os valores corretos das propriedades.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
