---
id: 694d8ddd106340237efb0c80
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

O método `splice()` pode modificar arrays adicionando ou removendo elementos em qualquer posição, inclusive no meio. Como `splice()` retorna um array contendo os elementos removidos, você pode usá-lo para trocar dois elementos em um array sem modificar o original. Isso pode ser feito em uma linha usando a seguinte técnica:

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

Aqui está um exemplo:

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

A técnica, aplicada na terceira `line`, funciona da seguinte forma:

- `splice(3, 1, copyArray[1])` remove o elemento no índice `3` (`55`)

- Insere o elemento de `copyArray[1]` (`97`) no índice `3`

- `splice()` retorna `[55]`, um array contendo o elemento removido

- `[0]` extrai `55` desse array

- Esse valor (`55`) é atribuído de volta a `copyArray[1]`, completando a troca

Use esta técnica com o `updatedCrew` array para trocar os astronautas em `fromIndex` e `toIndex`.

# --hints--

Você deve realizar a troca fazendo uma chamada ao `splice()` method.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

Você deve fazer uma chamada a `splice()` com o array `updatedCrew` para realizar a troca. Você deve usar a técnica de uma linha.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

Sua `swapCrewMembers` função não deve modificar o `crew` array de entrada original.

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
