---
id: 5f3313e74582ad9d063e3a38
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте елемент `head` в межах елемента `html`, щоб можна було додати елемент `title`. Текстом елемента `title` має бути `Cafe Menu`.

# --hints--

Ви повинні мати початковий тег `<head>`.

```js
assert.match(code, /<head>/i);
```

Ви повинні мати кінцевий тег `</head>`.

```js
assert.match(code, /<head>/i);
```

Ви повинні мати початковий тег `<title>`.

```js
assert.match(code, /<title>/i);
```

Ви повинні мати кінцевий тег `</title>`.

```js
assert.match(code, /<\/title>/i);
```

Елемент `title` повинен бути вкладеним в елементі `head`.

```js
assert.match(code, /<head>\s*<title>.*<\/title>\s*<\/head>/si);
```

Елемент `title` повинен містити текст `Cafe Menu`. Можливо, доведеться перевірити написання.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
