---
id: 5f33071498eb2472b87ddee4
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цієї майстерні ви вивчите основи CSS (каскадних таблиць стилів), створюючи меню кафе.

Start by adding your `<!DOCTYPE html>` and `html` element with a `lang` attribute set to `en`.

# --hints--

Ви повинні мати `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

У вас має бути початковий `<html lang="en">` з англійською мовою.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Ви повинні мати кінцевий тег `</html>`.

```js
assert.match(code, /<\/html>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
