---
id: 5f3313e74582ad9d063e3a38
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Add a `head` element within the `html` element, so you can add a `title` element. Текстом елемента `title` має бути `Cafe Menu`.

# --hints--

Ви повинні мати початковий тег `<head>`.

```js
assert(code.match(/<head>/i));
```

Ви повинні мати кінцевий тег `</head>`.

```js
assert(code.match(/<head>/i));
```

Ви повинні мати початковий тег `<title>`.

```js
assert(code.match(/<title>/i));
```

Ви повинні мати кінцевий тег `</title>`.

```js
assert(code.match(/<\/title>/i));
```

Your `title` element should be nested in your `head` element.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Елемент `title` повинен містити текст `Cafe Menu`. You may need to check your spelling.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
