---
id: 66c8ba975ee7230e29f6c4ad
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Um zu prüfen, für welchen Kredit ein Benutzer basierend auf dem `annualIncome` und `creditScore` qualifiziert ist, müssen Sie eine `if/else if`-Anweisung oder einen ternären Ausdruck direkt in der `getLoanMessage`-Funktion verwenden. Anschließend geben Sie die passende Nachricht im Block jeder Bedingung zurück.

Beginnen Sie mit dem Duplex-Kredit und prüfen Sie, ob `annualIncome` größer oder gleich `minIncomeForDuplex` UND ob `creditScore` größer oder gleich `minCreditScoreForDuplex` ist.

Wenn diese Bedingung wahr ist, dann ist der Antragsteller für einen Duplex-Kredit und die anderen Kredite berechtigt. Geben Sie also innerhalb der Prüfung den String `"You qualify for a duplex, condo, and car loan."` zurück.

# --hints--

Ihre `getLoanMessage`-Funktion sollte einen String zurückgeben.

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

Ihre `getLoanMessage`-Funktion sollte den String `"You qualify for a duplex, condo, and car loan."` zurückgeben.

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

Ihre `getLoanMessage`-Funktion sollte `undefined` zurückgeben, wenn das jährliche Einkommen und die Kreditwürdigkeit des Antragstellers die Anforderungen für einen Duplex-Kredit nicht erfüllen.

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
