---
id: 66c8ba975ee7230e29f6c4ac
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Wenn der Benutzer für einen Kredit berechtigt ist, möchten Sie ihm eine Nachricht in der Konsole anzeigen.

Dafür erstellen Sie eine Funktion, in der Sie einige Prüfungen durchführen, die zurückgeben, für welchen Kredit der Antragsteller berechtigt ist.

Erstellen Sie eine leere `getLoanMessage`-Funktion mit `annualIncome` und `creditScore` als Parameter.

# --hints--

Sie sollten eine `getLoanMessage`-Funktion erstellen

```js
assert.isFunction(getLoanMessage)
```

Ihre `getLoanMessage`-Funktion sollte `annualIncome` und `creditScore` als Parameter haben.

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` sollte eine leere Funktion sein.

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
