---
id: 66c8ba975ee7230e29f6c4af
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Wenn das Jahreseinkommen des Antragstellers größer oder gleich `minIncomeForCondo` ist UND seine Kreditwürdigkeit größer oder gleich `minCreditScoreForCondo` ist, dann qualifiziert er sich für einen Wohnungs- und Autokredit.

Überprüfen Sie, ob das in der `getLoanMessage`-Funktion zutrifft. Wenn ja, geben Sie den String `"You qualify for a condo and car loan."` zurück.

# --hints--

Ihre `getLoanMessage`-Funktion sollte einen String zurückgeben, wenn der Antragsteller sich für einen Wohnungs- und Autokredit qualifiziert.

```js
assert.isString(getLoanMessage(45000, 680));
```

Ihre `getLoanMessage`-Funktion sollte den String `"You qualify for a condo and car loan."` zurückgeben.

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

Ihre `getLoanMessage`-Funktion sollte `undefined` zurückgeben, wenn das Jahreseinkommen und die Kreditwürdigkeit des Antragstellers die Anforderungen für einen Wohnungsdarlehen nicht erfüllen.

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
