---
id: 66c8ba975ee7230e29f6c4b0
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Jetzt sollten Sie prüfen, ob der Antragsteller nur für einen Autokredit qualifiziert ist. Wenn ja, geben Sie den String `"You qualify for a car loan."` zurück.

# --hints--

Ihre `getLoanMessage`-Funktion sollte einen String zurückgeben, wenn der Antragsteller für einen Autokredit qualifiziert ist.

```js
assert.isString(getLoanMessage(30000, 650));
```

Ihre `getLoanMessage`-Funktion sollte den String `"You qualify for a car loan."` zurückgeben.

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

Ihre `getLoanMessage`-Funktion sollte `undefined` zurückgeben, wenn das Jahreseinkommen und die Kreditwürdigkeit des Antragstellers nicht die Anforderungen für einen Autokredit erfüllen.

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 
  
--fcc-editable-region--
}
```
