---
id: 66c8ba41a77db20a93f9d7a1
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop überprüfen Sie JavaScript-Bedingungen, indem Sie eine App zur Überprüfung der Kreditwürdigkeit erstellen.

Die App prüft, ob der Benutzer für einen Duplex-, Auto- oder Eigentumswohnungsdarlehen basierend auf seinem Jahreseinkommen und seiner Kreditwürdigkeit berechtigt ist. Eine Kreditwürdigkeit ist eine Zahl, die darstellt, wie gut Sie mit geliehenem Geld umgehen.

Um zu beginnen, erstellen Sie die folgenden Variablen und Werte.

| Variable Name | Wert   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

Sie sollten eine `minIncomeForDuplex`-Variable erstellen, die auf `60000` gesetzt ist.

```js
assert.equal(minIncomeForDuplex, 60000);
```

Sie sollten eine `minCreditScoreForDuplex` erstellen, die auf `700` gesetzt ist.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

Sie sollten eine `minIncomeForCar`-Variable erstellen, die auf `30000` gesetzt ist.

```js
assert.equal(minIncomeForCar, 30000);
```

Sie sollten eine `minCreditScoreForCar`-Variable erstellen, die auf `650` gesetzt ist.

```js
assert.equal(minCreditScoreForCar, 650);
```

Sie sollten eine `minIncomeForCondo`-Variable erstellen, die auf `45000` gesetzt ist.

```js
assert.equal(minIncomeForCondo, 45000);
```

Sie sollten eine `minCreditScoreForCondo`-Variable erstellen, die auf `680` gesetzt ist.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
