---
id: 66c8ba975ee7230e29f6c4b2
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Jetzt ist es Zeit, Ihre `getLoanMessage`-Funktion zu testen.

Verwenden Sie die folgende Tabelle, um `4`-Variablen und deren Werte zu erstellen:

| Variablenname | Wert    |
| ----------- | ------- |
| `duplexLoanMsg`       | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`   | `getLoanMessage(65000, 690)`   |
| `carLoanMsg`   | `getLoanMessage(45000, 660)`   |
| `noLoanMsg`   | `getLoanMessage(25000, 550)`   |

Anschließend geben Sie jede Variable in der Konsole aus, um die Nachrichten zu sehen.

Damit ist Ihr Projekt zur Überprüfung der Kreditwürdigkeit abgeschlossen!

# --hints--

Sie sollten eine Variable namens `duplexLoanMsg` haben.

```js
assert.isNotNull(duplexLoanMsg);
```

Ihre `duplexLoanMsg`-Variable sollte auf das Ergebnis von `getLoanMessage(85000, 850)` gesetzt werden.

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

Sie sollten die `duplexLoanMsg`-Variable in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

Sie sollten eine Variable namens `condoLoanMsg` haben.

```js
assert.isNotNull(condoLoanMsg);
```

Ihre `condoLoanMsg`-Variable sollte auf das Ergebnis von `getLoanMessage(65000, 690)` gesetzt werden.

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

Sie sollten die `condoLoanMsg`-Variable in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

Sie sollten eine Variable namens `carLoanMsg` haben.

```js
assert.isNotNull(carLoanMsg);
```

Ihre `carLoanMsg`-Variable sollte auf das Ergebnis von `getLoanMessage(45000, 660)` gesetzt werden.

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

Sie sollten die `carLoanMsg`-Variable in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

Sie sollten eine Variable namens `noLoanMsg` haben.

```js
assert.isNotNull(noLoanMsg);
```

Ihre `noLoanMsg`-Variable sollte auf das Ergebnis von `getLoanMessage(25000, 550)` gesetzt werden.

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

Sie sollten die `noLoanMsg`-Variable in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
