---
id: 66c8ba975ee7230e29f6c4b1
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Wenn das Jahreseinkommen und die Kreditwürdigkeit des Antragstellers unter `minIncomeForCar` und `minCreditScoreForCar` liegen, qualifiziert er sich für keinen Kredit. Geben Sie daher den String `"You don't qualify for any loans."` zurück.

# --hints--

Ihre `getLoanMessage`-Funktion sollte einen String zurückgeben, wenn sich der Antragsteller für keinen Kredit qualifiziert.

```js
assert.isString(getLoanMessage(25000, 550));
```

Ihre `getLoanMessage`-Funktion sollte den String `"You don't qualify for any loans."` zurückgeben, wenn der Antragsteller die Anforderungen für keinen Kredit erfüllt.

```js
assert.strictEqual(getLoanMessage(15000, 600), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(25000, 550), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(20000, 500), "You don't qualify for any loans.");
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
