---
id: 6808fe91a5cb0c3355f23027
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 슈퍼히어로 애플리케이션 폼을 만들어 React에서 폼을 다루는 방법을 배우세요.

일부 보일러플레이트가 제공되었으니 나머지 React 컴포넌트만 작성하면 됩니다.

`div` 요소를 `className`가 `form-wrap`인 상태로 반환하는 것부터 시작하세요. `div` 안에 `h2` 요소를 만들고 텍스트는 `Superhero Application Form`로 하세요. 그 아래에 `p` 요소를 만들고 텍스트는 `Please complete all fields`로 하세요.

# --hints--

`div` 요소를 `className`가 `form-wrap`인 상태로 반환해야 합니다.

```js
assert.exists(document?.querySelector("div.form-wrap"));
```

`h2` 요소 안에 `div` 요소를 만들어야 합니다.

```js
assert.exists(document?.querySelector(".form-wrap > h2"));
```

`h2` 요소는 텍스트가 `Superhero Application Form`여야 합니다.

```js
const h2El = document?.querySelector("h2")
assert.equal(h2El?.textContent, "Superhero Application Form");
```

`p` 요소 안에 `div` 요소를 만들어야 합니다. `h2` 다음에 와야 하는지 확인하세요.

```js
assert.exists(document?.querySelector(".form-wrap > h2 + p"));
```

`p` 요소는 텍스트가 `Please complete all fields`여야 합니다.

```js
const pEl = document?.querySelector("p")
assert.equal(pEl?.textContent, "Please complete all fields");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  --fcc-editable-region--
  
  --fcc-editable-region--

};
```
