---
id: 680900675ae3d54ee19590c3
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

폼 자체를 만들기 전에 필요한 상태 변수를 생성하세요. 이 변수들은 `heroName`, `realName`, `powerSource`, 그리고 `powers`입니다.

지금은 `heroName`와 `realName`에 대한 변수와 설정 함수를 만드세요. 두 변수 모두 초기값은 빈 문자열이어야 합니다.

# --hints--

배열 구조 분해 구문을 사용해 `heroName` 상태 변수와 `setHeroName` 설정 함수를 만드세요.

```js
assert.match(code, /(const|let|var)\s+\[\s*heroName\s*,\s*setHeroName\s*\]/);
```

`heroName`와 `setHeroName`는 `useState` 훅을 사용해야 합니다.

```js
  const abuseState = __helpers.spyOn(React, "useState");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.SuperheroForm);

  assert.isAtLeast(abuseState.calls.length, 1);
```

`useState` 훅을 사용해 `heroName`를 만들고 초기값은 빈 문자열이어야 합니다.

```js
  const abuseState = __helpers.spyOn(React, "useState");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.SuperheroForm);
  
  assert.strictEqual(abuseState.calls[0]?.[0], "");
```

배열 구조 분해 구문을 사용해 `realName` 상태 변수와 `setRealName` 설정 함수를 만드세요.

```js
assert.match(code, /(const|let|var)\s+\[\s*realName\s*,\s*setRealName\s*\]/);
```

`realName`와 `setRealName`는 `useState` 훅을 사용해야 합니다.

```js
  const abuseState = __helpers.spyOn(React, "useState");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.SuperheroForm);

  assert.isAtLeast(abuseState.calls.length, 2);
```

`useState` 훅을 사용해 `realName`를 만들고 초기값은 빈 문자열이어야 합니다.

```js
  const abuseState = __helpers.spyOn(React, "useState");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.SuperheroForm);

  assert.strictEqual(abuseState.calls[1]?.[0], "");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  --fcc-editable-region--
  
  --fcc-editable-region--

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
    </div>
  )
};
```
