---
id: 680900675ae3d54ee19590ca
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

setter 함수에 맞추려면 `value`에 `onChange` 속성과 `select` 이벤트를 붙이고, `value` 요소 각각에 `option` 속성을 붙여야 합니다.

과정이 길죠? 그래서 가장 좋은 방법은 옵션 값들의 배열을 만들어서 그 배열을 map으로 돌며 `option` 요소들을 생성하는 것입니다.

컴포넌트 상단에 `powerSourceOptions` 배열이 추가되어 있으니 확인하세요.

우선 `select` 요소 안의 옵션들을 첫 번째 것만 남기고 모두 제거하세요. 다음 단계에서 배열을 map으로 돌릴 겁니다.

# --hints--

코드에 첫 번째 옵션 요소가 있어야 합니다.

```js
const optionEl = document?.querySelectorAll("select > option");
assert.equal(optionEl[0]?.textContent, "Select one");
```

첫 번째 옵션을 제외한 모든 옵션 요소를 제거해야 합니다.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.lengthOf(optionEls, 1)
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const powerSourceOptions = [
    'Bitten by a strange creature',
    'Radioactive exposure',
    'Science experiment',
    'Alien heritage',
    'Ancient artifact discovery',
    'Other'
  ];

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          <label>
            Hero Name
            <input
              type='text'
              value={heroName}
              onChange={e => setHeroName(e.target.value)}
            />
          </label>
          <label>
            Real Name
            <input
              type='password'
              value={realName}
              onChange={e => setRealName(e.target.value)}
            />
          </label>
        </div>
        <label className='section column'>
          How did you get your powers?
          <select>
          --fcc-editable-region--
            <option value=''>
              Select one
            </option>
            <option>
              Bitten by a strange creature
            </option>
            <option>
              Radioactive exposure
            </option>
            <option>
              Science experiment
            </option>
            <option>
              Alien heritage
            </option>
            <option>
              Ancient artifact discovery
            </option>
            <option>
              Other
            </option>
            --fcc-editable-region--
          </select>
        </label>
      </form>
    </div>
  )
};
```
