---
id: 680900675ae3d54ee19590c5
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

레이블과 입력을 만들 차례입니다. 폼 안에 `label`를 만들고 텍스트로 `Hero Name`를 넣은 다음 그 안에 `input` 타입 `text` 요소를 만드세요.

`input`용 `Hero Name`를 `heroName` 상태 변수에 연결하려면 `value` 속성을 `heroName`로 설정하세요.

# --hints--

`label` 요소를 만들어야 합니다.

```js
assert.exists(document?.querySelector("label"));
```

`label` 요소의 텍스트는 `Hero Name`이어야 합니다.

```js
const labelEl = document.querySelector("label")
assert.equal(labelEl?.textContent, "Hero Name");
```

`input` 요소 안에 `text` 타입 `label`를 만들어야 합니다.

```js
const inputEl = document.querySelector("label > input");
assert.equal(inputEl?.getAttribute("type"), "text");
```

`input` 요소의 값은 `heroName` 상태 변수로 설정해야 합니다.

```js
assert.match(code, /value\s*=\s*\{\s*heroName\s*\}/)
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </div>
      </form>
    </div>
  )
};
```
