---
id: 7a97b79c096ecb000d03bd36
title: Build an Nth Fibonacci Number Calculator
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Obiettivo:** Soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**User Stories:**

1. Dovresti creare una funzione chiamata `fibonacci`.
2. Dovresti definire un array chiamato `sequence` all'interno della funzione `fibonacci`, inizializzato con i valori `[0, 1]`.
3. La funzione `fibonacci` dovrebbe accettare un parametro, un numero intero positivo `n`.
4. Chiamare `fibonacci(n)` dovrebbe usare un approccio di programmazione dinamica per calcolare e restituire il numero `n`-esimo della sequenza di Fibonacci, dove ogni numero è la somma dei due numeri precedenti.
5. Ogni numero calcolato nella posizione `n` della sequenza di Fibonacci dovrebbe essere archiviato nell'array `sequence` all'indice `n - 1`.

# --hints--

Dovresti avere una funzione chiamata `fibonacci`.

```js
assert.isFunction(fibonacci);
```

La tua funzione `fibonacci` dovrebbe prendere un solo parametro.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

Dovresti avere un array chiamato `sequence` all'interno della funzione `fibonacci` inizializzato a `[0, 1]`.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` should return `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` should return `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` should return `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` should return `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` should return `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` should return `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` should return `610`.

```js
assert.equal(fibonacci(15), 610);
```

Non dovresti usare la ricorsione nel tuo codice.

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
