---
id: 7a97b79c096ecb000d03bd36
title: Costruisci un calcolatore del numero di Fibonacci n-esimo
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una funzione chiamata `fibonacci`.
2. Dovresti definire un array chiamato `sequence` all'interno della funzione `fibonacci`, inizializzato con i valori `[0, 1]`.
3. La funzione `fibonacci` dovrebbe accettare un parametro, un numero intero positivo `n`.
4. Chiamare `fibonacci(n)` dovrebbe usare un approccio di programmazione dinamica per calcolare e restituire il numero `n`-esimo della sequenza di Fibonacci, dove ogni numero è la somma dei due numeri precedenti.
5. Ogni numero calcolato nella posizione `n` della sequenza di Fibonacci dovrebbe essere archiviato nell'array `sequence` all'indice `n - 1`.

# --hints--

Dovresti avere una funzione chiamata `fibonacci`.

```js
assert.isFunction(fibonacci);
```

La tua funzione `fibonacci` dovrebbe prendere un solo parametro.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

Dovresti avere un array chiamato `sequence` all'interno della funzione `fibonacci` inizializzato a `[0, 1]`.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` dovrebbe restituire `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` dovrebbe restituire `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` dovrebbe restituire `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` dovrebbe restituire `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` dovrebbe restituire `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` dovrebbe restituire `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` dovrebbe restituire `610`.

```js
assert.equal(fibonacci(15), 610);
```

Non dovresti usare la ricorsione nel tuo codice.

```js
const bodyMatch = __helpers.removeJSComments(code).match(/function\s+fibonacci\s*\([^)]*\)\s*\{([\s\S]*)\}/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
