---
id: 587d774c367417b2b2512a9d
title: 了解 Alt 文本留空的情景
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9P4t2'
forumTopicId: 301019
dashedName: know-when-alt-text-should-be-left-blank
---

# --description--

在上一个挑战中，我们了解到 `img` 标签必须有一个 `alt` 属性。 但是，有时图像通过它们的描述文本被归类，或者只用于装饰。 在这些情况下， `alt` 文本可能是多余的或没有必要的。

在图片已经有了文字说明，或者仅仅为了美化页面的情况下，`img` 仍然需要一个 `alt` 属性，但可以设置为空字符串。 这是一个示例：

```html
<img src="visualDecoration.jpeg" alt="">
```

比如，背景图片通常起装饰作用。 但这些图片通常都是通过 CSS 规则而非 HTML 引入的，因此屏幕阅读器毋需读取。

**注意：** 对于有标题的图片，依然需要添加 `alt` 文本，因为这样有助于搜索引擎记录图片内容。

# --instructions--

Camper Cat 已经为他网站的博客部分编写了一个骨架页面。他计划在两篇文章之间添加一张装饰性的武士刀图像作为视觉分隔。为 `img` 标签添加一个 `alt` 属性，并将其设置为空字串。（注意图像的 `src` 并未链接到实际文件——不用担心显示中没有出现刀剑。）

# --hints--

`img` 标签应具有 `alt` 属性。

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

`alt` 的属性值应为空。

```js
assert.isEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg" alt="">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```
