---
id: 67d83df6f82eda3868dd2a84
title: RPG 캐릭터 만들기
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

이 실습에서는 RPG 모험용 캐릭터를 만드는 작은 앱을 만들어 Python 기본기를 연습합니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `create_character`라는 이름의 함수를 만들어야 합니다.
1. 함수는 순서대로 캐릭터 이름과 세 가지 능력치: 힘, 지능, 매력을 인수로 받아야 합니다.
1. 캐릭터 이름을 검증해야 합니다:
   - 캐릭터 이름이 문자열이 아니면 함수는 `The character name should be a string`을 반환해야 합니다.
   - 캐릭터 이름이 빈 문자열이면 함수는 `The character should have a name`를 반환해야 합니다.
   - 캐릭터 이름이 10자보다 길면 함수는 `The character name is too long`을 반환해야 합니다.
   - 캐릭터 이름에 공백이 포함되면 함수는 `The character name should not contain spaces`를 반환해야 합니다.
1. 능력치도 검증해야 합니다:
   - 하나 이상의 능력치가 정수가 아니면 함수는 `All stats should be integers`를 반환해야 합니다.
   - 하나 이상의 능력치가 1보다 작으면 함수는 `All stats should be no less than 1`을 반환해야 합니다.
   - 하나 이상의 능력치가 4보다 크면 함수는 `All stats should be no more than 4`을 반환해야 합니다.
   - 모든 능력치 합이 7과 다르면 함수는 `The character should start with 7 points`을 반환해야 합니다.
1. 모든 값이 검증을 통과하면 함수는 네 줄짜리 문자열을 반환해야 합니다:
   - 첫 번째 줄은 캐릭터 이름을 포함해야 합니다.
   - 2~4번째 줄은 능력치 약어 `STR`, `INT`, `CHA` (이 순서대로)로 시작하고, 그 다음에 공백, 그리고 능력치 값만큼의 꽉 찬 점(`●`)과 10이 되도록 빈 점(`○`)이 이어져야 합니다. 예를 들어 힘이 3이면 꽉 찬 점 3개 다음에 빈 점 7개가 와야 합니다. 점들은 편집기에 제공되어 있습니다.

`create_character('ren', 4, 2, 1)`가 반환해야 하는 문자열 예시는 다음과 같습니다:

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

참고: `str`와 `int`는 능력치의 일반적인 약어지만, Python에서 예약 키워드이므로 변수 이름으로 사용하면 안 됩니다.

# --hints--

`create_character`라는 이름의 함수를 만들어야 합니다.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

`create_character`가 첫 번째 인수로 문자열이 아닌 값을 받으면 `The character name should be a string`를 반환해야 합니다.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

`create_character`가 첫 번째 인수로 문자열을 받으면 `The character name should be a string`를 반환하지 않아야 합니다.

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

`create_character`가 첫 번째 인수로 빈 문자열을 받으면 `The character should have a name`를 반환해야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

`create_character`가 첫 번째 인수로 빈 문자열이 아닌 값을 받으면 `The character should have a name`를 반환하지 않아야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

`create_character`가 첫 번째 인수로 10자보다 긴 문자열을 받으면 `The character name is too long`를 반환해야 합니다.

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

`create_character` 함수는 10자보다 길지 않은 문자열에 대해 너무 길다고 말하면 안 됩니다.

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

`create_character`가 첫 번째 인수로 공백이 포함된 문자열을 받으면 `The character name should not contain spaces`를 반환해야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

`create_character`가 첫 번째 인수로 공백이 없는 문자열을 받으면 `The character name should not contain spaces`를 반환하지 않아야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

`create_character`가 두 번째, 세 번째, 네 번째 인수 중 하나라도 정수가 아니면 `All stats should be integers`를 반환해야 합니다.

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

`create_character`가 두 번째, 세 번째, 네 번째 인수가 모두 정수이면 `All stats should be integers`를 반환하지 않아야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

`create_character`가 두 번째, 세 번째, 네 번째 인수 중 하나라도 `1`보다 작으면 `All stats should be no less than 1`를 반환해야 합니다.

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

`create_character`가 두 번째, 세 번째, 네 번째 인수가 모두 `1` 이상이면 `All stats should be no less than 1`를 반환하지 않아야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

`create_character`가 두 번째, 세 번째, 네 번째 인수 중 하나라도 `4`보다 크면 `All stats should be no more than 4`를 반환해야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

`create_character`가 두 번째, 세 번째, 네 번째 인수가 모두 `4` 이하이면 `All stats should be no more than 4`를 반환하지 않아야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

`create_character`가 두 번째, 세 번째, 네 번째 인수 합이 `7`가 아니면 `The character should start with 7 points`를 반환해야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

`create_character`가 두 번째, 세 번째, 네 번째 인수 합이 `7`이면 `The character should start with 7 points`를 반환하지 않아야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)`는 `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`를 반환해야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

`create_character`가 유효한 값으로 호출되면 요구한 대로 캐릭터 능력치를 출력해야 합니다.

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
