---
id: 66b102ee0bdbad34a9f85ba0
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop lernen Sie weiter über Strings, indem Sie einen Lehrer-Chatbot erstellen.

Beginnen Sie damit, eine `console`-Anweisung hinzuzufügen, mit der Nachricht `"Hi there!"`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Sie sollten eine `console`-Anweisung haben.

```js
assert.isAbove(logSpy.calls.length, 0);
```

Ihre Konsolenanweisung sollte die Nachricht `"Hi there!"` ausgeben.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
