---
id: 66b102ee0bdbad34a9f85ba0
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop lernen Sie weiter über Strings, indem Sie einen Lehrer-Chatbot erstellen.

Beginnen Sie damit, eine `console`-Anweisung hinzuzufügen, mit der Nachricht `"Hi there!"`.

# --hints--

Sie sollten eine `console`-Anweisung haben.

```js
assert.match(code, /console\.log(.*)/);
```

Ihre Konsolenanweisung sollte die Nachricht `"Hi there!"` ausgeben.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
