---
id: 66b59be6ab830800c4df9146
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Jetzt ist es Zeit, die Variablen zu verwenden, die Sie im vorherigen Schritt erstellt haben.

In den vorherigen Lektionen haben Sie gelernt, wie man mit Template-Literalen wie diesem arbeitet:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

Beginnen Sie damit, eine `sentence`-Variable zu erstellen.

Verwenden Sie die Syntax für Template-Literale, um den String `Today, you will learn about [topic variable goes here] in [subject variable goes here].` der `sentence`-Variable zuzuweisen.

Sie werden die `[topic variable goes here]`- und `[subject variable goes here]`-Platzhalter durch die `topic`- und `subject`-Variablen sowie die `${}`-Syntax ersetzen.

Zum Schluss geben Sie die `sentence`-Variable in der Konsole aus.

# --hints--

Sie sollten eine `sentence`-Variable haben.

```js
assert.isNotNull(sentence);
```

Ihre `sentence`-Variable sollte ein String sein.

```js
assert.isString(sentence);
```

Sie sollten die Syntax für Template-Literale verwenden, um den String `Today, you will learn about [topic variable goes here] in [subject variable goes here].` der `sentence`-Variable zuzuweisen. Schauen Sie sich das Beispiel an, wenn Sie Hilfe brauchen.

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

Sie sollten die `sentence`-Variable in der Konsole ausgeben.

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
