---
id: 66b5928b86e907fdfbf34e56
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Jetzt ist es Zeit, den Namen des Bots festzulegen.

Erstellen Sie eine Variable namens `botName` und weisen Sie ihr den String-Wert `"teacherBot"` zu.

# --hints--

Sie sollten eine Variable namens `botName` haben.

```js
assert.isNotNull(botName);
```

Ihre `botName`-Variable sollte ein String sein.

```js
assert.isString(botName);
```

Ihre `botName`-Variable sollte den Wert `"teacherBot"` haben.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
