---
id: 66b59829dba144ff1351220f
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Jetzt ist es Zeit, eine Begrüßung mit der `botName`-Variable zu erstellen.

In vorherigen Lektionen haben Sie gelernt, wie man Strings mit Template-Literalen wie folgt verknüpft:

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

Beginnen Sie damit, eine Variable namens `greeting` zu erstellen.

Weisen Sie anschließend mit der Syntax für Template-Literale einen String zu, der `My name is` sagt, gefolgt von der `botName`-Variable und endet mit einem Punkt (`.`).

Geben Sie zuletzt die `greeting`-Variable in der Konsole aus.

# --hints--

Sie sollten eine Variable namens `greeting` haben.

```js
assert.isNotNull(greeting);
```

Sie sollten Ihrer `greeting`-Variable einen String zuweisen.

```js
assert.isString(greeting);
```

Sie sollten Template-Literale verwenden, um den String `My name is` mit der `botName`-Variable gefolgt von einem Punkt (`.`) zu verknüpfen.

```js
assert.equal(greeting, "My name is teacherBot.");
```

Sie sollten eine `console`-Anweisung haben, die die `greeting`-Variable in der Konsole ausgibt.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
