---
id: 66b59b12a745e10011158f7b
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Der nächste Schritt besteht darin, einige weitere Variablen zu erstellen, die in zukünftigen Bot-Nachrichten verwendet werden.

Erstellen Sie eine Variable namens `subject` und weisen Sie ihr den String-Wert `"JavaScript"` zu.

Erstellen Sie dann eine Variable namens `topic` und weisen Sie ihr den String-Wert `"strings"` zu.

# --hints--

Sie sollten eine Variable namens `subject` haben.

```js
assert.isNotNull(subject);
```

Sie sollten Ihrer Variable `subject` einen String zuweisen.

```js
assert.isString(subject);
```

Sie sollten den String `"JavaScript"` Ihrer Variable `subject` zuweisen.

```js
assert.equal(subject, "JavaScript");
```

Sie sollten eine Variable namens `topic` haben.

```js
assert.isNotNull(topic);
```

Sie sollten Ihrer Variable `topic` einen String zuweisen.

```js
assert.isString(topic);
```

Sie sollten den String `"strings"` Ihrer Variable `topic` zuweisen.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
