---
id: 66b6d482bbb9e12f2e5ee1ae
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Für diesen nächsten Teil des Projekts wird der Bot den Umgang mit der String-Eigenschaft `length` lehren.

Beginnen Sie damit, eine neue Variable namens `strLengthIntro` zu erstellen.

Weisen Sie dann mit der Syntax für Template Literale den String `Here is an example of using the length property on the word [subject].` der Variable `strLengthIntro` zu.

Ersetzen Sie `[subject]` durch die Variable `subject`, so wie Sie es zuvor gemacht haben.

Geben Sie abschließend die Variable `strLengthIntro` in der Konsole aus.

# --hints--

Sie sollten eine Variable namens `strLengthIntro` haben.

```js
assert.isNotNull(strLengthIntro);
```

Ihre Variable `strLengthIntro` sollte ein String sein.

```js
assert.isString(strLengthIntro);
```

Sie sollten die Syntax für Template Literale verwenden, um den String `Here is an example of using the length property on the word [subject goes here].` der Variable `strLengthIntro` zuzuweisen.

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

Sie sollten die Variable `strLengthIntro` in der Konsole ausgeben.

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
