---
id: 66b6fdb76441c738719039fa
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie Sie auf das letzte Zeichen in einem String zugreifen, und zwar so:

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` gibt Ihnen immer die letzte Indexnummer für einen String.

Erstellen Sie eine neue Variable namens `lastCharacter` und weisen Sie ihr den Wert des letzten Zeichens in der Variable `subject` zu.

Geben Sie dann den Wert der Variable `lastCharacter` in der Konsole aus.

# --hints--

Sie sollten eine Variable namens `lastCharacter` haben.

```js
assert.isNotNull(lastCharacter);
```

Sie sollten der Variable `subject` den Wert des letzten Zeichens in der Variable `lastCharacter` zuweisen. Wenn Sie Hilfe brauchen, sehen Sie sich das Beispiel an.

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

Sie sollten den Wert der Variable `lastCharacter` in der `console` ausgeben.

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
