---
id: 66b6e80d6c3f0b329c360283
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Der nächste Teil dieses Workshops besteht darin, das Zugreifen auf Zeichen in einem String zu überprüfen.

Beginnen Sie damit, die Nachricht `Here is an example of accessing the first letter in the word [subject].` in der Konsole auszugeben.

Denken Sie daran, `[subject]` durch die `subject`-Variable zu ersetzen und die richtige Syntax für Template Literals zu verwenden, wie Sie es in den vorherigen Schritten getan haben.

# --hints--

Ihre `console`-Anweisung sollte die Nachricht `Here is an example of accessing the first letter in the word [subject].` ausgeben

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+first\s+letter\s+in\s+the\s+word\s+\${subject}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

--fcc-editable-region--

--fcc-editable-region--
```
