---
id: 66b6e62423e8a031d6c1f03d
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Jetzt ist es Zeit, die Länge des `topic` Strings zu ermitteln.

Sie können Template-Literale innerhalb von `console` Anweisungen wie folgt verwenden:

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

Beginnen Sie damit, die Nachricht `Here is an example of using the length property on the word [topic].` in der Konsole auszugeben.

Denken Sie daran, `[topic]` durch die `topic` Variable zu ersetzen und die richtige Template-Literal-Syntax zu verwenden, wie Sie es in den vorherigen Schritten getan haben.

Fügen Sie dann eine zweite `console.log` Anweisung hinzu, die die Länge des `topic` Strings in der Konsole ausgibt.

# --hints--

Sie sollten die Template-Literal-Syntax verwenden, um die Nachricht `Here is an example of using the length property on the word [topic].` an die `console` auszugeben.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

Sie sollten eine zweite `console` Anweisung haben, um die Länge der `topic` Variable an die `console` auszugeben.

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
