---
id: 66b6f586767a1534f3097353
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Jetzt ist es Zeit, auf den zweiten Buchstaben der `subject`-Variable zuzugreifen.

Beginnen Sie damit, eine `console`-Anweisung hinzuzufügen, die die Nachricht `Here is an example of accessing the second letter in the word [subject].` ausgibt. Denken Sie daran, `[subject]` durch den tatsächlichen Wert der `subject`-Variable zu ersetzen und die korrekte Syntax für Template-Literale zu verwenden.

Fügen Sie dann eine weitere `console`-Anweisung hinzu, die den zweiten Buchstaben der `subject`-Variable mit Klammernotation und der korrekten Indexnummer ausgibt.

# --hints--

Ihre `console`-Anweisung sollte die Nachricht `Here is an example of accessing the second letter in the word [subject].` ausgeben.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

Sie sollten eine weitere `console`-Anweisung haben, die den zweiten Buchstaben der `subject`-Variable mit Klammernotation und dem Index `1` ausgibt.

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
