---
id: 66b6ffb42b88e33943788abf
title: Schritt 14
challengeType: 1
dashedName: step-14
---

# --description--

Für den letzten Teil des Workshops überprüfen Sie, wie Sie die Indexposition eines Substrings in einem String finden. Denken Sie daran, dass ein Substring ein Teil eines Strings ist.

Beginnen Sie damit, eine Variable namens `learningIsFunSentence` zu erstellen und ihr den String-Wert `"Learning is fun."` zuzuweisen.

# --hints--

Sie sollten eine Variable namens `learningIsFunSentence` haben.

```js
assert.isNotNull(learningIsFunSentence);
```

Ihre `learningIsFunSentence` sollte den Wert eines Strings enthalten.

```js
assert.isString(learningIsFunSentence);
```

Sie sollten den Wert von `"Learning is fun."` der Variable `learningIsFunSentence` zuweisen.

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
