---
id: 66b6efddeca35833cd6f0b03
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie Sie auf Zeichen in einem String zugreifen, und zwar so:

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

Denken Sie daran, dass die Indexnummern bei `0` beginnen, sodass der erste Buchstabe in einem String immer am Index `0` steht.

Beginnen Sie damit, eine weitere `console`-Anweisung hinzuzufügen.

Geben Sie innerhalb der `console`-Anweisung den ersten Buchstaben der `subject`-Variablen mit der Klammernotation und der richtigen Indexnummer aus.

# --hints--

Ihre `console`-Anweisung sollte den ersten Buchstaben der `subject`-Variablen mit der Klammernotation und dem Index `0` ausgeben.

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
