---
id: 66b6e39031393e30f2c48d0f
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Um die Länge eines Strings zu ermitteln, können Sie die <dfn>length</dfn>-Eigenschaft verwenden. Diese Eigenschaft gibt die Anzahl der Zeichen in einem String zurück.

Hier ist ein Beispiel:

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

Beginnen Sie damit, die `length`-Eigenschaft zu verwenden, um die Länge des `subject`-Strings zu ermitteln und diesen Wert in der Konsole auszugeben.

# --hints--


Sie sollten die `length`-Eigenschaft verwenden, um die Länge des `subject`-Strings zu ermitteln.

```js
assert.match(code, /subject\.length/);
```

Sie sollten die Länge des `subject`-Strings in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
