---
id: 66b7137348cfb53fd3ec6c73
title: Schritt 17
challengeType: 1
dashedName: step-17
---

# --description--

Als NĂ¤chstes fĂĽgen Sie eine neue `console`-Anweisung hinzu, die das Ergebnis der Anwendung der `indexOf`-Methode auf die `learningIsFunSentence`-Variable ausgibt, um die Position des Teilstrings `"fun"` zu finden.

Unterhalb dieser `console`-Anweisung fĂĽgen Sie eine neue `console`-Anweisung hinzu, die das Ergebnis der Anwendung der `indexOf`-Methode ausgibt, um die Position des Teilstrings `"learning"` zu finden.

Beachten Sie, was die letzte `console`-Anweisung ausgibt.

# --hints--

Ihre `console` sollte die Indexposition des Teilstrings `"fun"` aus der `learningIsFunSentence`-Variable ausgeben.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(fun)\1\s*\)\s*\)/);
```

Ihre zweite `console`-Anweisung sollte die Indexposition des Teilstrings `"learning"` aus der `learningIsFunSentence`-Variable ausgeben.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

console.log(learningIsFunSentence.indexOf("Learning"));

--fcc-editable-region--

--fcc-editable-region--
```
