---
id: 680a431692a5f863f726c438
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

In questo workshop, costruirai un *cifrario di Cesare*. Questa è una delle tecniche più semplici per criptare un testo, che consiste nel sostituire ogni lettera del testo in chiaro con la lettera che si trova a un numero fisso di posizioni più avanti nell’alfabeto. Per esempio, con uno spostamento di 5, `a` verrebbe sostituito da `f`, `b` da `g` e così via.

Per implementare questo cifrario, dovrai creare una nuova versione del tuo alfabeto che inizi dalla posizione indicata dallo spostamento. Come hai imparato in una lezione precedente, puoi estrarre una parte di una stringa usando lo slicing delle stringhe:

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

Crea una variabile chiamata `shifted_alphabet` e usa la sintassi dello slicing per assegnarle la porzione di `alphabet` che inizia all’indice `shift`. Poi, chiama la funzione integrata `print()` per stampare `shifted_alphabet` sul terminale e guarda il risultato.

# --hints--

Dovresti dichiarare una nuova variabile chiamata `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

Dovresti assegnare `alphabet[shift:]` a `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

Dovresti chiamare la funzione `print()` passando `shifted_alphabet` come argomento.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
