---
id: 680a836235461fce1ff53b3f
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Il metodo `translate()` prende come argomento la tabella di traduzione generata da `str.maketrans()`. Viene chiamato su una stringa e restituisce una copia della stringa originale in cui i caratteri sono stati sostituiti in base alla tabella di traduzione:

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

Chiama il metodo `translate()` su `text` passando `translation_table` come argomento e assegna il risultato a una variabile chiamata `encrypted_text`.

# --hints--

Dovresti avere una variabile chiamata `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Dovresti assegnare `text.translate(translation_table)` a `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
