---
id: 680a549f60f5a3fedfc0edd9
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Come puoi vedere dall'output, l'alfabeto spostato inizia con la lettera `f` perché `shift` ha il valore `5`. Ma ora le prime cinque lettere dell'alfabeto – `a`, `b`, `c`, `d` e `e` – mancano nell'alfabeto spostato, quindi dovrai aggiungerle alla fine dell'alfabeto spostato.

L'operatore `+` viene usato per unire due o più stringhe insieme in un processo chiamato concatenazione, così:

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

Modifica l'assegnazione esistente della variabile `shifted_alphabet`: usa la sintassi di slicing per estrarre la porzione mancante all'inizio di `alphabet` e concatenala a `alphabet[shift:]`.

Per ricordarti, `sentence[start:stop]` restituisce i caratteri di `sentence` dalla posizione `start` (inclusa) fino a `stop` (esclusa).

# --hints--

Non dovresti modificare direttamente la variabile `alphabet`.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

Dovresti assegnare la concatenazione di `alphabet[shift:]` e della porzione mancante all'inizio di `alphabet` alla variabile `shifted_alphabet` sulla stessa riga in cui è stata dichiarata. Usa `shift` per specificare dove fermare lo slicing.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
