---
id: 680a3d4c065557260de2ff33
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop, creerai un programma che implementa un cifrario di Cesare, un metodo di crittografia che sposta le lettere nell'alfabeto per codificare i messaggi.
Inizia creando una variabile chiamata `shift` e assegnando il valore `5` alla tua nuova variabile.

# --hints--

Dovresti dichiarare una variabile chiamata `shift`. Assicurati di posizionare il nome della variabile all'inizio della riga.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

Dovresti assegnare il valore `5` alla tua variabile `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
