---
id: 680a41ef5bfdc059cfd9d4b5
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Dichiara un'altra variabile chiamata `alphabet` e assegna a questa variabile la stringa `abcdefghijklmnopqrstuvwxyz`.

# --hints--

Dovresti dichiarare una variabile chiamata `alphabet`. Fai attenzione a mettere il nome della variabile all'inizio della riga.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

Dovresti assegnare la stringa `abcdefghijklmnopqrstuvwxyz` alla tua variabile `alphabet`. Ricorda di usare virgolette singole o doppie per racchiudere la stringa.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
