---
id: 680a5e7f3d47715196b12ca0
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Ora che hai l'intero alfabeto spostato, rimuovi la chiamata `print(shifted_alphabet)`.

# --hints--

Dovresti rimuovere `print(shifted_alphabet)` dal tuo codice.

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
print(shifted_alphabet)
--fcc-editable-region--
```
