---
id: 680a3d4c065557260de2ff33
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Come ricorderai dalle lezioni precedenti, in Python dichiari una variabile scrivendo il nome della variabile sul lato sinistro dell'operatore di assegnazione (`=`) e il valore da assegnare sul lato destro:

```py
variable_name = value
```

Crea una variabile chiamata `shift` e assegna il valore `5` alla tua nuova variabile.

# --hints--

Dovresti dichiarare una variabile chiamata `shift`. Fai attenzione a mettere il nome della variabile all'inizio della riga.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

Dovresti assegnare il valore `5` alla tua variabile `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
