---
id: 680a7e200400089967b8b1f4
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Il metodo `str.maketrans()` prende due stringhe di uguale lunghezza e restituisce una tabella di traduzione che associa ogni carattere della prima stringa al carattere corrispondente della seconda stringa. Ogni carattere nella tabella di traduzione è archiviato come un valore Unicode ordinale, un numero che identifica univocamente il carattere.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

Imparerai di più sul tipo di struttura restituita da `str.maketrans()` più avanti nel curriculum. Per ora, crea una tabella di traduzione che associa i caratteri in `alphabet` ai caratteri in `shifted_alphabet` e assegnala a una variabile chiamata `translation_table`.

# --hints--

Dovresti dichiarare una variabile chiamata `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

Dovresti chiamare `str.maketrans()` passando `alphabet` e `shifted_alphabet` come argomenti, e assegnare il risultato a `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
