---
id: 680a41ef5bfdc059cfd9d4b5
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Nelle lezioni precedenti, hai imparato i diversi tipi di dati che puoi archiviare in una variabile. Hai appena assegnato un valore intero. Ora devi assegnare una stringa, che è una sequenza di caratteri racchiusa tra virgolette singole o doppie:

```py
string_1 = 'I am a string'
string_2 = "I am also a string"
```

Dichiara un'altra variabile chiamata `alphabet` e assegna a questa variabile la stringa `abcdefghijklmnopqrstuvwxyz`.

# --hints--

Dovresti dichiarare una variabile chiamata `alphabet`. Fai attenzione a mettere il nome della variabile all'inizio della riga.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

Dovresti assegnare la stringa `abcdefghijklmnopqrstuvwxyz` alla tua variabile `alphabet`. Ricorda di usare virgolette singole o doppie per racchiudere la stringa.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
