---
id: 680a82806c4c6bc3567f2708
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Dichiara una nuova variabile chiamata `text` e assegnale la stringa `hello world`. Questa sarà la citazione da criptare.

# --hints--

Dovresti dichiarare una variabile chiamata `text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

Dovresti assegnare la stringa `hello world` a `text`. Ricorda di racchiudere la stringa tra virgolette singole o virgolette doppie e fai attenzione alle maiuscole e minuscole.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
