---
id: 6810891bd6de8a87a833df42
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Come puoi vedere dall'output, il messaggio è stato criptato. Il passaggio successivo sarà rendere il tuo codice riutilizzabile nel caso volessi criptare messaggi diversi.

Per farlo, dovrai creare una funzione. Per ricordarti, ecco come creare una funzione chiamata `spam` che stampa `Spam!` sul terminale:

```py
def spam():
    print('Spam!')
```

Crea una funzione chiamata `caesar`. Inserisci tutto il codice esistente all'interno del corpo della funzione. Fai attenzione a mantenere lo stesso livello di indentazione per tutte le righe all'interno del corpo della funzione.

# --hints--

Dovresti avere una funzione chiamata `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

Dovresti spostare tutto il codice scritto finora all'interno del corpo della funzione `caesar`. Assicurati di mantenere lo stesso livello di indentazione per tutte le righe all'interno del corpo della funzione...

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
