---
id: 6818e9c7c8dc5694e8f64fc3
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Ora, testa la funzione `caesar` chiamandola con la stringa `freeCodeCamp` e il numero `3` come argomenti. Assegna la chiamata della funzione a una variabile chiamata `encrypted_text`.

# --hints--

Dovresti avere una variabile chiamata `encrypted_text` nell'ambito globale.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Dovresti assegnare `caesar('freeCodeCamp', 3)` alla variabile `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
