---
id: 6818eb3f98c0bca5af7c8d03
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Ora il tuo codice è riutilizzabile. Tuttavia, la funzione `caesar` stampa un messaggio sul terminale e restituisce `None` per impostazione predefinita. Per dimostrarlo, stampa `encrypted_text` alla fine del tuo codice.

# --hints--

Dovresti stampare `encrypted_text` all’esterno del corpo della funzione.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
