---
id: 6818b80519e6c6c47a8bb6e8
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Come hai imparato in una lezione precedente, le funzioni possono avere parametri, che sono variabili a cui si può fare riferimento all’interno della funzione. Ecco una funzione `sum` con due parametri, `num1` e `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

Il messaggio da criptare e lo spostamento sono ancora codificati direttamente nella tua funzione. Dai alla tua funzione due parametri chiamati `text` e `shift`, ed elimina le dichiarazioni delle variabili `text` e `shift` dal corpo della funzione `caesar`.

# --hints--

La tua funzione `caesar` dovrebbe avere due parametri, `text` e `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

Non dovresti avere una variabile `text` dichiarata all’interno del corpo della funzione `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

Non dovresti avere una variabile `shift` dichiarata all’interno del corpo della funzione `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
