---
id: 6818ed1adc249bb91d373729
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Nelle lezioni precedenti, hai imparato che l'istruzione `return` serve a restituire un valore da una funzione, così puoi usarlo altrove nel tuo codice:

```py
def spam():
    return 'Spam!'
```

Rimuovi la chiamata `print(encrypted_text)` dalla tua funzione. Poi, elimina la dichiarazione della variabile `encrypted_text` e restituisci direttamente `text.translate(translation_table)` dalla tua funzione.

# --hints--

Non dovresti avere `print(encrypted_text)` all'interno della funzione `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

Non dovresti avere una variabile `encrypted_text` all'interno della funzione `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

La tua funzione `caesar` dovrebbe restituire `text.translate(translation_table)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
