---
id: 6818fcb250f34e62cab32c39
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

Avrai notato che, anche se il messaggio è criptato, i caratteri maiuscoli sono rimasti intatti. Questo succede perché la tabella di traduzione non include le lettere maiuscole.

Per risolvere il problema, dovrai modificare le stringhe passate alla tua chiamata `str.maketrans()` usando il metodo `upper()`:

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

Aggiorna la tua chiamata `str.maketrans()` concatenando a ogni argomento la versione maiuscola dello stesso argomento.

# --hints--

Dovresti modificare gli argomenti della tua chiamata `str.maketrans()` concatenando a ogni argomento la versione maiuscola dello stesso argomento.

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
