---
id: 6819be9055828d65c813eec5
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

La funzione `isinstance()` restituisce `True` se il suo primo argomento è un’istanza del secondo argomento, altrimenti `False`:

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

Sostituisci la condizione attuale della tua istruzione `if` con una chiamata a `isinstance()`. Passa `shift` come primo argomento e `int` come secondo argomento.

# --hints--

La tua istruzione `if` dovrebbe usare `isinstance(shift, int)` come condizione.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if True:
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
