---
id: 6819bc0637b80256a33adccc
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Ora che hai implementato le funzionalità di base del cifrario, è il momento di aggiungere un po’ di validazione. Per questo, ti servirà un’istruzione `if`. Ecco un promemoria della sintassi per un’istruzione `if`:

```py
if condition:
    # code to run when condition is true
```

All’inizio del corpo della tua funzione, crea un’istruzione `if`. Per ora, usa `True` come condizione e, all’interno del corpo dell’istruzione `if`, restituisci la stringa `Shift must be an integer value.`

# --hints--

Dovresti avere un’istruzione `if` all’interno della tua funzione `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

La tua istruzione `if` dovrebbe usare `True` come condizione.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Dovresti restituire la stringa `Shift must be an integer value.` dalla tua istruzione `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
