---
id: 6819cf38880d8cc9b9f6af7a
title: Passaggio 21
challengeType: 20
dashedName: step-21
---

# --description--

Python ti permette di specificare un valore predefinito per i parametri di una funzione, creando una funzione che può essere chiamata con meno argomenti o senza argomenti. Ecco come creare una funzione con un parametro `name` che ha un valore predefinito:

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

Aggiungi un terzo parametro chiamato `encrypt` alla tua funzione e assegna a questo un valore predefinito di `True`. 

# --hints--

Dovresti aggiungere un terzo parametro chiamato `encrypt` con un valore predefinito di `True` alla tua funzione. Non modificare i parametri esistenti.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
