---
id: 6819ca464db219ab86f37dc4
title: Passaggio 20
challengeType: 20
dashedName: step-20
---

# --description--

Come hai già verificato, lo shift passato per criptare il testo dovrebbe essere un intero positivo. Tuttavia, non può superare `25` (l’ultimo indice di `alphabet`).

Aggiungi una seconda condizione all’istruzione `if` che controlli se `shift` è maggiore di `25`. Ricorda che l’operazione logica OR in Python si implementa tramite l’operatore `or`.

Aggiorna anche il messaggio restituito in `Shift must be an integer between 1 and 25.`

# --hints--

Dovresti usare l’operatore `or` per aggiungere una seconda condizione dopo `shift < 1` che controlli se `shift` è maggiore di `25`. Non dovresti modificare la condizione esistente.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

La tua seconda istruzione `if` dovrebbe restituire `Shift must be an integer between 1 and 25.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
