---
id: 6819c5c9faf2548e556163d5
title: Passaggio 18
challengeType: 20
dashedName: step-18
---

# --description--

Come ricordi dalle lezioni precedenti, l'operatore `not` è un operatore logico unario che nega un'espressione:

```py
print(not True) # False
print(not False) # True
```

Usa l'operatore `not` per correggere la condizione dell'istruzione `if` in modo che la tua funzione ritorni `Shift must be an integer value.` quando `shift` *non* è un intero.

# --hints--

Dovresti usare l'operatore `not` per negare `isinstance(shift, int)` nella tua istruzione `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
