---
id: 681a1d60e037127e7c8e0230
title: Passaggio 22
challengeType: 20
dashedName: step-22
---

# --description--

Userai il parametro aggiunto nel passaggio precedente per determinare se la funzione debba criptare il testo passato (comportamento predefinito, `encrypt=True`) oppure se debba decriptare un messaggio criptato.

Crea un'istruzione `if` che verifichi se `encrypt` non è truthy. All'interno della nuova istruzione `if`, imposta `shift` su `- shift`. Questo è necessario per permettere che lo spostamento avvenga nella direzione opposta rispetto al processo di crittografia.

# --hints--

Dovresti impostare `shift` su `- shift` quando `encrypt` non è truthy.

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
