---
id: 681a2a6cf95e82ca30866364
title: Passaggio 23
challengeType: 20
dashedName: step-23
---

# --description--

Dichiara due funzioni chiamate `encrypt` e `decrypt`, entrambe con i parametri `text` e `shift`.

Userai `encrypt` per il processo di crittografia e `decrypt` per la decrittografia, etichettando le due azioni con un nome descrittivo.

Ritorna una chiamata `caesar` passando `text` e `shift` da entrambe le tue nuove funzioni, ma assicurati di passare anche `False` come terzo argomento per la funzione `decrypt`.

# --hints--

Dovresti avere una funzione chiamata `encrypt` con due parametri, `text` e `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

La tua funzione `encrypt` dovrebbe ritornare `caesar(text, shift)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

Dovresti avere una funzione chiamata `decrypt` con due parametri, `text` e `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

La tua funzione `decrypt` dovrebbe ritornare `caesar(text, shift, False)`.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
