---
id: 68f18724d71db1dda8d6fec7
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將建置一個 HTML 影片播放器。已為你提供 HTML 樣板程式碼。

創建一個 `h1` 元素並為它設定文字 `Working with the HTML Video Element`。

# --hints--

你應該有一個 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素應該包含文字 `Working with the HTML Video Element`。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Working with the HTML Video Element');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
