---
id: 68f1fac259a4afabb478b960
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

你可能已經注意到你沒有連結到實際的 `video`。你會在下一階段這麼做。當談到 `video` 檔案型別時，瀏覽器支援會有所不同。為了因應這點，你可以在 `video` 元素內使用 `source` 元素，瀏覽器會選擇第一個相容的 `source`。

這裡是一個 `source` 元素的範例：

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

`source` 元素是空元素，因此沒有結束標籤。 

在你的 `video` 元素內添加一個 `source` 元素。 


# --hints--

你應該在你的 `video` 元素內有一個 `source` 元素。

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
