---
id: 68f1f993460f06a64c033061
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

`poster` 屬性是影片的縮圖影像。想像你在 YouTube 上觀看的影片。它會在影片下載時顯示。如果省略此屬性，則在下載階段會顯示第一個影片影格。

現在，為你的 `video` 元素添加值為 `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` 的 `poster` 屬性。

# --hints--

你的 `video` 元素應該有 `poster` 屬性，且其值為 `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`。

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
