---
id: 68f1f6f97af9749a42260a43
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

`controls` 屬性為 `video` 元素提供播放控制，包括播放、倒帶和音量控制。 

`controls` 屬性是一個布林值屬性，且不需要值。

為你的 `video` 元素添加 `controls` 屬性。

現在你應該會看到頁面上顯示 `video` 元素。

# --hints--

你的 `video` 元素應該有 `controls` 屬性。

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
