---
id: 68f1eb9ad0a7575e4dbc31d5
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

接著，在 `h1` 下方創建一個 `video` 元素。在接下來的幾個步驟中，你將添加必要的屬性，使影片播放器能夠運作。 

# --hints--

你應該有一個 `video` 元素。

```js
assert.exists(document.querySelector('video'));
```

新的 `video` 元素應該位於 `h1` 之下。

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
