---
id: 68f1f371fefd448b4f70b898
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在先前的課程中，你學習了 `video` 元素可用的不同屬性。`width` 屬性決定影片的寬度（以像素為單位）。

為你的 `video` 元素添加值為 `640` 的 `width` 屬性。

# --hints--

你的 `video` 元素應該有屬性 `width`，其值為 `640`。

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
