---
id: 68f1f57723debf9383c75947
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

`loop` 屬性會在播放完成後重新啟動影片。想像一個會重複播放的網路迷因。省略 `loop` 屬性會讓影片只播放一次。

`loop` 屬性是一個布林值屬性，且不需要值。

為你的 `video` 元素添加 `loop` 屬性。

# --hints--

你的 `video` 元素應該有 `loop` 屬性。

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
