---
id: 68f1f8c248306da204655117
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

`muted` 屬性會在初始播放時靜音音訊。如果你啟用了 `controls`，使用者將能夠取消靜音音訊。省略 `muted` 屬性會在初始播放時播放音訊。

`muted` 屬性是一個布林值屬性，且不需要值。

為你的 `video` 元素添加 `muted` 屬性。

# --hints--

你的 `video` 元素應該有 `muted` 屬性。

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
