---
id: 68f27139bddf9f6bf2821acf
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

你將添加的最後一個 `source` 元素將是用於 `video/quicktime` MIME 型別。

QuickTime 是由 Apple 生成的可擴充多媒體架構，支援播放、串流、編碼和轉碼多種數位媒體格式。由於不如 MP4 格式普及，你可能需要它來支援舊有應用程式。

在你的第三個 `source` 元素下方，添加第四個 `source` 元素，並為它設定 `src` 屬性，值為 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`，以及 `type` 屬性，值為 `video/quicktime`。

恭喜你！你完成了 HTML 影片播放器工作坊。

# --hints--

你應該有第四個 `source` 元素。

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

你的第四個 `source` 元素應該有一個 `src` 屬性。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

你的第四個 `source` 元素應該有一個 `src` 屬性，值為 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

你的第四個 `source` 元素應該有一個 `type` 屬性。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

你的第四個 `source` 元素應該有一個 `type` 屬性，其值為 `video/quicktime`。

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
