---
id: 68f26d52b13e0f5b6e2d2e09
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

另一個常見的 MIME 型別是 `video/webm` MIME 型別。

WebM 是由 Google 開發的開放原始碼音訊視訊媒體檔案格式，主要設計用於網頁媒體內容。它支援像是 VP8、VP9 和 AV1 的影片編碼器，以及 Vorbis 和 Opus 的音訊編碼器，使其成為 HTML5 影片和音訊元素的熱門選擇。

在你的第一個 `source` 元素下方，添加另一個 `source` 元素，並為它設定一個 `src` 屬性，值為 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`，以及一個 `type` 屬性，值為 `video/webm`。

# --hints--

你應該有第二個 `source` 元素。

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

你的第二個 `source` 元素應該有一個 `src` 屬性。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

你的第二個 `source` 元素應該有一個 `src` 屬性，其值為 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

你的第二個 `source` 元素應該有一個 `type` 屬性。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

你的第二個 `source` 元素應該有一個 `type` 屬性，其值為 `video/webm`。

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
