---
id: 68f26a5f79e97c5032c10642
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

你使用了副檔名為 `mp4` 的影片檔案，你需要告訴瀏覽器，讓它知道如何讀取該檔案。

你將使用 `type` 屬性來指定 `video/mp4` MIME 型別。

MIME（多用途網際網路郵件延伸功能）是一種標準，用來描述除 ASCII 文字以外的其他形式文件，例如音訊、影片和影像。

MP4，正式名稱為 MPEG-4 Part 14，是一種數位多媒體容器格式。它廣泛用於儲存影片和音訊，但也可以包含其他資料型別，如字幕和靜態影像。`MP4` 檔案設計用於透過網際網路串流，並且與許多裝置和平台相容。

現在，為你的元素添加 `type` 屬性和值 `video/mp4`。

# --hints--

你的 `source` 元素應該有一個 `type` 屬性。

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

你的 `source` 元素應該有一個 `type` 屬性，其值為 `video/mp4`。

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--
      
--fcc-editable-region--
    >
  </video>
</body>
</html>
```
