---
id: 683ec95f2f8781355556ff7b
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

在 `binary search` 中，有最低可能的索引和最高可能的索引。

最低可能的索引表现为当前被查找的值部分中的最左边界，在本例中为 `search_list`。另一方面，最高可能的索引表现为被查找值中的最右边界。

为了考虑这两个因素，你必须考虑被查找的整个列表。因此，在 `binary_search` 函数内部，定义一个值为 `0` 的 `low` 变量，以及一个值为被查找列表最后索引的 `high` 变量。

# --hints--

你应该在你的 `binary_search` 函数中有一个名为 `low` 的变量。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

你的 `low` 变量应设置为 `0`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

你应该在你的 `binary_search` 函数中有一个名为 `high` 的变量。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

你的 `high` 变量应设置为 `len(search_list) - 1`。

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []
    
--fcc-editable-region--
```
