---
id: 683ec95f2f8781355556ff84
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

如果 `elif` 中的条件为真，则删除 `pass` 关键字，并通过将 `1` 添加到 `mid` 变量来更新 `low` 变量的值。

这将把查找范围扩展到列表中当前查找区域的右半部分，因为如果 `value` 大于 `value_at_middle`，则意味着 `value` 必须在当前查找区域的右半部分。

# --hints--

你应该移除 `pass` 关键字并将 `low` 变量更新为 `mid + 1`。

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid + 1")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        elif value > value_at_middle:
            pass
--fcc-editable-region--
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
```
