---
id: 683ec95f2f8781355556ff80
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

接下来要做的是创建一个条件，用于查看目标 `value` 是否在中间。 

在 `while` 循环内，创建一个 `if` 语句，用于查看目标 `value` 是否与 `value_at_middle` 相等。如果相等，返回 `path_to_target` 变量。

# --hints--

你应该创建一个 `if` 语句来查看 `value` 是否与 `value_at_middle` 相等。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

你应该在 `if` 语句内返回 `path_to_target` 变量。

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
--fcc-editable-region--
        
--fcc-editable-region--
```
