---
id: 683ec95f2f8781355556ff7a
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

你正在创建的二分查找算法不仅会返回 `True` 或 `False` 来表示是否找到该值，还会返回查找该值所经过的路径。为此，你将有一个 `path_to_target` 变量来跟踪路径。

在你的 `binary_search` 函数内，将 `pass` 关键字替换为 `path_to_target` 变量，并将其初始化为空列表，作为起点。

# --hints--

你应该在你的 `binary_search` 函数中创建一个设置为空列表的 `path_to_target` 变量。别忘了移除 `pass` 关键字。

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    pass
--fcc-editable-region--
```
