---
id: 683ec95f2f8781355556ff7e
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

现在你已经计算出 `mid` 索引，你需要实际从 `search_list` 中该索引处获取值。这个值是你将与作为查找目标的 `value` 参数进行比较的。

创建一个 `value_at_middle` 变量，并将其赋值为位于 `mid` 索引处的 `search_list` 中的元素。 

# --hints--

你应该创建一个 `value_at_middle` 变量，设置为 `search_list[mid]`。

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2

        
--fcc-editable-region--
```
