---
id: 683ec95f2f8781355556ff82
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

请注意，该函数接受一个名为 `search_list` 的列表和你要查找的 `value`。

现在，使用 `binary_search([1, 2, 3, 4, 5], 3)` 和 `binary_search([1, 2, 3, 4, 5, 9], 4)` 调用函数，并立即打印这些调用。

# --hints--

你应该调用函数 `binary_search([1, 2, 3, 4, 5], 3)` 并立即打印它。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

你应该调用函数 `binary_search([1, 2, 3, 4, 5, 9], 4)` 并立即打印它。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        break
    return []
          
--fcc-editable-region--

--fcc-editable-region--
```
