---
id: 683ec95f2f8781355556ff7d
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

要开始这个循环，你需要查找当前查找空间的中间索引。

你可以通过计算 `low` 和 `high` 变量的平均值来获得中间索引。你必须使用地板除法在将两个值相加后获取平均值，这样答案总是向下取整的整数，因为索引必须是整数。

从你的 `while` 循环中移除 `pass` 关键字，并创建一个 `mid` 变量，设置为 `low` 和 `high` 的平均值。

# --hints--

你应该从 `while` 循环中移除 `pass` 关键字，并创建一个 `mid` 变量，设置为 `(low + high) // 2`。

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        pass
--fcc-editable-region--
```
