---
id: 683ec95f2f8781355556ff7f
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

现在，将 `value_at_middle` 包含在 `path_to_target` 列表中，以跟踪查找过程中所采取的步骤，无论它是否是目标。

# --hints--

你应该使用 `append()` 方法为你的 `path_to_target` 列表添加 `value_at_middle`。

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid] \\npath_to_target.append(value_at_middle)")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        
--fcc-editable-region--
```
