---
id: 683ec95f2f8781355556ff7c
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

既然你已经定义了查找的边界，现在是创建将执行二分查找算法的循环的时候了。

二分查找通过反复缩小查找空间来工作。只要存在有效的元素范围可供查看，这个过程就会继续。

你可以用一个 `while` 循环来表达，只要你的 `low` 指针小于或相等于你的 `high` 指针，该循环就会继续。

如果 `low` 变得大于 `high`，则表示查找空间已变为空，且 `value` 不在列表中。

在 `binary_search` 函数内，创建一个 `while` 循环，条件是检查 `low` 是否小于或相等于 `high`。在循环内，暂时添加 `pass` 关键字作为占位符。

# --hints--

你应该创建一个 `while` 循环，当 `low` 小于或等于 `high` 时运行。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
--fcc-editable-region--
    
--fcc-editable-region--
```
