---
id: 68416aab85dd00a7852cdeb6
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

你可以看到最后一次调用打印了一个空列表。为了进一步指定该值不在被查找的列表中，更新 `return` 语句并返回一个空列表和消息 `Value not found`。

# --hints--

你应该更新 `return` 语句以返回 `[], 'Value not found'`。

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").has_return("[], 'Value not found'")`
            )
        )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

--fcc-editable-region--
    return []
--fcc-editable-region--

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
print(binary_search([1, 3, 5, 9, 14, 22], 10))
```
