---
id: 684169b00a073c9b9f0a703c
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

你有没有注意到第二次函数调用现在在列表中有值 `3, 5, 4`？这表明二分查找正在按预期工作。

事情是这样发生的：算法首先将 `3` 作为初始列表的中间值进行查看。由于 `4` 大于 `3`，查找转向右半部分。然后它将 `5` 作为新的中间值进行检查。因为 `4` 小于 `5`，查找向左移动，最终确定 `4` 作为最终范围的中间值。

要再次测试该函数，使用 `[1, 3, 5, 9, 14, 22], 10` 调用它并立即打印该调用。这是一个不会找到值的情况。

# --hints--

你应该调用函数 `binary_search([1, 3, 5, 9, 14, 22], 10)` 并立即打印它。

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return []
    
print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
