---
id: 6976db7ecfa770bf21307b20
title: Jenga injini ya orodha ya nyimbo iliyorekebishwa
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

Katika maabara hii, utajenga programu inayotengeneza orodha moja ya nyimbo iliyorekebishwa kutoka kwa orodha nyingi za nyimbo zilizowasilishwa na wasikilizaji.

Kila msikilizaji hutoa orodha ya nyimbo wanazotaka kusikiliza. Baadhi ya nyimbo zinaweza kuonekana zaidi ya mara moja, na baadhi ya wasanii wanaweza kuonekana mara nyingi sana. Kazi yako ni kupitia orodha hizi hatua kwa hatua: ziunganishe kuwa orodha moja, pima kila wimbo, ondoa nyimbo zinazojirudia, weka kikomo kwa mara ngapi msanii mmoja anaonekana, kisha tengeneza mpangilio wa mwisho wa kucheza.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuunda kitendakazi kinachoitwa `flattenPlaylists` kinachokubali safu ya orodha za nyimbo ambapo kila orodha ni safu ya vitu vyenye sifa zifuatazo: `trackId`, `artist`, `title`, `votes`, `bpm`. Ikiwa ingizo si safu, `flattenPlaylists` inapaswa kurudisha safu tupu. Mfano wa orodha umepewa kwa ajili yako. Unaweza kutumia mfano huu kujaribu kitendakazi chako.

2. `flattenPlaylists` inapaswa kurudisha safu tambarare ya vitu vya nyimbo, ambapo kila kitu kina sifa zote za nyimbo asili pamoja na sifa ya `source` iliyowekwa kuwa safu yenye kielezo cha orodha na kielezo cha wimbo kinachoonyesha wapi wimbo ulitokea.

3. Unapaswa kuunda kitendakazi kinachoitwa `scoreTracks` kinachokubali safu ya vitu vya nyimbo kama ilivyo kurudishwa na `flattenPlaylists` (kila moja ikiwa na sifa za `trackId`, `artist`, `title`, `votes`, `bpm`, na `source`) na kurudisha safu mpya ya vitu vya nyimbo, kila moja ikiwa na sifa ya `score` iliyoongezwa kwa kutumia fomula: `votes * 10 - Math.abs(bpm - 120)`.

4. Unapaswa kuunda kitendakazi kinachoitwa `dedupeTracks` kinachokubali safu ya vitu vya nyimbo kama ilivyo kurudishwa na `scoreTracks` na kurudisha safu mpya yenye maingizo ya `trackId` yanayojirudia yameondolewa, ukihifadhi tu tukio la kwanza la kila moja.

5. Unapaswa kuunda kitendakazi kinachoitwa `enforceArtistQuota` kinachokubali safu ya vitu vya nyimbo kama ilivyo kurudishwa na `dedupeTracks` na nambari inayowakilisha idadi ya juu ya mara msanii mmoja anaruhusiwa kuonekana. Kitendakazi kinapaswa kurudisha safu mpya ambapo msanii yeyote haonekani zaidi ya mara iliyotolewa, ukihifadhi maonyesho ya mapema zaidi.

6. Unapaswa kuunda kitendakazi kinachoitwa `buildSchedule` kinachokubali safu ya vitu vya nyimbo kama ilivyo kurudishwa na `enforceArtistQuota` na kurudisha safu mpya ya vitu vya `{ slot, trackId }`, ambapo `slot` ni kielezo cha kuanzia 1 kinachoonyesha nafasi ya kila wimbo katika mpangilio wa matangazo.

7. Unapaswa kuunda kitendakazi kinachoitwa `remixPlaylist` kinachokubali safu ya orodha za nyimbo na idadi ya juu ya mara msanii anaruhusiwa kuonekana. Kitendakazi kinapaswa kurudisha ratiba ya mwisho ya matangazo kama safu ya vitu vya `{ slot, trackId }`, kwa kuita `flattenPlaylists`, `scoreTracks`, `dedupeTracks`, `enforceArtistQuota`, na `buildSchedule` kwa mpangilio.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `flattenPlaylists`.

```js
assert.isFunction(flattenPlaylists);
```

Unapaswa kurudisha safu tupu kutoka `flattenPlaylists` wakati ingizo si safu.

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

Kila wimbo unaorudishwa na `flattenPlaylists` unapaswa kuwa na uwanja wa `source` ambao ni safu inayojumuisha kielezo cha orodha na kielezo cha wimbo.

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

Unapaswa kuwa na kitendakazi kinachoitwa `scoreTracks`.

```js
assert.isFunction(scoreTracks)
```

Kila wimbo unaorudishwa na `scoreTracks` unapaswa kuwa na uwanja wa nambari wa `score`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

Unapaswa kuhesabu `score` kwa kutumia BPM lengwa ya `120` na fomula hii: `votes * 10 - Math.abs(bpm - 120)`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

Unapaswa kuwa na kitendakazi kinachoitwa `dedupeTracks`.

```js
assert.isFunction(dedupeTracks)
```

Wakati maingizo ya `trackId` yanaporudiwa, `dedupeTracks` inapaswa kuhifadhi tu tukio la kwanza la wimbo.

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

Unapaswa kuwa na kitendakazi kinachoitwa `enforceArtistQuota`.

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota` inapaswa kuhakikisha hakuna msanii anayeonekana zaidi ya `maxPerArtist` mara kwa kuondoa nyimbo ziizozidi huku ikihifadhi zile za mapema.

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

Unapaswa kuwa na kitendakazi kinachoitwa `buildSchedule`.

```js
assert.isFunction(buildSchedule);
```

`buildSchedule` inapaswa kurudisha safu ya vitu vyenye muundo wa `{ slot, trackId }`, ambapo `slot` huanza kwa `1`.

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

Unapaswa kuwa na kitendakazi kinachoitwa `remixPlaylist`.

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist` inapaswa kuita vitendakazi msaada kwa mpangilio ili kuzalisha ratiba ya mwisho.

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
