---
id: 612e804c54d5e7308d7ebe56
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在你的 `.keys` 元素中，為你的七個 `div` 元素添加。將它們全部設為 `key` 類別。

# --hints--

你應該創建七個新的 `div` 元素。

```js
const divDivDiv = document.querySelectorAll('div');
assert(divDivDiv?.length === 9);
```

你的七個新的 `div` 元素應該位於你的 `.keys` 元素內。

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].length === 7);
assert([...keys?.children].every(child => child?.tagName === 'DIV'));
```

你的七個新的 `div` 元素都應該將 `class` 設定為 `key`。

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].every(child => child?.classList?.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
