---
id: 612e9d142affc44a453655db
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

現在是為按鍵本身設計樣式的時候。為 `.key` 元素創建一個 `class` 選擇器。將 `background-color` 設定為值 `#ffffff`，`position` 屬性設定為 `relative`，`width` 屬性設定為 41px，`height` 屬性設定為 175px。

# --hints--

你應該有一個 `.key` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key'));
```

你的 `.key` 選擇器應該有設定 `background-color` 屬性為 `#ffffff`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor === 'rgb(255, 255, 255)');
```

你的 `.key` 選擇器應該將 `position` 屬性設為 `relative`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.position === 'relative');
```

你的 `.key` 選擇器應該設定 `width` 屬性為 `41px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.width === '41px');
```

你的 `.key` 選擇器應該設定 `height` 屬性為 `175px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.height === '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
