---
id: 612e6afc009b450a437940a1
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

從基本的 HTML 結構開始。為你的文件添加 `DOCTYPE` 宣告式以及 `html`、`head`、`body` 和 `title` 元素。

將此頁面的語言設定為 English。將 `title` 設定為 `Piano`。

# --hints--

你的程式碼應該包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你應該在 `DOCTYPE` 引用後面加上一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該定義文件型別為 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你應該在 `DOCTYPE` 宣告式的型別後面以 `>` 關閉。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素應該有一個開啟標籤，並且帶有 `lang` 屬性且值為 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 宣告應該放在你的 HTML 開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你應該有一個開啟的 `head` 頁籤。

```js
assert(code.match(/<head\s*>/i));
```

你應該有一個結束的 `head` 頁籤。

```js
assert(code.match(/<\/head\s*>/i));
```

你應該有一個開啟的 `body` 頁籤。

```js
assert(code.match(/<body\s*>/i));
```

你應該有一個結束的 `body` 頁籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是同層元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素應該位於 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該位於 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你的程式碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的專案應該有一個標題為 `Piano`。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

請記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
