---
id: 612e89d254fe5d3df7d6693d
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

既然你已經重設了 `html` 盒模型，你也需要將此設定傳遞給其中的元素。為此，你可以將 `box-sizing` 屬性設定為 `inherit`，這會告訴標的元素使用與父元素相同的值。

你還需要標的偽元素，偽元素是跟隨選擇器的特殊關鍵字。你將使用的兩個偽元素是 `::before` 和 `::after` 偽元素。

`::before` 選擇器創建一個偽元素，該元素是所選元素的第一個子元素，而 `::after` 選擇器創建一個偽元素，該元素是所選元素的最後一個子元素。這些偽元素通常用於創建裝飾性內容，你稍後會在這個專案中看到。

現在，創建一個 CSS 選擇器以 `*` 為標的所有元素，並包含偽元素 `::before` 和 `::after`。將 `box-sizing` 屬性設定為 `inherit`。

# --hints--

你應該有一個 `*, ::before, ::after` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

你的 `*, ::before, ::after` 選擇器應該將 `box-sizing` 屬性設為 `inherit`。

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
