---
id: 612e77aba7ca691f598feb02
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

為你的頁面添加兩個 `meta` 頁籤，一個用於最佳化你的頁面以適用於行動裝置，另一個用於指定頁面接受的 `charset`。

# --hints--

你應該為你的頁面添加兩個 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

你應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

一個 `meta` 元素應該設定 `name` 為 `viewport`，並將 `content` 設定為 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一個 `meta` 元素應該將 `charset` 屬性設為 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
