---
id: 612e78af05201622d4bab8aa
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

是時候開始製作鋼琴了。在你的 `body` 元素中創建一個 `div` 元素，並將 `id` 設定為 `piano`。

# --hints--

你應該創建一個新的 `div` 元素。

```js
const div = document.querySelector('div');
assert.exists(div);
```

你的 `div` 應該位於你的 `body` 元素內。

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

你的 `div` 應該將 `id` 設定為 `piano`。

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
