---
id: 612eb4893b63c75bb9251ddf
title: 第 22 步
challengeType: 0
dashedName: step-22
---

# --description--

開始透過建立 `.logo` 選擇器來設計標誌。將 `width` 設為 `200px`，`position` 設為 `absolute`，並將 `top` 設為 `23px`。

# --hints--

你應該有一個 `.logo` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.logo'));
```

你的 `.logo` 選擇器應該有一個設定為 `200px` 的 `width` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.logo')?.width === '200px');
```

你的 `.logo` 選擇器應該有一個 `position` 屬性設為 `absolute`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.logo')?.position === 'absolute');
```

你的 `.logo` 選擇器應該有一個設定為 `23px` 的 `top` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.logo')?.top === '23px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <img class="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg" alt="freeCodeCamp Logo" />
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

.key.black--key::after {
  background-color: #1d1e22;
  content: "";
  position: absolute;
  left: -18px;
  width: 32px;
  height: 100px;
}

--fcc-editable-region--

--fcc-editable-region--
```
