---
id: 587d8247367417b2b2512c38
title: 使用 helmet.frameguard() 降低点击劫持的风险
challengeType: 2
forumTopicId: 301582
dashedName: mitigate-the-risk-of-clickjacking-with-helmet-frameguard
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下公式化简介项目构建的。

你的页面可能会在未经你同意的情况下被放入 `<frame>` 或 `<iframe>` 中。这可能导致点击劫持攻击等问题。点击劫持是一种欺骗用户与页面交互的技术，用户以为是在操作一个页面，实际上却是另一个页面。通过在恶意上下文中以 iframe 方式执行你的页面，可以实现这种攻击。在这种情况下，黑客可以在你的页面上覆盖一层隐层。隐藏的按钮可能被用来运行恶意脚本。这个中间件设置了 X-Frame-Options 头部，用来限制谁可以将你的网站放入框架中。它有三种模式：DENY、SAMEORIGIN 和 ALLOW-FROM。

我们不需要让我们的应用程序可以被嵌入。

# --instructions--

使用 `helmet.frameguard()` 时应传递配置对象 `{action: 'deny'}`。

# --hints--

应正确加载 helmet.frameguard() 中间件

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(
  data.appStack,
  'frameguard',
  'helmet.frameguard() middleware is not mounted correctly'
);
```

helmet.frameguard() 'action' 应该设置为 “DENY”

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.property(data.headers, 'x-frame-options');
assert.equal(data.headers['x-frame-options'], 'DENY');
```

