---
id: 587d8248367417b2b2512c3b
title: 使用 helment.ieNoOpen() 防止 IE 打开不受信任的 HTML
challengeType: 2
forumTopicId: 301584
dashedName: prevent-ie-from-opening-untrusted-html-with-helmet-ienoopen
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下公式化简介项目构建的。

一些网络应用会提供不受信任的 HTML 供下载。某些版本的 Internet Explorer 默认会在你的网站上下文中打开这些 HTML 文件。这意味着不受信任的 HTML 页面可能会在你的页面上下文中开始执行恶意操作。这个中间件会将 X-Download-Options 头部设置为 noopen，从而防止 IE 用户在受信任的网站上下文中执行下载内容。

# --instructions--

应正确加载 `helmet.ieNoOpen()` 中间件

# --hints--

helmet.ieNoOpen() 中间件应正确安装。

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'ienoopen');
assert.equal(data.headers['x-download-options'], 'noopen');
```

