---
id: 587d8248367417b2b2512c3a
title: 用 helmet.noSniff() 避免推断出响应的 MIME 类型。
challengeType: 2
forumTopicId: 301574
dashedName: avoid-inferring-the-response-mime-type-with-helmet-nosniff
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下公式化简介项目构建的。

浏览器可以使用内容或 MIME 嗅探来覆盖响应的 `Content-Type` 标头，使用隐式内容类型猜测和处理数据。 虽然这在某些情况下可能很方便，但也可能导致一些危险的攻击。 该中间件将 `X-Content-Type-Options` 标头设置为 `nosniff`，指示浏览器不要绕过所提供的 `Content-Type`。

# --instructions--

在你的服务器上使用 `helmet.noSniff()` 方法。

# --hints--

helmet.noSniff() 中间件应该被正确安装

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nosniff');
assert.equal(data.headers['x-content-type-options'], 'nosniff');
```

