---
id: 587d8247367417b2b2512c37
title: 使用 helmet.hidePoweredBy() 隐藏潜在的危险信息
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下公式化简介项目构建的。

黑客如果发现你的网站由 Express 提供支持，就可能利用 Express/Node 中已知的漏洞。默认情况下，Express 会在每个请求中发送 `X-Powered-By: Express`。使用 `helmet.hidePoweredBy()` 中间件来移除 X-Powered-By 头部。

# --hints--

应正确地安装 helmet.hidePoweredBy() 中间件

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```

