---
id: 6a89c69c91bc7365fbd4b733
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Додайте базовий випадок у вашу `countup` функцію. Замініть `pass` на інструкцію `if`, яка перевіряє, чи `number` менше за `1`. Всередині інструкції `if` поверніть порожній список.

# --hints--

Ваша `countup` функція повинна повертати порожній список, коли `number` менше за `1`.

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

Ваша `countup` функція не повинна повертати порожній список, коли `number` більше або дорівнює `1`.

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
