---
id: 6a89c69c111694edc6d1c628
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Перевірте свою `countup` функцію, передавши їй `5` всередині виклику `print()`. Вихідні дані мають бути `[1, 2, 3, 4, 5]`.

На цьому практичне заняття з функції countup завершено.

# --hints--

Ви маєте зробити виклик `countup(5)` всередині функції `print()`.

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
