---
id: 6a89c69c8dbfda74151fa0b6
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите `countup` функцію, яка використовує рекурсію для створення списку чисел від `1` до заданого числа.

Почніть із визначення функції з назвою `countup` з одним параметром на ім'я `number`. Додайте ключове слово `pass` у тіло функції, щоб функція була дійсною, поки вона порожня.

# --hints--

У вас має бути функція з назвою `countup` з параметром на ім'я `number`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

Ваша `countup` функція має містити ключове слово `pass`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
