---
id: 6a89c69c99298438d6fbe60b
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Тепер додайте рекурсивний випадок. Після інструкції `if` створіть змінну з назвою `count_list`. Призначте їй результат виклику `countup` з аргументом `number - 1`.

Кожен виклик використовує менше число, тому рекурсія рухається до базового випадку.

# --hints--

Ви повинні призначити `countup(number - 1)` змінній з назвою `count_list`.

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
