---
id: 6a89c69c6af682856558d78b
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Рекурсивний виклик повертає список, що містить менші числа. Використовуйте `append()`, щоб додати `number` до `count_list`, а потім поверніть `count_list`.

# --hints--

Ви повинні додати `number` до `count_list`.

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

Ви повинні повернути `count_list` з вашої функції `countup`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

Ваша функція `countup` має повертати `[1, 2, 3]`, коли її викликають з `3`.

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
