---
id: 67172b43f84bcd2dec238a3d
title: Einen Buchorganisator erstellen
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten ein Array von Objekten namens `books` haben, wobei jedes Objekt im Array einen String `title`, einen weiteren String `authorName` und eine Zahl `releaseYear` enthalten sollte.

2. Ihr `books`-Array sollte mindestens drei Objekte enthalten.

3. Sie sollten eine Callback-Funktion namens `sortByYear` haben, die zwei Bücher als Parameter zum Sortieren des Arrays akzeptiert.

4. Die `sortByYear`-Funktion sollte `-1` zurückgeben, wenn das `releaseYear` des ersten Buchs kleiner ist als das des zweiten Buchs.

5. Die `sortByYear`-Funktion sollte `1` zurückgeben, wenn das `releaseYear` des ersten Buchs größer ist als das des zweiten Buchs.

6. Die `sortByYear`-Funktion sollte `0` zurückgeben, wenn beide `releaseYear`-Werte gleich sind.

7. Sie sollten Bücher, die nach einem bestimmten Jahr wie 1950 geschrieben wurden, aus dem `books`-Array herausfiltern und das gefilterte Array in einem neuen Array namens `filteredBooks` speichern.

8. Sie sollten die Bücher im `filteredBooks`-Array entsprechend ihrem `releaseYear` in aufsteigender Reihenfolge sortieren. Sie haben in einer früheren Lektion gelernt, dass die `sort()`-Methode das Array an Ort und Stelle sortiert. Das bedeutet, dass das `filteredBooks`-Array verändert wird.

# --hints--

Sie sollten eine Funktion `sortByYear` in Ihrem Code haben.

```js
assert.isFunction(sortByYear);
```

Ihre `sortByYear`-Funktion sollte zwei Parameter annehmen.

```js
assert.lengthOf(sortByYear, 2);
```

Ihre `sortByYear`-Funktion sollte `-1` zurückgeben, wenn das `releaseYear` des `book1`-Objekts kleiner ist als das des `book2`-Objekts, `1`, wenn das `releaseYear` des `book1`-Objekts größer ist als das des `book2`-Objekts, und `0` in allen anderen Fällen.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

Sie sollten ein Array `books` in Ihrem Code haben.

```js
assert.isArray(books)
```

Ihr `books`-Array sollte mindestens drei Buchobjekte enthalten.

```js
assert.isAtLeast(books.length, 3)
```

Ihr `books`-Array sollte Objekte enthalten, die jeweils einen String `title`, einen weiteren String `authorName` und eine Zahl `releaseYear` haben.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

Sie sollten ein Array `filteredBooks` in Ihrem Code haben.

```js
assert.isArray(filteredBooks)
```

Das `filteredBooks`-Array sollte einige der Bücher aus dem `books`-Array herausgefiltert enthalten und nicht leer sein.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

Das `filteredBooks`-Array sollte nur Bücher aus dem ursprünglichen `books`-Array enthalten.

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original => 
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

Das `filteredBooks`-Array sollte nur Bücher enthalten, die an oder vor einem bestimmten Jahr veröffentlicht wurden.

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

Sie sollten die höhere Ordnung Funktion `sort` aufrufen, indem Sie die Callback-Funktion `sortByYear` auf das `filteredBooks`-Array anwenden.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
