---
id: 587d7fb4367417b2b2512bfd
title: Adicionar palavras-chave ao seu package.json
challengeType: 2
forumTopicId: 301526
dashedName: add-keywords-to-your-package-json
---

# --description--

O campo `keywords` é onde você pode descrever seu projeto usando palavras-chave relacionadas. Aqui está um exemplo:

```json
"keywords": [ "descriptive", "related", "words" ],
```

Como você pode ver, este campo está estruturado como um array de frases cercadas por aspas duplas.

# --instructions--

Adicione uma array de strings adequadas para o campo `keywords` no arquivo package.json do seu projeto.

Uma das palavras-chave deve ser "freecodecamp".

# --hints--

package.json deve ter uma chave "keywords" válida

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.keywords, '"keywords" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

O campo "keywords" deve ser um array

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.isArray(data.keywords, '"keywords" is not an array');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

O campo "keywords" devem incluir "freecodecamp"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.include(
        data.keywords,
        'freecodecamp',
        '"keywords" does not include "freecodecamp"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
