---
id: 587d7fb5367417b2b2512c04
title: Remover um pacote de suas dependências
challengeType: 2
forumTopicId: 301530
dashedName: remove-a-package-from-your-dependencies
---

# --description--

Você agora já testou algumas maneiras de gerenciar as dependências do seu projeto usando a seção de dependências do package.json.Você também incluiu pacotes externos, adicionando-os ao arquivo e até dizendo ao npm quais tipos de versões você quer, usando caracteres especiais como o til ou o circunflexo.

Mas e se você quisesse remover um pacote externo do qual você não precisa mais?Você já deve ter adivinhado: apenas remova o par chave-valor correspondente a esse pacote das dependências.

Este mesmo método também se aplica à remoção de outros campos no seu package.json.

# --instructions--

Remova o pacote `@freecodecamp/example` das dependências.

**Nota:** Certifique-se de que você tem a quantidade correta de vírgulas após removê-la.

# --hints--

`"dependencies"` não deve incluir `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.notProperty(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" still includes "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
