---
id: 684cd027744c2a0520ba4d58
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS usa clases de utilidad que se aplican directamente al marcado. Estas clases incluyen estilos para fuentes, colores, diseños Flexbox, Grid y más. 

Para aplicar un color de fondo en Tailwind CSS, puedes usar el prefijo `bg`, seguido del nombre del color y un número. El número representa la tonalidad del color.

Aquí tienes un ejemplo de cómo aplicar un color de fondo azul oscuro a un elemento `paragraph`.

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

En tu etiqueta `div` de apertura, agrega una clase con la clase de utilidad `bg-red-100` que representa un color de fondo rojo claro.

# --hints--

Tu etiqueta de apertura `div` debe tener un atributo class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

El atributo class debe tener la clase de utilidad `bg-red-100`.

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
