---
id: 684cc59eed2706fa50e8933c
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, aprender치s c칩mo comenzar a trabajar con los conceptos b치sicos de Tailwind CSS construyendo un componente de mensaje de error. 

Durante los primeros pasos, te enfocar치s en construir la estructura inicial de HTML. 

Dentro del elemento `body`, agrega un elemento `div`. Luego, dentro de ese elemento `div`, agrega un elemento `p`.

# --hints--

Debe tener un elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Debes tener un elemento `p`.

```js
assert.exists(document.querySelector("p"));
```

Tu elemento `p` debe estar dentro del elemento `div`.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
