---
id: 684ccbdc1ef89e042b744e09
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

La siguiente parte del taller es agregar el botón de cierre. 

Debajo de tu elemento `p`, agrega un elemento `button` con un `aria-label` establecido en `"Dismiss"`. Dentro de ese elemento `button`, agrega un elemento `span` con un atributo `aria-hidden`. El texto del elemento `span` debe ser `&times;`. 

*NOTA*: Dado que este taller se centrará únicamente en trabajar con estilos de Tailwind CSS, el `button` no tendrá ningún tipo de interactividad añadida. 

# --hints--

Debería tener un elemento `button`.

```js
assert.exists(document.querySelector("button"));
```
 
Tu elemento `button` debe tener un atributo `aria-label`.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

Tu elemento `button` debe tener un atributo `aria-label` establecido en `"Dismiss"`.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

Tu elemento `button` debe tener un elemento `span` anidado dentro.

```js
assert.exists(document.querySelector("button span"));
```

Tu elemento `span` debe tener un atributo `aria-hidden`.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

Tu elemento `span` debe tener el texto `&times;`.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```
