---
id: 684cd59ec94b5e063bf776f0
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Sería bueno tener un borde alrededor de ese fondo rojo claro. 

Para establecer el `border-width` en Tailwind CSS, puedes usar la clase de utilidad `border-number` donde `number` representará la ligereza o grosor de tu borde. 

Aquí tienes un ejemplo de cómo aplicar un `border-width` de `4px` en todos los lados:

```html
<p class="border-4">This is how to set the border-width.</p>
```

Dentro de tu etiqueta de apertura `div`, agrega la clase de utilidad correcta para un `border-width` establecido en `2px` en todos los lados.

# --hints--

Tu etiqueta de apertura `div` debe tener una clase `border-2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
