---
id: 684cff0bf0a4ca0ad7e42990
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

El siguiente paso sería redondear las esquinas del borde. Para establecer la propiedad `border-radius` usando Tailwind CSS, puedes usar una clase de utilidad con el siguiente patrón: `rounded-size`, donde `size` representa el radio. 

Aquí tienes un ejemplo de cómo aplicar un radio de borde extra grande a un elemento:

```html
<p class="rounded-xl">This is border radius.</p>
```

Dentro de tu etiqueta `div` de apertura, aplica un radio de borde con un tamaño de `md`. 

# --hints--

Tu etiqueta de apertura `div` debe tener la clase de utilidad `rounded-md`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
