---
id: 684cca09aac0fc035db4099f
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro de tu elemento `p`, agrega el texto `Error! Something went wrong. Please try again.` donde `Error!` debe estar envuelto dentro de un elemento `strong`.

# --hints--

Debes tener un elemento `strong` dentro de tu elemento `p`.

```js
assert.exists(document.querySelector("p strong"));
```

Tu elemento `strong` debe tener el texto `Error!`.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

Tu elemento `p` debe tener el texto correcto que termina con `Something went wrong. Please try again.`

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
