---
id: 684d0e11094c3e0fe1fb98f5
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

En este momento, el elemento párrafo está apilado encima del elemento `button` de cierre. Pero sería bueno que el elemento `button` estuviera en la misma fila que el párrafo. Aquí es donde entra Flexbox.

Para establecer el elemento `div` en un diseño Flexbox en Tailwind CSS, puedes usar la clase de utilidad `flex`.

# --hints--

Tu etiqueta de apertura `div` debe tener la clase de utilidad `flex`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("flex"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
