---
id: 684d009994a3840b98c9aa3a
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

En este momento, el ancho del elemento `div` abarca todo el ancho de la página. Para aplicar anchos en Tailwind CSS, tienes varias opciones, como usar el patrón `w-number` para anchos fijos o el patrón `w-fraction` para porcentajes.  

Aquí tienes un ejemplo de cómo aplicar un ancho establecido en `66%` de su contenedor padre:

```html
<p class="w-2/3">This is a percentage width.</p>
```

Dentro de tu etiqueta de apertura `div`, aplica un ancho establecido en `50%` de su contenedor padre.

# --hints--

La etiqueta de apertura `div` debe tener la clase de utilidad `w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
