---
id: 684d0eeb9c81e5109e52fe8b
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

El elemento `div` está empezando a verse mejor. Pero ahora el `button` de cerrar está un poco demasiado cerca del elemento párrafo. 

Para crear espacio entre los elementos flex, puedes usar el siguiente patrón: `gap-number`. `number` representa el espacio fijo entre filas y columnas en una estructura flex.

Aquí tienes un ejemplo de cómo aplicar un `gap` a un elemento:

```html
<p class="flex gap-8">This is gap.</p>
```

Dentro de tu etiqueta `div` de apertura, usa la clase de utilidad correcta para establecer el `gap` en `4`.

# --hints--

Tu etiqueta de apertura `div` debe tener la clase de utilidad `gap-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
