---
id: 684d11ffd15a531225d081e9
title: Paso 17
challengeType: 0
dashedName: step-17
---

# --description--

Ahora que los estilos se han aplicado al elemento `div`, es momento de aplicar algunos estilos al elemento `p`. 

Actualmente, el color del texto para el elemento párrafo es negro. Pero sería bueno si el color fuera un rojo oscuro en su lugar.

En un paso anterior, aprendiste cómo trabajar con colores de fondo así: `bg-red-100`. Puedes hacer algo similar con el color del texto siguiendo este patrón: `text-color-shade`. 

Aquí tienes un ejemplo de cómo aplicar un color de texto azul a un elemento:

```html
<p class="text-blue-500">This has blue text.</p>
```

Comienza agregando un atributo class a tu elemento `p`. Luego, aplica la clase de utilidad correcta para establecer un color de texto rojo con el tono configurado en `700`.

# --hints--

Tu elemento `p` debe tener un atributo class.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

Tu etiqueta de apertura `p` debe tener la clase de utilidad `text-red-700`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
