---
id: 684d07cbfa43a20d6a002ed7
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Los estilos para el elemento `div` están empezando a tomar forma. Pero sería bueno que hubiera algo de espacio alrededor del texto. 

Para aplicar `padding` a un elemento usando Tailwind CSS, puedes usar el siguiente patrón: `p-number`, donde `number` representa los píxeles que quieres aplicar en los cuatro lados. 

Aquí tienes un ejemplo de cómo aplicar `padding` de `8px` a un elemento:

```html
<p class="p-8">This is padding.</p>
```

Dentro de tu etiqueta de apertura `div`, usa la clase de utilidad correcta para aplicar `padding` de `4px`.

# --hints--

Tu etiqueta de apertura `div` debe tener la clase de utilidad `p-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
