---
id: 684d1942a1904a14a3709c6c
title: Paso 20
challengeType: 0
dashedName: step-20
---

# --description--

El estilo para tu componente de mensaje de error está casi completo, pero queda una cosa por arreglar. Como configuraste diferentes tamaños de fuente para el párrafo y el texto del botón, la alineación vertical entre esos elementos flex está desajustada. 

Para posicionar elementos flex a lo largo del eje transversal de un contenedor usando Tailwind CSS, puedes usar el siguiente patrón: `items-value`. `value` puede ser cualquiera de `start`, `end`, `center`, `baseline` y más. 

Dentro de tu etiqueta `div` de apertura, usa la clase de utilidad correcta para establecer la propiedad `align-items` en `center`.

¡Y con ese último cambio, tu componente de mensaje de error está completo!

# --hints--

Tu etiqueta de apertura `div` debe tener la clase de utilidad `items-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
