---
id: 684d03df2a0ddf0c735865f0
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Ahora ves que el ancho del elemento es la mitad del tamaño del `body`. Esto funciona bien para dispositivos grandes pero no tan bien para dispositivos pequeños.

Tailwind CSS sigue un enfoque mobile-first en el diseño. Así que si quieres crear diseños responsivos para otros dispositivos, puedes usar las variantes responsivas de Tailwind como `md` y `lg` de esta manera:

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

El siguiente ejemplo establecerá el ancho en `66%` del elemento padre para tamaños de dispositivo de `1024px` o más. 

Dentro de tu etiqueta `div` de apertura, modifica tu `w-1/2` para que solo se aplique a tamaños de dispositivo de `768px` o más. Puedes anteponer esta clase con la variante `md`. 

# --hints--

Tu etiqueta de apertura `div` debe tener la clase de utilidad `md:w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

Ya no deberías tener la clase `w-1/2` en tu etiqueta `div` de apertura.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
