---
id: 684d17a3b350bb13d3c19c0f
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Ahora es momento de aplicar algunos estilos al elemento `button`. 

Comienza agregando un atributo class a tu elemento `button`. Luego, aplica la clase de utilidad correcta para establecer un color de texto rojo con el tono configurado en `700`.

Luego, dentro de tu etiqueta de apertura `button`, usa la clase de utilidad correcta para establecer el `font-size` en `2xl`, que representa `1.5rem` o (`24px`).

# --hints--

Tu elemento `button` debe tener un atributo class.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

Tu etiqueta de apertura `button` debe tener la clase de utilidad `text-red-700`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

Tu etiqueta de apertura `button` debe tener la clase de utilidad `text-2xl`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
