---
id: 684d0d369e2b3b0f288d2c20
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Sería bueno centrar horizontalmente el elemento `div`. Para aplicar `margin` a dos lados a la vez, puedes usar los siguientes patrones: `mx-number` o `my-number`. 

`mx` establece el margen para la izquierda y derecha, mientras que `my` establece el margen para la parte superior e inferior.

Aquí tienes un ejemplo de cómo aplicar los márgenes superior e inferior a un elemento:

```html
<p class="my-8">This is margin.</p>
```

Dentro de tu etiqueta `div` de apertura, usa la clase de utilidad correcta para establecer el margen izquierdo y derecho en `auto`.

# --hints--

Tu etiqueta de apertura `div` debe tener la clase de utilidad `mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
