---
id: 684d09cad4ec900e41db2e74
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

En este momento, el elemento `div` está colocado en la esquina superior izquierda de la pantalla de vista previa. Pero sería bueno agregar algo de espacio en la parte superior del elemento. Para aplicar `margin` a un elemento usando Tailwind CSS, puedes usar el siguiente patrón: `m-number`, donde `number` representa los píxeles fijos que se aplican a los cuatro lados.

Si quieres aplicar margen a solo un lado, puedes usar cualquiera de los siguientes patrones: `mt-number` (margin top), `mb-number` (margin bottom), `mr-number` (margin right) y `ml-number` (margin left).

Aquí tienes un ejemplo de cómo aplicar `margin-bottom` a un elemento:

```html
<p class="mb-8">This is margin.</p>
```

Dentro de tu etiqueta `div` de apertura, usa la clase de utilidad correcta para aplicar `margin-top` de `4px`.

# --hints--

La etiqueta de apertura `div` debe tener la clase de utilidad `mt-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
