---
id: 684d105e488dbd1166a717ae
title: Paso 16
challengeType: 0
dashedName: step-16
---

# --description--

Si aumentas el tamaño de la ventana de vista previa, notarás que todos los elementos siguen alineados a la izquierda en el contenedor. Sería bueno que los elementos flex estuvieran centrados. 

Para posicionar elementos flex a lo largo del eje principal de un contenedor con Tailwind CSS, puedes usar el siguiente patrón: `justify-value`. `value` puede ser `start`, `end`, `center`, `between` y más. Cualquier valor válido usado en CSS puro con `justify-content` puede usarse aquí con Tailwind CSS. 

Dentro de tu etiqueta `div` de apertura, usa la clase de utilidad correcta para establecer la propiedad `justify-content` en `center`.

# --hints--

Tu etiqueta de apertura `div` debe tener la clase de utilidad `justify-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
