---
id: 684d160a2370b01310ae0fda
title: Paso 18
challengeType: 0
dashedName: step-18
---

# --description--

El nuevo color rojo se ve bien, pero sería mejor si el texto fuera más grande. Para establecer el `font-size` de un elemento usando Tailwind CSS, puedes usar el siguiente patrón: `text-value`. 

Aquí tienes un ejemplo de cómo establecer el `font-size` para un elemento a `0.875rem` o (`14px`):

```html
<p class="text-sm">This is for font-size.</p>
```

Dentro de tu elemento `p` de apertura, usa la clase de utilidad correcta para establecer el `font-size` en `xl`, que representa `1.25rem` o (`20px`).

# --hints--

Tu etiqueta de apertura `p` debe tener la clase de utilidad `text-xl`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
