---
id: 68501d32b69f9d7756334ffc
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

Si amplías la ventana de vista previa, deberías ver el elemento `div` centrado horizontalmente en la página. Pero para dispositivos más pequeños, el ancho del elemento `div` abarca todo el ancho de la página.  

Actualiza la clase `mx-auto` para que solo se aplique en dispositivos `md` y mayores.

# --hints--

Tu etiqueta de apertura `div` debe tener la clase de utilidad `md:mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

Ya no deberías tener la clase `mx-auto` en tu etiqueta `div` de apertura.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
