---
id: 69a56b5069ca99f7317e6e19
title: 화물 명세서 검증기 만들기
challengeType: 26
dashedName: lab-cargo-manifest-validator
---

# --description--

이 실습에서는 JavaScript를 사용해 화물 명세서를 표준화하고 검증합니다. 화물 명세서는 일반적으로 운송되는 물품(예: 선박이나 기차로 운송되는 물품)을 나열하고 해당 물품에 대한 세부 정보를 포함하는 문서입니다.

각 화물 명세서는 다음 속성을 가진 객체로 나타냅니다:

- `containerId`: 관련 화물 컨테이너를 식별하는 양의 정수입니다.
- `destination`: 화물의 목적지를 나타내는 공백을 제거한 후 비어 있지 않은 문자열입니다.
- `weight`: 화물의 무게를 나타내는 양수입니다.
- `unit`: 화물 무게 속성의 단위를 설명하는 문자열로, `"kg"` (킬로그램) 또는 `"lb"` (파운드) 중 하나입니다.
- `hazmat`: 위험 물질 취급이 필요한지 여부를 나타내는 불리언 값입니다.

화물 명세서 객체 예시:

```js
{
  containerId: 1,
  destination: "Monterey, California, USA",
  weight: 831,
  unit: "lb",
  hazmat: false
}
```

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `normalizeUnits` 매개변수를 가진 `manifest`라는 이름의 함수를 구현하세요.
    - 이 함수는 원본 명세서 객체를 변경하지 않고, `weight`가 킬로그램으로 표준화되고 `unit`이 `"kg"`로 설정된 새 객체를 항상 반환해야 합니다.
    - 명세서 객체의 무게가 파운드(`unit: "lb"`)로 표현된 경우, 대략적인 변환 1 lb = 0.45 kg을 사용해 `weight`을 킬로그램으로 변환하고 `unit`를 적절히 업데이트해야 합니다.
    - 무게가 이미 킬로그램(`unit: "kg"`)으로 표현된 경우, `weight`과 `unit`는 변경하지 않아야 합니다.
2. `validateManifest` 매개변수를 가진 `manifest`라는 이름의 함수를 구현하세요.
    - 이 함수는 원본 명세서 객체를 변경하지 않고 항상 새 객체를 반환해야 합니다.
    - 입력 명세서가 유효하면(누락되거나 잘못된 속성이 없으면) 빈 객체를 반환해야 합니다.
    - 입력 명세서가 유효하지 않으면, 누락된 속성은 `"Missing"` 값을, 잘못된 속성은 `"Invalid"` 값을 가진 항목을 포함하는 객체를 반환해야 합니다.

    입력 객체가 `destination` 속성이 누락되고 `weight` 속성이 잘못된 경우 반환 예시:

    ```js
    {
      destination: "Missing",
      weight: "Invalid"
    }
    ```

3. `processManifest` 매개변수를 가진 `manifest`라는 이름의 함수를 구현하세요. 이 함수는 다음을 출력해야 합니다:
    - 명세서 객체가 유효하면 `Validation success: ${containerId}`를 출력한 후, 킬로그램 단위로 변환된 명세서 무게를 `Total weight: ${weight} kg` 형식으로 출력하세요. 이 변환에는 `normalizeUnits()`를 사용하세요.
    - 명세서 객체가 유효하지 않으면 `Validation error: ${containerId}`를 출력한 후, `validateManifest()` 함수를 호출해 반환된 객체를 출력하세요.
    **참고:** 두 경우 모두 `console.log()` 호출이 두 번 있어야 합니다.

# --before-each--

```js
const _validManifests = [
  {
    containerId: 4,
    destination: "Monterey, California, USA",
    weight: 831,
    unit: "lb",
    hazmat: false
  },
  {
    containerId: 5,
    destination: "Montreal, Quebec, Canada",
    weight: 151,
    unit: "kg",
    hazmat: false
  }
];

const _invalidManifests = [
  {
    containerId: -6,
    destination: 123,
    weight: 0,
    unit: "pounds",
    hazmat: true
  },
  {
    containerId: 0,
    weight: -21,
    unit: "KG"
  },
  {
    containerId: "eight",
    destination: "Guadalajara, Jalisco, Mexico",
    weight: 9001,
    hazmat: "no"
  },
  {
    destination: " ",
    weight: NaN,
    unit: "kg",
    hazmat: 1
  },
  {
    containerId: 10,
    destination: " ",
  },
];

const _testObj = {
  containerId: 1.51,
  destination: "Lavender Town",
  weight: NaN,
};

console.log = () => {};
```

# --hints--

`normalizeUnits` 매개변수를 가진 `manifest`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(normalizeUnits);
const funcString = normalizeUnits.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

`normalizeUnits()`를 인수로 `{ containerId: 68, destination: "Salinas", weight: 101, unit: "lb", hazmat: true }`를 호출하면 원본 입력을 변경하지 않고 새 객체 `{ containerId: 68, destination: "Salinas", weight: 45.45, unit: "kg", hazmat: true }`를 반환해야 합니다.

```js
const testObj = {
  containerId: 68,
  destination: "Salinas", 
  weight: 101,
  unit: "lb",
  hazmat: true
};

const expected = {
  containerId: 68,
  destination: "Salinas",
  weight: 45.45,
  unit: "kg",
  hazmat: true
};

const copy = normalizeUnits(testObj);

assert.isObject(copy);
assert.deepEqual(copy, expected, "normalizeUnits() did not return the expected normalized object");
assert.notStrictEqual(copy, testObj, "normalizeUnits() should return a new object.");
assert.strictEqual(testObj.containerId, 68, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, "Salinas", "Original object should not be mutated.");
assert.strictEqual(testObj.weight, 101, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "lb", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, true, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 5, "Original object should not be mutated.");
```

`normalizeUnits` 함수는 입력 명세서 객체의 복사본을 반환해야 하며, `weight`를 킬로그램으로 표준화하고 `unit`를 `"kg"`로 설정해야 합니다. 무게 변환에는 대략적인 변환 `1 lb = 0.45 kg`를 사용하세요.

```js
_validManifests.forEach((obj, i) => {
  const result = normalizeUnits(obj);
  assert.isObject(result, "normalizeUnits should return an object");
  assert.notStrictEqual(result, obj, "normalizeUnits must return a new object, not modify the original");
  const expectedWeight = obj.unit === "lb" ? obj.weight * 0.45 : obj.weight;
  const expected = {
    ...obj,
    weight: expectedWeight,
    unit: "kg"
  };

  assert.deepEqual(
    result,
    expected,
    `normalizeUnits did not return the expected normalized object for manifest at index ${i}`
  );
});
```

`normalizeUnits` 함수는 원본을 변경하지 않고 입력 `manifest` 객체의 새 복사본을 반환해야 합니다.

```js
const original = {
  containerId: 151,
  destination: "Saffron City",
  weight: 151,
  unit: "lb",
  hazmat: false
};

const expected = {
  containerId: 151,
  destination: "Saffron City",
  weight: 67.95,
  unit: "kg",
  hazmat: false
};

const copy = normalizeUnits(original);
assert.isObject(copy);
assert.deepEqual(copy, expected);
assert.notStrictEqual(copy, original, "normalizeUnits() should return a new object.");
assert.strictEqual(original.weight, 151, "Original object should not be mutated.");
assert.strictEqual(original.unit, "lb", "Original object should not be mutated.");
assert.equal(Object.keys(original).length, 5, "Original object should not be mutated.");
```

`validateManifest` 매개변수를 가진 `manifest`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(validateManifest);
const funcString = validateManifest.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

`validateManifest()`를 인수로 `{ containerId: 1, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }`를 호출하면 새 객체 `{}`를 반환해야 합니다.

```js
const testObj = {
  containerId: 1,
  destination: "Santa Cruz", 
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {};

const copy = validateManifest(testObj);

assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for a valid manifest");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, 1, "Original object should not be mutated");
assert.strictEqual(testObj.destination, "Santa Cruz", "Original object should not be mutated");
assert.strictEqual(testObj.weight, 304, "Original object should not be mutated");
assert.strictEqual(testObj.unit, "kg", "Original object should not be mutated");
assert.strictEqual(testObj.hazmat, false, "Original object should not be mutated");
```

입력 명세서 객체가 유효하면 `validateManifest` 함수는 빈 객체 `{}`를 반환해야 합니다.

```js
const solution = {}
for (const obj of _validManifests) {
  assert.deepEqual(
    validateManifest(obj),
    solution,
    "validateManifest() should return an empty object for valid manifests."
  );
}
```

`validateManifest()`를 인수로 `{}`를 호출하면 원본 입력을 변경하지 않고 새 객체 `{ containerId: "Missing", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }`를 반환해야 합니다.

```js
const testObj = {

};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.equal(Object.keys(testObj).length, 0, "validateManifest() should return a new object, not mutate the original")
```

`validateManifest()`를 인수로 `{ containerId: null, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }`를 호출하면 원본 입력을 변경하지 않고 새 객체 `{ containerId: "Invalid" }`를 반환해야 합니다.

```js
const testObj = {
  containerId: null,
  destination: "Santa Cruz",
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {
  containerId: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for null containerId");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
```

`validateManifest()`를 인수로 `{ containerId: 0, destination: 405, weight: -84, unit: "pounds", hazmat: "no" }`를 호출하면 원본 입력을 변경하지 않고 새 객체 `{ containerId: "Invalid", destination: "Invalid", weight: "Invalid", unit: "Invalid", hazmat: "Invalid" }`를 반환해야 합니다.

```js
const testObj = {
  containerId: 0,
  destination: 405,
  weight: -84,
  unit: "pounds",
  hazmat: "no"
};

const expected = {
  containerId: "Invalid",
  destination: "Invalid",
  weight: "Invalid",
  unit: "Invalid",
  hazmat: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, 0, "Original object should not be mutated");
assert.strictEqual(testObj.destination, 405, "Original object should not be mutated");
assert.strictEqual(testObj.weight, -84, "Original object should not be mutated");
assert.strictEqual(testObj.unit, "pounds", "Original object should not be mutated");
assert.strictEqual(testObj.hazmat, "no", "Original object should not be mutated");
```

`validateManifest()`를 인수로 `{ containerId: -2 }`를 호출하면 원본 입력을 변경하지 않고 새 객체 `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }`를 반환해야 합니다.

```js
const testObj = {
  containerId: -2,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, -2, "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

`validateManifest()`를 인수로 `{ containerId: 3.50 }`를 호출하면 원본 입력을 변경하지 않고 새 객체 `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }`를 반환해야 합니다. 정수 값을 검증하려면 `Number.isInteger()`를 사용할 수 있습니다.

```js
const testObj = {
  containerId: 3.50,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, 3.50, "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

`validateManifest()`를 인수로 `{ destination: "  " }`를 호출하면 원본 입력을 변경하지 않고 새 객체 `{ containerId: "Missing", destination: "Invalid", weight: "Missing", unit: "Missing", hazmat: "Missing" }`를 반환해야 합니다. 문자열에서 공백을 제거하려면 `String.trim()`를 사용할 수 있습니다.

```js
const testObj = {
  destination: "  ",
};

const expected = {
  containerId: "Missing",
  destination: "Invalid",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.destination, "  ", "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

`validateManifest()`를 인수로 `{ weight: NaN }`를 호출하면 원본 입력을 변경하지 않고 새 객체 `{ containerId: "Missing", destination: "Missing", weight: "Invalid", unit: "Missing", hazmat: "Missing" }`를 반환해야 합니다. NaN 값을 검증하려면 `Number.isNaN()`를 사용할 수 있습니다.

```js
const testObj = {
  weight: NaN,
};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Invalid",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.ok(Number.isNaN(testObj.weight), "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

입력 명세서 객체가 유효하지 않으면 `validateManifest` 함수는 누락되거나 잘못된 속성을 설명하는 객체를 반환해야 합니다.

```js
const solutions = [
  { containerId: 'Invalid', destination: 'Invalid', weight: 'Invalid', unit: 'Invalid' },
  { containerId: 'Invalid', destination: 'Missing', weight: 'Invalid', unit: 'Invalid', hazmat: 'Missing' },
  { containerId: 'Invalid', unit: 'Missing', hazmat: 'Invalid' },
  { containerId: 'Missing', destination: 'Invalid', weight: 'Invalid', hazmat: 'Invalid' },
  { destination: 'Invalid', weight: 'Missing', unit: 'Missing', hazmat: 'Missing' }
];

let i = 0;
for (const obj of _invalidManifests) {
  const result = validateManifest(obj);
  assert.deepEqual(
    result,
    solutions[i],
    "validateManifest() should return an object with errors for invalid manifests."
  );
  i += 1;
}
```

`validateManifest` 함수는 원본을 변경하지 않고 새 객체를 반환해야 합니다.

```js
const invalidOriginal = {
  containerId: -151,
  destination: "",
  weight: -151,
  unit: "kilograms",
  hazmat: "idk"
};

const validOriginal = {
  containerId: 151,
  destination: "Pallet Town",
  weight: 151,
  unit: "kg",
  hazmat: false
};

const res0 = validateManifest(invalidOriginal);
assert.isObject(res0);
assert.notStrictEqual(res0, invalidOriginal, "validateManifest() should return a new object.");
assert.strictEqual(invalidOriginal.containerId, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.destination, "", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.weight, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.unit, "kilograms", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.hazmat, "idk", "Original object should not be mutated.");
assert.equal(Object.keys(invalidOriginal).length, 5, "Original object should not be mutated.");

const res1 = validateManifest(validOriginal);
assert.isObject(res1);
assert.notStrictEqual(res1, validOriginal, "validateManifest() should return a new object.");
assert.strictEqual(validOriginal.containerId, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.destination, "Pallet Town", "Original object should not be mutated.");
assert.strictEqual(validOriginal.weight, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.unit, "kg", "Original object should not be mutated.");
assert.strictEqual(validOriginal.hazmat, false, "Original object should not be mutated.");
assert.equal(Object.keys(validOriginal).length, 5, "Original object should not be mutated.");
```

`processManifest` 매개변수를 가진 `manifest`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(processManifest);
const funcString = processManifest.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

`processManifest()`를 인수로 `{ containerId: 55, destination: "Carmel", weight: 400, unit: "lb", hazmat: false }`를 호출하면 먼저 `"Validation success: 55"`를 출력한 후 `"Total weight: 180 kg"`를 출력해야 합니다.

```js
const testObj = {
  containerId: 55,
  destination: "Carmel",
  weight: 400,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);

  const expectedCalls = [
    ["Validation success: 55"],
    ["Total weight: 180 kg"]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation success message");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the normalized weight in kg");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

입력 명세서 객체가 유효하면 `processManifest` 함수는 먼저 성공 메시지 `Validation success: ${containerId}`를 출력해야 합니다.

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ]
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

입력 명세서 객체가 유효하면 `processManifest` 함수는 `normalizeUnits()`를 사용해 킬로그램으로 표준화한 후 `Total weight: ${weight} kg`를 출력해야 합니다.

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ]
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

입력 명세서 객체가 유효하면 `processManifest` 함수는 객체의 `containerId`와 함께 성공 메시지를 출력한 후, `weight`를 사용해 킬로그램 단위로 변환된 객체의 `normalizeUnits()`를 출력해야 합니다. `console.log()` 호출이 두 번 있어야 합니다.

```js
const originalNormalize = normalizeUnits;
let called = false;
try {
  normalizeUnits = function (...args) {
    called = true;
    return originalNormalize(...args);
  };

  processManifest(_validManifests[0]);

  assert.isTrue(
    called,
    "processManifest should call normalizeUnits() when the manifest is valid"
  );

} catch (err) {
  assert.fail(err);
} finally {
  normalizeUnits = originalNormalize;
}

const spy = __helpers.spyOn(console, "log");

try {
  for (const obj of _validManifests) {
    processManifest(obj);
  }
  const expectedCalls = [
    ["Validation success: 4"],
    ["Total weight: 373.95 kg"],
    ["Validation success: 5"],
    ["Total weight: 151 kg"],
  ]
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

`processManifest()`를 인수로 `{ containerId: -88, destination: "Soledad", weight: NaN }`를 호출하면 먼저 `Validation error: -88`를 출력한 후 객체 `{ containerId: "Invalid", weight: "Invalid", unit: "Missing", hazmat: "Missing" }`를 출력해야 합니다.

```js
const testObj = {
  containerId: -88,
  destination: "Soledad",
  weight: NaN
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: -88"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

`processManifest()`를 인수로 `{ destination: "Watsonville", hazmat: true }`를 호출하면 먼저 `Validation error: undefined`를 출력한 후 객체 `{ containerId: "Missing", weight: "Missing", unit: "Missing" }`를 출력해야 합니다.

```js
const testObj = {
  destination: "Watsonville",
  hazmat: true
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: undefined"],
    [{
        containerId: "Missing",
        weight: "Missing",
        unit: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

입력 명세서 객체가 유효하지 않으면 `processManifest` 함수는 먼저 오류 메시지 `Validation error: ${containerId}`를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);   
} finally {
  spy.restore();
}
```

입력 명세서 객체가 유효하지 않으면 `processManifest` 함수는 원본 명세서 객체로 `validateManifest()`를 호출해 반환된 객체도 출력해야 합니다. `console.log()`를 반환된 객체에 직접 호출하세요.

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest should call validateManifest() when the manifest is invalid"
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);   
} finally {
  spy.restore();
}
```

입력 명세서 객체가 유효하지 않으면 `processManifest` 함수는 객체의 `containerId`와 함께 오류 메시지를 출력한 후, 입력 객체로 `validateManifest()`를 호출해 반환된 객체를 출력해야 합니다. 이 과정에서 `console.log()` 호출이 두 번 있어야 합니다.

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest should call validateManifest() when the manifest is invalid"
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);   
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function normalizeUnits(manifest) {
  const normalized = {
    ...manifest
  }
  if (normalized.unit === "lb") {
    normalized.weight = normalized.weight * 0.45;
    normalized.unit = "kg";
  }
  return normalized;
}

function validateManifest(manifest) {
  const result = {}

  if (manifest.containerId === undefined) {
    result.containerId = "Missing";
  } else if (
    typeof manifest.containerId !== "number" ||
    !Number.isInteger(manifest.containerId) ||
    manifest.containerId <= 0
  ) {
    result.containerId = "Invalid";
  }
  if (manifest.destination === undefined) {
    result.destination = "Missing";
  } else if (typeof manifest.destination !== "string" || manifest.destination.trim() === "") {
    result.destination = "Invalid";
  }
  if (manifest.weight === undefined) {
    result.weight = "Missing";
  } else if (typeof manifest.weight !== "number" || manifest.weight <= 0 || Number.isNaN(manifest.weight)) {
    result.weight = "Invalid";
  }
  if (manifest.unit === undefined) {
    result.unit = "Missing";
  } else if (manifest.unit !== "kg" && manifest.unit !== "lb") {
    result.unit = "Invalid";
  }
  if (manifest.hazmat === undefined) {
    result.hazmat = "Missing";
  } else if (typeof manifest.hazmat !== "boolean") {
    result.hazmat = "Invalid";
  }

  return result;
}

function processManifest(manifest) {
  const validation = validateManifest(manifest);

  if (Object.keys(validation).length > 0) {
    console.log(`Validation error: ${manifest.containerId}`);
    console.log(validation);
  } else {
    const normalized = normalizeUnits(manifest);
    console.log(`Validation success: ${manifest.containerId}`);
    console.log(`Total weight: ${normalized.weight} kg`);
  }
}
```
