---
id: 669aff9f5488f1bea056416d
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本工作坊中，你將通過創建一個專門介紹貓 Whiskers 先生的博客頁面來練習使用語義化 HTML。

開始項目，添加 `<!DOCTYPE html>`，添加一個 `html` 元素，其 `lang` 屬性值爲 `en`。

請記住，你在上一模塊中學過如何創建類似的基本 HTML 模板。

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

# --hints--

你應該有 `<!DOCTYPE html>`。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你應該有一個 `html` 開始標籤，語言設置爲英語。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

You should have a closing `html` tag.

```js
assert.match(code, /<\/html>/i);
```

你的 `DOCTYPE` 應放在 `html` 元素之前。

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
