---
id: 669fd6fd12918e3de87854d4
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

為你的部落格，應該有一種方式讓使用者能夠導覽頁面上的不同區段。

`nav` 元素用於提供導覽聯結，連結到文件中的其他區段或網站中的其他區段。你經常會看到 `nav` 元素用於選單或目錄。 

這裡是一個使用 `nav` 元素的範例：

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

在你的 `figure` 元素下方，添加一個 `nav` 元素，並在其中巢狀一個 `ul` 元素。

在 `ul` 元素內，為你的 `ul` 添加三個 `li` 元素。

# --hints--

你的 `nav` 元素應該有一個開啟的頁籤。 

```js
assert.match(code, /<nav\>/);
```

你的 `nav` 元素應該有一個結束標籤。 

```js
assert.match(code, /<\/nav\>/);
```

你的 `nav` 元素應該位於你的 `header` 元素內。

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

你的 `ul` 元素應該有一個開啟的標籤。 

```js
assert.match(code, /<ul\>/);
```

你的 `ul` 元素應該有一個結束標籤。 

```js
assert.match(code, /<\/ul\>/);
```

你的 `ul` 元素應該位於你的 `nav` 元素內。

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

你的 `ul` 元素應該包含三個 `li` 元素。

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
