---
id: 669fd003cb89ee3c2402e041
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

在這個介紹性內容中，需要展示一張 Whiskers 先生的圖片，並配上說明。

在 `h1` 元素下方，首先添加一個 `figure` 元素。

在 `figure` 元素中，添加一個 `img` 元素。

`img` 的 `src` 屬性值應爲 `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`，`alt` 文本應爲 `"a cat in the garden"`。

在你的 `img` 元素下方，添加一個 `figcaption</code`，文本爲 `Mr. Whiskers in the Garden`。

# --hints--

你的 `figure` 元素應該有一個開始標籤。

```js
assert.match(code, /<figure\>/);
```

你的 `figure` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/figure\>/);
```

你的 `figure` 元素應該在 `header` 元素中。

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

你應該有一個 `img` 元素。

```js
assert.exists(document.querySelector("img"));
```

你的 `img` 元素應在 `figure` 元素中。

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

你的 `img` 元素應該有一個 `src` 屬性，設置爲 `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`。

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

你的 `img` 元素應該有一個 `alt` 屬性，設置爲 `"a cat in the garden"`。

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

你的 `figcaption` 元素應該有一個開始標籤。

```js
assert.match(code, /<figcaption\>/);
```

你的 `figcaption` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/figcaption\>/);
```

你的 `figcaption` 元素應在 `figure` 元素中。

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

你的 `figcaption` 元素的文本應爲 `Mr. Whiskers in the Garden`。

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

你的 `figcaption` 應在 `img` 元素下方。 你把它們的順序寫錯了。

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--

--fcc-editable-region--
    </header>
  </body>
</html>
```
