---
id: 669fde2081f65141ad703fe4
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

頁面的第一個部分將是關於部分。 該部分將介紹 Whiskers 先生，讓用戶瞭解本博客的內容。

在你的 `main` 元素中，添加一個 `section` 元素，其 `id` 屬性設置爲 `"about"`。

在 `section` 元素中，添加一個 `h2`，其文本爲 `About`。

# --hints--

你的 `section` 元素應該有一個開始標籤。

```js
assert.match(code, /<section.*>/);
```

你的 `section` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/section\>/);
```

你的 `section` 元素應該有一個 `id`，設置爲 `"about"`。

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

你的 `section` 元素應該在 `main` 元素中。

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

你的 `h2` 元素應該有一個開始標籤。

```js
assert.match(code, /<h2\>/);
```

你的 `h2` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/h2\>/);
```

你的 `h2` 元素的文本應爲 `About`。 仔細檢查你的拼寫。

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

你的 `h2` 元素應該在 `section` 元素中。

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
</html>
```
