---
id: 669fcb5106eeb13ab7afcd09
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this workshop, you will practice working with semantic HTML by building a blog page dedicated to Mr. Whiskers the cat.

你要創建的第一個部分是<dfn>頁頭</dfn>。

`header` 元素用於表示介紹性內容，如頁面導航和其他介紹性信息。

下面是一個使用 `header` 元素的示例：

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

在 `body` 元素內，添加一個 `header` 元素。

# --hints--

你應該有一個 `header` 開始標籤。

```js
assert.match(code, /<header>/i);
```

你應該有一個 `header` 結束標籤。

```js
assert.match(code, /<\/header>/i);
```

你的 `header` 元素應該在 `body` 元素中。

```js
assert.match(code, /<\/header>[.\n\s]*<\/body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```
