---
id: 669fd003cb89ee3c2402e041
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在這段介紹內容中，你會想要顯示一張 Mr. Whiskers 的影像，並附上說明文字。

在 `h1` 元素下方，開始為你的內容添加一個 `figure` 元素。

在 `figure` 元素內，添加一個 `img` 元素。

`img` 的 `src` 屬性應該有一個值為 `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`，而 `alt` 文字應該有一個值為 `"a cat in the garden"`。

在你的 `img` 元素下方，為你的 `figcaption` 添加文字 `Mr. Whiskers in the Garden`。

# --hints--

你的 `figure` 元素應該有一個開啟的頁籤。 

```js
assert.match(code, /<figure\>/);
```

你的 `figure` 元素應該有一個結束標籤。 

```js
assert.match(code, /<\/figure\>/);
```

你的 `figure` 元素應該位於你的 `header` 元素內。

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

你應該有一個 `img` 元素。

```js
assert.exists(document.querySelector("img"));
```

你的 `img` 元素應該位於你的 `figure` 元素內。

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

你的 `img` 元素應該有一個 `src` 屬性設定為 `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`。

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

你的 `img` 元素應該設定 `alt` 屬性為 `"a cat in the garden"`。

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

你的 `figcaption` 元素應該有一個開啟的頁籤。 

```js
assert.match(code, /<figcaption\>/);
```

你的 `figcaption` 元素應該有一個結束標籤。 

```js
assert.match(code, /<\/figcaption\>/);
```

你的 `figcaption` 元素應該位於你的 `figure` 元素內。

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

你的 `figcaption` 元素應該有 `Mr. Whiskers in the Garden` 的文字。

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

你的 `figcaption` 應該在 `img` 元素下方。你把它們的順序弄錯了。

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
