---
id: 669fdfc9a5e3da42d2376609
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

既然你已經為你的 about 區段添加內容，試著點擊 `About` 聯結，看看頁面是否跳轉到該區段。

部落格頁面的下一個區段將會是一個討論 Mr. Whiskers 的文章列表。

為你的 `section` 元素添加另一個 `section`，並將 `id` 設定為 `"posts"`。

在 `section` 元素內，為你的 `section` 元素添加一個 `h2` 元素，並且文字為 `Posts`。

# --hints--

你的程式碼應該有一個開啟的 `section` 頁籤。

```js
assert.match(code, /<section.*>/i);
```

你的程式碼應該有一個結束的 `section` 頁籤。

```js
assert.match(code, /<\/section>/i);
```

你的 `section` 元素應該設定一個 `id` 為 `"posts"`。

```js
assert.match(code, /<section.*id="posts".*>/i);
```

你的 `section` 元素應該位於你的 `main` 元素內。

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

你的 `h2` 元素應該在 `section` 元素內。

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

你的 `h2` 元素應該有 `Posts` 的文字。

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
