---
id: 669fde2081f65141ad703fe4
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

頁面上的第一個區段將是關於區段。該區段將介紹 Mr. Whiskers 並讓使用者了解這個部落格的內容。 

在你的 `main` 元素內，為你的 `section` 元素添加 `id` 屬性並設為 `"about"`。 

在 `section` 元素內，為你的 `section` 添加一個帶有 `About` 文字的 `h2`。

# --hints--

你的 `section` 元素應該有一個開啟的頁籤。 

```js
assert.match(code, /<section.*>/);
```

你的 `section` 元素應該有一個結束標籤。 

```js
assert.match(code, /<\/section\>/);
```

你的 `section` 元素應該設定 `id` 為 `"about"`。

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

你的 `section` 元素應該位於你的 `main` 元素內。

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

你的 `h2` 元素應該有一個開啟的標籤。 

```js
assert.match(code, /<h2\>/);
```

你的 `h2` 元素應該有一個結束標籤。 

```js
assert.match(code, /<\/h2\>/);
```

你的 `h2` 元素應該有 `About` 的文字。請再次確認你的拼寫。

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

你的 `h2` 元素應該位於你的 `section` 元素內。

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
