---
id: 669fcb5106eeb13ab7afcd09
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個專屬於貓咪 Mr. Whiskers 的部落格頁面來練習使用語意化 HTML。 

你將建置的第一個區段是頁面 <dfn>header</dfn>。 

`header` 元素用於表述像頁面導覽和其他介紹性資訊等的介紹內容。 

以下是使用 `header` 元素的範例：

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

在 `body` 元素內，為你的 `header` 元素添加。

# --hints--

你應該有一個開啟的 `header` 頁籤。

```js
assert.match(code, /<header>/i);
```

你應該有一個結束的 `header` 頁籤。

```js
assert.match(code, /<\/header>/i);
```

你的 `header` 元素應該在 `body` 元素內。

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
