---
id: 669fd6fd12918e3de87854d4
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

用戶應該有辦法瀏覽你的博客頁面上的不同部分。

`nav` 元素用於提供指向文檔中其他部分或網站中其他部分的導航鏈接。 你經常會看到用於菜單或目錄的 `nav` 元素。

下面是使用 `nav` 元素的示例：

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

在 `figure` 元素下方，添加一個 `nav` 元素，在其中嵌套一個 `ul` 元素。

在 `ul` 元素中，添加三個 `li` 元素。

# --hints--

你的 `nav` 元素應有一個開始標籤。

```js
assert.match(code, /<nav\>/);
```

你的 `nav` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/nav\>/);
```

你的 `nav` 元素應該在 `header` 元素中。

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

你的 `ul` 元素應該有一個開始標籤。

```js
assert.match(code, /<ul\>/);
```

你的 `ul` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/ul\>/);
```

你的 `ul` 元素應在 `nav` 元素中。

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

你的 `ul` 元素應包含三個 `li` 元素。

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--

--fcc-editable-region--
    </header>
  </body>
</html>
```
