---
id: 669fdfc9a5e3da42d2376609
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

現在你已添加了關於部分，請嘗試點擊 `About` 鏈接，查看頁面是否能跳轉到該部分。

博客頁面的下一部分將是有關 Whiskers 先生的博文列表。

添加另一個 `section` 元素，其 `id` 設置爲 `"posts"`。

在 `section` 元素中，添加一個 `h2` 元素，其文本爲 `Posts`。

# --hints--

你的代碼應該有一個 `section` 開始標籤。

```js
assert.match(code, /<section.*>/i);
```

你的代碼應該有一個 `section` 結束標籤。

```js
assert.match(code, /<\/section>/i);
```

你的 `section` 元素應該有一個 `id`，設置爲 `"posts"`。

```js
assert.match(code, /<section.*id="posts".*>/i);
```

你的 `section` 元素應該在 `main` 元素中。

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

你的 `h2` 元素應該在 `section` 元素中。

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

你的 `h2` 元素的文本應爲 `Posts`。

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
  </body>
</html>
```
