---
id: 669fdc11c9b0263fe0814a7a
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在每個 `li` 元素內，你需要有一個錨點元素。

對於第一個錨點元素，文字應該是 `About`，且 `href` 屬性值應該是 `"#about"`。`about` 前面的井字號符號表述一個 `id` 名稱，該名稱將稍後在專案中添加。

對於第二個錨點元素，文字應該是 `Posts`，且 `href` 屬性值應該是 `"#posts"`。

對於第三個錨點元素，文字應該是 `Contact`，且 `href` 屬性值應該是 `"#contact"`。

# --hints--

你的第一個 `li` 元素內的第一個錨點元素應該包含 `About` 的文字。

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

你的第一個錨點元素應該有一個 `href` 屬性設定為 `"#about"`。

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

你的第二個 `li` 元素內的第二個錨點元素應該包含 `Posts` 的文字。

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

你的第二個錨點元素應該有一個 `href` 屬性設定為 `"#posts"`。

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

你的第三個 `li` 元素內的第三個錨點元素應該包含 `Contact` 的文字。

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

你的第三個錨點元素應該有一個 `href` 屬性設定為 `"#contact"`。

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
