---
id: 669fdd3965f36f40af9615e5
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

現在你已經完成建置頁首，接著你需要開始添加你的主要內容。 

在你的 `header` 下方，添加一個 `main` 元素。 

# --hints--

你的 `main` 元素應該有一個開啟的頁籤。 

```js
assert.match(code, /<main\>/);
```

你的 `main` 元素應該有一個結束標籤。 

```js
assert.match(code, /<\/main\>/);
```

你的 `main` 元素應該位於你的 `header` 元素下方。

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
