---
id: 669fdc11c9b0263fe0814a7a
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

在每個 `li` 元素中，都需要有一個錨元素。

第一個錨元素的文本應爲 `About`，`href` 屬性值應爲 `"#about"`。 `about` 前面的哈希符號代表一個 `id` 名稱，稍後將添加到項目中。

第二個錨元素的文本應爲 `Posts`，`href` 屬性值應爲 `"#posts"`。

第三個錨元素的文本應爲 `Contact`，`href` 屬性值應爲 `"#contact"`。

# --hints--

你的第一個錨元素應在第一個 `li` 元素中包含 `About` 文本。

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

你的第一個錨元素應具有 `href` 屬性，設置爲 `"#about"`。

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

你的第二個錨元素應在第二個 `li` 元素中包含 `Posts` 文本。

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

你的第二個錨元素應具有 `href` 屬性，設置爲 `"#posts"`。

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

你的第三個錨元素應在第三個 `li` 元素中包含 `Contact` 文本。

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

你的第三個錨元素應具有 `href` 屬性，設置爲 `"#contact"`。

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
