---
id: 66a334a1a7cca6354999f9bf
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

在第一篇部落格文章中，你將使用 <dfn>article</dfn> 元素。

`article` 元素表述網頁上自包含的內容。

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

在 `h2` 元素下方，添加一個 `article` 元素。

在 `article` 元素內，為你的 `h3` 元素添加文字 `Mr. Whiskers' First Day Home`。

這裡使用 `h3` 的原因是維護標題元素的正確結構階層體系很重要。由於文章的副標題是 `h2` 元素，階層體系中下一個層級將是 `h3`。

# --hints--

你的程式碼應該有一個開啟的 `article` 頁籤。

```js
assert.match(code, /<article>/i);
```

你的程式碼應該有一個結束的 `article` 頁籤。

```js
assert.match(code, /<\/article>/i);
```

你的 `article` 元素應該位於你的 `section` 元素內，且 `id` 設為 `posts`。

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

你的程式碼應該有一個開啟的 `h3` 頁籤。

```js
assert.match(code, /<h3>/i);
```

你的程式碼應該有一個結束的 `h3` 頁籤。

```js
assert.match(code, /<\/h3>/i);
```

你的 `h3` 元素應該有 `Mr. Whiskers' First Day Home` 的文字。

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

你的 `h3` 元素應該位於你的 `article` 元素內。

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
