---
id: 66a33a00514b40393a983c24
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

為你的第三篇部落格文章添加另一個 `article` 元素。

在 `article` 元素內，為你的 `article` 添加一個 `h3` 元素，並將文字設為 `Mr. Whiskers' First Birthday Party`。

在你的 `h3` 元素下方，添加兩段 lorem ipsum 文字。

# --hints--

你應該在你的 `section` 元素中有第三個 `article` 元素，且其 `id` 設為 `posts`。

```js
const articleElement = document.querySelector('#posts article:nth-of-type(3)');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

你應該在 `article` 元素內有一個 `h3` 元素，其文字為 `Mr. Whiskers' First Birthday Party`。

```js
const h3Element = document.querySelector('#posts article:nth-of-type(3) h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Birthday Party");
```

你的 `h3` 元素應該位於你的 `article` 元素內。

```js
const h3Element = document.querySelector('#posts article:nth-of-type(3) h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

你應該在 `article` 元素內有兩個段落元素。

```js
const paragraphElements = document.querySelectorAll('#posts article:nth-of-type(2) p');
assert.strictEqual(paragraphElements.length, 2);
```

你的第一個段落元素應該包含 lorem ipsum 文字。請再次確認拼寫和標點符號錯誤。 

```js
const paragraphElement = document.querySelector('#posts article:nth-of-type(3) p:first-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim(), loremIpsum);
```

你的第二個段落元素應該包含 lorem ipsum 文字。請仔細檢查拼寫和標點符號錯誤。 

```js
const paragraphElement = document.querySelector('#posts article:nth-of-type(3) p:last-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim(), loremIpsum);
```

你的段落應該放在 `h3` 元素之後。

```js
let h3 = document.querySelectorAll('#posts article:nth-of-type(3) > p ~ h3');
assert.isEmpty(h3);
const paragraphs = document.querySelectorAll('#posts article:nth-of-type(3) > h3 ~ p');
assert.lengthOf(paragraphs, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>

        <article>
          <h3>Mr. Whiskers' First Day Home</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
        <article>
          <h3>Mr. Whiskers' First Bath</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
