---
id: 66a49f685961e997e337cab1
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

`header` 將負責顯示博客的主標題、圖像和頁面導航。

在 `header` 元素中，添加一個 `h1`，其文本爲 `Welcome to Mr. Whiskers' Blog Page!`。

# --hints--

你應該有一個 `h1` 開始標籤。

```js
assert.match(code, /<h1>/i);
```

你應該有一個 `h1` 結束標籤。

```js
assert.match(code, /<\/h1>/i);
```

你的 `h1` 元素的文本應爲 `Welcome to Mr. Whiskers' Blog Page!`。 仔細檢查是否有拼寫錯誤。

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

你的 `h1` 元素應該嵌套在你的 `header` 元素中。

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--

--fcc-editable-region--
    </header>
  </body>
</html>
```
