---
id: 66a49f685961e997e337cab1
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

`header` 將負責顯示部落格的主標題、影像和頁面導覽。

在 `header` 元素內，為你的 `h1` 添加文字 `Welcome to Mr. Whiskers' Blog Page!`。

# --hints--

你應該有一個開啟的 `h1` 頁籤。

```js
assert.match(code, /<h1>/i);
```

你應該有一個結束的 `h1` 頁籤。

```js
assert.match(code, /<\/h1>/i);
```

你的 `h1` 元素應該包含 `Welcome to Mr. Whiskers' Blog Page!` 的文字。請仔細檢查是否有任何拼寫錯誤。

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

你的 `h1` 元素應該巢狀在你的 `header` 元素內。

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
