---
id: 66a7e72adf226c02626715a3
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

在 `h2` 元素下方，添加一個段落元素，文本爲 `Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.`。

在段落元素下方，添加另一個段落元素，文本爲 `His playful nature and boundless energy keeps me on my toes. I love him so much.`

# --hints--

第一個段落的文本應爲 `Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.` 仔細檢查你的拼寫。

```js
const pElement = document.querySelector('body p:first-of-type');
const text = "Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers."
assert.strictEqual(pElement?.innerText.trim(), text);
```

你的第一個段落應嵌套在 `section` 元素內。

```js
const pElement = document.querySelector('#about p:first-of-type');
assert.strictEqual(pElement?.parentElement.tagName, 'SECTION');
```

你應該在 `section` 元素中添加第二個段落。

```js
const pElement = document.querySelectorAll('#about p');
assert.strictEqual(pElement?.length, 2);
```

第二個段落的文本應爲 `His playful nature and boundless energy keeps me on my toes. I love him so much.` 仔細檢查你的拼寫。

```js
const pElement = document.querySelector('#about p:last-of-type');
const text = "His playful nature and boundless energy keeps me on my toes. I love him so much."
assert.strictEqual(pElement?.innerText.trim(), text);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
      --fcc-editable-region--

      --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
