---
id: 5f33294a6af5e9188dbdb8f3
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafundisha misingi ya CSS (Cascading Style Sheets) kwa kujenga menyu ya kafé.

Tuanze kwa kuongeza baadhi ya maudhui ya menyu. Ongeza kipengele cha `main` ndani ya kipengele kilichopo cha `body`. Hatimaye kitakuwa na taarifa za bei kuhusu kahawa na vitafunwa vinavyotolewa na kafé.

# --hints--

Msimbo wako unapaswa kuwa na lebo ya kufungua ya `<main>`.

```js
assert.match(code, /<main>/i);
```

Msimbo wako unapaswa kuwa na lebo ya kufunga ya `</main>`.

```js
assert.match(code, /<\/main>/i);
```

Haupasi kubadilisha kipengele chako cha `body`. Hakikisha hujafuta lebo yako ya kufunga kwa bahati mbaya.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Leboa ya `main` inapaswa kuwa ndani ya lebo ya `body`.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
