---
id: 5f332b23c2045fb843337579
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Ili kuwajulisha wageni kuwa mkahawa ulianzishwa mwaka `2020`, ongeza kipengele cha `p` chini ya kipengele cha `h1` chenye maandishi `Est. 2020`.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `<p>`.

```js
assert.match(code, /<p>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `</p>`.

```js
assert.match(code, /<\/p>/i);
```

Hupaswi kubadilisha kipengele chako cha `h1` kilichopo. Hakikisha hujafuta lebo ya kufunga.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

Kipengele chako cha `p` kinapaswa kuwa chini ya kipengele chako cha `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi `Est. 2020`.

```js
assert.equal(document.querySelector("p")?.innerText.trim(), "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
