---
id: 5f332a88dc25a0fd25c7687a
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Jina la mkahawa ni `CAMPER CAFE`. Kwa hiyo, ongeza kipengele cha `h1` ndani ya kipengele chako cha `main`. Mpe jina la mkahawa kwa herufi kubwa ili liweze kuonekana vizuri.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `<h1>`.

```js
assert.match(code, /<h1>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `</h1>`.

```js
assert.match(code, /<\/h1>/i);
```

Hupaswi kubadilisha kipengele chako kilichopo cha `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Kipengele chako cha `h1` kinapaswa kupangwa ndani ya kipengele chako cha `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi ya `CAMPER CAFE`.

```js
assert.equal(document.querySelector("h1")?.innerText.trim(), "CAMPER CAFE");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
