---
id: 5f344fad8bf01691e71a30eb
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Hadi sasa, umekuwa na udhibiti mdogo juu ya uwasilishaji na muonekano wa maudhui yako. Ili kubadilisha hilo, ongeza kipengele cha `style` ndani ya kipengele cha `head`.

# --hints--

Msimbo wako unapaswa kuwa na lebo ya kufungua ya `<style>`.

```js
assert.match(code, /<style\s*>/i);
```

Msimbo wako unapaswa kuwa na lebo ya kufunga ya `</style>`.

```js
assert.match(code, /<\/style\s*>/);
```

Kipengele chako cha `style` kinapaswa kupangwa ndani ya kipengele chako cha `head`.

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
