---
id: 5f344fc1520b6719f2e35605
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Kutakuwa na sehemu mbili kwenye menyu, moja kwa ajili ya kahawa na moja kwa ajili ya vitafunwa. Ongeza kipengele cha `section` ndani ya kipengele cha `main` ili uwe na mahali pa kuweka kahawa zote zinazopatikana.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `<section>`.

```js
assert.match(code, /<section\s*>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `</section>`.

```js
assert.match(code, /<\/section\s*>/i);
```

Hupaswi kubadilisha kipengele chako cha `main` kilichopo. Hakikisha hujafuta lebo ya kufunga.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Kipengele chako cha `section` kinapaswa kuwa ndani ya kipengele chako cha `main`.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert.strictEqual(section.parentElement, main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
