---
id: 5f356ed69db0a491745e2bb6
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

Ili kutumia mtindo wa darasa la vitu kwenye kipengele cha `div`, ondoa sifa ya kipengele ya `id` na ongeza sifa ya kipengele ya `class` kwenye lebo ya kufungua ya kipengele cha `div`. Hakikisha umeweka thamani ya darasa la vitu kuwa `menu`.

# --hints--

`div` yako inapaswa bado kuonyesha. Hakikisha hujapotosha lebo ya `<div>`.

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

Kipengele chako cha `div` kinapaswa kuwa na darasa la vitu la `menu`.

```js
assert.isTrue(document.querySelector('div').classList.contains('menu'));
```

Kipengele chako cha `div` hakipaswi tena kuwa na kitambulisho cha `menu`.

```js
assert.lengthOf(document.querySelectorAll('div#menu'), 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

