---
id: 5f356ed63e0fa262326eef05
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Sasa tumia kichaguzi cha `#menu` kilichopo kuweka rangi ya nyuma ya kipengele cha `div` kuwa `burlywood`.

# --hints--

Unapaswa kuweka sifa ya `background-color` kuwa `burlywood`.

```js
const hasBackgroundColor = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'burlywood');
assert.isTrue(hasBackgroundColor);
```

Kichaguzi chako cha `#menu` kinapaswa kuwa na rangi ya nyuma ya burlywood.

```js
const divBackground = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('background-color');
assert.equal(divBackground, 'burlywood');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

#menu {
  width: 300px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

