---
id: 5f35e5c4321f818cdc4bed30
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

Sasa mambo yanaonekana sawa, ni wakati wa kuanza kuongeza vitu vya menyu.

Ongeza kipengele tupu cha `article` chini ya kichwa cha sehemu cha `Coffee`. Kitakuwa na ladha na bei ya kila kahawa unayotoa kwa sasa.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `<article>`.

```js
assert.match(code,/<article>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `</article>`.

```js
assert.match(code, /<\/article>/i);
```

Haupasi kubadilisha kipengele kilichopo cha `h2`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 1);
```

Kipengele chako cha `article` kinapaswa kuja baada ya kipengele chako cha `h2`.

```js
const article = document.querySelector('article');
assert.equal(article?.previousElementSibling?.tagName, 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

