---
id: 5f356ed6199b0cdef1d2be8f
title: Hatua 22
challengeType: 0
dashedName: step-22
---

# --description--

Hadi sasa umekuwa ukitumia vichaguzi vya aina na kitambulisho (id) kupamba vipengele. Hata hivyo, ni kawaida zaidi kutumia kichaguzi tofauti kupamba vipengele vyako.

Umejifunza jinsi ya kufanya kazi na vichaguzi vya `class` katika mafundisho ya awali kama haya:

```css
.class-name {
  styles
}
```

Badilisha kichaguzi cha `#menu` kilichopo kuwa kichaguzi cha darasa la vitu kwa kubadilisha `#menu` na darasa la vitu linaloitwa `.menu`.

# --hints--

Unapaswa kuwa na kichaguzi cha darasa la vitu cha `.menu`.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert.exists(hasMenu);
```

Haupasi kuwa na kichaguzi cha `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert.notExists(hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
--fcc-editable-region--
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
