---
id: 5f3c866daec9a49519871816
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

Vipengele vya `article` kawaida huwa na vipengele vingi vinavyohusiana na taarifa. Katika kesi hii, kitakuwa na ladha ya kahawa na bei ya ladha hiyo.

Panga ndani vipengele viwili vya `p` ndani ya kipengele chako cha `article`. Kipengele cha kwanza kiwe na maandishi `French Vanilla`, na cha pili kiwe na maandishi `3.00`.

# --hints--

Haupaswi kubadilisha kipengele kilichopo cha `article`.

```js
assert.lengthOf(document.querySelectorAll('article'),1);
```

Kipengele chako cha `article` kinapaswa kuwa na vipengele viwili vya `p`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphChildren,2);
```

Kipengele chako cha kwanza cha `p` kinapaswa kuwa na maandishi `French Vanilla`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const firstP = paragraphChildren?.[0];
assert.match(firstP?.innerText.trim(), /French Vanilla/i);
```

Kipengele chako cha pili cha `p` kinapaswa kuwa na maandishi `3.00`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const secondP = paragraphChildren?.[1];
assert.match(secondP?.innerText.trim(), /3\.00/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
--fcc-editable-region--
            
--fcc-editable-region--
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
