---
id: 5f3c866dd0d0275f01d4d847
title: Step 35
challengeType: 0
dashedName: step-35
---

# --description--

Vizuri, hiyo haikufanya kazi. Kupamba vipengele vya `p` kama `inline-block` na kuviweka kwenye mistari tofauti huunda nafasi ya ziada upande wa kulia wa kipengele cha kwanza cha `p`, na kusababisha kipengele cha pili kusogea kwenye mstari unaofuata.

Njia moja ya kurekebisha hilo ni kufanya upana wa kila kipengele cha `p` kuwa kidogo chini ya `50%`. Kwa hivyo, badilisha thamani ya `width` kuwa `49%` kwa kila darasa la vitu ili kuona kinachotokea.

# --hints--

Unapaswa kuweka sifa ya `width` kuwa `49%` katika kichaguzi chako cha `.flavor`.

```js
const flavorWidth = new __helpers.CSSHelp(document).getStyle('.flavor')?.getPropertyValue('width');
assert.equal(flavorWidth, '49%');
```

Unapaswa kuweka sifa ya `width` kuwa `49%` katika kichaguzi chako cha `.price`.

```js
const priceWidth = new __helpers.CSSHelp(document).getStyle('.price')?.getPropertyValue('width');
assert.equal(priceWidth, '49%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

--fcc-editable-region--
.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
--fcc-editable-region--
```

