---
id: 6735a7370e0ae93a4577c689
title: تصميم تخطيط صحيفة
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

في هذا المختبر، ستصمم تخطيط صحيفة باستخدام CSS Grid، بما في ذلك مفاهيم مثل صفوف الشبكة وأعمدة الشبكة.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تحتوي صفحتك على عنصر `main` مع الصنف `newspaper-layout`.
1. يجب أن يتضمن `.newspaper-layout` عنصر `header` مع الصنف `title` يحتوي على `h1` لعرض اسم الصحيفة.
1. يجب أن يتضمن `.newspaper-layout` عنصر `article` مع الصنف `feature-article` للمقالة الإخبارية الرئيسية.
1. يجب أن يتضمن `.feature-article` عنصر `h2` لعنوان المقالة يليه عنصر `p` لمحتوى المقالة.
1. يجب إضافة ثلاثة عناصر `article` أخرى لمقالات أصغر، مع الأصناف `small-article1`، `small-article2`، و`small-article3`.
1. يجب أن يتضمن كل من المقالات الصغيرة عنصر `h3` لعنوان المقالة يليه عنصر `p` لمحتوى المقالة.
1. يجب أن يتضمن `.newspaper-layout` عنصر `article` مع الصنف `secondary-article` لقسم إخباري إضافي.
1. يجب أن يتضمن `.secondary-article` عنصر `h2` لعنوان المقالة يليه عنصر `p` لمحتوى المقالة.
1. يجب أن يتضمن `.newspaper-layout` عنصر `figure` مع الصنف `cover-image` لعرض صورة تمثل محتوى الصحيفة.
1. يجب أن تحتوي العناصر ذات الأصناف `title`، `feature-article`، `secondary-article`، `cover-image`، `small-article1`، `small-article2`، و`small-article3` على خاصية `grid-area` بقيمة اسم الصنف نفسه.
1. يجب أن يستخدم `.newspaper-layout` شبكة CSS Grid مع خاصية `grid-template-areas` لتعريف ترتيب الأقسام:
   - يجب أن يمتد `.title` عبر الصف العلوي.
   - يجب أن يمتد `.feature-article` عبر عمودين في الصف الثاني، مع وجود `.cover-image` في العمود الثالث.
   - يجب أن يمتد `.secondary-article` عبر عمودين في الصف الثالث، مع وجود `.cover-image` في العمود الثالث.
   - يجب أن تظهر المقالات الصغيرة الثلاث في الصف الرابع.
1. يجب أن يحتوي `.newspaper-layout` على خاصية `grid-template-columns` لتقسيم المساحة إلى ثلاثة أعمدة متساوية.
1. يجب تعيين خاصية `grid-template-rows` في `.newspaper-layout` إلى `auto` للصف الأول وتقسيم المساحة المتبقية إلى أجزاء متساوية للصفوف الأخرى.
1. يجب إضافة فجوة بين عناصر الشبكة.
1. تأكد من أن الصورة داخل `.cover-image` تناسب المساحة المخصصة لها عن طريق تعيين خاصية `max-width` إلى `100%`.

**ملاحظة:** تأكد من ربط ورقة الأنماط في ملف HTML الخاص بك وتطبيق CSS.

# --hints--

يجب أن يكون لديك عنصر `main`.

```js
assert.exists(document.querySelector('main'));
```

يجب أن يحتوي عنصر `main` على `class` باسم `newspaper-layout`.

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

يجب أن يحتوي عنصر `header` على `class` باسم `title` داخل عنصر `.newspaper-layout` الخاص بك.

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

يجب أن يحتوي عنصر `header` على عنصر `h1`.

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

يجب أن يحتوي عنصر `.title` على بعض النصوص.

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

يجب أن يكون لديك عنصر `article` لكل مقالة مع الأصناف `feature-article` و`small-article1` و`small-article2` و`small-article3` و`secondary-article` على التوالي.

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

يجب أن يتضمن `.feature-article` عنصر `h2` يليه عنصر `p`.

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

يجب أن تتضمن `.small-article1` و`.small-article2` و`.small-article3` عنصر `h3` يليه عنصر `p`.

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

يجب أن يتضمن `.secondary-article` عنصر `h2` يليه عنصر `p`.

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

يجب أن يكون لديك عنصر رسم بياني (`figure`) مع الصنف `cover-image` داخل `.newspaper-layout`.

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

يجب أن يحتوي `.cover-image` على عنصر `img`.

```js
assert.exists(document.querySelector(".cover-image img"));
```

يجب ضبط خاصية `display` لعنصر `.newspaper-layout` على `grid`.

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

يجب أن يحتوي عنصر `.newspaper-layout` على خاصية `grid-template-areas` بحيث يمتد `title` عبر الأعمدة الثلاثة في الصف الأول.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

يجب أن يحتوي الصف الثاني من قالب الشبكة على `feature-article` يمتد عبر عمودين و`cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

يجب أن يحتوي الصف الثالث من قالب الشبكة على `secondary-article` يمتد عبر عمودين و`cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

يجب أن يحتوي الصف الرابع من قالب الشبكة على `small-article1` و`small-article2` و`small-article3`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

يجب أن يستخدم محدد `.newspaper-layout` خاصية `grid-template-columns` لتقسيم المساحة إلى ثلاثة أجزاء متساوية. تذكر أنه يمكنك استخدام دالة `repeat` أو تحديد كل وحدة كسرية يدويًا.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

يجب أن يحتوي محدد `.newspaper-layout` على خاصية `grid-template-rows` مضبوطة على `auto 1fr 1fr 1fr`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

يجب إضافة فجوة بين عناصر الشبكة.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

يجب أن يكون لـ `.title` خاصية `grid-area` بقيمة `title`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

يجب أن يكون لـ `.feature-article` خاصية `grid-area` بقيمة `feature-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

يجب أن يكون لـ `.cover-image` خاصية `grid-area` بقيمة `cover-image`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

يجب أن يكون لـ `.secondary-article` خاصية `grid-area` بقيمة `secondary-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

يجب أن يكون لـ `.small-article1` خاصية `grid-area` بقيمة `small-article1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

يجب أن يكون لـ `.small-article2` خاصية `grid-area` بقيمة `small-article2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

يجب أن يكون لـ `.small-article3` خاصية `grid-area` بقيمة `small-article3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

يجب أن يكون لـ `img` داخل `.cover-image` خاصية `max-width` بقيمة `100%`.

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
