---
id: 685976a92ffab8b2dc744b72
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `.tabs` 元素内，添加一个 `div` 元素，其 `role` 属性设置为 `"tablist"`，`aria-labelledby` 属性设置为 `"tabs-title"`。  

这里需要 `aria-labelledby` 属性来将标签（页）列表与可见标签关联，帮助像屏幕阅读器这样的辅助技术向用户宣布标签（页）列表的用途。

# --hints--

你的 `.tabs` 元素内应包含一个 `div` 元素。

```js
assert.exists(document.querySelector(".tabs div"));
```

你的 `div` 元素应具有设置为 `"tablist"` 的 `role` 属性。

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

你的 `div` 元素应具有设置为 `"tabs-title"` 的 `aria-labelledby` 属性。

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
