---
id: 68599e9259c0eec3e78c52fd
title: 第 6 步
challengeType: 0
dashedName: step-6
---

# --description--

默认情况下，其他标签（页）不应被选中。你可以通过为第二个和第三个 `button` 元素添加 `aria-selected` 属性并将其值设置为 `"false"` 来实现。

# --hints--

你的第二个 `button` 元素应具有 `aria-selected` 属性。

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[1].hasAttribute("aria-selected")
);
```

你的第二个 `button` 元素应具有设置为 `"false"` 的 `aria-selected` 属性。

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-selected"), "false");
``` 

你的第三个 `button` 元素应具有 `aria-selected` 属性。

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[2].hasAttribute("aria-selected")
);
```

你的第三个 `button` 元素应具有设置为 `"false"` 的 `aria-selected` 属性。

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-selected"), "false");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true">🌍 Earth</button>
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
