---
id: 68599633407374c0d574587b
title: 第 3 步
challengeType: 0
dashedName: step-3
---

# --description--

下一步是添加负责在不同星球及其事实之间切换的按钮。

在 `tablist` 元素内，添加三个 `button` 元素，每个元素的 `role` 属性设置为 `"tab"`。第一个 `button` 元素的文本应为 `🌍 Earth`。第二个 `button` 的文本应为 `🪐 Saturn`。第三个 `button` 元素的文本应为 `🔴 Mars`。

# --hints--

你应该在 `tablist` 元素内有三个 `button` 元素。

```js
const btns = document.querySelectorAll("div[role='tablist'] button");
assert.lengthOf(btns, 3);
```

你的每个 `button` 元素都应具有设置为 `"tab"` 的 `role` 属性。

```js
const btns = document.querySelectorAll("div[role='tablist'] button[role='tab']");
assert.lengthOf(btns, 3);
```

你的第一个 `button` 元素应包含文本 `🌍 Earth`。

```js
const firstBtn = document.querySelector("[role='tab']");
assert.equal(firstBtn.textContent, "🌍 Earth");
```

你的第二个 `button` 元素应包含文本 `🪐 Saturn`。

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.textContent, "🪐 Saturn")
```

你的第三个 `button` 元素应该有 `🔴 Mars` 文本。

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2]
assert.equal(thirdBtn.textContent, "🔴 Mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
