---
id: 685db0bd1280154ae85bcf9f
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将使用超文本标记语言、CSS 和 JavaScript 构建一个可访问的行星标签（页）列表。已为你提供了超文本标记语言公式化和 CSS。

在这最初的几个步骤中，你将构建标签（页）列表的剩余超文本标记语言。

首先创建一个 `div` 元素，其类设置为 `"tabs"`。在该 `div` 元素内部，添加一个 `h2` 元素，其 `id` 属性设置为 `"tabs-title"`，文本内容为 `Planets`。

# --hints--

你应该有一个 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

你的 `div` 元素应设置类为 `"tabs"`。

```js
assert.exists(document.querySelector("div.tabs"));
```

你的 `div` 元素内应包含一个 `h2` 元素。

```js
assert.exists(document.querySelector(".tabs h2"));
```

你的 `h2` 元素应具有设置为 `"tabs-title"` 的 `id` 属性。

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

你的 `h2` 元素应具有文本内容 `Planets`。

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
