---
id: 69f0a37c3bf194a7c9c62492
title: 장비 대출 원장 만들기
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

회사 서비스 데스크에서 제공하는 IT 하드웨어 장비를 관리하는 원장 애플리케이션을 만드는 소프트웨어 개발자 역할입니다. 원장은 자산 태그를 키로 하고 값은 객체인 객체입니다:

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 각 장비가 자산 태그를 키로 하고 `type`, `status`, `borrower`(`name`과 `email` 포함), `dueDate` 프로퍼티를 가진 객체로 원장을 모델링하세요.

1. 원장 객체를 받아 원본을 변경하지 않고 복제하는 `checkoutDevice(ledger, assetTag, borrower)` 함수를 구현하세요. 복제본에서 `borrower` 객체의 `name`과 `email`을 설정하고, `status`를 `"CheckedOut"`으로 업데이트하며, `{ ledger: updatedLedger, message: confirmationString }` 형태의 객체를 반환하세요. 여기서 `message`는 자산 태그와 `borrower` 객체의 `name`을 포함해야 합니다.

   - `assetTag`가 원장에 없으면, 함수는 원장을 변경하지 않고 자산 태그를 포함하며 찾을 수 없음을 설명하는 `message`와 함께 반환해야 합니다.
   - 장비가 이미 `"CheckedOut"` 상태이면, 함수는 원장을 변경하지 않고 현재 대출자를 유지하며 자산 태그를 포함하고 장비가 이미 대출 중임을 설명하는 `message`와 함께 반환해야 합니다.

1. 원장 객체를 받아 원본을 변경하지 않고 복제하는 `checkinDevice(ledger, assetTag)` 함수를 구현하세요. 복제본에서 `borrower` 객체의 `name`과 `email`을 지우고, `dueDate`를 `""`로 초기화하며, `status`를 `"CheckedIn"`으로 업데이트하고, `{ ledger: updatedLedger, message: confirmationString }` 형태의 객체를 반환하세요. 여기서 `message`는 자산 태그를 포함해야 합니다.

   - `assetTag`가 원장에 없으면, 함수는 원장을 변경하지 않고 자산 태그를 포함하며 찾을 수 없음을 설명하는 `message`와 함께 반환해야 합니다.

1. 원장 객체를 받아 연체된 장비 객체들의 배열을 반환하는 `listOverdueDevices(ledger, today)` 함수를 구현하세요. 반환 배열은 `dueDate` 기준 오름차순으로 정렬되어야 합니다. 장비가 연체된 조건은 `status`가 `"CheckedOut"`이고 `dueDate`가 `today`보다 엄격히 이전일 때입니다. `today`와 `dueDate`는 월/일/연도 순서의 문자열이며, 월과 일이 0으로 채워져 있지 않아도 됩니다(예: `"9/5/2025"`와 `"09/05/2025"` 모두 유효합니다).

1. 코드 어디에서도 JavaScript `Date` 객체를 사용하지 마세요.

1. 원장 객체를 JSON 문자열로 변환하는 `serializeLedger(ledger)` 함수와 JSON 문자열을 다시 JavaScript 객체로 변환하는 `loadLedger(json)` 함수를 구현하세요.

# --hints--

`checkoutDevice`라는 이름의 함수를 만드세요.

```js
assert.isFunction(checkoutDevice);
```

`checkinDevice`라는 이름의 함수를 만드세요.

```js
assert.isFunction(checkinDevice);
```

`listOverdueDevices`라는 이름의 함수를 만드세요.

```js
assert.isFunction(listOverdueDevices);
```

`serializeLedger`라는 이름의 함수를 만드세요.

```js
assert.isFunction(serializeLedger);
```

`loadLedger`라는 이름의 함수를 만드세요.

```js
assert.isFunction(loadLedger);
```

`checkoutDevice` 함수는 원본 원장을 변경하지 않아야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

`checkoutDevice` 함수는 대출자의 `name`과 `email`을 설정하고 `status`를 `"CheckedOut"`으로 설정해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

`checkoutDevice` 함수는 자산 태그와 대출자의 `name`을 포함하는 확인 메시지를 반환해야 합니다.

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

`checkoutDevice` 함수는 반환된 원장에서 관련 없는 장비와 변경되지 않은 필드를 유지해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

`checkoutDevice` 함수는 자산 태그를 찾지 못했을 때 원장을 변경하지 않고 자산 태그를 포함하는 메시지를 반환해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

`checkoutDevice` 함수는 장비가 이미 대출 중일 때 원장을 변경하지 않고 현재 대출자를 유지하며 자산 태그를 포함하는 메시지를 반환해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

`checkinDevice` 함수는 원본 원장을 변경하지 않아야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

`checkinDevice` 함수는 대출자의 `name`과 `email`을 지우고 `dueDate`를 `""`로 초기화하며 `status`를 `"CheckedIn"`으로 설정해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

`checkinDevice` 함수는 자산 태그를 포함하는 확인 메시지를 반환해야 합니다.

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

`checkinDevice` 함수는 반환된 원장에서 관련 없는 장비와 변경되지 않은 필드를 유지해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

`checkinDevice` 함수는 자산 태그를 찾지 못했을 때 원장을 변경하지 않고 자산 태그를 포함하는 메시지를 반환해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

`listOverdueDevices` 함수는 배열을 반환해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

`listOverdueDevices` 함수는 `status`가 `"CheckedOut"`인 연체 장비만 반환해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

`listOverdueDevices` 함수는 연체 장비를 `dueDate` 기준 오름차순으로 정렬해 반환해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

`listOverdueDevices` 함수는 월이나 일이 0으로 채워져 있지 않아도 올바르게 정렬해야 합니다.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

코드 어디에서도 JavaScript `Date` 객체를 사용하지 마세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

`serializeLedger` 함수는 전달받은 원장을 나타내는 JSON 문자열을 반환해야 합니다.

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

`loadLedger` 함수는 JSON 문자열이 나타내는 JavaScript 객체를 반환해야 합니다.

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
