---
id: 6989f2e35e6941e13e445156
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

現在你將實作主要的 BFS 迴圈。創建一個 `while` 迴圈，只要 `queue` 不為空，也就是 `queue.length` 大於 `0`，就持續執行。

在迴圈內，將 `queue` 輸出到主控台。

# --hints--

你應該創建一個條件為 `while` 的 `queue.length > 0` 迴圈。

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

你應該在 `queue` 迴圈內將 `while` 輸出到主控台。

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--
  
--fcc-editable-region--
    
  return result;
}
```
