---
id: 6989f711e57008f3a23482f2
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

你的函式已經完成。請再次測試它，這次記得輸出 `genParentheses(2)` 而不是 `genParentheses(1)`。

# --hints--

你應該輸出 `genParentheses(2)`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*genParentheses\s*\(\s*2\s*\)\s*\)/);
```

你的程式碼中不應該有 `console.log(genParentheses(1))`。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\s*\(\s*genParentheses\s*\(\s*1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
      if (closesUsed < opensUsed) {
        queue.push([current + ')', opensUsed, closesUsed + 1]);
      }
    }
  }
  return result;
}

--fcc-editable-region--
console.log(genParentheses(1));
--fcc-editable-region--
```
