---
id: 6989eb65324d97775c4965fe
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將實作一個函式（程式），使用廣度優先搜尋（BFS）方法來產生所有有效的小括號組合。例如，兩對小括號的有效組合是 `(())` 和 `()()`。

首先建立一個名為 `genParentheses` 的函式，並帶有一個參數 `pairs`。目前先從函式傳回一個空陣列。

# --hints--

你應該定義一個名為 `genParentheses` 的函式。

```js
assert.isFunction(genParentheses);
```

你的 `genParentheses` 函式應該有一個名為 `pairs` 的參數。

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

你的函式應該傳回一個空陣列。

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
