---
id: 6989f1ba358acba651d5ee6b
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

現在你要設定用來儲存結果的資料結構。創建一個名為 `result` 的變數，並將它初始化為空陣列。這個陣列將用來儲存你產生的所有有效小括號組合。

更新你的傳回述語，改為傳回 `result`，而不是空陣列。

# --hints--

你應該在你的 `result` 函式中宣告一個名為 `genParentheses` 的變數。

```js
assert.match(genParentheses.toString(), /result/);
```

你應該將 `result` 初始化為空陣列。

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

你應該在函式結尾傳回 `result`。

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
