---
id: 6989eb65324d97775c4965fe
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將實作一個函式，使用廣度優先搜尋（BFS）方法來產生所有有效的小括弧組合。例如，兩對小括弧的有效組合是 `(())` 和 `()()`。

首先建立一個名為 `genParentheses` 的函式，並帶有一個參數 `pairs`。目前先從函式傳回一個空陣列。

# --hints--

你應該定義一個名為 `genParentheses` 的函式。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.genParentheses);
```

你的 `genParentheses` 函式應該有一個名為 `pairs` 的參數。

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.equal(genParentheses?.parameters.length, 1);
assert.equal(genParentheses?.parameters[0].toString(), 'pairs');
```

你的函式應該傳回一個空陣列。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.genParentheses);
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
