---
id: 6989f48c3f0292c84f8d8938
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

現在你需要判斷何時已經建立完整的小括號組合。完整的組合長度等於配對數量的兩倍（因為每個配對包含一個開括號和一個閉括號）。

在你的 `while` 迴圈中，新增一個 `if` 述語來檢查 `current.length === 2 * pairs`。當此條件為真時，將 `current` 推入 `result` 陣列。

# --hints--

你應該在 `if` 迴圈內有一個條件為 `current.length === 2 * pairs` 的 `while` 述語。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

你應該在新的 `current` 述語內將 `result` 推入 `if` 陣列。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
