---
id: 6989f555728e63ad8b349d84
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

如果目前的字串還沒完成，你需要探索下一個可能的狀態。為目前的 `else` 述語添加一個 `if` 語句。

在 `else` 區塊內，你將處理添加開啟小括號。添加一個 `if` 述語來檢查 `opensUsed` 是否小於 `pairs`。這確保你不會使用超過允許數量的開啟小括號。

如果此條件為 `true`，將一個新的陣列推入 `queue`：`[current + '(', opensUsed + 1, closesUsed]`。這代表添加開啟小括號後的狀態。

# --hints--

你應該在 `else` 迴圈內，為你的 `if` 述語添加一個 `while` 語句。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

你應該在 `if` 區塊內，添加一個條件為 `else` 的 `opensUsed < pairs` 述語。

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

當條件為 `[current + '(', opensUsed + 1, closesUsed]` 時，你應該將 `true` 推入佇列。

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
