---
id: 6989f207306332c06b0add39
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

對於廣度優先搜尋方法，你將使用佇列來追蹤在建置小括號組合時的不同狀態。每個狀態將以一個包含三個元素的陣列來表述：

- 目前正在建立的字串
- 目前使用的開啟小括號數量
- 目前使用的結束小括弧數量

創建一個名為 `queue` 的變數，並將其初始化為一個包含一個陣列 `['', 0, 0]` 的陣列。這代表起始狀態，字串為空且尚未使用任何小括號。

# --hints--

你應該在你的 `queue` 函式中宣告一個名為 `genParentheses` 的變數。

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.queue);
```

你應該將 `queue` 初始化為一個包含陣列 `['', 0, 0]` 的陣列。

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { queue } = genParentheses?.variables ?? {};
assert.isFalse(queue?.value.isEmpty());
const queueValue = eval(queue?.value.toString());
assert.deepEqual(queueValue, [['', 0, 0]]);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  
  return result;
}
```
