---
id: 6989f788362592a7238b1f94
title: 第 13 步
challengeType: 1
dashedName: step-13
---

# --description--

現在你不需要再記錄佇列了。所以請從你的 `console.log(queue)` 迴圈中移除 `while`。

# --hints--

你的 `console.log(queue)` 迴圈中不應該有 `while`。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
--fcc-editable-region--
    console.log(queue);
--fcc-editable-region--

    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
      if (closesUsed < opensUsed) {
        queue.push([current + ')', opensUsed, closesUsed + 1]);
      }
    }
  }
  return result;
}

console.log(genParentheses(2));
```
