---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

現在你需要處理添加閉合小括號。關鍵規則是你只能在維持平衡的情況下添加閉合小括號，也就是說，使用的開啟小括號數量必須多於閉合小括號。

仍在 `else` 區塊內，添加第二個 `if` 述語來檢查 `closesUsed` 是否小於 `opensUsed`。

如果此條件為真，將另一個新的陣列推入 `queue`：`[current + ')', opensUsed, closesUsed + 1]`。這代表添加閉合小括號後的狀態。

# --hints--

你應該在 `if` 區塊內有第二個 `closesUsed < opensUsed` 述語，條件是 `else`。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

當條件為真時，你應該將 `[current + ')', opensUsed, closesUsed + 1]` 推入佇列。

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
