---
id: 6989f3632f8b8bdf168aa321
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

在你的 `while` 迴圈中，使用 `queue.shift()` 來移除並取得佇列中的第一個元素。這樣可以實作出 BFS 特有的先進先出（FIFO）行為。

將這個陣列解構賦值成三個變數：`current`、`opensUsed` 和 `closesUsed`。

# --hints--

你應該使用 `queue.shift()` 來取得佇列中的第一個元素。

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

你應該將結果解構賦值成三個變數：`current`、`opensUsed` 和 `closesUsed`。

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return result;
}
```
