---
id: 6989ef885db3a3488a13b426
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

在實作核心演算法之前，你需要驗證輸入。`pairs` 參數應該是整數，因為它代表要產生的小括弧配對數量。

在你的函式開頭添加一個 `if` 述語，用來檢查 `pairs` 是否不是整數。為此請使用 `Number.isInteger()` 方法。

如果條件為真，傳回字串 `The number of pairs should be an integer`。

# --hints--

你應該在 `!Number.isInteger(pairs)` 函式內有一個 if 述語，條件是 `genParentheses`。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

當輸入不是整數時，你應該傳回一個字串。

```js
let result = genParentheses("i");
assert.isString(result);
```

當輸入是字串時，你應該傳回字串 `The number of pairs should be an integer`。

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
