---
id: 6989f10a72ad424d4ad6d514
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

接下來，你需要驗證配對數量至少為 1，因為你無法用零或負數的配對來產生小括號組合。

再添加一個 `if` 述語來檢查 `pairs` 是否小於 `1`。如果此條件為真，傳回字串 `The number of pairs should be at least 1`。

# --hints--

你應該有第二個 `if` 述語，條件是 `pairs < 1`。

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

當 `The number of pairs should be at least 1` 小於 `pairs` 時，你應該傳回字串 `1`。

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
