---
id: 66bba6fff611169359d9d36a
title: Costruisci una pagina email confidenziale
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `main` con un `id` di `email`.
1. Il tuo elemento `#email` dovrebbe avere `padding` di `50px`, un margine superiore di `50px`, un `width` di `500px` e un `border` che è `2px`.
1. La larghezza totale del tuo elemento `#email`, inclusi spaziature interne e bordi, dovrebbe essere `500px`.
1. Dovresti avere due elementi `div`, uno con un id di `confidential` e l'altro con un id di `top-secret`, all'interno del tuo elemento `#email`.
1. I tuoi elementi `#confidential` e `#top-secret` dovrebbero avere un `display` di `inline-block`.
1. I tuoi elementi `#confidential` e `#top-secret` dovrebbero avere un `padding`, un margine sinistro e un `border`.
1. L'elemento `#confidential` dovrebbe avere il testo `CONFIDENTIAL`.
1. L'elemento `#top-secret` dovrebbe avere il testo `TOP SECRET`.
1. I tuoi elementi `#confidential` e `#top-secret` dovrebbero essere ruotati usando una trasformazione CSS.
1. Dovresti avere almeno tre elementi paragrafo all'interno del tuo elemento `#email`.
1. Dovresti avere almeno tre elementi `span` con una classe `blurred`, all'interno dei tuoi paragrafi.
1. Dovresti avere un selettore `blurred` che sfoca l'elemento `3px` usando un filtro CSS.

# --hints--

Dovresti avere un elemento `main` con un `id` di `email`.

```js
assert.exists(document.querySelector('main#email'));
```

Dovresti avere un selettore `#email` che imposta la `padding` dei suoi elementi su tutti i lati a `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

Dovresti avere un selettore `#email` che imposta la `margin-top` dei suoi elementi a `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

Dovresti avere un selettore `#email` che imposta la `width` dei suoi elementi a `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

Il tuo elemento `#email` dovrebbe avere un bordo `2px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

Il tuo elemento `#email` dovrebbe avere un `box-sizing` di `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

Dovresti avere un elemento `div` con un `id` di `confidential` all'interno del tuo elemento `#email`.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

Dovresti avere un elemento `div` con un `id` di `top-secret` all'interno del tuo elemento `#email`.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

Il tuo elemento `#confidential` dovrebbe avere la sua `display` impostata a `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

Il tuo elemento `#top-secret` dovrebbe avere la sua `display` impostata a `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

Il tuo elemento `#confidential` dovrebbe avere un `padding` su tutti i lati.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Il tuo elemento `#top-secret` dovrebbe avere un `padding` su tutti i lati.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Il tuo elemento `#confidential` dovrebbe avere un margine sinistro.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

Il tuo elemento `#top-secret` dovrebbe avere un margine sinistro.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

Il tuo elemento `#confidential` dovrebbe avere un `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

Il tuo elemento `#top-secret` dovrebbe avere un `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

Il testo del tuo elemento `#confidential` dovrebbe essere `CONFIDENTIAL`. Hai omesso il testo o c'è un errore di battitura.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

Il testo del tuo elemento `#top-secret` dovrebbe essere `TOP SECRET`. Hai omesso il testo o c'è un errore di battitura.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

Il tuo elemento `#confidential` dovrebbe usare una trasformazione CSS `transform` per ruotare l'elemento.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

Il tuo elemento `#top-secret` dovrebbe usare una trasformazione CSS `transform` per ruotare l'elemento.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

Dovresti avere almeno tre elementi paragrafo all'interno del tuo elemento `#email`.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

Dovresti avere almeno tre elementi `span` con una classe `blurred` all'interno dei tuoi paragrafi.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

I tuoi elementi `.blurred` non dovrebbero essere vuoti.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

Dovresti avere un selettore `.blurred` che imposta la `filter` dei suoi elementi a `blur(3px)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
