---
id: 699f38937b053e4679878ff1
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Працюючи з чистим JavaScript, нічого не заважає вам викликати функцію `square` з рядком або масивом, що призведе до несподіваних результатів.

Щоб це показати, змініть виклик функції так, щоб використовувати рядок `"something"` замість числа `5`.

У наступному уроці ви навчитеся використовувати TypeScript, щоб додати типову безпеку у свій код і запобігти такого роду помилкам.

# --hints--

Ваш виклик функції `square` має мати аргумент `"something"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
