---
id: 699f458e47fb88c756fdba81
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Наразі параметр `numbers` має неявний тип `any`. Щоб це виправити, ви можете явно вказати тип параметра як масив чисел.

Нагадаємо з попередніх уроків, що масив чисел можна типізувати так:

```ts
const numbers: number[] = [1, 2, 3];
```

Оновіть вашу функцію `getAverage`, щоб явно типізувати параметр `numbers` як масив чисел.

# --hints--

Ваша функція `getAverage` має мати параметр з іменем `numbers`, який явно типізований як масив чисел.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
