---
id: 699f4ec939f10da62fb56c3e
title: Крок 15
challengeType: 1
dashedName: step-15
---

# --description--

Тепер настав час перевірити вашу функцію.

Створіть змінну з назвою `raisedResult` і призначте їй результат виклику `raiseTo` з аргументами `2` та `3`.

Потім виведіть `raisedResult` у консоль.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні мати змінну з назвою `raisedResult`.

```js
assert.isDefined(raisedResult);
```

Ваша змінна `raisedResult` має бути числом.

```js
assert.isNumber(raisedResult);
```

Ви повинні вивести `raisedResult` у консоль.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
