---
id: 699f3738921fb79a651ded56
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Тепер настав час викликати вашу функцію.

Створіть змінну з назвою `result` і призначте їй виклик функції `square` з аргументом `5`.

Нижче виведіть у консоль значення `result`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні мати змінну з назвою `result`.

```js
assert.isDefined(result);
```

Ваша змінна `result` має бути числом.

```js
assert.isNumber(result);
```

Ваша змінна `result` має бути квадратом `5`.

```js
assert.strictEqual(result, 25);
```

Ви повинні вивести у консоль змінну `result`.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
