---
id: 699f52220de88c8091032076
title: Крок 18
challengeType: 1
dashedName: step-18
---

# --description--

Якщо відкрити консоль, результат буде `NaN`, тому що другий аргумент для `exponent` ніколи не був наданий. Для цієї `raiseTo` функції потрібно зробити кілька оновлень, які це виправлять.

Першою зміною буде повернути `base ** 2`.

У наступному кроці ви побачите, як обробляти необов’язковий параметр `exponent`.

# --hints--

Ваша `raiseTo` функція має повертати число.

```js
assert.isNumber(raiseTo(4));
```

Ваша `raiseTo` функція має повертати значення `base`, піднесене до степеня `2`.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
