---
id: 699f3f636aa840b7201aa37c
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Якщо відкрити консоль, ви побачите таке повідомлення про помилку TypeScript:

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

Тепер, коли параметр явно типізований як число, TypeScript може виявити помилку, коли намагаються передати рядковий аргумент у функцію `square`.

Змініть аргумент, переданий у функцію `square`, на число `5`, щоб виправити помилку і побачити правильний вихід у консолі.

# --hints--

Ваш виклик функції `square` має містити аргумент `5` замість `"something"`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
