---
id: 699f52fac083214b6bf51de6
title: Крок 19
challengeType: 1
dashedName: step-19
---

# --description--

У цьому останньому кроці вам потрібно додати умовну перевірку для параметра `exponent` у функції `raiseTo`.

Якщо аргумент `exponent` надано, функція має повертати `base`, піднесене до степеня `exponent`. Якщо аргумент `exponent` не надано, вона має повертати `base ** 2`.

З цим останнім кроком ваш типобезпечний набір математичних інструментів буде завершено!

# --hints--

Ваша функція `raiseTo` має повертати `base ** exponent`, коли аргумент `exponent` надано.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

Ваша функція `raiseTo` має повертати `base ** 2`, коли аргумент `exponent` не надано.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
