---
id: 699f513e01eb1768cc5df726
title: Крок 17
challengeType: 1
dashedName: step-17
---

# --description--

Тепер, коли параметр `exponent` позначено як необов’язковий, ви можете викликати функцію з одним аргументом.

Оновіть свій виклик функції `raiseTo`, щоб передавати лише число `3`.

# --hints--

Вашу функцію `raiseTo` слід викликати з одним аргументом `3`.

```js
assert.match(
  code,
  /raiseTo\s*\(\s*3\s*\)/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}

--fcc-editable-region--
const raisedResult = raiseTo(2, 3);
--fcc-editable-region--
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
