---
id: 699f41b8da20f816b5c7dece
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Якщо ви відкриєте консоль, то побачите таке повідомлення про помилку:

```md
Expected 1 arguments, but got 2.
```

Щоб виправити помилку, оновіть виклик вашої функції `square`, щоб вона приймала лише один аргумент `5`.

# --hints--

Ваш виклик функції `square` повинен мати лише один аргумент `5`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
