---
id: 699f5063d2393c674b1f45dc
title: Крок 16
challengeType: 1
dashedName: step-16
---

# --description--

Як ви пам’ятаєте з попереднього кроку, якщо викликати функцію з неправильним числом аргументів, TypeScript видасть помилку.

Існує спосіб позначити параметри як необов’язкові, додавши `?` після імені параметра. Це дозволяє викликати функцію з меншою кількістю аргументів, ніж параметрів.

Ось приклад:

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

У цьому прикладі, якщо параметр `greeting` не передається під час виклику функції `greet`, він буде `undefined`, і функція поверне привітання за замовчуванням.

Оновіть свою функцію `raiseTo`, щоб зробити параметр `exponent` необов’язковим.

# --hints--

Ваша функція `raiseTo` має мати параметр `exponent`, позначений як необов’язковий.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
