---
id: 699f412114c87221f16a99e3
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Ще одна перевага використання TypeScript у тому, що він виявить помилки, якщо ви передасте забагато або замало аргументів у виклик функції.

Щоб це показати, оновіть свій виклик `square`, щоб він приймав два аргументи `5` і `10`.

# --hints--

Ваш виклик функції `square` має містити два аргументи `5` і `10`.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
