---
id: 699f4b903bf5427a599f791d
title: Крок 14
challengeType: 1
dashedName: step-14
---

# --description--

У останніх кількох кроках ви створите ще одну функцію, щоб продовжити практикувати анотації типів у функціях.

Почніть зі створення функції з назвою `raiseTo`, яка приймає два параметри `base` і `exponent`. Обидва параметри мають бути типу `number`. Функція має повертати результат піднесення `base` до степеня `exponent`. Тип повернення функції також має бути анотований як `number`.

# --hints--

Ви повинні мати функцію `raiseTo`.

```js
assert.isFunction(raiseTo);
```

Ваша функція `raiseTo` має мати два параметри, `base` і `exponent`.

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*=>/
);
```

Ваші параметри `base` і `exponent` мають бути обидва типу `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

Ваша функція `raiseTo` має повертати число.

```js
assert.isNumber(raiseTo(2, 3));
```

Ваша функція `raiseTo` має повертати результат піднесення `base` до степеня `exponent`.

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

Ваша функція `raiseTo` має мати анотацію типу повернення `number`.

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
