---
id: 699f1d72e906e4616155290d
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви продовжите вивчати анотації типів і те, як вони працюють з функціями, створюючи безпечний за типами набір математичних інструментів.

Почніть зі створення функції з назвою `square`, яка приймає параметр і повертає квадрат числа.

# --hints--

У вас має бути функція з назвою `square`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

Ваша функція `square` має мати один параметр.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

Ваша функція `square` має повертати квадрат числа.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
