---
id: 699f4a3e7b4c6bfa538cc4f7
title: Крок 13
challengeType: 1
dashedName: step-13
---

# --description--

Якщо ви відкриєте консоль, то побачите таке повідомлення про помилку:

```md
Type 'string' is not assignable to type 'number'.
```

Щоб виправити цю помилку, потрібно змінити повернення на `sum / numbers.length`.

# --hints--

Ваша функція `getAverage` має повертати число.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Ваша функція `getAverage` має повертати середнє значення чисел у масиві.

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```
