---
id: 699f4861dd5b5f164e1bda4a
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Ще одна перевага використання TypeScript — це гарантія того, що типи повернення дотримуються і відповідають вашим очікуванням. Іноді TypeScript може вивести тип повернення функції, але ви також можете явно його вказати.

Щоб проілюструвати це, змініть тип повернення функції `getAverage` на рядок `"something"`.

# --hints--

Ваша функція `getAverage` має повертати рядок.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

Ваша функція `getAverage` має повертати рядок `"something"`.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
