---
id: 699f434e294a9e11cf4f65f0
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Тепер настав час додати ще одного помічника з математики до вашого проєкту toolkit.

Почніть зі створення функції з назвою `getAverage`, яка приймає один параметр. Ваша функція має повертати середнє значення для масиву чисел.

У наступних кількох кроках ви додасте анотації типів для цієї функції.

# --hints--

У вас має бути функція з назвою `getAverage`.

```js
assert.isFunction(getAverage);
```

Ваша функція має приймати один параметр.

```js
assert.strictEqual(getAverage.length, 1);
```

Ваша функція має повертати число.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Ваша функція має повертати правильний результат для будь-якого масиву чисел.

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
