---
id: 699f4639ae0016caaf81b7c6
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Тепер настав час перевірити вашу функцію.

Почніть зі створення змінної з назвою `avgResult` і призначте їй виклик функції `getAverage`. Аргументом для цього виклику має бути `[2, 14, 26, 8]`.

Потім виведіть значення `avgResult` у консоль.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

У вас має бути змінна з назвою `avgResult`.

```js
assert.isDefined(avgResult);
```

Ви повинні призначити `getAverage([2, 14, 26, 8])` вашій змінній `avgResult`.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

Ваша змінна `avgResult` має бути числом.

```js
assert.isNumber(avgResult);
```

Ви повинні вивести `avgResult` у консоль.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
