---
id: 699f49669ea51a31d2ba89d8
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Як ви пам’ятаєте з попередніх уроків, ви можете явно вказати тип повернення для функції ось так:

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

У цьому прикладі тип повернення функції `circleArea` — це `number`. Якщо ви спробуєте повернути інший тип, TypeScript покаже помилку.

Оновіть свою функцію `getAverage`, щоб явно вказати тип повернення як `number`.

# --hints--

Ваша функція `getAverage` має мати тип повернення `number`.

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
