---
id: 699f3cb8d057f29954205d7d
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Як ви пам’ятаєте з попередніх уроків, ви можете додати анотації типів до параметрів функції ось так:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

У цьому прикладі TypeScript покаже помилку, якщо ви спробуєте викликати функцію `add` з аргументами, які не є числами.

Оновіть свою функцію `square`, щоб додати анотацію типу для параметра `num`. Тип має бути `number`.

# --hints--

Ви повинні додати анотацію типу до параметра `num` у функції `square`. Тип має бути `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
