---
id: 67f8ef57b9974052e5ef6b30
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Dentro del elemento `#results`, utiliza un operador ternario para verificar si la longitud del array `results` es mayor que cero. Si es así, recorre los elementos y muestra cada uno en un elemento `p` con la clase `result-item`. Asegúrate de incluir un atributo `key` con cada párrafo también.

Si `results` está vacío, renderiza un elemento `p` con el mensaje `No results found`.

# --hints--

El contenedor de `#results` debe contener una lista de las frutas coincidentes basadas en la consulta. Cada resultado debe ser un elemento `p` con la clase `result-item` y debe mostrar el nombre de la fruta.

```js
async () => {
  {
    const mockResults = ["Item 1", "Item 2"];
    React.useState = (arg) => {
      if (Array.isArray(arg)) return [mockResults, () => {}]; 
      return [null, () => {}];
    };
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const div = await __helpers.prepTestComponent(exports.FruitsSearch);
    const items = div.querySelectorAll("p.result-item");

    assert.equal(items.length, mockResults.length);
    assert.equal(items[0].textContent, "Item 1");
    assert.equal(items[1].textContent, "Item 2");

    assert.isTrue(items[0].classList.contains("result-item"));
    assert.isTrue(items[1].classList.contains("result-item"));
  }
}
```

El elemento `#results` debe contener un elemento `p` con el texto `No results found` cuando el array `results` está vacío.

```js
async () => {
  {
    React.useState = (arg) => {
      if (Array.isArray(arg)) return [[], () => {}];
      return [null, () => {}];
    };
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const div = await __helpers.prepTestComponent(exports.FruitsSearch);
    const noResultsEl = div.querySelector("#results p");
    assert.exists(noResultsEl);
    assert.equal(noResultsEl.textContent, "No results found");
  }
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Fruits Search</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { FruitsSearch } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<FruitsSearch />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#search-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#search-input {
  padding: 10px;
  width: 80%;
  border: 1px solid #ccc;
  border-radius: 5px;
  margin-bottom: 10px;
}

#results {
  text-align: left;
  max-height: 150px;
  overflow-y: auto;
}
.result-item {
  padding: 5px;
  border-bottom: 1px solid #ddd;
}
```

```jsx
const { useState, useEffect } = React;

export function FruitsSearch() {
  const [query, setQuery] = useState('');
  const [results, setResults] = useState([]);

  function handleSubmit(e) {
    e.preventDefault();
  }

  return (
    <div id="search-container">
      <form onSubmit={handleSubmit}>
        <label htmlFor="search-input">Search for fruits:</label>
        <input
          id="search-input"
          type="search"
          value={query}
          onChange={(e) => setQuery(e.target.value)}
        />
      </form>
      --fcc-editable-region--
      <div id="results">

      </div>
      --fcc-editable-region--
    </div>
  );
}
```
