---
id: 6716249b5405164036fd0b0d
title: 정렬 시각화 도구 만들기
challengeType: 25
dashedName: build-a-sorting-visualizer
demoType: onClick
---

# --description--

버블 정렬 알고리즘은 수열의 시작부터 인접한 두 요소를 비교하며 정수 수열을 정렬합니다. 첫 번째 요소가 두 번째 요소보다 크면 두 요소를 교환합니다. 그런 다음 다음 두 요소로 진행합니다. 수열의 마지막 요소에 도달하면 수열의 시작부터 새 사이클을 시작하며, 요소들이 정렬될 때까지 이 과정을 반복합니다. 교환 없이 한 사이클이 완료되면 알고리즘이 정지합니다.

이 실습에서는 모든 HTML과 CSS가 제공됩니다. JavaScript를 사용해 버블 정렬 시각화 도구를 완성하여 다섯 개 정수 배열을 정렬하는 데 필요한 각 단계를 시각화하세요.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `generateElement`라는 이름의 함수가 있어야 하며, `1`부터 `100`까지(포함) 무작위 정수를 반환해야 합니다.
1. `generateArray`라는 이름의 함수가 있어야 하며, `generateElement` 함수를 사용해 다섯 개 무작위 정수를 포함하는 배열을 반환해야 합니다.
1. `generateContainer`라는 이름의 함수가 있어야 하며, 빈 `div` 요소를 생성해 반환해야 합니다.
1. `fillArrContainer`라는 이름의 함수가 있어야 하며, 첫 번째 인수로 HTML 요소를, 두 번째 인수로 배열을 받습니다.
1. `fillArrContainer`는 첫 번째 매개변수로 요소를, 두 번째 매개변수로 정수 배열을 받아, 요소 안에 다섯 개의 `span` 요소를 채우고, 각 `span`은 배열의 정수 하나를 표시해야 합니다.
1. `isOrdered`라는 이름의 함수가 있어야 하며, 두 정수를 받아 첫 번째 정수가 두 번째 정수보다 작거나 같은지 여부를 불리언으로 반환해야 합니다.
1. `swapElements`라는 이름의 함수가 있어야 하며, 정수 배열과 숫자 인덱스를 받습니다.
1. `swapElements` 함수는 `isOrdered`가 `false`를 반환하면, 전달된 인덱스와 그 다음 요소를 교환하여 배열을 제자리에서 수정해야 합니다.
1. `highlightCurrentEls`라는 이름의 함수가 있어야 하며, HTML 요소와 숫자 인덱스를 받습니다.
1. `highlightCurrentEls` 함수는 주어진 요소의 자식 중 인덱스와 그 다음 인덱스에 해당하는 자식에게 `dashed` 스타일, `red` 색상, 그리고 원하는 너비를 설정해야 합니다.
1. `#generate-btn`을 클릭하면 `fillArrContainer` 함수를 사용해 `#starting-array`에 다섯 개의 `span` 요소를 채우고, 각 요소는 무작위 숫자를 텍스트로 가져야 합니다. 만약 존재한다면, `#array-container`의 다른 요소들은 제거해야 합니다.
1. `#sort-btn`을 클릭하면, `#array-container`에는 버블 정렬 알고리즘이 시작 배열을 정렬하는 데 필요한 각 단계마다 `div` 요소가 포함되어야 하며, 시작 배열을 나타내는 `div`와 정렬된 배열을 나타내는 `div`도 포함되어야 합니다. 지금까지 만든 함수들이 여기서 유용할 수 있습니다.
1. 각 `div`는 현재 정렬 상태의 배열을 나타내는 다섯 개 `span` 요소를 포함해야 합니다.
1. `#sort-btn`을 클릭하면 `#starting-array`는 초기 배열과 처음 두 정수가 강조된 시작 단계를 나타내야 합니다.
1. 각 정렬 단계마다 `highlightCurrentEls`를 사용해 비교 중인 두 숫자를 강조하고, 다음 단계에서는 `swapElements`를 사용해 두 숫자를 교환해야 합니다.

# --hints--

`generateElement`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(generateElement);
```

`generateElement` 함수는 `1`부터 `100`까지(포함) 무작위 정수를 반환해야 합니다.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    assert.strictEqual(generateElement(), 26);
    assert.strictEqual(generateElement(), 9);
} finally {
    randomMocker.restore();
}
```

`generateArray`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(generateArray)
```

`generateArray` 함수는 `generateElement` 함수를 사용해야 합니다.

```js
let flag = false;
const temp = generateElement;
generateElement = () => flag = true;
try {
    generateArray();
    assert.isTrue(flag);
} finally {
    generateElement = temp;
}
```

`generateArray` 함수는 `1`부터 `100`까지(포함) 다섯 개 무작위 정수를 포함하는 배열을 반환해야 합니다.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    assert.deepEqual(generateArray(), [26, 9, 58, 23, 38])
} finally {
    randomMocker.restore();
}
```

`generateContainer`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(generateContainer);
```

`generateContainer` 함수는 빈 `div` 요소를 반환해야 합니다.

```js
const div = generateContainer();
assert.equal(div.tagName, "DIV");
assert.isEmpty(div.children);
```

`fillArrContainer`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(fillArrContainer);
```

`fillArrContainer()`는 첫 번째 매개변수로 요소를, 두 번째 매개변수로 정수 배열을 받아, 요소 안에 다섯 개 `span` 요소를 채우고, 각 `span`은 배열의 정수 하나를 표시해야 합니다.

```js
const testDiv = document.createElement("div");
const testArr = [15, 98, 17, 5, 63]
fillArrContainer(testDiv, testArr);
const children = testDiv.children;
assert.lengthOf(children, 5);
Array.from(children).forEach((el, i) => {
    assert.equal(el.tagName, "SPAN");
    assert.equal(el.innerText.trim(), testArr[i])
})
```

`isOrdered`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(isOrdered);
```

`isOrdered` 함수는 두 정수를 받아 첫 번째 정수가 두 번째 정수보다 작거나 같은지 여부를 불리언으로 반환해야 합니다.

```js
assert.lengthOf(isOrdered, 2);
assert.isTrue(isOrdered(2, 60));
assert.isFalse(isOrdered(10, 3));
assert.isTrue(isOrdered(5, 5));
```

`swapElements`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(swapElements);
```

`swapElements` 함수는 정수 배열과 숫자 인덱스를 인수로 받아, 첫 번째 요소가 두 번째 요소보다 크면 전달된 배열에서 해당 인덱스와 다음 요소를 교환하여 제자리에서 수정해야 합니다.

```js
const testArr = [22, 4, 87, 47, 33];
swapElements(testArr, 0);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 1);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 2);
assert.deepEqual(testArr, [4, 22, 47, 87, 33]);
swapElements(testArr, 3);
assert.deepEqual(testArr, [4, 22, 47, 33, 87]);
```

`highlightCurrentEls`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(highlightCurrentEls);
```

`highlightCurrentEls` 함수는 지정된 요소의 자손 중 주어진 인덱스와 다음 인덱스에 해당하는 자손에게 `dashed`, `red`인 테두리를 너비는 원하는 값으로 설정해야 합니다.

```js
const testDiv = document.createElement("div");
document.querySelector("body").appendChild(testDiv)
for (let i = 0; i < 5; i++) {
    testDiv.appendChild(document.createElement("span"));
}
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|#F00|hsl\(0,\s*100%,\s*50%\))/;
const revertBorder = () => {
    for (const el of children) {
        el.style.border = "revert";
    }
}
const children = testDiv.children;

for (let i = 0; i < 3; i++) {
    highlightCurrentEls(testDiv, i);
    for (let j = 0; j < 5; j++) {
        let b = getComputedStyle(children[j]).border;
        if (j == i || j == i + 1) {            
            assert.match(b, redBorderRegex);
            assert.isAbove(parseFloat(b), 0);
        } else {
            assert.notMatch(b, redBorderRegex);
        }
    }
    revertBorder();
}
testDiv.remove();
```

`#generate-btn`를 클릭하면 `#starting-array`에 `span` 다섯 개를 채우고, 각 요소는 `1`부터 `100`까지 무작위 숫자를 텍스트로 가져야 합니다.

```js
const genBtn = document.querySelector("#generate-btn");
genBtn.click();
const children = document.querySelector("#starting-array").querySelectorAll("span");
assert.lengthOf(children, 5);
Array.from(children).forEach(el => {
    assert.equal(el.tagName, "SPAN");
    const num = Number(el.innerText.trim());
    assert.isAtMost(num, 100);
    assert.isAtLeast(num, 1);
})
```

`#starting-array`에 이미 생성된 배열이 있거나 `#array-container`에 정렬된 배열이 있으면, `#generate-btn`를 클릭할 때 `#array-container`의 다른 요소들은 제거되고 `#starting-array`만 새로 생성된 숫자로 남아야 합니다.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
genBtn.dispatchEvent(new Event("click"));

const prevNumbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));

sortBtn.dispatchEvent(new Event("click"));
genBtn.dispatchEvent(new Event("click"));

const container = document.querySelector("#array-container");
assert.lengthOf(container.children, 1);
const numbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));
assert.lengthOf(numbers, 5);
assert.isTrue(prevNumbers.some((num, index) => num !== numbers[index]))
```

`#sort-btn`를 클릭하면 `#array-container`에는 버블 정렬 알고리즘이 시작 배열을 정렬하는 데 필요한 단계 수만큼 `div` 요소가 포함되어야 하며, 시작 배열을 나타내는 `div`와 정렬된 배열을 나타내는 `div`도 포함되어야 합니다.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
// using randomMocker to be sure that the starting array requires 13 steps to be sorted
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.lengthOf(container.children, 13)
    Array.from(container.children).forEach(el => {assert.equal(el.tagName, "DIV")})
} finally {
    randomMocker.restore();
}
```

`#sort-btn`를 클릭하면 `#array-container` 내 각 `div`는 다섯 개 `span` 요소를 포함하고, 각 요소는 숫자를 텍스트로 가지며, 버블 정렬 알고리즘이 시작 배열을 정렬하는 데 필요한 단계를 나타내야 합니다.

```js
const finalArr = [9, 23, 26, 38, 58];
const arrays = [
    [26, 9, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 23, 58, 38],
    [9, 26, 23, 38, 58],
    [9, 26, 23, 38, 58],
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr
]
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.isNotEmpty(container.children);
    Array.from(container.children).forEach((el, i) => {
        Array.from(el.children).forEach((j, k) => {
            assert.strictEqual(Number(j.innerText.trim()), arrays[i][k])
        })
    })
} finally {
    randomMocker.restore();
}
```

`#sort-btn`를 클릭하면 `highlightCurrentEls` 함수를 사용해 각 단계에서 비교 중인 요소를 강조해야 합니다.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
let flag = false;
const temp = highlightCurrentEls;
highlightCurrentEls = () => flag = true;
try {
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    assert.isTrue(flag);
} finally {
    highlightCurrentEls = temp;
}
```

`#sort-btn`를 클릭하면 `#starting-array`는 `highlightCurrentEls`를 사용해 초기 배열과 처음 두 정수가 강조된 시작 단계를 나타내야 합니다.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");   
genBtn.dispatchEvent(new Event("click"));
sortBtn.dispatchEvent(new Event("click"));
const firstContainer = document.querySelector("#starting-array");
const children = firstContainer.children
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|hsl\(0,\s*100%,\s*50%\))/;
assert.match(getComputedStyle(children[0]).border, redBorderRegex);
assert.match(getComputedStyle(children[1]).border, redBorderRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

#sort-btn {
   display: none 
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js
const arrayContainer = document.getElementById("array-container");
const generateArrayBtn = document.getElementById("generate-btn");
const sortArrayBtn = document.getElementById("sort-btn");
const arrayLength = 5;
const minVal = 1;
const maxVal = 100;
const isStart = () => arrayContainer.children.length === 1;
const clearArrayContainer = () => { arrayContainer.innerHTML = '<div id="starting-array"></div>' };
const showSortBtn = () => {sortArrayBtn.style.display = "inline-block"};
const hideSortBtn = () => {sortArrayBtn.style.display = "none"};
const generateElement = () => Math.floor(Math.random() * maxVal + minVal);
const generateArray = () => Array.from({ length: arrayLength }, generateElement);
const fillArrContainer = (container, arr) => {
    container.innerHTML = "";
    arr.forEach(i => {
        const el = document.createElement("span");
        el.innerText = i;
        el.id = `number-${i}`
        container.appendChild(el);
    })
}
const isOrdered = (el1, el2) => el1 <= el2;

const swapElements = (arr, n = 0) => {
    if (n < arr.length - 1 && !isOrdered(arr[n], arr[n + 1])) {
        const temp = arr[n];
        arr[n] = arr[n + 1];
        arr[n + 1] = temp;
    }
}
const getLastChildren = () => arrayContainer.lastElementChild;
const getLastArr = () => {
    const els = Array.from(getLastChildren().children);
    const arr = els.map(el => Number(el.id.replace("number-", "")));
    return arr;
}
const generateContainer = () => {
    const container = document.createElement("div");
    arrayContainer.appendChild(container);
    return container;
}
const highlightCurrentEls = (container, n = 0) => {
    const children = container.children;
    children[n].style.border = "2px dashed red";
    children[n + 1].style.border = "2px dashed red";
}

const highlightSorted = () => {
    getLastChildren().style.border = "4px solid green";
}

const bubbleSort = () => {
    let swapped = true;
    while (swapped) {
        const startingArr = getLastArr();
        startingArr.forEach((_, i) => {
            if (i + 1 < arrayLength) {
                highlightCurrentEls(getLastChildren(), i);
                const arr = getLastArr();
                swapElements(arr, i);
                fillArrContainer(generateContainer(), arr);
            }
        })
        const lastArr = getLastArr()
        if (startingArr.every((el, i) => el === lastArr[i])) swapped = false;
    }
}


generateArrayBtn.addEventListener("click", () => {
    if (!isStart()) {
        clearArrayContainer();
    }
    fillArrContainer(document.getElementById("starting-array"), generateArray());
    showSortBtn();
});

sortArrayBtn.addEventListener("click", () => {
    bubbleSort();
    highlightSorted();
    hideSortBtn();
})
```
