---
id: 66da326c02141df538f29ba5
title: بناء صفحة الدفع
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تحتوي على عنصر `h1` بالنص `Checkout`.
1. يجب أن تحتوي على عنصرين من نوع `section` مباشرة بعد عنصر `h1`.
1. يجب أن تحتوي على عنصر `h2` بالنص `Your Cart` داخل القسم الأول.
1. يجب أن تحتوي على صورة لعنصر في القسم الأول مع نص بديل مناسب. يمكنك استخدام هذه الصورة إذا أردت: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. يجب أن تحتوي على عنصر `h2` بالنص `Payment Information` داخل القسم الثاني.
1. يجب أن تحتوي على عنصر `form` داخل القسم الثاني.
1. يجب أن يحتوي النموذج على مُدخَل له `id` و `name` باسم `card-name`، و `type` من نوع `text`، مع وجود `label` مرتبط به.
1. يجب أن يحتوي النموذج على مُدخَل له `id` و `name` باسم `card-number`، و `type` من نوع `text`، مع وجود `label` مرتبط به.
1. يجب أن تحتوي على الأقل على عنصرين من نوع `input` يحملان الخاصية `required`.
1. يجب أن تتضمن داخل كل `label` لكل مُدخَل مطلوب عنصر `span` يحتوي على النص `*` مع تعيين `aria-hidden` إلى `true`، ليُظهر أن الحقول المطلوبة محددة بصريًا.
1. يجب أن تحتوي على عنصر `p` يحتوي على نص مساعدة يشرح صيغة رقم البطاقة المطلوبة، ويوضع مباشرة بعد مُدخَل رقم البطاقة. يجب أن يحمل هذا العنصر `p` الخاصية `id` بقيمة `card-number-help` ويُشار إليه من مُدخَل رقم البطاقة باستخدام `aria-describedby`.

# --hints--

يجب أن يكون لديك عنصر `h1` يحتوي على النص `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

يجب أن يكون لديك عنصر واحد فقط من نوع `h1` في صفحتك.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

يجب أن يكون لديك على الأقل عنصران `section` بعد عنصر `h1`.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

يجب أن يكون لديك عنصر `h2` يحتوي على النص `Your Cart` داخل القسم الأول.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

يجب أن يكون لديك على الأقل صورة واحدة مع نص بديل في القسم الأول.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

يجب أن يكون لديك عنصر `h2` يحتوي على النص `Payment Information` داخل القسم الثاني.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

يجب أن يكون لديك عنصر `form` داخل عنصر `section` الثاني.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

يجب أن يكون لديك عنصر `input` له `id` و`name` باسم `card-name` داخل النموذج.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

يجب أن يكون لمُدخَل اسم البطاقة الخاص بك `type` من نوع `text`.

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

يجب أن يكون لديك عنصر `input` له `id` و`name` باسم `card-number` داخل النموذج.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

يجب أن يكون لمُدخَل رقم البطاقة الخاص بك `type` من نوع `text`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

يجب ربط جميع عناصر `input` التي ليست من نوع `type` يساوي `submit` بعنصر `label`.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

يجب أن يكون لديك على الأقل عنصران `input` يحملان الخاصية `required`.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

يجب أن تتضمن عنصر `span` يحتوي على النص `*` داخل عنصر `label` لكل مُدخَل مطلوب.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

يجب أن تكون الخاصية `aria-hidden` مضبوطة على `true` في عناصر `span` الخاصة بك.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

يجب أن يكون لديك عنصر `p` يحمل `id` باسم `card-number-help` مباشرة بعد مُدخَل رقم البطاقة.

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

يجب ألا يكون نص مساعدة رقم البطاقة فارغًا.

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

يجب أن يكون لمُدخَل رقم البطاقة الخاصية `aria-describedby` مضبوطة على `card-number-help`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
