---
id: 6735a7370e0ae93a4577c689
title: Спроектуйте макет газети
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

У цій лабораторній роботі ви спроектуєте макет газети за допомогою CSS Grid, включно з поняттями рядків і стовпців сітки.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ваша сторінка має містити `main` елемент із класом `newspaper-layout`.
1. `.newspaper-layout` має містити `header` з класом `title`, який містить `h1` для відображення назви газети.
1. `.newspaper-layout` має містити `article` з класом `feature-article` для головної новини.
1. `.feature-article` має містити `h2` елемент для заголовка статті, за яким іде `p` елемент для вмісту статті.
1. Ви маєте додати ще три `article` елементи для менших статей із класами `small-article1`, `small-article2` та `small-article3`.
1. Кожна з менших статей має містити `h3` елемент для заголовка статті, за яким іде `p` елемент для вмісту статті.
1. `.newspaper-layout` має містити `article` елемент із класом `secondary-article` для додаткового розділу новин.
1. `.secondary-article` має містити `h2` елемент для заголовка статті, за яким іде `p` елемент для вмісту статті.
1. `.newspaper-layout` має містити `figure` з класом `cover-image` для відображення зображення, що представляє вміст газети.
1. Елементи з класами `title`, `feature-article`, `secondary-article`, `cover-image`, `small-article1`, `small-article2` та `small-article3` мають мати властивість `grid-area`, встановлену на відповідну назву класу.
1. Ваш `.newspaper-layout` має використовувати CSS Grid із властивістю `grid-template-areas` для визначення розташування розділів:
   - `.title` має розтягуватися на верхній рядок.
   - `.feature-article` має охоплювати два стовпці у другому рядку, а `.cover-image` — у третьому стовпці.
   - `.secondary-article` має охоплювати два стовпці у третьому рядку, а `.cover-image` — у третьому стовпці.
   - Три маленькі статті мають розташовуватися у четвертому рядку.
1. `.newspaper-layout` має мати властивість `grid-template-columns`, що ділить простір на три рівні стовпці.
1. Ви маєте встановити властивість `.newspaper-layout` для `grid-template-rows` на `auto` для першого рядка та поділити решту простору на рівні частини для інших рядків.
1. Додайте горизонтальний розрив між елементами сітки.
1. Переконайтеся, що зображення всередині `.cover-image` поміщається у відведений простір, встановивши його властивість `max-width` на `100%`.

**Примітка:** Не забудьте підключити свій файл стилів у HTML і застосувати CSS.

# --hints--

Ви маєте мати `main` елемент.

```js
assert.exists(document.querySelector('main'));
```

Ваш `main` елемент має мати `class` зі значенням `newspaper-layout`.

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

Ви маєте мати `header` елемент із `class` зі значенням `title` всередині вашого `.newspaper-layout` елемента.

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

Ваш `header` елемент має містити `h1` елемент.

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

Ваш `.title` елемент має містити текст.

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

Ви маєте мати `article` елемент для кожної статті з класами `feature-article`, `small-article1`, `small-article2`, `small-article3` та `secondary-article` відповідно.

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

`.feature-article` має містити `h2` елемент, за яким іде `p` елемент.

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`, `.small-article2` та `.small-article3` мають містити `h3` елемент, за яким іде `p` елемент.

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

`.secondary-article` має містити `h2` елемент, за яким іде `p` елемент.

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

Ви маєте мати `figure` елемент із класом `cover-image` всередині вашого `.newspaper-layout`.

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

Ваш `.cover-image` має містити `img` елемент.

```js
assert.exists(document.querySelector(".cover-image img"));
```

Властивість `display` елемента `.newspaper-layout` має бути встановлена на `grid`.

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

Ваш `.newspaper-layout` елемент має мати властивість `grid-template-areas` зі значенням `title`, що охоплює всі три стовпці першого рядка.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

Другий рядок шаблону сітки має містити `feature-article`, що охоплює два стовпці, і `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

Третій рядок шаблону сітки має містити `secondary-article`, що охоплює два стовпці, і `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

Четвертий рядок шаблону сітки має містити `small-article1`, `small-article2` та `small-article3`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

Ваш селектор `.newspaper-layout` має використовувати властивість `grid-template-columns` для поділу простору на три рівні частини. Пам’ятайте, що можна використовувати функцію `repeat` або вручну визначити кожну дробову одиницю.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

Ваш селектор `.newspaper-layout` має мати властивість `grid-template-rows`, встановлену на `auto 1fr 1fr 1fr`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

Додайте горизонтальний розрив між елементами сітки.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title` має мати властивість `grid-area` зі значенням `title`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` має мати властивість `grid-area` зі значенням `feature-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` має мати властивість `grid-area` зі значенням `cover-image`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article` має мати властивість `grid-area` зі значенням `secondary-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1` має мати властивість `grid-area` зі значенням `small-article1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2` має мати властивість `grid-area` зі значенням `small-article2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3` має мати властивість `grid-area` зі значенням `small-article3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

`img` всередині `.cover-image` має мати властивість `max-width` зі значенням `100%`.

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
