---
id: 5e44414f903586ffb414c950
title: 確率計算機プロジェクトを作成する
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462364
dashedName: build-a-probability-calculator-project
---

# --description--

青いボールが5個、赤いボールが4個、緑のボールが2個入った帽子があるとします。4つのボールをランダムに引いたときに、少なくとも1つの赤いボールと2つの緑のボールが含まれる確率はどれくらいでしょうか？高度な数学を使って確率を計算することも可能ですが、より簡単な方法は、多数の実験を行うプログラムを書いて近似的な確率を推定することです。

このプロジェクトでは、帽子から特定のボールをランダムに引く確率の近似値を求めるプログラムを書きます。

まず、`Hat` クラスを `main.py` に作成してください。このクラスは、帽子に入っている各色のボールの数を指定する可変長の引数を受け取る必要があります。例えば、クラスのオブジェクトは以下のいずれかの方法で作成できます。

```py
hat1 = Hat(yellow=3, blue=2, green=6)
hat2 = Hat(red=5, orange=4)
hat3 = Hat(red=5, orange=4, black=1, blue=0, pink=2, striped=9)
```

帽子は常に少なくとも1つのボールを持って作成されます。作成時に帽子オブジェクトに渡された引数は `contents` のインスタンス変数に変換してください。`contents` は帽子の中の各ボールを表す文字列のリストである必要があります。リストの各要素は、その色のボール1つを表す色名の文字列です。例えば、帽子が `{'red': 2, 'blue': 1}` の場合、`contents` は `['red', 'red', 'blue']` となるはずです。

`Hat` クラスには、`draw` メソッドを実装してください。このメソッドは、帽子から引くボールの数を示す引数を受け取ります。このメソッドは `contents` からランダムにボールを取り除き、そのボールを文字列のリストとして返します。引いたボールは戻されず、置き換えなしの抽出実験のように扱います。引くボールの数が帽子の中のボールの数を超える場合は、すべてのボールを返してください。

次に、`experiment` 関数を `main.py` に作成してください（`Hat` クラスの中ではありません）。この関数は以下の引数を受け取ります。

- `hat`：帽子オブジェクト。関数内でコピーされるべきです。
- `expected_balls`：実験で引こうとする正確なボールのグループを示すオブジェクト。例えば、帽子から青いボール2個と赤いボール1個を引く確率を求める場合、`expected_balls` を `{'blue':2, 'red':1}` に設定します。
- `num_balls_drawn`：各実験で帽子から引くボールの数。
- `num_experiments`：実験の回数。（実験回数が多いほど近似確率は正確になります。）

`experiment` 関数は確率を返す必要があります。

例えば、黒いボール6個、赤いボール4個、緑のボール3個が入った帽子から5個のボールを引いたときに、少なくとも赤いボール2個と緑のボール1個が出る確率を求めたい場合、`N` 回の実験を行い、`M` 回だけ赤いボール2個以上と緑のボール1個以上が出た回数を数え、確率を `M/N` として推定します。各実験は、指定されたボールが入った帽子から始め、いくつかのボールを引き、引こうとしたボールが得られたかを確認する流れです。

上記の例に基づいて2000回の実験を行う場合、`experiment` 関数は次のように呼び出します。

```py
hat = Hat(black=6, red=4, green=3)
probability = experiment(hat=hat,
                  expected_balls={'red':2,'green':1},
                  num_balls_drawn=5,
                  num_experiments=2000)
```

出力は次のようになります。

```bash
0.356
```

これはランダムな抽出に基づいているため、コードを実行するたびに確率は少し異なります。

_ヒント：ファイルの先頭で既にインポートされているモジュールを使うことを検討してください。ファイル内でrandomのシードを初期化しないでください。_

注意：F12でブラウザのコンソールを開くと、テストの詳細な出力が見られます。

# --hints--

`hat` オブジェクトの作成は正しい内容を追加する必要があります。

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_class_contents(self):
        hat = probability_calculator.Hat(red=3,blue=2)
        actual = hat.contents
        expected = ["red","red","red","blue","blue"]
        self.assertEqual(actual, expected, 'Expected creation of hat object to add correct contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

`draw` メソッドは `hat` クラス内で、内容のアイテム数を減らす必要があります。

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)

class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw(self):
        hat = probability_calculator.Hat(red=5,blue=2)
        actual = hat.draw(2)
        expected = ['blue', 'red']
        self.assertEqual(actual, expected, 'Expected hat draw to return two random items from hat contents.')
        actual = len(hat.contents)
        expected = 5
        self.assertEqual(actual, expected, 'Expected hat draw to reduce number of items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

`draw` メソッドは、引くボールの数が帽子の中のボールの数より多い場合に正しく動作する必要があります。

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw_2(self):
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        actual = sorted(hat.draw(20))
        expected = sorted(['yellow', 'yellow', 'yellow', 'yellow', 'yellow', 'red', 'green', 'green', 'green', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'test'])
        self.assertEqual(actual, expected, 'Expected hat draw to return all items from hat contents.')
        actual = len(hat.contents)
        expected = 0
        self.assertEqual(actual, expected, 'Expected hat draw to leave no items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

`experiment` メソッドは異なる確率を返す必要があります。

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_prob_experiment(self):
        hat = probability_calculator.Hat(blue=3,red=2,green=6)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"blue":2,"green":1}, num_balls_drawn=4, num_experiments=1000)
        actual = probability
        expected = 0.272
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"yellow":2,"blue":3,"test":1}, num_balls_drawn=20, num_experiments=100)
        actual = probability
        expected = 1.0
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

# --seed--

## --seed-contents--

```py
import copy
import random

class Hat:
    pass

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):
    pass
```

# --solutions--

```py
import copy
import random

class Hat:
    def __init__(self, **hat):
        self.hat = hat
        contents = []
        for i in hat:
            for j in range(hat[i]):
                contents.append(i)           
        self.contents = contents
               

    def draw(self, number): 
        drawn = []
        if number >= len(self.contents):
            drawn.extend(self.contents)
            self.contents = []
        else:
            for i in range(number):
                drawn.append(
                    self.contents.pop(random.randrange(len(self.contents)))
                )                
        return drawn

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):    
    expected_balls_list = []
    drawn_list = []
    success = 0
    for i in expected_balls:
        for j in range(expected_balls[i]):
            expected_balls_list.append(i)
    for j in range(num_experiments):
        hat_copy = copy.deepcopy(hat)
        drawn_list.append(hat_copy.draw(num_balls_drawn))        
        exp_ball_list_copy = expected_balls_list[:]
        for k in range(len(drawn_list[j])):
            try:
                ind = exp_ball_list_copy.index(drawn_list[j][k])
                exp_ball_list_copy.pop(ind)
            except:
                continue
        if len(exp_ball_list_copy) == 0:
            success += 1        
    probability = success/num_experiments    
    return probability
```
