---
id: 66c051d13a6a20255a963695
title: بناء قائمة مهام مزخرفة
challengeType: 25
dashedName: build-a-stylized-to-do-list
demoType: onClick
---

# --description--

في هذا المختبر، ستتدرب على الأنماط المختلفة التي يمكن تطبيقها على الروابط عند التمرير فوقها، أو تركيزها، أو النقر عليها، أو زيارتها.

**الهدف:** أتمم قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك قائمة غير مرتبة واحدة بالصنف `todo-list`.
2. داخل القائمة غير المرتبة، يجب أن يكون لديك أربعة عناصر قائمة.
3. داخل كل عنصر قائمة، يجب أن يكون هناك:

- عنصر `input` من النوع `checkbox` مع تعيين `id` لقيمة فريدة.
   - عنصر `label` مع خاصية `for` تشير إلى `id` الخاص بعنصر `input` المقابل.
   - قائمة غير مرتبة بالصنف `sub-item`.
   - عنصر قائمة يحتوي على عنصر وسم الرابط. يجب أن يحمل وسم الرابط الصنف `sub-item-link`، وقيمة `href` صالحة، وقيمة `target` تجعل الرابط يفتح في تبويب جديد.

4. يجب ألا تحتوي عناصر `a` الخاصة بك على أي زخارف نصية.
5. يجب تعيين لون نص الروابط غير المزارة إلى لون تختاره.
6. عند زيارة الروابط، يجب أن يتغير اللون إلى لون آخر تختاره.
7. عند التمرير فوق الروابط، يجب أن يتغير اللون إلى لون آخر تختاره.
8. عند تركيز الروابط، يجب أن يكون هناك إطار ملون حول الرابط.
9. عند النقر على الروابط، يجب أن يتغير اللون إلى لون آخر تختاره.

**ملاحظة:** تأكد من ربط ورقة الأنماط في ملف HTML الخاص بك وتطبيق CSS.

# --hints--

يجب أن يكون لديك قائمة غير مرتبة واحدة بالصنف `todo-list`.

```js
assert.exists(document.querySelector('ul.todo-list'));
assert.lengthOf(document.querySelectorAll('ul.todo-list'), 1);
```

يجب أن يكون لديك أربعة عناصر قائمة داخل القائمة غير المرتبة.

```js
const li = document.querySelectorAll('ul.todo-list > li');
assert.lengthOf(li, 4);
```

يجب أن يحتوي عنصر `li` داخل `ul` ذي الصنف `todo-list` على عنصر `input` له `type` بقيمة `checkbox`.

```js
const checkboxes = document.querySelectorAll(
  'ul.todo-list li input[type="checkbox"]'
);
assert.lengthOf(checkboxes, 4);
```

يجب أن يحتوي عنصر `li` داخل `ul` بالصنف `todo-list` على عنصر `label`.

```js
const labels = document.querySelectorAll(
  'ul.todo-list li label'
);
assert.lengthOf(labels, 4);
```

يجب أن تحتوي جميع عناصر `input` على خاصية `id`.

```js
const inputs = document.querySelectorAll('input');
assert.isAbove(inputs.length, 0);
for (let input of inputs) {
  assert.isAbove(input?.id.length, 0);
}
```

يجب أن تحتوي جميع عناصر `label` على خاصية `for`.

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.htmlFor.length, 0);
}
```

يجب أن تحتوي جميع عناصر `label` على نص.

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.innerText.length, 0);
}
```

يجب أن تتطابق قيم خاصيتي `id` و`for` في أزواج عناصر `input` و`label` المقابلة.

```js
const labels = document.querySelectorAll("label");
assert.isNotEmpty(labels);

document.querySelectorAll("input").forEach((input, index) => {
  assert.equal(input.id, labels[index].htmlFor);
});
```

بعد عناصر `label`، يجب أن تكون هناك قائمة غير مرتبة بالصنف `sub-item`.

```js
const labelSbItems = document.querySelectorAll('ul.todo-list li label + ul.sub-item');
assert.lengthOf(labelSbItems, 4);
```

يجب أن يحتوي عنصر `li` داخل `ul` بالصنف `sub-item` على عنصر وسم الرابط بالصنف `sub-item-link`.

```js
const allItems = document.querySelectorAll(
  'ul.todo-list li label + ul.sub-item li a.sub-item-link'
);
assert.lengthOf(allItems, 4);
```

يجب أن تحتوي جميع عناصر وسم الرابط على خاصية `href` صالحة.

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.href.length, 0);
}
```

يجب أن تحتوي جميع عناصر وسم الرابط على نص.

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.innerText.length, 0);
}
```

يجب أن يحتوي كل عنصر `a` على خاصية `target` بقيمة `_blank`.

```js
const anchors = document.querySelectorAll('a');
assert.isAbove(anchors.length, 0);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

يجب ألا تحتوي عناصر `a` الخاصة بك على أي زخارف نصية.

```js
assert.isTrue(
  new __helpers.CSSHelp(document).getStyle('a.sub-item-link')
    ?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('a')?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('.sub-item-link')
      ?.textDecoration === 'none'
);
```

يجب تعيين لون نص الروابط غير المزارة إلى لون تختاره.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:link')?.color
);
```

يجب أن يتغير لون الروابط عند التمرير فوقها.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:hover')?.color
);
```

يجب أن يتغير لون الروابط عند النقر عليها.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:active')?.color
);
```

يجب أن يكون للروابط إطار عند تركيزها.

```js
const css = new __helpers.CSSHelp(document);
assert.exists(
  css.getStyle('.sub-item-link:focus')?.outline ||
    css.getStyle('a:focus')?.outline ||
    css.getStyle('a.sub-item-link:focus')?.outline ||
    css.getStyle('.sub-item-link:focus')?.border ||
    css.getStyle('a:focus')?.border ||
    css.getStyle('a.sub-item-link:focus')?.border
);
```

يجب أن يتغير لون الروابط عند زيارتها.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:visited')?.color
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Styled To-Do List</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Styled To-Do List</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div class="container">
      <h1>My To-Do List</h1>
      <ul class="todo-list">
        <li>
          <input type="checkbox" id="task1" />
          <label for="task1">Explore gaming keyboards</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link">Store link</a>
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task2" />
          <label for="task2">Finish the report</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >Request Access</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task3" />
          <label for="task3">View Phone's Warranty</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Receipts</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task4" />
          <label for="task4">Check Processor Specs</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Model Number</a
              >
            </li>
          </ul>
        </li>
      </ul>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: #f0f0f0;
  margin: 0;
  padding: 20px;
}

.container {
  max-width: 400px;
  margin: 0 auto;
  background-color: #fff;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
  color: #333;
}

.todo-list {
  list-style-type: none;
  padding: 0;
}

.todo-list li {
  margin-bottom: 15px;
  line-height: 1.6;
  padding: 10px;
  background-color: #e5f1e5;
  border-radius: 10px;
  position: relative;
}

a.sub-item-link {
  margin-left: 10px;
  text-decoration: none;
}

.sub-item-link:link {
  color: blue;
}

.sub-item-link:visited {
  color: purple;
}

.sub-item-link:hover {
  color: red;
}

.sub-item-link:focus {
  outline: 2px solid blue;
  outline-offset: 2px;
}

.sub-item-link:active {
  color: green;
}

.sub-item {
  list-style-type: square;
}
```
