---
id: 587d8248367417b2b2512c3b
title: helmet.ieNoOpen() で IE が信頼できない HTML を開くのを防ぐ
challengeType: 2
forumTopicId: 301584
dashedName: prevent-ie-from-opening-untrusted-html-with-helmet-ienoopen
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のスタータープロジェクトを基に構築されています。

一部のウェブアプリケーションは、信頼できない HTML をダウンロード用に配信します。Internet Explorer の一部のバージョンはデフォルトでは、その HTML ファイルをあなたのサイトのコンテキストで開きます。つまり、信頼できない HTML ページがあなたのページのコンテキストで悪意のある動作を始める可能性があります。このミドルウェアは X-Download-Options ヘッダーを noopen に設定します。これにより、IE ユーザーが信頼されたサイトのコンテキストでダウンロードを実行するのを防げます。

# --instructions--

サーバーで `helmet.ieNoOpen()` メソッドを使ってください。

# --hints--

helmet.ieNoOpen() ミドルウェアが正しくマウントされていることを確認してください

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'ienoopen');
assert.equal(data.headers['x-download-options'], 'noopen');
```
