---
id: 587d8249367417b2b2512c3e
title: helmet.noCache() でクライアント側キャッシュを無効化する
challengeType: 2
forumTopicId: 301576
dashedName: disable-client-side-caching-with-helmet-nocache
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のスタータープロジェクトを基に構築されています。

ウェブサイトのアップデートを公開するときに、ユーザーに常に新しいバージョンをダウンロードしてもらいたい場合、クライアントのブラウザでのキャッシュを無効化（試みる）ことができます。これは開発中にも役立ちます。キャッシュはパフォーマンス向上に効果があるため、実際に必要な場合のみこのオプションを使うことをお勧めします。

# --instructions--

サーバーで `helmet.noCache()` メソッドを使用してください。

# --hints--

helmet.noCache() ミドルウェアが正しくマウントされていることを確認してください

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nocache');
assert.equal(
  data.headers['cache-control'],
  'no-store, no-cache, must-revalidate, proxy-revalidate'
);
```
