---
id: 587d8248367417b2b2512c3a
title: helmet.noSniff() でレスポンスの MIME タイプの推測を防ぐ
challengeType: 2
forumTopicId: 301574
dashedName: avoid-inferring-the-response-mime-type-with-helmet-nosniff
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のスタータープロジェクトを基に構築されています。

ブラウザーはコンテンツや MIME スニッフィングを使って、レスポンスの `Content-Type` ヘッダーを上書きし、暗黙のコンテンツタイプを推測してデータを処理する場合があります。これは特定の状況では便利ですが、危険な攻撃につながることもあります。このミドルウェアは `X-Content-Type-Options` ヘッダーを `nosniff` に設定し、ブラウザーに提供された `Content-Type` をバイパスしないよう指示します。

# --instructions--

サーバーで `helmet.noSniff()` メソッドを使ってください。

# --hints--

helmet.noSniff() ミドルウェアが正しくマウントされていることを確認してください

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nosniff');
assert.equal(data.headers['x-content-type-options'], 'nosniff');
```
