---
id: 587d8247367417b2b2512c36
title: Helmet のインストールと require
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

このチャレンジでは、次のいずれかの方法でコードを書いて進めていきます。

- <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">この GitHub リポジトリ</a>をクローンして、ローカルでチャレンジを完了する。
- お好きなサイトビルダーを使ってプロジェクトを完成させる。GitHub リポジトリのすべてのファイルを必ず取り込んでください。

Helmet はさまざまな HTTP ヘッダーを設定することで、Express アプリのセキュリティを強化します。

# --instructions--

このレッスンのコードはすべて、`myApp.js` ファイルの、こちらで用意したコードの間に書いてください。追加したコードは変更したり削除したりしないでください。

Helmet バージョン `3.21.3` はすでにインストールされているので、`helmet` で `myApp.js` として require してください。

# --hints--

`helmet` バージョン `3.21.3` は `package.json` にあるはずです。

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```
