---
id: 587d8247367417b2b2512c39
title: helmet.xssFilter() でクロスサイトスクリプティング（XSS）攻撃のリスクを軽減する
challengeType: 2
forumTopicId: 301583
dashedName: mitigate-the-risk-of-cross-site-scripting-xss-attacks-with-helmet-xssfilter
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のスタータープロジェクトを基に構築されています。

クロスサイトスクリプティング（XSS）は、悪意のあるスクリプトが脆弱なページに注入される頻繁な攻撃の一種で、セッションCookieやパスワードなどの機密データを盗む目的があります。

XSS攻撃のリスクを下げる基本ルールはシンプルです：「ユーザーの入力を決して信用しないこと」。開発者として、外部からのすべての入力を常にサニタイズする必要があります。これにはフォームからのデータ、GETクエリURL、さらにはPOSTボディのデータも含まれます。サニタイズとは、危険な可能性のある文字（例：&lt;、>）を見つけてエンコードすることを意味します。

最新のブラウザは、より良いソフトウェア戦略を採用することでリスクを軽減するのに役立ちます。これらは多くの場合、HTTPヘッダーを通じて設定可能です。

X-XSS-Protection HTTPヘッダーは基本的な保護機能です。ブラウザはヒューリスティックフィルターを使って潜在的な注入スクリプトを検出します。このヘッダーが有効な場合、ブラウザはスクリプトコードを変更して無効化します。ただし、対応はまだ限定的です。

# --instructions--

サーバーに送信される入力をサニタイズするために `helmet.xssFilter()` を使用してください。

# --hints--

helmet.xssFilter() ミドルウェアが正しくマウントされていることを確認してください。

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'xXssProtection');
assert.property(data.headers, 'x-xss-protection');
```
