---
id: 587d8248367417b2b2512c3d
title: helmet.dnsPrefetchControl() で DNS プリフェッチを無効化する
challengeType: 2
forumTopicId: 301577
dashedName: disable-dns-prefetching-with-helmet-dnsprefetchcontrol
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のスタータープロジェクトを基に構築されています。

パフォーマンス向上のため、ほとんどのブラウザはページ内のリンクの DNS レコードをプリフェッチします。こうすることで、ユーザーがリンクをクリックしたときに目的の IP アドレスがすでに分かっている状態になります。しかし、これにより DNS サービスの過剰利用（数百万人が訪れる大規模なウェブサイトの場合など）、プライバシーの問題（盗聴者が特定のページにアクセスしていることを推測できる）、ページ統計の改ざん（一部のリンクが実際には訪問されていなくても訪問済みと表示される）などが起こる可能性があります。高いセキュリティが必要な場合は、パフォーマンスの低下を伴いますが DNS プリフェッチを無効化できます。

# --instructions--

サーバーで `helmet.dnsPrefetchControl()` メソッドを使用してください。

# --hints--

helmet.dnsPrefetchControl() ミドルウェアが正しくマウントされていることを確認してください

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'dnsPrefetchControl');
assert.equal(data.headers['x-dns-prefetch-control'], 'off');
```
