---
id: 587d8247367417b2b2512c37
title: helmet.hidePoweredBy() を使って潜在的に危険な情報を隠す
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> からクローンした以下のボイラープレートを基に構築されています。

Express/ノードを使っていることが分かると、ハッカーは既知の脆弱性を悪用する可能性があります。`X-Powered-By: Express` は Express からのすべてのリクエストにデフォルトで送信されます。`helmet.hidePoweredBy()` ミドルウェアを使って X-Powered-By ヘッダーを削除してください。

# --hints--

helmet.hidePoweredBy() ミドルウェアを正しくマウントしてください

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```
