---
id: 587d8249367417b2b2512c40
title: 「parent」helmet()ミドルウェアを使ったHelmetの設定
challengeType: 2
forumTopicId: 301575
dashedName: configure-helmet-using-the-parent-helmet-middleware
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に構築されています。

`app.use(helmet())`は、上で紹介したすべてのミドルウェアを自動的に含みますが、`noCache()`と`contentSecurityPolicy()`は含まれていません。ただし、必要に応じてこれらを有効にすることができます。また、設定オブジェクトを使って他のミドルウェアを個別に無効化したり設定したりすることも可能です。

**例：**

```js
app.use(helmet({
  frameguard: {         // configure
    action: 'deny'
  },
  contentSecurityPolicy: {    // enable and configure
    directives: {
      defaultSrc: ["'self'"],
      styleSrc: ['style.com'],
    }
  },
  dnsPrefetchControl: false     // disable
}))
```

学習目的とテストのしやすさのために、各ミドルウェアは個別に紹介しました。実際のプロジェクトでは、「parent」`helmet()`ミドルウェアを使うのが簡単に実装できます。

# --hints--

テストはありません - これは説明用のチャレンジです

```js
assert(true);
```
