---
id: 58a25bcff9fc0f352b528e7d
title: 非同期でパスワードをハッシュ化して比較する
challengeType: 2
forumTopicId: 301578
dashedName: hash-and-compare-passwords-asynchronously
---

# --description--

このプロジェクトは、<a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>からクローンした以下のスタータープロジェクトを基に構築されています。

ハッシュ化は計算負荷が高くなるよう設計されているため、ハッシュ化中に着信接続がブロックされないよう、サーバー上で非同期に行うことをお勧めします。パスワードを非同期にハッシュ化するには、単に次のように呼び出せばよいです。

```js
bcrypt.hash(myPlaintextPassword, saltRounds, (err, hash) => {
  /*Store hash in your db*/
});
```

# --instructions--

このハッシュ化関数をサーバーに追加してください（関数内で使う変数はすでに定義済みです）。コンソールにログを出力して確認しましょう！この時点で通常はハッシュをデータベースに保存します。

新しい入力がハッシュと同じデータかどうかを判別したい場合は、compare関数を使います。

```js
bcrypt.compare(myPlaintextPassword, hash, (err, res) => {
  /*res == true or false*/
});
```

ハッシュ化が完了するのを待ってからcompare関数を呼び出す必要があるので、既存のハッシュ関数内にこれを追加してください。ハッシュが完了したらログに出力し、compare内で'res'をコンソールにログしてください。コンソールにはハッシュと、その後に'true'が出力されるはずです。compare関数内の'myPlaintextPassword'を'someOtherPlaintextPassword'に変えると、falseが表示されるはずです。

```js
bcrypt.hash('passw0rd!', 13, (err, hash) => {
  console.log(hash);
  //$2a$12$Y.PHPE15wR25qrrtgGkiYe2sXo98cjuMCG1YwSI5rJW1DSJp0gEYS
  bcrypt.compare('passw0rd!', hash, (err, res) => {
    console.log(res); //true
  });
});

```

正しくできたと思ったらページを提出してください。

# --hints--

非同期のハッシュが生成され、正しく比較されていることを確認してください。

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash.*myPlaintextPassword( |),( |)saltRounds( |),( |).*err( |),( |)hash[^]*END_ASYNC/gi,
  'You should call bcrypt.hash on myPlaintextPassword and saltRounds and handle err and hash as a result in the callback'
);
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash[^]*bcrypt.compare.*myPlaintextPassword( |),( |)hash( |),( |).*err( |),( |)res[^]*}[^]*}[^]*END_ASYNC/gi,
  'Nested within the hash function should be the compare function comparing myPlaintextPassword to hash'
);
```
