---
id: 67fe626cc464f1238514cd8b
title: Passo 20
challengeType: 0
dashedName: step-20
---

# --description--

Antes da instrução `return` dentro da chamada `filteredItems.map()`, crie uma variável chamada `isChecked` e use o método `includes()` para verificar se o `item` atual está no array `selectedItems`.

Então, adicione uma prop `style` ao elemento `li` e use um operador ternário para definir a propriedade `textDecoration` como `line-through` se `isChecked` for verdadeiro ou `none` se for falso.Lembre-se de que você precisa envolver o operador ternário em chaves já que é uma expressão JavaScript.

Por exemplo, aqui está como você pode usar um operador ternário para definir condicionalmente a cor de um parágrafo em um componente simples:

```jsx
const MyComponent = () => {
  const isActive = true;

  return (
    <p style={{ color: isActive ? "green" : "black"}}>This text is conditionally styled.</p>
  );
};
```

# --hints--

Você deve criar uma variável chamada `isChecked` que usa o método `includes()` no array `selectedItems` para verificar se o `item` atual está no array.

```js
const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
const exports = {};
const a = eval(script);
const shoppingListString = exports.ShoppingList.toString();

assert.match(shoppingListString, /var\s*isChecked\s*=\s*selectedItems\.includes\s*\(\s*item\s*\)/);
```

Você deve adicionar uma prop `style` ao elemento `li`.

```js
  const testElem = await __helpers.prepTestComponent(window.index.ShoppingList);
  const li = testElem.querySelector('li');
  const key = li && Object.keys(li)?.find(key => key.startsWith("__reactProps"));
  const reactLi = li && li[key];

  assert.exists(reactLi.style);
```

Sua prop `style` deve incluir um operador ternário que define a propriedade `textDecoration` como `line-through` se `isChecked` for true ou `none` se for false.

```js
const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
const exports = {};
const a = eval(script);
const shoppingListString = exports.ShoppingList.toString();

assert.match(shoppingListString, /"li"\s*,\s*{(.|\n)*style:\s*{\s*textDecoration\s*:\s*isChecked\s*\?\s*("|')line-through("|')\s*:\s*("|')none("|')/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Shopping List</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main id="app-container"></main>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ShoppingList } from "./index.jsx";
      const appContainer = document.getElementById("app-container");
      const root = ReactDOM.createRoot(appContainer);
      root.render(<ShoppingList />);
    </script>
  </body>
</html>

```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

*,
*::before,
*::after {
  box-sizing: border-box;
}

body {
  font-family: Arial, sans-serif;
  line-height: 1.5;
  color: var(--dark-grey);
  background-color: var(--dark-grey);
}

main,
.container {
  display: flex;
  flex-direction: column;
  align-items: center;
}

.container {
  background-color: var(--light-grey);
  width: 90%;
  margin: 20px;
  padding: 10px;
}

h1 {
  color: var(--dark-grey);
}

form {
  text-align: center;
}

button {
  cursor: pointer;
}

button {
  cursor: pointer;
  width: 100px;
  margin: 3px;
  color: var(--dark-grey);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

input {
  color: var(--dark-grey);
  margin-left: 5px;
  padding: 3px;
}

li {
  text-align: left;
  margin: 10px 0;
}

.selected-item {
  font-weight: bold;
}

@media (min-width: 425px) {
  .container {
    width: 400px;
  }
}

```

```jsx
const { useState } = React;

const items = [
  "Apples",
  "Bananas",
  "Strawberries",
  "Blueberries",
  "Mangoes",
  "Pineapple",
  "Lettuce",
  "Broccoli",
  "Paper Towels",
  "Dish Soap",
];

export const ShoppingList = () => {
  const [query, setQuery] = useState("");
  const [selectedItems, setSelectedItems] = useState([]);

  const filteredItems = items.filter((item) => 
    item.toLowerCase().includes(query.toLowerCase())
  );

  const toggleItem = (item) => {
    setSelectedItems((prev) =>
      prev.includes(item) ? prev.filter((i) => i !== item) : [...prev, item]
    );
  };

  return (
    <div className="container">
      <h1>Shopping List</h1>
      <form>
        <label htmlFor="search">Search for an item:</label>
        <input
          id="search"
          type="search"
          placeholder="Search..."
          aria-describedby="search-description"
          value={query}
          onChange={(e) => setQuery(e.target.value)}
        /> 
        <p id="search-description">Type to filter the list below:</p>
        <ul>
          {filteredItems.map((item) => {
            --fcc-editable-region--

            return (
              <li
                key={item}
              >
                <label>
                  <input
                    type="checkbox"
                    onChange={() => toggleItem(item)}
                  />
                  {item}
                </label>
              </li>
            );
            --fcc-editable-region--
          })}
        </ul>
      </form>
    </div>
  );
};

```
