---
id: 612e77aba7ca691f598feb02
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

Add two `meta` tags, one to optimize your page for mobile devices, and one to specify an accepted `charset` for the page.

# --hints--

应该在页面中添加两个 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

应该有两个 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

一个 `meta` 元素应将 `name` 设置为 `viewport`，并将 `content` 设置为 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一个 `meta` 元素应该将 `charset` 属性设置为 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
