---
id: 6733b5632fb854bc6c58dc0d
title: Hatua ya 2
challengeType: 1
dashedName: step-2
---

# --description--

Sasa tengeneza kitu kingine, chenye sifa tatu sawa, ili wakilishe aina nyingine ya lavender ya Kihispania inayoitwa `Pretty Polly`, na weka thamani yake katika kigezo kinachoitwa `prettyPolly`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `prettyPolly`.

```js
assert.exists(prettyPolly);
```

Unapaswa kuweka thamani ya kitu katika kigezo chako cha `prettyPolly`.

```js
assert.isObject(prettyPolly);
```

Kitu chako cha `prettyPolly` kinapaswa kuwa na sifa ya `commonName` iliyowekwa kuwa `"Spanish lavender"`.

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

Kitu chako cha `prettyPolly` kinapaswa kuwa na sifa ya `scientificName` iliyowekwa kuwa `"Lavandula stoechas"`.

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

Kitu chako cha `prettyPolly` kinapaswa kuwa na sifa ya `cultivar` iliyowekwa kuwa `"Pretty Polly"`.

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

Kitu chako cha `prettyPolly` kinapaswa kuwa na funguo tatu.

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```
