---
id: 673339d96c56d36e963e888e
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda orodha ya bustani ya mimea huku ukifanya mazoezi na vitu vya JavaScript vya Map na Set.

Ili kuwakilisha kila mmea katika orodha yako, utatumia vitu, kila kimoja chenye sifa tatu:

- `commonName`, jina la kawaida la mmea
- `scientificName`, jina la kisayansi (au jina la Kilatini) la mmea
- `cultivar`, aina ya mmea

Anza kwa kuunda kitu chenye sifa tatu zilizo hapo juu kwa `Spanish lavender` (jina la kisayansi `Lavandula stoechas`) la aina ya mmea inayoitwa `Ballerina` na weka thamani yake katika kigezo kinachoitwa `ballerina`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `ballerina`.

```js
assert.exists(ballerina);
```

Unapaswa kuweka thamani ya kitu katika kigezo chako cha `ballerina`.

```js
assert.isObject(ballerina);
```

Kitu chako cha `ballerina` kinapaswa kuwa na sifa ya `commonName` iliyowekwa kuwa `"Spanish lavender"`.

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

Kitu chako cha `ballerina` kinapaswa kuwa na sifa ya `scientificName` iliyowekwa kuwa `"Lavandula stoechas"`.

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

Kitu chako cha `ballerina` kinapaswa kuwa na sifa ya `cultivar` iliyowekwa kuwa `"Ballerina"`.

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

Kitu chako cha `ballerina` kinapaswa kuwa na funguo tatu.

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
