---
id: 6733b5e6f210c1bd1afdf8b3
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Vitu vya mimea vilivyobaki vimeongezwa kwa ajili yako. Sasa ni wakati wa kuunda orodha yako.

Kama vitu vya kawaida vya JavaScript, vitu vya `Map` vina mkusanyiko wa jozi za ufunguo-thamani. Ili kuunda kitu cha `Map`, unaweza kutumia kijenzi cha `Map`:

```js
const myMap = new Map();
```

Tangaza kigezo cha `catalog` na ukaweke thamani ya kitu cha `Map`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `catalog`.

```js
assert.exists(catalog);
```

Unapaswa kuweka thamani ya kitu cha `Map` kwenye kigezo chako cha `catalog`.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
