---
id: 6734f0ae9d1cf4dfd43a7805
title: Hatua 15
challengeType: 1
dashedName: step-15
---

# --description--

Ili kufuta kipengele kutoka kwenye kitu cha ramani, unaweza kutumia njia ya `delete`, ambayo hupokea ufunguo kama hoja na kufuta rekodi inayohusiana na ufunguo huo. Inarudisha thamani ya kweli au siyo kweli ikionyesha kama kufutwa kumefanikiwa.

Tengeneza kitendakazi kinachoitwa `removePlant` kinachopokea kitu cha mmea kama hoja yake na kukifuta kutoka kwenye orodha. Kifanye kirudishe `true` au `false` kulingana na kama mmea uliotajwa umefutwa kutoka kwenye orodha.

Baada ya hapo, andika matokeo ya kuita kitendakazi chako kipya ukiwa na mfuatano wa herufi `"ballerina"` kama hoja yake.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `removePlant`.

```js
assert.isFunction(removePlant);
```

Kitendakazi chako cha `removePlant` kinapaswa kupokea hoja moja.

```js
assert.lengthOf(removePlant, 1);
```

Kitendakazi chako cha `removePlant` kinapaswa kuita njia ya `delete` kwenye `catalog`.

```js
let flag = false;
const temp = catalog.delete;
try {
    catalog.delete = (p) => {flag = true}
    removePlant("");
    assert.isTrue(flag);
} finally {
    catalog.delete = temp;
}
```

`removePlant(item)` inapaswa kufuta `item` kutoka kwenye orodha yako.

```js
const testArr = [];
const temp = catalog.delete;
try {
    catalog.delete = (p) => {testArr.push(p)}
    removePlant("test");
    assert.deepEqual(testArr, ["test"]);
} finally {
    catalog.delete = temp;
}
```

Kitendakazi chako cha `removePlant` kinapaswa kurudisha thamani ya kweli au siyo kweli ikionyesha kama kufutwa kumefanikiwa.

```js
try {
    assert.isFalse(removePlant("test"));
    assert.isTrue(removePlant(royalCrown));
} finally {
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 })
}
```

Unapaswa kuandika `removePlant("ballerina")` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*removePlant\s*\(\s*("|')ballerina\1\s*\)\s*\)/);
```



# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

--fcc-editable-region--

--fcc-editable-region--

const sellPlants = (plant, size, potsNo) => {
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```
