---
id: 6734c10a327760a665b7d5b9
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Katika hatua chache zijazo, utafanya kazi kwenye kitendakazi kitakachokuwezesha kusahihisha katalogi wakati unauza mimea.

Tengeneza kitendakazi kinachoitwa `sellPlants` kinachopokea hoja tatu: kitu cha mmea, ukubwa wa sufuria na idadi ya sufuria za kuuza, kwa mpangilio huu.

Wakati idadi iliyotajwa ya sufuria ni kubwa kuliko sufuria zilizopo, fanya kitendakazi chako kirudishe `Not enough <pot-size> size pots for <item-name>. Only <pot-number> left.`. Badilisha `<pot-size>` na ukubwa wa sufuria, `<pot-number>` na sufuria zilizobaki kwa ukubwa huo, na `<item-name>` na jina la kisayansi la mmea likifuatiwa na nafasi na jina la cultivar lililozungushwa na nukuu moja moja.

Ili kujaribu kama kila kitu kinafanya kazi, andika matokeo ya kuita kitendakazi chako cha `sellPlants` kwa kitu cha `ballerina`, `"small"`, na `25` kwenye konsoli.

# --hints--

Haupasi kubadilisha ramani yako ya `catalog`.

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 20, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

Unapaswa kuwa na kitendakazi kinachoitwa `sellPlants`.

```js
assert.isFunction(sellPlants);
```

Kitendakazi chako cha `sellPlants` kinapaswa kupokea hoja tatu.

```js
assert.lengthOf(sellPlants, 3);
```

`sellPlants(ballerina, "small", 25)` inapaswa kurudisha `Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.`

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")

assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

Unapaswa kuandika `sellPlants(ballerina, "small", 25)` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*25\s*\)\s*\)/);
```

`sellPlants(ballerina, "small", 10)` inapaswa kurudisha `undefined`.

```js
assert.isUndefined(sellPlants(ballerina, "small", 10));
// prevent hardcoding
assert.isUndefined(sellPlants(hidcote, "large", 5));
assert.isUndefined(sellPlants(royalCrown, "medium", 5));
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog)
```
