---
id: 67349d6e726b2c3c7f7301e8
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Sifa ya `size` ya kitu cha Map inarudisha idadi ya vipengele vilivyomo kwenye ramani. Andika `size` ya `catalog` yako kwenye konsoli.

# --hints--

Unapaswa kuandika sifa ya `size` ya ramani yako ya `catalog` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\.\s*size\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
console.log(catalog);

--fcc-editable-region--
```
