---
id: 6734a772196245601e4b80e0
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Unapohitaji kupata kipengele kutoka kwenye ramani, unaweza kutumia njia ya `get`. Inachukua ufunguo kama hoja yake, na inarudisha thamani inayohusiana.

Tumia njia ya `get` kupata thamani ya ufunguo wa `ballerina` kutoka kwenye katalogi yako na weka thamani hiyo kwenye kigezo kinachoitwa `getBallerina`. Kisha, andika kigezo chako kipya kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `getBallerina`.

```js
assert.exists(getBallerina);
```

Kigezo chako cha `getBallerina` kinapaswa kuwa na thamani ya `catalog.get(ballerina)`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+getBallerina\s*=\s*catalog\s*\.\s*get\s*\(\s*ballerina\s*\)/);
```

Unapaswa kuandika kigezo chako cha `getBallerina` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*getBallerina\s*\)/);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
```
