---
id: 6734a2385c52c54d916a8b86
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Ongeza jozi nyingine ya ufunguo-thamani kwenye katalogi yako. Tumia tofauti yako ya `prettyPolly` kama ufunguo na fanya thamani iwe kitu chenye sifa za `small`, `medium`, na `large` ambazo thamani zao ni nambari kamili tena.

# --hints--

Unapaswa kuweka kitu chako cha `prettyPolly` kama ufunguo wa ramani yako ya `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Thamani ya ufunguo wa `prettyPolly` katika ramani ya `catalog` inapaswa kuwa kitu chenye sifa za `small`, `medium`, na `large`.

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

Sifa ya `small` ya kitu kinachohusiana na ufunguo wa `prettyPolly` katika ramani ya `catalog` inapaswa kuwa nambari kamili chanya.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

Sifa ya `medium` ya kitu kinachohusiana na ufunguo wa `prettyPolly` katika ramani ya `catalog` inapaswa kuwa nambari kamili chanya.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

Sifa ya `large` ya kitu kinachohusiana na ufunguo wa `prettyPolly` katika ramani ya `catalog` inapaswa kuwa nambari kamili chanya.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Bado unapaswa kuwa na ufunguo wa `ballerina` katika ramani yako ya `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Ramani yako ya `catalog` inapaswa kuwa na vipengele viwili.

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
