---
id: 6734a606e5968a5a19b77171
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Endelea kuongeza mimea iliyobaki kwenye orodha yako kama ulivyofanya katika hatua zilizopita.

# --hints--

Unapaswa kuweka kitu chako cha `willowVale` kama ufunguo wa ramani yako ya `catalog`.

```js
assert.isTrue(catalog.has(willowVale));
```

Thamani ya ufunguo wa `willowVale` katika ramani ya `catalog` inapaswa kuwa kitu chenye vigezo `small`, `medium`, na `large`, kila kikiwa na thamani ya nambari kamili chanya.

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Unapaswa kuweka kitu chako cha `hidcote` kama ufunguo wa ramani yako ya `catalog`.

```js
assert.isTrue(catalog.has(hidcote));
```

Thamani ya ufunguo wa `hidcote` katika ramani ya `catalog` inapaswa kuwa kitu chenye vigezo `small`, `medium`, na `large`, kila kikiwa na thamani ya nambari kamili chanya.

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Unapaswa kuweka kitu chako cha `imperialGem` kama ufunguo wa ramani yako ya `catalog`.

```js
assert.isTrue(catalog.has(imperialGem));
```

Thamani ya ufunguo wa `imperialGem` katika ramani ya `catalog` inapaswa kuwa kitu chenye vigezo `small`, `medium`, na `large`, kila kikiwa na thamani ya nambari kamili chanya.

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Unapaswa kuweka kitu chako cha `royalCrown` kama ufunguo wa ramani yako ya `catalog`.

```js
assert.isTrue(catalog.has(royalCrown));
```

Thamani ya ufunguo wa `royalCrown` katika ramani ya `catalog` inapaswa kuwa kitu chenye vigezo `small`, `medium`, na `large`, kila kikiwa na thamani ya nambari kamili chanya.

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Bado unapaswa kuwa na ufunguo wa `ballerina` katika ramani yako ya `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Bado unapaswa kuwa na ufunguo wa `prettyPolly` katika ramani yako ya `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Ramani yako ya `catalog` inapaswa kuwa na vipengele sita.

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
