---
id: 67349ba230cff237c0f2bdf2
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Vitu vya ramani hukumbuka mpangilio ambao funguo zinaingizwa kwenye ramani. Ili kuweka jozi ya funguo-thamani katika kitu chako cha ramani, unaweza kutumia njia ya `set`. Ifuatayo inaweka funguo `"name"` kwa thamani `"Jerry"` katika kitu cha ramani cha `myMap`.

```js
myMap.set("name", "Jerry");
```

Weka funguo `"ballerina"` kwa thamani ya kitu chako cha `ballerina` katika ramani yako ya `catalog`. Kisha, andika orodha yako kwenye konsoli.

Kumbuka kwamba ili kuona kitu chako cha ramani ipasavyo utahitaji kufungua konsoli ya kivinjari.

# --hints--

Unapaswa kuweka funguo `"ballerina"` kwa thamani ya kitu chako cha `ballerina` katika ramani yako ya `catalog`.

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

Ramani yako ya `catalog` inapaswa kuwa na jozi moja tu ya funguo-thamani.

```js
assert.lengthOf(catalog, 1);
```

Unapaswa kuandika ramani yako ya `catalog` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```
