---
id: 6734a13a29acdf47f2b8e161
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Tofauti na vitu vya kawaida, ambavyo ufunguo wake lazima daima uwe mfuatano wa herufi, vitu vya ramani vinakuwezesha kutumia aina yoyote ya data kama funguo. Katika katalogi yako, utatumia jozi za ufunguo-thamani za kitu-kitu, ukihusisha kila kitu chako cha mmea na kitu chenye idadi ya sufuria zinazopatikana kwa ukubwa tofauti.

Badilisha wito wako wa `set` ili utumie kitu chako cha `ballerina` kama funguo. Kwa thamani, tumia kitu chenye sifa tatu, `small`, `medium`, na `large`, kila moja ikiwa na thamani ya nambari kamili chanya.

# --hints--

Unapaswa kuweka kitu chako cha `ballerina` kama funguo la ramani yako ya `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Thamani ya funguo la `ballerina` katika ramani ya `catalog` inapaswa kuwa kitu chenye sifa `small`, `medium`, na `large`.

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

Sifa ya `small` ya kitu kinachohusishwa na funguo la `ballerina` katika ramani ya `catalog` inapaswa kuwa nambari kamili chanya.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

Sifa ya `medium` ya kitu kinachohusishwa na funguo la `ballerina` katika ramani ya `catalog` inapaswa kuwa nambari kamili chanya.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

Sifa ya `large` ya kitu kinachohusishwa na funguo la `ballerina` katika ramani ya `catalog` inapaswa kuwa nambari kamili chanya.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Ramani yako ya `catalog` inapaswa kuwa na jozi moja tu ya ufunguo-thamani.

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
