---
id: 67350acfd2006f2777295b96
title: Hatua 19
challengeType: 1
dashedName: step-19
---

# --description--

Njia ya `keys` inapatikana moja kwa moja kwenye kitu cha `Map` na inarudisha `MapIterator` inayokuwezesha kurudia juu ya funguo za ramani yako.

Tengeneza kitendakazi kinachoitwa `displayCatalog`. Ndani ya kitendakazi chako, tengeneza mzunguko wa `for...of` unaorudia juu ya funguo katika `catalog` na kuandika kila moja yao kwenye konsoli.

Kisha, fanya wito la `displayCatalog` na uone matokeo kwenye konsoli.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `displayCatalog`.

```js
assert.isFunction(displayCatalog);
```

Unapaswa kuita kitendakazi chako cha `displayCatalog`.

```js
assert.match(__helpers.removeJSComments(code), /displayCatalog\s*\(\s*\)/);
```

Kitendakazi chako cha `displayCatalog` kinapaswa kuita njia ya `keys` kwenye `catalog`.

```js
const catalogCopy = new Map(catalog);
let flag = false;
const temp = catalog.keys;
try {
    catalog.keys = () => {
        flag = true;
        return catalogCopy.keys()
    };
    displayCatalog();    
    assert.isTrue(flag);
} finally {
    catalog.keys = temp;
}
```

Kitendakazi chako cha `displayCatalog` kinapaswa kuandika kila funguo iliyomo katika `catalog` kwenye konsoli.

```js
const testArr = [];
const temp = console.log;
try {
    console.log = obj => {testArr.push(obj)};
    displayCatalog();    
    assert.isNotEmpty(testArr);
    const k = catalog.keys();
    for (let i = 0; i < catalog.size; i++) {
        assert.strictEqual(testArr[i], k.next().value)
    }
} finally {
    console.log = temp;
}
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

const removePlant = plant => catalog.delete(plant);

--fcc-editable-region--

--fcc-editable-region--
```
