---
id: bad87fee1348bd9aedf08807
title: Googleフォントのインポート
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9MRsJ'
forumTopicId: 18200
dashedName: import-a-google-font
---

# --description--

ほとんどのオペレーティングシステムにある一般的なフォントを指定することに加えて、ウェブサイトで使うための標準外のカスタムウェブフォントを指定することもできます。インターネット上には多くのウェブフォントのソースがあります。この例ではGoogle Fontsライブラリに焦点を当てます。

Google Fontsは無料のウェブフォントライブラリで、フォントのURLを参照することでCSSで使うことができます。

では、Googleフォントをインポートして適用してみましょう（国によってはGoogleがブロックされている場合があるので、その場合はこのチャレンジをスキップしてください）。

Googleフォントをインポートするには、Google FontsライブラリからフォントのURLをコピーしてHTMLに貼り付けます。このチャレンジでは`Lobster`フォントをインポートします。次のコードスニペットをコピーして、コードエディターの先頭（`style`要素の前）に貼り付けてください。

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
```

これでCSSで`Lobster`をFAMILY_NAMEとして使うことで`Lobster`フォントを使用できます。以下の例のようにしてください。

```css
font-family: FAMILY_NAME, GENERIC_NAME;
```

GENERIC_NAMEはオプションで、指定したフォントが利用できない場合のフォールバックフォントです。これは次のチャレンジで扱います。

ファミリ名は大文字小文字を区別し、名前にスペースがある場合は引用符で囲む必要があります。例えば、`"Open Sans"`フォントを使うには引用符が必要ですが、`Lobster`フォントを使う場合は不要です。

# --instructions--

`Lobster`フォントをウェブページにインポートしてください。その後、要素セレクターを使って`Lobster`要素の`font-family`を`h2`に設定してください。

# --hints--

`Lobster`フォントをインポートする必要があります。

```js
assert.exists(document.querySelector('link[href*="googleapis" i]'));
```

`h2`要素は`Lobster`フォントを使うべきです。

```js
const h2 = document.querySelector('h2'); 
const fontFamily = window.getComputedStyle(h2)['font-family']; 
assert.match(fontFamily, /lobster/i);
```

フォントを変更するには`h2`要素セレクターのみを使うべきです。

```js
assert.match(__helpers.removeHtmlComments(code), /\s*[^\.]h2\s*\{\s*font-family\s*:\s*('|"|)Lobster\1\s*(,\s*('|"|)[a-z -]+\3\s*)?(;\s*\}|\})/gi);
```

`p`要素は引き続き`monospace`フォントを使うべきです。

```js
const paragraphElement = document.querySelector('p');
const fontFamily = window.getComputedStyle(paragraphElement)['font-family']; 
assert.match(fontFamily, /monospace/i);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }

  h2 {
    font-family: Lobster;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
