---
id: bad87fee1348bd9aede08718
title: 要素の色付けにRGB値を使う
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkp2fr'
forumTopicId: 18369
dashedName: use-rgb-values-to-color-elements
---

# --description--

CSSで色を表すもう一つの方法は`RGB`値を使うことです。

黒の`RGB`値はこのようになります：

```css
rgb(0, 0, 0)
```

白の`RGB`値はこのようになります：

```css
rgb(255, 255, 255)
```

16進コードのように6桁の16進数を使う代わりに、`RGB`では各色の明るさを0から255の数字で指定します。

計算すると、1色あたりの2桁は16×16で256の値になります。つまり、0から数え始める`RGB`は16進コードと全く同じ数の値を持っています。

以下は`body`の背景をRGBコードでオレンジに変える例です。

```css
body {
  background-color: rgb(255, 165, 0);
}
```

# --instructions--

`body`要素の背景色の16進コードを黒のRGB値：`rgb(0, 0, 0)`に置き換えましょう。

# --hints--

`body`要素は黒い背景色になっているはずです。

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

`rgb`を使って`body`要素に黒い背景色を付けましょう。

```js
assert.match(__helpers.removeCssComments(code), /rgb\s*\(\s*0\s*,\s*0\s*,\s*0\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #F00;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: rgb(0, 0, 0);
  }
</style>
```
