---
id: bad87fee1348bd9aedf08805
title: CSSセレクターを使って要素にスタイルを適用する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKMBT2'
forumTopicId: 18349
dashedName: use-css-selectors-to-style-elements
---

# --description--

CSSには、ページ上の要素の見た目を変えるために使える数百ものCSSプロパティがあります。

`<h2 style="color: red;">CatPhotoApp</h2>`に入力したときは、インラインCSSを使ってその個々の`h2`要素にスタイルを適用していました。インラインCSSとはCascading Style Sheetsのことです。

これは要素のスタイルを指定する方法の一つですが、CSSを適用するもっと良い方法があります。

コードの先頭に、次のような`style`ブロックを作成してください。

```html
<style>
</style>
```

そのスタイルブロックの中で、すべての`h2`要素に対する<dfn>CSSセレクター</dfn>を作成できます。例えば、すべての`h2`要素を赤にしたい場合は、次のようなスタイルルールを追加します。

```html
<style>
  h2 {
    color: red;
  }
</style>
```

各要素のスタイルルールの周りに開き波括弧（`{`）と閉じ波括弧（`}`）の両方があることが重要です。また、要素のスタイル定義が開きと閉じのstyleタグの間にあることを確認してください。最後に、各スタイルルールの末尾にセミコロンを付けることも忘れないでください。

# --instructions--

`h2`要素のstyle属性を削除し、代わりにCSSの`style`ブロックを作成してください。すべての`h2`要素を青にするために必要なCSSを追加してください。

# --hints--

`style`属性は`h2`要素から削除してください。

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

`style`要素を作成してください。

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
assert.isAtLeast(document.querySelectorAll('style:not(.fcc-hide-header)').length, 1);
```

`h2`要素は青色にしてください。

```js
const h2Element = document.querySelector('h2');
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

スタイルシートの`h2`宣言は、セミコロンと閉じ括弧があって有効である必要があります。

```js
assert.match(__helpers.removeCssComments(code), /h2\s*\{\s*color\s*:.*;\s*\}/g);
```

すべての`style`要素は有効で終了タグがある必要があります。

```js
assert.match(__helpers.removeHtmlComments(code), /<\/style>/g);

const closingElementLength = __helpers.removeHtmlComments(code).match(/<\/style>/g).length;
const openingElementsLength = __helpers.removeHtmlComments(code).match(/<style((\s)*((type|media|scoped|title|disabled)="[^"]*")?(\s)*)*>/g).length;
assert.strictEqual(closingElementLength, openingElementsLength);
```

# --seed--

## --seed-contents--

```html
<h2 style="color: red;">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  h2 {
    color: blue;
  }
</style>
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
