---
id: bad87fee1348bd9aedf08756
title: スタイルの優先順位をつける
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ8wnHv'
forumTopicId: 18258
dashedName: prioritize-one-style-over-another
---

# --description--

時々、HTML要素に複数のスタイルが適用されて互いに競合することがあります。

例えば、`h1` 要素は同時に緑色とピンク色の両方にはなれません。

テキストをピンクにするクラスを作成し、それを要素に適用したらどうなるか見てみましょう。クラスは`body` 要素の`color: green;` CSSプロパティを*オーバーライド*できるでしょうか？

# --instructions--

`pink-text` という名前のCSSクラスを作成し、要素の色をピンクにしてください。

`h1` 要素に`pink-text` クラスを付けてください。

# --hints--

`h1` 要素には`pink-text` クラスを付けてください。

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

`<style>` には`pink-text` CSSクラスを付けて、`color` を変更してください。

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*\{\s*color\s*:\s*.+\s*;?\s*\}/g);
```

`h1` 要素はピンク色になっているはずです。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
</style>
<h1>Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```
