---
id: bad87fee1348bd9aedf07756
title: Important を使って他のすべてのスタイルをオーバーライドする
challengeType: 0
videoUrl: 'https://scrimba.com/c/cm24rcp'
forumTopicId: 18249
dashedName: override-all-other-styles-by-using-important
---

# --description--

やった！インラインスタイルが`style`要素内のすべてのCSS宣言をオーバーライドすることを証明しました。

でも待ってください。CSSをオーバーライドする最後の方法があります。これは最も強力な方法です。でもその前に、なぜCSSをオーバーライドしたいのかを話しましょう。

多くの場合、CSSライブラリを使います。これらは意図せず自分のCSSをオーバーライドしてしまうことがあります。だから、特定の要素に絶対に特定のCSSを適用したい場合は、`!important`を使うことができます。

`pink-text`クラス宣言までさかのぼりましょう。`pink-text`クラスは後からのクラス宣言、id宣言、インラインスタイルによってオーバーライドされていましたね。

# --instructions--

ピンク色のテキスト要素のcolor宣言にキーワード`!important`を追加して、`h1`要素が確実にピンクになるようにしましょう。

やり方の例は以下の通りです。

```css
color: red !important;
```

# --hints--

あなたの`h1`要素はクラス`pink-text`を持っているはずです。

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

あなたの`h1`要素はクラス`blue-text`を持っているはずです。

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

あなたの`h1`要素は`id`が`orange-text`であるはずです。

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

あなたの`h1`要素はインラインスタイルで`color: white`を持っているはずです。

```js
const commentessCode = __helpers.removeHtmlComments(code);
assert.match(commentessCode, /<h1.*style/gi);
assert.match(commentessCode, /<h1.*style.*color\s*?:/gi);
```

あなたの`pink-text`クラス宣言は他のすべての宣言をオーバーライドするために`!important`キーワードを持っているはずです。

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*?\{[\s\S]*?color:.*pink.*!important\s*;?[^\.]*\}/g);
```

あなたの`h1`要素はピンク色であるはずです。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink !important;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
