---
id: bad87fee1348bd9aecf08806
title: CSSクラスを使って要素にスタイルを適用する
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MvDtV'
forumTopicId: 18337
dashedName: use-a-css-class-to-style-an-element
---

# --description--

クラスはHTML要素に追加できる再利用可能なスタイルです。

以下はCSSクラスの宣言例です：

```html
<style>
  .blue-text {
    color: blue;
  }
</style>
```

`blue-text`タグ内に`<style>`というCSSクラスを作成したのがわかります。HTML要素にクラスを適用するにはこのようにします：`<h2 class="blue-text">CatPhotoApp</h2>`。CSSの`style`要素ではクラス名はピリオドで始まりますが、HTML要素のclass属性ではピリオドは含みません。

# --instructions--

`style`要素内で、`h2`セレクターを`.red-text`に変更し、colorの値を`blue`から`red`に更新してください。

`h2`要素に`class`属性を付け、その値を`red-text`にしてください。

# --hints--

`h2`要素は赤色である必要があります。

```js
const h2Element = document.querySelector('h2'); 
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

`h2`要素には`red-text`クラスが付いている必要があります。

```js
assert.isTrue(document.querySelector('h2').classList.contains('red-text'));
```

スタイルシートには`red-text`クラスが宣言され、そのcolorは`red`に設定されている必要があります。

```js
assert.match(__helpers.removeCssComments(code), /\.red-text\s*\{\s*color\s*:\s*red;?\s*\}/g);
```

`style="color: red"`要素に`h2`のようなインラインスタイル宣言を使わないでください。

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

# --seed--

## --seed-contents--

```html
<style>
  h2 {
    color: blue;
  }
</style>

<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  .red-text {
    color: red;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
