---
id: bad87fee1348bd9aedf06756
title: インラインスタイルでクラス宣言をオーバーライドする
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDRha'
forumTopicId: 18252
dashedName: override-class-declarations-with-inline-styles
---

# --description--

id宣言が、`style`要素のCSS内でどこに宣言されていてもクラス宣言をオーバーライドすることは証明済みです。

CSSをオーバーライドする他の方法もあります。インラインスタイルを覚えていますか？

# --instructions--

インラインスタイルを使って、`h1`要素を白にしてみてください。インラインスタイルは次のように書きます。

```html
<h1 style="color: green;">
```

`blue-text`要素には`pink-text`と`h1`のクラスを残してください。

# --hints--

`h1`要素には`pink-text`のクラスが付いているはずです。

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

`h1`要素には`blue-text`のクラスが付いているはずです。

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

`h1`要素には`orange-text`のidが付いているはずです。

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

`h1`要素にはインラインスタイルが付いているはずです。

```js
assert.exists(document.querySelector('h1[style]'));
```

`h1`要素は白であるべきです。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
