---
id: bad82fee1322bd9aedf08721
title: 絶対単位と相対単位の理解
challengeType: 0
videoUrl: 'https://scrimba.com/c/cN66JSL'
forumTopicId: 301089
dashedName: understand-absolute-versus-relative-units
---

# --description--

これまでのチャレンジでは、要素のマージンやパディングをピクセル（`px`）で設定してきました。ピクセルは長さの単位の一種で、ブラウザーにアイテムのサイズや間隔を指示します。`px`に加えて、CSSにはさまざまな長さの単位オプションがあります。

長さの単位には大きく分けて絶対単位と相対単位があります。絶対単位は物理的な長さの単位に結びついています。例えば、`in`や`mm`はそれぞれインチやミリメートルを指します。絶対長さの単位は画面上の実際の測定値に近いですが、画面の解像度によって多少の違いがあります。

相対単位は、`em`や`rem`のように、他の長さの値に対して相対的です。例えば、`em`は要素のフォントサイズに基づいています。これを`font-size`プロパティ自体に使う場合は、親要素の`font-size`に対して相対的です。

**注意：** ビューポートのサイズに結びついた相対単位のオプションがいくつかあります。これらはレスポンシブウェブデザインの原則のセクションで扱います。

# --instructions--

クラス`padding`の要素に`red-box`プロパティを追加し、`1.5em`に設定してください。

# --hints--

クラス`red-box`には`padding`プロパティが必要です。

```js
const redBox =document.querySelector('.red-box'); 
const style = window.getComputedStyle(redBox); 

assert.notEqual(style['padding-top'], '0px');
assert.notEqual(style['padding-right'], '0px');
assert.notEqual(style['padding-bottom'], '0px');
assert.notEqual(style['padding-left'], '0px');
```

クラス`red-box`は要素に1.5emの`padding`を与えるべきです。

```js
assert.match(__helpers.removeCssComments(code), /\.red-box\s*?{[\s\S]*padding\s*:\s*?1\.5em/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;

  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;
    padding: 1.5em;
  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```
