---
id: bad87fee1348bd8aedf06756
title: ID属性のスタイル指定でクラス宣言をオーバーライドする
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpDhB'
forumTopicId: 18251
dashedName: override-class-declarations-by-styling-id-attributes
---

# --description--

ブラウザがCSSを宣言された順に上から下へ読み込むことを証明しました。つまり、競合があった場合は、最後に宣言されたCSSが使われます。`blue-text`要素のクラスで`pink-text`を`h1`の前に置いても、使われるのは宣言順であって、使われる順序ではないことに注意してください！

でも、まだ終わりではありません。CSSをオーバーライドする他の方法もあります。id属性を覚えていますか？

`pink-text`と`blue-text`のクラスをオーバーライドして、`h1`要素をオレンジ色にしましょう。`h1`要素にidを付けて、そのidをスタイル指定することで実現します。

# --instructions--

`h1`要素に`id`属性として`orange-text`を付けてください。idスタイルは次のように書きます：

```html
<h1 id="orange-text">
```

`blue-text`と`pink-text`のクラスは`h1`要素に残してください。

`orange-text`要素に`style`のid用のCSS宣言を作成してください。例は次の通りです：

```css
#brown-text {
  color: brown;
}
```

**注意：** このCSSを`pink-text`クラスの上でも下でもどちらに宣言しても問題ありません。`id`属性が常に優先されるためです。

# --hints--

`h1`要素には`pink-text`クラスが付いているはずです。

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

`h1`要素には`blue-text`クラスが付いているはずです。

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

`h1`要素には`orange-text`のidが付いているはずです。

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'),'orange-text');
```

`h1`要素は1つだけであるべきです。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`orange-text`のidにはCSS宣言があるはずです。

```js
assert.match(__helpers.removeCssComments(code), /#orange-text\s*{/gi);
```

`h1`には`style`属性が付いていないはずです。

```js
assert.notMatch(__helpers.removeHtmlComments(code), /<h1.*style.*>/gi);
```

`h1`要素はオレンジ色であるべきです。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 165, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
  #orange-text {
    color: orange;
  }  
</style>
<h1 id="orange-text"  class="pink-text blue-text">Hello World!</h1>
```
