---
id: bad87fee1348bd9aedf08726
title: 特定の色に16進コードを使う
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

CSSで色を表す方法が他にもあることをご存知ですか？そのうちの一つが16進コード、略して16進数コードです。

私たちは通常、0から9の記号を使う<dfn>10進数</dfn>（基数10の数字）を使います。<dfn>16進数</dfn>（または<dfn>hex</dfn>）は基数16の数字です。つまり16種類の異なる記号を使います。10進数と同様に、0～9は0から9の値を表します。さらにA、B、C、D、E、Fは10から15の値を表します。合わせて0からFまでが16進数の1桁を表し、合計16通りの値が可能です。<a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">16進数についてはこちら</a>をご覧ください。

CSSでは、赤（R）、緑（G）、青（B）の各成分に対して2桁ずつ、合計6桁の16進数で色を表せます。例えば、`#000000` は黒で、最も低い値でもあります。<a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">RGBカラーモデルについてはこちら</a>をご覧ください。

```css
body {
  color: #000000;
}
```

# --instructions--

`black`要素のbackground-colorにある`body`という単語を、その16進数コード表現である`#000000`に置き換えてください。

# --hints--

`body`要素は`background-color`の色であるべきです。

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

色「黒」の単語`black`の代わりに16進コードを使うべきです。

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
