---
id: bad87fee1348bd9aedf08736
title: HTMLのbody要素にスタイルを適用する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB77PHW'
forumTopicId: 18313
dashedName: style-the-html-body-element
---

# --description--

ここから新たに始めて、CSSの継承について話しましょう。

すべてのHTMLページには`body`要素があります。

# --instructions--

`body`要素がここに存在することを、`background-color`を黒に設定することで証明できます。

これを行うには、`style`要素に次の内容を追加してください。

```css
body {
  background-color: black;
}
```

# --hints--

`body`要素には`background-color`を黒に設定してください。

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color'];

assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

CSSルールは、開き括弧と閉じ括弧の両方を正しく書いてください。

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

CSSルールはセミコロンで終わらせてください。

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

# --seed--

## --seed-contents--

```html
<style>

</style>
```

# --solutions--

```html
<style>
body {
  background-color: black;
}
</style>
```
