---
id: bad87fee1348bd9aedf08808
title: フォントのフォールバック方法を指定する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cpVKBfQ'
forumTopicId: 18304
dashedName: specify-how-fonts-should-degrade
---

# --description--

すべてのブラウザで利用可能なデフォルトのフォントがいくつかあります。これらの総称フォントファミリーには`monospace`、`serif`、`sans-serif`が含まれます。

あるフォントが利用できない場合、ブラウザに別のフォントへ「フォールバック」するよう指示できます。

例えば、ある要素に`Helvetica`フォントを使い、`sans-serif`が利用できない場合は`Helvetica`フォントにフォールバックさせたい場合、次のように指定します。

```css
p {
  font-family: Helvetica, sans-serif;
}
```

総称フォントファミリー名は大文字・小文字を区別しません。また、CSSのキーワードなので引用符は不要です。

# --instructions--

まず、`monospace`フォントを`h2`要素に適用して、`Lobster`と`monospace`の2つのフォントを指定してください。

前のチャレンジでは、`Lobster`フォントを`link`タグでインポートしました。今回はGoogle Fontsからの`Lobster`フォントのインポートを、これまでに学んだHTMLコメントを使ってコメントアウトし、利用できないようにしてください。そうすると`h2`要素が`monospace`フォントにフォールバックするのがわかります。

**注意:** お使いのコンピューターに`Lobster`フォントがインストールされている場合、ブラウザがそのフォントを見つけられるためフォールバックは起きません。

# --hints--

h2要素には`Lobster`フォントを使うべきです。

```js
const h2Element = document.querySelector('h2');
const fontFamily = window.getComputedStyle(h2Element)['font-family']; 
assert.match(fontFamily, /^"?lobster/i);
```

h2要素は`monospace`が利用できない場合、`Lobster`フォントにフォールバックするべきです。

```js
assert.match(__helpers.removeCssComments(code), /\s*h2\s*\{\s*font-family\s*\:\s*(\'|"|)Lobster\1\s*,\s*monospace\s*;?\s*\}/gi);
```

Googleからの`Lobster`フォントの呼び出しは、`<!--`を前に置いてコメントアウトしてください。

```js
assert.match(code, /<!--[^fc]/gi);
```

コメントは`-->`で閉じてください。

```js
assert.match(code, /[^fc]-->/gi);
```

# --seed--

## --seed-contents--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<!--<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">-->
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster, monospace;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
