---
id: bad88fee1348bd9aedf08825
title: 要素のパディングを調整する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cED8ZC2'
forumTopicId: 301083
dashedName: adjust-the-padding-of-an-element
---

# --description--

ここで少しの間、Cat Photo App を置いておいて、HTMLのスタイリングについてもっと学びましょう。

すでに気づいているかもしれませんが、すべてのHTML要素は基本的に小さな長方形です。

3つの重要なプロパティが、それぞれのHTML要素を囲むスペースを制御しています：`padding`、`border`、そして`margin`です。

要素の`padding`は、その要素のコンテンツと`border`の間のスペースの量を制御します。

ここでは、青いボックスと赤いボックスが黄色いボックスの中にネストされているのが見えます。赤いボックスは青いボックスよりも多くの`padding`を持っていることに注意してください。

青いボックスの`padding`を増やすと、テキストとその周りのボーダーとの距離（`padding`）が増えます。

# --instructions--

青いボックスの`padding`を赤いボックスのものと同じに変更してください。

# --hints--

`blue-box`クラスは要素に`20px`の`padding`を与えるはずです。

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
