---
id: bad87fee1348bd9aedf08746
title: body要素からスタイルを継承する
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bmdtR'
forumTopicId: 18204
dashedName: inherit-styles-from-the-body-element
---

# --description--

すべてのHTMLページに`body`要素があり、その`body`要素もCSSでスタイルを設定できることを証明しました。

`body`要素は他のHTML要素と同じようにスタイルを設定でき、ほかのすべての要素は`body`要素のスタイルを継承することを覚えておいてください。

# --instructions--

まず、テキストが`h1`の`Hello World`要素を作成してください。

次に、ページ上のすべての要素に`green`の色を設定するために、`color: green;`要素のスタイル宣言に`body`を追加しましょう。

最後に、`body`要素に`monospace`のフォントファミリーを設定するために、`font-family: monospace;`要素のスタイル宣言に`body`を追加してください。

# --hints--

`h1`要素を作成してください。

```js
assert.isNotEmpty(document.querySelectorAll('h1'));
```

`h1`要素のテキストは`Hello World`にしてください。

```js
assert.match(
  document.querySelector('h1').textContent,
  /hello world/i
);
```

`h1`要素には終了タグが必要です。

```js
const commentlessCode = __helpers.removeHtmlComments(code);
assert.match(commentlessCode, /<\/h1>/g);
assert.match(commentlessCode, /<h1/g);
assert.lengthOf(commentlessCode.match(/<\/h1>/g), commentlessCode.match(/<h1/g).length);
```

`body`要素には`color`プロパティとして`green`を設定してください。

```js
const bodyElement = document.querySelector('body');
const color = window.getComputedStyle(bodyElement)['color']; 
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

`body`要素には`font-family`プロパティとして`monospace`を設定してください。

```js
const bodyElement = document.querySelector('body');
const fontFamily = window.getComputedStyle(bodyElement)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

`h1`要素は`monospace`要素からフォント`body`を継承します。

```js
const h1Element = document.querySelector('h1');
const fontFamily = window.getComputedStyle(h1Element)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

`h1`要素は`green`要素から色`body`を継承します。

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }

</style>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }

</style>
<h1>Hello World!</h1>
```
