---
id: bad87fee1348bd9aedf08719
title: 短縮された16進コードを使う
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpKAm'
forumTopicId: 18338
dashedName: use-abbreviated-hex-code
---

# --description--

多くの人は、1600万色以上の色の可能性に圧倒されてしまいます。そして16進コードを覚えるのは難しいです。幸い、短縮することができます。

例えば、赤の16進コード`#FF0000`は`#F00`のように短縮できます。この短縮形は赤、緑、青それぞれに1桁ずつ割り当てられています。

これにより、可能な色の総数は約4,000色に減ります。しかしブラウザは`#FF0000`と`#F00`をまったく同じ色として解釈します。

# --instructions--

さあ、短縮された16進コードを使って正しい要素に色を付けてみましょう。

<table><tbody><tr><th>色</th><th>短縮16進コード</th></tr><tr><td>シアン</td><td><code>#0FF</code></td></tr><tr><td>緑</td><td><code>#0F0</code></td></tr><tr><td>赤</td><td><code>#F00</code></td></tr><tr><td>フクシア</td><td><code>#F0F</code></td></tr></tbody></table>

# --hints--

テキストが`h1`の`I am red!`要素には`color`の赤を指定してください。

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

赤の色には16進コード`#FF0000`の代わりに短縮された16進コードを使う必要があります。

```js
assert.match(code,/\.red-text\s*?{\s*?color\s*:\s*?#F00\s*?;?\s*?}/gi);
```

テキストが`h1`の`I am green!`要素には`color`の緑を指定してください。

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

緑の色には16進コード`#00FF00`の代わりに短縮された16進コードを使う必要があります。

```js
assert.match(__helpers.removeCssComments(code), /\.green-text\s*?{\s*?color\s*:\s*?#0F0\s*?;?\s*?}/gi);
```

テキストが`h1`の`I am cyan!`要素には`color`のシアンを指定してください。

```js
const cyanText = document.querySelector('.cyan-text');
const color = window.getComputedStyle(cyanText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 255)');
```

シアンの色には16進コード`#00FFFF`の代わりに短縮された16進コードを使う必要があります。

```js
assert.match(__helpers.removeCssComments(code), /\.cyan-text\s*?{\s*?color\s*:\s*?#0FF\s*?;?\s*?}/gi);
```

テキストが`h1`の`I am fuchsia!`要素には`color`のフクシアを指定してください。

```js
const fuchsiaText = document.querySelector('.fuchsia-text');
const color = window.getComputedStyle(fuchsiaText)['color']; 
assert.strictEqual(color,'rgb(255, 0, 255)');
```

フクシアの色には16進コード`#FF00FF`の代わりに短縮された16進コードを使う必要があります。

```js
assert.match(__helpers.removeCssComments(code), /\.fuchsia-text\s*?{\s*?color\s*:\s*?#F0F\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: #000000;
  }
  .fuchsia-text {
    color: #000000;
  }
  .cyan-text {
    color: #000000;
  }
  .green-text {
    color: #000000;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #F00;
  }
  .fuchsia-text {
    color: #F0F;
  }
  .cyan-text {
    color: #0FF;
  }
  .green-text {
    color: #0F0;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```
