---
id: bad87fee1348bd9aedf08721
title: 16進コードを使って色を混ぜる
challengeType: 0
videoUrl: 'https://scrimba.com/c/cK89PhP'
forumTopicId: 18359
dashedName: use-hex-code-to-mix-colors
---

# --description--

復習すると、16進コードは6桁の16進数の数字を使って色を表し、赤（R）、緑（G）、青（B）の各コンポーネントに2桁ずつ割り当てられています。

これら3つの純粋な色（赤、緑、青）から、それぞれの量を変えることで1600万色以上の色を作り出せます！

例えば、オレンジは純粋な赤に少し緑を混ぜて、青は使いません。16進コードでは、これは`#FFA500`に対応します。

数字の`0`は16進コードで最も低い数字で、色がまったくないことを表します。

数字の`F`は16進コードで最も高い数字で、最大の明るさを表します。

# --instructions--

`style`要素内の色の単語を正しい16進コードに置き換えてください。

<table><tbody><tr><th>色</th><th>16進コード</th></tr><tr><td>Dodger Blue</td><td><code>#1E90FF</code></td></tr><tr><td>緑</td><td><code>#00FF00</code></td></tr><tr><td>オレンジ</td><td><code>#FFA500</code></td></tr><tr><td>赤</td><td><code>#FF0000</code></td></tr></tbody></table>

# --hints--

テキストが`h1`の`I am red!`要素には`color`の赤を設定してください。

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

色「赤」の`hex code`は単語`red`の代わりに使うべきです。

```js
assert.match(code, /\.red-text\s*?{\s*?color\s*:\s*?(#FF0000|#F00)\s*?;?\s*?}/gi);
```

テキストが`h1`の`I am green!`要素には`color`の緑を設定してください。

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

色「緑」の`hex code`は単語`green`の代わりに使うべきです。

```js
assert.match(code, /\.green-text\s*?{\s*?color\s*:\s*?(#00FF00|#0F0)\s*?;?\s*?}/gi);
```

テキストが`h1`の`I am dodger blue!`要素には`color`のdodger blueを設定してください。

```js
const blueText = document.querySelector('.dodger-blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(30, 144, 255)');
```

色「dodger blue」の`hex code`は単語`dodgerblue`の代わりに使うべきです。

```js
assert.match(code, /\.dodger-blue-text\s*?{\s*?color\s*:\s*?#1E90FF\s*?;?\s*?}/gi);
```

テキストが`h1`の`I am orange!`要素には`color`のオレンジを設定してください。

```js
const orangeText = document.querySelector('.orange-text');
const color = window.getComputedStyle(orangeText)['color']; 

assert.strictEqual(color, 'rgb(255, 165, 0)');
```

色「オレンジ」の`hex code`は単語`orange`の代わりに使うべきです。

```js
assert.match(code, /\.orange-text\s*?{\s*?color\s*:\s*?#FFA500\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: black;
  }
  .green-text {
    color: black;
  }
  .dodger-blue-text {
    color: black;
  }
  .orange-text {
    color: black;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #FF0000;
  }
  .green-text {
    color: #00FF00;
  }
  .dodger-blue-text {
    color: #1E90FF;
  }
  .orange-text {
    color: #FFA500;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```
