---
id: 6989eb65324d97775c4965fe
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستنفذ دالة تولد كل التركيبات الصحيحة من الأقواس العادية باستخدام طريقة البحث بالاتساع أولاً (BFS). على سبيل المثال، التركيبات الصحيحة لزوجين من الأقواس هي `(())` و `()()`.

ابدأ بإنشاء دالة باسم `genParentheses` تحتوي على معلمة واحدة هي `pairs`. في الوقت الحالي، أرجع مصفوفة فارغة من الدالة.

# --hints--

يجب أن تعرف دالة باسم `genParentheses`.

```js
assert.isFunction(genParentheses);
```

يجب أن تحتوي دالتك `genParentheses` على معلمة واحدة باسم `pairs`.

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

يجب أن تُرجع دالتك مصفوفة فارغة.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
