---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

الآن عليك التعامل مع إضافة الأقواس العادية المغلقة. القاعدة الأساسية هي أنه يمكنك إضافة قوس عادي مغلق فقط إذا حافظت على التوازن، أي يجب أن يكون هناك أقواس عادية مفتوحة أكثر من الأقواس المغلقة المستخدمة.

لا تزال داخل كتلة `else`، أضف جملة `if` ثانية للتحقق مما إذا كانت قيمة `closesUsed` أقل من `opensUsed`.

إذا تحقق هذا الشرط، ادفع مصفوفة جديدة أخرى إلى `queue`: `[current + ')', opensUsed, closesUsed + 1]`. هذا يمثل الحالة بعد إضافة قوس عادي مغلق.

# --hints--

يجب أن تحتوي كتلة `else` على جملة `if` ثانية بالشرط `closesUsed < opensUsed`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

يجب أن تدفع `[current + ')', opensUsed, closesUsed + 1]` إلى قائمة الانتظار عندما يكون الشرط صحيحًا.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
