---
id: 6989f2e35e6941e13e445156
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

الآن ستنفذ الحلقة الرئيسية لـ BFS. أنشئ حلقة `while` تستمر طالما أن `queue` غير فارغة، أي أن `queue.length` أكبر من `0`.

داخل الحلقة، سجّل `queue` في وحدة التحكم.

# --hints--

يجب أن تنشئ حلقة `while` بشرط `queue.length > 0`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

يجب أن تسجل `queue` في وحدة التحكم داخل حلقة `while` الخاصة بك.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
--fcc-editable-region--
  
--fcc-editable-region--
  return result;
}
```
