---
id: 6989f207306332c06b0add39
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

بالنسبة لطريقة البحث بالاتساع أولاً، ستستخدم قائمة انتظار لتتبع الحالات المختلفة أثناء بناء تراكيب الأقواس العادية. كل حالة ستُمثّل كمصفوفة تحتوي على ثلاثة عناصر:

- السلسلة النصية الحالية التي تُبنى
- عدد الأقواس العادية المفتوحة المستخدمة حتى الآن
- عدد الأقواس العادية المغلقة المستخدمة حتى الآن

أنشئ متغيرًا باسم `queue` وابدأ تهيئته بمصفوفة تحتوي على مصفوفة واحدة: `['', 0, 0]`. هذا يُمثل الحالة الابتدائية بسلسلة نصية فارغة وعدد صفر من الأقواس المستخدمة.

# --hints--

يجب إعلان متغير باسم `queue` داخل دالة `genParentheses` الخاصة بك.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.queue);
```

يجب تهيئة `queue` بمصفوفة تحتوي على المصفوفة `['', 0, 0]`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { queue } = genParentheses?.variables ?? {};
assert.isFalse(queue?.value.isEmpty());
const queueValue = eval(queue?.value.toString());
assert.deepEqual(queueValue, [['', 0, 0]]);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  
  return result;
}
```
