---
id: 6989ef885db3a3488a13b426
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

قبل تنفيذ الخوارزمية الأساسية، يجب التحقق من صحة المُدخَل. يجب أن تكون معلمة `pairs` عددًا صحيحًا، لأنها تمثل عدد أزواج الأقواس العادية التي ستُولد.

أضف عبارة `if` في بداية دالتك للتحقق مما إذا كان `pairs` ليس عددًا صحيحًا. استخدم الطريقة `Number.isInteger()` لذلك.

إذا تحقق الشرط، أرجع السلسلة النصية `The number of pairs should be an integer`.

# --hints--

يجب أن تحتوي دالة `genParentheses` على جملة if الشرطية بالشرط `!Number.isInteger(pairs)`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

يجب أن تُرجع سلسلة نصية عندما لا يكون المُدخَل عددًا صحيحًا.

```js
let result = genParentheses("i");
assert.isString(result);
```

يجب أن تُرجع السلسلة النصية `The number of pairs should be an integer` عندما يكون المُدخَل سلسلة نصية.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  
  return [];
}
```

