---
id: 6989f10a72ad424d4ad6d514
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

بعد ذلك، يجب التحقق من أن عدد الأزواج لا يقل عن واحد، إذ لا يمكنك توليد تركيبات الأقواس العادية إذا كان عدد الأزواج صفرًا أو سالبًا.

أضف جملة `if` أخرى للتحقق مما إذا كان `pairs` أقل من `1`. إذا تحقق هذا الشرط، أرجع السلسلة النصية `The number of pairs should be at least 1`.

# --hints--

يجب أن تحتوي على جملة `if` ثانية بالشرط `pairs < 1`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

يجب أن ترجع السلسلة النصية `The number of pairs should be at least 1` عندما يكون `pairs` أقل من `1`.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  
  return [];
}
```
