---
id: 6989f555728e63ad8b349d84
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

إذا لم تكتمل السلسلة النصية الحالية بعد، يجب استكشاف الحالات الممكنة التالية. أضف بند `else` لعبارة `if` الحالية.

داخل كتلة `else`، ستتعامل مع إضافة الأقواس العادية المفتوحة. أضف عبارة `if` للتحقق مما إذا كان `opensUsed` أقل من `pairs`. هذا يضمن ألا تستخدم أكثر من عدد الأقواس المفتوحة المسموح به.

إذا كان هذا الشرط `true`، ادفع مصفوفة جديدة إلى `queue`: `[current + '(', opensUsed + 1, closesUsed]`. هذا يمثل الحالة بعد إضافة قوس عادي مفتوح.

# --hints--

يجب أن يكون لديك بند `else` لعبارة `if` داخل حلقة `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

يجب أن يكون لديك عبارة `if` داخل كتلة `else` بالشرط `opensUsed < pairs`.

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

يجب أن تدفع `[current + '(', opensUsed + 1, closesUsed]` إلى قائمة الانتظار عندما يكون الشرط `true`.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
