---
id: 6989f1ba358acba651d5ee6b
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

الآن ستُعد هيكل البيانات لتخزين نتائجك. أنشئ متغيرًا باسم `result` وابدأه بمصفوفة فارغة. ستخزن هذه المصفوفة كل تراكيب الأقواس العادية الصالحة التي تولدها.

حدّث جملة الإرجاع لتُرجع `result` بدلًا من مصفوفة فارغة.

# --hints--

يجب إعلان متغير باسم `result` في دالة `genParentheses` الخاصة بك.

```js
assert.match(genParentheses.toString(), /result/);
```

يجب أن تبدأ `result` بمصفوفة فارغة.

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

يجب أن تُرجع `result` في نهاية دالتك.

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
