---
id: 6989f48c3f0292c84f8d8938
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

الآن عليك تحديد متى تكون قد بنيت تركيبة كاملة من الأقواس العادية. التركيبة الكاملة طولها يساوي ضعف عدد الأزواج (لأن كل زوج يضيف قوسًا افتتاحيًا وقوسًا إغلاقًا).

داخل حلقة `while`، أضف جملة `if` للتحقق مما إذا كان `current.length === 2 * pairs`. عندما يكون هذا الشرط صحيحًا، أضف `current` إلى مصفوفة `result`.

# --hints--

يجب أن تحتوي حلقة `while` على جملة `if` بالشرط `current.length === 2 * pairs`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

يجب أن تضيف `current` إلى مصفوفة `result` داخل جملة `if` الجديدة.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
