---
id: 6989f3632f8b8bdf168aa321
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

داخل حلقة `while`، استخدم `queue.shift()` لإزالة أول عنصر من الطابور والحصول عليه. هذا ينفذ سلوك الأول يدخل أول يخرج (FIFO) المميز لخوارزمية BFS.

فكك هذه المصفوفة إلى ثلاثة متغيرات: `current`، و`opensUsed`، و`closesUsed`.

# --hints--

يجب أن تستخدم `queue.shift()` للحصول على أول عنصر من الطابور.

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

يجب أن تفكك النتيجة إلى ثلاثة متغيرات: `current`، و`opensUsed`، و`closesUsed`.

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}
```
