---
id: 698c55e7aaacfae48a227459
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`h2` 要素の下に、テキスト `audio` を持つ新しい `Cruising for a Musing` 要素を追加してください。

`h2` の下に、テキスト `p` を持つ `Artist: Quincy Larson` 要素を追加してください。

# --hints--

2つ目の `h2` 要素があるはずです。

```js
assert.exists(document.querySelector('h2:nth-of-type(2)'));
```

2つ目の `h2` 要素にはテキスト `Cruising for a Musing` が含まれているはずです。

```js
const h2 = document.querySelector('h2:nth-of-type(2)');
assert.strictEqual(h2?.textContent.trim(), 'Cruising for a Musing');
```

2つ目の `p` 要素があるはずです。

```js
assert.exists(document.querySelector('p:nth-of-type(2)'));
```

2つ目の `p` 要素にはテキスト `Artist: Quincy Larson` が含まれているはずです。

```js
const p = document.querySelector('p:nth-of-type(2)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
