---
id: 698c58745b73289bb6c89ece
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

最後の曲のタイトルは `Scratching the Surface`、アーティストは Quincy Larson、使用するファイルは `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3` です。

最初の2曲と同じパターンで、最後の曲を追加してミュージックプレーヤーを完成させてください。

# --hints--

3つ目の `h2` 要素があるはずです。

```js
assert.exists(document.querySelector('h2:nth-of-type(3)'));
```

3つ目の `h2` 要素にはテキスト `Scratching the Surface` が含まれているはずです。

```js
const h2 = document.querySelector('h2:nth-of-type(3)');
assert.strictEqual(h2?.textContent.trim(), 'Scratching the Surface');
```

`p` 要素があるはずです。

```js
assert.exists(document.querySelector('p:nth-of-type(3)'));
```

`p` 要素にはテキスト `Artist: Quincy Larson` が含まれているはずです。

```js
const p = document.querySelector('p:nth-of-type(3)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

3つ目の `audio` 要素があるはずです。

```js
assert.exists(document.querySelector('audio:nth-of-type(3)'));
```

新しい `audio` 要素は `p` 要素の下にあるはずです。

```js
assert.exists(document.querySelector('p:nth-of-type(3) + audio:nth-of-type(3)'));
```

`audio` 要素には `src` 属性があるはずです。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('src'));
```

`audio` 要素には値 `src` の `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3` 属性があるはずです。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

`audio` 要素には `loop` 属性があるはずです。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('loop'));
```

`audio` 要素には `controls` 属性があるはずです。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

  <h2>Scratching the Surface</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3" loop controls></audio>
  
</body>
</html>
```
