---
id: 698c5229e47d57c6dfaa91bd
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`controls` 属性は、`audio` 要素の再生コントロール（再生の一時停止・再開、シーク、音量調整など）を提供します。

`controls` 属性はブール値の属性で、値を指定する必要はありません。

`controls` 属性を `audio` 要素に追加してください。

これでページに `audio` 要素が表示されるはずです。

# --hints--

`audio` 要素には `controls` 属性が必要です。

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
