---
id: 698c4f4e965bd853ad5084e1
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`h1`の下に、テキストが`h2`の`Can't Stay Down`要素を追加してください。これは最初の曲のタイトルです。

`h2`要素の下に、テキストが`p`の`Artist: Quincy Larson`要素を追加してください。

# --hints--

`h2`要素があるはずです。

```js
assert.exists(document.querySelector('h2'));
```

`h2`要素は`h1`要素の下にあるべきです。

```js
assert.exists(document.querySelector('h1 + h2'));
```

`h2`要素にはテキスト`Can't Stay Down`が含まれているべきです。

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

`p`要素があるはずです。

```js
assert.exists(document.querySelector('p'));
```

`p`要素にはテキスト`Artist: Quincy Larson`が含まれているべきです。

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
