---
id: 698c4fe24f1f1a17be74a23e
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

次に、`audio` 要素の下に `p` 要素を作成してください。次のいくつかのステップで、音声要素が音楽を再生するために必要な属性を追加していきます。

# --hints--

`audio` 要素があるはずです。

```js
assert.exists(document.querySelector('audio'));
```

新しい `audio` 要素は `p` 要素の下に配置してください。

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
