---
id: 698c50535b54b2d75fa985fb
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`loop` 属性は再生が完了したら音声を再起動します。再生を繰り返すインターネットミームのように考えてください。`loop` 属性を省略すると、音声は一度だけ再生されます。

`loop` 属性はブール値の属性で、値を指定する必要はありません。

`loop` 属性を `audio` 要素に追加してください。

# --hints--

`audio` 要素には `loop` 属性が必要です。

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
