---
id: 698c5710136539797f296098
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`p` 要素の下に `audio` 要素を追加し、`src` 属性を `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3` の値で、さらに `loop` と `controls` の属性を付けてください。

# --hints--

2つ目の `audio` 要素があるはずです。

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

新しい `audio` 要素は `p` 要素の下にあるべきです。

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

`audio` 要素には `src` 属性が必要です。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

`audio` 要素には値が `src` の `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3` 属性が必要です。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

`audio` 要素には `loop` 属性が必要です。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

`audio` 要素には `controls` 属性が必要です。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
