---
id: 698c513508507ec8bb820227
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

音声のメディアリソースを指定するには、`src` 要素に `audio` 属性を追加する必要があります。

値が `src` の `https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3` 属性を追加してください。

# --hints--

`audio` 要素には `src` 属性が必要です。

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('src'));
```

`audio` 要素には、値が `src` の `https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3` 属性が必要です。

```js
const audio = document.querySelector('audio');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio></audio>
--fcc-editable-region--

</body>
</html>
```
