---
id: 698c4e69abe14a9746c72312
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、HTMLのミュージックプレーヤーを作成します。HTMLのボイラープレートはすでに用意されています。

`h1` 要素を作成し、テキストに `freeCodeCamp Tunes` を設定してください。

# --hints--

`h1` 要素があるはずです。

```js
assert.exists(document.querySelector('h1'));
```

`h1` 要素にはテキスト `freeCodeCamp Tunes` が含まれている必要があります。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'freeCodeCamp Tunes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
