---
id: 6565a536ba1f9f25bd30e88b
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在此項目中，你將實現 Luhn 算法。 該算法是各種識別驗證號碼的公式。

首先聲明一個名爲 `main` 的函數，它將作爲程序的入口點。 使用 `pass` 關鍵字以避免錯誤。

# --hints--

你的代碼中應該有 `def main():` 。

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

你的代碼中應該有 `pass` 。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
