---
id: 6565a536ba1f9f25bd30e88b
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個專案中，你將實作 `Luhn Algorithm`。這個演算法是一種用來驗證各種識別號碼的公式。

首先宣告一個名為 `main` 的函式（程式），這將作為程式的進入點。使用 `pass` 關鍵字以避免錯誤。

# --hints--

你應該在你的程式碼中有 `def main():`。

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

你應該在你的程式碼中有 `pass`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
