---
id: 656879c6f35a85a59c06b3a7
title: 步驟 10
challengeType: 20
dashedName: step-10
---

# --description--

你之前已使用索引運算符 `[]` 訪問過字符串的元素（字符）。 你還可以使用索引運算符通過 `string[start:stop:step]` 訪問字符串中的一系列字符：

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

其中 `start` 是起始索引（包含），`stop` 是結束索引（不包含），`step` 是要跳過的字符數。 如果未指定，`step` 默認爲 1。

創建一個名爲 `card_number_reversed` 的變量，並將 `card_number` 的前 4 個字符的值賦給它。

# --hints--

`verify_card_number` 函數中應該有 `card_number_reversed = card_number[0:4]`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
