---
id: 656878585631369a6b2d2191
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

定義翻譯本身不會翻譯字串。必須對要翻譯的字串呼叫 `translate` 方法，並將翻譯表作為引數：

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

創建一個名為 `translated_card_number` 的變數，並將呼叫 `translate` 方法於 `card_number` 並以 `card_translation` 作為引數的結果指定給它。

# --hints--

你應該在 `main` 中創建一個 `translated_card_number` 變數。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

你應該將 `translated_card_number` 指定為 `card_number.translate(card_translation)` 的值。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
