---
id: 65687b48f2201ab32e06c37c
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

目前，你的腳本會丟出 `TypeError`，因為你正嘗試將字串加到整數。你可以透過在將 `digit` 變數加到 `sum_of_odd_digits` 之前，使用內建的 `int` 函式將其轉換為整數來修正此問題：

```python
my_string = '123'
my_int = int(my_string)
```

在將 `digit` 變數添加到 `sum_of_odd_digits` 之前，先將其轉換為整數。然後，將 `print` 呼叫移動到 `verify_card_number` 函式的結尾，以列印 `sum_of_odd_digits` 的值。

# --hints--

你應該在 `for` 迴圈中包含 `sum_of_odd_digits += int(digit)`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

你應該在 `verify_card_number` 函式（程式）的結尾有 `print(sum_of_odd_digits)`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

你不應該再有 `print(digit)`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
