---
id: 656879a66338b2a461d5d307
title: 步驟 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

Luhn算法如下：

1. 從右到左，每隔一個數字的值加倍；如果乘積大於 9，則將乘積的數字相加。
2. 取所有數字的總和。
3. 如果所有數字之和是10的倍數，則該數字有效；否則無效。

假設一個帳號“7992739871”的示例將添加一個校驗位，使其具有 7992739871x 的形式：

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

將 `pass` 語句替換爲變量 `sum_of_odd_digits` 和值 `0`。

# --hints--

`verify_card_number` 函數中應該有 `sum_of_odd_digits = 0`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

你不應在 `verify_card_number` 函數中使用 `pass`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
