---
id: 656877f71bba2b97acedf9af
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

Python 內建了可以協助我們進行字串操作的類別。其中一個是 `str` 類別。它有一個名為 `maketrans` 的方法，可以幫助我們創建翻譯表。這個表可以用來替換字串中的字元：

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

上述內容在對字串進行呼叫時，會將所有 `t` 字元替換為 `c`，並將所有 `l` 字元替換為 `b`。

創建一個名為 `card_translation` 的變數，並指定一個翻譯表來將所有 `-` 和 ` ` 字元取代為空字串。

# --hints--

你應該在 `main` 中創建一個 `card_translation` 變數。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

你應該為 `card_translation` 指定一個值為 `str.maketrans({'-': '', ' ': ''})`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
