---
id: 656879c6f35a85a59c06b3a7
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

你之前已經使用索引運算子 `[]` 存取過字串的元素（字元）。你也可以使用索引運算子來存取字串中一個字元範圍，語法為 `string[start:stop:step]`：

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

其中 `start` 是起始索引（包含），`stop` 是結束索引（不包含），`step` 是要跳過的字元數量。如果未指定，`step` 預設為 1。

創建一個名為 `card_number_reversed` 的變數，並將其指定為 `card_number` 的前 4 個字元的值。

# --hints--

你應該在 `verify_card_number` 函式（程式）中有 `card_number_reversed = card_number[0:4]`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
