---
id: 656877c202f44995ed952608
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

將 `pass` 述語取代為名為 `card_number` 且值為 `'4111-1111-4555-1142'` 的變數。

# --hints--

你應該在 `main` 函式（程式）中有 `card_number = '4111-1111-4555-1142'`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

你不應該在你的程式碼中有 `pass`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
