---
id: 656879613a96aba21cbe80b9
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

在你的 `main` 函式中，呼叫 `verify_card_number` 函式，並將 `translated_card_number` 變數作為引數傳入。

# --hints--

你應該在 `main` 函式中包含 `verify_card_number(translated_card_number)`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +verify_card_number\(\s*translated_card_number\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

--fcc-editable-region--
main()
```
