---
id: 65687db6a9fcf3c548a6c1b7
title: 步驟 28
challengeType: 20
dashedName: step-28
---

# --description--

演算法的一部分是從右邊開始將每隔一個數字加倍。如果加倍的結果大於或相等於 `10`，將兩個數字相加。例如，如果數字是 `6`，加倍後得到 `12`。將 `1` 和 `2` 相加得到 `3`。你可以使用*整數除法*來取得第一個數字，並使用取模運算子（`%`）來取得第二個數字：

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

整數除法的結果是除法的商，向下取整到最接近的整數。

在 `if` 述語中，將 `number` 指定為 `number // 10`（整數除法）的結果加上 `number` 與 `10` 的餘數。

# --hints--

你應該在 `if` 述語中有 `number = number // 10 + number % 10`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
