---
id: 656879e1c749d7a6c5eba2d2
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

從 `verify_card_number` 函式中移除 `print` 呼叫。

# --hints--

你不應該在 `verify_card_number` 函式中有 `print` 呼叫。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;
        const no_comments = __helpers.python.removeComments(function_body);
        assert.notMatch(no_comments, /print\(/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    print(card_number_reversed)
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
