---
id: 656879a66338b2a461d5d307
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

Luhn 演算法如下：

1. 從右至左，將每隔一位的數字加倍；如果乘積大於 9，則將乘積的各位數字相加。
2. 將所有數字相加。
3. 如果所有數字的總和是 10 的倍數，則該數字有效；否則無效。

假設一個帳號範例為 "7992739871"，將會添加一個檢查碼，使其成為 7992739871x 的形式：

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

將 `pass` 述語取代為變數 `sum_of_odd_digits` 並賦值為 `0`。

# --hints--

你應該在 `verify_card_number` 函式中有 `sum_of_odd_digits = 0`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

你不應該在 `verify_card_number` 函式（程式）中使用 `pass`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
