---
id: 656877f71bba2b97acedf9af
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

Python 帶有內置的類，可以幫助我們進行字符串操作。 其中之一是 `str` 類。 它有一個名爲 `maketrans` 的方法可以幫助我們創建翻譯表。 該表可用於替換字符串中的字符：

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

當在字符串上調用時，上面的代碼將用 `c` 替換所有 `t` 字符，用 `b` 替換所有 `l` 字符。

創建一個名爲 `card_translation` 的變量，併爲其分配一個轉換表，以將所有 `-` 和 `` 字符替換爲空字符串。

# --hints--

你應該在 `main` 中創建一個 `card_translation` 變量。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

你應該爲 `card_translation` 分配一個值 `str.maketrans({'-': '', ' ': ''})`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'

--fcc-editable-region--
```
