---
id: 6569b831a07d7154c793301b
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

定義一個函式（程式）`verify_card_number`，帶有一個參數 `card_number`，並使用關鍵字 `pass` 使該函式不執行任何操作。

# --hints--

你應該有一個名為 `verify_card_number` 的函式（程式），並帶有一個名為 `card_number` 的參數。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_parameters } = verify_card_number;
        assert.equal(function_parameters.trim(), "card_number");
    }
})
```

你應該在 `verify_card_number` 函式的主體中使用 `pass` 關鍵字。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_body } = verify_card_number;
        assert.match(function_body, / +pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

main()
```
