---
id: 68dbc6b7e8f794a2ae69bc74
title: "implementare l'algoritmo delle n regine"
challengeType: 27
dashedName: implement-the-n-queens-algorithm
---

# --description--

Il problema delle N regine ti chiede di posizionare N regine su una scacchiera N×N in modo che nessuna regina possa attaccarne un'altra (nessuna condivide la stessa riga, colonna o diagonale).

Ad esempio, se c'è una scacchiera 4x4, una disposizione valida è:

```md
[1, 3, 0, 2]
```

Questo significa che nella riga 0 la regina è posizionata nella colonna 1; nella riga 1 la regina è posizionata nella colonna 3; nella riga 2 la regina è posizionata nella colonna 0; e nella riga 3 la regina è posizionata nella colonna 2.

Visivamente, questa disposizione appare così:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Dove `Q` rappresenta una regina e `.` rappresenta una casella vuota.

In questo laboratorio, implementerai il risolutore del problema delle N regine usando l'approccio della ricerca in profondità.


**Obiettivo**: Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `dfs_n_queens`.
2. La funzione dovrebbe accettare esattamente un argomento: un intero `n`.
3. Se `n` è minore di `1`, la funzione dovrebbe restituire una lista vuota (`[]`).
4. La funzione dovrebbe restituire una lista di soluzioni; ogni soluzione è a sua volta una lista di lunghezza `n`, dove l'elemento all'indice `i` è l'indice di colonna (basato su 0) della regina nella riga `i`.

# --hints--

Dovresti avere una funzione chiamata `dfs_n_queens` che prende un argomento.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs_n_queens)
  sig = inspect.signature(dfs_n_queens)
  assert len(sig.parameters) == 1
`) })
```

Se `n` è minore di `1`, la funzione dovrebbe restituire una lista vuota.

```js
({ test: () => runPython(`
  assert dfs_n_queens(0) == []
  assert dfs_n_queens(-1) == []
  assert dfs_n_queens(-5) == []
`) })
```

La funzione dovrebbe restituire una lista di soluzioni, dove ogni soluzione è una lista di lunghezza `n`.

```js
({ test: () => runPython(`
  result = dfs_n_queens(4)
  assert isinstance(result, list)
  for solution in result:
    assert isinstance(solution, list)
    assert len(solution) == 4
`) })
```

`dfs_n_queens(1)` dovrebbe restituire `[[0]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(1) == [[0]]
`) })
```

`dfs_n_queens(2)` dovrebbe restituire `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(2) == []
`) })
```

`dfs_n_queens(3)` dovrebbe restituire `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(3) == []
`) })
```

`dfs_n_queens(4)` dovrebbe restituire `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(4) == [[1, 3, 0, 2], [2, 0, 3, 1]]
`) })
```

`dfs_n_queens(5)` dovrebbe restituire `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`len(dfs_n_queens(5))` dovrebbe essere `10`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(5)) == 10
`) })
```

`len(dfs_n_queens(8))` dovrebbe essere `92`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(8)) == 92
`) })
```

`dfs_n_queens` dovrebbe restituire il risultato corretto.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(7)) == 40
  assert dfs_n_queens(6) == [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs_n_queens(n):
    if n < 1:
        return []

    results = []

    def is_safe(queens, row, col):
        for r, c in enumerate(queens):
            if c == col or abs(row - r) == abs(col - c):
                return False
        return True

    def dfs(queens):
        row = len(queens)
        if row == n:
            results.append(queens[:])
            return
        for col in range(n):
            if is_safe(queens, row, col):
                queens.append(col)
                dfs(queens)
                queens.pop()

    dfs([])
    return results
```
