---
id: a39963a4c10bc8b4d4f06d7e
title: 값 제거 함수 구현하기
challengeType: 26
dashedName: implement-a-value-remover-function
---

# --description--

이 실습에서는 첫 번째 인수로 초기 배열을 받고, 그 뒤에 제거할 값을 나타내는 하나 이상의 추가 인수를 받는 함수를 만드세요.

**목표**: 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 배열과 하나 이상의 추가 인수를 받는 `destroyer` 함수를 만들어야 합니다.
1. `destroyer` 함수는 첫 번째 인수 배열에서 이후 인수 중 하나와 일치하는 모든 요소를 제외한 새 배열을 반환해야 합니다.
1. 함수는 인수 개수가 정해지지 않아도 받아야 합니다.

# --hints--

`destroyer` 함수를 만들어야 합니다.

```js
assert.isFunction(destroyer);
```

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)`는 `[1, 1]`를 반환해야 합니다.

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)`는 `[1, 5, 1]`를 반환해야 합니다.

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)`는 `[1]`를 반환해야 합니다.

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)`는 `["hamburger"]`를 반환해야 합니다.

```js
assert.deepEqual(destroyer(["tree", "hamburger", 53], "tree", 53), ["hamburger"]);
```

`destroyer( ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan" )`는 `[12, 92, 65]`를 반환해야 합니다.

```js
assert.deepEqual(
  destroyer(
    ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"],
    "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan"
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function destroyer(arr, ...valsToRemove) {
  return arr.filter(item => !valsToRemove.includes(item));
}
```
