---
id: 61329d52e5010e08d9b9d66b
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Un altro importante elemento `meta` per l'accessibilitĂ  e la SEO Ă¨ la definizione `description`. Il valore dell'attributo `content` viene usato dai motori di ricerca per fornire una descrizione della tua pagina.

Aggiungi un elemento `meta` con l'attributo `name` impostato su `description` e dagli un utile attributo `content`.

# --hints--

Dovresti aggiungere un nuovo elemento `meta` al `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

Dovresti dare al `meta` un attributo `name` di `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

Dovresti dare al `meta` un attributo `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

Il valore dell'attributo `content` non dovrebbe superare i 165 caratteri. _Questa Ă¨ la lunghezza massima della descrizione secondo Google._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
