---
id: 613e275749ebd008e74bb62e
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Attualmente, il `img` sta assumendo la sua dimensione predefinita, che è troppo grande. CSS ha una funzione `max` che restituisce il valore più grande di un insieme di valori separati da virgola. Per esempio:

```css
img {
  width: max(250px, 25vw);
}
```

Nell'esempio sopra, la larghezza dell'immagine sarà di 250px se la larghezza della viewport è inferiore a 1000 pixel. Se la larghezza della viewport è maggiore di 1000 pixel, la larghezza dell'immagine sarà 25vw. Questo perché 25vw corrisponde al 25% della larghezza della viewport.

Scala l'immagine usando il suo `id` come selettore, e imposta il `width` per essere il massimo tra `10rem` o `18vw`. 

# --hints--

Dovresti usare il selettore `#logo` per mirare all'elemento `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

Dovresti dare al `img` un `width` di `max(10rem, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
