---
id: 6140827cff96e906bd38fc2b
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Come descritto nella <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">guida di stile freeCodeCamp</a>, il logo dovrebbe mantenere un rapporto d'aspetto di `35 / 4` e avere una spaziatura interna attorno al testo.

Per prima cosa, cambia `background-color` in `#0a0a23` così puoi vedere il logo. Poi, usa la proprietà `aspect-ratio` per impostare il rapporto d'aspetto desiderato su `35 / 4`. Infine, aggiungi una spaziatura interna `padding` di `0.4rem` tutt'intorno.

# --hints--

Dovresti dare a `#logo` una spaziatura interna `background-color` di `#0a0a23`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

Dovresti usare la proprietà `aspect-ratio`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

Non dovresti usare la proprietà `height`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

Non dovresti modificare la proprietà `width`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

Dovresti dare a `img` una spaziatura interna `aspect-ratio` di `35 / 4`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

Dovresti dare a `img` una spaziatura interna `padding` di `0.4rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--
#logo {
  width: max(10rem, 18vw);

}
--fcc-editable-region--

```
