---
id: a39963a4c10bc8b4d4f06d7e
title: Реалізувати функцію для видалення значень
challengeType: 26
dashedName: implement-a-value-remover-function
---

# --description--

У цій лабораторній роботі ви створите функцію, яка приймає ініціальний масив як перший аргумент, а потім один або кілька додаткових аргументів, що представляють значення для видалення.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути `destroyer` функція, яка приймає масив і один або кілька додаткових аргументів.
1. `destroyer` функція має повертати новий масив, виключаючи всі елементи з першого аргументу, які співпадають з будь-яким із наступних аргументів.
1. Функція має приймати невизначену кількість аргументів.

# --hints--

У вас має бути `destroyer` функція.

```js
assert.isFunction(destroyer);
```

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)` має повертати `[1, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` має повертати `[1, 5, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)` має повертати `[1]`.

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)` має повертати `[]`.

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)` має повертати `["hamburger"]`.

```js
assert.deepEqual(destroyer(["tree", "hamburger", 53], "tree", 53), ["hamburger"]);
```

`destroyer( ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan" )` має повертати `[12, 92, 65]`.

```js
assert.deepEqual(
  destroyer(
    ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"],
    "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan"
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function destroyer(arr, ...valsToRemove) {
  return arr.filter(item => !valsToRemove.includes(item));
}
```
