---
id: 6852dd8486c7ee0171322950
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Ti servirà una variabile per contenere il codice segreto mentre lo decodifichi, quindi rimuovi `pass` e crea una variabile chiamata `secret_code` all'interno della funzione assegnandole una stringa vuota.

# --hints--

Dovresti avere una variabile `secret_code` all'interno della funzione `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

La variabile `secret_code` dovrebbe avere come valore una stringa vuota.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

Non dovresti più avere `pass` nel tuo codice.

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
