---
id: 6852db755c9c1274d75aae9a
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop creerai un estrattore di pin, le cifre del pin sono nascoste in ogni riga di una poesia.

Inizia creando una funzione chiamata `pin_extractor` con un parametro `poem`. Una funzione vuota genera un errore, quindi aggiungi momentaneamente `pass` all'interno della funzione per renderla valida.

# --hints--

Dovresti avere una funzione chiamata `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

La funzione `pin_extractor` dovrebbe avere un parametro chiamato `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
