---
id: 68531fbb934d735415053274
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora che hai aggiunto una chiamata a una funzione puoi vedere eventuali errori creati all'interno della funzione nel terminale, e puoi sempre usare `print` per vedere il valore delle cose.

C'è una cifra del pin nascosta in ogni riga, quindi all'interno della funzione usa il metodo `split` per dividere la stringa in una lista di righe. Dividi le righe sul carattere di nuova riga (`\n`) e assegna la lista risultante a una variabile chiamata `lines`.

# --hints--

Dovresti avere una variabile chiamata `lines` all'interno della tua funzione `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

La variabile `lines` dovrebbe avere un valore di `poem.split('\n')`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
