---
id: 68530231c05b93324e071104
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Mentre lavori sulla funzione, sarebbe utile avere una chiamata alla funzione per scopi di debugging, quindi crea una variabile `poem` esterna alla funzione che contenga questa poesia:

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

Puoi usare una stringa multilinea per questo.

Poi, chiama la funzione con la variabile `poem` come argomento.

# --hints--

Dovresti avere una variabile globale chiamata `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

La variabile `poem` dovrebbe avere un valore uguale a `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

Dovresti chiamare la funzione `pin_extractor` con `poem` come argomento.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
