---
id: 68553423d775cc60134ac8bb
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Devi lavorare su ogni riga in modo indipendente: crea un loop su `lines` che usa `line` come variabile del loop. All’interno del loop, stampa `line`.

# --hints--

Dovresti avere un loop `for` nella tua funzione `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

Il loop dovrebbe iterare su `lines`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

Il loop dovrebbe avere `line` come variabile del loop (`for line in lines`).

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

Dovresti avere `print(line)` all’interno del loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
