---
id: 6855918e87c33c6979d21d3f
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Ora la funzione fa un buon lavoro nell’estrarre il codice dalla poesia, ma c’è un caso limite da considerare: la riga della poesia potrebbe essere più corta del previsto, e questo farebbe generare un errore alla funzione (puoi provare a rimuovere la parola `bright` dalla poesia).

Metti la riga `secret_code += str(len(words[line_index]))` in un’istruzione `if` che verifichi che ci siano abbastanza parole nella lista `words`.

# --hints--

Dentro il loop dovresti avere un’istruzione `if`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

La condizione dell’istruzione if dovrebbe essere `len(words) > line_index`.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

Dentro l’istruzione `if` dovresti avere `secret_code += str(len(words[line_index]))`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
