---
id: 685535bf527b2c77236b3a9a
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

La cifra n-esima del pin è nascosta come la lunghezza della parola n-esima nella riga n-esima.

Per trovare la lunghezza della parola n-esima, il passaggio successivo è dividere la riga della poesia in una lista di parole: all'interno del loop, crea una variabile `words` e assegnale il valore di `line` diviso in parole usando il metodo `split`.

Poi, aggiungi una chiamata `print` con `words` come suo argomento.

# --hints--

Dovresti avere una variabile `words` all'interno del loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

La variabile `words` dovrebbe avere un valore di `line.split()`.

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

Dovresti avere `print(words)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
