---
id: 685578aae65e4106ad2b4300
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Come hai imparato in una lezione precedente, la funzione `enumerate` permette di tenere traccia dell'indice mentre si itera su un iterabile:

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

La cifra n-esima del codice segreto proviene dalla parola n-esima della riga n-esima, quindi devi sapere qual è il numero di ogni riga.

Modifica il loop in modo che iteri su `enumerate(lines)`, e aggiungi un'altra variabile di loop prima di `line` chiamata `line_index`.

Inoltre, aggiorna la chiamata a `print` in `print(line_index, line)`.

# --hints--

Il loop dovrebbe iterare su `enumerate(lines)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

Il loop dovrebbe avere `line_index, line` come variabile di loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

Dovresti avere `print(line_index, line)` all'interno del loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
