---
id: 68592730aaf34b61ae279168
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

La lunghezza restituita da `len` è un intero, devi convertirla in una stringa usando la funzione `str` per poterla aggiungere alla stringa `secret_code`.

Aggiorna la chiamata `print` per stampare la lunghezza come stringa.

# --hints--

Dovresti avere `print(str(len(words[line_index])))` alla fine del loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(len(words[line_index]))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
