---
id: 684aaf9ec670c68d20efd0d0
title: Побудова менеджера конфігурації користувача
challengeType: 27
dashedName: build-a-user-configuration-manager
saveSubmissionToDB: true
---

# --description--

У цій лабораторній роботі ви створите менеджер конфігурації користувача, який дозволяє користувачам керувати своїми налаштуваннями, такими як тема, мова та сповіщення. Ви реалізуєте функції для додавання, оновлення, видалення та перегляду налаштувань користувача.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні визначити функцію з іменем `add_setting` з двома параметрами, які представляють словник налаштувань і кортеж, що містить пару ключ-значення.

1. Функція `add_setting` повинна:
    
   - Перетворити ключ і значення на нижній регістр.
   - Якщо налаштування з таким ключем існує, повернути `Setting '[key]' already exists! Cannot add a new setting with this name.`.
   - Якщо налаштування з таким ключем не існує, додати пару ключ-значення до заданого словника налаштувань і повернути `Setting '[key]' added with value '[value]' successfully!`.
   - Повідомлення, що повертаються, повинні містити ключ і значення у нижньому регістрі.

1. Ви повинні визначити функцію з іменем `update_setting` з двома параметрами, які представляють словник налаштувань і кортеж, що містить пару ключ-значення.

1. Функція `update_setting` повинна:
       
   - Перетворити ключ і значення на нижній регістр.
   - Якщо налаштування з таким ключем існує, оновити його значення у заданому словнику налаштувань і повернути: `Setting '[key]' updated to '[value]' successfully!`.
   - Якщо налаштування з таким ключем не існує, повернути `Setting '[key]' does not exist! Cannot update a non-existing setting.`.
   - Повідомлення, що повертаються, повинні містити ключ і значення у нижньому регістрі.

1. Ви повинні визначити функцію з іменем `delete_setting` з двома параметрами, які представляють словник налаштувань і ключ.

1. Функція `delete_setting` повинна:
    
   - Перетворити переданий ключ на нижній регістр.
   - Якщо налаштування з таким ключем існує, видалити пару ключ-значення із заданого словника налаштувань і повернути `Setting '[key]' deleted successfully!`.
   - Якщо налаштування з таким ключем не існує, повернути `Setting not found!`.
   - Повідомлення, що повертаються, повинні містити ключ у нижньому регістрі.

1. Ви повинні визначити функцію з іменем `view_settings` з одним параметром, який представляє словник налаштувань.

1. Функція `view_settings` повинна:
   
   - Повернути `No settings available.`, якщо заданий словник налаштувань порожній.
   - Якщо словник містить будь-які налаштування, повернути рядок, що відображає налаштування. Рядок повинен починатися з `Current User Settings:`, за яким ідуть пари ключ-значення, кожна з нової лінії, а ключ повинен бути з великої літери. Наприклад, `view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})` має повернути:

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. Для тестування коду ви повинні створити словник з іменем `test_settings`, щоб зберігати деякі налаштування користувача.

# --hints--

Ви повинні створити словник з іменем `test_settings` і додати до нього деякі значення.

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

Ви повинні визначити функцію з іменем `add_setting`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

Функція `add_setting` повинна мати два параметри.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting` має перетворювати ключ на нижній регістр.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting` має перетворювати значення на нижній регістр.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))` має повертати повідомлення про помилку `Setting 'theme' already exists! Cannot add a new setting with this name.`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))` має додавати нову пару ключ-значення і повертати повідомлення про успіх `Setting 'volume' added with value 'high' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting` має правильно додавати задану пару ключ-значення до словника.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

Ви повинні визначити функцію з іменем `update_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

Функція `update_setting` повинна мати два параметри.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

Функція `update_setting` повинна перетворювати ключ на нижній регістр.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

Функція `update_setting` повинна перетворювати значення на нижній регістр.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))` має оновлювати існуючий ключ і повертати повідомлення про успіх `Setting 'theme' updated to 'dark' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))` має повертати повідомлення про помилку `Setting 'volume' does not exist! Cannot update a non-existing setting.`, коли ключ не існує.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting` має правильно оновлювати задану пару ключ-значення у словнику.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

Ви повинні визначити функцію з іменем `delete_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

Функція `delete_setting` повинна мати два параметри.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting` має перетворювати ключ на нижній регістр.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')` має видаляти існуючий ключ і повертати повідомлення про успіх `Setting 'theme' deleted successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting` має повертати повідомлення про помилку `Setting not found!`, коли ключ не існує.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting` має правильно видаляти заданий ключ зі словника.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

Ви повинні визначити функцію з іменем `view_settings`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

Функція `view_settings` повинна мати один параметр.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings` має повертати повідомлення `No settings available.`, якщо заданий словник порожній.

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings` має повертати відформатовані налаштування для непорожнього словника.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings` має робити першу літеру кожного імені налаштування великою.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings` має відображати правильні результати і закінчуватися символом нового рядка.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "english", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: english\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "english",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
