---
id: 616971b3cd990186b66c99a1
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

外部CSSファイルの設定ができたので、ページのスタイルを設定し始めましょう。

復習として、段落要素をターゲットにして右寄せにする方法は次の通りです。

```css
p {
  text-align: right;
}
```

新しいCSSルールを作成し、`h1`要素をターゲットにして、`text-align`プロパティを`center`に設定してください。

# --hints--

`h1`セレクターを使う必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

`h1`のCSSルールには、`text-align`プロパティを`center`に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
