---
id: 61695197ac34f0407e339882
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

これまでのプロジェクトで見てきたように、ウェブページは`DOCTYPE html` 宣言で始まり、その後に`html` 要素が続きます。

ドキュメントの先頭に`DOCTYPE html` 宣言を追加し、その後に`html` 要素を追加してください。`html` 要素には、`lang` 属性を設定し、その値を`en` にしてください。

# --hints--

コードには`DOCTYPE` 参照が含まれている必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れてください。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプは`html` に設定してください。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

`DOCTYPE` 宣言はタイプの後に`>` で閉じてください。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 要素の開始タグには`lang` 属性が`en` として設定されている必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/gi));
```

`DOCTYPE` 宣言はHTMLの最初に置いてください。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
