---
id: 61695ab9f6ffe951c16d03dd
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`head` 要素の中に `html` 要素をネストしてください。`head` 要素の直後に `body` 要素を追加してください。

# --hints--

開始 `head` タグがあるはずです。

```js
assert(code.match(/<head\s*>/i));
```

終了 `head` タグがあるはずです。

```js
assert(code.match(/<\/head\s*>/i));
```

開始 `body` タグがあるはずです。

```js
assert(code.match(/<body\s*>/i));
```

終了 `body` タグがあるはずです。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` の開始タグは1つだけのはずです。

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

`head` の終了タグは1つだけのはずです。

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

`body` の開始タグは1つだけのはずです。

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

`body` の終了タグは1つだけのはずです。

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

`head` 要素は空であるべきです。

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

`body` 要素は空であるべきです。

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

`body` 要素は `head` 要素の後に配置されているべきです。

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

`head` 要素は `html` 要素の中にあるべきです。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は `html` 要素の中にあるべきです。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
