---
id: 616965351e74d4689eb6de30
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

ウェブページには複数の`meta`要素を配置できます。各`meta`要素は、他のHTML要素では表現できないページに関する情報を追加します。

`meta`の中にもう1つ`head`要素を追加してください。`name`属性を`"viewport"`に、`content`属性を`"width=device-width, initial-scale=1.0"`に設定して、すべてのデバイスでページが同じように見えるようにしましょう。

# --hints--

`meta`要素が2つあるはずです。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

新しい`meta`要素は空要素で、終了タグ`</meta>`を持ってはいけません。

```js
assert.notMatch(code, /<\/meta>/i);
```

新しい`meta`要素には、`name`属性を`"viewport"`に、`content`属性を`"width=device-width, initial-scale=1.0"`に設定してください。

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

2つの`meta`要素は`head`要素の中にあるべきです。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
