---
id: 61695d1fbc003856628bf561
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`charset` 属性は、ドキュメントで使用される文字エンコーディングを指定します。`utf-8`（Unicode Transformation Format – 8-bit）は、電子通信で使われる文字エンコーディングの標準です。

`head` 要素の中に、`meta` 属性が `charset` に設定された `"utf-8"` 要素をネストしてください。

# --hints--

`meta` 要素は1つだけにしてください。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

`meta` 要素は空要素である必要があり、終了タグ `</meta>` を持ってはいけません。

```js
assert.notMatch(code, /<\/meta>/i);
```

`meta` 要素には、`charset` 属性が `"utf-8"` の値に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

`meta` 要素は `head` 要素の中にある必要があります。

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
