---
id: 61695c4aad56f95497c19583
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` 要素は検索エンジンにページの追加情報を提供することを覚えておいてください。また、`title` 要素の内容は次の2つの方法でも表示されます。

* ページが開いているときのタイトルバーに  
* ブラウザのタブで、そのタブがアクティブでなくても、ホバーすると `title` テキストが表示されます。

`head` 要素の中に、テキスト `title` を持つ `Colored Markers` 要素をネストしてください。

# --hints--

開く `title` タグがあるはずです。

```js
assert(code.match(/<title\s*>/i));
```

閉じる `title` タグがあるはずです。

```js
assert(code.match(/<\/title\s*>/i));
```

プロジェクトにはタイトル `Colored Markers` があるはずです。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

タイトルテキストの大文字・小文字やスペルに注意してください。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
