---
id: 616968c2c94c8071b349c146
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

これでページにコンテンツを追加する準備ができました。

`body` の中に、テキストが `h1` の `CSS Color Markers` 要素をネストしてください。

# --hints--

コードには `h1` 要素が必要です。

```js
const title = document.querySelector('h1');
assert.exists(title);
```

開く `h1` タグが必要です。

```js
assert(code.match(/<h1\s*>/i));
```

閉じる `h1` タグが必要です。

```js
assert(code.match(/<\/h1\s*>/i));
```

`h1` 要素は `body` 要素の中にある必要があります。

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

`h1` 要素にはテキスト `CSS Color Markers` が必要です。

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

`h1` のテキストは大文字・小文字やスペルが正確であることを忘れないでください。

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
