---
id: 616d595270d902f0e7086e18
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

3つの別々のマーカー `div` 要素がありますが、それらは1つの大きな長方形のように見えます。各要素が見やすくなるように、間にスペースを追加してください。

ショートハンドの `margin` プロパティに2つの値がある場合、最初の値が `margin-top` と `margin-bottom` に設定され、2つ目の値が `margin-left` と `margin-right` に設定されます。

`.marker` の CSS ルールで、`margin` プロパティを `10px auto` に設定してください。

# --hints--

`.marker` の CSS ルールに `margin` プロパティを `10px auto` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
