---
id: 616d3a67ccf800ad94ec89ae
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

これから、最終的に色のマーカーとしてスタイルを適用する要素を追加します。

まず、`body` 要素の中に `div` 要素を追加し、その `class` 属性を `container` に設定してください。`div` 要素が `h1` 要素の下にあることを確認しましょう。

# --hints--

`div` の開始タグがあるはずです。

```js
assert(code.match(/<div.*>/i));
```

`div` の終了タグがあるはずです。

```js
assert(code.match(/<\/div\s*>/i));
```

`div` 要素は `body` 要素の中にあるはずです。

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

`div` 要素は `h1` 要素の後にあるはずです。

```js
assert.exists(document.querySelector('h1 + div'));
```

`div` 要素は `class` 属性が `container` に設定されているはずです。

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
