---
id: 616d4a84b756d9c4b8255093
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

マーカーに色を付ける時間です。要素に色を付ける方法の一つは、`black`、`cyan`、または `yellow` のような<dfn>カラーキーワード</dfn>を使うことです。

復習として、クラス `freecodecamp` をターゲットにする方法はこちらです：

```css
.freecodecamp {
  
}
```

クラス `marker` をターゲットにする新しい CSS ルールを作成し、その `background-color` プロパティを `red` に設定してください。

**注意：** CSS を追加してもすぐには変化が見られません。

# --hints--

クラス `marker` をターゲットにするクラスセレクターを作成してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

`.marker` の CSS ルールには、`background-color` プロパティを `red` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
