---
id: 616d50b93ba424d6282c99cf
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

背景色は適用されましたが、マーカー `div` 要素には内容がないため、デフォルトでは高さがありません。

`.marker` の CSS ルールで、`height` プロパティを `25px` に、`width` プロパティを `200px` に設定してください。

# --hints--

`.marker` の CSS ルールには、`width` プロパティを `200px` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

`.marker` の CSS ルールには、`height` プロパティを `25px` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
