---
id: 616d525007b8c5d8b3308b1c
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

マーカーはページの中央に配置すると見た目が良くなります。簡単に中央にする方法は `margin` <dfn>省略形プロパティ</dfn> を使うことです。

前のプロジェクトでは、`margin-top` や `margin-left` のようなプロパティで要素のマージンエリアを個別に設定しました。`margin` 省略形プロパティを使うと、複数のマージンエリアを同時に設定できます。

マーカーをページの中央に配置するには、`margin` プロパティを `auto` に設定してください。これにより `margin-top`、`margin-right`、`margin-bottom`、`margin-left` がすべて `auto` に設定されます。

# --hints--

`.marker` の CSS ルールに `margin` プロパティを `auto` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
