---
id: 61764c602bee6974e7790f35
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

マーカーに異なる色を付けるには、それぞれにユニークなクラスを追加する必要があります。複数のクラスは、`class` 属性にクラス名をスペースで区切って並べることで、1つの要素に追加できます。例えば、次の例は `animal` 要素に `dog` と `div` の両方のクラスを追加しています。

```html
<div class="animal dog">
```

HTML要素に複数のクラスを追加すると、先に記述したクラスのスタイルが後に記述したクラスによって上書きされる場合があります。

まず、最初のマーカー `one` 要素にクラス `div` を追加してください。


# --hints--

最初のマーカー `one` 要素にクラス `div` を追加してください。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

最初のマーカー `div` には、クラス `marker` と `one` を付ける必要があります。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
