---
id: 617b8e0d93a8d10d9a90e720
title: ステップ 23
challengeType: 0
dashedName: step-23
---

# --description--

マーカーの`background-color`がまだ赤いままであることに気づきましたか。これは、`rgb`関数の赤の値を`255`の最大値、つまり100%の赤に設定し、緑と青の値を`0`に設定したためです。

次に、`rgb`関数を使って他の色を設定しましょう。

`.two`のCSSルールでは、`rgb`関数を使って`background-color`を緑の最大値に設定し、他の値は`0`に設定してください。そして、`.three`のCSSルールでは、`rgb`関数を使って`background-color`を青の最大値に設定し、他の値は`0`に設定してください。

# --hints--

`.two`のCSSルールでは、`green`プロパティを設定するのに`background-color`カラーキーワードを使わないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

`.two`のCSSルールには、`background-color`プロパティが`rgb(0, 255, 0)`に設定されている必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

`.three`のCSSルールでは、`blue`プロパティを設定するのに`background-color`カラーキーワードを使わないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

`.three`のCSSルールには、`background-color`プロパティが`rgb(0, 0, 255)`に設定されている必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
