---
id: 617b954d9f4f6217a749380e
title: ステップ 25
challengeType: 0
dashedName: step-25
---

# --description--

マーカーと、それが入っている`container`要素の端との間に、もう少し縦のスペースを追加しましょう。

`.container`のCSSルールで、ショートハンドの`padding`プロパティを使って上下に`10px`のパディングを追加し、左右のパディングは`0`に設定してください。これは、以前使ったショートハンドの`margin`プロパティと同じように動作します。

# --hints--

`background-color`のCSSルールから`.container`プロパティとその値を削除しないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

`.container`のCSSルールには、`padding`プロパティが`10px 0`に設定されている必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.padding === '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
