---
id: 618a00ed1ca871a2b3aca0cb
title: ステップ 35
challengeType: 0
dashedName: step-35
---

# --description--

カラーホイールの基本色、二次色、三次色をすべて学んだので、他の色彩理論の概念やそれらがデザインに与える影響を理解しやすくなります。

まず、CSSルールの `.one`、`.two`、`.three` にある `rgb` 関数の値を調整して、各要素の `background-color` を純粋な黒に設定してください。`rgb` 関数は加法混色モデルを使っており、色は黒から始まり、赤、緑、青の値が増えるにつれて変化します。

# --hints--

`.one` の CSSルールには `background-color` プロパティを `rgb(0, 0, 0)` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

`.two` の CSSルールには `background-color` プロパティを `rgb(0, 0, 0)` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

`.three` の CSSルールには `background-color` プロパティを `rgb(0, 0, 0)` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
