---
id: 618a0927005553b74bfa05ae
title: ステップ 36
challengeType: 0
dashedName: step-36
---

# --description--

カラーホイールは、似た色や<dfn>色相</dfn>が近くにあり、異なる色は離れている円です。例えば、純粋な赤はバラ色とオレンジの色相の間にあります。

カラーホイール上で互いに反対に位置する2つの色は<dfn>補色の色</dfn>と呼ばれます。補色の2色を組み合わせると灰色になりますが、隣り合わせに置くと強い視覚的コントラストを生み、より明るく見えます。

`rgb` CSSルールの`.one`関数で、純粋な赤を作るために赤の値を`255`の最大値に設定してください。`rgb` CSSルールの`.two`関数で、シアンを作るために緑と青の値を`255`の最大値に設定してください。

# --hints--

`.one` CSSルールには`background-color`プロパティを`rgb(255, 0, 0)`に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

`.two` CSSルールには`background-color`プロパティを`rgb(0, 255, 255)`に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
