---
id: 619b761916dac02643940022
title: ステップ 46
challengeType: 0
dashedName: step-46
---

# --description--

CSSで要素に色を適用する非常に一般的な方法は、<dfn>16進数</dfn>または16進値を使うことです。16進値は複雑に聞こえますが、実際にはRGB値の別の形式にすぎません。

16進カラー値は`#`文字で始まり、0-9とA-Fの6文字を取ります。最初の2文字は赤を表し、2番目の2文字は緑を表し、3番目の2文字は青を表します。例えば、`#4B5320`です。

`.green`クラスセレクターで、`background-color`プロパティを赤が`00`、緑が`FF`、青が`00`の16進カラーコードに設定してください。

# --hints--

`.green`のCSSルールに`background-color`プロパティが`#00FF00`に設定されている必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
