---
id: 61a49d15bdbb5e57cc6fd280
title: ステップ 54
challengeType: 0
dashedName: step-54
---

# --description--

カラーストップは、グラデーションラインに沿って色が配置される位置を細かく調整できます。これは `px` やパーセンテージのような長さの単位で、`linear-gradient` 関数内の色の後に続きます。

例えば、この赤と黒のグラデーションでは、赤から黒への変化がグラデーションラインの90％の位置で起こるため、赤が利用可能なスペースの大部分を占めています。

```css
linear-gradient(90deg, red 90%, black);
```

`linear-gradient` 関数内で、最初の赤の色引数の後に `75%` のカラーストップを追加してください。他の色引数にはカラーストップを追加しないでください。

# --hints--

あなたの `.red` CSSルールには、`background` プロパティが `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))` に設定されている必要があります。

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
