---
id: 61a489b8579e87364b2d2cdb
title: ステップ 49
challengeType: 0
dashedName: step-49
---

# --description--

CSSで単色を設定する方法はいくつか学びましたが、要素にカラーの遷移、つまり<dfn>グラデーション</dfn>を使うこともできます。

グラデーションとは、ある色が別の色に変わっていくことです。CSSの`linear-gradient`関数を使うと、遷移の方向や使う色を指定できます。

覚えておくべきことは、`linear-gradient`関数は実際には`image`要素を作成し、通常は画像を値として受け取れる`background`プロパティと組み合わせて使われることです。

`.red`のCSSルールで、`background-color`プロパティを`background`に変更してください。

# --hints--

`.red`のCSSルールには、値が`background`の`rgb(255, 0, 0)`プロパティがあるはずです。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
