---
id: 61a5bfe091060f1d6a629dd0
title: ステップ 64
challengeType: 0
dashedName: step-64
---

# --description--

カラーストップがなくても、緑のマーカーの色が赤のマーカーと同じポイントで変化していることに気づいたかもしれません。最初の色は開始位置（0%）、2番目は中央（50%）、最後はグラデーションラインの終わり（100%）にあります。

`linear-gradient` 関数はこれらの値を自動的に計算し、デフォルトではグラデーションラインに沿って色を均等に配置します。

`.red` の CSS ルールで、`linear-gradient` 関数から3つのカラーストップを削除してコードをすっきりさせてください。

# --hints--

`.red` の CSS ルールには `background` プロパティが `linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))` に設定されている必要があります。

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}
--fcc-editable-region--

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
