---
id: 61b0a44a6b865738ba49b9fb
title: ステップ 86
challengeType: 0
dashedName: step-86
---

# --description--

最後に、各マーカーにわずかな影を追加して、よりリアルに見えるようにします。

`box-shadow` プロパティを使うと、要素の周りに1つ以上の影を適用できます。基本的な構文は次の通りです。

```css
box-shadow: offsetX offsetY color;
```

`offsetX` と `offsetY` の値がどのように機能するか説明します：

* `offsetX` と `offsetY` はどちらも `px` やその他のCSS単位で数値を受け入れます
* 正の `offsetX` の値は影を右に移動させ、負の値は左に移動させます
* 正の `offsetY` の値は影を下に移動させ、負の値は上に移動させます
* `0` と `offsetX` のいずれかまたは両方にゼロ（`offsetY`）の値を使いたい場合、単位を付ける必要はありません。すべてのブラウザはゼロが変化なしを意味することを理解しています。

影の高さと幅は、影を適用する要素の高さと幅で決まります。さらに、影の広がりを調整するためのオプションの `spreadRadius` の値も使えます。これは後で詳しく説明します。

まずは赤いマーカーにシンプルな影を追加しましょう。

`.red` のCSSルールに、`box-shadow` プロパティを追加し、`5px` に `offsetX`、`5px` に `offsetY`、`red` に `color` の値を設定してください。

# --hints--

`.red` のCSSルールには、`box-shadow` 省略形プロパティがあり、値は `5px 5px red` であるべきです。

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
