---
id: 61b093329e7fc020b43b1bb5
title: ステップ 75
challengeType: 0
dashedName: step-75
---

# --description--

`rgb` 関数を使って色を設定する方法はすでに知っていますね。`rgb` 色にアルファチャンネルを追加するには、代わりに `rgba` 関数を使います。

`rgba` 関数は `rgb` 関数と同じように動作しますが、アルファチャンネル用に `0` から `1.0` までのもう1つの数値を取ります。

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

アルファチャンネルは `hsl` と `hex` の色でも使えます。すぐにその方法も見ていきます。

`.sleeve` ルールでは、`rgba` 関数を使って `background-color` プロパティを純粋な白で不透明度50％に設定してください。

# --hints--

`.sleeve` の CSS ルールには、`background-color` プロパティが `rgba(255, 255, 255, 0.5)` に設定されているはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor === 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
