---
id: 65f055b9190fc41ca35549b8
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

Python 提供了各种方法，它们的名称中包含前导和尾随双下划线。 你可能已经熟悉一些，例如 `__init__` 和 `__str__`。 这些遵循 `__<name>__` 命名模式的方法被称为特殊方法、魔术方法或 *dunder* （代表双下划线）方法。

在类中定义特殊方法会影响该类的行为。 它们在特定情况下在后台被调用（例如，实例化期间为 `__init__`，打印对象或将对象传递给 `str()` 时为 `__str__`）。 在这个项目中，你将学习一些最常用的特殊方法。

现在，将 `x` 分配给 `Vector` 对象的 `x` 属性。

# --hints--

你应该将 `x` 分配给 `self.x`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
