---
id: 65f058e06f34fd1f0ee6e55d
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

接下来，在 `Vector` 类中，创建一个空的 `norm` 方法并为其提供 `self` 参数。

# --hints--

你的 `Vector` 类应该有一个 `norm` 方法。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

你的 `norm` 方法应该有一个 `self` 参数。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
