---
id: 65f0694fb2296f3caadf8347
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

向量 $\mathbf{a}$ 的长度或范数通常表示为 $\\| \mathbf{a} \\|$。 它可以计算为其各分量平方和的平方根：
\\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2}
\\]

计算向量范数并从 `norm` 方法返回结果。

# --hints--

你的 `norm` 方法应该使用提供的公式来计算并返回一个表示当前向量实例范数的数字。 不要近似该值。

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
