---
id: 65f06a8e5a57673d700c79c3
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

在 `Vector` 类之外，创建 `Vector` 的实例，传递整数 `2` 和 `3` 作为 `x` 和 `y` 值，并将其分配给名为 `v1` 的变量。

# --hints--

你应该声明一个变量 `v1` 并为其分配 `Vector(2, 3)`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
