---
id: 65f01969115f933073b6be03
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

矢量是具有长度（或大小）和方向的对象，它不能用单个数字来表示。 在物理学中，矢量通常用于表示力、速度、加速度和其他量。

在这个项目中，你将构建一个向量空间，即一个集合，其中元素之间定义了一系列操作。 你很快就会了解所有细节。

现在，通过声明一个名为 `Vector` 的空类来启动项目。

# --hints--

你应该声明一个名为 `Vector` 的空类。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
