---
id: 65f056a405239e1dc4cc2854
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

现在，将 `y` 分配给 `Vector` 对象的 `y` 属性。

# --hints--

你应该将 `y` 分配给 `self.y`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").find_variable("self.y").is_equivalent("self.y = y")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
--fcc-editable-region--
```
