---
id: 65f07c9b1ffb814d856dcffc
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

通过打印该方法返回的值来测试 `norm` 是否按预期工作。

# --hints--

你应该在 `v1` 上调用 `norm()` 并打印结果。

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1.norm())")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
--fcc-editable-region--
```
