---
id: 65f03d9f92eac9183a4d3281
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

一个向量可以由欧几里得平面中的两个坐标 `x` 和 `y` 来定义。 轴的原点和点 `(x, y)` 之间的距离将是其长度或范数。 并且向量方向将指向 `(x, y)`。

<img alt="a 2-dimensional vector of coordinates (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

在 `Vector` 类中，创建一个 `__init__` 方法并为其提供三个参数，`self`、`x` 和 `y`。

# --hints--

你应该在 `Vector` 类中定义一个 `__init__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

你的 `__init__` 方法应该采用三个参数：`self`、`x` 和 `y`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
