---
id: 65f3fe07cc763212efe91285
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

一个向量可以有 `n` 个维度（分量）。 这是一个三维向量的表示：

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

到目前为止，你已经创建了一个二维向量。 你希望能够表示具有不同维数的向量，而无需针对每种特定情况重写必要的代码。 为此，你将使用继承。

首先将 `Vector` 类重命名为 `R2Vector` ，以指定该类将处理二维向量。 记住也要修改 `v1` 的实例。

# --hints--

你应该将 `Vector` 类重命名为 `R2Vector`。 记住也要修改 `v1` 的实例。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
