---
id: 65f3fa60a93b84110b3f1708
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

你想要返回一个有意义的字符串表示形式。 例如，在 `v1` 的情况下，你想要返回一个包含具有向量分量值的元组的字符串：`(2, 3)`。

从 `__str__` 方法中，返回一个表示向量的字符串，该字符串作为包含按顺序排列的向量分量的元组。 然后，走出 `Vector` 类并打印 `v1` 来检查结果。

# --hints--

你的 `__str__` 方法应该返回一个字符串，该字符串表示向量作为按顺序包含向量分量的元组。

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

你应该打印 `v1`。

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
