---
id: 65f3fa097f26b510db6c710b
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

太好了，`norm` 正常工作。现在，如果你尝试打印 `v1`，你将得到对象的默认字串表现（类似 `<__main__.Vector object at 0x11eb778>`）。

在 `Vector` 类中，声明一个空的 `__str__` 方法来实现可读的字符串表示。 记得给它一个 `self` 参数。

注意，直到 `__str__` 返回字符串时才打印 `v1`，否则你将得到 `TypeError`，因为 `__str__` 必须始终返回字符串。

# --hints--

你应该在 `Vector` 类中定义一个 `__str__` 方法。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

你的 `__str__` 方法应该有一个 `self` 参数。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
