---
id: 65f407ea37ad6e181b90462e
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

你可以像之前那样将每个 `i` 参数赋值给 `self.i`。虽然在这种情况下你只需重复几行代码，避免这种重复的一种方法是使用 `super()` 函数。`super()` 使你能够隐式地参考父类：`super().__init__(x, y)` 调用父类的 `__init__` 方法。

向你的 `__init__` 方法添加 `super().__init__(x, y)` 调用。

# --hints--

你应该在 `__init__` 方法中有一个 `super().__init__(x, y)` 调用。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
