---
id: 65f4535bbdb28d436ff3ddc9
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

从输出中可以看到，`__dict__` 包含实例属性的值。 你不需要明确地将 `self.x` 和 `self.y` 的平方相加，而是要迭代存储在 `__dict__` 中的值来计算范数的值。

在 `norm` 方法主体中，将括号中的内容替换为生成器表达式，该表达式将 `self.__dict__.values()` 中的每个值 `val` 提升为 `2` 次方。 另外，将该生成器表达式作为参数传递给 `sum` 函数，以便在计算平方根之前将所有平方相加。

# --hints--

你应该从 `norm()` 方法返回 `sum(val**2 for val in self.__dict__.values())**0.5`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in self.__dict__.values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)

```
