---
id: 65f4520e363e2642f8112e33
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

正如你所见，有些事情不太对劲。 从 `R2Vector` 继承的 `norm` 和 `__str__` 方法无法适应三维向量。 它们的实施必须更加灵活。

Python 中的每个对象都有一个特殊的属性，名为 `__dict__`，它是一个存储对象属性的字典。

删除现有的 `print` 调用。 然后，打印 `v1` 和 `v2` 向量的 `__dict__` 属性来查看它们是什么样子。

# --hints--

你的代码中不应该有 `print(v2)` 和 `print(v2.norm())`。

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

你应该打印 `v1` 和 `v2` 的 `__dict__` 属性。

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
