---
id: 65f412c208c3791fee305acf
title: Step 17
challengeType: 20
dashedName: step-25
---

# --description--

在 Python 中，你可以通过将 `*` 作为附加参数添加到函数或方法签名来强制使用仅限关键字参数。 修改两个 `__init__` 方法，添加 `*` 作为第二个参数（在 `self` 之后）。 此后放置的每个参数都需要在函数/方法调用中使用关键字参数。

这意味着你也需要修改 `super().__init__(x, y)` 调用。 通过赋予 `x` 值 `x`，以及赋予 `y` 值 `y` 来实现。

最后，使用关键字参数修改 `v1` 和 `v2` 的实例。

# --hints--

你应该通过在两个 `__init__` 方法中添加 `*` 作为第二个参数来强制执行关键字参数。

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

你应该修改 `super.__init__(x, y)` 调用以使用关键字参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

你应该将 `v1` 的赋值修改为 `R2Vector(x=2, y=3)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

你应该将 `v2` 的赋值修改为 `R3Vector(x=2, y=2, z=3)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
