---
id: 65f40051d6b09a139f253e8e
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

继承使你能够从现有类定义一个类。 新的类（称为 child）继承了现有类（称为 parent）的所有方法和属性。

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass
    
Oak().sprout() # Output: Making new leaves!
```

在上面的例子中，子类 `Oak` 继承自 `Tree`，并从父类 `Tree` 继承了 `sprout` 方法。

创建一个名为 `R3Vector` 的新类，并按照上面的示例使其从 `R2Vector` 类继承。

# --hints--

你应该定义一个新类 `R3Vector`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

你的新类 `R3Vector` 应该继承自 `R2Vector`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
