---
id: 65f8720685ec351abef26740
title: 步骤 24
challengeType: 20
dashedName: step-24
---

# --description--

要获得正确的字符串表示，请调用 `tuple()` 构造函数并将你在上一步中编写的生成器表达式作为参数传递给它。

# --hints--

你应该将 `__str__` 方法返回的表达式传递给 `tuple()` 构造函数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").has_return("tuple(getattr(self, i) for i in vars(self))")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return (getattr(self, i) for i in vars(self))
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
