---
id: 65f8709620b2ce1a62608f5a
title: 步骤-23
challengeType: 20
dashedName: step-23
---

# --description--

当你需要从字符串输入开始动态访问某些属性时，内置的 `getattr()` 函数就是你所需要的。 它将对象作为其第一个参数，并将包含属性名称的字符串作为其第二个属性。

开始修复 `__str__` 方法，方法是将 `__str__()` 返回的字符串替换为一个生成器表达式，该生成器表达式遍历对象属性并为每个属性 `i` 调用 `getattr()`。

# --hints--

你应该返回一个迭代 `vars(self)` 的生成器表达式。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

你应该返回一个使用 `i` 作为迭代变量的生成器表达式。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

你应该返回一个生成器表达式，该表达式为 `vars(self)` 中的每个项目 `i` 调用 `getattr(self, i)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
