---
id: 65f8749b13774b1d2e4a7fba
title: 步骤 27
challengeType: 20
dashedName: 步骤-27
---

# --description--

`__str__` 方法返回对象的人类可读的字符串表示形式，而 `__repr__` 方法应该返回实例化该对象所需的字符串。

在 `R2Vector` 类中添加一个带有 `self` 参数的 `__repr__` 方法。

# --hints--

你的 `R2Vector` 类应该有一个 `__repr__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__repr__")`)) })
```

你的 `__repr__` 方法应该有一个 `self` 参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--

--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
