---
id: 65f876d17832001e8e1abb05
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

声明一个变量 `arg_list`。 为其提供列表推导的值，该列表推导针对 `vars(self).items()` 中的每个键值对迭代 `key` 和 `val`，并在每次迭代时计算 f 字符串 `f'{key}={val}'`。

# --hints--

你应该声明一个变量 `arg_list`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_variable("arg_list")`)) })
```

你应该分配一个列表推导，从 `vars(self).items()` 迭代到 `arg_list`。

```js
({ test: () => runPython(`
import ast
var = _Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list")
assert isinstance(var.tree.value, ast.ListComp), "Expected arg_list to be a list comprehension"
assert (actual := var.find_comp_iters()[0].is_equivalent("vars(self).items()")), f"Expected vars(self).items(), got {actual}"
`) })
```

分配给 `arg_list` 的列表推导应该使用 `key` 和 `val` 来迭代 `vars(self).items()`。

```js
({ test: () => assert(runPython(`
_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_targets()[0].is_equivalent("key, val")
`)) })
```

分配给 `arg_list` 的列表推导应该为 `vars(self).items()` 中的每个键值对计算字符串 `f'{key}={val}'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_expr().is_equivalent("f'{key}={val}'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
