---
id: 65f9b17a638f0d0dcce8c354
title: 步骤 38
challengeType: 20
dashedName: step-38
---

# --description--

要创建向量空间，你需要定义向量在几种情况下的行为方式。 可以添加向量，形成新的向量。

可以实现特殊方法 `__add__` 来覆盖使用 `+` 运算符将两个对象相加时默认发生的情况。

现在，尝试添加两个 `R2Vector` 或 `R3Vector` 实例将引发异常。 在 `R2Vector` 类中创建一个空的 `__add__` 方法，并为其提供两个参数：`self` 和 `other`。

# --hints--

你应该在 `R2Vector` 类中定义一个 `__add__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

你的 `__add__` 方法应该有两个参数，`self` 和 `other`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--

--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
