---
id: 65f9cb678070ca1668898c70
title: 步骤 44
challengeType: 20
dashedName: step-44
---

# --description--

同样，当两个对象相加时，会默认调用 `__add__` 方法，而当进行减法时，则会隐式调用 `__sub__` 方法。

现在，定义一个空的 `__sub__` 方法并给出两个参数：`self` 和 `other`。 在你的新方法中，创建一个 if 语句来检查 `self` 和 `other` 是否不属于同一个类，并返回 `NotImplemented`，就像你之前所做的那样。

# --hints--

你应该在 `R2Vector` 类中定义一个 `__sub__` 方法。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__sub__")
    `))
})
```

你的 `__sub__` 方法应该有两个参数，`self` 和 `other`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").has_args("self, other")
    `))
})
```

你应该创建一个 `if` 语句来检查 `self` 和 `other` 是否不属于同一类。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
    `))
})
```

你应该从 `if` 主体中返回 `NotImplemented`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
