---
id: 65f93b67169a9c703264458a
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

由于该方法应返回字符串以在继承时实例化 `R2Vector` 和 `R3Vector` 的对象，因此你无法构建指定类名的字符串。

你可以使用 `__class__.__name__` 访问类的名称。 向 `__repr__` 方法添加 `return` 语句并返回实例化对象所需的字符串。

# --hints--

你应该使用 `__class__.__name__` 来构建实例化向量对象所需的 f 字符串并从 `__repr__` 方法返回它。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_return("f'{self.__class__.__name__}({args})'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
