---
id: 65f9bc10a9fb1612e066e9e8
title: 步骤 42
challengeType: 20
dashedName: step-42
---

# --description--

当添加两个向量时，一个向量的每个分量都会添加到另一个向量的相同分量中。 例如，将 `(1, 2)` 和 `(2, 4)` 相加会生成第三个向量 `(3, 6)`，其中 `3` 是 `1` 和 `2` 之和，而 `6` 是 `2` 和 `4` 之和。

`kwargs` 字典将包含实例化两个向量相加的同一类的新向量所需的键值对。

将空字典转换为字典推导式，迭代 `vars(self)` 并为每个键（`i`）创建一个键值对，其中键与当前迭代的键相同，其值是两个操作数的 `i` 属性值的总和。

# --hints--

你应该将分配给 `kwargs` 的空字典转换为对 `vars(self)` 进行迭代的字典推导。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

分配给 `kwargs` 的字典理解应该使用变量 `i` 来迭代 `vars(self)`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

分配给 `kwargs` 的字典理解应该使用 `i` 作为键。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

分配给 `kwargs` 的字典推导式应该将 `self` 的 `i` 属性与 `other` 的 `i` 属性之和分配给 `vars(self)` 中每个 `i` 的 `i` 键。 为此，请使用 `getattr()` 函数。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
