---
id: 65f9b9710cca621244d3bde1
title: 步骤 40
challengeType: 20
dashedName: step-40
---

# --description--

在 Python 中，`NotImplemented` 是一个特殊值，用于表示某个操作未针对特定情况实现。

`NotImplemented` 不会立即引发异常。 相反，它发出信号询问另一个操作数如何执行操作。 如果请求无法满足，则默认返回 `TypeError`。

因为你只希望能够对属于同一类的两个向量进行求和，所以从上一步中创建的 `if` 语句返回 `NotImplemented`。

# --hints--

你应该从 `if` 主体中返回 `NotImplemented`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
