---
id: 65f9c2c2006feb1428ad2d4c
title: 第 43 步
challengeType: 20
dashedName: step-43
---

# --description--

你可以使用 `**` 运算符将字典解包为关键字参数：

```py
def spam(a, b):
    return a + b

my_dict = {a: 11, b: 4}

spam(**my_dict) # 15
```

通过使用 `__class__` 并传递 `**kwargs` 作为参数来返回当前类的实例。

# --hints--

你应该从 `__add__` 方法返回 `self.__class__(**kwargs)`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").has_return("self.__class__(**kwargs)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
