---
id: 65f9dd6e5a08af19c196c2df
title: 步骤 45
challengeType: 20
dashedName: step-45
---

# --description--

一个向量与另一个向量相减的结果是通过计算它们各个分量的差值而得到的。 例如，从 `(7, 3)` 中减去 `(2, 4)` 会生成第三个向量 `(5, -1)`，其中 `5` 是 `7` 和 `2` 之间的差值，而 `-1` 是 `3` 和 `4` 之间的差值。

就像之前在 `__add__` 方法中所做的那样，在 `if` 语句后声明一个 `kwargs` 变量。 为其分配一个字典推导，该字典推导遍历对象属性并为每个键 `i` 创建一个键值对，其中键与当前迭代的键相同，其值是两个操作数的 `i` 属性值之间的差值。

# --hints--

你应该声明一个变量 `kwargs` 并为其分配一个遍历对象属性的字典理解。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

分配给 `kwargs` 的字典理解应该使用变量 `i` 来迭代对象属性。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

分配给 `kwargs` 的字典理解应该使用 `i` 作为键。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

分配给 `kwargs` 的字典理解应该将 `self` 的 `i` 属性与 `other` 的 `i` 属性之间的差异分配给每个 `i` 的 `i` 键。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
