---
id: 6659875502b6d7765498f324
title: 步骤 26
challengeType: 20
dashedName: step-26
---

# --description--

使用 f 字符串修改你的 `print` 调用，以便它们分别产生输出 `v1 = (2, 3)` 和 `v2 = (2, 2, 3)`。

# --hints--

你应该使用 f 字符串来插入 `v1` 并生成输出 `v1 = (2, 3)`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v1 = {v1}')")`)) })
```

你应该使用 f 字符串来插入 `v2` 并生成输出 `v2 = (2, 2, 3)`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v2 = {v2}')")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())
--fcc-editable-region--
```
