---
id: 666832e427d70bc5219dc62a
title: 步骤 36
challengeType: 20
dashedName: step-36
---

# --description--

从输出中可以看出，虽然你在类中定义了 `__getattr__`，但是此方法尚未被调用。 发生这种情况是因为默认属性访问是通过 `__getattribute__` 进行的。 因此，在终端上就可以看到打印的属性值。

现在修改最后两行代码，以便在两个 `print()` 调用中访问 `v1` 的 `z` 属性。 这次，你将访问 `v1` 没有的属性。 结果，`__getattribute__` 将引发错误并调用 `__getattr__`。

# --hints--

你应该修改代码，以使用点运算符和 `getattr()` 函数打印 `v1` 的 `z` 属性。

```js
({ test: () => assert(runPython(`
n = _Node(_code)
(n.has_call("print(v1.z)") and n.has_call("print(getattr(v1, 'z'))")) or n.has_call("print(v1.z, getattr(v1, 'z'))") or n.has_call("print(getattr(v1, 'z'), v1.z)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __getattr__(self, attr):                
        return 'calling __getattr__'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')
--fcc-editable-region--
print(v1.x)
print(getattr(v1, 'x'))
--fcc-editable-region--
```
