---
id: 666827a6fd0dbaafe8330ea6
title: 步骤35
challengeType: 20
dashedName: step-35
---

# --description--

`__getattr__` 是另一种在访问属性方面发挥作用的特殊方法。

`__getattribute__` 的默认实现是当请求的属性不是实例属性或不在类树中时引发 `AttributeError`。

在这种情况下，如果类定义了 `__getattr__`，则会调用它。 当通常的属性访问过程失败时，你可以将其视为一种后备。

将 `__getattribute__` 方法变为 `__getattr__`，并将该方法返回的字符串修改为 `'calling __getattr__'`。

# --hints--

你应该在 `R2Vector` 类中定义一个名为 `__getattr__` 的方法，它有两个参数，`self` 和 `attr`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_args("self, attr")`)) })
```

你的 `__getattr__` 方法应该返回字符串 `'calling __getattr__'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_return("'calling __getattr__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __getattribute__(self, attr):                
        return 'calling __getattribute__'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

print(v1.x)
print(getattr(v1, 'x'))

```
