---
id: 66682150151af29efec9727d
title: 步骤 33
challengeType: 20
dashedName: step-33
---

# --description--

每当你尝试访问实例属性时，都会在后台调用 `__getattribute__` 方法。 如果在实例级别未找到该属性，则该方法将在类级别搜索它，并最终在父类中搜索它。

在你的类中，定义一个带有两个参数 `self` 和 `attr` 的 `__getattribute__` 方法，并使其返回字符串 `'calling __getattribute__'`。 你将暂时覆盖默认实现，只是为了了解属性访问是如何工作的。

# --hints--

你应该在 `R2Vector` 类中定义一个名为 `__getattribute__` 的方法，它有两个参数，`self` 和 `attr`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_args("self, attr")`)) })
```

你的 `__getattribute__` 方法应该返回字符串 `'calling __getattribute__'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_return("'calling __getattribute__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

```
