---
id: 68e4276483b0c71658aab1cd
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Ora è il momento di vedere il risultato.

Inizia aggiungendo un `console.log()` con la stringa `"Result of trimming whitespace from both ends:"`.

Poi, sotto quel `console.log()`, aggiungi un altro `console.log()` con la variabile `cleanedInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti registrare la stringa `"Result of trimming whitespace from both ends:"` nella console.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

Dovresti registrare la variabile `cleanedInput` nella console.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
