---
id: 68e4276483b0c71658aab1cc
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Il tuo valore `userInput` ha degli spazi extra all’inizio e alla fine della stringa. Per creare una nuova stringa con quegli spazi rimossi, puoi usare il metodo `trim()`.  

Ecco un esempio:

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

Crea una variabile chiamata `cleanedInput` e assegnale il risultato dell’uso del metodo `trim()` sulla variabile `userInput`.

# --hints--

Dovresti avere una variabile `cleanedInput`.

```js
assert.exists(cleanedInput);
```

La tua variabile `cleanedInput` dovrebbe essere una stringa.

```js
assert.isString(cleanedInput);
```

Dovresti assegnare `userInput.trim()` alla tua variabile `cleanedInput`.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
