---
id: 68e4276483b0c71658aab1ce
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Ci saranno momenti in cui vorrai rimuovere solo gli spazi bianchi all'inizio della stringa. È qui che il metodo `trimStart()` torna utile.

Ecco un esempio:

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

Ricorda che le stringhe sono immutabili. Questo metodo creerà una nuova stringa con il risultato, senza modificare quella originale.

Crea una variabile `trimmedStart` e assegnale il risultato dell'uso del metodo `trimStart()` sulla variabile `userInput`.

# --hints--

Dovresti avere una variabile `trimmedStart`.

```js
assert.exists(trimmedStart);
```

La tua variabile `trimmedStart` dovrebbe essere una stringa.

```js
assert.isString(trimmedStart);
```

Dovresti assegnare `userInput.trimStart()` alla tua variabile `trimmedStart`.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
