---
id: 68e4276483b0c71658aab1ca
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop, eserciterai l’uso di vari metodi per stringhe utilizzati per formattare stringhe.

Inizia creando una variabile chiamata `userInput` e assegnale il seguente valore: `"   Hello World!   "`.

# --hints--

Dovresti avere una variabile `userInput`.

```js
assert.exists(userInput);
```

La tua variabile `userInput` dovrebbe essere una stringa.

```js
assert.isString(userInput);
```

Dovresti assegnare `"   Hello World!   "` alla tua variabile `userInput`. Assicurati che ci siano tre spazi prima di `H` e tre spazi dopo `!`.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
