---
id: 68e4276483b0c71658aab1d0
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Ci saranno momenti in cui dovrai creare una nuova stringa con una diversa capitalizzazione.

Per creare una nuova stringa con tutte le lettere convertite in maiuscolo, puoi usare il metodo `toUpperCase()`.

Ecco un esempio:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

Crea una variabile chiamata `upperCaseInput` e assegnale il risultato dell’uso del metodo `toUpperCase()` sulla variabile `cleanedInput`.

# --hints--

Dovresti avere una variabile `upperCaseInput`.

```js
assert.exists(upperCaseInput);
```

La tua variabile `upperCaseInput` dovrebbe essere una stringa.

```js
assert.isString(upperCaseInput);
```

Dovresti assegnare `cleanedInput.toUpperCase()` alla tua variabile `upperCaseInput`.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
