---
id: 68e4276483b0c71658aab1d3
title: Passaggio 14
challengeType: 1
dashedName: step-14
---

# --description--

Come ricorderai dalle lezioni e dai workshop precedenti, puoi estrarre porzioni di una stringa usando il metodo `slice()`.

Ecco un promemoria:

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

Dovresti modificare l'assegnazione `camelCasedVersion` esistente sostituendo la stringa vuota con `lowercaseWord.slice(0, 5)`.

Ora dovresti vedere la parola `"camel"` nella console.

# --hints--

La tua variabile `camelCasedVersion` non dovrebbe essere una stringa vuota.

```js
assert.isNotEmpty(camelCasedVersion);
```

Dovresti assegnare `lowercaseWord.slice(0, 5)` alla variabile `camelCasedVersion`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

Non dovresti assegnare la stringa letterale `"camel"` a `camelCasedVersion`. Rileggi le istruzioni per l'approccio corretto.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
