---
id: 68e4276483b0c71658aab1d1
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Se hai bisogno di creare una nuova stringa in cui tutte le lettere sono in minuscolo, puoi usare il metodo `toLowerCase()`.  

Inizia creando una variabile chiamata `lowerCaseInput` e assegnale il risultato dell'uso del metodo `toLowerCase()` sulla variabile `cleanedInput`. 

Poi stampa la stringa `"Result of using the toLowerCase() method:"` nella console.

Infine, aggiungi un'altra `console.log()` per la variabile `lowerCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti avere una variabile `lowerCaseInput`.

```js
assert.exists(lowerCaseInput);
```

La tua variabile `lowerCaseInput` dovrebbe essere una stringa.

```js
assert.isString(lowerCaseInput);
```

Dovresti assegnare `cleanedInput.toLowerCase()` alla tua variabile `lowerCaseInput`.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

Dovresti stampare la stringa `"Result of using the toLowerCase() method:"` nella console.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

Dovresti stampare la variabile `lowerCaseInput` nella console.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
