---
id: 68e4276483b0c71658aab1d2
title: Passaggio 13
challengeType: 1
dashedName: step-13
---

# --description--

Ci saranno momenti in cui dovrai usare più metodi per creare una nuova stringa. Nei prossimi passaggi, userai alcuni dei metodi che hai imparato finora per creare una nuova stringa che usa la convenzione di denominazione camelCase.

Inizia creando una variabile chiamata `lowercaseWord` e assegnale la stringa `"camelcase"`.

Sotto quella variabile `lowercaseWord`, crea un'altra variabile chiamata `camelCasedVersion` e assegnale per ora una stringa vuota.

Sotto quella variabile `camelCasedVersion`, aggiungi un `console.log()` con la stringa `"Camel cased version:"` seguita da un altro `console.log()` per la variabile `camelCasedVersion`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti avere una variabile chiamata `lowercaseWord`.

```js
assert.exists(lowercaseWord);
```

La tua `lowercaseWord` dovrebbe essere una stringa.

```js
assert.isString(lowercaseWord);
```

Dovresti assegnare la stringa `"camelcase"` alla tua variabile `lowercaseWord`.

```js
assert.equal(lowercaseWord, "camelcase");
```

Dovresti avere una variabile `camelCasedVersion`.

```js
assert.exists(camelCasedVersion);
```

La tua variabile `camelCasedVersion` dovrebbe essere una stringa.

```js
assert.isString(camelCasedVersion);
```

Dovresti assegnare una stringa vuota alla tua variabile `camelCasedVersion`.

```js
assert.isEmpty(camelCasedVersion);
```

Dovresti stampare la stringa `"Camel cased version:"` nella console.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

Dovresti stampare la variabile `camelCasedVersion` nella console.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
