---
id: 68f1179eaea3bb59ba434386
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Per eliminare gli spazi bianchi alla fine di una stringa, puoi usare il metodo `trimEnd()`.

Ecco un esempio:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

Crea una variabile chiamata `trimmedEnd` e assegnale il risultato dell'uso del metodo `trimEnd()` sulla variabile `userInput`.

# --hints--

Dovresti avere una variabile `trimmedEnd`.

```js
assert.exists(trimmedEnd);
```

La tua variabile `trimmedEnd` dovrebbe essere una stringa.

```js
assert.isString(trimmedEnd);
```

Dovresti assegnare `userInput.trimEnd()` alla tua variabile `trimmedEnd`.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
