---
id: 68f3eeaa65614107ecd5f1f2
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Sotto il tuo primo `console.log()`, aggiungi un altro `console.log()` per registrare la variabile `userInput` nella console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti registrare la variabile `userInput` nella console.

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
