---
id: 68f3f3d65d51d80ac0b412fc
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Ora è il momento di vedere i risultati della rimozione degli spazi bianchi all'inizio della stringa.

Sotto la tua variabile `trimmedStart`, aggiungi una `console.log()` con la stringa `"After using the trimStart() method, leading spaces removed:"`.

Poi, sotto quella `console.log()`, aggiungi un'altra `console.log()` con la variabile `trimmedStart`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti stampare la stringa `"After using the trimStart() method, leading spaces removed:"` nella console.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

Dovresti stampare la variabile `trimmedStart` nella console.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
