---
id: 68f3f890010c870e17bbe754
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Ora è il momento di vedere i risultati.

Sotto la tua variabile `upperCaseInput`, aggiungi un `console.log()` con la stringa `"Result of using the toUpperCase() method:"`.

Poi, sotto quel `console.log()`, aggiungi un altro `console.log()` con la variabile `upperCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti registrare la stringa `"Result of using the toUpperCase() method:"` nella console.

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

Dovresti registrare la variabile `upperCaseInput` nella console.

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
