---
id: 68f3f5ca1d2a490c97f14f03
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Ora è il momento di vedere i risultati del taglio degli spazi bianchi alla fine della stringa.

Sotto la tua variabile `trimmedEnd`, aggiungi una `console.log()` con la stringa `"After using the trimEnd() method, trailing spaces removed:"`.

Poi, sotto quella `console.log()`, aggiungi un'altra `console.log()` con la variabile `trimmedEnd`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti stampare la stringa `"After using the trimEnd() method, trailing spaces removed:"` nella console.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

Dovresti stampare la variabile `trimmedEnd` nella console.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
