---
id: 68f40925d810ae14b2d5a169
title: Passaggio 16
challengeType: 1
dashedName: step-16
---

# --description--

In questo momento la tua variabile `camelCasedVersion` contiene il valore di `camelC`. Ma il risultato desiderato è `camelCase`. 

Per ottenere le lettere `ase` rimanenti, puoi usare il metodo `slice()` con un indice negativo. Ricorda che usare indici negativi indica al metodo `slice()` di estrarre dalla fine. 

Ecco un esempio di promemoria:

```js
const str = "abcd";
str.slice(-2); // cd
```

Usa l’operatore `+` per concatenare `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` con il risultato dell’uso del metodo `slice()` su `lowercaseWord`. Passa il numero `-3` al metodo `slice()`.

E con quest’ultima modifica, il tuo workshop è completo!

# --hints--

Dovresti concatenare `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` con `lowercaseWord.slice(-3)`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

Non dovresti assegnare la stringa letterale `"camelCase"` a `camelCasedVersion`. Rileggi le istruzioni per l’approccio corretto.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
