---
id: 68f405c8c9df44132c8fe1bb
title: Passaggio 15
challengeType: 1
dashedName: step-15
---

# --description--

Come ricorderai dalle lezioni precedenti, la convenzione di denominazione camelCase inizia con una lettera minuscola per la prima parola, e ogni parola successiva inizia con una lettera maiuscola seguita da lettere minuscole.

Ecco alcuni esempi di promemoria:

```md
freeCodeCamp
camelCase
learningIsFun
```

La seconda parola nella variabile `lowercaseWord` è `"case"`. Per consultare `c` in quella parola, puoi usare `lowercaseWord[5]`.

Aggiorna l’assegnazione `camelCasedVersion` per usare l’operatore `+` per concatenare `lowercaseWord.slice(0, 5)` con il risultato dell’uso del metodo corretto per convertire le stringhe in maiuscolo su `lowercaseWord[5]`.

Ora dovresti vedere `camelC` nella console.

# --hints--

Dovresti concatenare `lowercaseWord.slice(0, 5)` con `lowercaseWord[5].toUpperCase()`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

Non dovresti assegnare la stringa letterale `"camelC"` a `camelCasedVersion`. Rileggi le istruzioni per l’approccio corretto.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
