---
id: 6140d4bc9db3c81c51a09ab7
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

패턴을 계속해서, 다음 `.cabin` 요소들을 선택하고 특정 규칙을 적용하세요:

- 두 번째 `.cabin`는 `right` 속성을 `17%`로, `top` 속성을 `93.5%`로 설정해야 합니다.
- 세 번째 `.cabin`는 `right` 속성을 `67%`로, `top` 속성을 `93.5%`로 설정해야 합니다.
- 네 번째 `.cabin`는 `left` 속성을 `-8.5%`로, `top` 속성을 `50%`로 설정해야 합니다.
- 다섯 번째 `.cabin`는 `left` 속성을 `17%`로, `top` 속성을 `7%`로 설정해야 합니다.
- 여섯 번째 `.cabin`는 `right` 속성을 `17%`로, `top` 속성을 `7%`로 설정해야 합니다.

# --hints--

`.cabin:nth-of-type(2)` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)'));
```

`.cabin:nth-of-type(2)` 선택자는 `right` 속성을 `17%`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.right === '17%');
```

`.cabin:nth-of-type(2)` 선택자는 `top` 속성을 `93.5%`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.top === '93.5%');
```

`.cabin:nth-of-type(3)` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)'));
```

`.cabin:nth-of-type(3)` 선택자는 `right` 속성을 `67%`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.right === '67%');
```

`.cabin:nth-of-type(3)` 선택자는 `top` 속성을 `93.5%`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.top === '93.5%');
```

`.cabin:nth-of-type(4)` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)'));
```

`.cabin:nth-of-type(4)` 선택자는 `left` 속성을 `-8.5%`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.left === '-8.5%');
```

`.cabin:nth-of-type(4)` 선택자는 `top` 속성을 `50%`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.top === '50%');
```

`.cabin:nth-of-type(5)` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)'));
```

`.cabin:nth-of-type(5)` 선택자는 `left` 속성을 `17%`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.left === '17%');
```

`.cabin:nth-of-type(5)` 선택자는 `top` 속성을 `7%`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.top === '7%');
```

`.cabin:nth-of-type(6)` 선택자가 있어야 합니다.

```js
const def = (s) => new __helpers.CSSHelp(document).getStyle(s)
assert(def('.cabin:nth-of-type(6)') || def('.cabin:last-of-type'));
```

`.cabin:nth-of-type(6)` 선택자는 `right` 속성을 `17%`로 설정해야 합니다.

```js
const right = (s) => new __helpers.CSSHelp(document).getStyle(s)?.right
assert.equal(right('.cabin:nth-of-type(6)') || right('.cabin:last-of-type'),'17%');
```

`.cabin:nth-of-type(6)` 선택자는 `top` 속성을 `7%`로 설정해야 합니다.

```js
const top = (s) => new __helpers.CSSHelp(document).getStyle(s)?.top
assert.equal(top('.cabin:nth-of-type(6)') || top('.cabin:last-of-type'),'7%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
--fcc-editable-region--

--fcc-editable-region--
```
