---
id: 6140c9d35015ae0ba0c250e8
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`div` 요소 안에 `body`를 추가하고 `class` 값을 가진 `wheel`를 지정하세요.

새로 만든 `div` 안에 `span` 값으로 설정된 `class`를 가진 `line` 요소 6개와 `div` 값으로 설정된 `class`를 가진 `cabin` 요소 6개를 추가하세요.

# --hints--

`div` 요소 안에 새로운 `body`를 만들어야 합니다.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

`div` 요소 안에 있는 `body`는 `class` 값을 가진 `wheel`를 가져야 합니다.

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

`.wheel` 요소 안에는 `span` 요소 6개가 있어야 합니다.

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

`span` 요소 안에 있는 `.wheel` 6개 요소는 `class` 값을 가진 `line`를 가져야 합니다.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

`.wheel` 요소 안에는 `div` 요소 6개가 있어야 합니다.

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

`div` 요소 안에 있는 `.wheel` 6개 요소는 `class` 값을 가진 `cabin`를 가져야 합니다.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
