---
id: 6140df547f09402144e40b92
title: 20단계
challengeType: 0
dashedName: step-20
---

# --description--

`animation-name` 속성은 `@keyframes` 규칙을 CSS 선택자에 연결하는 데 사용됩니다. 이 속성의 값은 `@keyframes` 규칙 이름과 일치해야 합니다. `.wheel` 선택자에 `animation-name` 속성을 `wheel`로 설정하세요.

`animation-duration` 속성은 애니메이션이 완료되는 데 걸리는 시간을 설정하는 데 사용됩니다. 시간은 초(`s`) 또는 밀리초(`ms`) 단위로 지정해야 합니다. `.wheel` 선택자에 `animation-duration` 속성을 `10s`로 설정하세요.

# --hints--

`.wheel` 선택자에 `animation-name` 속성을 `wheel`로 설정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationName === 'wheel');
```

`.wheel` 선택자에 `animation-duration` 속성을 `10s`로 설정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationDuration === '10s');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}
--fcc-editable-region--

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(2) {
  right: 17%;
  top: 93.5%;
}
.cabin:nth-of-type(3) {
  right: 67%;
  top: 93.5%;
}
.cabin:nth-of-type(4) {
  left: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(5) {
  left: 17%;
  top: 7%;
}
.cabin:nth-of-type(6) {
  right: 17%;
  top: 7%;
}

@keyframes wheel {
   0% {
     transform: rotate(0deg);
   }
   100% {
     transform: rotate(360deg);
   }
}
```
