---
id: 6140c7e645d8e905819f1dd4
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

표준 보일러플레이트로 시작하세요. `DOCTYPE` 선언, 언어가 영어로 설정된 `html` 요소, `head`와 `body` 요소를 추가하세요.

올바른 `meta`에 대한 `charset` 요소, `title` 요소, 그리고 `link` 파일에 대한 `./styles.css` 요소를 추가하세요.

`title`을 `Ferris Wheel`으로 설정하세요.

# --hints--

코드에 `DOCTYPE` 선언이 포함되어야 합니다.

```js
assert.match(code, /<!DOCTYPE/gi);
```

`DOCTYPE` 선언 뒤에 공백을 포함해야 합니다.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

문서 유형을 `html`으로 정의해야 합니다.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

`DOCTYPE` 선언을 유형 뒤에 `>`로 닫아야 합니다.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

`html` 요소는 `lang` 값을 가진 `en` 속성을 가진 여는 태그가 있어야 합니다.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

`html` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/html\s*>/);
```

`DOCTYPE` 선언은 HTML의 맨 앞에 있어야 합니다.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

여는 `head` 태그가 있어야 합니다.

```js
assert.match(code, /<head\s*>/i);
```

닫는 `head` 태그가 있어야 합니다.

```js
assert.match(code, /<\/head\s*>/i);
```

여는 `body` 태그가 있어야 합니다.

```js
assert.match(code, /<body\s*>/i);
```

닫는 `body` 태그가 있어야 합니다.

```js
assert.match(code, /<\/body\s*>/i);
```

`body` 요소는 `head` 요소 바로 다음 첫 번째 요소여야 합니다.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

`head`과 `body` 요소는 형제 요소여야 합니다.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

`head` 요소는 `html` 요소 내부의 첫 번째 요소여야 합니다.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

`head` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

코드에 `meta` 요소가 있어야 합니다.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

`meta` 요소는 `charset` 값을 가진 `UTF-8` 속성을 가져야 합니다.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

`meta` 요소는 `head` 요소 안에 있어야 합니다.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

코드에 `title` 요소가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.exists(title);
```

프로젝트 제목은 `Ferris Wheel`이어야 합니다.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

제목의 대소문자와 철자가 중요하다는 점을 기억하세요.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

코드에 `link` 요소가 있어야 합니다.

```js
assert.match(code, /<link/);
```

하나의 `link` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector('head > link'));
```

`link` 요소는 `rel` 값을 가진 `stylesheet` 속성을 가져야 합니다.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 요소는 `href` 값을 가진 `styles.css` 속성을 가져야 합니다.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
