---
id: 6140cbeec34e970dfe75e710
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`.wheel` 요소에 대한 선택자를 만드세요. `border`을 `2px solid black`로, `border-radius`을 `50%`로, `margin-left`를 `50px`로 설정하는 것부터 시작하세요.

# --hints--

`.wheel` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

`.wheel` 선택자에는 `border` 속성이 `2px solid black`으로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

`.wheel` 선택자에는 `border-radius` 속성이 `50%`으로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

`.wheel` 선택자에는 `margin-left` 속성이 `50px`으로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
