---
id: 617ace7d831f9c16a569b38a
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`.wheel` 선택자에 `max-height`과 `max-width` 속성을 모두 `500px`로 설정하세요.

# --hints--

`.wheel` 선택자에 `max-height` 속성을 `500px`로 설정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxHeight === '500px');
```

`.wheel` 선택자에 `max-width` 속성을 `500px`로 설정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxWidth === '500px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
}
--fcc-editable-region--
```
