---
id: 66bb6a9c2dd58b73cd759034
title: Crie uma aplicação de quadro de notas
challengeType: 14
dashedName: build-a-gradebook-app
---

# --description--

**Objetivo:** atender às histórias de usuário abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve ter uma função chamada `getAverage` que recebe um array de pontuações de testes como parâmetro e retorna a pontuação média.
1. Você deve ter uma função chamada `getGrade` que recebe a pontuação de um estudante como parâmetro e retorna uma string que representa uma nota com base na pontuação. Aqui estão as pontuações e suas notas correspondentes:

    | Faixa de Pontuação | Nota   |
    | ------------------ | ------ |
    | `100`              | `"A+"` |
    | `90 - 99`          | `"A"`  |
    | `80 - 89`          | `"B"`  |
    | `70 - 79`          | `"C"`  |
    | `60 - 69`          | `"D"`  |
    | `0 - 59`           | `"F"`  |

1. Você deve ter uma função chamada `hasPassingGrade`, que recebe uma pontuação como parâmetro e retorna `true` ou `false` dependendo do fato de a pontuação corresponder a uma nota de aprovação.
1. A função `hasPassingGrade` deve usar a função `getGrade` para obter a nota, e usá-la para determinar se a nota é de aprovação. Uma nota de aprovação é qualquer coisa diferente de `"F"`.
1. Você deve ter uma função chamada `studentMsg`, que recebe um array de pontuações e uma pontuação de estudante como parâmetros. A função deve retornar uma string com o formato:

    - `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`, se o estudante passou no curso.
    - `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`, se o aluno não passou no curso.

    Substitua `average-goes-here` pela média dos totais de pontuações e `grade-goes-here` pela nota do aluno. Use `getAverage` para obter a pontuação média e `getGrade` para obter a nota do estudante.

# --hints--

Você deve ter uma função chamada `getAverage`.

```js
assert.isFunction(getAverage);
```

A função `getAverage` deve retornar um número.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` deve retornar `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` deve retornar `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` deve retornar `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

A função `getAverage` deve retornar a pontuação média das pontuações de teste.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

Você deve ter uma função chamada `getGrade`.

```js
assert.isFunction(getGrade);
```

A função `getGrade` deve retornar `"A+"` se a pontuação for `100`.

```js
assert.strictEqual(getGrade(100), "A+");
```

A função `getGrade` deve retornar `"A"` se a pontuação estiver entre `90` e `99`.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

A função `getGrade` deve retornar `"B"` se a pontuação estiver entre `80` e `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

A função `getGrade` deve retornar `"C"` se a pontuação estiver entre `70` e `79`.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

A função `getGrade` deve retornar `"D"` se a pontuação estiver entre `60` e `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

A função `getGrade` deve retornar `"F"` se a pontuação estiver entre `0` e `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


Você deve ter uma função chamada `hasPassingGrade`.

```js
assert.isFunction(hasPassingGrade);
```

A função `hasPassingGrade` deve retornar um valor booleano.

```js
assert.isBoolean(hasPassingGrade(100));
```

A função `hasPassingGrade` deve retornar `true` se a nota for um `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

A função `hasPassingGrade` deve retornar `false` se a nota for um `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

A função `hasPassingGrade` deve retornar `true` para todas as notas de aprovação.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

Você deve ter uma função chamada `studentMsg`.

```js
assert.isFunction(studentMsg);
```

A chamada da função `studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` deve retornar a seguinte mensagem: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

A chamada de função de `studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` deve retornar a seguinte mensagem: `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

A função `studentMsg` deve retornar a mensagem correta com base na pontuação do estudante e na média da turma.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
