---
id: 5cc0bd7a49b71cb96132e54c
title: Створіть генератор діапазону чисел
challengeType: 26
dashedName: build-a-range-of-numbers-generator
---

# --description--

У цій лабораторній роботі ви створите `rangeOfNumbers` функцію, яка генерує масив чисел у заданому діапазоні.

Приклади:

- `rangeOfNumbers(3, 9)` повертає `[3, 4, 5, 6, 7, 8, 9]`
- `rangeOfNumbers(5, 5)` повертає `[5]`

Вимоги:

- Використовуйте рекурсію (функція має викликати сама себе)
- Без циклів або методів масиву (`for`, `while`, `Array.from`, `forEach`, `map`, `filter`, `reduce`)

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви маєте створити функцію з іменем `rangeOfNumbers`, яка приймає два параметри: `startNum` та `endNum`.
1. Функція має повертати масив цілих чисел, який починається з числа, представленого параметром `startNum`, і закінчується числом, представленим параметром `endNum` (включно).
1. `startNum` завжди буде меншим або рівним `endNum`.
1. Ваша функція має використовувати рекурсію, викликаючи сама себе. Вона не повинна використовувати жодного синтаксису циклів (`for`, `while`, `Array.from()`) або функцій вищого порядку, таких як `forEach`, `map`, `filter` чи `reduce`.
1. Функція має обробляти базовий випадок, коли `startNum` дорівнює `endNum`, повертаючи масив, що містить лише це одне число.
1. Для рекурсивного випадку функція має викликати сама себе з модифікованими параметрами, щоб побудувати масив, а потім додати поточне число до результату.
1. Функція не повинна покладатися на глобальні змінні для кешування або побудови масиву.

# --hints--

Ви маєте мати функцію з іменем `rangeOfNumbers`.

```js
assert.isFunction(rangeOfNumbers);
```

Функція `rangeOfNumbers` має приймати два параметри.

```js
assert.match(rangeOfNumbers.toString(), /rangeOfNumbers\s*\(\s*\w+\s*,\s*\w+\s*\)/);
```

Ваша функція має повертати масив.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

Ваш код не повинен використовувати жодного синтаксису циклів (`for`, `while`, `Array.from()`) або функцій вищого порядку, таких як `forEach`, `map`, `filter` чи `reduce`.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

`rangeOfNumbers` має використовувати рекурсію (викликати сама себе) для розв’язання цього завдання.

```js
assert.isAtLeast(
  __helpers.removeJSComments(rangeOfNumbers.toString()).match(/rangeOfNumbers\s*\(.+\)/g)?.length,
  2
);
```

`rangeOfNumbers(1, 5)` має повертати `[1, 2, 3, 4, 5]`.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)` має повертати `[6, 7, 8, 9]`.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)` має повертати `[4]`.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

`rangeOfNumbers(10, 15)` має повертати `[10, 11, 12, 13, 14, 15]`.

```js
assert.deepStrictEqual(rangeOfNumbers(10, 15), [10, 11, 12, 13, 14, 15]);
```

`rangeOfNumbers(2, 8)` має повертати `[2, 3, 4, 5, 6, 7, 8]`.

```js
assert.deepStrictEqual(rangeOfNumbers(2, 8), [2, 3, 4, 5, 6, 7, 8]);
```

Глобальні змінні не повинні використовуватися для кешування масиву.

```js
rangeOfNumbers(1, 3);
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
