---
id: 68e6698b29e562d21b1d260c
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви створите список основних веб-браузерів. Для вас уже підготовлено шаблонний код HTML.

Почніть із додавання заголовка на свою сторінку з текстом `List of Major Web Browsers` за допомогою елемента `h1` всередині елемента `body`.

# --hints--

У вас має бути елемент `h1` всередині елемента `body`.

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

Ваш елемент `h1` має містити текст `List of Major Web Browsers`.

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
