---
id: 68ecfa0f03d88990bb96a33c
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Як ви пам’ятаєте з попереднього уроку, списки описів використовують для подання термінів і визначень у впорядкованому та легкому для читання форматі.

Ось приклад:

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

Під елементом `h1` створіть елемент `dl`. Він міститиме список браузерів.

# --hints--

У вас має бути елемент `dl`.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

Ваш елемент `dl` має бути розташований під елементом `h1`.

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        --fcc-editable-region--
        
        --fcc-editable-region--
    </body> 
</html>
```
