---
id: 68ecfa1003d88990bb96a33d
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Описові списки складаються зі списку термінів і деталей. Елемент `dt` представляє описовий термін, тоді як елемент `dd` представляє описову деталь.

Перший браузер, який ви додасте до свого описового списку, буде для Google Chrome.

Всередині вашого елемента `dl` додайте елемент `dt`, що містить текст `Google Chrome`.

Під вашим елементом `dt` додайте елемент `dd` з таким текстом:

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

У вас має бути елемент `dt` всередині вашого елемента `dl`.

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

Ваш елемент `dt` має містити текст `Google Chrome`.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

У вас має бути елемент `dd` під вашим елементом `dt`.

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

Ваш елемент `dd` має містити текст `This is a free web browser developed by Google and first released in 2008.`

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```
