---
id: 68ecfa1003d88990bb96a33e
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Тепер настав час додати другий браузер до списку.

Під вашим елементом `Google Chrome` додайте ще один елемент `dt`, який містить текст `Firefox`.

Під вашим елементом `dt` додайте елемент `dd` з таким текстом:

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

У вашому коді має бути два елементи `dt`.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

Другий елемент `dt` має містити текст `Firefox`.

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

У вашому коді має бути два елементи `dd`.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

Другий елемент `dd` має містити текст `This is a free web browser developed by the Mozilla Corporation and first created in 2004.`

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
