---
id: 69038fe7da59ca23a5eba82d
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Останній браузер, який ви додасте до списку, буде для браузера Arc.

Додайте ще один `dt` елемент, що містить текст `Arc`.

Під вашим `dt` елементом додайте `dd` елемент із таким текстом:

```md
This is a free Chromium based web browser first released in 2023 by The Browser Company.
```

З цим останнім додаванням ваш список браузерів завершено!

# --hints--

У вашому коді має бути п’ять `dt` елементів.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 5);
```

Ваш п’ятий `dt` елемент має містити текст `Arc`.

```js
const dt = document.querySelectorAll('dt')[4];
assert.strictEqual(dt?.textContent.trim(), "Arc");
```

У вашому коді має бути п’ять `dd` елементів.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 5);
```

Ваш п’ятий `dd` елемент має містити текст `This is a free Chromium based web browser first released in 2023 by The Browser Company.`.

```js
const ddElement = document.querySelectorAll('dd')[4];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free Chromium based web browser first released in 2023 by The Browser Company.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            <dt>Brave</dt>
            <dd>This is a free web browser first released in 2016 that is based on the Chromium web browser.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>List of Browsers and Descriptions</title>
</head>
<body>
    <h1>List of Major Web Browsers</h1>
    <dl>
        <dt>Google Chrome</dt>
        <dd>This is a free web browser developed by Google and first released in 2008.</dd>

        <dt>Firefox</dt>
        <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

        <dt>Safari</dt>
        <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

        <dt>Brave</dt>
        <dd>This is a free web browser first released in 2016 that is based on the Chromium web browser.</dd>

        <dt>Arc</dt>
        <dd>This is a free Chromium based web browser first released in 2023 by The Browser Company.</dd>
    </dl>
</body>
</html>
```
