---
id: 5895f70bf9fc0f352b528e64
title: usa i poteri di un motore di template
challengeType: 2
forumTopicId: 301567
dashedName: use-a-template-engines-powers
---

# --description--

Una delle funzionalità più importanti nell’usare un motore di template è poter passare variabili dal server al file di template prima di renderizzarlo in HTML.

Nel tuo file Pug, puoi usare una variabile facendo riferimento al nome della variabile come `#{variable_name}` in linea con altro testo su un elemento oppure usando un segno di uguale sull’elemento senza spazi come `p=variable_name` che assegna il valore della variabile al testo dell’elemento p.

Pug si basa sull’uso degli spazi bianchi e delle tabulazioni per mostrare gli elementi annidati e ridurre la quantità di codice necessaria per creare un sito bello.

Prendi per esempio questo codice Pug: 
 
```pug
head
  script(type='text/javascript').
    if (foo) bar(1 + 5);
body
  if youAreUsingPug
      p You are amazing
    else
      p Get on it!
```

Il codice sopra genera il seguente HTML:

```html
<head>
  <script type="text/javascript">
    if (foo) bar(1 + 5);
  </script>
</head>
<body>
  <p>You are amazing</p>
</body>
```

Il tuo file `index.pug` incluso nel progetto usa le variabili `title` e `message`.

Passale dal server al file Pug aggiungendo un oggetto come secondo argomento alla chiamata `res.render` con le variabili e i loro valori. Dai a `title` il valore `Hello` e a `message` il valore `Please log in`.

Dovrebbe apparire così:

```javascript
res.render('index', { title: 'Hello', message: 'Please log in' });
```

Ora aggiorna la pagina e dovresti vedere quei valori renderizzati nella vista nel punto corretto come definito nel tuo file `index.pug`!

Invia la pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi consultare il <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#use-a-template-engines-power-2" target="_blank" rel="noopener noreferrer nofollow">progetto completato fino a questo punto</a>.

# --hints--

Pug dovrebbe renderizzare correttamente le variabili.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /pug-variable("|')>Please log in/gi,
    'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
  );
```
