---
id: 5895f700f9fc0f352b528e63
title: Configurare un motore di template
challengeType: 2
forumTopicId: 301564
dashedName: set-up-a-template-engine
---

# --description--

Per affrontare queste sfide dovrai scrivere il tuo codice usando uno dei seguenti metodi:

- Clona <a href="https://github.com/freeCodeCamp/boilerplate-advancednode/" target="_blank" rel="noopener noreferrer nofollow">questo repository GitHub</a> e completa queste sfide in locale.
- Usa un generatore di siti a tua scelta per completare il progetto. Assicurati di includere tutti i file dal nostro repository GitHub.

Un motore di template ti permette di usare file template statici (come quelli scritti in *Pug*) nella tua app. A runtime, il motore di template sostituisce le variabili in un file template con valori reali forniti dal tuo server. Poi trasforma il template in un file HTML statico che viene inviato al client. Questo approccio rende più semplice progettare una pagina HTML e consente di mostrare variabili nella pagina senza dover fare una chiamata API dal client.

`pug@~3.0.0` è già stato installato ed è elencato come dipendenza nel tuo file `package.json`.

Express deve sapere quale motore di template stai usando. Usa il metodo `set` per assegnare `pug` come valore della proprietà `view engine`:

```javascript
app.set('view engine', 'pug');
```

Dopodiché, aggiungi un altro metodo `set` che imposta la proprietà `views` della tua `app` per puntare alla directory `./views/pug`. Questo dice a Express di renderizzare tutte le viste relative a quella directory.

Infine, usa `res.render()` nella rotta per la tua pagina principale, passando `index` come primo argomento. Questo renderizzerà il template `pug`.

Se tutto è andato come previsto, la pagina principale della tua app non sarà più vuota. Invece, mostrerà un messaggio che indica che hai renderizzato con successo il template Pug!

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-a-template-engine-1" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

# --hints--

Pug dovrebbe essere una dipendenza.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'pug',
    'Your project should list "pug" as a dependency'
  );
```

Il motore di vista dovrebbe essere Pug.

```js
  const url = new URL("/_api/app", code);
  const res = await fetch(url);
  const app = await res.json();
  assert.equal(app?.settings?.['view engine'], "pug");
```

Dovresti impostare la proprietà `views` dell’applicazione su `./views/pug`.

```js
  const url = new URL("/_api/app", code);
  const res = await fetch(url);
  const app = await res.json();
  assert.equal(app?.settings?.views, "./views/pug");
```

Usa il metodo corretto di ExpressJS per renderizzare la pagina indice dalla risposta.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
      assert.match(
        data,
        /FCC Advanced Node and Express/gi,
        'You successfully rendered the Pug template!'
      );
```

Pug dovrebbe funzionare.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
      assert.match(
        data,
        /pug-success-message/gi,
        'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
      );
```
