---
id: 5895f70cf9fc0f352b528e65
title: Configura Passport
challengeType: 2
forumTopicId: 301565
dashedName: set-up-passport
---

# --description--

È ora di configurare *Passport* così potrai finalmente permettere a un utente di registrarsi o accedere a un account. Oltre a Passport, userai Express-session per gestire le sessioni. Express-session offre molte funzionalità avanzate, ma per ora userai solo le basi. Usare questo middleware salva l'id della sessione come cookie nel client e ci permette di consultare i dati della sessione usando quell'id sul server. In questo modo, tieni le informazioni personali dell’account fuori dal cookie usato dal client per comunicare al tuo server che i client sono autenticati e conservi la *chiave* per accedere ai dati archiviati sul server.

`passport@~0.4.1` e `express-session@~1.17.1` sono già installati e sono entrambi elencati come dipendenze nel tuo file `package.json`.

Dovrai configurare le impostazioni della sessione e inizializzare Passport. Per prima cosa, crea le variabili `session` e `passport` per richiedere rispettivamente `express-session` e `passport`.

Poi, configura la tua app Express per usare la sessione definendo le seguenti opzioni:

```javascript
app.use(session({
  secret: process.env.SESSION_SECRET,
  resave: true,
  saveUninitialized: true,
  cookie: { secure: false }
}));
```

Assicurati di aggiungere `SESSION_SECRET` al tuo file `.env` e di dargli un valore casuale. Questo serve a calcolare l’hash usato per criptare il tuo cookie!

Dopo aver fatto tutto questo, indica alla tua app express di **usare** `passport.initialize()` e `passport.session()`.

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-passport-3" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

# --hints--

Passport e Express-session dovrebbero essere dipendenze.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport',
    'Your project should list "passport" as a dependency'
  );
  assert.property(
    packJson.dependencies,
    'express-session',
    'Your project should list "express-session" as a dependency'
  );
```

Le dipendenze dovrebbero essere correttamente richieste.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport("|')/,
    'You should have required passport'
  );
  assert.match(
    data,
    /require.*("|')express-session("|')/,
    'You should have required express-session'
  );
```

L’app Express dovrebbe usare le nuove dipendenze.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /passport\.initialize/, 'Your express app should use "passport.initialize()"');
  assert.match(data, /passport\.session/, 'Your express app should use "passport.session()"');
```

Sessione e segreto della sessione dovrebbero essere configurati correttamente.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /secret *:\s*process\.env(\.SESSION_SECRET|\[(?<q>"|')SESSION_SECRET\k<q>\])/,
    'Your express app should have express-session set up with your secret as process.env.SESSION_SECRET'
  );
```
