---
id: 5895f70cf9fc0f352b528e66
title: Serializzazione di un oggetto utente
challengeType: 2
forumTopicId: 301563
dashedName: serialization-of-a-user-object
---

# --description--

La serializzazione e la deserializzazione sono concetti importanti per quanto riguarda l'autenticazione. Serializzare un oggetto significa convertire il suo contenuto in una piccola *chiave* che può poi essere deserializzata nell'oggetto originale. Questo ci permette di sapere chi ha comunicato con il server senza dover inviare i dati di autenticazione, come nome utente e password, a ogni richiesta di una nuova pagina.

Per configurare tutto correttamente, devi avere una funzione di serializzazione e una funzione di deserializzazione. In Passport, queste possono essere create con:

```javascript
passport.serializeUser(cb);
passport.deserializeUser(cb);
```

La funzione di callback passata a `serializeUser` viene chiamata con due argomenti: l'intero oggetto utente e una callback usata da passport.

La callback si aspetta due argomenti: un errore, se presente, e una chiave unica per identificare l'utente che deve essere restituita nella callback. Userai il `_id` dell'utente nell'oggetto. Questo è garantito come unico, poiché viene generato da MongoDB.

Allo stesso modo, `deserializeUser` viene chiamata con due argomenti: la chiave unica e una funzione di callback.

Questa callback si aspetta due argomenti: un errore, se presente, e l'intero oggetto utente. Per ottenere l'intero oggetto utente, esegui una query di ricerca per un `_id` Mongo, come mostrato di seguito:


```javascript
passport.serializeUser((user, done) => {
  done(null, user._id);
});

passport.deserializeUser((id, done) => {
  myDataBase.findOne({ _id: new ObjectID(id) }, (err, doc) => {
    done(null, null);
  });
});
```

Aggiungi le due funzioni sopra al tuo server. La classe `ObjectID` proviene dal pacchetto `mongodb`. `mongodb@~3.6.0` è già stata aggiunta come dipendenza. Dichiara questa classe con:

```javascript
const { ObjectID } = require('mongodb');
```

Il `deserializeUser` lancerà un errore finché non configurerai la connessione al database. Quindi, per ora, commenta la chiamata a `myDatabase.findOne` e chiama semplicemente `done(null, null)` nella funzione di callback `deserializeUser`.

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#serialization-of-a-user-object-4" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

# --hints--

Dovresti serializzare correttamente l'oggetto utente.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.serializeUser/gi,
    'You should have created your passport.serializeUser function'
  );
  assert.match(
    data,
    /null,\s*user._id/gi,
    'There should be a callback in your serializeUser with (null, user._id)'
  );
```

Dovresti deserializzare correttamente l'oggetto utente.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.deserializeUser/gi,
    'You should have created your passport.deserializeUser function'
  );
  assert.match(
    data,
    /null,\s*null/gi,
    'There should be a callback in your deserializeUser with (null, null) for now'
  );
```

MongoDB dovrebbe essere una dipendenza.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'mongodb',
    'Your project should list "mongodb" as a dependency'
  );
```

MongoDB dovrebbe essere correttamente importato includendo ObjectId.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')mongodb\1/gi,
    'You should have required mongodb'
  );
  assert.match(
    data,
    /new ObjectID.*id/gi,
    'Even though the block is commented out, you should use new ObjectID(id) for when we add the database'
  );
```
