---
id: 5895f70df9fc0f352b528e6a
title: Crea un nuovo middleware
challengeType: 2
forumTopicId: 301551
dashedName: create-new-middleware
---

# --description--

Così com’è, qualsiasi utente può semplicemente andare a `/profile` sia che sia autenticato o meno, digitando l’URL. Vuoi evitare questo controllando prima se l’utente è autenticato prima di renderizzare la pagina del profilo. Questo è l’esempio perfetto di quando creare un middleware.

La sfida qui è creare la funzione middleware `ensureAuthenticated(req, res, next)`, che controllerà se un utente è autenticato chiamando il metodo `isAuthenticated` di Passport sul `request` che verifica se `req.user` è definito. Se lo è, allora `next()` dovrebbe essere chiamato. Altrimenti, puoi semplicemente rispondere alla richiesta con un reindirizzamento alla tua homepage per il login.

Un’implementazione di questo middleware è:

```javascript
function ensureAuthenticated(req, res, next) {
  if (req.isAuthenticated()) {
    return next();
  }
  res.redirect('/');
};
```

Crea la funzione middleware sopra, poi passa `ensureAuthenticated` come middleware alle richieste per la pagina del profilo prima dell’argomento della richiesta GET:

```javascript
app
 .route('/profile')
 .get(ensureAuthenticated, (req,res) => {
    res.render('profile');
 });
```

Invia la tua pagina quando pensi di averla fatta correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#create-new-middleware-8" target="_blank" rel="noopener noreferrer nofollow">controllare il progetto completato fino a questo punto</a>.

# --hints--

Il middleware `ensureAuthenticated` dovrebbe essere implementato e allegato alla rotta `/profile`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /ensureAuthenticated[^]*req.isAuthenticated/,
    'Your ensureAuthenticated middleware should be defined and utilize the req.isAuthenticated function'
  );
  assert.match(
    data,
    /profile[^]*get[^]*ensureAuthenticated/,
    'Your ensureAuthenticated middleware should be attached to the /profile route'
  );
```

Una richiesta GET non autenticata a `/profile` dovrebbe correttamente reindirizzare a `/`.

```js
  const url = new URL("/profile", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/,
    'An attempt to go to the profile at this point should redirect to the homepage since we are not logged in'
  );
```
