---
id: 5895f70cf9fc0f352b528e67
title: implementare la serializzazione di un utente passport
challengeType: 2
forumTopicId: 301556
dashedName: implement-the-serialization-of-a-passport-user
---

# --description--

Non stai caricando un vero oggetto utente perché il database non è configurato. Connettiti al database una sola volta, all'avvio del server, e mantieni una connessione persistente per l'intero ciclo di vita dell'app. Per farlo, aggiungi la stringa di connessione del tuo database (per esempio: `mongodb+srv://<username>:<password>@cluster0-jvwxi.mongodb.net/?retryWrites=true&w=majority`) alla variabile d'ambiente `MONGO_URI`. Questa viene usata nel file `connection.js`.

*Se hai problemi a configurare un database gratuito su MongoDB Atlas, dai un'occhiata a questo <a href="https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/" target="_blank" rel="noopener noreferrer nofollow">tutorial</a>.*

Ora vuoi connetterti al database e poi iniziare ad ascoltare le richieste. Lo scopo è non permettere richieste prima che il database sia connesso o in caso di errore del database. Per farlo, racchiudi la tua serializzazione e le rotte dell'app nel seguente codice:

```javascript
myDB(async client => {
  const myDataBase = await client.db('database').collection('users');

  // Be sure to change the title
  app.route('/').get((req, res) => {
    // Change the response to render the Pug template
    res.render('index', {
      title: 'Connected to Database',
      message: 'Please login'
    });
  });

  // Serialization and deserialization here...

  // Be sure to add this...
}).catch(e => {
  app.route('/').get((req, res) => {
    res.render('index', { title: e, message: 'Unable to connect to database' });
  });
});
// app.listen out here...
```

Assicurati di togliere il commento dal codice `myDataBase` in `deserializeUser` e modifica il tuo `done(null, null)` per includere il `doc`.

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implement-the-serialization-of-a-passport-user-5" target="_blank" rel="noopener noreferrer nofollow">controllare il progetto completato fino a questo punto</a>.

# --hints--

La connessione al database dovrebbe essere presente.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Connected to Database/gi,
    'You successfully connected to the database!'
  );
```

La deserializzazione dovrebbe ora usare correttamente il DB e `done(null, null)` dovrebbe essere chiamato con il `doc`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /null,\s*doc/gi,
    'The callback in deserializeUser of (null, null) should be altered to (null, doc)'
  );
```
