---
id: 5895f70ef9fc0f352b528e6b
title: Come mettere insieme un profilo
challengeType: 2
forumTopicId: 301554
dashedName: how-to-put-a-profile-together
---

# --description--

Ora che puoi assicurarti che l’utente che sta consultando `/profile` sia autenticato, puoi usare le informazioni contenute in `req.user` nella tua pagina.

Passa un oggetto contenente la proprietà `username` e il valore `req.user.username` come secondo argomento per il metodo `render` della vista profilo.

Poi, vai alla tua vista `profile.pug` e aggiungi la seguente riga sotto l’elemento `h1` esistente, allo stesso livello di indentazione:

```pug
h2.center#welcome Welcome, #{username}!
```

Questo crea un elemento `h2` con la classe `center` e l’id `welcome` contenente il testo `Welcome, ` seguito dal nome utente.

Inoltre, in `profile.pug`, aggiungi un collegamento che punta alla rotta `/logout`, che ospiterà la logica per disautenticare un utente:

```pug
a(href='/logout') Logout
```

Invia la tua pagina quando pensi di averla fatta correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-put-a-profile-together-9" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

# --hints--

Dovresti aggiungere correttamente una variabile di render Pug a `/profile`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /username:( |)req.user.username/,
    'You should be passing the variable username with req.user.username into the render function of the profile page'
  );
```
