---
id: 58966a17f9fc0f352b528e6d
title: Registrazione di nuovi utenti
challengeType: 2
forumTopicId: 301561
dashedName: registration-of-new-users
---

# --description--

Ora devi permettere a un nuovo utente del tuo sito di registrare un account. Nel `res.render` per la home page aggiungi una nuova variabile all’oggetto passato — `showRegistration: true`. Quando aggiorni la pagina, dovresti vedere il modulo di registrazione che è già stato creato nel tuo file `index.pug`. Questo modulo è configurato per fare una **POST** su `/register`, quindi crea quella rotta e fai in modo che aggiunga l’oggetto utente al database seguendo la logica qui sotto.

La logica della rotta di registrazione dovrebbe essere la seguente:

1. Registrare il nuovo utente  
2. Autenticare il nuovo utente  
3. Reindirizzare a `/profile`

La logica del passaggio 1 dovrebbe essere la seguente:

1. Fare una query al database con `findOne`  
2. Se c’è un errore, chiamare `next` con l’errore  
3. Se viene restituito un utente, reindirizzare alla home  
4. Se non viene trovato un utente e non si verificano errori, allora `insertOne` nel database con username e password. Finché non si verificano errori, chiamare `next` per passare al passaggio 2, autenticando il nuovo utente, per cui hai già scritto la logica nella tua rotta `POST /login`.

```js
app.route('/register')
  .post((req, res, next) => {
    myDataBase.findOne({ username: req.body.username }, (err, user) => {
      if (err) {
        next(err);
      } else if (user) {
        res.redirect('/');
      } else {
        myDataBase.insertOne({
          username: req.body.username,
          password: req.body.password
        },
          (err, doc) => {
            if (err) {
              res.redirect('/');
            } else {
              // The inserted document is held within
              // the ops property of the doc
              next(null, doc.ops[0]);
            }
          }
        )
      }
    })
  },
    passport.authenticate('local', { failureRedirect: '/' }),
    (req, res, next) => {
      res.redirect('/profile');
    }
  );
```

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#registration-of-new-users-11" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

**NOTA:** Da questo punto in poi, possono sorgere problemi legati all’uso del browser *picture-in-picture*. Se stai usando un IDE online che offre un’anteprima dell’app all’interno dell’editor, è consigliato aprire questa anteprima in una nuova scheda.

# --hints--

Dovresti avere una rotta `/register` e mostrare un modulo di registrazione nella home page.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /showRegistration:( |)true/gi,
    'You should be passing the variable showRegistration as true to your render function for the homepage'
  );
  assert.match(
    data,
    /register[^]*post[^]*findOne[^]*username:( |)req.body.username/gi,
    'You should have a route that accepts a POST request on /register that queries the db with findOne and the query being username: req.body.username'
  );
```

La registrazione dovrebbe funzionare.

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttp = new XMLHttpRequest();
  xhttp.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      test(this);
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttp.open('POST', url + '/register', true);
  xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
  xhttp.send(`username=${user}&password=${user}`);
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Register should work, and redirect successfully to the profile.'
    );
  }
```

Il login dovrebbe funzionare.

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttpReg = new XMLHttpRequest();
  xhttpReg.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      login();
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttpReg.open('POST', url + '/register', true);
  xhttpReg.setRequestHeader(
    'Content-type',
    'application/x-www-form-urlencoded'
  );
  xhttpReg.send(`username=${user}&password=${user}`);
  function login() {
    const xhttp = new XMLHttpRequest();
    xhttp.onreadystatechange = function () {
      if (this.readyState == 4 && this.status == 200) {
        test(this);
      } else {
        throw new Error(`${this.status} ${this.statusText}`);
      }
    };
    xhttp.open('POST', url + '/login', true);
    xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
    xhttp.send(`username=${user}&password=${user}`);
  }
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Login should work if previous test was done successfully and redirect successfully to the profile.'
    );
    assert.match(
      data,
      new RegExp(user, 'g'),
      'The profile should properly display the welcome to the user logged in'
    );
  }
```

Il logout dovrebbe funzionare.

```js
  $.ajax({
    url: code + '/logout',
    type: 'GET',
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(data, /Home/gi, 'Logout should redirect to home');
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```

Il profilo non dovrebbe più funzionare dopo il logout.

```js
  $.ajax({
    url: code + '/profile',
    type: 'GET',
    crossDomain: true,
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(
        data,
        /Home/gi,
        'Profile should redirect to home when we are logged out now again'
      );
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```
