---
id: 589690e6f9fc0f352b528e6e
title: Pulisci il tuo progetto con i moduli
challengeType: 2
forumTopicId: 301549
dashedName: clean-up-your-project-with-modules
---

# --description--

Al momento, tutto quello che hai è nel tuo file `server.js`. Questo può portare a un codice difficile da gestire e poco espandibile. Crea 2 nuovi file: `routes.js` e `auth.js`

Entrambi devono iniziare con il seguente codice:

```js
module.exports = function (app, myDataBase) {

}
```

Ora, in cima al file del server, importa questi file così: `const routes = require('./routes.js');` Subito dopo aver stabilito una connessione riuscita con il database, istanzia ciascuno di essi così: `routes(app, myDataBase)`

Infine, prendi tutte le rotte nel tuo server e incollale nei nuovi file, rimuovendole dal file del server. Prendi anche la funzione `ensureAuthenticated`, dato che è stata creata appositamente per il routing. Ora dovrai aggiungere correttamente le dipendenze usate, come `const passport = require('passport');`, in cima, sopra la riga di export nel file `routes.js`.

Continua ad aggiungerle finché non ci sono più errori e il file del server non contiene più rotte (**eccetto la rotta nel blocco catch**)!

Fai lo stesso nel file `auth.js` con tutto ciò che riguarda l'autenticazione, come la serializzazione e la configurazione della strategia locale, ed eliminali dal file del server. Assicurati di aggiungere le dipendenze e di chiamare `auth(app, myDataBase)` nel server nello stesso punto.

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#clean-up-your-project-with-modules-2" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata a un esempio del progetto completato</a>.

# --hints--

I moduli dovrebbero essere presenti.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require\s*\(('|")\.\/routes(\.js)?\1\)/gi,
    'You should have required your new files'
  );
  assert.match(
    data,
    /client\s*\.db[^]*routes/gi,
    'Your new modules should be called after your connection to the database'
  );
```
