---
id: 58965611f9fc0f352b528e6c
title: Disconnettere un utente
challengeType: 2
forumTopicId: 301560
dashedName: logging-a-user-out
---

# --description--

Creare la logica per il logout è semplice. La rotta dovrebbe semplicemente disautenticare l'utente e reindirizzare alla pagina principale invece di renderizzare una vista.

In passport, disautenticare un utente è facile come chiamare `req.logout()` prima di reindirizzare. Aggiungi questa rotta `/logout` per farlo:

```js
app.route('/logout')
  .get((req, res) => {
    req.logout();
    res.redirect('/');
});
```

Avrai notato che non stai gestendo le pagine mancanti (404). Il modo comune per gestirlo in Node è con il seguente middleware. Vai avanti e aggiungilo dopo tutte le altre rotte:

```js
app.use((req, res, next) => {
  res.status(404)
    .type('text')
    .send('Not Found');
});
```

Invia la tua pagina quando pensi di averla fatta correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#logging-a-user-out-10" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

# --hints--

`req.logout()` dovrebbe essere chiamato nella tua rotta `/logout`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /req.logout/gi,
    'You should be calling req.logout() in your /logout route'
  );
```

`/logout` dovrebbe reindirizzare alla pagina principale.

```js
  const url = new URL("/logout", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/gi,
    'When a user logs out they should be redirected to the homepage'
  );
```
