---
id: 589a8eb3f9fc0f352b528e72
title: "Implementazione dell'autenticazione sociale III"
challengeType: 2
forumTopicId: 301558
dashedName: implementation-of-social-authentication-iii
---

# --description--

La parte finale della strategia consiste nel gestire il profilo restituito da GitHub. Dobbiamo caricare l'oggetto utente dal database se esiste, oppure crearne uno se non esiste, e popolare i campi dal profilo, quindi restituire l'oggetto utente. GitHub ci fornisce un *id* unico all'interno di ogni profilo che possiamo usare per cercare e serializzare l'utente (già implementato). Di seguito trovi un esempio di implementazione che puoi usare nel tuo progetto — va inserito nella funzione che è il secondo argomento della nuova strategia, subito sotto dove `console.log(profile);` è attualmente:

```js
myDataBase.findOneAndUpdate(
  { id: profile.id },
  {
    $setOnInsert: {
      id: profile.id,
      username: profile.username,
      name: profile.displayName || 'John Doe',
      photo: profile.photos[0].value || '',
      email: Array.isArray(profile.emails)
        ? profile.emails[0].value
        : 'No public email',
      created_on: new Date(),
      provider: profile.provider || ''
    },
    $set: {
      last_login: new Date()
    },
    $inc: {
      login_count: 1
    }
  },
  { upsert: true, new: true },
  (err, doc) => {
    return cb(null, doc.value);
  }
);
```

`findOneAndUpdate` ti permette di cercare un oggetto e aggiornarlo. Se l'oggetto non esiste, verrà inserito e reso disponibile alla funzione di callback. In questo esempio, impostiamo sempre `last_login`, incrementiamo `login_count` di `1` e popolare la maggior parte dei campi solo quando viene inserito un nuovo oggetto (nuovo utente). Nota l'uso di valori predefiniti. A volte un profilo restituito non avrà tutte le informazioni compilate o l'utente le terrà private. In questo caso, gestisci la situazione per evitare un errore.

Ora dovresti poter effettuare il login nella tua app. Provalo!

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-iii-5" target="_blank" rel="noopener noreferrer nofollow">dare un'occhiata al progetto completato fino a questo punto</a>.

# --hints--

La configurazione della strategia GitHub dovrebbe essere completa.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /GitHubStrategy[^]*myDataBase/gi,
    'Strategy should use now use the database to search for the user'
  );
  assert.match(
    data,
    /GitHubStrategy[^]*cb/gi,
    'Strategy should return the callback function "cb"'
  );
```
