---
id: 589a69f5f9fc0f352b528e70
title: "Implementazione dell'autenticazione social"
challengeType: 2
forumTopicId: 301559
dashedName: implementation-of-social-authentication
---

# --description--

Il percorso base che questo tipo di autenticazione seguirà nella tua app è:

1.  L'utente clicca un pulsante o un collegamento che lo porta alla tua rotta per autenticarsi usando una strategia specifica (ad esempio GitHub).
2.  La tua rotta chiama `passport.authenticate('github')` che li reindirizza a GitHub.
3.  La pagina su cui l'utente arriva, su GitHub, gli permette di effettuare il login se non lo ha già fatto. Poi gli chiede di approvare l'accesso al suo profilo da parte della tua app.
4.  L'utente viene quindi riportato alla tua app a un URL di callback specifico con il suo profilo se ha dato l'approvazione.
5.  Ora è autenticato, e la tua app dovrebbe verificare se si tratta di un profilo già esistente o salvarlo nel database se non lo è.

Le strategie con OAuth richiedono di avere almeno un *Client ID* e un *Client Secret*, che sono un modo per il servizio di verificare chi sta facendo la richiesta di autenticazione e se è valida. Questi si ottengono dal sito con cui stai implementando l'autenticazione, come GitHub, e sono unici per la tua app—**NON DEVONO ESSERE CONDIVISI** e non dovrebbero mai essere caricati in un repository pubblico o scritti direttamente nel codice. Una pratica comune è metterli nel file `.env` e riferirsi a essi così: `process.env.GITHUB_CLIENT_ID`. Per questa sfida userai la strategia GitHub.

Segui <a href="https://www.freecodecamp.org/news/how-to-set-up-a-github-oauth-application/" target="_blank" rel="noopener noreferrer nofollow">queste istruzioni</a> per ottenere il tuo *Client ID e Secret* da GitHub. Imposta l'URL della homepage sulla tua homepage (**non l'URL del codice del progetto**), e imposta l'URL di callback allo stesso URL della homepage con `/auth/github/callback` aggiunto alla fine. Salva il client ID e il client secret nel file `.env` del tuo progetto come `GITHUB_CLIENT_ID` e `GITHUB_CLIENT_SECRET`.

Nel file `routes.js`, aggiungi `showSocialAuth: true` alla rotta della homepage, dopo `showRegistration: true`. Ora crea 2 rotte che accettano richieste GET: `/auth/github` e `/auth/github/callback`. La prima dovrebbe solo chiamare passport per autenticare `'github'`. La seconda dovrebbe chiamare passport per autenticare `'github'` con un reindirizzamento in caso di fallimento a `/`, e poi se ha successo reindirizzare a `/profile` (simile al tuo ultimo progetto).

Un esempio di come dovrebbe essere `/auth/github/callback` è simile a come hai gestito un login normale:

```js
app.route('/login')
  .post(passport.authenticate('local', { failureRedirect: '/' }), (req,res) => {
    res.redirect('/profile');
  });
```

Invia la tua pagina quando pensi di averla fatta correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-3" target="_blank" rel="noopener noreferrer nofollow">controllare il progetto fino a questo punto</a>.

# --hints--

La rotta `/auth/github` dovrebbe essere corretta.

```js
  try {
    const res = await fetch(code + '/_api/routes.js');
    if (res.ok) {
      const data = await res.text();
      assert.match(
          data.replace(/\s/g, ''),
          /passport.authenticate.*?github/g,
          'Route auth/github should only call passport.authenticate with github'
        );
    } else {
      throw new Error(res.statusText);
    }
    const res2 = await fetch(code + '/_api/app-stack');
    if (res2.ok) {
      const data2 = JSON.parse(await res2.json());
      const dataLayer = data2.find(layer => layer?.route?.path === '/auth/github');
      assert.deepInclude(dataLayer?.route, { methods: {get: true}, path: "/auth/github"});
      assert.deepInclude(dataLayer?.route?.stack?.[0], {method: "get", name: "authenticate"});
    } else {
      throw new Error(res2.statusText);
    }
  } catch (err) {
    throw new Error(err);
  }
```

La rotta `/auth/github/callback` dovrebbe essere corretta.

```js
  try {
    const res = await fetch(code + '/_api/routes.js');
    if (res.ok) {
      const data = await res.text();
      assert.match(
        data.replace(/\s/g, ''),
        /failureRedirect:("|')\/\1/g,
        'Route auth/github/callback should accept a get request and call passport.authenticate for github with a failure redirect to home'
      );
    } else {
      throw new Error(res.statusText);
    }
    const res2 = await fetch(code + '/_api/app-stack');
    if (res2.ok) {
      const data2 = JSON.parse(await res2.json());
      const dataLayer = data2.find(layer => layer?.route?.path === '/auth/github/callback');
      assert.deepInclude(dataLayer?.route, { methods: {get: true}, path: "/auth/github/callback"});
      assert.deepInclude(dataLayer?.route?.stack?.[0], {method: "get", name: "authenticate"});
    } else {
      throw new Error(res2.statusText);
    }
  } catch (err) {
    throw new Error(err);
  }
```
