---
id: 589a69f5f9fc0f352b528e71
title: "Implementazione dell'autenticazione social II"
challengeType: 2
forumTopicId: 301557
dashedName: implementation-of-social-authentication-ii
---

# --description--

L'ultima parte per configurare l'autenticazione con GitHub è creare la strategia stessa. `passport-github@~1.1.0` è già stata aggiunta come dipendenza, quindi richiamala nel tuo file `auth.js` come `GitHubStrategy` in questo modo: `const GitHubStrategy = require('passport-github').Strategy;`. Non dimenticare di richiamare e configurare `dotenv` per usare le tue variabili d'ambiente.

Per configurare la strategia GitHub, devi dire a Passport di usare un'istanza di `GitHubStrategy`, che accetta 2 argomenti: un oggetto (contenente `clientID`, `clientSecret` e `callbackURL`) e una funzione da chiamare quando un utente viene autenticato con successo, che determinerà se l'utente è nuovo e quali campi salvare inizialmente nell'oggetto database dell'utente. Questo è comune a molte strategie, ma alcune potrebbero richiedere più informazioni come indicato nel README GitHub di quella specifica strategia. Per esempio, Google richiede anche uno *scope* che determina che tipo di informazioni la tua richiesta chiede di ricevere e chiede all'utente di approvare tale accesso.

La strategia che stai implementando autentica gli utenti usando un account GitHub e token OAuth 2.0. L'id client e il segreto ottenuti creando un'applicazione sono forniti come opzioni quando crei la strategia. La strategia richiede anche una callback `verify`, che riceve il token di accesso e un token di aggiornamento opzionale, oltre a `profile` che contiene il profilo GitHub dell'utente autenticato. La callback `verify` deve chiamare `cb` fornendo un utente per completare l'autenticazione.

Ecco come dovrebbe apparire la tua nuova strategia a questo punto:

```js
passport.use(new GitHubStrategy({
  clientID: process.env.GITHUB_CLIENT_ID,
  clientSecret: process.env.GITHUB_CLIENT_SECRET,
  callbackURL: /*INSERT CALLBACK URL ENTERED INTO GITHUB HERE*/
},
  function(accessToken, refreshToken, profile, cb) {
    console.log(profile);
    //Database logic here with callback containing your user object
  }
));
```

La tua autenticazione non avrà ancora successo, e in realtà lancerà un errore senza la logica del database e la callback, ma dovrebbe stampare il tuo profilo GitHub nella console se provi!

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-ii-4" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

# --hints--

`passport-github` deve essere aggiunta come dipendenza.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport-github',
    'Your project should list "passport-github" as a dependency'
  );
```

`passport-github` deve essere richiamata.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport-github("|')/gi,
    'You should have required passport-github'
  );
```

La strategia GitHub dovrebbe essere configurata correttamente finora.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport\.use.*new GitHubStrategy/gis,
    'Passport should use a new GitHubStrategy'
  );
  assert.match(
    data,
    /callbackURL:\s*("|').*("|')/gi,
    'You should have a callbackURL'
  );
  assert.match(
    data,
    /process\.env(\.GITHUB_CLIENT_SECRET|\[(?<q>"|')GITHUB_CLIENT_SECRET\k<q>\])/g,
    'You should use process.env.GITHUB_CLIENT_SECRET'
  );
  assert.match(
    data,
    /process\.env(\.GITHUB_CLIENT_ID|\[(?<q>"|')GITHUB_CLIENT_ID\k<q>\])/g,
    'You should use process.env.GITHUB_CLIENT_ID'
  );
```
