---
id: 589fc832f9fc0f352b528e79
title: Inviare e mostrare messaggi di chat
challengeType: 2
forumTopicId: 301562
dashedName: send-and-display-chat-messages
---

# --description--

È ora di permettere ai client di inviare un messaggio di chat al server per trasmetterlo a tutti i client! Nel tuo file `client.js` dovresti già vedere un blocco di codice che gestisce l'invio del modulo del messaggio.

```js
$('form').submit(function() {
  /*logic*/
});
```

All'interno del codice di invio del modulo, dovresti emettere un evento dopo aver definito `messageToSend` ma prima di cancellare il riquadro di testo `#m`. L'evento dovrebbe chiamarsi `'chat message'` e i dati dovrebbero essere semplicemente `messageToSend`.

```js
socket.emit('chat message', messageToSend);
```

Ora, sul tuo server, dovresti ascoltare il socket per l'evento `'chat message'` con i dati chiamati `message`. Una volta ricevuto l'evento, dovrebbe emettere l'evento `'chat message'` a tutti i socket usando `io.emit`, inviando un oggetto dati contenente `username` e `message`.

In `client.js`, dovresti ora ascoltare l'evento `'chat message'` e, quando ricevuto, aggiungere un elemento di lista a `#messages` con il nome utente, due punti e il messaggio!

A questo punto, la chat dovrebbe essere completamente funzionante e inviare messaggi a tutti i client!

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#send-and-display-chat-messages-11" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

# --hints--

Il server dovrebbe ascoltare `'chat message'` e poi emetterlo correttamente.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*io.emit.*('|")chat message('|").*username.*message/s,
    'Your server should listen to the socket for "chat message" then emit to all users "chat message" with name and message in the data object'
  );
```

Il client dovrebbe gestire e mostrare correttamente i nuovi dati dall'evento `'chat message'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*messages.*li/s,
    'You should append a list item to #messages on your client within the "chat message" event listener to display the new message'
  );
```
