---
id: 589fc831f9fc0f352b528e75
title: Comunicare emettendo
challengeType: 2
forumTopicId: 301550
dashedName: communicate-by-emitting
---

# --description--

<dfn>Emit</dfn> è il modo più comune di comunicare che userai. Quando emetti qualcosa dal server verso 'io', invii il nome di un evento e i dati a tutti i socket connessi. Un buon esempio di questo concetto è emettere il conteggio attuale degli utenti connessi ogni volta che un nuovo utente si collega!

Inizia aggiungendo una variabile per tenere traccia degli utenti, appena prima del punto in cui stai attualmente ascoltando le connessioni.

```js
let currentUsers = 0;
```

Ora, quando qualcuno si connette, dovresti incrementare il conteggio prima di emettere il conteggio. Quindi, vorrai aggiungere l'incremento all'interno del listener di connessione.

```js
++currentUsers;
```

Infine, dopo aver incrementato il conteggio, dovresti emettere l'evento (ancora all'interno del listener di connessione). L'evento dovrebbe chiamarsi 'user count', e i dati dovrebbero essere semplicemente il `currentUsers`.

```js
io.emit('user count', currentUsers);
```

Ora, puoi implementare un modo per il tuo client di ascoltare questo evento! Simile ad ascoltare una connessione sul server, userai la parola chiave `on`.

```js
socket.on('user count', function(data) {
  console.log(data);
});
```

Ora, prova a caricare la tua app, autenticarti, e dovresti vedere nella console del client '1' che rappresenta il conteggio attuale degli utenti! Prova a caricare più client e autenticarti per vedere il numero aumentare.

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#communicate-by-emitting-7" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

# --hints--

`currentUsers` dovrebbe essere definito.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /currentUsers/s,
    'You should have variable currentUsers defined'
  );
```

Il server dovrebbe emettere il conteggio attuale degli utenti a ogni nuova connessione.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.emit.*('|")user count('|").*currentUsers/s,
    'You should emit "user count" with data currentUsers'
  );
```

Il tuo client dovrebbe ascoltare l'evento `'user count'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user count('|")/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
