---
id: 589fc831f9fc0f352b528e77
title: Autenticazione con Socket.IO
challengeType: 2
forumTopicId: 301548
dashedName: authentication-with-socket-io
---

# --description--

Attualmente, non puoi determinare chi è connesso al tuo web socket. Mentre `req.user` contiene l'oggetto utente, questo accade solo quando il tuo utente interagisce con il web server, e con i web socket non hai `req` (richiesta) e quindi nessun dato utente. Un modo per risolvere il problema di sapere chi è connesso al tuo web socket è analizzare e decodificare il cookie che contiene la sessione passport e poi deserializzarlo per ottenere l'oggetto utente. Per fortuna, esiste un pacchetto su NPM proprio per questo che trasforma un'attività una volta complessa in qualcosa di semplice!

`passport.socketio@~3.7.0`, `connect-mongo@~3.2.0` e `cookie-parser@~1.4.5` sono già stati aggiunti come dipendenze. Richiamali come `passportSocketIo`, `MongoStore` e `cookieParser` rispettivamente. Inoltre, dobbiamo inizializzare un nuovo archivio in memoria, da `express-session` che abbiamo richiesto in precedenza. Dovrebbe apparire così:

```js
const MongoStore = require('connect-mongo')(session);
const URI = process.env.MONGO_URI;
const store = new MongoStore({ url: URI });
```

Ora dobbiamo solo dire a Socket.IO di usarlo e impostare le opzioni. Assicurati che questo venga aggiunto prima del codice socket esistente e non nel listener di connessione esistente. Per il tuo server, dovrebbe apparire così:

```js
io.use(
  passportSocketIo.authorize({
    cookieParser: cookieParser,
    key: 'express.sid',
    secret: process.env.SESSION_SECRET,
    store: store,
    success: onAuthorizeSuccess,
    fail: onAuthorizeFail
  })
);
```

Nota che configurare l'autenticazione Passport per Socket.IO è molto simile al modo in cui abbiamo configurato il middleware `session` per l'API. Questo perché sono pensati per usare lo stesso metodo di autenticazione — ottenere l'id della sessione da un cookie e convalidarlo.

In precedenza, quando abbiamo configurato il middleware `session`, non abbiamo impostato esplicitamente il nome del cookie per la sessione (`key`). Questo perché il pacchetto `session` usava il valore predefinito. Ora che abbiamo aggiunto un altro pacchetto che ha bisogno di consultare lo stesso valore dai cookie, dobbiamo impostare esplicitamente il valore `key` in entrambi gli oggetti di configurazione.

Assicurati di aggiungere `key` con il nome del cookie al middleware `session` che corrisponde alla chiave di Socket.IO. Inoltre, aggiungi il riferimento `store` alle opzioni, vicino a dove impostiamo `saveUninitialized: true`. Questo è necessario per dire a Socket.IO a quale sessione fare riferimento.

<hr>

Ora, definisci le funzioni di callback `success` e `fail`:

```js
function onAuthorizeSuccess(data, accept) {
  console.log('successful connection to socket.io');

  accept(null, true);
}

function onAuthorizeFail(data, message, error, accept) {
  if (error) throw new Error(message);
  console.log('failed connection to socket.io:', message);
  accept(null, false);
}
```

L'oggetto utente è ora accessibile sul tuo oggetto socket come `socket.request.user`. Per esempio, ora puoi aggiungere quanto segue:

```js
console.log('user ' + socket.request.user.username + ' connected');
```

Questo stamperà nella console del server chi si è connesso!

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi  <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#authentication-with-socketio-9" target="_blank" rel="noopener noreferrer nofollow">controllare il progetto fino a questo punto</a>.

# --hints--

`passport.socketio` dovrebbe essere una dipendenza.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport.socketio',
    'Your project should list "passport.socketio" as a dependency'
  );
```

`cookie-parser` dovrebbe essere una dipendenza.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'cookie-parser',
    'Your project should list "cookie-parser" as a dependency'
  );
```

passportSocketIo dovrebbe essere richiamato correttamente.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require\((['"])passport\.socketio\1\)/gi,
    'You should correctly require and instantiate "passport.socketio"'
  );
```

passportSocketIo dovrebbe essere configurato correttamente.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io\.use\(\s*\w+\.authorize\(/,
    'You should register "passport.socketio" as socket.io middleware and provide it correct options'
  );
```
