---
id: 589fc832f9fc0f352b528e78
title: Annunciare nuovi utenti
challengeType: 2
forumTopicId: 301546
dashedName: announce-new-users
---

# --description--

Molte chat room sono in grado di annunciare quando un utente si connette o si disconnette e poi mostrare questa informazione a tutti gli utenti connessi nella chat. Dal momento che stai già emettendo un evento alla connessione e alla disconnessione, dovrai solo modificare questo evento per supportare questa funzionalità. Il modo più logico per farlo è inviare 3 dati con l’evento: il nome utente dell’utente che si è connesso/disconnesso, il conteggio attuale degli utenti e se quel nome utente si è connesso o disconnesso.

Cambia il nome dell’evento in `'user'` e passa un oggetto contenente i campi `username`, `currentUsers` e `connected` (che sarà `true` in caso di connessione, o `false` per la disconnessione dell’utente inviato). Assicurati di modificare entrambi gli eventi `'user count'` e imposta quello di disconnessione per inviare `false` per il campo `connected` invece di `true` come nell’evento emesso alla connessione.

```js
io.emit('user', {
  username: socket.request.user.username,
  currentUsers,
  connected: true
});
```

Ora il tuo client avrà tutte le informazioni necessarie per mostrare correttamente il conteggio attuale degli utenti e annunciare quando un utente si connette o si disconnette! Per gestire questo evento sul client dovremmo ascoltare `'user'`, quindi aggiornare il conteggio attuale degli utenti usando jQuery per cambiare il testo di `#num-users` in `'{NUMBER} users online'`, oltre ad aggiungere un `<li>` alla lista non ordinata con id `messages` con `'{NAME} has {joined/left} the chat.'`.

Un’implementazione di questo potrebbe essere la seguente:

```js
socket.on('user', data => {
  $('#num-users').text(data.currentUsers + ' users online');
  let message =
    data.username +
    (data.connected ? ' has joined the chat.' : ' has left the chat.');
  $('#messages').append($('<li>').html('<b>' + message + '</b>'));
});
```

Invia la tua pagina quando pensi di averla fatta correttamente. Se incontri errori, puoi consultare <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135/3#announce-new-users-10" target="_blank" rel="noopener noreferrer nofollow">il progetto completato fino a questo punto</a>.

# --hints--

L’evento `'user'` dovrebbe essere emesso con `username`, `currentUsers` e `connected`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  // Regex is lenient to match both `username` and `name` as the key on purpose.
  assert.match(
    data,
    /io.emit.*('|")user\1.*name.*currentUsers.*connected/s,
    'You should have an event emitted named user sending name, currentUsers, and connected'
  );
```

Il client dovrebbe gestire e mostrare correttamente i nuovi dati dall’evento `'user'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*num-users/s,
    'You should change the text of "#num-users" within on your client within the "user" event listener to show the current users connected'
  );
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*messages.*li/s,
    'You should append a list item to "#messages" on your client within the "user" event listener to announce a user came or went'
  );
```
