---
id: 589fc831f9fc0f352b528e76
title: Gestire una disconnessione
challengeType: 2
forumTopicId: 301552
dashedName: handle-a-disconnect
---

# --description--

Potresti aver notato che finora hai solo aumentato il conteggio degli utenti. Gestire la disconnessione di un utente è altrettanto semplice quanto gestire la connessione iniziale, solo che devi ascoltarla su ogni singolo socket invece che sull'intero server.

Per farlo, aggiungi un altro listener all'interno del listener `'connect'` esistente che ascolti l'evento `'disconnect'` sul socket senza passare dati. Puoi testare questa funzionalità semplicemente registrando nella console che un utente si è disconnesso.

```js
socket.on('disconnect', () => {
  /*anything you want to do on disconnect*/
});
```

Per assicurarti che i client abbiano continuamente il conteggio aggiornato degli utenti attuali, dovresti diminuire `currentUsers` di 1 quando avviene la disconnessione e poi emettere l'evento `'user count'` con il conteggio aggiornato.

**Nota:** Proprio come `'disconnect'`, tutti gli altri eventi che un socket può emettere al server dovrebbero essere gestiti all'interno del listener di connessione dove abbiamo definito 'socket'.

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#handle-a-disconnect-8" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

# --hints--

Il server dovrebbe gestire l'evento disconnect da un socket.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /socket.on.*('|")disconnect('|")/s, '');
```

Il tuo client dovrebbe ascoltare l'evento `'user count'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user count('|")/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
