---
id: 589fc830f9fc0f352b528e74
title: "Configura l'ambiente"
challengeType: 2
forumTopicId: 301566
dashedName: set-up-the-environment
---

# --description--

Le sfide seguenti utilizzeranno il file `chat.pug`. Quindi, nel tuo file `routes.js`, aggiungi una rotta GET che punti a `/chat` che usa `ensureAuthenticated` e renderizza `chat.pug`, passando `{ user: req.user }` come argomento alla risposta. Ora, modifica la tua rotta `/auth/github/callback` esistente per impostare `req.session.user_id = req.user.id` e reindirizzare a `/chat`.

`socket.io@~2.3.0` è già stata aggiunta come dipendenza, quindi richiedila/istanziala nel tuo server come segue con `http` (integrata in Node.js):

```javascript
const http = require('http').createServer(app);
const io = require('socket.io')(http);
```

Ora che il server *http* è montato sull'*express app*, devi ascoltare dal server *http*. Cambia la riga con `app.listen` in `http.listen`.

La prima cosa da gestire è ascoltare una nuova connessione dal client. La parola chiave <dfn>on</dfn> fa proprio questo: ascolta un evento specifico. Richiede 2 argomenti: una stringa contenente il nome dell'evento emesso e una funzione con cui vengono passati i dati. Nel caso del nostro listener di connessione, usa `socket` per definire i dati nel secondo argomento. Un socket è un singolo client connesso.

Per ascoltare le connessioni al tuo server, aggiungi quanto segue all'interno della connessione al database:

```javascript
io.on('connection', socket => {
  console.log('A user has connected');
});
```

Ora, per far connettere il client, devi solo aggiungere quanto segue al tuo `client.js` che viene caricato dalla pagina dopo che ti sei autenticato:

```js
/*global io*/
let socket = io();
```

Il commento sopprime l'errore che normalmente vedresti poiché 'io' non è definito nel file. Hai già aggiunto un CDN affidabile alla libreria Socket.IO nella pagina in `chat.pug`.

Ora prova a caricare la tua app e autenticarti e dovresti vedere nella console del server `A user has connected`.

**Nota:** `io()` funziona solo quando ci si connette a un socket ospitato sulla stessa url/server. Per connettersi a un socket esterno ospitato altrove, useresti `io.connect('URL');`.

Invia la tua pagina quando pensi di aver fatto tutto correttamente. Se incontri errori, puoi <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-the-environment-6" target="_blank" rel="noopener noreferrer nofollow">dare un’occhiata al progetto completato fino a questo punto</a>.

# --hints--

`socket.io` dovrebbe essere una dipendenza.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'socket.io',
    'Your project should list "socket.io" as a dependency'
  );
```

Dovresti richiedere e istanziare correttamente `http` come `http`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /http.*=.*require.*('|")http\1/s,
    'Your project should list "http" as a dependency'
  );
```

Dovresti richiedere e istanziare correttamente `socket.io` come `io`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.*=.*require.*('|")socket.io\1.*http/s,
    'You should correctly require and instantiate socket.io as io.'
  );
```

Socket.IO dovrebbe ascoltare le connessioni.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.on.*('|")connection\1.*socket/s,
    'io should listen for "connection" and socket should be the 2nd arguments variable'
  );
```

Il tuo client dovrebbe connettersi al tuo server.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.*=.*io/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
