---
id: 614389f601bb4f611db98563
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Unter Ihrem `.author`-Element erstellen Sie ein neues `div`-Element mit der Klasse `social-icons`.

Fügen Sie fünf `a`-Elemente innerhalb dieses neuen `div` ein und geben Sie ihnen die folgenden `href`-Attribute.

- Das erste `a`-Element sollte ein `href` mit dem Wert `https://www.facebook.com/freecodecamp` haben.
- Das zweite `a`-Element sollte ein `href` mit dem Wert `https://twitter.com/freecodecamp` haben.
- Das dritte `a`-Element sollte ein `href` mit dem Wert `https://instagram.com/freecodecamp` haben.
- Das vierte `a`-Element sollte ein `href` mit dem Wert `https://www.linkedin.com/school/free-code-camp` haben.
- Das fünfte `a`-Element sollte ein `href` mit dem Wert `https://www.youtube.com/freecodecamp` haben.

# --hints--

Sie sollten ein neues `div`-Element erstellen.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

Ihr neues `div`-Element sollte nach Ihrem `.author`-Element kommen.

```js
assert(document.querySelector('.author')?.nextElementSibling?.localName === 'div');
```

Ihr neues `div`-Element sollte die Klasse `social-icons` haben.

```js
assert(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

Ihr `.social-icons`-Element sollte fünf `a`-Elemente enthalten.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.length === 5);
```

Ihr erstes `a`-Element sollte ein `href` mit dem Wert `https://www.facebook.com/freecodecamp` haben.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href')?.includes('https://www.facebook.com/freecodecamp'));
```

Ihr zweites `a`-Element sollte ein `href` mit dem Wert `https://twitter.com/freecodecamp` haben.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

Ihr drittes `a`-Element sollte ein `href` mit dem Wert `https://instagram.com/freecodecamp` haben.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

Ihr viertes `a`-Element sollte ein `href` mit dem Wert `https://www.linkedin.com/school/free-code-camp` haben.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

Ihr fünftes `a`-Element sollte ein `href` mit dem Wert `https://www.youtube.com/freecodecamp` haben.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
