---
id: 61437d575fb98f57fa8f7f36
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Beginnen Sie mit Ihrem Standard-HTML-Boilerplate-Code. Fügen Sie eine `DOCTYPE`-Deklaration, ein `html`-Element hinzu, das angibt, dass diese Seite auf Englisch ist, ein `head`-Element und ein `body`-Element.

Fügen Sie innerhalb des `<meta>`-Elements ein `charset`-Tag mit dem passenden `<meta>` und ein `head`-Tag für die mobile Responsivität hinzu.

# --hints--

Ihr Code sollte die `DOCTYPE`-Referenz enthalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Sie sollten nach der `DOCTYPE`-Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Sie sollten den Dokumenttyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Sie sollten die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Sie sollten ein öffnendes `<html>`-Tag haben, das `lang` von `en` besitzt.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ihr `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Ihre `DOCTYPE`-Deklaration sollte am Anfang Ihres HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Sie sollten ein öffnendes `<head>`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Sie sollten ein schließendes `</head>`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Sie sollten ein öffnendes `<body>`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Sie sollten ein schließendes `</body>`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Die `head`- und `body`-Elemente sollten Geschwister sein.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Das `head`-Element sollte innerhalb des `html`-Elements sein.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Das `body`-Element sollte innerhalb des `html`-Elements sein.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Sie sollten zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Ein `meta`-Element sollte ein `name` mit dem Wert `viewport` und `content` mit dem Wert `width=device-width, initial-scale=1.0` haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Das andere `meta`-Element sollte das `charset`-Attribut mit dem Wert `UTF-8` besitzen.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
