---
id: 614e0e588f0e8a772a8a81a6
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Ihr Bild nimmt derzeit viel Platz ein. Um besser sehen zu können, woran Sie arbeiten, fügen Sie dem `width`-Element ein `img`-Attribut mit dem Wert `400` hinzu.

Sie werden dieses später entfernen, wenn Sie am CSS gearbeitet haben.

# --hints--

Ihr `img`-Element sollte ein `width`-Attribut mit dem Wert `400` haben.

```js
assert(document.querySelector('img')?.getAttribute('width') === '400');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
