---
id: 689f7996426c1534fa5bea66
title: 建置一個高爾夫球分數轉換器
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

在高爾夫遊戲中，每個洞都有一個 `par`，代表高爾夫球手預期需要的平均 `strokes` 數，以將球擊入洞中完成該洞。根據你的 `strokes` 超過或低於 `par` 的程度，會有不同的暱稱。

在本實驗中，你將撰寫一個函式（程式），將 `par` 和 `strokes` 轉換為它們的暱稱。 

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `golfScore` 的函式（程式）。
1. `golfScore` 應該接受兩個數字引數，分別是球場的標準桿數和擊球次數。
1. `golfScore` 應該傳回一個字串。
1. `golfScore` 應該在 `strokes` 為 `1` 時傳回 `"Hole-in-one!"`。
1. 如果 `strokes` 小於或相等於 `par` 減 `2`，`golfScore` 應該傳回 `"Eagle"`。
1. `golfScore` 應該在 `strokes` 與 `par` 減 1 相等時傳回 `"Birdie"`。
1. `golfScore` 應該在 `strokes` 與 `par` 相等時傳回 `"Par"`。
1. `golfScore` 應該在 `strokes` 等於 `par` 加 1 時傳回 `"Bogey"`。
1. 如果 `strokes` 相等於 `par` 加上 `2`，`golfScore` 應該傳回 `"Double Bogey"`。
1. 如果 `strokes` 大於或相等於 `par` 加上 `3`，`golfScore` 應該傳回 `"Go Home!"`。

# --hints--

你應該創建一個名為 `golfScore` 的函式（程式）。

```js
assert.isFunction(golfScore);
```

`golfScore` 應該接受兩個參數。

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` 應該傳回一個字串

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)` 應該傳回字串 `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` 應該傳回字串 `Hole-in-one!`

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` 應該傳回字串 `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` 應該傳回字串 `Hole-in-one!`

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` 應該傳回字串 `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` 應該傳回字串 `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` 應該傳回字串 `Birdie`

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` 應該傳回字串 `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` 應該傳回字串 `Birdie`

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` 應該傳回字串 `Par`

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` 應該傳回字串 `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` 應該傳回字串 `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` 應該傳回字串 `Bogey`

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` 應該傳回字串 `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` 應該傳回字串 `Bogey`

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` 應該傳回字串 `Double Bogey`

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` 應該傳回字串 `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` 應該傳回字串 `Double Bogey`

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` 應該傳回字串 `Go Home!`

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` 應該傳回字串 `Go Home!`

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` 應該傳回字串 `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
