---
id: ac6993d51946422351508a41
title: 문자열 자르기 알고리즘 구현하기
challengeType: 26
dashedName: truncate-a-string
---

# --description--

이 실습에서는 문자열을 특정 길이로 자르는 연습을 합니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 두 개의 인수를 받는 함수 `truncateString`가 있어야 합니다. 첫 번째는 문자열, 두 번째는 숫자입니다.
1. 문자열 길이가 주어진 숫자보다 길면, 문자열 길이를 주어진 숫자와 같게 줄이도록 자르고, 잘린 문자열 끝에 `...`를 추가해야 합니다.
1. 문자열 길이가 주어진 숫자와 같거나 짧으면, 문자열을 변경하지 않고 반환해야 합니다.

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)`는 문자열 `A-tisket...`를 반환해야 합니다.

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)`는 문자열 `Peter Piper...`를 반환해야 합니다.

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)`는 문자열 `A-tisket a-tasket A green and yellow basket`를 반환해야 합니다.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)`는 문자열 `A-tisket a-tasket A green and yellow basket`를 반환해야 합니다.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)`는 문자열 `A...`를 반환해야 합니다.

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)`는 문자열 `Ab...`를 반환해야 합니다.

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
