---
id: 5d822fd413a79914d39e98ca
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Füge öffnende und schließende `html`-Tags unter `DOCTYPE` hinzu, so dass du einen Platz hast, um mit der Eingabe des Codes zu beginnen. Achte darauf, die Sprache auf Englisch zu setzen.

# --hints--

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dein `html`-Element sollte ein öffnendes Tag enthalten sowie ein `lang`-Attribut mit dem Wert `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein abschließendes Tag enthalten.

```js
assert(code.match(/<\/html\s*>/));
```

Deine `html`-Tags sollten in der richtigen Reihenfolge stehen.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Du solltest nur ein `html`-Element haben.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
