---
id: 6434552bcc0a951a0a99df3b
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Функція `range` очікує чисел, але значення `start` та `end` будуть рядками (зокрема, вони будуть окремими символами, наприклад `A`).

Перетворіть значення `start` та `end` у виклику `range()` на числа використовуючи метод `.charCodeAt()` на них, передавши число `0` як аргумент цьому методу.

# --hints--

Ви повинні використати метод `.charCodeAt()`.

```js
assert.match(code, /\.charCodeAt\(/);
```

Викличте метод `.charCodeAt()` на `start`.

```js
assert.match(code, /start\.charCodeAt\(/);
```

Передайте `0` методу `.charCodeAt()` з `start`.

```js
assert.match(code, /start\.charCodeAt\(\s*0\s*\)/);
```

Викличте метод `.charCodeAt()` на `end`.

```js
assert.match(code, /end\.charCodeAt\(/);
```

Передайте `0` методу `.charCodeAt()` з `end`.

```js
assert.match(code, /end\.charCodeAt\(\s*0\s*\)/);
```

Використайте методи `.charCodeAt()` одразу у виклику `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start, end);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
