---
id: 64345c560591891f64976f7a
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Тепер, коли допоміжні функції завершені, повернувшись у обробник подій `onload` оголосіть змінну `letters`. Присвойте до неї результат виклику `charRange()` з літерами `"A"` та `"J"` як аргументами.

# --hints--

Оголосіть змінну `letters`.

```js
assert.match(code, /(?:let|const|var)\s+letters/);
```

Використайте `const`, щоб оголосити змінну `letters`.

```js
assert.match(code, /const\s+letters/);
```

Присвоїте виклик `charRange()` змінній `letters`.

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(/);
```

Передайте `"A"` як перший аргумент до виклику `charRange()`.

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1/);
```

Передайте `"J"` як другий аргумент до виклику `charRange()`.

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1\s*,\s*('|"|`)J\2\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
