---
id: 68e4417b7317835208145323
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Далі створіть нову змінну з назвою `dogsOnlySentence` і призначте їй `exampleSentence.replace("cats", "dogs")`.

Потім під цією змінною виведіть у консоль `"Replacing all occurrences of cats with dogs:"`. Під цим `console.log()` додайте ще один `console.log()` для змінної `dogsOnlySentence`.

Тепер ви повинні побачити в консолі текст `"I love dogs and cats are so much fun!"`. Але це не той результат, який потрібен. У наступному кроці ви навчитеся, як це виправити, щоб текст читався як `"I love dogs and dogs are so much fun!"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні мати змінну `dogsOnlySentence`.

```js
assert.exists(dogsOnlySentence);
```

Ваша змінна `dogsOnlySentence` має бути рядком.

```js
assert.isString(dogsOnlySentence);
```

Ви повинні призначити `exampleSentence.replace("cats", "dogs")` вашій змінній `dogsOnlySentence`.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

Ви повинні вивести `"Replacing all occurrences of cats with dogs:"` у консоль.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

Ви повинні вивести змінну `dogsOnlySentence` у консоль.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
