---
id: 68e42de6396f5c49d90d1951
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

У попередньому уроці ви дізналися про метод `replace()`, який використовується для пошуку підрядка в рядку та заміни його на інше значення.

Ось приклад:

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

Пам’ятайте, що рядки є незмінними, тобто оригінальний `text` у цьому прикладі не змінюється. Натомість створюється новий рядок.

Створіть нову змінну з назвою `replacedString` і призначте їй `originalString.replace("cats", "dogs")`.

# --hints--

У вас має бути змінна з назвою `replacedString`.

```js
assert.exists(replacedString);
```

Ваша змінна `replacedString` має бути рядком.

```js
assert.isString(replacedString);
```

Ви маєте призначити `originalString.replace("cats", "dogs")` змінній `replacedString`.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
