---
id: 68e42b62d7093345c6975f22
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Під вашим змінним `originalString` додайте `console.log()` зі рядком `"Original string:"`. 

Потім під цим `console.log()` додайте ще один `console.log()` зі змінною `originalString`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні вивести рядок `"Original string:"` у консоль.

```js
assert.equal(getLogs()[0], "Original string:");
```

Ви повинні вивести змінну `originalString` у консоль. 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
