---
id: 68e43e884ef09d4d2ab05631
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Тепер настав час побачити результат використання методу `replace()`.

Почніть із додавання `console.log()` зі значенням `"After using the replace() method:"`.

Потім нижче цього `console.log()` додайте ще один `console.log()` з змінною `replacedString`.

Погляньте на консоль, і ви повинні побачити, що нове речення каже `"I love dogs."` замість оригінального `"I love cats."`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні вивести `"After using the replace() method:"` у консоль.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

Ви повинні вивести змінну `replacedString` у консоль.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
