---
id: 68e44abea678a65b877864dc
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Було б добре побачити це повторене слово у реченні. Створіть нову змінну з назвою `newSentence` і призначте їй результат конкатенації слова `"I "` зі змінною `repeatedLove`, а потім додайте рядок `" learning."`. Потім виведіть у консоль змінну `newSentence`. Ви можете обрати використання оператора `+` або шаблонних літералів для конкатенації рядків.

Тепер ви повинні побачити текст `"I love love love  learning."`.

*ПРИМІТКА*: Після останнього `"love"` є зайвий пробіл, і це правильно. У останньому кроці ви видалите цей зайвий пробіл.

# --hints--

У вас має бути змінна з назвою `newSentence`.

```js
assert.exists(newSentence);
```

Ваша змінна `newSentence` має бути рядком.

```js
assert.isString(newSentence);
```

Ви маєте конкатенувати слово `"I "` зі змінною `repeatedLove`, а потім додати рядок `" learning."`. Цей результат має бути призначений змінній `newSentence`. Будьте уважні з пробілами.

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

Ви НЕ повинні призначати літеральний рядок `"I love love love  learning."` змінній `newSentence`. Будь ласка, перечитайте інструкції для правильного підходу.

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
