---
id: 68e43f94365a964f26c88f78
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Тепер час попрацювати з новим методом рядка.

Почніть зі створення нової змінної з назвою `exampleSentence` і призначте їй рядок `"I love cats and cats are so much fun!"`.

Потім нижче цієї змінної виведіть у консоль `"Original sentence:"`. Потім нижче цього `console.log()`, додайте ще один `console.log()` зі змінною `exampleSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні мати змінну `exampleSentence`.

```js
assert.exists(exampleSentence);
```

Ваша змінна `exampleSentence` має бути рядком.

```js
assert.isString(exampleSentence);
```

Ви повинні призначити рядок `"I love cats and cats are so much fun!"` змінній `exampleSentence`.

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

Ви повинні вивести `"Original sentence:"` у консоль.

```js
assert.equal(getLogs()[4], "Original sentence:");
```

Ви повинні вивести змінну `exampleSentence` у консоль.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
