---
id: 68e4465fba0bd256868b7deb
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Тепер настав час попрацювати з новим методом рядка.

Почніть зі створення змінної з назвою `learningSentence` і призначте їй рядок `"I love learning!"`.

Потім під цією змінною виведіть у консоль `"Original learning sentence:"`. Під `console.log()` додайте ще один `console.log()` для змінної `learningSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні мати змінну `learningSentence`.

```js
assert.exists(learningSentence);
```

Ваша змінна `learningSentence` має бути рядком.

```js
assert.isString(learningSentence);
```

Ви повинні призначити рядок `"I love learning!"` змінній `learningSentence`.

```js
assert.equal(learningSentence, "I love learning!");
```

Ви повинні вивести `"Original learning sentence:"` у консоль.

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

Ви повинні вивести змінну `learningSentence` у консоль.

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
