---
id: 68e44519c769bf545bb3f22a
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Щоб отримати бажаний вихідний результат `"I love dogs and dogs are so much fun!"`, вам потрібно буде використовувати метод `replaceAll()` замість методу `replace()`. Цей метод повертає новий рядок для всіх співпадінь підрядка.

Оновіть `exampleSentence.replace("cats", "dogs")`, щоб використовувати метод `replaceAll()` замість цього. Тепер ви повинні побачити правильний текст у консолі.

# --hints--

Вам слід оновити лінію коду `exampleSentence.replace("cats", "dogs")`, щоб використовувати метод `replaceAll()`.

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
