---
id: 68e41084bb275539f72de9fe
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з методами `replace()`, `replaceAll()` та `repeat()`, створюючи застосунок для трансформації рядків.

Для початку створіть змінну з назвою `originalString` і призначте їй значення `"I love cats."`.

# --hints--

У вас має бути змінна з назвою `originalString`.

```js
assert.exists(originalString);
```

Змінна `originalString` має бути рядком.

```js
assert.isString(originalString);
```

Ви маєте призначити рядок `"I love cats."` змінній `originalString`.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
