---
id: 68e44c305f1f935d949d76bb
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Як згадувалося в попередньому кроці, після останнього `"love"` у реченні `"I love love love  learning."` є зайвий пробіл. Щоб видалити цей зайвий пробіл, ви можете використати метод `trimEnd()`.

Наприкінці методу `"love ".repeat(3)` додайте виклик методу `trimEnd()`. Методи можна ланцюжити таким чином:

```js
.firstMethod().secondMethod()
```

Тепер, коли ви перевірите консоль, має з’явитися, що зайвий пробіл було видалено.

І з цією останньою зміною ваше практичне заняття завершено!

# --hints--

Ви повинні ланцюжити метод `trimEnd()` до методу `"love ".repeat(3)`.

```js
assert.equal(repeatedLove, "love ".repeat(3).trimEnd());
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--
const repeatedLove = "love ".repeat(3)
--fcc-editable-region--
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```

# --solutions--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3).trimEnd();
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```
