---
id: 68e448ec4ac4205968c0a805
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Було б добре, якби слово `"love"` повторювалося тричі в реченні `"I love learning!"`.

Як ви пам’ятаєте з попереднього уроку, ви можете повторити рядок певну кількість разів за допомогою методу `repeat()`. Ось приклад:

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

Оскільки рядки незмінні, цей метод не змінить оригінальний рядок. Він поверне новий рядок із повтореним вмістом.

Створіть змінну з назвою `repeatedLove` і призначте їй `"love ".repeat(3)`. Потім виведіть у консоль змінну `repeatedLove`, щоб побачити результат.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

У вас має бути змінна з назвою `repeatedLove`.

```js
assert.exists(repeatedLove);
```

Ваша змінна `repeatedLove` має бути рядком.

```js
assert.isString(repeatedLove);
```

Ви маєте призначити `"love ".repeat(3)` змінній `repeatedLove`.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

Ви маєте вивести змінну `repeatedLove` у консоль.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
