---
id: 6140c7e645d8e905819f1dd4
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin with the standard boilerplate. Add your `DOCTYPE` declaration, your `html` element with the language set to English, your `head` and `body` elements.

Adicione seu elemento `meta` com o `charset` correto, o elemento `title` e um elemento `link` que vincule o arquivo `./styles.css`.

Defina o título `title` para `Ferris Wheel`.

# --hints--

O código deve conter a declaração de `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE/gi);
```

Você deve incluir um espaço após a declaração de `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Você deve definir o tipo de documento para `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Você deve fechar a declaração `DOCTYPE` com um `>` após o tipo.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

O elemento `html` deve ter uma tag de abertura com um atributo `lang` de `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert.match(code, /<\/html\s*>/);
```

A declaração `DOCTYPE` deve ser no início do HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Você deve acrescentar uma tag de abertura `head`.

```js
assert.match(code, /<head\s*>/i);
```

Você deve ter uma tag de fechamento para `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

Você deve ter uma tag de abertura para `body`.

```js
assert.match(code, /<body\s*>/i);
```

Você deve ter uma tag de fechamento para `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

O elemento `body` deve ser o primeiro elemento depois do elemento `head`.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Os elementos `head` e `body` devem ser irmãos.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

O elemento `head` deve ser o primeiro elemento dentro do elemento `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

O elemento `head` deve estar dentro do elemento `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

O elemento `body` deve ser filho do elemento `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Your code should have a `meta` element.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Your `meta` element should have a `charset` attribute with the value `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

O elemento `meta` deve estar dentro do elemento `head`.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your project should have a title of `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Remember, the casing and spelling matter for the title.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Your code should have a `link` element.

```js
assert.match(code, /<link/);
```

Você deve ter um elemento `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Your `link` element should be within your `head` element.

```js
assert.exists(document.querySelector('head > link'));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
