---
id: 66f686b8ebdb982fa8e14330
title: 构建一个日期转换程序
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

在本实验中，你将编写一个日期转换 `Program`，将给定日期转换为不同的形式。例如，当前日期 `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` 将以以下 2 种形式格式化：

- `(MM/DD/YYYY)：9/27/2024`。
- （年 月 日）：2024 年 9 月 27 日。

**目标：**实现以下用户需求并通过所有测试以完成实验。 

**用户故事：**

1. 你应该创建一个名为 `currentDate` 的变量，并使用 `Date` 对象赋值为当前日期和时间。
2. 你应该创建一个名为 `currentDateFormat` 的变量，并赋值为字串 `Current Date and Time: [current date]`。将 `[current date]` 替换为 `currentDate` 变量的结果。
3. 你应该将 `currentDateFormat` 的值记录到控制台。
4. 你应该创建一个名为 `formatDateMMDDYYYY` 的函数，该函数以日期对象作为参数。你可以随意命名这个参数。
5. 你的 `formatDateMMDDYYYY` 函数应将当前日期转换为 `month/day/year` 形式并返回字串 `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`。
6. 你应该创建一个名为 `formatDateLong` 的函数，该函数以日期对象作为参数。
7. 你的 `formatDateLong` 函数应将当前日期转换为 `Month Day, Year` 形式并返回字串 `Formatted Date (Month Day, Year): [formatted date]`。

**注意：** 为了测试通过，确保在格式化日期时使用 `en-US` 作为局部的。

# --hints--

你应该有一个名为 `currentDate` 的变量，用于保存使用 `Date` 对象的当前日期和时间。

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

你应该有一个名为 `currentDateFormat` 的变量，用于保存格式为 `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>` 的当前日期。

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

你应该将 `currentDateFormat` 的值记录到控制台。

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

你应该有一个函数 `formatDateMMDDYYYY` 

```js
assert.isFunction(formatDateMMDDYYYY);
```

函数 `formatDateMMDDYYYY` 应该接受一个单一参数。

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

当日期对象包含 `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)` 时，函数 `formatDateMMDDYYYY` 应该返回 `Formatted Date (MM/DD/YYYY): 9/27/2024`。

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

你应该有一个函数 `formatDateLong` 

```js
assert.isFunction(formatDateLong);
```

函数 `formatDateLong` 应该只接受一个参数。

```js
assert.lengthOf(formatDateLong, 1);
```

当日期对象保存 `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)` 时，函数 `formatDateLong` 应该返回 `Formatted Date (Month Day, Year): September 27, 2024`。

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```

