---
id: 699f3cb8d057f29954205d7d
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Como recuerdas de lecciones anteriores, puedes agregar anotaciones de tipo a los parámetros de una función así:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

En este ejemplo, TypeScript mostrará un error si intentas llamar a la función `add` con argumentos que no sean números.

Actualiza tu función `square` para incluir una anotación de tipo para el parámetro `num`. El tipo debe ser `number`.

# --hints--

Debes agregar una anotación de tipo al parámetro `num` en la función `square`. El tipo debe ser un `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
