---
id: 699f434e294a9e11cf4f65f0
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Ahora es momento de agregar otra función matemática a tu proyecto toolkit.

Comienza creando una función llamada `getAverage` que reciba un parámetro. Tu función debe devolver el promedio de un arreglo de números.

En los siguientes pasos, agregarás las anotaciones de tipo para esta función.

# --hints--

Debes tener una función llamada `getAverage`.

```js
assert.isFunction(getAverage);
```

Tu función debe recibir un parámetro.

```js
assert.strictEqual(getAverage.length, 1);
```

Tu función debe devolver un número.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Tu función debe devolver el resultado correcto para cualquier arreglo de números.

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
