---
id: 699f38937b053e4679878ff1
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Cuando trabajas con JavaScript puro, nada impide que llames a la función `square` con una cadena o un arreglo, lo que causaría resultados inesperados.

Para ilustrar esto, cambia la llamada a tu función para usar la cadena `"something"` en lugar del número `5`.

En la siguiente lección, aprenderás a usar TypeScript para añadir seguridad de tipos a tu código y evitar este tipo de errores.

# --hints--

Tu llamada a la función `square` debe tener un argumento de `"something"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
