---
id: 699f41b8da20f816b5c7dece
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Si abres la consola, deberías ver el siguiente mensaje de error:

```md
Expected 1 arguments, but got 2.
```

Para resolver el error, actualiza tu llamada a la función `square` para que tome solo un argumento de `5`.

# --hints--

Tu llamada a la función `square` debe tener solo un argumento de `5`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
