---
id: 699f52220de88c8091032076
title: Paso 18
challengeType: 1
dashedName: step-18
---

# --description--

Si abres la consola, el resultado es `NaN` porque nunca se proporcionó el segundo argumento para `exponent`. Hay un par de actualizaciones necesarias para esta función `raiseTo` que solucionarán esto.

El primer cambio es devolver `base ** 2`.

En el siguiente paso, verás cómo manejar el parámetro opcional `exponent`.

# --hints--

Tu función `raiseTo` debe devolver un número.

```js
assert.isNumber(raiseTo(4));
```

Tu función `raiseTo` debe devolver el valor de `base` elevado a la potencia de `2`.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
