---
id: 699f4a3e7b4c6bfa538cc4f7
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

Si abres la consola, verás el siguiente mensaje de error:

```md
Type 'string' is not assignable to type 'number'.
```

Para resolver este error, necesitas cambiar el return a `sum / numbers.length`. 

# --hints--

Tu función `getAverage` debe devolver un número.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Tu función `getAverage` debe devolver el promedio de los números en el arreglo.

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```
