---
id: 699f412114c87221f16a99e3
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Otro beneficio de usar TypeScript es que detectará errores si pasas demasiados o muy pocos argumentos a una llamada de función.

Para ilustrar esto, actualiza tu llamada a la función `square` para que tome dos argumentos `5` y `10`.

# --hints--

Tu llamada a la función `square` debe tener dos argumentos `5` y `10`.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
