---
id: 699f513e01eb1768cc5df726
title: Paso 17
challengeType: 1
dashedName: step-17
---

# --description--

Ahora que el parámetro `exponent` está marcado como opcional, puedes llamar a la función con solo un argumento.

Actualiza tu llamada a la función `raiseTo` para que solo pase el número `3`.

# --hints--

Tu función `raiseTo` debe ser llamada con solo un argumento `3`.

```js
assert.match(
  code,
  /raiseTo\s*\(\s*3\s*\)/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}

--fcc-editable-region--
const raisedResult = raiseTo(2, 3);
--fcc-editable-region--
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
