---
id: 699f52fac083214b6bf51de6
title: Paso 19
challengeType: 1
dashedName: step-19
---

# --description--

En este último paso, necesitarás agregar una comprobación condicional para el parámetro `exponent` en la función `raiseTo`.

Si se proporciona el argumento `exponent`, la función debe devolver `base` elevado a la potencia de `exponent`. Si no se proporciona el argumento `exponent`, debe devolver `base ** 2`.

¡Con ese último paso, tu conjunto de herramientas matemáticas con seguridad de tipos estará completo!

# --hints--

Tu función `raiseTo` debe devolver `base ** exponent` cuando se proporcione el argumento `exponent`.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

Tu función `raiseTo` debe devolver `base ** 2` cuando no se proporcione el argumento `exponent`.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
