---
id: 699f4ec939f10da62fb56c3e
title: Paso 15
challengeType: 1
dashedName: step-15
---

# --description--

Ahora es momento de probar tu función.

Crea una variable llamada `raisedResult` y asígnale el resultado de llamar a `raiseTo` con los argumentos `2` y `3`.

Luego, muestra `raisedResult` en la consola.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes tener una variable llamada `raisedResult`.

```js
assert.isDefined(raisedResult);
```

Tu variable `raisedResult` debe ser un número.

```js
assert.isNumber(raisedResult);
```

Debes mostrar `raisedResult` en la consola.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
