---
id: 699f1d72e906e4616155290d
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, continuarás aprendiendo sobre anotaciones de tipo y cómo funcionan con las funciones creando un conjunto de herramientas matemáticas con tipos seguros.

Comienza creando una función llamada `square` que acepte un parámetro y devuelva el cuadrado de un número.

# --hints--

Deberías tener una función llamada `square`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

Tu función `square` debe tener un parámetro.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

Tu función `square` debe devolver el cuadrado de un número.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
