---
id: 699f4861dd5b5f164e1bda4a
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

Otro beneficio de usar TypeScript es asegurarte de que los tipos de retorno se cumplan y sean lo que esperas. A veces TypeScript puede inferir el tipo de retorno de una función, pero también puedes especificarlo explícitamente.

Para ilustrar esto, cambia el tipo de retorno de la función `getAverage` a la cadena `"something"`. 

# --hints--

Tu función `getAverage` debe devolver un string.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

Tu función `getAverage` debe devolver la cadena `"something"`.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
