---
id: 699f4b903bf5427a599f791d
title: Paso 14
challengeType: 1
dashedName: step-14
---

# --description--

En los Ăºltimos pasos, construirĂ¡s una funciĂ³n mĂ¡s para seguir practicando las anotaciones de tipos con funciones.

Comienza creando una funciĂ³n llamada `raiseTo` que tome dos parĂ¡metros llamados `base` y `exponent`. Ambos parĂ¡metros deben ser del tipo `number`. La funciĂ³n debe devolver el resultado de elevar `base` a la potencia de `exponent`. El tipo de retorno de la funciĂ³n tambiĂ©n debe estar anotado como `number`.

# --hints--

Debes tener una funciĂ³n `raiseTo`.

```js
assert.isFunction(raiseTo);
```

Tu funciĂ³n `raiseTo` debe tener dos parĂ¡metros, `base` y `exponent`.

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*=>/
);
```

Tus parĂ¡metros `base` y `exponent` deben ser ambos del tipo `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

Tu funciĂ³n `raiseTo` debe devolver un nĂºmero.

```js
assert.isNumber(raiseTo(2, 3));
```

Tu funciĂ³n `raiseTo` debe devolver el resultado de elevar `base` a la potencia de `exponent`.

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

Tu funciĂ³n `raiseTo` debe tener una anotaciĂ³n de tipo de retorno `number`.

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
