---
id: 699f3f636aa840b7201aa37c
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Si abres la consola, deberías ver el siguiente mensaje de error de TypeScript:

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

Ahora que el parámetro está explícitamente tipado como un número, TypeScript puede detectar el error de intentar pasar un argumento de tipo cadena a la función `square`.

Cambia el argumento pasado a la función `square` al número `5` para corregir el error y ver la salida correcta en la consola.

# --hints--

Tu llamada a la función `square` debe tener un argumento de `5` en lugar de `"something"`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
