---
id: 699f4639ae0016caaf81b7c6
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Ahora es momento de probar tu función.

Comienza creando una variable llamada `avgResult` y asignándole la llamada a la función `getAverage`. El argumento para esa llamada debe ser `[2, 14, 26, 8]`.

Luego, muestra el valor de `avgResult` en la consola.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Deberías tener una variable llamada `avgResult`.

```js
assert.isDefined(avgResult);
```

Deberías asignar `getAverage([2, 14, 26, 8])` a tu variable `avgResult`.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

Tu variable `avgResult` debe ser un número.

```js
assert.isNumber(avgResult);
```

Debes mostrar `avgResult` en la consola.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
