---
id: 699f458e47fb88c756fdba81
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

En este momento, el parámetro `numbers` tiene un tipo `any` implícito. Para corregir esto, puedes tipar explícitamente el parámetro como un arreglo de números.

Recuerda de lecciones anteriores que puedes tipar un arreglo de números así:

```ts
const numbers: number[] = [1, 2, 3];
```

Actualiza tu función `getAverage` para tipar explícitamente el parámetro `numbers` como un arreglo de números.

# --hints--

Tu función `getAverage` debe tener un parámetro llamado `numbers` que esté tipado explícitamente como un arreglo de números.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
