---
id: 699f49669ea51a31d2ba89d8
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

Como recuerdas de lecciones anteriores, puedes especificar explícitamente el tipo de retorno de una función así:

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

En este ejemplo, el tipo de retorno de la función `circleArea` es `number`. Si intentas devolver un tipo diferente, TypeScript mostrará un error.

Actualiza tu función `getAverage` para que el tipo de retorno sea explícitamente `number`. 

# --hints--

Tu función `getAverage` debe tener un tipo de retorno `number`.

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
