---
id: 699f3738921fb79a651ded56
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Ahora es momento de llamar a tu función.

Crea una variable llamada `result` y asígnale la llamada a la función `square` con un argumento `5`.

Debajo de eso, muestra en la consola el valor de `result`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Deberías tener una variable llamada `result`.

```js
assert.isDefined(result);
```

Tu variable `result` debe ser un número.

```js
assert.isNumber(result);
```

Tu variable `result` debe ser el cuadrado de `5`.

```js
assert.strictEqual(result, 25);
```

Debes mostrar en la consola la variable `result`.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
