---
id: 699f5063d2393c674b1f45dc
title: Paso 16
challengeType: 1
dashedName: step-16
---

# --description--

Como recuerdas de un paso anterior, si llamas a una función con un número incorrecto de argumentos, TypeScript lanzará un error.

Hay una forma de marcar parámetros como opcionales agregando un `?` después del nombre del parámetro. Esto te permite llamar a la función con menos argumentos que parámetros.

Aquí tienes un ejemplo:

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

En este ejemplo, si el parámetro `greeting` no se proporciona al llamar a la función `greet`, será `undefined`, y la función devolverá un saludo predeterminado.

Actualiza tu función `raiseTo` para hacer que el parámetro `exponent` sea opcional.

# --hints--

Tu función `raiseTo` debe tener el parámetro `exponent` marcado como opcional.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
