---
id: 587d7dbf367417b2b2512bba
title: 使用 @each 來對列表中的項目進行映射
challengeType: 0
forumTopicId: 301461
dashedName: use-each-to-map-over-items-in-a-list
---

# --description--

上一個挑戰展示了 `@for` 指令如何使用起始值和結束值來執行特定次數的迴圈。Sass 也提供了 `@each` 指令，該指令會對列表或映射中的每個項目進行迴圈。在每次迭代中，變數會被指定為列表或映射中的當前值。

```scss
@each $color in blue, red, green {
  .#{$color}-text {color: $color;}
}
```

映射的語法略有不同。以下是一個範例：

```scss
$colors: (color1: blue, color2: red, color3: green);

@each $key, $color in $colors {
  .#{$color}-text {color: $color;}
}
```

請注意，`$key` 變數是用來取用映射中的鍵。否則，編譯後的 CSS 會包含 `color1`、`color2`……。上述兩個程式碼範例都會轉換成以下 CSS：

```scss
.blue-text {
  color: blue;
}

.red-text {
  color: red;
}

.green-text {
  color: green;
}
```

# --instructions--

撰寫一個 `@each` 指令，遍歷一個列表：`blue, black, red`，並將每個變數指定到 `.color-bg` 類別，其中 `color` 部分會根據每個項目變更為相應的顏色。每個類別也應該將 `background-color` 設定為相應的顏色。

# --hints--

你的程式碼應該使用 `@each` 指令。

```js
assert(code.match(/@each /g));
```

你的 `.blue-bg` 類別應該有藍色的 `background-color`。

```js
assert($('.blue-bg').css('background-color') == 'rgb(0, 0, 255)');
```

你的 `.black-bg` 類別應該有黑色的 `background-color`。

```js
assert($('.black-bg').css('background-color') == 'rgb(0, 0, 0)');
```

你的 `.red-bg` 類別應該有紅色的 `background-color`。

```js
assert($('.red-bg').css('background-color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

# --solutions--

```html
<style type='text/scss'>

  @each $color in blue, black, red {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

---

```html
<style type='text/scss'>

  $colors: (color1: blue, color2: black, color3: red);

  @each $key, $color in $colors {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```
