---
id: 587d7dbd367417b2b2512bb6
title: 用 Mixins 創建可重用 CSS
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

In Sass, a <dfn>mixin</dfn> is a group of CSS declarations that can be reused throughout the style sheet. The definition starts with the `@mixin` at-rule, followed by a custom name. You apply the mixin using the `@include` at-rule.

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

編譯爲：

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

你的 mixin 還可以接受參數，從而允許定製它們的行爲。 使用 mixin 時需要參數。

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

你可以通過賦予參數默認值來使參數成爲可選的。

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

編寫一個名爲 `shape` 的 mixin，並賦予其 3 個參數：`$w`、`$h` 和 `$bg-color`。

使用 `shape` 混合爲 `#square` 元素提供 `50px` 的寬度和高度，以及 `red` 的背景顏色。 對於 `#rect-a` 元素，添加寬度 `100px`、高度 `50px` 和背景顏色 `blue`。 最後，爲 `#rect-b` 元素添加寬度 `50px`、高度 `100px` 和背景顏色 `orange`。

# --hints--

你應該聲明一個名爲 `shape` 的 mixin，它有 3 個參數：`$w`、`$h` 和 `$bg-color`。

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

你的 mixin 應該包含一個使用 `$w` 參數的 `width` 屬性。

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

你的 mixin 應該包含一個使用 `$h` 參數的 `height` 屬性。

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

你的 mixin 應該包含一個使用 `$bg-color` 參數的 `background-color` 屬性。

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

你應該使用 `@include` 關鍵字將 `#square` 選擇器內的樣式替換爲對 `shape` 混合的調用。 設置寬度和高度爲 `50px`，背景顏色爲 `red`。

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

你應該使用 `@include` 關鍵字將 `#rect-a` 選擇器內的樣式替換爲對 `shape` 混合的調用。 設置寬度爲 `100px`、高度爲 `50px`、背景顏色爲 `blue`。

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

你應該使用 `@include` 關鍵字將 `#rect-b` 選擇器內的樣式替換爲對 `shape` 混合的調用。 設置寬度爲 `50px`、高度爲 `100px`、背景顏色爲 `orange`。

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
