---
id: 587d7dbd367417b2b2512bb5
title: 使用 Sass 嵌套 CSS
challengeType: 0
forumTopicId: 301457
dashedName: nest-css-with-sass
---

# --description--

Sass 允許巢狀 CSS 規則，這是組織樣式表的有用方法。

通常，每個元素會在不同的行上作為標的來進行樣式設定，如下所示：

```scss
article {
  height: 200px;
}

article p {
  color: white;
}

article ul {
  color: blue;
}
```

對於大型專案，CSS 檔案會有許多行和規則。這時，巢狀結構可以透過將子樣式規則放在相應的父元素內，幫助你組織程式碼：

```scss
article {
  height: 200px;

  p {
    color: white;
  }

  ul {
    color: blue;
  }
}

```

# --instructions--

使用上述的巢狀技巧重新組織 `.blog-post` 元素兩個子元素的 CSS 規則。為了測試目的，`h1` 應該放在 `p` 元素之前。

# --hints--

你的程式碼應該重新組織 CSS 規則，使 `h1` 和 `p` 嵌套在 `.blog-post` 父元素中。

```js
assert(
  code.match(
    /\.blog-post\s*?{\s*?h1\s*?{\s*?text-align:\s*?center;\s*?color:\s*?blue;\s*?}\s*?p\s*?{\s*?font-size:\s*?20px;\s*?}\s*?}/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  .blog-post {

  }
  h1 {
    text-align: center;
    color: blue;
  }
  p {
    font-size: 20px;
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  .blog-post {
    h1 {
      text-align: center;
      color: blue;
    }
    p {
      font-size: 20px;
    }
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```
