---
id: 587d7dbd367417b2b2512bb6
title: 使用 Mixins 創建可重複使用的 CSS
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

在 Sass 中，<dfn>mixin</dfn> 是一組可以在整個樣式表中重複使用的 CSS 宣告。定義以 `@mixin` 這個 at-rule 開頭，後面接著一個訂製名稱。你使用 `@include` 這個 at-rule 來套用 mixin。

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

編譯為：

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

你的 mixins 也可以接受引數，這讓它們的行為能夠被自訂。使用 mixin 時必須提供引數。

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

你可以透過給參數預設值來使引數變成選用。

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

撰寫一個名為 `shape` 的 mixin，並給它 3 個參數：`$w`、`$h` 和 `$bg-color`。

使用 `shape` mixin 為 `#square` 元素設定寬度和高度為 `50px`，以及背景顏色為 `red`。為 `#rect-a` 元素添加寬度為 `100px`、高度為 `50px`，以及背景顏色為 `blue`。最後，為 `#rect-b` 元素添加寬度為 `50px`、高度為 `100px`，以及背景顏色為 `orange`。

# --hints--

你應該宣告一個名為 `shape` 的 mixin，並帶有 3 個參數：`$w`、`$h` 和 `$bg-color`。

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

你的 mixin 應該包含使用 `$w` 參數的 `width` 屬性。

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

你的 mixin 應該包含使用 `$h` 參數的 `height` 屬性。

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

你的 mixin 應該包含使用 `$bg-color` 參數的 `background-color` 屬性。

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

你應該將 `#square` 選擇器內的樣式取代為使用 `@include` 關鍵字對 `shape` mixin 的呼叫。設定寬度和高度為 `50px`，以及背景顏色為 `red`。

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

你應該將 `#rect-a` 選擇器內的樣式取代為使用 `@include` 關鍵字呼叫 `shape` 混入。設定寬度為 `100px`、高度為 `50px`，以及背景顏色為 `blue`。

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

你應該將 `#rect-b` 選擇器內的樣式取代為使用 `@include` 關鍵字呼叫 `shape` 混入。設定寬度為 `50px`、高度為 `100px`，以及背景顏色為 `orange`。

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
