---
id: 587d7dbf367417b2b2512bbc
title: 將你的樣式拆分為較小的零件
challengeType: 0
forumTopicId: 301459
dashedName: split-your-styles-into-smaller-chunks-with-partials
---

# --description--

<dfn>Partials</dfn> 在 Sass 中是獨立的檔案，用來存放 CSS 程式碼的零件。這些檔案會被匯入並用於其他 Sass 檔案中。這是將相似程式碼群組成模組以保持組織良好的絕佳方法。

部分檔案的名稱以底線（`_`）字元開頭，這告訴 Sass 它是 CSS 的一小段，不會將其轉換成 CSS 檔案。此外，Sass 檔案以 `.scss` 副檔名結尾。若要將部分檔案中的程式碼匯入到另一個 Sass 檔案，請使用 `@import` 指令。

例如，如果你的所有 mixin 都存檔在名為 `"_mixins.scss"` 的零件中，並且它們需要在 `main.scss` 檔案中使用，這就是在主檔案中使用它們的方法：

```scss
@import 'mixins'
```

請注意，`import` 述語中不需要底線和副檔名－Sass 會理解它是一個零件。一旦零件被匯入到檔案中，所有變數、混入（mixins）和其他程式碼都可以使用。

# --instructions--

撰寫一個 `@import` 述語，將名為 `_variables.scss` 的零件匯入到 main.scss 檔案中。

# --hints--

你的程式碼應該使用 `@import` 指令，且檔案名稱中不應包含底線。

```js
assert(code.match(/@import\s+?('|")variables\1/gi));
```

# --seed--

## --seed-contents--

```html
<!-- The main.scss file -->
```

# --solutions--

```html
@import 'variables'
```
