---
id: 587d7dbf367417b2b2512bbb
title: 使用 @while 套用樣式直到符合條件
challengeType: 0
forumTopicId: 301454
dashedName: apply-a-style-until-a-condition-is-met-with-while
---

# --description--

`@while` 指令是一個具有類似 JavaScript `while` 迴圈功能的選項。它會持續產生 CSS 規則直到條件達成。

`@for` 挑戰提供了一個創建簡單網格系統的範例。這也可以用 `@while` 來實現。

```scss
$x: 1;
@while $x < 13 {
  .col-#{$x} { width: 100%/12 * $x;}
  $x: $x + 1;
}
```

首先，定義一個變數 `$x` 並將其設為 1。接著，使用 `@while` 指令在 `$x` 小於 13 時創建網格系統。設定 `width` 的 CSS 規則後，將 `$x` 增加 1 以避免無窮迴圈。

# --instructions--

使用 `@while` 來創建一系列具有不同 `font-sizes` 的類別。

應該有 5 個不同的類別，從 `text-1` 到 `text-5`。然後將 `font-size` 設定為 `15px` 乘以目前的索引數字。務必避免無窮迴圈！

# --hints--

你的程式碼應該使用 `@while` 指令。

```js
assert(code.match(/@while /g));
```

你的程式碼應該使用一個從索引 1 開始的索引變數。

```js
assert(code.match(/\$.*:\s*?1;/gi));
```

你的程式碼應該遞增計數器變數。

```js
assert(code.match(/\$(.*)\s*?:\s*\$\1\s*\+\s*1\s*;/gi));
```

你的 `.text-1` 類別應該有 `font-size` 為 `15px`。

```js
assert($('.text-1').css('font-size') == '15px');
```

你的 `.text-2` 類別應該有 `font-size` 為 `30px`。

```js
assert($('.text-2').css('font-size') == '30px');
```

你的 `.text-3` 類別應該有 `font-size` 為 `45px`。

```js
assert($('.text-3').css('font-size') == '45px');
```

你的 `.text-4` 類別應該有 `font-size` 為 `60px`。

```js
assert($('.text-4').css('font-size') == '60px');
```

你的 `.text-5` 類別應該有 `font-size` 為 `75px`。

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>
  $x: 1;
  @while $x < 6 {
    .text-#{$x}{
      font-size: 15px * $x;
    }
    $x: $x + 1;
  }
</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
