---
id: 587d7dbd367417b2b2512bb4
title: 使用 Sass 變數儲存資料
challengeType: 0
forumTopicId: 301460
dashedName: store-data-with-sass-variables
---

# --description--

Sass 與 CSS 不同的一個特性是它使用 `variables`。它們被宣告並設定來儲存資料，類似於 JavaScript。

在 JavaScript 中，變數是使用 `let` 和 `const` 關鍵字定義的。在 Sass 中，變數以 `$` 開頭，後面接變數名稱。

以下是幾個範例：

```scss
$main-fonts: Arial, sans-serif;
$headings-color: green;
```

並且使用這些變數：

```scss
h1 {
  font-family: $main-fonts;
  color: $headings-color;
}
```

變數有用的一個例子是當多個元素需要相同顏色時。如果該顏色被更改，唯一需要編輯程式碼的地方就是變數值。

# --instructions--

創建一個變數 `$text-color` 並將其設為 `red`。然後將 `.blog-post` 和 `h2` 的 `color` 屬性值更改為 `$text-color` 變數。

# --hints--

你的程式碼應該宣告一個 Sass 變數 `$text-color`，其值為 `red`。

```js
assert(code.match(/\$text-color\s*:\s*?red\s*;/g));
```

你的程式碼應該使用 `$text-color` 變數來更改 `.blog-post` 和 `h2` 項目的 `color`。

```js
assert(code.match(/color\s*:\s*\$text-color\s*;?/g));
```

你的 `.blog-post` 元素的 `color` 應該是紅色。

```js
assert($('.blog-post').css('color') == 'rgb(255, 0, 0)');
```

你的 `h2` 元素應該有紅色的 `color`。

```js
assert($('h2').css('color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>


  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: red;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  $text-color: red;

  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: $text-color;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```
