---
id: 587d7dbe367417b2b2512bb8
title: 使用 @if 和 @else 為你的樣式添加邏輯
challengeType: 0
forumTopicId: 301463
dashedName: use-if-and-else-to-add-logic-to-your-styles
---

# --description--

Sass 中的 `@if` 指令用於測試特定情況－它的運作方式就像 JavaScript 中的 `if` 述語。

```scss
@mixin make-bold($bool) {
  @if $bool == true {
    font-weight: bold;
  }
}
```

就像在 JavaScript 中一樣，`@else if` 和 `@else` 指令用來測試更多條件：

```scss
@mixin text-effect($val) {
  @if $val == danger {
    color: red;
  }
  @else if $val == alert {
    color: yellow;
  }
  @else if $val == success {
    color: green;
  }
  @else {
    color: black;
  }
}
```

# --instructions--

創建一個名為 `border-stroke` 的 mixin，該 mixin 接受一個參數 `$val`。該 mixin 應該使用 `@if`、`@else if` 和 `@else` 指令檢查以下條件：

```scss
light - 1px solid black
medium - 3px solid black
heavy - 6px solid black
```

如果 `$val` 參數值不是 `light`、`medium` 或 `heavy`，則 `border` 屬性應設為 `none`。

# --hints--

你的程式碼應該宣告一個名為 `border-stroke` 的 mixin，該 mixin 有一個名為 `$val` 的參數。

```js
assert(code.match(/@mixin\s+?border-stroke\s*?\(\s*?\$val\s*?\)\s*?{/gi));
```

你的 mixin 應該有一個 `@if` 述語來檢查 `$val` 是否為 `light`，並將 `border` 設定為 `1px solid black`。

```js
assert(
  code.match(
    /@if\s+?\$val\s*?===?\s*?light\s*?{\s*?border\s*?:\s*?1px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

你的 mixin 應該有一個 `@else if` 述語來檢查 `$val` 是否為 `medium`，並將 `border` 設定為 `3px solid black`。

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?medium\s*?{\s*?border\s*?:\s*?3px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

你的 mixin 應該有一個 `@else if` 述語來檢查 `$val` 是否為 `heavy`，並將 `border` 設定為 `6px solid black`。

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?heavy\s*?{\s*?border\s*?:\s*?6px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

你的 mixin 應該有一個 `@else` 述語來將 `border` 設為 `none`。

```js
assert(code.match(/@else\s*?{\s*?border\s*?:\s*?none\s*?;\s*?}/gi));
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```

# --solutions--

```html
<style type='text/scss'>
  @mixin border-stroke($val) {
    @if $val == light {
      border: 1px solid black;
    }
    @else if $val == medium {
      border: 3px solid black;
    }
    @else if $val == heavy {
      border: 6px solid black;
    }
    @else {
      border: none;
    }
  }


  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```
