---
id: 587d7dbe367417b2b2512bb9
title: 使用 @for 來創建 Sass 迴圈
challengeType: 0
forumTopicId: 301462
dashedName: use-for-to-create-a-sass-loop
---

# --description--

`@for` 指令會在迴圈中添加樣式，非常類似 JavaScript 中的 `for` 迴圈。

`@for` 有兩種用法：「start through end」或「start to end」。主要差異在於「start **to** end」*不包含*結束數字作為計數的一部分，而「start **through** end」*包含*結束數字作為計數的一部分。

這裡是一個從 **through** 到結束的範例：

```scss
@for $i from 1 through 12 {
  .col-#{$i} { width: 100%/12 * $i; }
}
```

`#{$i}` 零件是將變數（`i`）與文字結合成字串的語法。當 Sass 檔案轉換成 CSS 時，看起來像這樣：

```scss
.col-1 {
  width: 8.33333%;
}

.col-2 {
  width: 16.66667%;
}

...

.col-12 {
  width: 100%;
}
```

這是一種強大的方式來創建網格佈局。現在你有十二個可用作為 CSS 類別的行寬選項。

# --instructions--

撰寫一個 `@for` 指令，該指令帶有一個從 1 **到** 6 的變數 `$j`。

它應該創建 5 個名為 `.text-1` 到 `.text-5` 的類別，每個類別的 `font-size` 設定為 15px 乘以索引。

# --hints--

你的程式碼應該使用 `@for` 指令。

```js
assert(code.match(/@for /g));
```

你的 `.text-1` 類別應該有 15px 的 `font-size`。

```js
assert($('.text-1').css('font-size') == '15px');
```

你的 `.text-2` 類別應該有 30px 的 `font-size`。

```js
assert($('.text-2').css('font-size') == '30px');
```

你的 `.text-3` 類別應該有 45px 的 `font-size`。

```js
assert($('.text-3').css('font-size') == '45px');
```

你的 `.text-4` 類別應該有 60px 的 `font-size`。

```js
assert($('.text-4').css('font-size') == '60px');
```

你的 `.text-5` 類別應該有 75px 的 `font-size`。

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>

@for $i from 1 through 5 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

---

```html
<style type='text/scss'>

@for $i from 1 to 6 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
