---
id: 587d7fa5367417b2b2512bbd
title: 將一組 CSS 樣式擴展到另一個元素
challengeType: 0
forumTopicId: 301456
dashedName: extend-one-set-of-css-styles-to-another-element
---

# --description--

Sass 有一個名為 `extend` 的特性，可以輕鬆地從一個元素借用 CSS 規則，並在另一個元素上進行建置。

例如，下面的 CSS 規則區段為 `.panel` 類別設定樣式。它具有 `background-color`、`height` 和 `border`。

```scss
.panel{
  background-color: red;
  height: 70px;
  border: 2px solid green;
}
```

現在你想要另一個名為 `.big-panel` 的面板。它具有與 `.panel` 相同的基本屬性，但還需要 `width` 和 `font-size`。雖然可以複製並貼上 `.panel` 的初始 CSS 規則，但隨著你添加更多類型的面板，程式碼會變得重複。`extend` 指令是一種簡單的方法，可以重複使用為一個元素撰寫的規則，然後為另一個元素添加更多規則：

```scss
.big-panel{
  @extend .panel;
  width: 150px;
  font-size: 2em;
}
```

`.big-panel` 將會擁有與 `.panel` 相同的屬性，並且附加新的樣式。

# --instructions--

建立一個繼承自 `.info` 且 `background-color` 設為洋紅色的類別 `.info-important`。

# --hints--

你的 `info-important` 類別應該設定 `background-color` 為 `magenta`。

```js
assert(
  code.match(
    /\.info-important\s*?{[\s\S]*background-color\s*?:\s*?magenta\s*?;[\s\S]*}/gi
  )
);
```

你的 `info-important` 類別應該使用 `@extend` 來繼承 `info` 類別的樣式。

```js
assert(
  code.match(/\.info-important\s*?{[\s\S]*@extend\s*?.info\s*?;[\s\S]*/gi)
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }




</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }
  .info-important{
    @extend .info;
    background-color: magenta;
  }



</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```
