---
id: 642df9df4b5216350de7b0d2
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Вам знадобиться функція для створення діапазону чисел.

Оголосіть порожню функцію `range`, яка приймає параметри `start` та `end`. Використайте конструктор `Array()` та неявно поверніть порожній масив.

# --hints--

Оголосіть змінну `range`.

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

Використайте `const`, щоб оголосити змінну `range`.

```js
assert.match(code, /const\s+range/);
```

Змінна `range` має бути функцією.

```js
assert.isFunction(range);
```

Функція `range` повинна використовувати стрілковий синтаксис.

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

Функція `range` спочатку має приймати параметр `start`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

Функція `range` має приймати параметр `end` другим.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

Функція `range` має використати неявне повернення. Пам’ятайте, що це означає, що ви не будете використовувати фігурні дужки.

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

Функція `range` має використовувати конструктор `Array()`. Примітивним конструкторам не потрібне ключове слово `new`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

Ви не повинні нічого передавати у конструктор `Array()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
