---
id: 642db8c409d9991d0b3b2f0d
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Ваш проєкт починається з базового HTML-контейнера та відповідного CSS. Ваше перше завдання — програмно згенерувати клітинки для вашої електронної таблиці.

Глобальний об’єкт `window` представляє вікно браузера (або вкладку). Він має властивість `onload`, яка дозволяє визначити поведінку, коли вікно завантажило всю сторінку, включно зі стилями та скриптами.

Почніть із призначення властивості `onload` об’єкта `window` стрілковою функцією без параметрів. У функції оголосіть змінну `container` і призначте їй значення, отримане за допомогою методу отримання елемента за `id` об’єкта `"container"`.

# --hints--

Ви повинні отримати доступ до властивості `onload` об’єкта `window`.

```js
assert.match(code, /window\.onload/);
```

Ви повинні призначити властивість `onload` функції.

```js
assert.isFunction(window.onload);
```

Використовуйте стрілковий синтаксис.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

Ваша функція `onload` не повинна приймати жодних параметрів.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

Оголосіть змінну `container` у вашій функції `onload`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

Змінна `container` має бути оголошена за допомогою `const`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

Використовуйте `document.getElementById()`.

```js
assert.match(code, /document\.getElementById\(/);
```

Отримайте елемент за `id` об’єкта `"container"`.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

Призначте елемент `#container` змінній `container`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
