---
id: 642db8c409d9991d0b3b2f0d
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Проєкт починається з базового HTML контейнера та трохи відповідного CSS. Вашим першим завданням буде програмно генерувати клітинки для електронної таблиці.

Глобальний об’єкт `window` представляє вікно переглядача (або вкладку). Він має властивість `onload`, яка дозволяє визначити поведінку, коли вікно завантажило всю сторінку, включно з таблицею стилів та сценаріями.

Почніть із встановлення властивості `onload` з `window` на стрілкову функцію без параметрів. У функції оголосіть змінну `container` та призначте до неї значення отримання елемента за `id` зі значенням `"container"`.

# --hints--

Отримайте доступ до властивості `onload` об’єкта `window`.

```js
assert.match(code, /window\.onload/);
```

Встановіть властивість `onload` на функцію.

```js
assert.isFunction(window.onload);
```

Використайте стрілковий синтаксис.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

Функція `onload` не повинна приймати жодного параметра.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

Оголосіть змінну `container` у функції `onload`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

Змінна `container` має бути оголошеною за допомогою `const`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

Використайте `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

Отримайте елемент з `id` зі значенням `"container"`.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

Ви повинні присвоїти елемент `#container` до `container`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
