---
id: 643715013330824ecaa70442
title: Крок 23
challengeType: 0
dashedName: step-23
---

# --description--

Ваша наступна функція має обчислити медіанне значення масиву чисел. Почніть з оголошення `median` стрілкової функції, яка приймає `nums` параметр.

У функції оголосіть `sorted` змінну і призначте їй значення сортування копії масиву `nums`.

Ви повинні використати метод `slice()` для створення поверхневої копії масиву.

# --hints--

Ви повинні оголосити `median` змінну.

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

Ви повинні використати `const` для оголошення вашої `median` змінної.

```js
assert.match(code, /const\s+median/);
```

Ваша `median` змінна має бути функцією.

```js
assert.isFunction(median);
```

Ваша `median` функція має бути стрілковою.

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

Ваша `median` функція має мати `nums` параметр.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

Ваша `median` функція не повинна використовувати неявний return.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

Ваша `median` функція має мати `sorted` змінну.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

Ви повинні використати `const` для оголошення вашої `sorted` змінної.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

Ви повинні використати `.slice()` для призначення копії масиву `nums` змінній `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

Ви повинні додати метод `.sort()` до вашого методу `.slice()`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

Ви повинні передати функцію зворотного виклику до вашого методу `sort`, щоб правильно відсортувати числа за зростанням. Використайте неявний return для зрозумілості.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
