---
id: 64496df724dd3716a71fe971
title: Крок 26
challengeType: 0
dashedName: step-26
---

# --description--

Властивості об’єктів складаються з пар ключ/значення. Ви можете використовувати скорочені імена властивостей при оголошенні літерала об’єкта. Коли ви використовуєте синтаксис скороченої властивості, ім’я змінної стає ключем властивості, а її значення — значенням властивості.

У наступному прикладі оголошується об’єкт `user` з властивостями `userId`, `firstName` та `loggedIn`.

```js
const userId = 1;
const firstName = "John";
const loggedIn = true;

const user = {
  userId,
  firstName,
  loggedIn,
};

console.log(user); // { userId: 1, firstName: 'John', loggedIn: true }
```

Щоб відстежувати всі функції вашої електронної таблиці, оголосіть об’єкт `spreadsheetFunctions`. Використовуючи синтаксис скороченого запису, призначте `sum`, `average` та `median` як властивості об’єкта `spreadsheetFunctions`.

# --hints--

Ви повинні оголосити змінну `spreadsheetFunctions`.

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

Ви повинні використати `const` для оголошення змінної `spreadsheetFunctions`.

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

Ваша змінна `spreadsheetFunctions` має бути об’єктом.

```js
assert.isObject(spreadsheetFunctions);
```

Ваш об’єкт `spreadsheetFunctions` має мати властивість `sum`.

```js
assert.property(spreadsheetFunctions, "sum");
```

Ваша властивість `sum` має бути вашою функцією `sum`.

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

Ваш об’єкт `spreadsheetFunctions` має мати властивість `average`.

```js
assert.property(spreadsheetFunctions, "average");
```

Ваша властивість `average` має бути вашою функцією `average`.

```js
assert.equal(spreadsheetFunctions?.average, average);
```

Ваш об’єкт `spreadsheetFunctions` має мати властивість `median`.

```js
assert.property(spreadsheetFunctions, "median");
```

Ваша властивість `median` має бути вашою функцією `median`.

```js
assert.equal(spreadsheetFunctions?.median, median);
```

Ви повинні використати синтаксис деструктуризації, щоб призначити свої властивості.

```js
const objectText = code.replace(/.*const\s+spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
