---
id: 68599af28ead38c22590a89d
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Serve un modo per collegare i pulsanti al contenuto che rappresentano. Questo si fa usando l'attributo `aria-controls` su ogni pulsante. Più avanti aggiungerai il contenuto e i corrispondenti attributi `id` agli elementi di contenuto.

Inizia aggiungendo l'attributo `aria-controls` a ciascuno degli elementi `button`. Il valore per l'attributo `aria-controls` del primo pulsante dovrebbe essere impostato su `"panel-earth"`. L'attributo `aria-controls` del secondo pulsante dovrebbe essere impostato su `"panel-saturn"`. L'attributo `aria-controls` del terzo pulsante dovrebbe essere impostato su `"panel-mars"`.

# --hints--

Ognuno dei tuoi elementi `button` dovrebbe avere un attributo `aria-controls`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

Il primo elemento `button` dovrebbe avere un attributo `aria-controls` impostato su `"panel-earth"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
``` 

Il secondo elemento `button` dovrebbe avere un attributo `aria-controls` impostato su `"panel-saturn"`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

Il terzo elemento `button` dovrebbe avere un attributo `aria-controls` impostato su `"panel-mars"`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab">🌍 Earth</button>
        <button role="tab">🪐 Saturn</button>
        <button role="tab">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
