---
id: 68599633407374c0d574587b
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Il passaggio successivo è aggiungere i pulsanti che saranno responsabili di passare tra i diversi pianeti e i loro fatti.

All'interno dell'elemento `tablist`, aggiungi tre elementi `button` ciascuno con un attributo `role` impostato su `"tab"`. Il primo elemento `button` dovrebbe avere il testo `🌍 Earth`. Il secondo `button` dovrebbe avere il testo `🪐 Saturn`. Il terzo elemento `button` dovrebbe avere il testo `🔴 Mars`.  

# --hints--

Dovresti avere tre elementi `button` all'interno dell'elemento `tablist`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button");
assert.lengthOf(btns, 3);
```

Ognuno dei tuoi elementi `button` dovrebbe avere un attributo `role` impostato su `"tab"`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[role='tab']");
assert.lengthOf(btns, 3);
```

Il tuo primo elemento `button` dovrebbe avere il testo `🌍 Earth`.

```js
const firstBtn = document.querySelector("[role='tab']");
assert.equal(firstBtn.textContent, "🌍 Earth");
```

Il tuo secondo elemento `button` dovrebbe avere il testo `🪐 Saturn`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.textContent, "🪐 Saturn")
```

Il tuo terzo elemento `button` dovrebbe avere il testo `🔴 Mars`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2]
assert.equal(thirdBtn.textContent, "🔴 Mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./style.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
