---
id: 685976a92ffab8b2dc744b72
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

All'interno del tuo elemento `.tabs`, aggiungi un elemento `div` con un attributo `role` impostato su `"tablist"` e un attributo `aria-labelledby` impostato su `"tabs-title"`.  

L'attributo `aria-labelledby` è necessario qui per associare il tablist a un'etichetta visibile, aiutando le tecnologie assistive come i lettori di schermo a comunicare agli utenti lo scopo del tablist.

# --hints--

Dovresti avere un elemento `div` all'interno del tuo elemento `.tabs`.

```js
assert.exists(document.querySelector(".tabs div"));
```

Il tuo elemento `div` dovrebbe avere un attributo `role` impostato su `"tablist"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

Il tuo elemento `div` dovrebbe avere un attributo `aria-labelledby` impostato su `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
