---
id: 68599f9df77791c4a311cc5f
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

L'ultimo elemento da aggiungere a ciascuno degli elementi `button` sono gli attributi `id`. Questi verranno usati per collegare gli elementi `button` ai loro pannelli di contenuto corrispondenti.

Il primo elemento `button` dovrebbe avere un attributo `id` con valore `"tab-earth"`, il secondo dovrebbe avere un attributo `id` con valore `"tab-saturn"`, e il terzo dovrebbe avere un attributo `id` con valore `"tab-mars"`.

# --hints--

Il tuo primo elemento `button` dovrebbe avere un attributo `id`.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[0].hasAttribute("id")
);
```

Il tuo primo elemento `button` dovrebbe avere un attributo `id` impostato su `"tab-earth"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("id"), "tab-earth");
```

Il tuo secondo elemento `button` dovrebbe avere un attributo `id`.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[1].hasAttribute("id")
);
```

Il tuo secondo elemento `button` dovrebbe avere un attributo `id` impostato su `"tab-saturn"`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("id"), "tab-saturn");
```

Il tuo terzo elemento `button` dovrebbe avere un attributo `id`.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[2].hasAttribute("id")
);
```

Il tuo terzo elemento `button` dovrebbe avere un attributo `id` impostato su `"tab-mars"`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("id"), "tab-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-earth" aria-selected="true">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
