---
id: 6859a1270ac255c571352189
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Ora è il momento di creare il markup per ciascuno dei pannelli che mostreranno i fatti su ogni pianeta.

Inizia aggiungendo tre elementi `div`, ognuno con un attributo `role` impostato su `"tabpanel"`. All'interno di ogni elemento `div` ci deve essere un elemento `p`.

# --hints--

Dovresti avere tre elementi `div` con un attributo `role` impostato su `"tabpanel"`.

```js
assert.lengthOf(document.querySelectorAll('div[role="tabpanel"]'), 3);
```

Ogni elemento `tabpanel` dovrebbe contenere un elemento `p` al suo interno.

```js
assert.lengthOf(document.querySelectorAll('div[role="tabpanel"] p'), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
