---
id: 685db0bd1280154ae85bcf9f
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop costruirai una tablist di pianeti accessibile usando HTML, CSS e JavaScript. Il codice boilerplate HTML e il CSS sono già forniti.

In questi primi passaggi, costruirai il resto dell'HTML per la tablist.

Inizia creando un elemento `div` con una classe impostata su `"tabs"`. All'interno di quell'elemento `div`, aggiungi un elemento `h2` con un attributo `id` impostato su `"tabs-title"` e il contenuto testuale `Planets`.

# --hints--

Dovresti avere un elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Il tuo elemento `div` dovrebbe avere una classe impostata su `"tabs"`.

```js
assert.exists(document.querySelector("div.tabs"));
```

Il tuo elemento `div` dovrebbe contenere un elemento `h2` al suo interno.

```js
assert.exists(document.querySelector(".tabs h2"));
```

Il tuo elemento `h2` dovrebbe avere un attributo `id` impostato su `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

Il tuo elemento `h2` dovrebbe avere il contenuto testuale `Planets`.

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
