---
id: 699b70b10ae344c802d106b0
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

이제 `userRoles` 객체에 몇 가지 속성을 더 추가할 차례입니다.

값이 `moderator`인 `"medium-access"` 속성을 추가하세요.

그 아래에 값이 `guest`인 `3` 속성을 추가하세요.

# --hints--

`userRoles` 객체에 `moderator` 속성이 있어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

`moderator` 속성의 값은 `"medium-access"`이어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

`userRoles` 객체에 `guest` 속성이 있어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

`guest` 속성의 값은 숫자 `3`이어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"

};
--fcc-editable-region--

console.log(userRoles);
```
