---
id: 699b6dd1a2961bd38491ef2e
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

`profile` 객체는 엄격한 키와 타입을 가졌습니다. 정의되지 않은 속성을 추가하거나 호환되지 않는 타입을 할당하려 하면 TypeScript가 오류를 발생시켰습니다.

`userRoles` 객체가 그렇게 엄격하지 않고 더 유연한 키를 가졌으면 좋겠습니다.

이전 수업에서 기억하듯이, `Record` 타입을 사용하면 이 결과를 얻을 수 있습니다:

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

이 `studentGrades` 예제에서 TypeScript는 모든 과목을 미리 알 필요가 없습니다. 이렇게 하면 나중에 `Art` 또는 `Music` 같은 과목을 추가할 수 있습니다. TypeScript는 각 키의 값이 숫자인지만 신경 씁니다.

`userRoles` 객체를 업데이트하여 인라인 `Record` 타입을 추가하세요. 각 키가 임의의 문자열이고, 각 값은 문자열이어야 하도록 `<string, string>` 으로 설정하세요.

# --hints--

`userRoles` 객체에 인라인 `Record` 타입이 `Record<string, string>` 으로 설정되어 있어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
