---
id: 699b61ca8c9cb9bc1968364e
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

콘솔을 열면 다음과 같은 오류 메시지가 표시될 것입니다:

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

순수 JavaScript에서는 `profile` 객체에 원하는 속성과 값을 마음대로 추가할 수 있었습니다. 하지만 TypeScript에서는 컴파일러가 객체에 정의한 속성(`username`, `age`, `isLoggedIn` 이 경우)을 가진 것만 허용하는지 검사합니다. `mood` 가 타입에 포함되지 않았기 때문에 TypeScript는 이를 잘못된 속성으로 간주하고 오류를 발생시킵니다.

코드에서 `mood: null` 를 제거하세요.

# --hints--

코드에 더 이상 `mood: null` 가 없어야 합니다.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
