---
id: 699b5370a43e299b4673bcaa
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

`profile` 객체 프로퍼티를 작성하면서 콘솔에서 확인하면 도움이 됩니다.

`console.log`를 추가하여 `profile`를 콘솔에 출력하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`profile`를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
