---
id: 699b68f94a0d154c2a7a5ce0
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

`profile` 객체가 완성되었으니, 이제 사용자 역할을 나타내는 새 객체를 만들어야 합니다.

`userRoles`이라는 새 변수를 만들고 객체를 할당하세요. 객체에는 다음 속성과 값이 포함되어야 합니다.

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

`userRoles`라는 변수가 있어야 합니다.

```js
assert.isDefined(userRoles);
```

`userRoles` 변수는 객체여야 합니다.

```js
assert.isObject(userRoles);
```

`userRoles` 객체에는 `admin` 속성이 있어야 합니다.

```js
assert.property(userRoles, "admin");
```

`admin` 속성의 값은 `"full-access"`이어야 합니다.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

`userRoles` 객체에는 `editor` 속성이 있어야 합니다.

```js
assert.property(userRoles, "editor");
```

`editor` 속성의 값은 `"limited-access"`이어야 합니다.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

`userRoles` 객체에는 `viewer` 속성이 있어야 합니다.

```js
assert.property(userRoles, "viewer");
```

`viewer` 속성의 값은 `"read-only"`이어야 합니다.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
