---
id: 699b6d91cde15459d3619c4a
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

`profile` 객체처럼 `userRoles` 객체도 콘솔에서 확인할 수 있으면 좋겠습니다.

`console.log()`를 추가하여 `userRoles` 객체를 출력하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`userRoles`를 콘솔에 출력해야 합니다.

```js
assert.equal(getLogs()[1], userRoles);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

--fcc-editable-region--

--fcc-editable-region--
```
