---
id: 699b5b7cb68b6d35afca61b7
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

`profile` 객체에 이렇게 속성을 계속 추가할 수는 있지만, 허용할 속성의 유형을 제한하는 것이 좋습니다. `profile` 객체의 형태를 미리 정의하는 편이 더 낫습니다.

이전 수업에서 다음과 같은 객체 유형을 다루는 방법을 배웠습니다:

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

이 예제에서 `composer`는 인라인 객체 유형을 가진 객체입니다. 이 객체는 `name`, `instrument`, `isActive` 세 개의 필수 속성을 가집니다. 처음 두 속성은 `string` 유형이어야 하고 `isActive`은 `boolean` 유형이어야 합니다.

기존 `profile` 객체를 인라인 객체 유형을 포함하도록 업데이트하세요. `username`은 `string` 유형이어야 하고, `age`는 `number` 유형, `isLoggedIn`는 `boolean` 유형이어야 합니다.

# --hints--

`profile` 객체는 `username` 속성이 `string`로 설정된 인라인 객체 유형이어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

`profile` 객체는 `age` 속성이 `number`로 설정된 인라인 객체 유형이어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

`profile` 객체는 `isLoggedIn` 속성이 `boolean`로 설정된 인라인 객체 유형이어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}
--fcc-editable-region--

console.log(profile);
```
