---
id: 699b62c99ec8c63470313365
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

이제 `profile` 객체에 또 다른 속성을 추가할 차례입니다.

인라인 객체 타입을 업데이트하여 `bio` 속성을 포함하고, 그 타입을 `string`로 설정하세요.

# --hints--

인라인 객체 타입을 업데이트하여 `bio` 속성을 포함하고, 그 값을 `string`로 설정하세요.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
