---
id: 699b64096301fd75e3efeb1b
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

콘솔을 다시 열면 다른 오류 메시지가 나타납니다:

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

이전 수업에서 기억하듯이, 속성 이름 옆에 `?` 기호를 사용해 속성을 선택 사항으로 표시할 수 있습니다:

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

이 예제에서 `numberOfAlbums`은 선택 사항으로 표시되어 TypeScript가 이 경우 오류를 발생시키지 않습니다.

`bio` 속성을 선택 사항으로 표시하도록 업데이트하세요.

# --hints--

`bio` 속성은 선택 사항으로 표시되어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
