---
id: 699b65caf17a8e8d340af27d
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

이제 `profile` 객체에 마지막 속성을 추가할 시간입니다.

이전 수업에서 배열 타입을 다루는 방법을 배웠습니다. `string[]` 또는 `number[]` 같은 다양한 배열 타입에 대해 배웠죠:

```ts
const scores: number[] = [95, 87, 100];
```

이 예제에서 `scores` 배열은 숫자만 허용합니다. 만약 `string`나 `null` 같은 다른 타입을 추가하려 하면 TypeScript가 오류를 발생시킵니다.

`profile` 객체를 업데이트하여 `programmingLanguages` 속성을 `string[]` 타입으로 추가하세요. 객체 리터럴 안에서 이 속성의 값을 `["JavaScript", "Python", "C++"]`로 설정하세요.

# --hints--

`programmingLanguages` 객체 안에 `profile` 속성이 있어야 합니다.

```js
assert.property(profile, "programmingLanguages");
```

`programmingLanguages` 속성의 값은 `["JavaScript", "Python", "C++"]`이어야 합니다.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

`programmingLanguages` 속성은 `string[]` 타입이어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};
--fcc-editable-region--

console.log(profile);
```
