---
id: 699b4e284291adcfcb90df47
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 원시 타입과 객체 및 배열 타입을 다루는 연습을 하면서 타입 안전한 사용자 프로필을 만들어 보세요.

먼저 `profile`라는 변수를 만들고 객체를 할당하세요. 객체에는 다음 세 가지 속성과 값이 있어야 합니다:

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

`profile`라는 변수가 있어야 합니다.

```js
assert.exists(profile);
```

`profile` 변수는 객체여야 합니다.

```js
assert.isObject(profile);
```

`username` 객체에 `profile` 속성이 있어야 합니다.

```js
assert.property(profile, "username");
```

`username` 속성의 값은 `"codeLearner"`이어야 합니다.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

`profile` 객체에 `age` 속성이 있어야 합니다.

```js
assert.property(profile, "age");
```

`age` 속성의 값은 `25`이어야 합니다.

```js
assert.propertyVal(profile, "age", 25);
```

`profile` 객체에 `isLoggedIn` 속성이 있어야 합니다.

```js
assert.property(profile, "isLoggedIn");
```

`isLoggedIn` 속성의 값은 `false`이어야 합니다.

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
