---
id: 699b71d98c820ac7ab9eea2a
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

콘솔을 다시 열면 새로운 오류 메시지가 나타납니다:

```md
Type 'number' is not assignable to type 'string'.
```

이는 숫자 값을 가진 `guest` 속성을 추가하려고 시도했기 때문입니다.

TypeScript 오류를 없애려면 `guest` 속성 값을 `"read-only"`로 변경하세요.

마지막 변경으로 사용자 프로필 워크숍이 완료됩니다!

# --hints--

`guest` 객체의 `userRoles` 속성 값은 `"read-only"`여야 합니다.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
