---
id: 6859ae3ef795c4c9d4badf9c
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

이제 탭 목록에 기능을 추가할 차례입니다. 첫 번째 단계는 탭 목록에서 모든 탭과 패널 요소에 접근하는 것입니다.

`tabs`라는 변수를 만들고 `role` 속성이 `tab`로 설정된 모든 요소를 쿼리한 결과를 할당하세요.

그다음 `panels`라는 변수를 만들고 `role` 속성이 `tabpanel`로 설정된 모든 요소를 쿼리한 결과를 할당하세요.

여러 요소를 한 번에 선택하려면 `querySelectorAll` 메서드를 사용할 수 있다는 점을 기억하세요.

# --hints--

`tabs`라는 변수가 있어야 합니다.

```js
assert.isDefined(tabs);
```

`document.querySelectorAll('[role="tab"]')` 변수에 `tabs`를 할당해야 합니다.

```js
assert.deepEqual(
  Array.from(tabs),
  Array.from(document.querySelectorAll('[role="tab"]'))
);
```

`panels`라는 변수가 있어야 합니다.

```js
assert.isDefined(panels);
```

`document.querySelectorAll('[role="tabpanel"]')` 변수에 `panels`를 할당해야 합니다.

```js
assert.deepEqual(
  Array.from(panels),
  Array.from(document.querySelectorAll('[role="tabpanel"]'))
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      <div id="panel-earth" role="tabpanel" aria-labelledby="tab-earth">
        <p>
          Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It's the only planet in the solar system known to harbor life.
        </p>
      </div>
      <div id="panel-saturn" role="tabpanel" aria-labelledby="tab-saturn" hidden>
        <p>
          Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It's a gas giant with dozens of moons orbiting it.
        </p>
      </div>
      <div id="panel-mars" role="tabpanel" aria-labelledby="tab-mars" hidden>
        <p>
          Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It's a key focus for exploration in the search for past or present life.
        </p>
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
