---
id: 68599e9259c0eec3e78c52fd
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

기본적으로 다른 탭들은 선택되지 않아야 합니다. 두 번째와 세 번째 `aria-selected` 요소에 `button` 속성을 추가하고 그 값을 `"false"`로 설정하면 됩니다.

# --hints--

두 번째 `button` 요소에는 `aria-selected` 속성이 있어야 합니다.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[1].hasAttribute("aria-selected")
);
```

두 번째 `button` 요소에는 `aria-selected` 속성이 있고 그 값은 `"false"`로 설정되어야 합니다.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-selected"), "false");
``` 

세 번째 `button` 요소에는 `aria-selected` 속성이 있어야 합니다.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[2].hasAttribute("aria-selected")
);
```

세 번째 `button` 요소에는 `aria-selected` 속성이 있고 그 값은 `"false"`로 설정되어야 합니다.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-selected"), "false");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true">🌍 Earth</button>
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
