---
id: 6859a1270ac255c571352189
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

이제 각 행성에 대한 정보를 표시할 패널의 마크업을 만들 차례입니다.

`div` 요소 세 개를 추가하고 각각에 `role` 속성을 `"tabpanel"`로 설정하세요. 각 `div` 요소 안에는 `p` 요소가 있어야 합니다.

# --hints--

`div` 요소 세 개가 있어야 하며, 각각에 `role` 속성이 `"tabpanel"`로 설정되어 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div[role="tabpanel"]'), 3);
```

각 `tabpanel` 요소 안에는 `p` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div[role="tabpanel"] p'), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
