---
id: 68599f9df77791c4a311cc5f
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

각 `button` 요소에 추가할 마지막 항목은 `id` 속성입니다. 이 속성들은 `button` 요소를 해당 콘텐츠 패널과 연결하는 데 사용됩니다.

첫 번째 `button` 요소는 `id`가 `"tab-earth"`여야 하고, 두 번째는 `id`가 `"tab-saturn"`여야 하며, 세 번째는 `id`가 `"tab-mars"`여야 합니다.

# --hints--

첫 번째 `button` 요소에는 `id` 속성이 있어야 합니다.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[0].hasAttribute("id")
);
```

첫 번째 `button` 요소에는 `id` 속성이 `"tab-earth"`로 설정되어야 합니다.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("id"), "tab-earth");
```

두 번째 `button` 요소에는 `id` 속성이 있어야 합니다.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[1].hasAttribute("id")
);
```

두 번째 `button` 요소에는 `id` 속성이 `"tab-saturn"`로 설정되어야 합니다.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("id"), "tab-saturn");
```

세 번째 `button` 요소에는 `id` 속성이 있어야 합니다.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[2].hasAttribute("id")
);
```

세 번째 `button` 요소에는 `id` 속성이 `"tab-mars"`로 설정되어야 합니다.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("id"), "tab-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-earth" aria-selected="true">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
