---
id: 68599633407374c0d574587b
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

다음 단계는 서로 다른 행성과 그에 관한 사실들을 전환하는 역할을 할 버튼을 추가하는 것입니다.

`tablist` 요소 안에 `button` 속성이 `role`로 설정된 `"tab"` 요소 세 개를 추가하세요. 첫 번째 `button` 요소에는 `🌍 Earth`라는 텍스트가 있어야 합니다. 두 번째 `button` 요소에는 `🪐 Saturn`라는 텍스트가 있어야 합니다. 세 번째 `button` 요소에는 `🔴 Mars`라는 텍스트가 있어야 합니다.  

# --hints--

`button` 요소 안에 `tablist` 요소 세 개가 있어야 합니다.

```js
const btns = document.querySelectorAll("div[role='tablist'] button");
assert.lengthOf(btns, 3);
```

각 `button` 요소에는 `role`로 설정된 `"tab"` 속성이 있어야 합니다.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[role='tab']");
assert.lengthOf(btns, 3);
```

첫 번째 `button` 요소에는 `🌍 Earth`라는 텍스트가 있어야 합니다.

```js
const firstBtn = document.querySelector("[role='tab']");
assert.equal(firstBtn.textContent, "🌍 Earth");
```

두 번째 `button` 요소에는 `🪐 Saturn`라는 텍스트가 있어야 합니다.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.textContent, "🪐 Saturn")
```

세 번째 `button` 요소에는 `🔴 Mars`라는 텍스트가 있어야 합니다.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2]
assert.equal(thirdBtn.textContent, "🔴 Mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./style.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
