---
id: 685976a92ffab8b2dc744b72
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`.tabs` 요소 안에 `div` 요소를 추가하고, `role` 속성을 `"tablist"`으로, `aria-labelledby` 속성을 `"tabs-title"`로 설정하세요.  

`aria-labelledby` 속성은 탭리스트를 눈에 보이는 레이블과 연결하는 데 필요하며, 화면 읽기 프로그램 같은 보조 기술이 탭리스트의 목적을 사용자에게 알리는 데 도움을 줍니다.

# --hints--

`div` 요소 안에 `.tabs` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector(".tabs div"));
```

`div` 요소에 `role` 속성을 `"tablist"`로 설정하세요.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

`div` 요소에 `aria-labelledby` 속성을 `"tabs-title"`로 설정하세요.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
