---
id: 68599af28ead38c22590a89d
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

버튼을 나타내는 콘텐츠와 연결하는 방법이 필요합니다. 이는 각 버튼에 `aria-controls` 속성을 사용해서 합니다. 나중에 콘텐츠와 해당하는 `id` 속성을 콘텐츠 요소에 추가할 것입니다.

먼저 각 `aria-controls` 요소에 `button` 속성을 추가하세요. 첫 번째 버튼의 `aria-controls` 속성 값은 `"panel-earth"`로 설정해야 합니다. 두 번째 버튼의 `aria-controls` 속성 값은 `"panel-saturn"`로 설정해야 합니다. 세 번째 버튼의 `aria-controls` 속성 값은 `"panel-mars"`로 설정해야 합니다.

# --hints--

각 `button` 요소에는 `aria-controls` 속성이 있어야 합니다.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

첫 번째 `button` 요소에는 `aria-controls` 속성이 `"panel-earth"`로 설정되어야 합니다.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
``` 

두 번째 `button` 요소에는 `aria-controls` 속성이 `"panel-saturn"`로 설정되어야 합니다.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

세 번째 `button` 요소에는 `aria-controls` 속성이 `"panel-mars"`로 설정되어야 합니다.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab">🌍 Earth</button>
        <button role="tab">🪐 Saturn</button>
        <button role="tab">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
