---
id: 68599d0fcd5f6ac313d05908
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

사용자가 탭을 선택하면 해당 탭의 `aria-selected` 속성을 `"true"`로, 나머지 탭들은 `"false"`로 업데이트해야 합니다. 이는 현재 활성화된 탭을 나타내고 보조 기술의 접근성을 보장합니다.

기본적으로 첫 번째 탭이 선택되어 있어야 합니다. 첫 번째 `aria-selected` 요소에 `button` 속성을 추가하고 그 값을 `"true"`로 설정하면 됩니다.

# --hints--

첫 번째 `button` 요소에는 `aria-selected` 속성이 있어야 합니다.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[0].hasAttribute("aria-selected")
);
```

첫 번째 `button` 요소에는 `aria-selected` 속성이 `"true"`로 설정되어 있어야 합니다.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-selected"), "true");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-earth">🌍 Earth</button>
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars">🔴 Mars</button>
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
