---
id: 685db0bd1280154ae85bcf9f
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 HTML, CSS, JavaScript를 사용해 접근 가능한 행성 탭리스트를 만드세요. HTML 보일러플레이트와 CSS는 제공되었습니다.

처음 몇 단계에서는 탭리스트의 나머지 HTML을 작성할 것입니다.

`div` 요소를 만들고 클래스는 `"tabs"`로 설정하세요. 그 `div` 요소 안에 `h2` 요소를 추가하고 `id` 속성은 `"tabs-title"`로 설정하며 텍스트 내용은 `Planets`로 하세요.

# --hints--

`div` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("div"));
```

`div` 요소의 클래스는 `"tabs"`로 설정하세요.

```js
assert.exists(document.querySelector("div.tabs"));
```

`div` 요소 안에 `h2` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector(".tabs h2"));
```

`h2` 요소에 `id` 속성이 `"tabs-title"`로 설정되어야 합니다.

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

`h2` 요소의 텍스트 내용은 `Planets`이어야 합니다.

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
