---
id: 689f7996426c1534fa5bea66
title: Erstellen Sie einen Golf-Score-Übersetzer
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

Im Golfspiel hat jedes Loch ein `par`, was die durchschnittliche Anzahl von `strokes` bedeutet, die ein Golfer voraussichtlich benötigt, um den Ball ins Loch zu bringen und das Spiel zu beenden. Je nachdem, wie weit Ihre `par` über oder unter `strokes` liegen, gibt es einen anderen Spitznamen.

In diesem Labor schreiben Sie eine Funktion, die die `par` und `strokes` in ihren Spitznamen übersetzt.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und lassen Sie alle Tests erfolgreich durchlaufen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `golfScore` erstellen.
1. `golfScore` sollte zwei numerische Argumente annehmen, nämlich den Par für den Kurs und die Anzahl der gemachten Striche.
1. `golfScore` sollte einen String zurückgeben.
1. `golfScore` sollte `"Hole-in-one!"` zurückgeben, wenn `strokes` `1` ist.
1. `golfScore` sollte `"Eagle"` zurückgeben, wenn `strokes` kleiner oder gleich `par` minus `2` ist.
1. `golfScore` sollte `"Birdie"` zurückgeben, wenn `strokes` gleich `par` minus `1` ist.
1. `golfScore` sollte `"Par"` zurückgeben, wenn `strokes` gleich `par` ist.
1. `golfScore` sollte `"Bogey"` zurückgeben, wenn `strokes` gleich `par` plus `1` ist.
1. `golfScore` sollte `"Double Bogey"` zurückgeben, wenn `strokes` gleich `par` plus `2` ist.
1. `golfScore` sollte `"Go Home!"` zurückgeben, wenn `strokes` größer oder gleich `par` plus `3` ist.

# --hints--

Sie sollten eine Funktion mit dem Namen `golfScore` erstellen.

```js
assert.isFunction(golfScore);
```

`golfScore` sollte zwei Parameter annehmen.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` sollte einen String zurückgeben.

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)` sollte den String `Hole-in-one!` zurückgeben.

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` sollte den String `Hole-in-one!` zurückgeben.

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` sollte den String `Hole-in-one!` zurückgeben.

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` sollte den String `Hole-in-one!` zurückgeben.

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` sollte den String `Eagle` zurückgeben.

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` sollte den String `Eagle` zurückgeben.

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` sollte den String `Birdie` zurückgeben.

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` sollte den String `Birdie` zurückgeben.

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` sollte den String `Birdie` zurückgeben.

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` sollte den String `Par` zurückgeben.

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` sollte den String `Par` zurückgeben.

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` sollte den String `Par` zurückgeben.

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` sollte den String `Bogey` zurückgeben.

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` sollte den String `Bogey` zurückgeben.

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` sollte den String `Bogey` zurückgeben.

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` sollte den String `Double Bogey` zurückgeben.

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` sollte den String `Double Bogey` zurückgeben.

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` sollte den String `Double Bogey` zurückgeben.

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` sollte den String `Go Home!` zurückgeben.

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` sollte den String `Go Home!` zurückgeben.

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` sollte den String `Go Home!` zurückgeben.

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
