---
id: a5229172f011153519423690
title: 建置一個奇數費波那契數列總和計算器
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

在本實驗中，你將建置一個奇數 Fibonacci 總和計算器，用來計算所有小於或相等於給定正整數的奇數 Fibonacci 數字的總和。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事**

1. 你應該有一個 `sumFibs` 函式（程式），它接受一個數字作為引數。
1. `sumFibs` 函式應該傳回所有小於或等於給定數字的奇數 Fibonacci 數字的總和。
1. 費波那契數列從 `0` 和 `1` 開始，每個後續的數字都是前兩個數字的總和。
1. 只有奇數的 Fibonacci 數字應該被加到總和中。

# --hints--

你應該有一個 `sumFibs` 函式（程式）。

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` 應該傳回一個數字。

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` 應該傳回 `1785`。

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` 應該傳回 `4613732`。

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` 應該傳回 `5`。

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` 應該傳回 `60696`。

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` 應該傳回 `135721`。

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
