---
id: 6814d8e1516e86b171929de4
title: "تحدٍّ 1: توازن الحروف المتحركة"
challengeType: 29
dashedName: challenge-1
---

# --description--

اعطِ سلسلة نصية، حدّد ما إذا كان عدد الحروف المتحركة في النصف الأول من السلسلة يساوي عدد الحروف المتحركة في النصف الثاني.

- يمكن أن تحتوي السلسلة النصية على أي حروف.
- تُعتبر الحروف `a`، `e`، `i`، `o`، و `u`، سواء كانت بأحرف كبيرة أو أحرف صغيرة، حروفًا متحركة.
- إذا كان عدد الحروف في السلسلة فرديًا، تجاهل الحرف الأوسط.

# --hints--

`is_balanced("racecar")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("racecar"), True)`)
}})
```

`is_balanced("Lorem Ipsum")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("Lorem Ipsum"), True)`)
}})
```

`is_balanced("Kitty Ipsum")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("Kitty Ipsum"), False)`)
}})
```

`is_balanced("string")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("string"), False)`)
}})
```

`is_balanced(" ")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced(" "), True)`)
}})
```

`is_balanced("abcdefghijklmnopqrstuvwxyz")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("abcdefghijklmnopqrstuvwxyz"), False)`)
}})
```

`is_balanced("123A#b!E&*456-o.U")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("123A#b!E&*456-o.U"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_balanced(s):

    return s
```

# --solutions--

```py
def is_balanced(s):
    vowels = set("aeiouAEIOU")
    n = len(s)
    half = n // 2

    first_half = s[:half]
    second_half = s[-half:]

    def count_vowels(sub):
        return sum(1 for char in sub if char in vowels)

    return count_vowels(first_half) == count_vowels(second_half)
```
