---
id: 681cb1afdab50c87ddb2e516
title: "تحدٍّ 5: نص Jbelmud"
challengeType: 29
dashedName: challenge-5
---

# --description--

بالنظر إلى سلسلة نصية، أعد نسخة مشوشة من تلك السلسلة حيث تُحوّل كل كلمة باستخدام القيود التالية:

- تبقى الحروف الأولى والأخيرة من الكلمات في مكانها
- تُرتب جميع الحروف بين الحرف الأول والأخير أبجديًا
- لن تحتوي السلاسل النصية المُدخلة على علامات ترقيم، وستكون كلها أحرف صغيرة.

# --hints--

`jbelmu("hello world")` يجب أن تُعيد `"hello wlord"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("hello world"), "hello wlord")`)
}})
```

`jbelmu("i love jumbled text")` يجب أن تُعيد `"i love jbelmud text"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("i love jumbled text"), "i love jbelmud text")`)
}})
```

`jbelmu("freecodecamp is my favorite place to learn to code")` يجب أن تُعيد `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe")`)
}})
```

`jbelmu("the quick brown fox jumps over the lazy dog")` يجب أن تُعيد `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def jbelmu(text):

    return text
```

# --solutions--

```py
def jbelmu(text):
    words = text.split()
    jumbled = []

    for word in words:
        if len(word) <= 3:
            jumbled.append(word)
        else:
            first = word[0]
            last = word[-1]
            middle = ''.join(sorted(word[1:-1]))
            jumbled.append(first + middle + last)
    
    return ' '.join(jumbled)
```
