---
id: 681cb1afdab50c87ddb2e517
title: "تحدٍّ 6: مدقق الكلمات المتشابهة"
challengeType: 29
dashedName: challenge-6
---

# --description--

معطيتان سلسلتا نصية، حدّد إذا كانتا كلمات متشابهة (تحتويان على نفس الأحرف بأي ترتيب).

- تجاهل حالة الأحرف والمسافات البيضاء.

# --hints--

يجب أن تُرجع الدالة `are_anagrams("listen", "silent")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("listen", "silent"), True)`)
}})
```

يجب أن تُرجع الدالة `are_anagrams("School master", "The classroom")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("School master", "The classroom"), True)`)
}})
```

يجب أن تُرجع الدالة `are_anagrams("A gentleman", "Elegant man")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("A gentleman", "Elegant man"), True)`)
}})
```

يجب أن تُرجع الدالة `are_anagrams("Hello", "World")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("Hello", "World"), False)`)
}})
```

يجب أن تُرجع الدالة `are_anagrams("apple", "banana")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("apple", "banana"), False)`)
}})
```

يجب أن تُرجع الدالة `are_anagrams("cat", "dog")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("cat", "dog"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def are_anagrams(str1, str2):

    return str1
```

# --solutions--

```py
def are_anagrams(str1, str2):
    def clean(s):
        return sorted(s.replace(" ", "").lower())
    
    return clean(str1) == clean(str2)
```
