---
id: 681cb1b0dab50c87ddb2e518
title: "تحدٍّ 7: المجموع المستهدف"
challengeType: 29
dashedName: challenge-7
---

# --description--

معطى مصفوفة من الأرقام وعدد صحيح مستهدف، ابحث عن رقمين فريدين في المصفوفة مجموعهما يساوي القيمة المستهدفة. أرجع مصفوفة تحتوي على مؤشرات هذين الرقمين، أو `"Target not found"` إذا لم يوجد رقمان مجموعهما يساوي الهدف.

- يجب أن تحتوي المصفوفة المرجعة على المؤشرات بترتيب تصاعدي.

# --hints--

`find_target([2, 7, 11, 15], 9)` يجب أن تُرجع `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([2, 7, 11, 15], 9), [0, 1])`)
}})
```

`find_target([3, 2, 4, 5], 6)` يجب أن تُرجع `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([3, 2, 4, 5], 6), [1, 2])`)
}})
```

`find_target([1, 3, 5, 6, 7, 8], 15)` يجب أن تُرجع `[4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 6, 7, 8], 15), [4, 5])`)
}})
```

`find_target([1, 3, 5, 7], 14)` يجب أن تُرجع `'Target not found'`.

```js

({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 7], 14), "Target not found")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_target(arr, target):

    return arr
```

# --solutions--

```py
def find_target(arr, target):
    for i in range(len(arr)):
        for j in range(i + 1, len(arr)):
            if arr[i] + arr[j] == target:
                return [i, j]
    return 'Target not found'
```
