---
id: 681cb1a2dab50c87ddb2e514
title: "تحدٍّ 3: متتالية فيبوناتشي"
challengeType: 29
dashedName: challenge-3
---

# --description--

متتالية فيبوناتشي هي سلسلة أعداد يكون كل عدد فيها مجموع العددين السابقين له. عند البدء بـ `0` و `1`، تكون أول 10 أعداد في المتتالية هي `0`، `1`، `1`، `2`، `3`، `5`، `8`، `13`، `21`، `34`.

معطى مصفوفة تحتوي على أول رقمين من متتالية فيبوناتشي، وعدد صحيح يمثل طول المتتالية، أرجع مصفوفة تحتوي على المتتالية بالطول المعطى.

- يجب أن تتعامل دالتك مع متتاليات بأي طول أكبر من أو يساوي صفر.
- إذا كان الطول صفرًا، أرجع مصفوفة فارغة.
- لاحظ أن الأعداد الابتدائية جزء من المتتالية.

# --hints--


`fibonacci_sequence([0, 1], 20)` يجب أن تُرجع `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181])`)
}})
```

`fibonacci_sequence([21, 32], 1)` يجب أن تُرجع `[21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([21, 32], 1), [21])`)
}})
```

`fibonacci_sequence([0, 1], 0)` يجب أن تُرجع `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 0), [])`)
}})
```

`fibonacci_sequence([10, 20], 2)` يجب أن تُرجع `[10, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([10, 20], 2), [10, 20])`)
}})
```

`fibonacci_sequence([123456789, 987654321], 5)` يجب أن تُرجع `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541])`)
}})
```

# --seed--

## --seed-contents--

```py
def fibonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def fibonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2]
        sequence.append(next_value)

    return sequence
```
