---
id: 681cb1b0dab50c87ddb2e519
title: "التحدي 8: مضروب"
challengeType: 29
dashedName: challenge-8
---

# --description--

أعطِ عددًا صحيحًا من صفر إلى 20، أرجع مضروب ذلك العدد. المضروب هو حاصل ضرب كل الأعداد بين 1 والعدد المعطى.

- مضروب الصفر هو 1.

# --hints--

`factorial(0)` يجب أن تُرجع `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(0), 1)`)
}})
```

`factorial(5)` يجب أن تُرجع `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(5), 120)`)
}})
```

`factorial(20)` يجب أن تُرجع `2432902008176640000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(20), 2432902008176640000)`)
}})
```

# --seed--

## --seed-contents--

```py
def factorial(n):

    return n
```

# --solutions--

```py
def factorial(n):
    result = 1
    for i in range(1, n + 1):
        result *= i
    return result
```
