---
id: 681cb1b0dab50c87ddb2e51b
title: "تحدٍّ 10: 3 ضربات"
challengeType: 29
dashedName: challenge-10
---

# --description--

بالنظر إلى عدد صحيح بين 1 و 10,000، أرجع عدد الأعداد من 1 حتى ذلك العدد التي يحتوي مربعها على رقم 3 واحد على الأقل.

# --hints--

يجب أن تُرجع الدالة `squares_with_three(1)` القيمة `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(1), 0)`)
}})
```

يجب أن تُرجع الدالة `squares_with_three(10)` القيمة `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(10), 1)`)
}})
```

يجب أن تُرجع الدالة `squares_with_three(100)` القيمة `19`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(100), 19)`)
}})
```

يجب أن تُرجع الدالة `squares_with_three(1000)` القيمة `326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(1000), 326)`)
}})
```

يجب أن تُرجع الدالة `squares_with_three(10000)` القيمة `4531`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(10000), 4531)`)
}})
```

# --seed--

## --seed-contents--

```py
def squares_with_three(n):

    return n
```

# --solutions--

```py
def squares_with_three(n):
    count = 0
    for i in range(1, n + 1):
        square = i * i
        if '3' in str(square):
            count += 1
    return count
```
