---
id: 681cb1b0dab50c87ddb2e51a
title: "تحدٍّ 9: مجموع المربعات"
challengeType: 29
dashedName: challenge-9
---

# --description--

أعطِ عددًا صحيحًا موجبًا يصل إلى 1000، أرجع مجموع جميع الأعداد الصحيحة المربعة من 1 حتى ذلك العدد.

# --hints--

`sum_of_squares(5)` يجب أن تُرجع `55`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(5), 55)`)
}})
```

`sum_of_squares(10)` يجب أن تُرجع `385`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(10), 385)`)
}})
```

`sum_of_squares(25)` يجب أن تُرجع `5525`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(25), 5525)`)
}})
```

`sum_of_squares(500)` يجب أن تُرجع `41791750`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(500), 41791750)`)
}})
```

`sum_of_squares(1000)` يجب أن تُرجع `333833500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(1000), 333833500)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_of_squares(n):

    return n
```

# --solutions--

```py
def sum_of_squares(n):
    sum = ((n)*(n+1)*(2*n+1))//6
    return sum
```
