---
id: 6821ebce237de8297eaee790
title: "تحدٍّ 14: معركة الشخصيات"
challengeType: 29
dashedName: challenge-14
---

# --description--

معطيتان سلسلتا نص تمثلان جيشك وجيش الخصم، يقاتل كل حرف من جيشك الحرف في نفس الموضع من جيش الخصم وفق القواعد التالية:

- الأحرف `a-z` لها قوة من `1-26`، على التوالي.
- الأحرف `A-Z` لها قوة من `27-52`، على التوالي.
- الأرقام `0-9` لها قوة تساوي قيمتها الظاهرة.
- جميع الأحرف الأخرى قيمتها صفر.
- يمكن لكل حرف أن يخوض معركة واحدة فقط.

في كل معركة، ينتصر الحرف الأقوى. الجيش الذي يحقق انتصارات أكثر يفوز بالحرب. أرجع القيم التالية:

- `"Opponent retreated"` إذا كان جيشك يحتوي على أحرف أكثر من جيش الخصم.
- `"We retreated"` إذا كان جيش الخصم يحتوي على أحرف أكثر من جيشك.
- `"We won"` إذا فاز جيشك بمعارك أكثر.
- `"We lost"` إذا فاز جيش الخصم بمعارك أكثر.
- `"It was a tie"` إذا تعادل الجيشان في عدد المعارك التي فاز بها كل منهما.

# --hints--

`battle("Hello", "World")` يجب أن تُرجع `"We lost"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` يجب أن تُرجع `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` يجب أن تُرجع `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` يجب أن تُرجع `"Opponent retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` يجب أن تُرجع `"We retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` يجب أن تُرجع `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` يجب أن تُرجع `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```
