---
id: 68216eb60f957572e7c340c4
title: "تحدٍّ 11: وتيرة الميل"
challengeType: 29
dashedName: challenge-11
---

# --description--

بما أن لديك عدد الأميال التي جُريت، ووقتًا بصيغة `"MM:SS"` (دقائق:ثوانٍ) استغرقته لقطع تلك الأميال، أرجع سلسلة نصية تمثل متوسط الوقت الذي استغرقه قطع كل ميل بصيغة `"MM:SS"`.

- أضف أصفارًا بادئة عند الحاجة.

# --hints--

`mile_pace(3, "24:00")` يجب أن تُرجع `"08:00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(3, "24:00"), "08:00")`)
}})
```

`mile_pace(1, "06:45")` يجب أن تُرجع `"06:45"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(1, "06:45"), "06:45")`)
}})
```

`mile_pace(2, "07:00")` يجب أن تُرجع `"03:30"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(2, "07:00"), "03:30")`)
}})
```

`mile_pace(26.2, "120:35")` يجب أن تُرجع `"04:36"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(26.2, "120:35"), "04:36")`)
}})
```

# --seed--

## --seed-contents--

```py
def mile_pace(miles, duration):

    return miles
```

# --solutions--

```py
def mile_pace(miles, duration):
    minutes, seconds = map(int, duration.split(":"))
    total_seconds = minutes * 60 + seconds
    avg_seconds_per_mile = total_seconds / miles

    avg_minutes = int(avg_seconds_per_mile // 60)
    avg_seconds = round(avg_seconds_per_mile % 60)

    return f"{avg_minutes:02d}:{avg_seconds:02d}"
```
