---
id: 6821ebda237de8297eaee792
title: "تحدٍّ 16: عكس الأقواس العادية"
challengeType: 29
dashedName: challenge-16
---

# --description--

أعطِ سلسلة نصية تحتوي على أقواس عادية مسكنة بشكل صحيح، أرجع النسخة المفككة من السلسلة باستخدام القواعد التالية:

- يجب عكس كل الأحرف داخل كل زوج من الأقواس العادية.
- يجب إزالة الأقواس العادية من النتيجة النهائية.
- إذا كانت الأقواس العادية مسكنة، يجب عكس الزوج الداخلي أولًا، ثم تضمين نتيجته في عكس الزوج الخارجي.
- افترض أن كل الأقواس العادية متوازنة بشكل متساوٍ ومسكنة بشكل صحيح.

# --hints--

`decode("(f(b(dc)e)a)")` يجب أن تُرجع `"abcdef"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` يجب أن تُرجع `"Can you read this?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` يجب أن تُرجع `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```
