---
id: 68216ef80f957572e7c340c5
title: "تحدٍّ 12: فك شفرة الرسالة"
challengeType: 29
dashedName: challenge-12
---

# --description--

أعطِ سلسلة نصية لرسالة سرية، وعددًا صحيحًا يمثل عدد الحروف التي استُخدمت لتحويل الرسالة وترميزها، أرجع السلسلة النصية المفكوكة.

- العدد الموجب يعني أن الرسالة تحولت للأمام في الأبجدية.
- العدد السالب يعني أن الرسالة تحولت للخلف في الأبجدية.
- حالة الحروف مهمة، يجب أن تحتفظ الأحرف المفكوكة بحالتها كما في الأحرف المشفرة.
- لا تُفكك الأحرف غير الأبجدية.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` يجب أن تُرجع `"This is a secret message."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.")`)
}})
```

`decode("Byffi Qilfx!", 20)` يجب أن تُرجع `"Hello World!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Byffi Qilfx!", 20), "Hello World!")`)
}})
```

`decode("Zqd xnt njzx?", -1)` يجب أن تُرجع `"Are you okay?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Zqd xnt njzx?", -1), "Are you okay?")`)
}})
```

`decode("oannLxmnLjvy", 9)` يجب أن تُرجع `"freeCodeCamp"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("oannLxmnLjvy", 9), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message, shift):

    return message
```

# --solutions--

```py
def decode(message, shift):
    decoded_message = []
    for char in message:
        if char.isalpha():
            base = ord('a') if char.islower() else ord('A')
            new_char = chr((ord(char) - base - shift) % 26 + base)
            decoded_message.append(new_char)
        else:
            decoded_message.append(char)
    return ''.join(decoded_message)
```
