---
id: 6821ebd4237de8297eaee791
title: "تحدٍّ 15: تنسيق سنام الجمل"
challengeType: 29
dashedName: challenge-15
---

# --description--

بالنظر إلى سلسلة نصية، أعد إصدارها بتنسيق سنام الجمل باستخدام القواعد التالية:

- الكلمات في سلسلة النص المعطاة مفصولة بحرف واحد أو أكثر من المجموعة التالية: مسافة (` `)، شرطة (`-`)، أو شرطة سفلية (`_`). اعتبر أي متتالية من هذه كفاصل للكلمات.
- يجب أن تكون الكلمة الأولى كلها بأحرف صغيرة.
- يجب أن تبدأ كل كلمة لاحقة بحرف كبير، مع بقية الحروف بأحرف صغيرة.
- يجب إزالة كل المسافات والفواصل.

# --hints--

`to_camel_case("hello world")` يجب أن تُعيد `"helloWorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("hello world"), "helloWorld")`)
}})
```

`to_camel_case("HELLO WORLD")` يجب أن تُعيد `"helloWorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("HELLO WORLD"), "helloWorld")`)
}})
```

`to_camel_case("secret agent-X")` يجب أن تُعيد `"secretAgentX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("secret agent-X"), "secretAgentX")`)
}})
```

`to_camel_case("FREE cODE cAMP")` يجب أن تُعيد `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("FREE cODE cAMP"), "freeCodeCamp")`)
}})
```

`to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` يجب أن تُعيد `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_camel_case(s):

    return s
```

# --solutions--

```py
import re
def to_camel_case(s):
    words = re.split(r'[_\- ]+', s)

    camel = [
        words[0].lower() if words else ''
    ] + [
        word.capitalize() for word in words[1:]
    ]

    return ''.join(camel)
```
