---
id: 6821ebee237de8297eaee796
title: "تحدٍّ 20: العناصر المكررة في المصفوفة"
challengeType: 29
dashedName: challenge-20
---

# --description--

بالنظر إلى مصفوفة من الأعداد الصحيحة، أعد مصفوفة من الأعداد الصحيحة التي تظهر أكثر من مرة في المصفوفة الأصلية، مرتبة تصاعديًا. إذا لم تظهر أي قيمة أكثر من مرة، أعد مصفوفة فارغة.

- تضمّن نسخة واحدة فقط من كل قيمة في المصفوفة المعادة.

# --hints--

`find_duplicates([1, 2, 3, 4, 5])` يجب أن تُعيد `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([1, 2, 3, 4, 5]), [])`)
}})
```

`find_duplicates([1, 2, 3, 4, 1, 2])` يجب أن تُعيد `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([1, 2, 3, 4, 1, 2]), [1, 2])`)
}})
```

`find_duplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` يجب أن تُعيد `[-6, 0, 2, 4, 5, 23]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_duplicates(arr):

    return arr
```

# --solutions--

```py
def find_duplicates(arr):
    seen = set()
    duplicates = set()

    for num in arr:
        if num in seen:
            duplicates.add(num)
        else:
            seen.add(num)

    return sorted(duplicates)
```
