---
id: 6821ebf8237de8297eaee798
title: "تحدٍّ 22: متتالية Tribonacci"
challengeType: 29
dashedName: challenge-22
---

# --description--

متتالية Tribonacci هي سلسلة من الأعداد حيث يكون كل عدد هو مجموع الأعداد الثلاثة السابقة له. عند البدء بـ `0`، `0` و `1`، تكون أول 10 أعداد في المتتالية هي `0`، `0`، `1`، `1`، `2`، `4`، `7`، `13`، `24`، `44`.

معطى مصفوفة تحتوي على أول ثلاثة أعداد من متتالية Tribonacci، وعدد صحيح يمثل طول المتتالية، أرجع مصفوفة تحتوي على المتتالية بالطول المعطى.

- يجب أن تتعامل دالتك مع متتاليات بأي طول أكبر من أو يساوي صفر.
- إذا كان الطول صفرًا، أرجع مصفوفة فارغة.
- لاحظ أن الأعداد الابتدائية جزء من المتتالية.

# --hints--

`tribonacci_sequence([0, 0, 1], 20)` يجب أن تُرجع `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513])`)
}})
```

`tribonacci_sequence([21, 32, 43], 1)` يجب أن تُرجع `[21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([21, 32, 43], 1), [21])`)
}})
```

`tribonacci_sequence([0, 0, 1], 0)` يجب أن تُرجع `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 0), [])`)
}})
```

`tribonacci_sequence([10, 20, 30], 2)` يجب أن تُرجع `[10, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 2), [10, 20])`)
}})
```

`tribonacci_sequence([10, 20, 30], 3)` يجب أن تُرجع `[10, 20, 30]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 3), [10, 20, 30])`)
}})
```

`tribonacci_sequence([123, 456, 789], 8)` يجب أن تُرجع `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134])`)
}})
```

# --seed--

## --seed-contents--

```py
def tribonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def tribonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return [start_sequence[0], start_sequence[1]]
    if length == 3:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2] + sequence[-3]
        sequence.append(next_value)

    return sequence
```
