---
id: 6821ebc9237de8297eaee78f
title: "تحدٍّ 13: عدد أولي غير طبيعي"
challengeType: 29
dashedName: challenge-13
---

# --description--

بما أن العدد صحيح، حدّد إذا كان ذلك العدد عددًا أوليًا أو عددًا أوليًا سالبًا.

- العدد الأولي هو عدد صحيح موجب أكبر من 1 لا يقبل القسمة إلا على 1 ونفسه.
- العدد الأولي السالب هو النسخة السالبة من العدد الأولي الموجب.
- لا يُعتبر `1` و `0` أعدادًا أولية.

# --hints--

يجب أن تُعيد الدالة `is_unnatural_prime(1)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(1), False)`)
}})
```

يجب أن تُعيد الدالة `is_unnatural_prime(-1)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-1), False)`)
}})
```

يجب أن تُعيد الدالة `is_unnatural_prime(19)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(19), True)`)
}})
```

يجب أن تُعيد الدالة `is_unnatural_prime(-23)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-23), True)`)
}})
```

يجب أن تُعيد الدالة `is_unnatural_prime(0)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(0), False)`)
}})
```

يجب أن تُعيد الدالة `is_unnatural_prime(97)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(97), True)`)
}})
```

يجب أن تُعيد الدالة `is_unnatural_prime(-61)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-61), True)`)
}})
```

يجب أن تُعيد الدالة `is_unnatural_prime(99)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(99), False)`)
}})
```

يجب أن تُعيد الدالة `is_unnatural_prime(-44)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-44), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_unnatural_prime(n):

    return n
```

# --solutions--

```py
def is_unnatural_prime(n):
    abs_n = abs(n)

    if abs_n <= 1:
        return False

    for i in range(2, int(abs_n ** 0.5) + 1):
        if abs_n % i == 0:
            return False

    return True
```
