---
id: 6821ebe4237de8297eaee794
title: "تحدٍّ 18: الثاني الأفضل"
challengeType: 29
dashedName: challenge-18
---

# --description--

معطاة مصفوفة من الأعداد الصحيحة تمثل سعر حواسيب محمولة مختلفة، وعدد صحيح يمثل ميزانيتك، أرجع:

1. الحاسوب المحمول الثاني الأغلى إذا كان ضمن ميزانيتك، أو
2. الحاسوب المحمول الأغلى الذي يقع ضمن ميزانيتك، أو
3. `0` إذا لم يكن هناك حواسيب محمولة ضمن ميزانيتك.

- يجب تجاهل الأسعار المكررة.

# --hints--

`get_laptop_cost([1500, 2000, 1800, 1400], 1900)` يجب أن تُرجع `1800`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900)` يجب أن تُرجع `1800`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 2200)` يجب أن تُرجع `1899`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 2200), 1899)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 1000)` يجب أن تُرجع `0`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 1000), 0)`)
}})
```

`get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` يجب أن تُرجع `1400`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_laptop_cost(laptops, budget):

    return laptops
```

# --solutions--

```py
def get_laptop_cost(laptops, budget):
    laptops = sorted(set(laptops), reverse=True)

    if len(laptops) >= 2 and budget >= laptops[1]:
        return laptops[1]
    if not laptops or budget < laptops[-1]:
        return 0

    for price in laptops[2:]:
        if budget >= price:
            return price

    return 0
```
