---
id: 6821ebdf237de8297eaee793
title: "تحدٍّ 17: ترتيب العمليات غير المرتب"
challengeType: 29
dashedName: challenge-17
---

# --description--

معطى مصفوفة من الأعداد الصحيحة ومصفوفة من عوامل التشغيل كسلاسل نصية، طبّق العمليات على الأعداد بالتتابع من اليسار إلى اليمين. كرر العمليات حسب الحاجة حتى تُستخدم كل الأعداد. أرجع النتيجة النهائية.

على سبيل المثال، معطى `[1, 2, 3, 4, 5]` و`['+', '*']`، أرجع نتيجة تقييم `1 + 2 * 3 + 4 * 5` من اليسار إلى اليمين متجاهلًا ترتيب العمليات القياسي.

- عوامل التشغيل الصالحة هي `+`، `-`، `*`، `/`، و`%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` يجب أن تُرجع `3`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3)`)
}})
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` يجب أن تُرجع `38`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38)`)
}})
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` يجب أن تُرجع `60`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60)`)
}})
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` يجب أن تُرجع `30`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30)`)
}})
```

`evaluate([33, 11, 29, 13], ['/', '-'])` يجب أن تُرجع `-2`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([33, 11, 29, 13], ['/', '-']), -2)`)
}})
```

# --seed--

## --seed-contents--

```py
def evaluate(numbers, operators):

    return numbers
```

# --solutions--

```py
def do_math(a, b, operator):
    if operator == '+':
        return a + b
    elif operator == '-':
        return a - b
    elif operator == '*':
        return a * b
    elif operator == '/':
        return a / b
    else:
        return a % b

def evaluate(numbers, operators):
    total = numbers[0]

    for i in range(1, len(numbers)):
        operator = operators[(i - 1) % len(operators)]
        total = do_math(total, numbers[i], operator)

    return total
```
