---
id: 6821ebfd237de8297eaee799
title: "تحدٍّ 23: تحويل RGB إلى Hex"
challengeType: 29
dashedName: challenge-23
---

# --description--

بالنظر إلى سلسلة لون CSS `rgb(r, g, b)`، أرجع ما يعادلها بالنظام الست عشري.

إليك بعض الأمثلة على المخرجات لمُدخَل معين:

| المُدخَل   | المخرجات   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- اجعل كل الأحرف أحرف صغيرة.
- أرجع `#` متبوعًا بستة أحرف. لا تستخدم أي قيم مختصرة.

# --hints--

`rgb_to_hex("rgb(255, 255, 255)")` يجب أن تُرجع `"#ffffff"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(255, 255, 255)"), "#ffffff")`)
}})
```

`rgb_to_hex("rgb(1, 11, 111)")` يجب أن تُرجع `"#010b6f"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(1, 11, 111)"), "#010b6f")`)
}})
```

`rgb_to_hex("rgb(173, 216, 230)")` يجب أن تُرجع `"#add8e6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(173, 216, 230)"), "#add8e6")`)
}})
```

`rgb_to_hex("rgb(79, 123, 201)")` يجب أن تُرجع `"#4f7bc9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(79, 123, 201)"), "#4f7bc9")`)
}})
```

# --seed--

## --seed-contents--

```py
def rgb_to_hex(rgb):

    return rgb
```

# --solutions--

```py
def rgb_to_hex(rgb):
    import re
    match = re.findall(r'\d+', rgb)
    r, g, b = [max(0, min(255, int(x))) for x in match[:3]]
    return f'#{r:02x}{g:02x}{b:02x}'
```
