---
id: 6821ebea237de8297eaee795
title: "تحدٍّ 19: ضوء الشموع"
challengeType: 29
dashedName: challenge-19
---

# --description--

معطى عدد صحيح يمثل عدد الشموع التي تبدأ بها، وعدد صحيح يمثل عدد الشموع المحترقة اللازمة لصنع شمعة جديدة، أعد عدد الشموع التي ستستخدمها بعد صنع وحرق أكبر عدد ممكن.

على سبيل المثال، إذا كان لديك 7 شموع ويحتاج صنع شمعة جديدة إلى شمعتين محترقتين:

1. احرق 7 شموع لتحصل على 7 بقايا،
2. أعد تدوير 6 بقايا لصنع 3 شموع جديدة (تبقى شمعة واحدة)،
3. احرق 3 شموع لتحصل على 3 بقايا إضافية (المجموع 4)،
4. أعد تدوير 4 بقايا لصنع 2 شمعة جديدة،
5. احرق 2 شمعة لتحصل على 2 بقايا،
6. أعد تدوير 2 بقايا لصنع شمعة جديدة واحدة،
7. احرق شمعة واحدة.

ستكون قد حرقت 13 شمعة في المجموع في هذا المثال.

# --hints--

يجب أن تُعيد الدالة `burn_candles(7, 2)` القيمة `13`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(7, 2), 13)`)
}})
```

يجب أن تُعيد الدالة `burn_candles(10, 5)` القيمة `12`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(10, 5), 12)`)
}})
```

يجب أن تُعيد الدالة `burn_candles(20, 3)` القيمة `29`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(20, 3), 29)`)
}})
```

يجب أن تُعيد الدالة `burn_candles(17, 4)` القيمة `22`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(17, 4), 22)`)
}})
```

يجب أن تُعيد الدالة `burn_candles(2345, 3)` القيمة `3517`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(2345, 3), 3517)`)
}})
```

# --seed--

## --seed-contents--

```py
def burn_candles(candles, leftovers_needed):

    return candles
```

# --solutions--

```py
def burn_candles(candles, leftovers_needed):
    total_burned = 0
    unused_leftovers = 0

    while candles > 0:
        total_burned += candles
        leftovers = candles + unused_leftovers
        candles = leftovers // leftovers_needed
        unused_leftovers = leftovers % leftovers_needed

    return total_burned
```
