---
id: 68adce01c0e1144d0a90295e
title: "تحدٍّ 29: منشئ الاختصارات"
challengeType: 29
dashedName: challenge-29
---

# --description--

بالنظر إلى سلسلة نصية تحتوي على كلمة واحدة أو أكثر، أرجع اختصارًا للكلمات باستخدام القيود التالية:

- يجب أن يتكون الاختصار من الحرف الأول من كل كلمة بحرف كبير، ما لم يُذكر خلاف ذلك.
- يجب أن يتجاهل الاختصار الحرف الأول من هذه الكلمات إلا إذا كانت الكلمة الأولى في السلسلة النصية المعطاة: `a`، `for`، `an`، `and`، `by`، و`of`.
- يجب إرجاع حروف الاختصار بالترتيب الذي تظهر به.
- يجب ألا يحتوي الاختصار على أي فراغات.

# --hints--

`build_acronym("Search Engine Optimization")` يجب أن تُرجع `"SEO"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Search Engine Optimization"), "SEO")`)
}})
```

`build_acronym("Frequently Asked Questions")` يجب أن تُرجع `"FAQ"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Frequently Asked Questions"), "FAQ")`)
}})
```

`build_acronym("National Aeronautics and Space Administration")` يجب أن تُرجع `"NASA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("National Aeronautics and Space Administration"), "NASA")`)
}})
```

`build_acronym("Federal Bureau of Investigation")` يجب أن تُرجع `"FBI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Federal Bureau of Investigation"), "FBI")`)
}})
```

`build_acronym("For your information")` يجب أن تُرجع `"FYI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("For your information"), "FYI")`)
}})
```

`build_acronym("By the way")` يجب أن تُرجع `"BTW"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("By the way"), "BTW")`)
}})
```

`build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` يجب أن تُرجع `"AUHWPOTIMSH"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH")`)
}})
```

# --seed--

## --seed-contents--

```py
def build_acronym(s):

    return s
```

# --solutions--

```py
def build_acronym(s):
    small_words = {"a", "for", "an", "and", "by", "of"}
    words = s.split()
    acronym = ""

    for i, word in enumerate(words):
        if i == 0 or word.lower() not in small_words:
            acronym += word[0].upper()

    return acronym
```
