---
id: 68adce01c0e1144d0a90295c
title: "تحدٍّ 28: محلل الأعداد الرومانية"
challengeType: 29
dashedName: challenge-28
---

# --description--

أعطِ سلسلة نصية تمثّل عددًا رومانيًا، أرجع قيمته الصحيحة.

تتكون الأعداد الرومانية من الرموز والقيم التالية:

| الرمز | القيمة |
|--------|-------|
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

- تُقرأ الأعداد من اليسار إلى اليمين. إذا ظهر عدد أصغر قبل عدد أكبر، تُطرح قيمته. وإلا، تُجمع القيم.

# --hints--

يجب أن تُرجع الدالة `parse_roman_numeral("III")` القيمة `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("III"), 3)`)
}})
```

يجب أن تُرجع الدالة `parse_roman_numeral("IV")` القيمة `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("IV"), 4)`)
}})
```

يجب أن تُرجع الدالة `parse_roman_numeral("XXVI")` القيمة `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("XXVI"), 26)`)
}})
```

يجب أن تُرجع الدالة `parse_roman_numeral("XCIX")` القيمة `99`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("XCIX"), 99)`)
}})
```

يجب أن تُرجع الدالة `parse_roman_numeral("CDLX")` القيمة `460`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("CDLX"), 460)`)
}})
```

يجب أن تُرجع الدالة `parse_roman_numeral("DIV")` القيمة `504`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("DIV"), 504)`)
}})
```

يجب أن تُرجع الدالة `parse_roman_numeral("MMXXV")` القيمة `2025`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("MMXXV"), 2025)`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_roman_numeral(numeral):

    return numeral
```

# --solutions--

```py
def parse_roman_numeral(numeral):
    roman_map = {
        "I": 1,
        "V": 5,
        "X": 10,
        "L": 50,
        "C": 100,
        "D": 500,
        "M": 1000
    }

    total = 0
    for i in range(len(numeral)):
        current = roman_map[numeral[i]]
        next_val = roman_map[numeral[i + 1]] if i + 1 < len(numeral) else 0

        if current < next_val:
            total -= current
        else:
            total += current

    return total
```
