---
id: 68adce01c0e1144d0a90295a
title: "تحدٍّ 27: تدوير المصفوفة"
challengeType: 29
dashedName: challenge-27
---

# --description--

معطى مصفوفة (مصفوفة من المصفوفات)، دوّر المصفوفة 90 درجة باتجاه عقارب الساعة وأعدها. على سبيل المثال، معطى `[[1, 2], [3, 4]]`، التي تبدو هكذا:

| 1 | 2 |
|---|---|
| 3 | 4 |

يجب أن تُعيد `[[3, 1], [4, 2]]`، التي تبدو هكذا:

| 3 | 1 |
|---|---|
| 4 | 2 |

# --hints--

`rotate([[1]])` يجب أن تُعيد `[[1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1]]), [[1]])`)
}})
```

`rotate([[1, 2], [3, 4]])` يجب أن تُعيد `[[3, 1], [4, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]])`)
}})
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` يجب أن تُعيد `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]])`)
}})
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` يجب أن تُعيد `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def rotate(matrix):

    return matrix
```

# --solutions--

```py
def rotate(matrix):
    n = len(matrix)
    result = [[0] * n for _ in range(n)]

    for i in range(n):
        for j in range(n):
            result[j][n - 1 - i] = matrix[i][j]

    return result
```
