---
id: 68adce01c0e1144d0a902956
title: "تحدٍّ 25: مكرر الحروف المتحركة"
challengeType: 29
dashedName: challenge-25
---

# --description--

أعطِ سلسلة نصية، أرجع إصدارًا جديدًا من السلسلة حيث تتكرر كل حرف متحرك مرة أكثر من الحرف المتحرك السابق الذي واجهته. على سبيل المثال، يجب أن يبقى الحرف المتحرك الأول في الجملة دون تغيير. يجب أن يظهر الحرف المتحرك الثاني مرتين متتاليتين. يجب أن يظهر الحرف المتحرك الثالث ثلاث مرات متتالية، وهكذا.

- الحروف `a`، `e`، `i`، `o`، و `u`، سواء كانت بأحرف كبيرة أو أحرف صغيرة، تُعتبر حروفًا متحركة.
- يجب أن يحتفظ الحرف المتحرك الأصلي بحالته.
- يجب أن تكون الحروف المتحركة المكررة بأحرف صغيرة.
- يجب أن تحتفظ جميع الأحرف غير المتحركة بحالتها الأصلية.

# --hints--

`repeat_vowels("hello world")` يجب أن تُرجع `"helloo wooorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` يجب أن تُرجع `"freeeCooodeeeeCaaaaamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` يجب أن تُرجع `"AEeIiiOoooUuuuu"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` يجب أن تُرجع `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
