---
id: 68adce01c0e1144d0a902958
title: "تحدٍّ 26: مدقق عنوان IPv4"
challengeType: 29
dashedName: challenge-26
---

# --description--

بناءً على سلسلة نصية، حدّد ما إذا كان عنوان IPv4 صالحًا. يتكون عنوان IPv4 الصالح من أربعة أعداد صحيحة مفصولة بنقاط (`.`). يجب أن يحقق كل عدد الشروط التالية:

- أن يكون بين 0 و255 شاملة.
- ألا يحتوي على أصفار بادئة (مثلاً 0 مسموح، لكن 01 غير مسموح).
- يُسمح فقط بالأحرف الرقمية.

# --hints--

`is_valid_ipv4("192.168.1.1")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.1.1"), True)`)
}})
```

`is_valid_ipv4("0.0.0.0")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("0.0.0.0"), True)`)
}})
```

`is_valid_ipv4("255.01.50.111")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("255.01.50.111"), False)`)
}})
```

`is_valid_ipv4("255.00.50.111")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("255.00.50.111"), False)`)
}})
```

`is_valid_ipv4("256.101.50.115")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("256.101.50.115"), False)`)
}})
```

`is_valid_ipv4("192.168.101.")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.101."), False)`)
}})
```

`is_valid_ipv4("192168145213")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192168145213"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_ipv4(ipv4):

    return ipv4
```

# --solutions--

```py
def is_valid_ipv4(ipv4):
    parts = ipv4.split(".")    

    if len(parts) != 4:
        return False

    for part in parts:
        if not part.isdigit():
            return False

        num = int(part)
        if num < 0 or num > 255:
            return False

        if len(part) > 1 and part.startswith("0"):
            return False

    return True
```
