---
id: 68af0687ef34c76c28ffa54b
title: "تحدٍّ 32: عكس الجملة"
challengeType: 29
dashedName: challenge-32
---

# --description--

أعطِ سلسلة نصية من الكلمات، أرجع سلسلة جديدة بالكلمات بترتيب عكسي. على سبيل المثال، يجب أن تكون الكلمة الأولى في نهاية السلسلة المعادة، والكلمة الأخيرة في بداية السلسلة المعادة.

- في السلسلة المعطاة، يمكن أن تفصل الكلمات بمسافة واحدة أو أكثر.
- يجب أن تحتوي السلسلة المعادة على مسافة واحدة فقط بين الكلمات.

# --hints--

`reverse_sentence("world hello")` يجب أن تُرجع `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("world hello"), "hello world")`)
}})
```

`reverse_sentence("push commit git")` يجب أن تُرجع `"git commit push"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("push commit git"), "git commit push")`)
}})
```

`reverse_sentence("npm  install   apt    sudo")` يجب أن تُرجع `"sudo apt install npm"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("npm  install   apt    sudo"), "sudo apt install npm")`)
}})
```

`reverse_sentence("import    default   function  export")` يجب أن تُرجع `"export function default import"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("import    default   function  export"), "export function default import")`)
}})
```

# --seed--

## --seed-contents--

```py
def reverse_sentence(sentence):

    return sentence
```

# --solutions--

```py
def reverse_sentence(sentence):

   return " ".join(sentence.split()[::-1])
```
