---
id: 68af0687ef34c76c28ffa54f
title: "تحدٍّ 34: الأعداد المفقودة"
challengeType: 29
dashedName: challenge-34
---

# --description--

معطاة مصفوفة من الأعداد الصحيحة من 1 إلى `n`، شاملة، أعد مصفوفة بكل الأعداد الصحيحة المفقودة بين 1 و`n` (حيث `n` هو أكبر رقم في المصفوفة المعطاة).

- قد تكون المصفوفة المعطاة غير مرتبة وقد تحتوي على أعداد مكررة.
- يجب أن تكون المصفوفة المعادة مرتبة تصاعديًا.
- إذا لم تكن هناك أعداد صحيحة مفقودة، أعد مصفوفة فارغة.

# --hints--

يجب أن تُعيد الدالة `find_missing_numbers([1, 3, 5])` القيمة `[2, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 3, 5]), [2, 4])`)
}})
```

يجب أن تُعيد الدالة `find_missing_numbers([1, 2, 3, 4, 5])` القيمة `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 2, 3, 4, 5]), [])`)
}})
```

يجب أن تُعيد الدالة `find_missing_numbers([1, 10])` القيمة `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9])`)
}})
```

يجب أن تُعيد الدالة `find_missing_numbers([10, 1, 10, 1, 10, 1])` القيمة `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9])`)
}})
```

يجب أن تُعيد الدالة `find_missing_numbers([3, 1, 4, 1, 5, 9])` القيمة `[2, 6, 7, 8]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([3, 1, 4, 1, 5, 9]), [2, 6, 7, 8])`)
}})
```

يجب أن تُعيد الدالة `find_missing_numbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])` القيمة `[11]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_missing_numbers(arr):

    return arr
```

# --solutions--

```py
def find_missing_numbers(arr):
    if not arr:
        return []
    
    n = max(arr)
    seen = set(arr)
    result = [i for i in range(1, n + 1) if i not in seen]
    return result
```
