---
id: 68b06e589bf2273243814771
title: "تحدٍّ 36: مُعدِّل منظم الحرارة"
challengeType: 29
dashedName: challenge-36
---

# --description--

بالنظر إلى درجة حرارة الغرفة الحالية ودرجة الحرارة المستهدفة، أرجع سلسلة نصية تشير إلى كيفية تعديل درجة حرارة الغرفة بناءً على هذه القيود:

- أرجع `"heat"` إذا كانت درجة الحرارة الحالية أقل من المستهدفة.
- أرجع `"cool"` إذا كانت درجة الحرارة الحالية أعلى من المستهدفة.
- أرجع `"hold"` إذا كانت درجة الحرارة الحالية مساوية للمستهدفة.

# --hints--

`adjust_thermostat(68, 72)` يجب أن تُرجع `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(68, 72), "heat")`)
}})
```

`adjust_thermostat(75, 72)` يجب أن تُرجع `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(75, 72), "cool")`)
}})
```

`adjust_thermostat(72, 72)` يجب أن تُرجع `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 72), "hold")`)
}})
```

`adjust_thermostat(-20.5, -10.1)` يجب أن تُرجع `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(-20.5, -10.1), "heat")`)
}})
```

`adjust_thermostat(100, 99.9)` يجب أن تُرجع `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(100, 99.9), "cool")`)
}})
```

`adjust_thermostat(0.0, 0.0)` يجب أن تُرجع `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(0.0, 0.0), "hold")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(temp, target):

    return temp
```

# --solutions--

```py
def adjust_thermostat(temp, target):
    if temp < target:
        return "heat"
    elif temp > target:
        return "cool"
    else:
        return "hold"
```
