---
id: 68b06e589bf227324381476f
title: "تحدٍّ 35: تكرار الكلمات"
challengeType: 29
dashedName: challenge-35
---

# --description--

أعطِ فقرة، أرجع مصفوفة تحتوي على ثلاث كلمات تتكرر أكثر من غيرها.

- الكلمات في الفقرة ستكون مفصولة بمسافات.
- تجاهل حالة الأحرف في الفقرة المعطاة. على سبيل المثال، اعتبر `Hello` و `hello` نفس الكلمة.
- تجاهل علامات الترقيم في الفقرة المعطاة. علامات الترقيم هي الفواصل (`,`)، والنقاط (`.`)، وعلامات التعجب (`!`).
- يجب أن تحتوي المصفوفة المرجعة على كلمات بأحرف صغيرة فقط.
- يجب أن تكون المصفوفة المرجعة مرتبة ترتيبًا تنازليًا بحيث تكون الكلمة الأكثر تكرارًا أولاً.

# --hints--

يجب أن تُرجع `get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` المصفوفة `["coding", "python", "in"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

يجب أن تُرجع `get_words("I like coding. I like testing. I love debugging!")` المصفوفة `["i", "like", "coding"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

يجب أن تُرجع `get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` المصفوفة `["debug", "test", "deploy"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```
