---
id: 68b06e589bf2273243814775
title: "تحدٍّ 38: مولد العنوان المختصر"
challengeType: 29
dashedName: challenge-38
---

# --description--

بالنظر إلى سلسلة نصية، أرجع إصدارًا مناسبًا للاستخدام في عنوان URL من السلسلة باستخدام القيود التالية:

- يجب أن تكون كل الأحرف أحرف صغيرة.
- يجب إزالة كل الأحرف التي ليست أحرفًا أو أرقامًا أو مسافات.
- يجب استبدال كل المسافات برمز المسافة المشفر في URL `%20`.
- يجب استبدال المسافات المتتالية برمز `%20` واحد فقط.
- يجب ألا تحتوي السلسلة المرجعة على `%20` في البداية أو النهاية.

# --hints--

`generate_slug("helloWorld")` يجب أن تُرجع `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("helloWorld"), "helloworld")`)
}})
```

`generate_slug("hello world!")` يجب أن تُرجع `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello world!"), "hello%20world")`)
}})
```

`generate_slug(" hello-world ")` يجب أن تُرجع `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug(" hello-world "), "helloworld")`)
}})
```

`generate_slug("hello  world")` يجب أن تُرجع `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello  world"), "hello%20world")`)
}})
```

`generate_slug("  ?H^3-1*1]0! W[0%R#1]D  ")` يجب أن تُرجع `"h3110%20w0r1d"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_slug(str):

    return str
```

# --solutions--

```py
import re
def generate_slug(s):
    cleaned = re.sub(r'[^a-zA-Z0-9 ]+', '', s)
    cleaned = re.sub(r'\s+', ' ', cleaned).strip()
    return cleaned.lower().replace(' ', '%20')
```
