---
id: 68b06e589bf2273243814777
title: "تحدٍّ 39: ملء الخزان"
challengeType: 29
dashedName: challenge-39
---

# --description--

بما أن حجم خزان الوقود، ومستوى الوقود الحالي، وسعر الجالون معروفة، أعد تكلفة ملء الخزان بالكامل.

- `tankSize` هو السعة الكلية للخزان بالجالونات.
- `fuelLevel` هو كمية الوقود الحالية في الخزان بالجالونات.
- `pricePerGallon` هو تكلفة جالون واحد من الوقود.
- يجب تقريب القيمة المرجعة إلى منزلتين عشريتين بالشكل: `"$d.dd"`.

# --hints--

`cost_to_fill(20, 0, 4.00)` يجب أن تُعيد `"$80.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` يجب أن تُعيد `"$17.50"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` يجب أن تُعيد `"$29.25"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` يجب أن تُعيد `"$0.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` يجب أن تُعيد `"$21.89"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
