---
id: 68b1f72371a5ac895ac70a06
title: "تحدٍّ 42: تخزين الفيديو"
challengeType: 29
dashedName: challenge-42
---

# --description--

معطى حجم فيديو، وحدة لحجم الفيديو، سعة قرص صلب، ووحدة للقرص الصلب، أرجع عدد الفيديوهات التي يمكن للقرص الصلب تخزينها باستخدام القيود التالية:

- وحدة حجم الفيديو يمكن أن تكون بايت (`"B"`)، كيلوبايت (`"KB"`)، ميغابايت (`"MB"`)، أو غيغابايت (`"GB"`).
- إذا لم تُعطَ واحدة من وحدات الفيديو السابقة، أرجع `"Invalid video unit"`.
- وحدة سعة القرص الصلب يمكن أن تكون غيغابايت (`"GB"`) أو تيرابايت (`"TB"`).
- إذا لم تُعطَ واحدة من وحدات القرص الصلب السابقة، أرجع `"Invalid drive unit"`.
- أرجع عدد الفيديوهات الكاملة التي يمكن للقرص استيعابها.
- استخدم التحويلات التالية:

| الوحدة | ما يعادلها |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |
| 1 TB |   1000 GB  |

على سبيل المثال، معطى `500`، `"MB"`، `100`، و`"GB"` كمعلمات، حدد كم عدد فيديوهات 500 ميغابايت يمكن أن تتسع على قرص صلب سعته 100 غيغابايت.

# --hints--

`number_of_videos(500, "MB", 100, "GB")` يجب أن تُرجع `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500, "MB", 100, "GB"), 200)`)
}})
```

`number_of_videos(1, "TB", 10, "TB")` يجب أن تُرجع `"Invalid video unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1, "TB", 10, "TB"), "Invalid video unit")`)
}})
```

`number_of_videos(2000, "MB", 100000, "MB")` يجب أن تُرجع `"Invalid drive unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(2000, "MB", 100000, "MB"), "Invalid drive unit")`)
}})
```

`number_of_videos(500000, "KB", 2, "TB")` يجب أن تُرجع `4000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500000, "KB", 2, "TB"), 4000)`)
}})
```

`number_of_videos(1.5, "GB", 2.2, "TB")` يجب أن تُرجع `1466`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1.5, "GB", 2.2, "TB"), 1466)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):

    return video_size
```

# --solutions--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):
    video_units = {"KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000}
    drive_units = {"GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000}

    if video_unit not in video_units:
        return "Invalid video unit"
    if drive_unit not in drive_units:
        return "Invalid drive unit"

    video_bytes = video_size * video_units[video_unit]
    drive_bytes = drive_size * drive_units[drive_unit]

    return int(drive_bytes // video_bytes)
```
