---
id: 68b1f72371a5ac895ac70a02
title: "تحدٍّ 40: تخزين الصور"
challengeType: 29
dashedName: challenge-40
---

# --description--

معطى حجم الصورة بالميغابايت (MB) وسعة القرص الصلب بالجيغابايت (GB)، أرجع عدد الصور التي يمكن للقرص الصلب تخزينها باستخدام القيود التالية:

- 1 جيغابايت يساوي 1000 ميغابايت.
- أرجع عدد الصور الكاملة التي يمكن للقرص تخزينها.

# --hints--

يجب أن تُرجع الدالة `number_of_photos(1, 1)` القيمة `1000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(1, 1), 1000)`)
}})
```

يجب أن تُرجع الدالة `number_of_photos(2, 1)` القيمة `500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(2, 1), 500)`)
}})
```

يجب أن تُرجع الدالة `number_of_photos(4, 256)` القيمة `64000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(4, 256), 64000)`)
}})
```

يجب أن تُرجع الدالة `number_of_photos(3.5, 750)` القيمة `214285`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 750), 214285)`)
}})
```

يجب أن تُرجع الدالة `number_of_photos(3.5, 5.5)` القيمة `1571`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 5.5), 1571)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_photos(photo_size_mb, drive_size_gb):

    return photo_size_mb
```

# --solutions--

```py
def number_of_photos(photo_size_mb, drive_size_gb):
    drive_size_mb = drive_size_gb * 1000
    return drive_size_mb // photo_size_mb
```
