---
id: 68b1f72371a5ac895ac70a04
title: "تحدٍّ 41: تخزين الملفات"
challengeType: 29
dashedName: challenge-41
---

# --description--

معطى حجم ملف، وحدة حجم الملف، وسعة القرص الصلب بالجيجابايت (GB)، أرجع عدد الملفات التي يمكن للقرص الصلب تخزينها باستخدام القيود التالية:

- وحدة حجم الملف يمكن أن تكون بايت (`"B"`)، كيلوبايت (`"KB"`)، أو ميغابايت (`"MB"`).
- أرجع عدد الملفات الكاملة التي يمكن للقرص استيعابها.
- استخدم التحويلات التالية:

| الوحدة | ما يعادلها |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

على سبيل المثال، معطى `500`، `"KB"`، و`1` كمعلمات، حدد كم عدد ملفات 500 كيلوبايت يمكن تخزينها على قرص صلب سعته 1 جيجابايت.

# --hints--

`number_of_files(500, "KB", 1)` يجب أن تُرجع `2000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(500, "KB", 1), 2000)`)
}})
```

`number_of_files(50000, "B", 1)` يجب أن تُرجع `20000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(50000, "B", 1), 20000)`)
}})
```

`number_of_files(5, "MB", 1)` يجب أن تُرجع `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(5, "MB", 1), 200)`)
}})
```

`number_of_files(4096, "B", 1.5)` يجب أن تُرجع `366210`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4096, "B", 1.5), 366210)`)
}})
```

`number_of_files(220.5, "KB", 100)` يجب أن تُرجع `453514`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(220.5, "KB", 100), 453514)`)
}})
```

`number_of_files(4.5, "MB", 750)` يجب أن تُرجع `166666`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4.5, "MB", 750), 166666)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_files(file_size, file_unit, drive_size_gb):

    return file_size
```

# --solutions--

```py
def number_of_files(file_size, file_unit, drive_size_gb):
    drive_size_bytes = drive_size_gb * 1000 * 1000 * 1000

    if file_unit == "B":
        file_size_bytes = file_size
    elif file_unit == "KB":
        file_size_bytes = file_size * 1000
    else:
        file_size_bytes = file_size * 1000 * 1000

    return int(drive_size_bytes // file_size_bytes)
```
