---
id: 68b1f72371a5ac895ac70a08
title: "تحدٍّ 43: الأرقام مقابل الحروف"
challengeType: 29
dashedName: challenge-43
---

# --description--

أعطِ سلسلة نصية، أرجع `"digits"` إذا كانت السلسلة تحتوي على أرقام أكثر من الحروف، و`"letters"` إذا كانت تحتوي على حروف أكثر من الأرقام، و`"tie"` إذا كان عدد الأرقام والحروف متساويًا.

- الأرقام تتكون من `0-9`.
- الحروف تتكون من `a-z` سواء بحروف كبيرة أو صغيرة.
- تجاهل أي حروف أخرى.

# --hints--

`digits_or_letters("abc123")` يجب أن تُرجع `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123"), "tie")`)
}})
```

`digits_or_letters("a1b2c3d")` يجب أن تُرجع `"letters"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("a1b2c3d"), "letters")`)
}})
```

`digits_or_letters("1a2b3c4")` يجب أن تُرجع `"digits"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("1a2b3c4"), "digits")`)
}})
```

`digits_or_letters("abc123!@#DEF")` يجب أن تُرجع `"letters"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123!@#DEF"), "letters")`)
}})
```

`digits_or_letters("H3110 W0R1D")` يجب أن تُرجع `"digits"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("H3110 W0R1D"), "digits")`)
}})
```

`digits_or_letters("P455W0RD")` يجب أن تُرجع `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("P455W0RD"), "tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def digits_or_letters(s):

    return s
```

# --solutions--

```py
def digits_or_letters(s):
    digit_count = 0
    letter_count = 0

    for char in s:
        if char.isdigit():
            digit_count += 1
        elif char.isalpha():
            letter_count += 1

    if digit_count > letter_count:
        return "digits"
    if letter_count > digit_count:
        return "letters"
    return "tie"
```
