---
id: 68b7687dded630607aceccb3
title: "تحدٍّ 49: محلل القسم العلوي لملف CSV"
challengeType: 29
dashedName: challenge-49
---

# --description--

بالنظر إلى السطر الأول من ملف قيم مفصولة بفواصل (CSV)، أرجع مصفوفة تحتوي على العناوين.

- السطر الأول من ملف CSV يحتوي على عناوين مفصولة بفواصل.
- أزل أي مسافات بيضاء زائدة في بداية أو نهاية كل عنوان.

# --hints--

`get_headings("name,age,city")` يجب أن تُرجع `["name", "age", "city"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("name,age,city"), ["name", "age", "city"])`)
}})
```

`get_headings("first name,last name,phone")` يجب أن تُرجع `["first name", "last name", "phone"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("first name,last name,phone"), ["first name", "last name", "phone"])`)
}})
```

`get_headings("username , email , signup date ")` يجب أن تُرجع `["username", "email", "signup date"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("username , email , signup date "), ["username", "email", "signup date"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_headings(csv):

    return csv
```

# --solutions--

```py
def get_headings(csv):
    return [h.strip() for h in csv.split(",")]
```
