---
id: 68b7cadffed0e75a517da673
title: "تحدٍّ 52: من الثنائي إلى العشري"
challengeType: 29
dashedName: challenge-52
---

# --description--

أعطِ سلسلة نصية تمثل رقمًا ثنائيًا، أرجع ما يعادله في النظام العشري كرقم.

يستخدم الرقم الثنائي فقط الرقمين `0` و `1` لتمثيل أي رقم. لتحويل الرقم الثنائي إلى عشري، اضرب كل رقم في قوة من `2` واجمع النتائج معًا. ابدأ بضرب الرقم الأيمن في `2^0`، والرقم التالي إلى اليسار في `2^1`، وهكذا. بعد ضرب كل الأرقام في قوة من `2`، اجمع النتائج معًا.

على سبيل المثال، الرقم الثنائي `101` يساوي `5` في النظام العشري لأن:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`to_decimal("101")` يجب أن تُرجع `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("101"), 5)`)
}})
```

`to_decimal("1010")` يجب أن تُرجع `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010"), 10)`)
}})
```

`to_decimal("10010")` يجب أن تُرجع `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("10010"), 18)`)
}})
```

`to_decimal("1010101")` يجب أن تُرجع `85`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010101"), 85)`)
}})
```

# --seed--

## --seed-contents--

```py
def to_decimal(binary):

    return binary
```

# --solutions--

```py
def to_decimal(binary):
    decimal = 0
    for i, bit_char in enumerate(reversed(binary)):
        bit = int(bit_char)
        decimal += bit * (2 ** i)
    return decimal
```
