---
id: 68b7cadffed0e75a517da675
title: "تحدٍّ 53: من عشري إلى ثنائي"
challengeType: 29
dashedName: challenge-53
---

# --description--

أعطِ عددًا صحيحًا غير سالب، أرجع تمثيله الثنائي كسلسلة نصية.

العدد الثنائي يستخدم فقط الرقمين `0` و `1` لتمثيل أي رقم. لتحويل عدد عشري إلى ثنائي، اقسم العدد على `2` مرارًا وسجل الباقي. كرر حتى يصبح العدد صفرًا. اقرأ البواقي من الأخير إلى الأول. على سبيل المثال، لتحويل `12` إلى ثنائي:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` في النظام الثنائي هو `1100`.

# --hints--

`to_binary(5)` يجب أن تُرجع `"101"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(5), "101")`)
}})
```

`to_binary(12)` يجب أن تُرجع `"1100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(12), "1100")`)
}})
```

`to_binary(50)` يجب أن تُرجع `"110010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(50), "110010")`)
}})
```

`to_binary(99)` يجب أن تُرجع `"1100011"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(99), "1100011")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_binary(decimal):

    return decimal
```

# --solutions--

```py
def to_binary(decimal):
    if decimal == 0:
        return "0"
    binary = ""
    while decimal > 0:
        binary = str(decimal % 2) + binary
        decimal //= 2
    return binary
```
