---
id: 68b7687dded630607aceccb1
title: "تحدٍّ 48: كاشف الرسائل المزعجة"
challengeType: 29
dashedName: challenge-48
---

# --description--

معطى رقم هاتف بالصيغة `"+A (BBB) CCC-DDDD"`، حيث تمثل كل حرف رقمًا كما يلي:

- `A` تمثل رمز الدولة ويمكن أن تكون أي عدد من الأرقام.
- `BBB` تمثل رمز المنطقة وستكون دائمًا ثلاثة أرقام.
- `CCC` و `DDDD` تمثلان الرقم المحلي وستكون دائمًا ثلاثة وأربعة أرقام على التوالي.

حدد ما إذا كان رقمًا مزعجًا بناءً على المعايير التالية:

- رمز الدولة أطول من رقمين أو لا يبدأ بصفر (`0`).
- رمز المنطقة أكبر من 900 أو أقل من 200.
- مجموع أول ثلاثة أرقام من الرقم المحلي يظهر ضمن آخر أربعة أرقام من الرقم المحلي.
- الرقم يحتوي على نفس الرقم مكررًا أربع مرات أو أكثر على التوالي (مع تجاهل أحرف التنسيق).

# --hints--

`is_spam("+0 (200) 234-0182")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (200) 234-0182"), False)`)
}})
```

`is_spam("+091 (555) 309-1922")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+091 (555) 309-1922"), True)`)
}})
```

`is_spam("+1 (555) 435-4792")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+1 (555) 435-4792"), True)`)
}})
```

`is_spam("+0 (955) 234-4364")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (955) 234-4364"), True)`)
}})
```

`is_spam("+0 (155) 131-6943")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (155) 131-6943"), True)`)
}})
```

`is_spam("+0 (555) 135-0192")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 135-0192"), True)`)
}})
```

`is_spam("+0 (555) 564-1987")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 564-1987"), True)`)
}})
```

`is_spam("+00 (555) 234-0182")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+00 (555) 234-0182"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_spam(number):

    return number
```

# --solutions--

```py
import re
def is_spam(number):
    digits = re.sub(r"\D", "", number)
    match = re.match(r"^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$", number)
    country_code, area_code, ccc, dddd = match.groups()

    if len(country_code) > 2 or not country_code.startswith("0"):
        return True

    area_num = int(area_code)
    if area_num > 900 or area_num < 200:
        return True

    sum_ccc = sum(int(d) for d in ccc)
    if str(sum_ccc) in dddd:
        return True

    if re.search(r"(\d)\1\1\1", digits):
        return True

    return False
```
