---
id: 68b7687dded630607aceccad
title: "تحدٍّ 46: ثاني أكبر قيمة"
challengeType: 29
dashedName: challenge-46
---

# --description--

معطاة مصفوفة، أعد الرقم المميز الثاني الأكبر.

# --hints--

يجب أن تُعيد الدالة `second_largest([1, 2, 3, 4])` القيمة `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([1, 2, 3, 4]), 3)`)
}})
```

يجب أن تُعيد الدالة `second_largest([20, 139, 94, 67, 31])` القيمة `94`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([20, 139, 94, 67, 31]), 94)`)
}})
```

يجب أن تُعيد الدالة `second_largest([2, 3, 4, 6, 6])` القيمة `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([2, 3, 4, 6, 6]), 4)`)
}})
```

يجب أن تُعيد الدالة `second_largest([10, -17, 55.5, 44, 91, 0])` القيمة `55.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([10, -17, 55.5, 44, 91, 0]), 55.5)`)
}})
```

يجب أن تُعيد الدالة `second_largest([1, 0, -1, 0, 1, 0, -1, 1, 0])` القيمة `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([1, 0, -1, 0, 1, 0, -1, 1, 0]), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def second_largest(arr):

    return arr
```

# --solutions--

```py
def second_largest(arr):
    unique = list(set(arr))
    unique.sort(reverse=True)
    return unique[1]
```
