---
id: 68b7cadffed0e75a517da677
title: "تحدٍّ 54: قوة كلمة المرور P@ssw0rd Str3ngth!"
challengeType: 29
dashedName: challenge-54
---

# --description--

بالنظر إلى سلسلة نصية لكلمة مرور، أعد `"weak"` أو `"medium"` أو `"strong"` بناءً على قوة كلمة المرور.

تُقيَّم كلمة المرور وفقًا للقواعد التالية:

- أن تكون على الأقل 8 أحرف.
- أن تحتوي على أحرف كبيرة وأحرف صغيرة.
- أن تحتوي على رقم واحد على الأقل.
- أن تحتوي على حرف خاص واحد على الأقل من هذه المجموعة: `!`، `@`، `#`، `$`، `%`، `^`، `&`، أو `*`.

أعد `"weak"` إذا استوفت كلمة المرور أقل من شرطين.
أعد `"medium"` إذا استوفت كلمة المرور شرطين أو ثلاثة.
أعد `"strong"` إذا استوفت كلمة المرور جميع الشروط الأربعة.

# --hints--

`check_strength("123456")` يجب أن تعيد `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("123456"), "weak")`)
}})
```

`check_strength("pass!!!")` يجب أن تعيد `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("pass!!!"), "weak")`)
}})
```

`check_strength("Qwerty")` يجب أن تعيد `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("Qwerty"), "weak")`)
}})
```

`check_strength("PASSWORD")` يجب أن تعيد `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD"), "weak")`)
}})
```

`check_strength("PASSWORD!")` يجب أن تعيد `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD!"), "medium")`)
}})
```

`check_strength("PassWord%^!")` يجب أن تعيد `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PassWord%^!"), "medium")`)
}})
```

`check_strength("qwerty12345")` يجب أن تعيد `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("qwerty12345"), "medium")`)
}})
```

`check_strength("S3cur3P@ssw0rd")` يجب أن تعيد `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("S3cur3P@ssw0rd"), "strong")`)
}})
```

`check_strength("C0d3&Fun!")` يجب أن تعيد `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("C0d3&Fun!"), "strong")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_strength(password):

    return password
```

# --solutions--

```py
import re
def check_strength(password):
    rules_met = 0

    if len(password) >= 8:
        rules_met += 1
    if re.search(r'[a-z]', password) and re.search(r'[A-Z]', password):
        rules_met += 1
    if re.search(r'\d', password):
        rules_met += 1
    if re.search(r'[!@#$%^&*]', password):
        rules_met += 1

    if rules_met < 2:
        return "weak"
    elif rules_met < 4:
        return "medium"
    else:
        return "strong"
```
