---
id: 68c1a929005bf54d342aa8d6
title: "تحدٍّ 59: أسبوع الفضاء اليوم 5: منطقة جولدي لوكس"
challengeType: 29
dashedName: challenge-59
---

# --description--

في اليوم الخامس من أسبوع الفضاء، ستحسب "منطقة جولدي لوكس" لنجم - المنطقة حول النجم حيث تكون الظروف "مناسبة تمامًا" لوجود الماء السائل.

معطى كتلة نجم، أرجع مصفوفة تحتوي على مسافات بداية ونهاية منطقة جولدي لوكس بوحدات فلكية.

لحساب منطقة جولدي لوكس:

1. احسب لمعان النجم برفع كتلته إلى القوة 3.5.
2. بداية المنطقة تساوي 0.95 مضروبة في الجذر التربيعي للمعان النجم.
3. نهاية المنطقة تساوي 1.37 مضروبة في الجذر التربيعي للمعان النجم.

- أرجع المسافات بعد تقريبها إلى منزلتين عشريتين.

على سبيل المثال، إذا كانت الكتلة `1`، أرجع `[0.95, 1.37]`.

# --hints--

`goldilocks_zone(1)` يجب أن تُرجع `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` يجب أن تُرجع `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` يجب أن تُرجع `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` يجب أن تُرجع `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` يجب أن تُرجع `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
