---
id: 68c1a929005bf54d342aa8d4
title: "تحدٍّ 57: أسبوع الفضاء اليوم 3: اتصل بالمنزل"
challengeType: 29
dashedName: challenge-57
---

# --description--

في اليوم الثالث من أسبوع الفضاء، تُعطى مصفوفة أرقام تمثل المسافات (بالكيلومترات) بينك وبين الأقمار الصناعية وكوكبك الأصلي في مسار اتصال. حدّد كم من الوقت سيستغرق وصول رسالة تُرسل عبر المسار إلى كوكب الوجهة باستخدام القيود التالية:

- القيمة الأولى في المصفوفة هي المسافة من موقعك إلى القمر الصناعي الأول.
- كل قيمة لاحقة، ما عدا الأخيرة، هي المسافة إلى القمر الصناعي التالي.
- القيمة الأخيرة في المصفوفة هي المسافة من القمر الصناعي السابق إلى كوكبك الأصلي.
- تنتقل الرسالة بسرعة 300,000 كم/ثانية.
- كل قمر صناعي تمر الرسالة **من خلاله** يضيف تأخير إرسال مقداره 0.5 ثانية.
- أرجع رقمًا مقربًا إلى 4 منازل عشرية، مع إزالة الأصفار الزائدة في النهاية.

# --hints--

`send_message([300000, 300000])` يجب أن تُرجع `2.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([300000, 300000]), 2.5)`)
}})
```

`send_message([384400, 384400])` يجب أن تُرجع `3.0627`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([384400, 384400]), 3.0627)`)
}})
```

`send_message([54600000, 54600000])` يجب أن تُرجع `364.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([54600000, 54600000]), 364.5)`)
}})
```

`send_message([1000000, 500000000, 1000000])` يجب أن تُرجع `1674.3333`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([1000000, 500000000, 1000000]), 1674.3333)`)
}})
```

`send_message([10000, 21339, 50000, 31243, 10000])` يجب أن تُرجع `2.4086`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([10000, 21339, 50000, 31243, 10000]), 2.4086)`)
}})
```

`send_message([802101, 725994, 112808, 3625770, 481239])` يجب أن تُرجع `21.1597`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([802101, 725994, 112808, 3625770, 481239]), 21.1597)`)
}})
```

# --seed--

## --seed-contents--

```py
def send_message(route):

    return route
```

# --solutions--

```py
def send_message(route):
    total_distance = sum(route)
    delay = (len(route) - 1) * 0.5
    time = total_distance / 300_000
    total = time + delay
    return round(total, 4)
```
