---
id: 68c1a929005bf54d342aa8d3
title: "تحدٍّ 56: أسبوع الفضاء اليوم الثاني: البحث عن كواكب خارج المجموعة الشمسية"
challengeType: 29
dashedName: challenge-56
---

# --description--

في اليوم الثاني من أسبوع الفضاء، تُعطى سلسلة نصية حيث يمثل كل حرف قراءة سطوع نجم. حدّد إذا كانت القراءات قد اكتشفت كوكبًا خارج المجموعة الشمسية باستخدام طريقة العبور. طريقة العبور هي عندما يمر كوكب أمام نجم، مما يقلل من السطوع المرصود له.

- قراءات السطوع تتكون فقط من الأحرف `0-9` و `A-Z` حيث تمثل كل قراءة القيم الرقمية التالية:
- الأحرف `0-9` تمثل مستويات السطوع من `0-9`.
- الأحرف `A-Z` تمثل مستويات السطوع من `10-35`.

يُعتبر النجم يحتوي على كوكب خارج المجموعة الشمسية إذا كانت أي قراءة واحدة أقل من أو تساوي 80% من متوسط جميع القراءات. على سبيل المثال، إذا كان متوسط سطوع النجم 10، يُعتبر أنه يحتوي على كوكب إذا كانت أي قراءة واحدة 8 أو أقل.

# --hints--

`has_exoplanet("665544554")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("665544554"), False)`)
}})
```

`has_exoplanet("FGFFCFFGG")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FGFFCFFGG"), True)`)
}})
```

`has_exoplanet("MONOPLONOMONPLNOMPNOMP")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("MONOPLONOMONPLNOMPNOMP"), False)`)
}})
```

`has_exoplanet("FREECODECAMP")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FREECODECAMP"), True)`)
}})
```

`has_exoplanet("9AB98AB9BC98A")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("9AB98AB9BC98A"), False)`)
}})
```

`has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_exoplanet(readings):

    return readings
```

# --solutions--

```py
def has_exoplanet(readings):
    values = [int(c, 36) for c in readings]
    average = sum(values) / len(values)
    threshold = average * 0.8
    for v in values:
        if v <= threshold:
            return True
    return False
```
