---
id: 68c1a929005bf54d342aa8d5
title: "تحدٍّ 58: أسبوع الفضاء اليوم الرابع: موقع الهبوط"
challengeType: 29
dashedName: challenge-58
---

# --description--

في اليوم الرابع من أسبوع الفضاء، تُعطى مصفوفة من الأرقام (مصفوفة من المصفوفات)، تمثل مواقع هبوط محتملة للمركبة الجوالة. ابحث عن موقع الهبوط الأكثر أمانًا بناءً على القواعد التالية:

- كل موقع في المصفوفة يحتوي على رقم من `0-9`، شاملة.
- أي `0` يمثل موقع هبوط محتمل.
- أي رقم غير `0` خطير جدًا للهبوط. كلما زاد الرقم، زادت الخطورة.
- الموقع الأكثر أمانًا هو خلية `0` التي تحيط بها خلايا (حتى 4 جيران، تجاهل الأقطار) بأقل مجموع خطر.
- تجاهل الجيران خارج الحدود (الزوايا والحواف لها جيران أقل فقط).
- أعد إحداثيات موقع الهبوط الأكثر أمانًا. سيكون هناك دائمًا موقع واحد فقط هو الأكثر أمانًا.

على سبيل المثال، إذا كانت المصفوفة:

```js
[
  [1, 0],
  [2, 0]
]
```

أعد `[0, 1]`، وهي إحداثيات الـ `0` في المصفوفة الأولى.

# --hints--

`find_landing_spot([[1, 0], [2, 0]])` يجب أن تُعيد `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 0], [2, 0]]), [0, 1])`)
}})
```

`find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` يجب أن تُعيد `[1, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1])`)
}})
```

`find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` يجب أن تُعيد `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2])`)
}})
```

`find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` يجب أن تُعيد `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_landing_spot(matrix):

    return matrix
```

# --solutions--

```py
def find_landing_spot(matrix):
    best_spot = None
    lowest_neighbor_sum = float('inf')

    for i in range(len(matrix)):
        for j in range(len(matrix[i])):
            if matrix[i][j] == 0:
                current_neighbor_sum = 0

                if i > 0:
                    current_neighbor_sum += matrix[i - 1][j]
                if j < len(matrix[i]) - 1:
                    current_neighbor_sum += matrix[i][j + 1]
                if i < len(matrix) - 1:
                    current_neighbor_sum += matrix[i + 1][j]
                if j > 0:
                    current_neighbor_sum += matrix[i][j - 1]

                if current_neighbor_sum < lowest_neighbor_sum:
                    lowest_neighbor_sum = current_neighbor_sum
                    best_spot = [i, j]

    return best_spot
```
