---
id: 68c497f3aaefc9fd9f1b0e26
title: "تحدٍّ 62: التحويل من النظام السداسي عشري إلى العشري"
challengeType: 29
dashedName: challenge-62
---

# --description--

أعطِ سلسلة نصية تمثل رقمًا في النظام السداسي عشري (القاعدة 16)، أرجع قيمته في النظام العشري (القاعدة 10) كعدد صحيح.

النظام السداسي عشري هو نظام أرقام يستخدم 16 رمزًا:

- تمثّل `0-9` القيم من `0` إلى `9`.
- تمثّل `A-F` القيم من `10` إلى `15`.

إليك جدول تحويل جزئي:

| سداسي عشري | عشري |
|:-----------:|:-------:|
| 0           | 0       |
| 1           | 1       |
| ...         | ...     |
| 9           | 9       |
| A           | 10      |
| ...         | ...     |
| F           | 15      |
| 10          | 16      |
| ...         | ...     |
| 9F          | 159     |
| A0          | 160     |
| ...         | ...     |
| FF          | 255     |
| 100         | 256     |

- السلسلة النصية ستحتوي فقط على الأحرف `0–9` و `A–F`.

# --hints--

يجب أن تُرجع الدالة `hex_to_decimal("A")` القيمة `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A"), 10)`)
}})
```

يجب أن تُرجع الدالة `hex_to_decimal("15")` القيمة `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("15"), 21)`)
}})
```

يجب أن تُرجع الدالة `hex_to_decimal("2E")` القيمة `46`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("2E"), 46)`)
}})
```

يجب أن تُرجع الدالة `hex_to_decimal("FF")` القيمة `255`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("FF"), 255)`)
}})
```

يجب أن تُرجع الدالة `hex_to_decimal("A3F")` القيمة `2623`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A3F"), 2623)`)
}})
```

# --seed--

## --seed-contents--

```py
def hex_to_decimal(hex):

    return hex
```

# --solutions--

```py
def hex_to_decimal(hex):

    return int(hex, 16)
```
