---
id: 68c497f3aaefc9fd9f1b0e25
title: "تحدٍّ 61: أسبوع الفضاء اليوم السابع: وقود الإطلاق"
challengeType: 29
dashedName: challenge-61
---

# --description--

في اليوم الأخير من أسبوع الفضاء، ستُعطى كتلة الحمولة بالكيلوغرام (kg) التي تريد إرسالها إلى المدار. حدّد كمية الوقود اللازمة لإرسال الحمولة إلى المدار باستخدام القواعد التالية:

- الصواريخ تحتاج إلى 1 كغ من الوقود لكل 5 كغ من الكتلة التي يجب رفعها.
- الوقود نفسه له كتلة. لذلك عند إضافة الوقود، تزداد الكتلة التي يجب رفعها، مما يتطلب مزيدًا من الوقود، وهذا يزيد الكتلة، وهكذا.
- لحساب إجمالي الوقود اللازم: ابدأ بكتلة الحمولة، احسب الوقود اللازم لها، أضف هذا الوقود إلى الكتلة الإجمالية، ثم احسب مرة أخرى. كرر هذه العملية حتى تصبح كمية الوقود الإضافية المطلوبة أقل من 1 كغ، ثم توقف.
- تجاهل كتلة الصاروخ نفسه. احسب فقط الوقود اللازم لرفع الحمولة ووقودها الخاص.

على سبيل المثال، إذا كانت كتلة الحمولة 50 كغ، ستحتاج إلى 10 كغ من الوقود لرفعها (الحمولة / 5)، مما يزيد الكتلة الإجمالية إلى 60 كغ، والتي تحتاج إلى 12 كغ للرفع (2 كغ إضافية)، مما يزيد الكتلة الإجمالية إلى 62 كغ، والتي تحتاج إلى 12.4 كغ للرفع - 0.4 كغ إضافية - وهي أقل من 1 كغ إضافية، لذا نتوقف هنا. الكتلة الإجمالية للرفع هي 62.4 كغ، منها 50 كغ هي الحمولة الأصلية و12.4 كغ وقود.

- أرجع كمية الوقود اللازمة مقربة إلى منزلة عشرية واحدة.

# --hints--

`launch_fuel(50)` يجب أن تُرجع `12.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(50), 12.4)`)
}})
```

`launch_fuel(500)` يجب أن تُرجع `124.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(500), 124.8)`)
}})
```

`launch_fuel(243)` يجب أن تُرجع `60.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(243), 60.7)`)
}})
```

`launch_fuel(11000)` يجب أن تُرجع `2749.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(11000), 2749.8)`)
}})
```

`launch_fuel(6214)` يجب أن تُرجع `1553.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(6214), 1553.4)`)
}})
```

# --seed--

## --seed-contents--

```py
def launch_fuel(payload):

    return payload
```

# --solutions--

```py
def launch_fuel(payload):
    total_mass = payload
    additional_fuel = total_mass / 5
    total_fuel = additional_fuel

    while additional_fuel >= 1:
        total_mass += additional_fuel
        additional_fuel = (total_mass / 5) - total_fuel
        total_fuel += additional_fuel

    return round(total_fuel, 1)
```
