---
id: 68c497f3aaefc9fd9f1b0e24
title: "تحدٍّ 60: أسبوع الفضاء اليوم 6: مرحلة القمر"
challengeType: 29
dashedName: challenge-60
---

# --description--

في اليوم السادس من أسبوع الفضاء، ستُعطى تاريخًا بصيغة `"YYYY-MM-DD"` وتحتاج إلى تحديد مرحلة القمر لذلك اليوم باستخدام القواعد التالية:

استخدم دورة قمرية مبسطة مدتها 28 يومًا، مقسمة إلى أربع مراحل متساوية:

- `"New"`: الأيام من 1 إلى 7
- `"Waxing"`: الأيام من 8 إلى 14
- `"Full"`: الأيام من 15 إلى 21
- `"Waning"`: الأيام من 22 إلى 28

بعد اليوم 28، تتكرر الدورة مع اليوم 1، وهو القمر الجديد.

- استخدم `"2000-01-06"` كمرجع للقمر الجديد (اليوم 1 من الدورة) لتحديد مرحلة اليوم المعطى.
- لن تُعطى تواريخ قبل تاريخ المرجع.
- أرجع المرحلة الصحيحة كسلسلة نصية.

**ملاحظة:** اليوم 1 يمثل يوم القمر الجديد، أي لم تمر أي أيام منذ القمر الجديد السابق.

# --hints--

`moon_phase("2000-01-12")` يجب أن تُرجع `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` يجب أن تُرجع `"Waxing"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` يجب أن تُرجع `"Full"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` يجب أن تُرجع `"Waning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` يجب أن تُرجع `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```
