---
id: 68cae5b538ff798bbd4da002
title: "تحدٍّ 64: من 24 إلى 12"
challengeType: 29
dashedName: challenge-64
---

# --description--

معطى سلسلة نصية تمثل وقتًا من اليوم بصيغة 24 ساعة `"HHMM"`، أرجع الوقت بصيغته المكافئة في 12 ساعة `"H:MM AM"` أو `"H:MM PM"`.

- سيكون المُدخَل دائمًا سلسلة نصية مكونة من أربعة أرقام بصيغة وقت 24 ساعة، من `"0000"` إلى `"2359"`.

# --hints--

`to_12("1124")` يجب أن تُرجع `"11:24 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1124"), "11:24 AM")`)
}})
```

`to_12("0900")` يجب أن تُرجع `"9:00 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0900"), "9:00 AM")`)
}})
```

`to_12("1455")` يجب أن تُرجع `"2:55 PM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1455"), "2:55 PM")`)
}})
```

`to_12("2346")` يجب أن تُرجع `"11:46 PM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("2346"), "11:46 PM")`)
}})
```

`to_12("0030")` يجب أن تُرجع `"12:30 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0030"), "12:30 AM")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_12(time):

    return time
```

# --solutions--

```py
def convert_hours(hours):
    if hours == 0:
        return 12
    elif hours > 12:
        return hours - 12
    else:
        return hours

def to_12(time):
    hours = int(time[:2])
    minutes = time[2:]
    period = "AM" if hours < 12 else "PM"

    return f"{convert_hours(hours)}:{minutes} {period}"
```
