---
id: 68cae5b538ff798bbd4da005
title: "تحدٍّ 67: مدقق البريد الإلكتروني"
challengeType: 29
dashedName: challenge-67
---

# --description--

معطاة سلسلة نصية، حدّد إذا كانت عنوان بريد إلكتروني صالح باستخدام القيود التالية:

- يجب أن تحتوي على رمز `@` واحد فقط.
- الجزء المحلي (قبل `@`):
  - يمكن أن يحتوي فقط على حروف (`a-z`، `A-Z`)، أرقام (`0-9`)، نقاط (`.`)، شرطات سفلية (`_`)، أو شرطات (`-`).
  - لا يمكن أن يبدأ أو ينتهي بنقطة.
- الجزء الخاص بالنطاق (بعد `@`):
  - يجب أن يحتوي على نقطة واحدة على الأقل.
  - يجب أن ينتهي بنقطة تليها حرفان على الأقل.
- لا يمكن أن يحتوي الجزء المحلي أو جزء النطاق على نقطتين متتاليتين.

# --hints--

يجب أن تُرجع الدالة `validate("a@b.cd")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("a@b.cd"), True)`)
}})
```

يجب أن تُرجع الدالة `validate("hell.-w.rld@example.com")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hell.-w.rld@example.com"), True)`)
}})
```

يجب أن تُرجع الدالة `validate(".b@sh.rc")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate(".b@sh.rc"), False)`)
}})
```

يجب أن تُرجع الدالة `validate("example@test.c0")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("example@test.c0"), False)`)
}})
```

يجب أن تُرجع الدالة `validate("freecodecamp.org")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("freecodecamp.org"), False)`)
}})
```

يجب أن تُرجع الدالة `validate("develop.ment_user@c0D!NG.R.CKS")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("develop.ment_user@c0D!NG.R.CKS"), True)`)
}})
```

يجب أن تُرجع الدالة `validate("hello.@wo.rld")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hello.@wo.rld"), False)`)
}})
```

يجب أن تُرجع الدالة `validate("hello@world..com")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hello@world..com"), False)`)
}})
```

يجب أن تُرجع الدالة `validate("develop..ment_user@c0D!NG.R.CKS")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("develop..ment_user@c0D!NG.R.CKS"), False)`)
}})
```

يجب أن تُرجع الدالة `validate("git@commit@push.io")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("git@commit@push.io"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def validate(email):

    return email
```

# --solutions--

```py
import re
def validate(email):
    if '..' in email:
        return False

    parts = email.split('@')
    if len(parts) != 2:
        return False

    local, domain = parts

    if local.startswith('.') or local.endswith('.'):
        return False
    if not re.match(r'^[a-zA-Z0-9._-]+$', local):
        return False

    if '.' not in domain:
        return False

    tld = domain.split('.')[-1]
    if len(tld) < 2 or not tld.isalpha():
        return False

    return True
```
