---
id: 68cae5b538ff798bbd4da007
title: "تحدٍّ 69: الجوارب المفقودة"
challengeType: 29
dashedName: challenge-69
---

# --description--

معطى عدد صحيح يمثل عدد أزواج الجوارب التي بدأت بها، وعدد صحيح آخر يمثل عدد دورات الغسيل التي مررت بها، أرجع عدد أزواج الجوارب الكاملة التي تملكها حاليًا باستخدام القيود التالية:

- كل دورتين غسيل، تفقد جوربًا واحدًا.
- كل ثلاث دورات غسيل، تجد جوربًا مفقودًا واحدًا.
- كل خمس دورات غسيل، يتلف جورب واحد ويجب رميه.
- كل عشر دورات غسيل، تشتري زوجًا جديدًا من الجوارب.
- لا يمكنك أن تمتلك أقل من صفر جوارب إجمالًا.
- يمكن أن تتداخل القواعد. على سبيل المثال، في دورة الغسيل العاشرة، ستفقد جوربًا واحدًا، وترمي جوربًا واحدًا، وتشتري زوجًا جديدًا من الجوارب.
- أرجع عدد أزواج الجوارب الكاملة.

# --hints--

`sock_pairs(2, 5)` يجب أن تُرجع `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(2, 5), 1)`)
}})
```

`sock_pairs(1, 2)` يجب أن تُرجع `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 2), 0)`)
}})
```

`sock_pairs(5, 11)` يجب أن تُرجع `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(5, 11), 4)`)
}})
```

`sock_pairs(6, 25)` يجب أن تُرجع `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(6, 25), 3)`)
}})
```

`sock_pairs(1, 8)` يجب أن تُرجع `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 8), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sock_pairs(pairs, cycles):

    return pairs
```

# --solutions--

```py
def sock_pairs(pairs, cycles):
    socks = pairs * 2

    for i in range(1, cycles + 1):
        if i % 2 == 0:
            socks -= 1
        if i % 3 == 0:
            socks += 1
        if i % 5 == 0:
            socks -= 1
        if i % 10 == 0:
            socks += 2

        if socks < 0:
            socks = 0

    return socks // 2
```
