---
id: 68cae5b538ff798bbd4da00a
title: "تحدٍّ 72: مُعدِّل منظم الحرارة 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

معطى درجة الحرارة الحالية لغرفة بالفهرنهايت ودرجة حرارة مستهدفة بالسلسلة النصية، أرجع سلسلة نصية توضح كيفية تعديل درجة حرارة الغرفة بناءً على هذه القيود:

- أرجع `"Heat: X degrees Fahrenheit"` إذا كانت درجة الحرارة الحالية أقل من المستهدفة. حيث `X` هو عدد الدرجات بالفهرنهايت اللازمة لتسخين الغرفة للوصول إلى الهدف، مقربًا إلى منزلة عشرية واحدة.
- أرجع `"Cool: X degrees Fahrenheit"` إذا كانت درجة الحرارة الحالية أعلى من المستهدفة. حيث `X` هو عدد الدرجات بالفهرنهايت اللازمة لتبريد الغرفة للوصول إلى الهدف، مقربًا إلى منزلة عشرية واحدة.
- أرجع `"Hold"` إذا كانت درجة الحرارة الحالية مساوية للهدف.

لتحويل درجة الحرارة من السلسلة النصية إلى الفهرنهايت، اضرب درجة الحرارة بالسلسلة النصية في 1.8 وأضف 32 إلى الناتج (`F = (C * 1.8) + 32`).

# --hints--

`adjust_thermostat(32, 0)` يجب أن تُرجع `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` يجب أن تُرجع `"Heat: 7.0 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` يجب أن تُرجع `"Cool: 7.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` يجب أن تُرجع `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` يجب أن تُرجع `"Heat: 12.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
