---
id: 68cae5b538ff798bbd4da009
title: "تحدٍّ 71: حاسبة البقشيش"
challengeType: 29
dashedName: challenge-71
---

# --description--

بما أن لديك سعر وجبتك ونسبة بقشيش مخصصة، أرجع مصفوفة تحتوي على ثلاثة قيم للبقشيش؛ 15%، 20%، والمبلغ المخصص.

- ستُعطى الأسعار بالتنسيق: `"$N.NN"`.
- ستُعطى نسب البقشيش المخصصة بهذا التنسيق: `"25%"`.
- أرجع المبالغ بنفس تنسيق `"$N.NN"`، مع تقريبها إلى منزلتين عشريتين.

على سبيل المثال، إذا كان سعر الوجبة `"$10.00"`، وقيمة البقشيش المخصصة `"25%"`، أرجع `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculate_tips("$10.00", "25%")` يجب أن تُرجع `["$1.50", "$2.00", "$2.50"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"])`)
}})
```

`calculate_tips("$89.67", "26%")` يجب أن تُرجع `["$13.45", "$17.93", "$23.31"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"])`)
}})
```

`calculate_tips("$19.85", "9%")` يجب أن تُرجع `["$2.98", "$3.97", "$1.79"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_tips(meal_price, custom_tip):

    return meal_price
```

# --solutions--

```py
def calculate_tips(meal_price, custom_tip):
    meal = float(meal_price[1:])
    custom_percent = float(custom_tip[:-1])

    tips_percent = [15, 20, custom_percent]
    tips = [f"${meal * p / 100:.2f}" for p in tips_percent]

    return tips
```
