---
id: 68cae5b538ff798bbd4da004
title: "تحدٍّ 66: مزيل وسوم HTML"
challengeType: 29
dashedName: challenge-66
---

# --description--

بالنظر إلى سلسلة نصية من شفرة HTML، أزل الوسوم وأعد المحتوى النصي العادي.

- ستحتوي السلسلة النصية المُدخلة على HTML صالح فقط.
- قد تكون وسوم HTML مسكنة.
- أزل الوسوم وأي خاصيات.

على سبيل المثال، يجب أن تُعيد `'<a href="#">Click here</a>'` القيمة `"Click here"`.

# --hints--

يجب أن تُعيد الدالة `strip_tags('<a href="#">Click here</a>')` القيمة `"Click here"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

يجب أن تُعيد الدالة `strip_tags('<p class="center">Hello <b>World</b>!</p>')` القيمة `"Hello World!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

يجب أن تُعيد الدالة `strip_tags('<img src="cat.jpg" alt="Cat">')` سلسلة فارغة (`""`).

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

يجب أن تُعيد الدالة `strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` القيمة `sectionsection`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```
