---
id: 68cae5b538ff798bbd4da003
title: "تحدٍّ 65: عد السلسلة النصية"
challengeType: 29
dashedName: challenge-65
---

# --description--

معطيتان سلسلتان نصيتان، حدّد كم مرة تظهر السلسلة النصية الثانية في الأولى.

- يمكن أن تتداخل سلسلة النمط داخل السلسلة الأولى. على سبيل المثال، تحتوي `"aaa"` على `"aa"` مرتين. الحرفان الأولان `a` والحرفان الثانيان.

# --hints--

`count('abcdefg', 'def')` يجب أن تُرجع `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('abcdefg', 'def'), 1)`)
}})
```

`count('hello', 'world')` يجب أن تُرجع `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('hello', 'world'), 0)`)
}})
```

`count('mississippi', 'iss')` يجب أن تُرجع `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('mississippi', 'iss'), 2)`)
}})
```

`count('she sells seashells by the seashore', 'sh')` يجب أن تُرجع `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('she sells seashells by the seashore', 'sh'), 3)`)
}})
```

`count('101010101010101010101', '101')` يجب أن تُرجع `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('101010101010101010101', '101'), 10)`)
}})
```

# --seed--

## --seed-contents--

```py
def count(text, parameter):

    return text
```

# --solutions--

```py
def count(text, pattern):
    if not pattern:
        return 0
    occurrences = 0

    for i in range(len(text) - len(pattern) + 1):
        if text[i:i+len(pattern)] == pattern:
            occurrences += 1

    return occurrences
```
