---
id: 68cae5b538ff798bbd4da006
title: "تحدٍّ 68: قناع بطاقة الائتمان"
challengeType: 29
dashedName: challenge-68
---

# --description--

بالنظر إلى سلسلة نصية من أرقام بطاقات الائتمان، أعد نسخة مقنعة منها باستخدام القيود التالية:

- تحتوي السلسلة النصية على أربع مجموعات من أربعة أرقام (`0-9`)، مع فصل كل مجموعة بمسافة واحدة أو شرطة واحدة (`-`).
- استبدل كل الأرقام، ما عدا الأربعة الأخيرة، بنجمة (`*`).
- اترك الأحرف المتبقية دون تغيير.

على سبيل المثال، إذا كانت القيمة `"4012-8888-8888-1881"` فأعد `"****-****-****-1881"`.

# --hints--

`mask("4012-8888-8888-1881")` يجب أن تُعيد `"****-****-****-1881"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("4012-8888-8888-1881"), "****-****-****-1881")`)
}})
```

`mask("5105 1051 0510 5100")` يجب أن تُعيد `"**** **** **** 5100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("5105 1051 0510 5100"), "**** **** **** 5100")`)
}})
```

`mask("6011 1111 1111 1117")` يجب أن تُعيد `"**** **** **** 1117"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("6011 1111 1111 1117"), "**** **** **** 1117")`)
}})
```

`mask("2223-0000-4845-0010")` يجب أن تُعيد `"****-****-****-0010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("2223-0000-4845-0010"), "****-****-****-0010")`)
}})
```

# --seed--

## --seed-contents--

```py
def mask(card):

    return card
```

# --solutions--

```py
def mask(card):
    if '-' in card:
        split = card.split('-')
        return f"****-****-****-{split[3]}"
    else:
        split = card.split(' ')
        return f"**** **** **** {split[3]}"
```
