---
id: 68cae5b538ff798bbd4da008
title: "تحدٍّ 70: مستخرج خاصيات HTML"
challengeType: 29
dashedName: challenge-70
---

# --description--

بالنظر إلى سلسلة نصية تمثل عنصر HTML صالح، أرجع خاصيات العنصر باستخدام المعايير التالية:

- سيُعطى عنصر واحد فقط.
- تكون الخاصيات بالشكل: `attribute="value"`.
- أرجع مصفوفة من السلاسل النصية، كل منها يحتوي على الخاصية والقيمة مفصولين بفاصلة، بالشكل التالي: `["attribute1, value1", "attribute2, value2"]`.
- أرجع الخاصيات بنفس الترتيب الذي تظهر به.
- إذا لم توجد خاصيات، أرجع مصفوفة فارغة.

# --hints--

`extract_attributes('<span class="red"></span>')` يجب أن تُرجع `["class, red"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<span class="red"></span>'), ["class, red"])`)
}})
```

`extract_attributes('<meta charset="UTF-8" />')` يجب أن تُرجع `["charset, UTF-8"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<meta charset="UTF-8" />'), ["charset, UTF-8"])`)
}})
```

`extract_attributes("<p>Lorem ipsum dolor sit amet</p>")` يجب أن تُرجع `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes("<p>Lorem ipsum dolor sit amet</p>"), [])`)
}})
```

`extract_attributes('<input name="email" type="email" required="true" />')` يجب أن تُرجع `["name, email", "type, email", "required, true"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"])`)
}})
```

`extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>')` يجب أن تُرجع `["id, submit", "class, btn btn-primary"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"])`)
}})
```

# --seed--

## --seed-contents--

```py
def extract_attributes(element):

    return element
```

# --solutions--

```py
import re
def extract_attributes(element):
    pattern = r'([\w-]+)="([^"]*)"'
    matches = re.findall(pattern, element)
    return [f"{attr}, {val}" for attr, val in matches]
```
