---
id: 68cae5b538ff798bbd4da001
title: "تحدٍّ 63: معركة الكلمات"
challengeType: 29
dashedName: challenge-63
---

# --description--

معطى جملتان تمثلان فريقك وفريقًا معارضًا، حيث تتقاتل كل كلمة من فريقك مع الكلمة المقابلة من الفريق المعارض، حدد الفريق الفائز باستخدام القواعد التالية:

- الجمل المعطاة ستحتوي دائمًا على نفس عدد الكلمات.
- الكلمات مفصولة بمسافة واحدة فقط وستحتوي على حروف فقط.
- قيمة كل كلمة هي مجموع قيم حروفها.
- الحروف من `a` إلى `z` تمثل القيم من `1` إلى `26`. على سبيل المثال، `a` قيمتها `1`، و`z` قيمتها `26`.
- الحرف الكبير يضاعف قيمة الحرف. على سبيل المثال، `A` قيمته `2`، و`Z` قيمته `52`.
- الكلمات تتقاتل بالترتيب: الكلمة الأولى من فريقك تتقاتل مع الكلمة الأولى من الفريق المعارض، وهكذا.
- تفوز الكلمة إذا كانت قيمتها أكبر من قيمة الكلمة المقابلة.
- الفريق الذي يملك عددًا أكبر من الكلمات الفائزة هو الفائز.

أرجع `"We win"` إذا كان فريقك هو الفائز، و`"We lose"` إذا خسر فريقك، و`"Draw"` إذا كان عدد الكلمات الفائزة متساويًا بين الفريقين.

# --hints--

`battle("hello world", "hello word")` يجب أن تُرجع `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` يجب أن تُرجع `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` يجب أن تُرجع `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` يجب أن تُرجع `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` يجب أن تُرجع `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` يجب أن تُرجع `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` يجب أن تُرجع `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
