---
id: 68d2ba1468508398389487cf
title: "تحدٍّ 74: الأغاني المفضلة"
challengeType: 29
dashedName: challenge-74
---

# --description--

هل تتذكر أجهزة iPod؟ صدر أول نموذج منها قبل 24 سنة اليوم، في 23 أكتوبر 2001.

معطى مصفوفة من كائنات الأغاني التي تمثل قائمة تشغيل iPod الخاصة بك، أرجع مصفوفة بعناوين أكثر أغنيتين تم تشغيلهما، مع وضع الأغنية الأكثر تشغيلًا أولًا.

- كل كائن يحتوي على خاصية `"title"` (سلسلة نصية)، وخاصية `"plays"` (عدد صحيح).

# --hints--

`favorite_songs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` يجب أن تُرجع `["Sync or Swim", "Earbud Blues"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"])`)
}})
```

`favorite_songs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` يجب أن تُرجع `["Clickwheel Love", "99 Downloads"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"])`)
}})
```

`favorite_songs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` يجب أن تُرجع `["Song B", "Song C"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"])`)
}})
```

# --seed--

## --seed-contents--

```py
def favorite_songs(playlist):

    return playlist
```

# --solutions--

```py
def favorite_songs(playlist):
    sorted_songs = sorted(playlist, key=lambda x: x["plays"], reverse=True)
    return [song["title"] for song in sorted_songs[:2]]
```
