---
id: 68d2ba1468508398389487ce
title: "تحدٍّ 73: تحدث بحكمة، يجب عليك"
challengeType: 29
dashedName: challenge-73
---

# --description--

بالنظر إلى جملة، أعد إصدارًا منها يبدو كنصيحة من معلم حكيم باستخدام القواعد التالية:

- الكلمات مفصولة بمسافة واحدة.
- ابحث عن أول ظهور لإحدى الكلمات التالية في الجملة: `"have"`، `"must"`، `"are"`، `"will"`، `"can"`.
- انقل كل الكلمات التي تسبق هذه الكلمة وتشملها إلى نهاية الجملة و:
  - حافظ على ترتيب الكلمات عند نقلها.
  - اجعلها كلها بأحرف صغيرة.
  - وأضف فاصلة ومسافة قبلها.
- اجعل الحرف الأول من الكلمة الجديدة الأولى في الجملة بحرف كبير.
- كل الجمل المعطاة ستنتهي بعلامة ترقيم واحدة. احتفظ بعلامة الترقيم الأصلية للجملة وانقلها إلى نهاية الجملة الجديدة.
- أعد الجملة الجديدة، مع التأكد من وجود مسافة واحدة بين كل كلمة وأخرى وعدم وجود مسافات في بداية أو نهاية الجملة.

على سبيل المثال، إذا كانت الجملة `"You must speak wisely."` فأعد `"Speak wisely, you must."`

# --hints--

`wise_speak("You must speak wisely.")` يجب أن تُعيد `"Speak wisely, you must."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You must speak wisely."), "Speak wisely, you must.")`)
}})
```

`wise_speak("You can do it!")` يجب أن تُعيد `"Do it, you can!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You can do it!"), "Do it, you can!")`)
}})
```

`wise_speak("Do you think you will complete this?")` يجب أن تُعيد `"Complete this, do you think you will?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("Do you think you will complete this?"), "Complete this, do you think you will?")`)
}})
```

`wise_speak("All your base are belong to us.")` يجب أن تُعيد `"Belong to us, all your base are."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("All your base are belong to us."), "Belong to us, all your base are.")`)
}})
```

`wise_speak("You have much to learn.")` يجب أن تُعيد `"Much to learn, you have."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You have much to learn."), "Much to learn, you have.")`)
}})
```

# --seed--

## --seed-contents--

```py
def wise_speak(sentence):

    return sentence
```

# --solutions--

```py
def wise_speak(sentence):
    triggers = ["have", "must", "are", "will", "can"]
    punctuation = sentence[-1]
    words = sentence[:-1].split(" ")

    index = next(i for i, w in enumerate(words) if w in triggers)

    to_move = [w.lower() for w in words[:index + 1]]
    remaining = words[index + 1:]

    if remaining:
        remaining[0] = remaining[0][0].upper() + remaining[0][1:]

    return " ".join(remaining) + ", " + " ".join(to_move) + punctuation
```
