---
id: 68d30845cc08266018fc46bd
title: "تحدٍّ 77: منسق المدة"
challengeType: 29
dashedName: challenge-77
---

# --description--

بالنظر إلى عدد صحيح من الثواني، أعد سلسلة نصية تمثل نفس المدة بالتنسيق `"H:MM:SS"`، حيث `"H"` هو عدد الساعات، و`"MM"` هو عدد الدقائق، و`"SS"` هو عدد الثواني. أعد الوقت باستخدام القواعد التالية:

- الثواني: يجب أن تكون دائمًا رقمين.
- الدقائق: يجب حذف الأصفار البادئة عندما لا تكون ضرورية. استخدم `"0"` إذا كانت المدة أقل من دقيقة واحدة.
- الساعات: يجب تضمينها فقط إذا كانت أكبر من صفر.

# --hints--

`format(500)` يجب أن يعيد `"8:20"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(500), "8:20")`)
}})
```

`format(4000)` يجب أن يعيد `"1:06:40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(4000), "1:06:40")`)
}})
```

`format(1)` يجب أن يعيد `"0:01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(1), "0:01")`)
}})
```

`format(5555)` يجب أن يعيد `"1:32:35"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(5555), "1:32:35")`)
}})
```

`format(99999)` يجب أن يعيد `"27:46:39"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(99999), "27:46:39")`)
}})
```

# --seed--

## --seed-contents--

```py
def format(seconds):

    return seconds
```

# --solutions--

```py
def format(seconds):
    h = seconds // 3600
    m = (seconds % 3600) // 60
    s = seconds % 60

    seconds_str = f"{s:02d}"
    minutes_str = str(m)

    if h > 0:
        hours_str = str(h)
        return f"{hours_str}:{minutes_str.zfill(2)}:{seconds_str}"
    else:
        return f"{minutes_str}:{seconds_str}"
```
