---
id: 68d30845cc08266018fc46bc
title: "تحدٍّ 76: الحمض النووي التكميلي"
challengeType: 29
dashedName: challenge-76
---

# --description--

معطاة سلسلة نصية تمثل متتالية DNA، أعد السلسلة التكاملية لها باستخدام القواعد التالية:

- يتكون DNA من الحروف `"A"`، `"C"`، `"G"`، و`"T"`.
- الحرفان `"A"` و`"T"` متكاملان.
- الحرفان `"C"` و`"G"` متكاملان.

على سبيل المثال، إذا كانت السلسلة `"ACGT"`، فأعد `"TGCA"`.

# --hints--

يجب أن تُعيد الدالة `complementary_dna("ACGT")` القيمة `"TGCA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ACGT"), "TGCA")`)
}})
```

يجب أن تُعيد الدالة `complementary_dna("ATGCGTACGTTAGC")` القيمة `"TACGCATGCAATCG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ATGCGTACGTTAGC"), "TACGCATGCAATCG")`)
}})
```

يجب أن تُعيد الدالة `complementary_dna("GGCTTACGATCGAAG")` القيمة `"CCGAATGCTAGCTTC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC")`)
}})
```

يجب أن تُعيد الدالة `complementary_dna("GATCTAGCTAGGCTAGCTAG")` القيمة `"CTAGATCGATCCGATCGATC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC")`)
}})
```

# --seed--

## --seed-contents--

```py
def complementary_dna(strand):

    return strand
```

# --solutions--

```py
def complementary_dna(strand):
    complements = { "A": "T", "T": "A", "C": "G", "G": "C" }
    return "".join(complements[n] for n in strand)
```
