---
id: 68d30fc57588d97fd3027b30
title: "تحدٍّ 79: Navigator"
challengeType: 29
dashedName: challenge-79
---

# --description--

في 28 أكتوبر 1994، أُصدر Netscape Navigator، مما ساعد ملايين الأشخاص على استكشاف الويب المبكر.

معطى مصفوفة من أوامر المتصفح التي نفذتها على Netscape Navigator، أعد الصفحة الحالية التي أنت عليها بعد تنفيذ كل الأوامر باستخدام القواعد التالية:

- تبدأ دائمًا على صفحة `"Home"`، التي لن تكون مدرجة في مصفوفة الأوامر.
- الأوامر الصالحة هي:
  - `"Visit Page"`: حيث `"Page"` هو اسم الصفحة التي تزورها. على سبيل المثال، `"Visit About"` ينقلك إلى صفحة `"About"`. عند زيارة صفحة جديدة، تأكد من حذف أي سجل تقدم لديك.
  - `"Back"`: ينقلك إلى الصفحة السابقة في سجل التصفح أو يبقيك على الصفحة الحالية إذا لم تكن هناك صفحة سابقة.
  - `"Forward"`: ينقلك إلى الأمام في السجل إلى الصفحة التي جئت منها أو يبقيك على الصفحة الحالية إذا لم تكن هناك صفحة للأمام.

على سبيل المثال، معطى `["Visit About Us", "Back", "Forward"]`، أعد `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` يجب أن يعيد `"About Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Back", "Forward"]), "About Us")`)
}})
```

`navigate(["Forward"])` يجب أن يعيد `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Forward"]), "Home")`)
}})
```

`navigate(["Back"])` يجب أن يعيد `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Back"]), "Home")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery"])` يجب أن يعيد `"Gallery"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery"]), "Gallery")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` يجب أن يعيد `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home")`)
}})
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` يجب أن يعيد `"Contact"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact")`)
}})
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` يجب أن يعيد `"Visit Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate(commands):

    return commands
```

# --solutions--

```py
def navigate(commands):
    history = ["Home"]
    current_index = 0

    for command in commands:
        if command.startswith("Visit "):
            history = history[:current_index + 1]
            history.append(command[6:])
            current_index += 1
        elif command == "Back" and current_index > 0:
            current_index -= 1
        elif command == "Forward" and current_index < len(history) - 1:
            current_index += 1

    return history[current_index]
```
