---
id: 68d30845cc08266018fc46be
title: "تحدٍّ 78: متتالية أعداد صحيحة"
challengeType: 29
dashedName: challenge-78
---

# --description--

أعطِ عددًا صحيحًا موجبًا، أرجع سلسلة نصية تحتوي كل الأعداد الصحيحة من `1` حتى العدد المعطى، بما في ذلك هذا العدد، بترتيب رقمي.

على سبيل المثال، إذا كان العدد `5`، أرجع `"12345"`.

# --hints--

`sequence(5)` يجب أن تُرجع `"12345"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(5), "12345")`)
}})
```

`sequence(10)` يجب أن تُرجع `"12345678910"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(10), "12345678910")`)
}})
```

`sequence(1)` يجب أن تُرجع `"1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(1), "1")`)
}})
```

`sequence(27)` يجب أن تُرجع `"123456789101112131415161718192021222324252627"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(27), "123456789101112131415161718192021222324252627")`)
}})
```

# --seed--

## --seed-contents--

```py
def sequence(n):

    return n
```

# --solutions--

```py
def sequence(n):
    result = ""
    for i in range(1, n + 1):
        result += str(i)
    return result
```
