---
id: 68e39ed6106dac2f0a98fd64
title: "تحدٍّ 82: SpOoKy~CaSe"
challengeType: 29
dashedName: challenge-82
---

# --description--

معطاة سلسلة نصية تمثل اسم متغير، حوّلها إلى "spooky case" باستخدام القيود التالية:

- استبدل كل الشرطات السفلية (`_`) والشرطات (`-`) بمد (`~`).
- اجعل الحرف الأول من السلسلة وأي حرف آخر بالتناوب بعده بحروف كبيرة. تجاهل حرف المد عند العد. اجعل بقية الحروف أحرف صغيرة.

على سبيل المثال، معطاة `hello_world`، أرجع `HeLlO~wOrLd`.

# --hints--

`spookify("hello_world")` يجب أن تُرجع `"HeLlO~wOrLd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("hello_world"), "HeLlO~wOrLd")`)
}})
```

`spookify("Spooky_Case")` يجب أن تُرجع `"SpOoKy~CaSe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("Spooky_Case"), "SpOoKy~CaSe")`)
}})
```

`spookify("TRICK-or-TREAT")` يجب أن تُرجع `"TrIcK~oR~tReAt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt")`)
}})
```

`spookify("c_a-n_d-y_-b-o_w_l")` يجب أن تُرجع `"C~a~N~d~Y~~b~O~w~L"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L")`)
}})
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` يجب أن تُرجع `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS")`)
}})
```

# --seed--

## --seed-contents--

```py
def spookify(boo):

    return boo
```

# --solutions--

```py
def spookify(boo):
    replaced = boo.replace("_", "~").replace("-", "~")
    
    result = []
    capitalize = True

    for char in replaced:
        if char == "~":
            result.append(char)
        else:
            result.append(char.upper() if capitalize else char.lower())
            capitalize = not capitalize

    return "".join(result)
```
