---
id: 68e39ed6106dac2f0a98fd63
title: "تحدٍّ 81: العدد الأولي الـ nth"
challengeType: 29
dashedName: challenge-81
---

# --description--

العدد الأولي هو عدد صحيح موجب أكبر من 1 لا يقبل القسمة إلا على 1 ونفسه فقط. أول خمسة أعداد أولية هي `2`، `3`، `5`، `7`، و`11`.

معطى عدد صحيح موجب `n`، أرجع العدد الأولي الـ `n`. على سبيل المثال، إذا كان `5` فأرجع العدد الأولي الخامس: `11`.

# --hints--

يجب أن تُرجع الدالة `nth_prime(5)` القيمة `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(5), 11)`)
}})
```

يجب أن تُرجع الدالة `nth_prime(10)` القيمة `29`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(10), 29)`)
}})
```

يجب أن تُرجع الدالة `nth_prime(16)` القيمة `53`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(16), 53)`)
}})
```

يجب أن تُرجع الدالة `nth_prime(99)` القيمة `523`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(99), 523)`)
}})
```

يجب أن تُرجع الدالة `nth_prime(1000)` القيمة `7919`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(1000), 7919)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_prime(n):

    return n
```

# --solutions--

```py
def nth_prime(n):
    primes = []
    num = 2

    while len(primes) < n:
        is_prime = True
        for i in range(2, int(num**0.5) + 1):
            if num % i == 0:
                is_prime = False
                break
        if is_prime:
            primes.append(num)
        num += 1

    return primes[n - 1]
```
