---
id: 68e39ed6106dac2f0a98fd66
title: "تحدٍّ 84: مصاب"
challengeType: 29
dashedName: challenge-84
---

# --description--

في 2 نوفمبر 1988، تم إطلاق أول دودة إنترنت كبيرة، وأصابت حوالي 10% من الحواسيب المتصلة بالإنترنت بعد يوم واحد فقط.

في هذا التحدٍّ، يُعطى لك عدد الأيام التي مرت منذ إطلاق دودة الإنترنت، ويجب عليك تحديد عدد الحواسيب المصابة باستخدام القواعد التالية:

- في اليوم 0، يُصاب أول حاسوب.
- في كل يوم لاحق، يتضاعف عدد الحواسيب المصابة.
- كل يوم ثالث، يُطبق تصحيح بعد انتشار الفيروس ويقلل عدد الحواسيب المصابة بنسبة 20%. قم بتقريب عدد الحواسيب التي تم تصحيحها إلى أقرب عدد صحيح للأعلى.

على سبيل المثال، في:

- اليوم 0: يصاب حاسوب واحد إجمالاً.
- اليوم 1: يصاب حاسوبان إجمالاً.
- اليوم 2: تصاب 4 حواسيب إجمالاً.
- اليوم 3: تصاب 8 حواسيب إجمالاً. ثم يُطبق التصحيح: 8 مصابين * 20% = 1.6 تم تصحيحهم. قرب 1.6 إلى 2. 8 حواسيب مصابة - 2 مصححة = 6 حواسيب مصابة إجمالاً بعد اليوم 3.

أعد عدد الحواسيب المصابة إجمالاً بعد مرور عدد الأيام المعطى.

# --hints--

`infected(1)` يجب أن تُعيد `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(1), 2)`)
}})
```

`infected(3)` يجب أن تُعيد `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(3), 6)`)
}})
```

`infected(8)` يجب أن تُعيد `152`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(8), 152)`)
}})
```

`infected(17)` يجب أن تُعيد `39808`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(17), 39808)`)
}})
```

`infected(25)` يجب أن تُعيد `5217638`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(25), 5217638)`)
}})
```

# --seed--

## --seed-contents--

```py
def infected(days):

    return days
```

# --solutions--

```py
import math
def infected(days):
    infected = 1
    for day in range(1, days + 1):
        infected *= 2
        if day % 3 == 0:
            patched = math.ceil(infected * 0.2)
            infected -= patched

    return infected
```
