---
id: 68e39ed6106dac2f0a98fd65
title: "تحدٍّ 83: التحقق من التوقيع"
challengeType: 29
dashedName: challenge-83
---

# --description--

معطى سلسلة نصية للرسالة، وسلسلة نصية للمفتاح السري، ورقم التوقيع، حدّد ما إذا كان التوقيع صحيحًا باستخدام هذه الطريقة في الترميز:

- الحروف في الرسالة والمفتاح السري لها القيم التالية:
  - الحروف من `a` إلى `z` لها القيم من `1` إلى `26` على التوالي.
  - الحروف من `A` إلى `Z` لها القيم من `27` إلى `52` على التوالي.
- جميع الأحرف الأخرى لا قيمة لها.
- احسب التوقيع بأخذ مجموع قيم الرسالة زائد مجموع قيم المفتاح السري.

على سبيل المثال، معطى الرسالة `"foo"` والمفتاح السري `"bar"`، سيكون التوقيع `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

تحقق مما إذا كان التوقيع المحسوب يطابق التوقيع المقدم.

# --hints--

`verify("foo", "bar", 57)` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("foo", "bar", 57), True)`)
}})
```

`verify("foo", "bar", 54)` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("foo", "bar", 54), False)`)
}})
```

`verify("freeCodeCamp", "Rocks", 238)` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("freeCodeCamp", "Rocks", 238), True)`)
}})
```

`verify("Is this valid?", "No", 210)` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Is this valid?", "No", 210), False)`)
}})
```

`verify("Is this valid?", "Yes", 233)` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Is this valid?", "Yes", 233), True)`)
}})
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def verify(message, key, signature):

    return message
```

# --solutions--

```py
def verify(message, key, signature):
    def charValue(ch) -> int:
        if 'a' <= ch <= 'z':
            return ord(ch) - ord('a') + 1
        elif 'A' <= ch <= 'Z':
            return ord(ch) - ord('A') + 27
        else:
            return 0

    def compute_sum(s):
        return sum(charValue(ch) for ch in s)

    total = compute_sum(message) + compute_sum(key)
    return total == signature
```
