---
id: 68ee9e3066cfd4eb2328e8a8
title: "تحدٍّ 89: عدّ الأوراق"
challengeType: 29
dashedName: challenge-89
---

# --description--

تحتوي مجموعة أوراق اللعب القياسية على 13 ورقة فريدة في كل نوع. معطى عدد صحيح يمثل عدد الأوراق التي ستختارها من المجموعة، أرجع عدد التوليفات الفريدة للأوراق التي يمكنك اختيارها.

- الترتيب لا يهم. اختيار الورقة A ثم الورقة B هو نفسه اختيار الورقة B ثم الورقة A.

على سبيل المثال، إذا كان العدد `52`، أرجع `1`. هناك توليفة واحدة فقط من 52 ورقة يمكن اختيارها من مجموعة أوراق مكونة من 52 ورقة. وإذا كان العدد `2`، أرجع `1326`، هناك 1326 توليفة من الأوراق يمكنك الحصول عليها عند اختيار ورقتين من المجموعة.

# --hints--

`combinations(52)` يجب أن تُرجع `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` يجب أن تُرجع `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` يجب أن تُرجع `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` يجب أن تُرجع `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` يجب أن تُرجع `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` يجب أن تُرجع `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```
