---
id: 68ee9e3066cfd4eb2328e8a6
title: "تحدٍّ 87: منشئ المصفوفة"
challengeType: 29
dashedName: challenge-87
---

# --description--

بما أن لديك عددين صحيحين (عدد الصفوف وعدد الأعمدة)، أرجع مصفوفة (مصفوفة من المصفوفات) مملوءة بالأصفار (`0`) بالحجم المطلوب.

على سبيل المثال، إذا أعطيت `2` و `3`، أرجع:

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`build_matrix(2, 3)` يجب أن تُرجع `[[0, 0, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(2, 3), [[0, 0, 0], [0, 0, 0]])`)
}})
```

`build_matrix(3, 2)` يجب أن تُرجع `[[0, 0], [0, 0], [0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(3, 2), [[0, 0], [0, 0], [0, 0]])`)
}})
```

`build_matrix(4, 3)` يجب أن تُرجع `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]])`)
}})
```

`build_matrix(9, 1)` يجب أن تُرجع `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def build_matrix(rows, cols):

    return rows
```

# --solutions--

```py
def build_matrix(rows, cols):

    return [[0 for _ in range(cols)] for _ in range(rows)]
```
