---
id: 68ee9e3066cfd4eb2328e8a7
title: "تحدٍّ 88: محدد أيام الأسبوع"
challengeType: 29
dashedName: challenge-88
---

# --description--

أعطِ تاريخًا كسلسلة نصية بالصيغة `YYYY-MM-DD`، وأرجع يوم الأسبوع.

أيام الإرجاع الصحيحة هي:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

تأكد من تجاهل المناطق الزمنية.

# --hints--

`get_weekday("2025-11-06")` يجب أن تُرجع `Thursday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2025-11-06"), "Thursday")`)
}})
```

`get_weekday("1999-12-31")` يجب أن تُرجع `Friday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("1999-12-31"), "Friday")`)
}})
```

`get_weekday("1111-11-11")` يجب أن تُرجع `Saturday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("1111-11-11"), "Saturday")`)
}})
```

`get_weekday("2112-12-21")` يجب أن تُرجع `Wednesday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2112-12-21"), "Wednesday")`)
}})
```

`get_weekday("2345-10-01")` يجب أن تُرجع `Monday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2345-10-01"), "Monday")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_weekday(date_string):

    return date_string
```

# --solutions--

```py
import datetime
def get_weekday(date_string):
    days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"]
    year, month, day = map(int, date_string.split("-"))
    date = datetime.date(year, month, day)

    return days[date.weekday() % 7 + 1 if date.weekday() != 6 else 0]
```
