---
id: 68ee9e3066cfd4eb2328e8a4
title: "تحدٍّ 85: عداد الكلمات"
challengeType: 29
dashedName: challenge-85
---

# --description--

معطاة سلسلة نصية لجملة، أرجع عدد الكلمات الموجودة في الجملة.

- الكلمات هي أي متتالية من الأحرف غير الفراغية وتُفصل بمسافة واحدة.

# --hints--

`count_words("Hello world")` يجب أن تُرجع `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("Hello world"), 2)`)
}})
```

`count_words("The quick brown fox jumps over the lazy dog.")` يجب أن تُرجع `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("The quick brown fox jumps over the lazy dog."), 9)`)
}})
```

`count_words("I like coding challenges!")` يجب أن تُرجع `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("I like coding challenges!"), 4)`)
}})
```

`count_words("Complete the challenge in JavaScript and Python.")` يجب أن تُرجع `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("Complete the challenge in JavaScript and Python."), 7)`)
}})
```

`count_words("The missing semi-colon crashed the entire internet.")` يجب أن تُرجع `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("The missing semi-colon crashed the entire internet."), 7)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_words(sentence):

    return sentence
```

# --solutions--

```py
def count_words(sentence):

    return len(sentence.split(' '))
```
