---
id: 68f6587287ad1f4ad39b0c7c
title: "تحدٍّ 90: حد الأحرف"
challengeType: 29
dashedName: challenge-90
---

# --description--

في هذا التحدٍّ، تحصل على سلسلة نصية وتحتاج إلى تحديد ما إذا كانت تناسب منشورًا على وسائل التواصل الاجتماعي. أرجع السلاسل النصية التالية بناءً على القواعد المعطاة:

- `"short post"` إذا كانت تناسب حد 40 حرفًا.
- `"long post"` إذا كانت أكثر من 40 حرفًا وتناسب حد 80 حرفًا.
- `"invalid post"` إذا كانت طويلة جدًا بحيث لا تناسب أيًا من الحدين.

# --hints--

`can_post("Hello world")` يجب أن تُرجع `"short post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("Hello world"), "short post")`)
}})
```

`can_post("This is a longer message but still under eighty characters.")` يجب أن تُرجع `"long post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This is a longer message but still under eighty characters."), "long post")`)
}})
```

`can_post("This message is too long to fit into either of the character limits for a social media post.")` يجب أن تُرجع `"invalid post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This message is too long to fit into either of the character limits for a social media post."), "invalid post")`)
}})
```

# --seed--

## --seed-contents--

```py
def can_post(message):

    return message
```

# --solutions--

```py
def can_post(message):
    if len(message) <= 40:
        return "short post"
    elif len(message) <= 80:
        return "long post"
    else:
        return "invalid post"
```
