---
id: 68f6587287ad1f4ad39b0c82
title: "تحدٍّ 96: هل هو عطلة نهاية الأسبوع؟"
challengeType: 29
dashedName: challenge-96
---

# --description--

أعطِ تاريخًا بصيغة `"YYYY-MM-DD"`، وأرجع عدد الأيام المتبقية حتى عطلة نهاية الأسبوع.

- يبدأ عطلة نهاية الأسبوع يوم السبت.
- إذا كان التاريخ المعطى يوم السبت أو الأحد، أعد `"It's the weekend!"`.
- وإلا، أعد `"X days until the weekend."`، حيث `X` هو عدد الأيام حتى يوم السبت.
- إذا كان `X` يساوي `1`، استخدم `"day"` (مفرد) بدلاً من `"days"` (جمع).
- تأكد من أن الحساب يتجاهل المنطقة الزمنية المحلية.

# --hints--

`days_until_weekend("2025-11-14")` يجب أن تُرجع `"1 day until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-11-14"), "1 day until the weekend.")`)
}})
```

`days_until_weekend("2025-01-01")` يجب أن تُرجع `"3 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-01-01"), "3 days until the weekend.")`)
}})
```

`days_until_weekend("2025-12-06")` يجب أن تُرجع `"It's the weekend!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-12-06"), "It's the weekend!")`)
}})
```

`days_until_weekend("2026-01-27")` يجب أن تُرجع `"4 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-01-27"), "4 days until the weekend.")`)
}})
```

`days_until_weekend("2026-09-07")` يجب أن تُرجع `"5 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-09-07"), "5 days until the weekend.")`)
}})
```

`days_until_weekend("2026-11-29")` يجب أن تُرجع `"It's the weekend!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-11-29"), "It's the weekend!")`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_weekend(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime

def days_until_weekend(date_string):
    date = datetime.strptime(date_string, "%Y-%m-%d").date()
    day_of_week = date.weekday()

    if day_of_week == 5 or day_of_week == 6:
        return "It's the weekend!"

    days_until_saturday = 5 - day_of_week
    return f"{days_until_saturday} day{'s' if days_until_saturday != 1 else ''} until the weekend."
```
