---
id: 68f6587287ad1f4ad39b0c80
title: "تحدٍّ 94: مولد توقيع البريد الإلكتروني"
challengeType: 29
dashedName: challenge-94
---

# --description--

بالنظر إلى سلاسل نصية لاسم شخص، واللقب، والشركة، أرجع توقيع بريد إلكتروني كسلسلة نصية واحدة باستخدام القواعد التالية:

- يجب أن يظهر الاسم أولًا، مسبوقًا ببادئة تعتمد على الحرف الأول من الاسم. للأسماء التي تبدأ بـ (بغض النظر عن حالة الحروف):
  - `A-I`: استخدم `>>` كبادئة.
  - `J-R`: استخدم `--` كبادئة.
  - `S-Z`: استخدم `::` كبادئة.
- يجب أن يتبع الاسم فاصلة ومسافة (`, `).
- يجب أن يتبع اللقب والشركة الفاصلة والمسافة، مفصولين بـ `" at "` (مع مسافات حولها).

على سبيل المثال، إذا كانت القيم المعطاة هي `"Quinn Waverly"`، و`"Founder and CEO"`، و`"TechCo"`، فأرجع `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generate_signature("Quinn Waverly", "Founder and CEO", "TechCo")` يجب أن تُرجع `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo")`)
}})
```

`generate_signature("Alice Reed", "Engineer", "TechCo")` يجب أن تُرجع `">>Alice Reed, Engineer at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo")`)
}})
```

`generate_signature("Tina Vaughn", "Developer", "example.com")` يجب أن تُرجع `"::Tina Vaughn, Developer at example.com"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com")`)
}})
```

`generate_signature("B. B.", "Product Tester", "AcmeCorp")` يجب أن تُرجع `">>B. B., Product Tester at AcmeCorp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp")`)
}})
```

`generate_signature("windstorm", "Cloud Architect", "Atmospheronics")` يجب أن تُرجع `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_signature(name, title, company):

    return name
```

# --solutions--

```py
def generate_signature(name, title, company):
    first_letter = name[0].upper()
    if first_letter in "ABCDEFGHI":
        prefix = ">>"
    elif first_letter in "JKLMNOPQR":
        prefix = "--"
    else:
        prefix = "::"

    return f"{prefix}{name}, {title} at {company}"
```
