---
id: 68f6587287ad1f4ad39b0c7e
title: "تحدٍّ 92: مستخرج الامتداد"
challengeType: 29
dashedName: challenge-92
---

# --description--

أعطِ سلسلة نصية تمثل اسم ملف، وأرجع امتداد الملف.

- الامتداد هو الجزء من اسم الملف الذي يأتي بعد آخر نقطة (`.`).
- إذا لم يحتوي اسم الملف على نقطة أو انتهى بنقطة، أرجع `"none"`.
- يجب إرجاع الامتداد كما هو، مع الحفاظ على حالة الأحرف.

# --hints--

`get_extension("document.txt")` يجب أن تُرجع `"txt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("document.txt"), "txt")`)
}})
```

`get_extension("README")` يجب أن تُرجع `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("README"), "none")`)
}})
```

`get_extension("image.PNG")` يجب أن تُرجع `"PNG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("image.PNG"), "PNG")`)
}})
```

`get_extension(".gitignore")` يجب أن تُرجع `"gitignore"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension(".gitignore"), "gitignore")`)
}})
```

`get_extension("archive.tar.gz")` يجب أن تُرجع `"gz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("archive.tar.gz"), "gz")`)
}})
```

`get_extension("final.draft.")` يجب أن تُرجع `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("final.draft."), "none")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_extension(filename):

    return filename
```

# --solutions--

```py
def get_extension(filename):
    last_dot = filename.rfind('.')
    
    if last_dot == -1 or last_dot == len(filename) - 1:
        return "none"
    
    return filename[last_dot + 1:]
```
