---
id: 68f6587287ad1f4ad39b0c84
title: "تحدٍّ 98: عد المستطيلات"
challengeType: 29
dashedName: challenge-98
---

# --description--

معطى عددان صحيحان موجبان يمثلان عرض وارتفاع مستطيل، حدّد عدد المستطيلات التي يمكن أن تتسع داخل المستطيل المعطى.

- عد فقط المستطيلات التي يكون عرضها وارتفاعها أعدادًا صحيحة.

على سبيل المثال، إذا كان المعطى `1` و`3`، أرجع `6`. ثلاثة مستطيلات 1x1، مستطيلان 1x2، ومستطيل واحد 1x3.

# --hints--

`count_rectangles(1, 3)` يجب أن تُرجع `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 3), 6)`)
}})
```

`count_rectangles(3, 2)` يجب أن تُرجع `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(3, 2), 18)`)
}})
```

`count_rectangles(1, 2)` يجب أن تُرجع `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 2), 3)`)
}})
```

`count_rectangles(5, 4)` يجب أن تُرجع `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(5, 4), 150)`)
}})
```

`count_rectangles(11, 19)` يجب أن تُرجع `12540`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(11, 19), 12540)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_rectangles(width, height):

    return width
```

# --solutions--

```py
def count_rectangles(width, height):
    total = 0
    for w in range(1, width + 1):
        for h in range(1, height + 1):
            total += (width - w + 1) * (height - h + 1)
    return total
```
