---
id: 68f6587287ad1f4ad39b0c7f
title: "تحدٍّ 93: الحروف المتحركة والحروف الساكنة"
challengeType: 29
dashedName: challenge-93
---

# --description--

أعطِ سلسلة نصية، وأرجع مصفوفة تحتوي على عدد الحروف المتحركة وعدد الحروف الساكنة في السلسلة.

- الحروف المتحركة تتكون من `a`، `e`، `i`، `o`، `u` بأي حالة.
- الحروف الساكنة تتكون من جميع الحروف الأخرى بأي حالة.
- تجاهل أي حروف غير الحروف الأبجدية.

على سبيل المثال، عند إعطاء `"Hello World"`، أرجع `[3, 7]`.

# --hints--

`count("Hello World")` يجب أن تُرجع `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("JavaScript")` يجب أن تُرجع `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("JavaScript"), [3, 7])`)
}})
```

`count("Python")` يجب أن تُرجع `[1, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Python"), [1, 5])`)
}})
```

`count("freeCodeCamp")` يجب أن تُرجع `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("freeCodeCamp"), [5, 7])`)
}})
```

`count("Hello, World!")` يجب أن تُرجع `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("The quick brown fox jumps over the lazy dog.")` يجب أن تُرجع `[11, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

# --seed--

## --seed-contents--

```py
def count(s):

    return s
```

# --solutions--

```py
def count(s):
    vowels = "aeiou"
    consonants = "bcdfghjklmnpqrstvwxyz"
    v = 0
    c = 0

    for char in s:
        ch = char.lower()
        if ch in vowels:
            v += 1
        if ch in consonants:
            c += 1

    return [v, c]
```
