---
id: 68f6587287ad1f4ad39b0c7d
title: "تحدٍّ 91: البحث عن كلمة"
challengeType: 29
dashedName: challenge-91
---

# --description--

معطى مصفوفة (مصفوفة من المصفوفات) من أحرف مفردة وكلمة للعثور عليها، أرجع مؤشرات البداية والنهاية للكلمة في المصفوفة.

- ستكون المصفوفة المعطاة مملوءة بجميع الأحرف الصغيرة (`a-z`).
- ستكون الكلمة التي تبحث عنها موجودة في المصفوفة مرة واحدة بالضبط.
- ستكون الكلمة التي تبحث عنها دائمًا على خط مستقيم في أحد الاتجاهات التالية:
  - من اليسار إلى اليمين
  - من اليمين إلى اليسار
  - من الأعلى إلى الأسفل
  - من الأسفل إلى الأعلى

على سبيل المثال، معطى المصفوفة:

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

والكلمة `"cat"`، أرجع:

```md
[[0, 1], [2, 1]]
```

حيث أن `[0, 1]` هي مؤشرات الحرف `"c"` (بداية الكلمة)، و`[2, 1]` هي مؤشرات الحرف `"t"` (نهاية الكلمة).

# --hints--

`find_word([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")` يجب أن تُرجع `[[0, 1], [2, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_word([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]])`)
}})
```

`find_word([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")` يجب أن تُرجع `[[0, 0], [0, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_word([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]])`)
}})
```

`find_word([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")` يجب أن تُرجع `[[3, 3], [0, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_word([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]])`)
}})
```

`find_word([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")` يجب أن تُرجع `[[1, 3], [1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_word([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_word(matrix, word):

    return matrix
```

# --solutions--

```py
def find_word(matrix, word):
    rows = len(matrix)
    cols = len(matrix[0])
    length = len(word)

    directions = [
        (0, 1),
        (0, -1),
        (1, 0),
        (-1, 0)
    ]

    for r in range(rows):
        for c in range(cols):
            for dr, dc in directions:
                match = True
                for i in range(length):
                    nr = r + dr * i
                    nc = c + dc * i
                    if nr < 0 or nr >= rows or nc < 0 or nc >= cols or matrix[nr][nc] != word[i]:
                        match = False
                        break
                if match:
                    return [[r, c], [r + dr * (length - 1), c + dc * (length - 1)]]
```
