---
id: 68f6587287ad1f4ad39b0c81
title: "تحدٍّ 95: إزاحة المصفوفة"
challengeType: 29
dashedName: challenge-95
---

# --description--

أعطِ مصفوفة وعددًا صحيحًا يمثل عدد المواقع التي تُزاح بها المصفوفة، أرجع المصفوفة بعد الإزاحة.

- العدد الصحيح الموجب يُزاح بالمصفوفة إلى اليسار.
- العدد الصحيح السالب يُزاح بالمصفوفة إلى اليمين.
- الإزاحة تدور حول المصفوفة.

على سبيل المثال، إذا كانت المصفوفة `[1, 2, 3]` والعدد `1`، ازح المصفوفة موقعًا واحدًا إلى اليسار، فترجع `[2, 3, 1]`.

# --hints--

`shift_array([1, 2, 3], 1)` يجب أن تُرجع `[2, 3, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([1, 2, 3], 1), [2, 3, 1])`)
}})
```

`shift_array([1, 2, 3], -1)` يجب أن تُرجع `[3, 1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([1, 2, 3], -1), [3, 1, 2])`)
}})
```

`shift_array(["alpha", "bravo", "charlie"], 5)` يجب أن تُرجع `["charlie", "alpha", "bravo"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"])`)
}})
```

`shift_array(["alpha", "bravo", "charlie"], -11)` يجب أن تُرجع `["bravo", "charlie", "alpha"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"])`)
}})
```

`shift_array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` يجب أن تُرجع `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4])`)
}})
```

# --seed--

## --seed-contents--

```py
def shift_array(arr, n):

    return arr
```

# --solutions--

```py
def shift_array(arr, n):
    length = len(arr)
    n = n % length
    if n < 0:
        n += length

    return arr[n:] + arr[:n]
```
