---
id: 68ffb91507a5b645769328c6
title: "تحدٍّ 103: المضاعف المشترك الأصغر"
challengeType: 29
dashedName: challenge-103
---

# --description--

بالنظر إلى عددين صحيحين، أعد المضاعف المشترك الأصغر (LCM) لهما.

المضاعف المشترك الأصغر لعددين هو أصغر عدد صحيح موجب يكون مضاعفًا لكلا العددين. على سبيل المثال، إذا كان العددان `4` و `6`، أعد `12` لأن:

- مضاعفات `4` هي `4`، `8`، `12` وهكذا.
- مضاعفات `6` هي `6`، `12`، `18` وهكذا.
- `12` هو أصغر عدد مضاعف لكليهما.

# --hints--

`lcm(4, 6)` يجب أن يعيد `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(4, 6), 12)`)
}})
```

`lcm(9, 6)` يجب أن يعيد `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(9, 6), 18)`)
}})
```

`lcm(10, 100)` يجب أن يعيد `100`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(10, 100), 100)`)
}})
```

`lcm(13, 17)` يجب أن يعيد `221`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(13, 17), 221)`)
}})
```

`lcm(45, 70)` يجب أن يعيد `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(45, 70), 630)`)
}})
```

# --seed--

## --seed-contents--

```py
def lcm(a, b):

    return a
```

# --solutions--

```py
def lcm(a, b):
    def gcd(x, y):
        return x if y == 0 else gcd(y, x % y)

    return abs(a * b) // gcd(a, b)
```
