---
id: 68ffb91507a5b645769328c5
title: "تحدٍّ 102: أطول كلمة"
challengeType: 29
dashedName: challenge-102
---

# --description--

أعطِ سلسلة نصية تمثل جملة، وأرجع أطول كلمة في الجملة.

- الكلمات مفصولة بمسافة واحدة.
- فقط الحروف (`a-z`، غير حساسة لحالة الحروف) تُحتسب في طول الكلمة.
- إذا وُجدت كلمات متعددة بنفس الطول، أرجع أول كلمة تظهر.
- أرجع الكلمة كما تظهر في السلسلة النصية المعطاة، مع إزالة علامات الترقيم.

# --hints--

`longest_word("The quick red fox")` يجب أن تُرجع `"quick"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("The quick red fox"), "quick")`)
}})
```

`longest_word("Hello coding challenge.")` يجب أن تُرجع `"challenge"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Hello coding challenge."), "challenge")`)
}})
```

`longest_word("Do Try This At Home.")` يجب أن تُرجع `"This"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Do Try This At Home."), "This")`)
}})
```

`longest_word("This sentence... has commas, ellipses, and an exclamation point!")` يجب أن تُرجع `"exclamation"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation")`)
}})
```

`longest_word("A tie? No way!")` يجب أن تُرجع `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("A tie? No way!"), "tie")`)
}})
```

`longest_word("Wouldn't you like to know.")` يجب أن تُرجع `"Wouldnt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Wouldn't you like to know."), "Wouldnt")`)
}})
```

# --seed--

## --seed-contents--

```py
def longest_word(sentence):

    return sentence
```

# --solutions--

```py
import re
def longest_word(sentence):
    words = sentence.split(" ")
    longest = ""

    for word in words:
        cleaned = re.sub(r'[^a-zA-Z]', '', word)
        if len(cleaned) > len(longest):
            longest = cleaned

    return longest
```
