---
id: 68ffb91507a5b645769328cc
title: "تحدٍّ 109: كم عمري مرة أخرى؟"
challengeType: 29
dashedName: challenge-109
---

# --description--

بالنظر إلى تاريخ ميلاد شخص ما بصيغة `YYYY-MM-DD`، أرجع عمر الشخص حتى 27 نوفمبر 2025.

- افترض أن جميع تواريخ الميلاد صحيحة وتقع قبل 27 نوفمبر 2025.
- أرجع العمر كعدد صحيح.
- تأكد من احتساب ما إذا كان الشخص قد احتفل بعيد ميلاده في 2025 أم لا.

# --hints--

`calculate_age("2000-11-20")` يجب أن تُرجع `25`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-11-20"), 25)`)
}})
```

`calculate_age("2000-12-01")` يجب أن تُرجع `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-12-01"), 24)`)
}})
```

`calculate_age("2014-10-25")` يجب أن تُرجع `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2014-10-25"), 11)`)
}})
```

`calculate_age("1994-01-06")` يجب أن تُرجع `31`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-01-06"), 31)`)
}})
```

`calculate_age("1994-12-14")` يجب أن تُرجع `30`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-12-14"), 30)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_age(birthday):

    return birthday
```

# --solutions--

```py
from datetime import datetime
def calculate_age(birthday):
    today = datetime(2025, 11, 27)
    birth_date = datetime.strptime(birthday, "%Y-%m-%d")

    age = today.year - birth_date.year

    if (today.month, today.day) < (birth_date.month, birth_date.day):
        age -= 1

    return age
```
