---
id: 69162d64f96574d9bb629efb
title: "تحدٍّ 110: خمن الكلمة"
challengeType: 29
dashedName: challenge-110
---

# --description--

معطيتان سلسلتا نصية بنفس الطول، كلمة سرية وتخمين، قارن التخمين بالكلمة السرية باستخدام القواعد التالية:

- تتكون الكلمة السرية والتخمين فقط من حروف كبيرة (`"A"` إلى `"Z"`);
- لكل حرف في التخمين، استبدله برقم حسب مدى مطابقته للكلمة السرية:
  - `"2"` إذا كان الحرف في الكلمة السرية وفي الموضع الصحيح.
  - `"1"` إذا كان الحرف في الكلمة السرية لكنه في موضع خاطئ.
  - `"0"` إذا لم يكن الحرف في الكلمة السرية.
- يمكن استخدام كل حرف في الكلمة السرية مرة واحدة على الأكثر.
- تُعيَّن التطابقات الدقيقة (`"2"`) أولًا، ثم تُعيَّن التطابقات الجزئية (`"1"`) من اليسار إلى اليمين للحروف المتبقية.
- إذا تكرر الحرف في التخمين، يمكن أن يتطابق فقط بعدد مرات ظهوره في الكلمة السرية.

على سبيل المثال، مع كلمة سرية `"APPLE"` وتخمين `"POPPA"`، أرجع `"10201"`:

الحرف الأول `"P"` ليس في الموضع الصحيح (`"1"`)، الحرف `"O"` غير موجود في الكلمة السرية (`"0"`)، الحرف الثاني `"P"` في الموضع الصحيح (`"2"`)، الحرف الثالث `"P"` صفر (`"0"`) لأن الحرفين `"P"` في الكلمة السرية قد استُخدما، والحرف `"A"` ليس في الموضع الصحيح (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` يجب أن تُرجع `"10201"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("APPLE", "POPPA"), "10201")`)
}})
```

`compare("REACT", "TRACE")` يجب أن تُرجع `"11221"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("REACT", "TRACE"), "11221")`)
}})
```

`compare("DEBUGS", "PYTHON")` يجب أن تُرجع `"000000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("DEBUGS", "PYTHON"), "000000")`)
}})
```

`compare("JAVASCRIPT", "TYPESCRIPT")` يجب أن تُرجع `"0000222222"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222")`)
}})
```

`compare("ORANGE", "ROUNDS")` يجب أن تُرجع `"110200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("ORANGE", "ROUNDS"), "110200")`)
}})
```

`compare("WIRELESS", "ETHERNET")` يجب أن تُرجع `"10021000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("WIRELESS", "ETHERNET"), "10021000")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare(word, guess):

    return word
```

# --solutions--

```py
def compare(word, guess):
    n = len(word)
    result = ["0"] * n
    used_secret = [False] * n

    for i in range(n):
        if guess[i] == word[i]:
            result[i] = "2"
            used_secret[i] = True

    for i in range(n):
        if result[i] == "0":
            for j in range(n):
                if not used_secret[j] and guess[i] == word[j]:
                    result[i] = "1"
                    used_secret[j] = True
                    break

    return "".join(result)
```
