---
id: 69162d64f96574d9bb629efd
title: "تحدٍّ 112: كاشف الذكاء الاصطناعي"
challengeType: 29
dashedName: challenge-112
---

# --description--

التحدٍّ اليوم مستوحى من إصدار ChatGPT في 30 نوفمبر 2022.

معطى سلسلة نصية تحتوي على جملة واحدة أو أكثر، حدّد ما إذا كان من المحتمل أن تكون مولدة بواسطة الذكاء الاصطناعي باستخدام القواعد التالية:

- تحتوي على شرطتين أو أكثر (`-`).
- تحتوي على مجموعتين أو أكثر من الأقواس (`()`). يمكن أن يكون هناك نص داخل الأقواس.
- تحتوي على ثلاث كلمات أو أكثر تحتوي على 7 أحرف أو أكثر.

- الكلمات مفصولة بمسافة واحدة فقط وتتكون من حروف فقط (`A-Z`). لا تشمل علامات الترقيم أو غير الحروف كجزء من الكلمة.

إذا استوفت الجملة المعطاة أيًا من القواعد أعلاه، أرجع `"AI"`، وإلا أرجع `"Human"`.

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` يجب أن تُرجع `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` يجب أن تُرجع `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` يجب أن تُرجع `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` يجب أن تُرجع `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` يجب أن تُرجع `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```
