---
id: 69162d64f96574d9bb629f04
title: "تحدٍّ 119: ضغط السلسلة النصية"
challengeType: 29
dashedName: challenge-119
---

# --description--

بالنظر إلى جملة سلسلة نصية، أعد نسخة مضغوطة من الجملة حيث تُستبدَل الكلمات المتكررة المتتالية بالكلمة متبوعة بعدد مرات تكرارها داخل أقواس عادية.

- تُضغط فقط التكرارات المتتالية.
- تُفصل الكلمات بمسافات مفردة.

على سبيل المثال، إذا كانت السلسلة النصية هي `"yes yes yes please"`، فأعد `"yes(3) please"`.

# --hints--

`compress_string("yes yes yes please")` يجب أن تُعيد `"yes(3) please"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("yes yes yes please"), "yes(3) please")`)
}})
```

`compress_string("I have have have apples")` يجب أن تُعيد `"I have(3) apples"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("I have have have apples"), "I have(3) apples")`)
}})
```

`compress_string("one one three and to the the the the")` يجب أن تُعيد `"one(2) three and to the(4)"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("one one three and to the the the the"), "one(2) three and to the(4)")`)
}})
```

`compress_string("route route route route route route tee tee tee tee tee tee")` يجب أن تُعيد `"route(6) tee(6)"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)")`)
}})
```

# --seed--

## --seed-contents--

```py
def compress_string(sentence):

    return sentence
```

# --solutions--

```py
def compress_string(sentence):
    words = sentence.split(" ")
    result = []
    count = 1

    for i in range(len(words)):
        if i < len(words) - 1 and words[i] == words[i + 1]:
            count += 1
        else:
            if count > 1:
                result.append(f"{words[i]}({count})")
            else:
                result.append(words[i])
            count = 1

    return " ".join(result)
```
