---
id: 69162d64f96574d9bb629f01
title: "تحدٍّ 116: عدّ التباديل"
challengeType: 29
dashedName: challenge-116
---

# --description--

بالنظر إلى سلسلة نصية، أرجع عدد التباديل المميزة التي يمكن تكوينها من حروفها.

- التبديل هو أي إعادة ترتيب لحروف السلسلة النصية.
- لا تحسب التباديل المكررة.
- إذا احتوت السلسلة على حروف مكررة، يجب حساب الترتيبات المكررة مرة واحدة فقط.
- السلسلة ستحتوي على حروف فقط (`A-Z`، `a-z`).

على سبيل المثال، عند إعطاء `"abb"`، أرجع `3` لأن هناك ثلاث طرق فريدة لترتيب الحروف: `"abb"`، `"bab"`، و`"bba"`.

# --hints--

`count_permutations("abb")` يجب أن تُرجع `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` يجب أن تُرجع `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` يجب أن تُرجع `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` يجب أن تُرجع `39916800`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
