---
id: 69162d64f96574d9bb629f03
title: "تحدٍّ 118: منسق التاريخ"
challengeType: 29
dashedName: challenge-118
---

# --description--

بالنظر إلى تاريخ بالصيغة `"Month day, year"`، أرجع التاريخ بالصيغة `"YYYY-MM-DD"`.

- سيكون الشهر المعطى هو الاسم الكامل للشهر بالإنجليزية. على سبيل المثال: `"January"`، `"February"`، وهكذا.
- في قيمة الإرجاع، أضف أصفارًا بادئة للشهر واليوم إذا لزم الأمر لضمان وجود رقمين.

على سبيل المثال، إذا كان التاريخ `"December 6, 2025"`، أرجع `"2025-12-06"`.

# --hints--

`format_date("December 6, 2025")` يجب أن تُرجع `"2025-12-06"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("December 6, 2025"), "2025-12-06")`)
}})
```

`format_date("January 1, 2000")` يجب أن تُرجع `"2000-01-01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("January 1, 2000"), "2000-01-01")`)
}})
```

`format_date("November 11, 1111")` يجب أن تُرجع `"1111-11-11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("November 11, 1111"), "1111-11-11")`)
}})
```

`format_date("September 7, 512")` يجب أن تُرجع `"512-09-07"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("September 7, 512"), "512-09-07")`)
}})
```

`format_date("May 4, 1950")` يجب أن تُرجع `"1950-05-04"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("May 4, 1950"), "1950-05-04")`)
}})
```

`format_date("February 29, 1992")` يجب أن تُرجع `"1992-02-29"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("February 29, 1992"), "1992-02-29")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_date(date_string):

    return date_string
```

# --solutions--

```py
def format_date(date_string):
    months = {
        "January": "01",
        "February": "02",
        "March": "03",
        "April": "04",
        "May": "05",
        "June": "06",
        "July": "07",
        "August": "08",
        "September": "09",
        "October": "10",
        "November": "11",
        "December": "12"
    }

    month, day, year = date_string.replace(",", "").split(" ")
    day = day.zfill(2)

    return f"{year}-{months[month]}-{day}"
```
