---
id: 69162d64f96574d9bb629f02
title: "تحدٍّ 117: الفرق المتماثل"
challengeType: 29
dashedName: challenge-117
---

# --description--

بالنظر إلى مصفوفتين، أعد مصفوفة جديدة تحتوي على الفرق المتماثل بينهما.

- الفرق المتماثل بين مجموعتين هو مجموعة القيم التي تظهر في إحدى المجموعتين فقط، وليس في كلتيهما.
- أعد القيم بالترتيب الذي تظهر به أول مرة في المصفوفات المُدخَلة.

# --hints--

`difference([1, 2, 3], [3, 4, 5])` يجب أن تُعيد `[1, 2, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5])`)
}})
```

`difference(["a", "b"], ["c", "b"])` يجب أن تُعيد `["a", "c"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"])`)
}})
```

`difference([1, "a", 2], [2, "b", "a"])` يجب أن تُعيد `[1, "b"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"])`)
}})
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` يجب أن تُعيد `[2, 4, 6, 8]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8])`)
}})
```

# --seed--

## --seed-contents--

```py
def difference(arr1, arr2):

    return arr1
```

# --solutions--

```py
def difference(arr1, arr2):
    diff = []

    for v in arr1:
        if v not in arr2 and v not in diff:
            diff.append(v)

    for v in arr2:
        if v not in arr1 and v not in diff:
            diff.append(v)

    return diff
```
