---
id: 69162d64f96574d9bb629eff
title: "تحدٍّ 114: من Camel إلى Snake"
challengeType: 29
dashedName: challenge-114
---

# --description--

معطاة سلسلة نصية بصيغة camel case، أعد النسخة بصيغة snake case باستخدام القواعد التالية:

- ستحتوي السلسلة النصية المُدخَلة على أحرف فقط (`A-Z` و `a-z`) وستبدأ دائمًا بحرف صغير.
- كل حرف كبير في سلسلة camel case يبدأ كلمة جديدة.
- حوّل كل الأحرف إلى أحرف صغيرة.
- افصل الكلمات بشرطة سفلية (`_`).

# --hints--

`to_snake("helloWorld")` يجب أن تُعيد `"hello_world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` يجب أن تُعيد `"my_variable_name"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` يجب أن تُعيد `"freecodecamp_daily_challenges"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
