---
id: 69162d64f96574d9bb629f00
title: "تحدٍّ 115: محول عنصر قائمة مرتبة في Markdown"
challengeType: 29
dashedName: challenge-115
---

# --description--

أعطِ سلسلة نصية تمثل عنصر قائمة مرتبة في Markdown، وأرجع السلسلة المكافئة بصيغة HTML.

يجب أن يكون عنصر القائمة المرتبة الصحيح في Markdown:

- يبدأ بصفر أو أكثر من الفراغات، يتبعها
- رقم (1 أو أكبر) ونقطة (`.`)، يتبعها
- فراغ واحد على الأقل، ثم
- نص عنصر القائمة.

إذا لم تكن السلسلة النصية بالصيغة الدقيقة أعلاه، أرجع `"Invalid format"`. وإلا، لف نص عنصر القائمة داخل وسم `li` وأرجع السلسلة.

على سبيل المثال، إذا كانت السلسلة `"1. My item"`، أرجع `"<li>My item</li>"`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML في السلاسل النصية عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`convert_list_item("1. My item")` يجب أن تُرجع `"<li>My item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` يجب أن تُرجع `"<li>Another item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` يجب أن تُرجع `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` يجب أن تُرجع `"<li>list item text</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` يجب أن تُرجع `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` يجب أن تُرجع `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```
