---
id: 69162d64f96574d9bb629efe
title: "تحدٍّ 113: الأميال إلى الكيلومترات"
challengeType: 29
dashedName: challenge-113
---

# --description--

بالنظر إلى مسافة بالأميال على شكل رقم، أعد المسافة المكافئة بالكيلومترات.

- سيكون المُدخَل دائمًا رقمًا غير سالب.
- تساوي الميل الواحد 1.60934 كيلومترًا.
- قرب النتيجة إلى منزلتين عشريتين.
- أزل الأصفار الزائدة غير الضرورية من النتيجة المقربة.

# --hints--

يجب أن تُعيد الدالة `convert_to_km(1)` القيمة `1.61`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(1), 1.61)`)
}})
```

يجب أن تُعيد الدالة `convert_to_km(21)` القيمة `33.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(21), 33.8)`)
}})
```

يجب أن تُعيد الدالة `convert_to_km(3.5)` القيمة `5.63`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(3.5), 5.63)`)
}})
```

يجب أن تُعيد الدالة `convert_to_km(0)` القيمة `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0), 0)`)
}})
```

يجب أن تُعيد الدالة `convert_to_km(0.621371)` القيمة `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0.621371), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_km(miles):

    return miles
```

# --solutions--

```py
def convert_to_km(miles):
    km = miles * 1.60934
    return round(km, 2)
```
