---
id: 691b559495c5cb5a37b9b487
title: "تحدٍّ 127: فحص السرعة"
challengeType: 29
dashedName: challenge-127
---

# --description--

بما أن السرعة التي تسير بها بالميل في الساعة (MPH) وحدود السرعة بالكيلومتر في الساعة (KPH)، حدّد ما إذا كنت تتجاوز السرعة وإذا كنت ستحصل على تحذير أو مخالفة.

- الميل الواحد يساوي 1.60934 كيلومتر.
- إذا كنت تسير بسرعة أقل من أو تساوي حد السرعة، أرجع `"Not Speeding"`.
- إذا كنت تسير بسرعة تزيد 5 كيلومتر في الساعة أو أقل فوق حد السرعة، أرجع `"Warning"`.
- إذا كنت تسير بسرعة تزيد أكثر من 5 كيلومتر في الساعة فوق حد السرعة، أرجع `"Ticket"`.

# --hints--

`speed_check(30, 70)` يجب أن تُرجع `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` يجب أن تُرجع `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` يجب أن تُرجع `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` يجب أن تُرجع `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` يجب أن تُرجع `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` يجب أن تُرجع `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```
