---
id: 691b559495c5cb5a37b9b480
title: "تحدٍّ 120: من باوند إلى كيلوجرام"
challengeType: 29
dashedName: challenge-120
---

# --description--

معطى وزن بوحدة الباوند كرقم، أرجع السلسلة النصية `"(lbs) pounds equals (kgs) kilograms."`.

- استبدل `"(lbs)"` بالرقم المُدخَل.
- استبدل `"(kgs)"` بالقيمة المحولة إلى كيلوجرام، مقربة إلى منزلتين عشريتين مع تضمين منزلتين عشريتين دائمًا في القيمة.
- 1 باوند يساوي 0.453592 كيلوجرام.
- إذا كان المُدخَل `1`، استخدم `"pound"` بدلًا من `"pounds"`.
- إذا كانت القيمة المحولة `1`، استخدم `"kilogram"` بدلًا من `"kilograms"`.

# --hints--

`convert_to_kgs(1)` يجب أن تُرجع `"1 pound equals 0.45 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(1), "1 pound equals 0.45 kilograms.")`)
}})
```

`convert_to_kgs(0)` يجب أن تُرجع `"0 pounds equals 0.00 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(0), "0 pounds equals 0.00 kilograms.")`)
}})
```

`convert_to_kgs(100)` يجب أن تُرجع `"100 pounds equals 45.36 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(100), "100 pounds equals 45.36 kilograms.")`)
}})
```

`convert_to_kgs(2.5)` يجب أن تُرجع `"2.5 pounds equals 1.13 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.5), "2.5 pounds equals 1.13 kilograms.")`)
}})
```

`convert_to_kgs(2.20462)` يجب أن تُرجع `"2.20462 pounds equals 1.00 kilogram."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.20462), "2.20462 pounds equals 1.00 kilogram.")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_kgs(lbs):

    return lbs
```

# --solutions--

```py
def convert_to_kgs(lbs):
    KG_PER_POUND = 0.453592
    kgs = f"{lbs * KG_PER_POUND:.2f}"

    pound_word = "pound" if lbs == 1 else "pounds"
    kilogram_word = "kilogram" if kgs == "1.00" else "kilograms"

    return f"{lbs} {pound_word} equals {kgs} {kilogram_word}."
```
