---
id: 691b5597f0f3e85a588a5c94
title: "تحدٍّ 130: لوحة الشطرنج"
challengeType: 29
dashedName: challenge-130
---

# --description--

معطى مصفوفة تحتوي على رقمين، الأول هو عدد الصفوف والثاني هو عدد الأعمدة، أرجع مصفوفة (مصفوفة من المصفوفات) مملوءة بحروف `"X"` و `"O"` بالحجم المعطى.

- يجب أن تتناوب الحروف مثل لوحة الشطرنج.
- يجب أن تكون الخانة في أعلى اليسار دائمًا `"X"`.

على سبيل المثال، إذا كانت المصفوفة `[3, 3]`، أرجع:

```json
[
  ["X", "O", "X"],
  ["O", "X", "O"],
  ["X", "O", "X"]
]
```

# --hints--

`create_board([3, 3])` يجب أن تُرجع `[["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(create_board([3, 3]), [["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]])`)
}})
```

`create_board([6, 1])` يجب أن تُرجع `[["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(create_board([6, 1]), [["X"], ["O"], ["X"], ["O"], ["X"], ["O"]])`)
}})
```

`create_board([2, 10])` يجب أن تُرجع `[["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(create_board([2, 10]), [["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]])`)
}})
```

`create_board([5, 4])` يجب أن تُرجع `[["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(create_board([5, 4]), [["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def create_board(dimensions):

    return dimensions
```

# --solutions--

```py
def create_board(dimensions):
    rows, cols = dimensions
    board = []

    for r in range(rows):
        row = []
        for c in range(cols):
            row.append("X" if (r + c) % 2 == 0 else "O")
        board.append(row)

    return board
```
