---
id: 691b559495c5cb5a37b9b483
title: "تحدٍّ 123: منشئ الأعداد الرومانية"
challengeType: 29
dashedName: challenge-123
---

# --description--

أعطِ عددًا صحيحًا، أرجع قيمته المكافئة بالأعداد الرومانية.

تستخدم الأعداد الرومانية هذه الرموز:

| الرمز | القيمة |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

تُكتب الأعداد الرومانية من الأكبر إلى الأصغر، من اليسار إلى اليمين باستخدام القواعد التالية:

- تُستخدم الإضافة عندما يتبع رمز رمز آخر له نفس القيمة أو قيمة أصغر. على سبيل المثال، يُكتب `18` كـ `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- تُستخدم الطرح عندما يظهر رمز أصغر قبل رمز أكبر، لتمثيل 4 أو 9 في أي منزلة. على سبيل المثال، يُكتب 19 كـ `XIX` (10 + (10 - 1)).
- لا يجوز تكرار أي رمز أكثر من ثلاث مرات متتالية. يُستخدم الطرح عندما تحتاج إلى كتابة رمز أكثر من ثلاث مرات متتالية. لذلك أكبر رقم يمكنك كتابته هو 3999.

إليك مثالًا آخر: إذا أعطيت `1464`، أرجع `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`to_roman(18)` يجب أن تُرجع `"XVIII"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(18), "XVIII")`)
}})
```

`to_roman(19)` يجب أن تُرجع `"XIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(19), "XIX")`)
}})
```

`to_roman(1464)` يجب أن تُرجع `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(1464), "MCDLXIV")`)
}})
```

`to_roman(2025)` يجب أن تُرجع `"MMXXV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(2025), "MMXXV")`)
}})
```

`to_roman(3999)` يجب أن تُرجع `"MMMCMXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(3999), "MMMCMXCIX")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_roman(num):

    return num
```

# --solutions--

```py
def to_roman(num):
    symbols = [
        ("M", 1000),
        ("CM", 900),
        ("D", 500),
        ("CD", 400),
        ("C", 100),
        ("XC", 90),
        ("L", 50),
        ("XL", 40),
        ("X", 10),
        ("IX", 9),
        ("V", 5),
        ("IV", 4),
        ("I", 1)
    ]

    result = ""
    for sym, val in symbols:
        while num >= val:
            result += sym
            num -= val
    return result
```
