---
id: 691b559495c5cb5a37b9b486
title: "تحدٍّ 126: اجعل الحروف كبيرة"
challengeType: 29
dashedName: challenge-126
---

# --description--

معطى سلسلة نصية بعنوان، أعد سلسلة جديدة مُنسقة بحالة العنوان باستخدام القواعد التالية:

- اجعل الحرف الأول من كل كلمة كبيرًا.
- اجعل جميع الحروف الأخرى في كل كلمة أحرفًا صغيرة.
- الكلمات مفصولة دائمًا بمسافة واحدة.

# --hints--

`title_case("hello world")` يجب أن تُعيد `"Hello World"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("hello world"), "Hello World")`)
}})
```

`title_case("the quick brown fox")` يجب أن تُعيد `"The Quick Brown Fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("the quick brown fox"), "The Quick Brown Fox")`)
}})
```

`title_case("JAVASCRIPT AND PYTHON")` يجب أن تُعيد `"Javascript And Python"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("JAVASCRIPT AND PYTHON"), "Javascript And Python")`)
}})
```

`title_case("AvOcAdO tOAst fOr brEAkfAst")` يجب أن تُعيد `"Avocado Toast For Breakfast"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast")`)
}})
```

# --seed--

## --seed-contents--

```py
def title_case(title):

    return title
```

# --solutions--

```py
def title_case(title):
    return " ".join(
        w[:1].upper() + w[1:].lower()
        for w in title.split(" ")
    )
```
