---
id: 691b559495c5cb5a37b9b482
title: "تحدٍّ 122: محلل النص العريض في Markdown"
challengeType: 29
dashedName: challenge-122
---

# --description--

أعطِ سلسلة نصية قد تحتوي على بعض النصوص العريضة في Markdown، أرجع سلسلة HTML المكافئة.

- النص العريض في Markdown هو أي نص يبدأ وينتهي بنجمتين (`**`) أو بشرطتين سفليتين (`__`).
- لا يمكن وجود مسافات بين النص والنجمتين أو الشرطتين السفليتين، لكن يمكن وجود
  مسافات داخل النص نفسه.
- حوّل كل النصوص العريضة إلى وسم HTML `b` وأرجع السلسلة النصية.

على سبيل المثال، إذا أعطيت `"**This is bold**"`، أرجع `"<b>This is bold</b>"`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML داخل السلاسل النصية عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parse_bold("**This is bold**")` يجب أن تُرجع `"<b>This is bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is bold**"), "<b>This is bold</b>")`)
}})
```

`parse_bold("__This is also bold__")` يجب أن تُرجع `"<b>This is also bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__This is also bold__"), "<b>This is also bold</b>")`)
}})
```

`parse_bold("**This is not bold **")` يجب أن تُرجع `"**This is not bold **"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is not bold **"), "**This is not bold **")`)
}})
```

`parse_bold("__ This is also not bold__")` يجب أن تُرجع `"__ This is also not bold__"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__ This is also not bold__"), "__ This is also not bold__")`)
}})
```

`parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog.")` يجب أن تُرجع `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_bold(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_bold(markdown):
    markdown = re.sub(r'\*\*(\S(?:.*?\S)?)\*\*', r'<b>\1</b>', markdown)
    markdown = re.sub(r'__(\S(?:.*?\S)?)__', r'<b>\1</b>', markdown)

    return markdown
```
