---
id: 691b559495c5cb5a37b9b488
title: "تحدٍّ 128: عدّ الحروف الساكنة"
challengeType: 29
dashedName: challenge-128
---

# --description--

معطاة سلسلة نصية ورقم مستهدف، حدّد ما إذا كانت السلسلة تحتوي بالضبط على العدد المستهدف من الحروف الساكنة.

- الحروف الساكنة هي كل الأحرف الأبجدية ما عدا `"a"`، `"e"`، `"i"`، `"o"`، و`"u"` بأي حالة.
- تجاهل الأرقام، علامات الترقيم، الفراغات، وأي حروف غير أبجدية عند العد.

# --hints--

يجب أن تُعيد الدالة `has_consonant_count("helloworld", 7)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("helloworld", 7), True)`)
}})
```

يجب أن تُعيد الدالة `has_consonant_count("eieio", 5)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("eieio", 5), False)`)
}})
```

يجب أن تُعيد الدالة `has_consonant_count("freeCodeCamp Rocks!", 11)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("freeCodeCamp Rocks!", 11), True)`)
}})
```

يجب أن تُعيد الدالة `has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24), False)`)
}})
```

يجب أن تُعيد الدالة `has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_consonant_count(text, target):

    return text
```

# --solutions--

```py
def has_consonant_count(text, target):
    vowels = {"a", "e", "i", "o", "u"}
    count = 0

    for ch in text.lower():
        if ch.isalpha() and ch not in vowels:
            count += 1

    return count == target
```
