---
id: 691b559495c5cb5a37b9b485
title: "تحدٍّ 125: لعبة الحياة"
challengeType: 29
dashedName: challenge-125
---

# --description--

معطى مصفوفة (مصفوفة من المصفوفات) تمثل الحالة الحالية في لعبة الحياة لـConway، أعد الحالة التالية للمصفوفة باستخدام هذه القواعد:

- كل خلية إما `1` (حية) أو `0` (ميتة).
- جيران الخلية هم حتى ثمانية خلايا محيطة بها (عموديًا، أفقيًا، وقطريًا).
- الخلايا على الحواف لها عدد أقل من ثمانية جيران.

قواعد تحديث كل خلية:

- أي خلية حية لها أقل من جارين أحياء تموت (قلة السكان).
- أي خلية حية لها جاران أو ثلاثة جيران أحياء تبقى على قيد الحياة.
- أي خلية حية لها أكثر من ثلاثة جيران أحياء تموت (زيادة السكان).
- أي خلية ميتة لها بالضبط ثلاثة جيران أحياء تصبح حية (تكاثر).

على سبيل المثال، معطى:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

أعد:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

تُحدّث كل خلية حسب عدد الجيران الأحياء. على سبيل المثال، تبقى `[0][0]` ميتة (لديها 2 جار حي)، وتبقى `[0][1]` حية (لديها 2 جار حي)، وتموت `[0][2]` (لديها 3 جيران أحياء)، وهكذا.

# --hints--

`game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` يجب أن تُعيد `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]])`)
}})
```

`game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` يجب أن تُعيد `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]])`)
}})
```

`game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` يجب أن تُعيد `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]])`)
}})
```

`game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` يجب أن تُعيد `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def game_of_life(grid):

    return grid
```

# --solutions--

```py
def game_of_life(grid):
    rows, cols = len(grid), len(grid[0])

    def count_live_neighbors(r, c):
        count = 0
        for i in range(r-1, r+2):
            for j in range(c-1, c+2):
                if 0 <= i < rows and 0 <= j < cols and (i != r or j != c):
                    count += grid[i][j]
        return count

    next_state = [row[:] for row in grid]

    for r in range(rows):
        for c in range(cols):
            live_neighbors = count_live_neighbors(r, c)

            if grid[r][c] == 1:
                if live_neighbors < 2 or live_neighbors > 3:
                    next_state[r][c] = 0
            else:
                if live_neighbors == 3:
                    next_state[r][c] = 1

    return next_state
```
