---
id: 691b559495c5cb5a37b9b489
title: "تحدٍّ 129: محلل الاقتباس في Markdown"
challengeType: 29
dashedName: challenge-129
---

# --description--

أعطِ سلسلة نصية تتضمن اقتباسًا في Markdown، وأرجع السلسلة المكافئة بصيغة HTML.

الاقتباس في Markdown هو أي سطر يحقق ما يلي:

- يبدأ بصفر أو أكثر من الفراغات
- يتبعه علامة أكبر من (`>`)
- ثم واحد أو أكثر من الفراغات
- وأخيرًا نص الاقتباس.

أرجع نص الاقتباس محاطًا بوسمي HTML `blockquote` الافتتاحي والختامي.

على سبيل المثال، إذا كانت السلسلة `"> This is a quote"`، فأرجع `<blockquote>This is a quote</blockquote>`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML داخل السلاسل النصية عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parse_blockquote("> This is a quote")` يجب أن تُرجع `"<blockquote>This is a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>")`)
}})
```

`parse_blockquote(" > This is also a quote")` يجب أن تُرجع `"<blockquote>This is also a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>")`)
}})
```

`parse_blockquote("  >    So  Is  This")` يجب أن تُرجع `"<blockquote>So  Is  This</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```
