---
id: 6925e2068081f40f549ced1b
title: "تحدٍّ 137: مولد ندفة الثلج"
challengeType: 29
dashedName: challenge-137
---

# --description--

معطاة سلسلة نصية متعددة الأسطر تستخدم أحرف السطر الجديد (`\n`) لتمثيل فاصل الأسطر، أعد سلسلة جديدة حيث تُعكس كل سطر أفقيًا وتُلحق بنهاية السطر الأصلي.

- عكس السطر يتم بعكس جميع حروفه، بما في ذلك الفراغات.

على سبيل المثال، معطاة `"* \n *\n* "`، التي تُسجل في وحدة التحكم كالتالي:

```sh
* 
 *
* 
```

أعد `"*  *\n ** \n*  *"`، التي تُسجل في وحدة التحكم كالتالي:

```sh
*  *
 ** 
*  *
```

انتبه جيدًا إلى الفراغات في السلاسل المعطاة والمعادة. تأكد من عدم حذف أي منها.

# --hints--

`generate_snowflake("* \n *\n* ")` يجب أن تعيد `"*  *\n ** \n*  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` يجب أن تعيد `"X=~~=X"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` يجب أن تعيد `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` يجب أن تعيد `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` يجب أن تعيد `"*  --  *\n * -- * \n*  --  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```
