---
id: 6925e2068081f40f549ced1d
title: "تحدٍّ 139: حجر، ورقة، مقص"
challengeType: 29
dashedName: challenge-139
---

# --description--

معطاة سلسلتا نصية، الأولى تمثل اللاعب 1 والثانية تمثل اللاعب 2، حدّد الفائز في مباراة حجر، ورقة، مقص.

- ستكون السلاسل النصية دائمًا `"Rock"`، `"Paper"`، أو `"Scissors"`.
- `"Rock"` تغلب `"Scissors"`.
- `"Paper"` تغلب `"Rock"`.
- `"Scissors"` تغلب `"Paper"`.

أرجع:

- `"Player 1 wins"` إذا فاز اللاعب 1.
- `"Player 2 wins"` إذا فاز اللاعب 2.
- `"Tie"` إذا اختار اللاعبان نفس الخيار.

# --hints--

`rock_paper_scissors("Rock", "Rock")` يجب أن تُرجع `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Rock"), "Tie")`)
}})
```

`rock_paper_scissors("Rock", "Paper")` يجب أن تُرجع `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Paper"), "Player 2 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Paper")` يجب أن تُرجع `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Paper"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Rock", "Scissors")` يجب أن تُرجع `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Scissors"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Scissors")` يجب أن تُرجع `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Scissors"), "Tie")`)
}})
```

`rock_paper_scissors("Scissors", "Rock")` يجب أن تُرجع `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Rock"), "Player 2 wins")`)
}})
```

# --seed--

## --seed-contents--

```py
def rock_paper_scissors(player1, player2):

    return player1
```

# --solutions--

```py
def rock_paper_scissors(player1, player2):
    if player1 == player2:
        return "Tie"
    
    if (player1 == "Rock" and player2 == "Scissors") or \
       (player1 == "Paper" and player2 == "Rock") or \
       (player1 == "Scissors" and player2 == "Paper"):
        return "Player 1 wins"
    else:
        return "Player 2 wins"
```
