---
id: 6925e2068081f40f549ced1c
title: "تحدٍّ 138: مجموع القواسم"
challengeType: 29
dashedName: challenge-138
---

# --description--

أعطِ عددًا صحيحًا موجبًا، أرجع مجموع جميع قواسمه.

- القاسم هو أي عدد صحيح يقسم العدد بالتساوي (الباقي هو `0`).
- احسب كل قاسم مرة واحدة فقط.

على سبيل المثال، إذا كان العدد `6`، أرجع `12` لأن قواسم `6` هي `1` و `2` و `3` و `6`، ومجموعها هو `12`.

# --hints--

`sum_divisors(6)` يجب أن تُرجع `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` يجب أن تُرجع `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` يجب أن تُرجع `56`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` يجب أن تُرجع `224`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` يجب أن تُرجع `806`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` يجب أن تُرجع `23520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```
