---
id: 6925e2068081f40f549ced1a
title: "تحدٍّ 136: محلل صور Markdown"
challengeType: 29
dashedName: challenge-136
---

# --description--

أعطِ سلسلة نصية لصورة في Markdown، وأرجع سلسلة HTML المكافئة.

صورة Markdown لها الصيغة التالية: `"![alt text](image_url)"`. حيث:

- `alt text` هو وصف الصورة (قيمة خاصية `alt`).
- `image_url` هو عنوان URL المصدر للصورة (قيمة خاصية `src`).

أرجع سلسلة وسم HTML `img` مع تعيين `src` إلى عنوان الصورة و`alt` إلى نص الوصف.

على سبيل المثال، إذا أعطيت `"![Cute cat](cat.png)"` أرجع `'<img src="cat.png" alt="Cute cat">'`;

- تأكد من أن الوسم، وترتيب الخصائص، والتباعد، واستخدام علامات الاقتباس هي نفسها كما في الأعلى.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML داخل السلاسل النصية عند تسجيل الرسائل — تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parse_image("![Cute cat](cat.png)")` يجب أن تُرجع `'<img src="cat.png" alt="Cute cat">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` يجب أن تُرجع `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` يجب أن تُرجع `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```
