---
id: 69272dcf1c24b44fd79137c4
title: "تحدٍّ 141: وقود الإقلاع"
challengeType: 29
dashedName: challenge-141
---

# --description--

بالنظر إلى عدد جالونات الوقود الموجودة حاليًا في طائرتك، وعدد اللترات المطلوبة من الوقود للوصول إلى وجهتك، حدد كم جالونًا إضافيًا من الوقود يجب أن تضيف.

- الجالون الواحد يساوي 3.78541 لترات.
- إذا كانت الطائرة تحتوي على وقود كافٍ بالفعل، أرجع `0`.
- يمكنك إضافة جالونات كاملة فقط.
- لا تدرج الأعداد العشرية في الرقم المرجع.

# --hints--

يجب أن تُرجع الدالة `fuel_to_add(0, 1)` القيمة `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(0, 1), 1)`)
}})
```

يجب أن تُرجع الدالة `fuel_to_add(5, 40)` القيمة `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(5, 40), 6)`)
}})
```

يجب أن تُرجع الدالة `fuel_to_add(10, 30)` القيمة `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(10, 30), 0)`)
}})
```

يجب أن تُرجع الدالة `fuel_to_add(896, 20500)` القيمة `4520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(896, 20500), 4520)`)
}})
```

يجب أن تُرجع الدالة `fuel_to_add(1000, 50000)` القيمة `12209`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(1000, 50000), 12209)`)
}})
```

# --seed--

## --seed-contents--

```py
def fuel_to_add(current_gallons, required_liters):

    return current_gallons
```

# --solutions--

```py
import math
def fuel_to_add(current_gallons, required_liters):
    liters_per_gallon = 3.78541
    current_liters = current_gallons * liters_per_gallon
    if current_liters >= required_liters:
        return 0
    liters_needed = required_liters - current_liters
    return math.ceil(liters_needed / liters_per_gallon)
```
