---
id: 69272dcf1c24b44fd79137c3
title: "تحدٍّ 140: SCREAMING_SNAKE_CASE"
challengeType: 29
dashedName: challenge-140
---

# --description--

أعطِ سلسلة نصية تمثّل اسم متغير، أرجع اسم المتغير بعد تحويله إلى SCREAMING_SNAKE_CASE.

ستُكتب أسماء المتغيرات المعطاة بأحد التنسيقات التالية:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

في التنسيقات السابقة، تُفصل الكلمات بشرطة سفلية (`_`) أو شرطة (`-`)، أو تبدأ كلمة جديدة بحرف كبير.

لتحويل الاسم إلى SCREAMING_SNAKE_CASE:

- اجعل كل الحروف كبيرة
- افصل الكلمات بشرطة سفلية (`_`)

# --hints--

`to_screaming_snake_case("userEmail")` يجب أن تُرجع `"USER_EMAIL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("userEmail"), "USER_EMAIL")`)
}})
```

`to_screaming_snake_case("UserPassword")` يجب أن تُرجع `"USER_PASSWORD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("UserPassword"), "USER_PASSWORD")`)
}})
```

`to_screaming_snake_case("user_id")` يجب أن تُرجع `"USER_ID"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user_id"), "USER_ID")`)
}})
```

`to_screaming_snake_case("user-address")` يجب أن تُرجع `"USER_ADDRESS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user-address"), "USER_ADDRESS")`)
}})
```

`to_screaming_snake_case("username")` يجب أن تُرجع `"USERNAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("username"), "USERNAME")`)
}})
```

`to_screaming_snake_case("my_variable_name")` يجب أن تُرجع `"MY_VARIABLE_NAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("my_variable_name"), "MY_VARIABLE_NAME")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_screaming_snake_case(variable_name):

    return variable_name
```

# --solutions--

```py
import re
def to_screaming_snake_case(variable_name):
    temp = re.sub(r'[-_]+', ' ', variable_name)
    temp = re.sub(r'([a-z0-9])([A-Z])', r'\1 \2', temp)
    words = temp.strip().split()
    return '_'.join(words).upper()
```
