---
id: 69272dcf1c24b44fd79137c6
title: "تحدٍّ 143: محلل النص المائل في Markdown"
challengeType: 29
dashedName: challenge-143
---

# --description--

أعطِ سلسلة نصية قد تتضمن بعض النصوص المائلة في Markdown، أعد السلسلة المكافئة بصيغة HTML.

- النص المائل في Markdown هو أي نص يبدأ وينتهي بنجمة واحدة (`*`) أو بشرطة سفلية واحدة (`_`).
- لا يمكن أن توجد مسافات بين النص والنجمة أو الشرطة السفلية، لكن يمكن أن توجد مسافات داخل النص نفسه.
- حوّل كل النصوص المائلة إلى وسم `i` في HTML وأعد السلسلة.

على سبيل المثال، إذا كانت السلسلة `"*This is italic*"`، أعد `"<i>This is italic</i>"`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML داخل السلاسل النصية عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parse_italics("*This is italic*")` يجب أن تعيد `"<i>This is italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is italic*"), "<i>This is italic</i>")`)
}})
```

`parse_italics("_This is also italic_")` يجب أن تعيد `"<i>This is also italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_This is also italic_"), "<i>This is also italic</i>")`)
}})
```

`parse_italics("*This is not italic *")` يجب أن تعيد `"*This is not italic *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is not italic *"), "*This is not italic *")`)
}})
```

`parse_italics("_ This is also not italic_")` يجب أن تعيد `"_ This is also not italic_"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_ This is also not italic_"), "_ This is also not italic_")`)
}})
```

`parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog.")` يجب أن تعيد `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_italics(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_italics(markdown):
    pattern = r'(\*|_)([^\s][^]*?[^\s])\1'
    pattern = r'(\*|_)([^\s].*?[^\s])\1'
    return re.sub(pattern, r'<i>\2</i>', markdown)
```
