---
id: 69272dcf1c24b44fd79137c5
title: "تحدٍّ 142: مجموع السلسلة النصية"
challengeType: 29
dashedName: challenge-142
---

# --description--

أعطِ سلسلة نصية تحتوي على أرقام وحروف أخرى، أرجع مجموع كل الأرقام في السلسلة.

- اعتبر الأرقام المتتالية رقمًا واحدًا. على سبيل المثال، `"13"` تُحسب كـ 13، وليس 1 + 3.
- تجاهل أي حروف غير أرقام.

# --hints--

`string_sum("3apples2bananas")` يجب أن تُرجع `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("3apples2bananas"), 5)`)
}})
```

`string_sum("10cats5dogs2birds")` يجب أن تُرجع `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("10cats5dogs2birds"), 17)`)
}})
```

`string_sum("125344")` يجب أن تُرجع `125344`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("125344"), 125344)`)
}})
```

`string_sum("a1b20c300")` يجب أن تُرجع `321`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a1b20c300"), 321)`)
}})
```

`string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` يجب أن تُرجع `1653`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653)`)
}})
```

# --seed--

## --seed-contents--

```py
def string_sum(s):

    return s
```

# --solutions--

```py
import re
def string_sum(s):
    numbers = re.findall(r'\d+', s)
    return sum(int(num) for num in numbers)
```
