---
id: 69306364df283fcaff2e1ada
title: "تحدٍّ 149: vOwElcAsE"
challengeType: 29
dashedName: challenge-149
---

# --description--

أعطِ سلسلة نصية، أرجع سلسلة جديدة تُحوّل فيها جميع الحروف المتحركة إلى أحرف كبيرة وباقي الحروف الأبجدية إلى أحرف صغيرة.

- الحروف المتحركة هي `"a"`، `"e"`، `"i"`، `"o"`، و`"u"` بأي حالة.
- يجب أن تبقى الأحرف غير الأبجدية دون تغيير.

# --hints--

`vowel_case("vowelcase")` يجب أن تُرجع `"vOwElcAsE"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("vowelcase"), "vOwElcAsE")`)
}})
```

`vowel_case("coding is fun")` يجب أن تُرجع `"cOdIng Is fUn"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("coding is fun"), "cOdIng Is fUn")`)
}})
```

`vowel_case("HELLO, world!")` يجب أن تُرجع `"hEllO, wOrld!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HELLO, world!"), "hEllO, wOrld!")`)
}})
```

`vowel_case("git cherry-pick")` يجب أن تُرجع `"gIt chErry-pIck"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("git cherry-pick"), "gIt chErry-pIck")`)
}})
```

`vowel_case("HEAD~1")` يجب أن تُرجع `"hEAd~1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HEAD~1"), "hEAd~1")`)
}})
```

# --seed--

## --seed-contents--

```py
def vowel_case(s):

    return s
```

# --solutions--

```py
def vowel_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char in vowels:
            result += char.upper()
        elif char.isalpha():
            result += char.lower()
        else:
            result += char

    return result
```
