---
id: 69306364df283fcaff2e1ad8
title: "تحدٍّ 147: حاسبة السنة الكبيسة"
challengeType: 29
dashedName: challenge-147
---

# --description--

معطى سنة صحيحة، حدّد ما إذا كانت سنة كبيسة.

تكون السنة كبيسة إذا استوفت القواعد التالية:

- السنة تقبل القسمة على 4 بدون باقي، و
- السنة لا تقبل القسمة على 100 بدون باقي، إلا إذا
- السنة تقبل القسمة على 400 بدون باقي.

# --hints--

يجب أن تُرجع `is_leap_year(2024)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2024), True)`)
}})
```

يجب أن تُرجع `is_leap_year(2023)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2023), False)`)
}})
```

يجب أن تُرجع `is_leap_year(2100)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2100), False)`)
}})
```

يجب أن تُرجع `is_leap_year(2000)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2000), True)`)
}})
```

يجب أن تُرجع `is_leap_year(1999)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(1999), False)`)
}})
```

يجب أن تُرجع `is_leap_year(2040)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2040), True)`)
}})
```

يجب أن تُرجع `is_leap_year(2026)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2026), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_leap_year(year):

    return year
```

# --solutions--

```py
def is_leap_year(year):
    if (year % 4 == 0 and year % 100 != 0) or (year % 400 == 0):
        return True
    else:
        return False
```
