---
id: 69306364df283fcaff2e1ad9
title: "تحدٍّ 148: ضغط الإطارات"
challengeType: 29
dashedName: challenge-148
---

# --description--

معطى مصفوفة بأربعة أرقام تمثل ضغوط الإطارات بوحدة psi لأربعة إطارات في مركبتك، ومصفوفة أخرى برقمين تمثلان الحد الأدنى والحد الأقصى للضغط لإطاراتك بوحدة bar، أرجع مصفوفة بأربعة سلاسل نصية تصف حالة كل إطار.

- 1 bar يساوي 14.5038 psi.

أرجع مصفوفة بالقيم التالية لكل إطار:

- `"Low"` إذا كان ضغط الإطار أقل من الحد الأدنى المسموح.
- `"Good"` إذا كان بين الحد الأدنى والحد الأقصى المسموح.
- `"High"` إذا كان أعلى من الحد الأقصى المسموح.

# --hints--

`tire_status([32, 28, 35, 29], [2, 3])` يجب أن تُرجع `["Good", "Low", "Good", "Low"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"])`)
}})
```

`tire_status([32, 28, 35, 30], [2, 2.3])` يجب أن تُرجع `["Good", "Low", "High", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"])`)
}})
```

`tire_status([29, 26, 31, 28], [2.1, 2.5])` يجب أن تُرجع `["Low", "Low", "Good", "Low"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"])`)
}})
```

`tire_status([31, 31, 30, 29], [1.5, 2])` يجب أن تُرجع `["High", "High", "High", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"])`)
}})
```

`tire_status([30, 28, 30, 29], [1.9, 2.1])` يجب أن تُرجع `["Good", "Good", "Good", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"])`)
}})
```

# --seed--

## --seed-contents--

```py
def tire_status(pressures_psi, range_bar):

    return pressures_psi
```

# --solutions--

```py
def tire_status(pressures_psi, range_bar):
    psi_to_bar = 1 / 14.5038
    min_bar, max_bar = range_bar
    result = []

    for psi in pressures_psi:
        bar = psi * psi_to_bar
        if bar < min_bar:
            result.append("Low")
        elif bar > max_bar:
            result.append("High")
        else:
            result.append("Good")
    return result
```
