---
id: 69306364df283fcaff2e1ad7
title: "تحدٍّ 146: المقعد الأيسر على الطاولة"
challengeType: 29
dashedName: challenge-146
---

# --description--

معطى مصفوفة 4x2 (مصفوفة من المصفوفات) تمثل ترتيب المقاعد لوجبة، حدّد عدد المقاعد التي يمكن للشخص الأيسر الجلوس عليها.

- لا يمكن للشخص الأيسر الجلوس حيث يجلس شخص أيمن في المقعد المجاور على اليسار مباشرةً.

في المصفوفة المعطاة:

- المقعد الذي يحتوي على `"R"` يشغله شخص أيمن.
- المقعد الذي يحتوي على `"L"` يشغله شخص أيسر.
- المقعد الذي يحتوي على `"U"` هو مقعد شاغر.
- المقاعد الشاغرة فقط متاحة للجلوس.
- المقاعد في الصف العلوي تواجه "للأسفل"، والمقاعد في الصف السفلي تواجه "للأعلى" (مثل الطاولة)، لذا اليسار واليمين نسبيان لاتجاه المقعد.
- المقاعد الزاوية لها مقعد واحد فقط بجانبها.

على سبيل المثال، في المصفوفة المعطاة:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

لا يمكن الجلوس في المقعد العلوي الأيسر لأن هناك شخصًا أيمن على اليسار. يمكن الجلوس في المقعدين المفتوحين الآخرين لأنه لا يوجد شخص أيمن على اليسار.

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` يجب أن تُرجع `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` يجب أن تُرجع `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` يجب أن تُرجع `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` يجب أن تُرجع `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` يجب أن تُرجع `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```
