---
id: 69306364df283fcaff2e1ad6
title: "تحدٍّ 145: العدد النوني في متتالية فيبوناتشي"
challengeType: 29
dashedName: challenge-145
---

# --description--

معطى عدد صحيح `n`، أرجع العدد `n` في متتالية فيبوناتشي.

متتالية فيبوناتشي هي سلسلة أعداد يكون كل عدد فيها مجموع العددين السابقين له. أول 10 أعداد في المتتالية هي `0`، `1`، `1`، `2`، `3`، `5`، `8`، `13`، `21`، `34`.

# --hints--

يجب أن تُرجع الدالة `nth_fibonacci(4)` القيمة `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(4), 2)`)
}})
```

يجب أن تُرجع الدالة `nth_fibonacci(10)` القيمة `34`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(10), 34)`)
}})
```

يجب أن تُرجع الدالة `nth_fibonacci(15)` القيمة `377`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(15), 377)`)
}})
```

يجب أن تُرجع الدالة `nth_fibonacci(40)` القيمة `63245986`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(40), 63245986)`)
}})
```

يجب أن تُرجع الدالة `nth_fibonacci(75)` القيمة `1304969544928657`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(75), 1304969544928657)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_fibonacci(n):

    return n
```

# --solutions--

```py
def nth_fibonacci(n):
    if n == 1:
        return 0
    if n == 2:
        return 1

    a, b = 0, 1
    for _ in range(3, n + 1):
        a, b = b, a + b
    return b
```
