---
id: 69373793f5a867f769cde139
title: "تحدٍّ 154: المعدل الطبيعي للحفرة"
challengeType: 29
dashedName: challenge-154
---

# --description--

معطى عددان صحيحان، المعدل الطبيعي لحفرة الغولف وعدد خطوط الرسم التي استغرقها اللاعب في تلك الحفرة، أرجع نتيجة اللاعب باستخدام مصطلحات الغولف.

أرجع:

- `"Hole in one!"` إذا استغرق خط رسم واحد.
- `"Eagle"` إذا استغرق خطي رسم أقل من المعدل الطبيعي.
- `"Birdie"` إذا استغرق خط رسم واحد أقل من المعدل الطبيعي.
- `"Par"` إذا استغرق نفس عدد خطوط الرسم كالمعدل الطبيعي.
- `"Bogey"` إذا استغرق خط رسم واحد أكثر من المعدل الطبيعي.
- `"Double bogey"` إذا استغرق خطي رسم أكثر من المعدل الطبيعي.

# --hints--

`golf_score(3, 3)` يجب أن تُرجع `"Par"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 3), "Par")`)
}})
```

`golf_score(4, 3)` يجب أن تُرجع `"Birdie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 3), "Birdie")`)
}})
```

`golf_score(3, 1)` يجب أن تُرجع `"Hole in one!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 1), "Hole in one!")`)
}})
```

`golf_score(5, 7)` يجب أن تُرجع `"Double bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 7), "Double bogey")`)
}})
```

`golf_score(4, 5)` يجب أن تُرجع `"Bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 5), "Bogey")`)
}})
```

`golf_score(5, 3)` يجب أن تُرجع `"Eagle"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 3), "Eagle")`)
}})
```

# --seed--

## --seed-contents--

```py
def golf_score(par, strokes):

    return par
```

# --solutions--

```py
def golf_score(par, strokes):
    if strokes == 1:
        return "Hole in one!"

    diff = strokes - par

    if diff == -2:
        return "Eagle"
    if diff == -1:
        return "Birdie"
    if diff == 0:
        return "Par"
    if diff == 1:
        return "Bogey"
    if diff == 2:
        return "Double bogey"
```
