---
id: 69373793f5a867f769cde13a
title: "تحدٍّ 155: زراعة المحصول"
challengeType: 29
dashedName: challenge-155
---

# --description--

معطى عدد صحيح يُمثّل حجم حقل مزرعتك، و`"acres"` أو `"hectares"` كوحدة لقياس حجم الحقل، ونوع المحصول، حدّد كم عدد النباتات من هذا النوع التي يمكنك زراعتها في الحقل.

- 1 فدان يساوي 4046.86 متر مربع.
- 1 هكتار يساوي 10,000 متر مربع.

إليك قائمة بالمحاصيل التي ستُعطى كمُدخَل ومساحة كل نبتة منها:

|المحصول|المساحة لكل نبتة|
|-|-|
|`"corn"`|1 متر مربع|
|`"wheat"`|0.1 متر مربع|
|`"soybeans"`|0.5 متر مربع|
|`"tomatoes"`|0.25 متر مربع|
|`"lettuce"`|0.2 متر مربع|

أرجع عدد النباتات التي تتسع في الحقل، مقربًا للأسفل لأقرب عدد صحيح.

# --hints--

يجب أن تُرجع `get_number_of_plants(1, "acres", "corn")` القيمة `4046`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(1, "acres", "corn"), 4046)`)
}})
```

يجب أن تُرجع `get_number_of_plants(2, "hectares", "lettuce")` القيمة `100000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(2, "hectares", "lettuce"), 100000)`)
}})
```

يجب أن تُرجع `get_number_of_plants(20, "acres", "soybeans")` القيمة `161874`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(20, "acres", "soybeans"), 161874)`)
}})
```

يجب أن تُرجع `get_number_of_plants(3.75, "hectares", "tomatoes")` القيمة `150000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(3.75, "hectares", "tomatoes"), 150000)`)
}})
```

يجب أن تُرجع `get_number_of_plants(16.75, "acres", "tomatoes")` القيمة `271139`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(16.75, "acres", "tomatoes"), 271139)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_of_plants(field_size, unit, crop):

    return field_size
```

# --solutions--

```py
import math
def get_number_of_plants(field_size, unit, crop):
    crop_space = {
        "corn": 1,
        "wheat": 0.1,
        "soybeans": 0.5,
        "tomatoes": 0.25,
        "lettuce": 0.2
    }

    field_in_sqm = field_size * 4046.86 if unit == "acres" else field_size * 10000

    return math.floor(field_in_sqm / crop_space[crop.lower()])
```
