---
id: 69373793f5a867f769cde135
title: "تحدٍّ 150: محلل قائمة غير مرتبة في Markdown"
challengeType: 29
dashedName: challenge-150
---

# --description--

بالنظر إلى سلسلة نصية لقائمة غير مرتبة صحيحة في Markdown، أعد السلسلة المكافئة بصيغة HTML.

تتكون القائمة غير المرتبة من عنصر قائمة واحد أو أكثر. يظهر عنصر القائمة الصحيح في سطر منفصل و:

- يبدأ بشرطة (`"-"`)، تليها
- مسافة واحدة على الأقل، ثم
- نص عنصر القائمة.

تُعطى القائمة كسلسلة نصية واحدة مع فواصل أسطر مفصولة بحرف السطر الجديد (`"\n"`). لا تضمن أحرف السطر الجديد في نص العنصر.

لف كل عنصر قائمة داخل وسم `li` في HTML، ولجميع العناصر داخل وسم `ul`.

على سبيل المثال، إذا أعطيت `"- Item A\n- Item B"`، أعد `"<ul><li>Item A</li><li>Item B</li></ul>"`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML في السلاسل النصية عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parse_unordered_list("- Item A\n- Item B")` يجب أن يعيد `"<ul><li>Item A</li><li>Item B</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- Item A\\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>")`)
}})
```

`parse_unordered_list("-  JavaScript\n-  Python")` يجب أن يعيد `"<ul><li>JavaScript</li><li>Python</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("-  JavaScript\\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>")`)
}})
```

`parse_unordered_list("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` يجب أن يعيد `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- 2 C Flour\\n- 1/2 C Sugar\\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>")`)
}})
```

`parse_unordered_list("- A-1\n- A-2\n- B-1")` يجب أن يعيد `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- A-1\\n- A-2\\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_unordered_list(markdown):

    return markdown
```

# --solutions--

```py
def parse_unordered_list(markdown):
    lines = markdown.split("\n")
    list_items = []

    for line in lines:
        list_items.append(f"<li>{line[1:].strip()}</li>")

    return f"<ul>{''.join(list_items)}</ul>"
```
