---
id: 69373793f5a867f769cde136
title: "تحدٍّ 151: مصفوفة مرتبة؟"
challengeType: 29
dashedName: challenge-151
---

# --description--

معطاة مصفوفة من الأرقام، حدّد ما إذا كانت الأرقام مرتبة ترتيبًا تصاعديًا، ترتيبًا تنازليًا، أو غير مرتبة.

إذا كانت المصفوفة المعطاة:

- مرتبة ترتيبًا تصاعديًا (من الأقل إلى الأعلى)، أرجع `"Ascending"`.
- مرتبة ترتيبًا تنازليًا (من الأعلى إلى الأقل)، أرجع `"Descending"`.
- غير مرتبة ترتيبًا تصاعديًا أو تنازليًا، أرجع `"Not sorted"`.

# --hints--

`is_sorted([1, 2, 3, 4, 5])` يجب أن تُرجع `"Ascending"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([1, 2, 3, 4, 5]), "Ascending")`)
}})
```

`is_sorted([10, 8, 6, 4, 2])` يجب أن تُرجع `"Descending"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([10, 8, 6, 4, 2]), "Descending")`)
}})
```

`is_sorted([1, 3, 2, 4, 5])` يجب أن تُرجع `"Not sorted"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([1, 3, 2, 4, 5]), "Not sorted")`)
}})
```

`is_sorted([3.14, 2.71, 1.61, 0.57])` يجب أن تُرجع `"Descending"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([3.14, 2.71, 1.61, 0.57]), "Descending")`)
}})
```

`is_sorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])` يجب أن تُرجع `"Ascending"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending")`)
}})
```

`is_sorted([0.4, 0.5, 0.3])` يجب أن تُرجع `"Not sorted"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([0.4, 0.5, 0.3]), "Not sorted")`)
}})
```

# --seed--

## --seed-contents--

```py
def is_sorted(arr):

    return arr
```

# --solutions--

```py
def is_sorted(arr):
    ascending = True
    descending = True

    for i in range(1, len(arr)):
        if arr[i] < arr[i - 1]:
            ascending = False
        if arr[i] > arr[i - 1]:
            descending = False

    if ascending:
        return "Ascending"
    if descending:
        return "Descending"
    return "Not sorted"
```
