---
id: 6939b873185d8e00d453563d
title: "تحدٍّ 158: تبديل المصفوفة"
challengeType: 29
dashedName: challenge-158
---

# --description--

أعطِ مصفوفة تحتوي على قيمتين، أرجع مصفوفة بالقيم بعد تبديل مواضعها.

على سبيل المثال، إذا كانت المصفوفة `["A", "B"]` فأرجع `["B", "A"]`.

# --hints--

يجب أن تُرجع الدالة `array_swap(["A", "B"])` القيمة `["B", "A"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["A", "B"]), ["B", "A"])`)
}})
```

يجب أن تُرجع الدالة `array_swap([25, 20])` القيمة `[20, 25]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([25, 20]), [20, 25])`)
}})
```

يجب أن تُرجع الدالة `array_swap([True, False])` القيمة `[False, True]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([True, False]), [False, True])`)
}})
```

يجب أن تُرجع الدالة `array_swap(["1", 1])` القيمة `[1, "1"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["1", 1]), [1, "1"])`)
}})
```

# --seed--

## --seed-contents--

```py
def array_swap(arr):

    return arr
```

# --solutions--

```py
def array_swap(arr):

    return [arr[1], arr[0]]
```
