---
id: 6939b873185d8e00d453563f
title: "تحدٍّ 160: حركات الفارس"
challengeType: 29
dashedName: challenge-160
---

# --description--

بالنظر إلى موقع الفارس على رقعة الشطرنج، أعد عدد المربعات الصالحة التي يمكن للفارس التحرك إليها.

رقعة الشطرنج القياسية هي 8x8، مع أعمدة معنونة من `A` إلى `H` (من اليسار إلى اليمين) وصفوف معنونة من `1` إلى `8` (من الأسفل إلى الأعلى). تبدو هكذا:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

يتحرك الفارس على شكل "L": مربعان في اتجاه واحد (أفقي أو عمودي)، ومربع واحد في الاتجاه العمودي عليه.

هذا يعني أن الفارس يمكنه التحرك إلى ثمانية مواقع محتملة كحد أقصى، ولكن أقل عندما يكون بالقرب من حواف الرقعة. على سبيل المثال، إذا كان الفارس عند `A1`، يمكنه فقط التحرك إلى `B3` أو `C2`.

# --hints--

`knight_moves("A1")` يجب أن تُعيد `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("A1"), 2)`)
}})
```

`knight_moves("D4")` يجب أن تُعيد `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("D4"), 8)`)
}})
```

`knight_moves("G6")` يجب أن تُعيد `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("G6"), 6)`)
}})
```

`knight_moves("B8")` يجب أن تُعيد `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("B8"), 3)`)
}})
```

`knight_moves("H3")` يجب أن تُعيد `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("H3"), 4)`)
}})
```

# --seed--

## --seed-contents--

```py
def knight_moves(position):

    return position
```

# --solutions--

```py
def knight_moves(position):
    col = ord(position[0]) - ord('A')
    row = int(position[1]) - 1

    moves = [
        (2, 1), (2, -1),
        (-2, 1), (-2, -1),
        (1, 2), (1, -2),
        (-1, 2), (-1, -2)
    ]

    valid = 0

    for dx, dy in moves:
        new_col = col + dx
        new_row = row + dy

        if 0 <= new_col < 8 and 0 <= new_row < 8:
            valid += 1

    return valid
```
