---
id: 6939b873185d8e00d453563c
title: "تحدٍّ 157: محلل روابط Markdown"
challengeType: 29
dashedName: challenge-157
---

# --description--

بالنظر إلى سلسلة نصية لرابط في Markdown، أعد السلسلة المكافئة في HTML.

صورة Markdown لها الصيغة التالية: `"[link_text](link_url)"`. أعد سلسلة وسم HTML `a` مع تعيين `href` إلى `link_url` و`link_text` كمحتوى الوسم.

على سبيل المثال، إذا كانت السلسلة `"[freeCodeCamp](https://freecodecamp.org/)"` فأعد `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML داخل السلاسل النصية عند تسجيل الرسائل — تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parse_link("[freeCodeCamp](https://freecodecamp.org/)")` يجب أن تعيد `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>')`)
}})
```

`parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` يجب أن تعيد `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>')`)
}})
```

`parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` يجب أن تعيد `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_link(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_link(markdown):
    match = re.match(r'\[(.*?)\]\((.*?)\)', markdown)
    text, url = match.groups()
    return f'<a href="{url}">{text}</a>'
```
