---
id: 6939b873185d8e00d453563b
title: "تحدٍّ 156: فردي أم زوجي؟"
challengeType: 29
dashedName: challenge-156
---

# --description--

أعطِ عددًا صحيحًا موجبًا، أرجع `"Odd"` إذا كان عددًا فرديًا، و`"Even"` إذا كان زوجيًا.

# --hints--

`odd_or_even(1)` يجب أن تُرجع `"Odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(1), "Odd")`)
}})
```

`odd_or_even(2)` يجب أن تُرجع `"Even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(2), "Even")`)
}})
```

`odd_or_even(13)` يجب أن تُرجع `"Odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(13), "Odd")`)
}})
```

`odd_or_even(196)` يجب أن تُرجع `"Even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(196), "Even")`)
}})
```

`odd_or_even(123456789)` يجب أن تُرجع `"Odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(123456789), "Odd")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even(n):

    return n
```

# --solutions--

```py
def odd_or_even(n):
    return "Even" if n % 2 == 0 else "Odd"
```
