---
id: 6939b873185d8e00d453563e
title: "تحدٍّ 159: الوتر عدد صحيح"
challengeType: 29
dashedName: challenge-159
---

# --description--

معطى عددان صحيحان موجبان يمثلان أطوال الضلعين القائمين (الجانبين القصيرين) لمثلث قائم، حدّد ما إذا كان الوتر عددًا صحيحًا.

يُحسب طول الوتر بجمع مربعات طولي الضلعين القائمين ثم أخذ الجذر التربيعي للمجموع (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

يجب أن تُعيد الدالة `is_integer_hypotenuse(3, 4)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(3, 4), True)`)
}})
```

يجب أن تُعيد الدالة `is_integer_hypotenuse(2, 3)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(2, 3), False)`)
}})
```

يجب أن تُعيد الدالة `is_integer_hypotenuse(5, 12)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(5, 12), True)`)
}})
```

يجب أن تُعيد الدالة `is_integer_hypotenuse(10, 10)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(10, 10), False)`)
}})
```

يجب أن تُعيد الدالة `is_integer_hypotenuse(780, 1040)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(780, 1040), True)`)
}})
```

يجب أن تُعيد الدالة `is_integer_hypotenuse(250, 333)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(250, 333), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_integer_hypotenuse(a, b):

    return a
```

# --solutions--

```py
import math
def is_integer_hypotenuse(a, b):
    total = a*a + b*b
    c = math.isqrt(total)
    return c*c == total
```
