---
id: 694596b0585c11170ac7c7fb
title: "تحدٍّ 163: حالة الحروف الساكنة"
challengeType: 29
dashedName: challenge-163
---

# --description--

أعطِ سلسلة نصية تمثّل اسم متغير، حوّلها إلى حالة الحروف الساكنة باستخدام القواعد التالية:

- يجب تحويل جميع الحروف الساكنة إلى أحرف كبيرة.
- يجب تحويل جميع الحروف المتحركة (`a`، `e`، `i`، `o`، `u` بأي حالة) إلى أحرف صغيرة.
- يجب تحويل جميع الشرطات (`-`) إلى شرطات سفلية (`_`).

# --hints--

`to_consonant_case("helloworld")` يجب أن تُعيد `"HeLLoWoRLD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("helloworld"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("HELLOWORLD")` يجب أن تُعيد `"HeLLoWoRLD"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("HELLOWORLD"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("_hElLO-WOrlD-")` يجب أن تُعيد `"_HeLLo_WoRLD_"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_hElLO-WOrlD-"), "_HeLLo_WoRLD_")`)
}})
```

`to_consonant_case("_~-generic_~-variable_~-name_~-here-~_")` يجب أن تُعيد `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_consonant_case(s):

    return s
```

# --solutions--

```py
def to_consonant_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char == "-":
            result += "_"
        elif char.isalpha():
            if char in vowels:
                result += char.lower()
            else:
                result += char.upper()
        else:
            result += char

    return result
```
