---
id: 694596b0585c11170ac7c7fa
title: "تحدٍّ 162: استهلاك الطاقة"
challengeType: 29
dashedName: challenge-162
---

# --description--

بما أن لديك عدد السعرات الحرارية المحروقة أثناء التمرين وعدد واط-ساعات الطاقة التي استخدمتها أجهزتك الإلكترونية خلال التمرين، حدّد أيهما استهلك طاقة أكثر.

لمقارنتها، حوّل القيمتين إلى جول باستخدام التحويلات التالية:

- 1 سعر حراري يساوي 4184 جول.
- 1 واط-ساعة يساوي 3600 جول.

أرجع:

- `"Workout"` إذا استهلك التمرين طاقة أكثر.
- `"Devices"` إذا استهلك الجهاز طاقة أكثر.
- `"Equal"` إذا استهلك كلاهما نفس كمية الطاقة.

# --hints--

`compare_energy(250, 50)` يجب أن تُرجع `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` يجب أن تُرجع `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` يجب أن تُرجع `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` يجب أن تُرجع `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` يجب أن تُرجع `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` يجب أن تُرجع `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```
