---
id: 694596b0585c11170ac7c7fc
title: "تحدٍّ 164: محلل شفرة مضمَّن في Markdown"
challengeType: 29
dashedName: challenge-164
---

# --description--

أعطِ سلسلة نصية من Markdown تتضمن كتلة أو أكثر من الشفرات المضمَّنة، أعد السلسلة المكافئة بصيغة HTML.

تستخدم كتل الشفرة المضمَّنة في Markdown علامة backtick واحدة (``` ` ```) في بداية ونهاية نص الكتلة.

أعد السلسلة المعطاة مع تحويل كل كتل الشفرة إلى وسوم HTML من نوع `code`.

على سبيل المثال، بالنسبة للسلسلة النصية ```"Use `let` to declare the variable."```، أعد `"Use <code>let</code> to declare the variable."`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML في السلاسل عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

```parse_inline_code("Use `let` to declare the variable.")``` يجب أن تُعيد `"Use <code>let</code> to declare the variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` to declare the variable."), "Use <code>let</code> to declare the variable.")`)
}})
```

```parse_inline_code("Use `let` or `const` to declare a variable.")``` يجب أن تُعيد `"Use <code>let</code> or <code>const</code> to declare a variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` or \`const\` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.")`)
}})
```

```parse_inline_code("Run `npm install` then `npm start`.")``` يجب أن تُعيد `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Run \`npm install\` then \`npm start\`."), "Run <code>npm install</code> then <code>npm start</code>.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_inline_code(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_inline_code(markdown):
    return re.sub(r"`([^`]+)`", r"<code>\1</code>", markdown)
```
