---
id: 694596b0585c11170ac7c7f9
title: "تحدٍّ 161: الشحن المجاني"
challengeType: 29
dashedName: challenge-161
---

# --description--

معطاة مصفوفة من السلاسل النصية تمثل العناصر في عربة التسوق الخاصة بك، ورقم يمثل الحد الأدنى لمبلغ الطلب للتأهل للشحن المجاني، حدّد إذا كانت العناصر في عربة التسوق الخاصة بك تؤهل للشحن المجاني.

ستحتوي المصفوفة المعطاة على عناصر من القائمة أدناه:

| العنصر | السعر |
| - | - |
| `"shirt"`	| 34.25 |
| `"jeans"`	| 48.50 |
| `"shoes"` |	75.00 |
| `"hat"`	| 19.95 |
| `"socks"` | 15.00 |
| `"jacket"` | 109.95 |

# --hints--

`gets_free_shipping(["shoes"], 50)` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["shoes"], 50), True)`)
}})
```

`gets_free_shipping(["hat", "socks"], 50)` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["hat", "socks"], 50), False)`)
}})
```

`gets_free_shipping(["jeans", "shirt", "jacket"], 75)` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["jeans", "shirt", "jacket"], 75), True)`)
}})
```

`gets_free_shipping(["socks", "socks", "hat"], 75)` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["socks", "socks", "hat"], 75), False)`)
}})
```

`gets_free_shipping(["shirt", "shirt", "jeans", "socks"], 100)` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["shirt", "shirt", "jeans", "socks"], 100), True)`)
}})
```

`gets_free_shipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200)` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def gets_free_shipping(cart, minimum):

    return cart
```

# --solutions--

```py
def gets_free_shipping(cart, minimum):
    prices = {
        "shirt": 34.25,
        "jeans": 48.50,
        "shoes": 75.00,
        "hat": 19.95,
        "socks": 15.00,
        "jacket": 109.95
    }

    total = 0

    for item in cart:
        total += prices[item]

    return total >= minimum
```
