---
id: 696655d24b614176d4c9b789
title: "تحدٍّ 166: مدقق الألوان الست عشرية"
challengeType: 29
dashedName: challenge-166
---

# --description--

معطاة سلسلة نصية، حدّد ما إذا كانت لون CSS ست عشري صالحًا. يجب أن يكون لون CSS الست عشري الصالح:

- يبدأ بـ `#`،
- ويتبعه إما 3 أو 6 أحرف ست عشرية.

الأحرف الست عشرية هي الأرقام من `0` إلى `9` والحروف من `a` إلى `f` (غير حساسة لحالة الحروف).

# --hints--

يجب أن تُرجع الدالة `is_valid_hex("#123")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#123"), True)`)
}})
```

يجب أن تُرجع الدالة `is_valid_hex("#123abc")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#123abc"), True)`)
}})
```

يجب أن تُرجع الدالة `is_valid_hex("#ABCDEF")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#ABCDEF"), True)`)
}})
```

يجب أن تُرجع الدالة `is_valid_hex("#0a1B2c")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#0a1B2c"), True)`)
}})
```

يجب أن تُرجع الدالة `is_valid_hex("#12G")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#12G"), False)`)
}})
```

يجب أن تُرجع الدالة `is_valid_hex("#1234567")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#1234567"), False)`)
}})
```

يجب أن تُرجع الدالة `is_valid_hex("#12 3")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#12 3"), False)`)
}})
```

يجب أن تُرجع الدالة `is_valid_hex("fff")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("fff"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_hex(s):

    return s
```

# --solutions--

```py
def is_valid_hex(s):
    if not s.startswith("#"):
        return False

    hex_part = s[1:]
    if len(hex_part) not in (3, 6):
        return False

    for char in hex_part:
        if not char.lower() in "0123456789abcdef":
            return False

    return True
```
