---
id: 696655d24b614176d4c9b78d
title: "تحدٍّ 170: يوم فردي أم زوجي"
challengeType: 29
dashedName: challenge-170
---

# --description--

بالنظر إلى طابع زمني (عدد المللي ثانية منذ بداية عصر Unix)، أعد:

- `"odd"` إذا كان يوم الشهر لذلك الطابع الزمني فرديًا.
- `"even"` إذا كان يوم الشهر لذلك الطابع الزمني زوجيًا.

على سبيل المثال، عند إعطاء `1769472000000`، وهو طابع زمني ليوم 27 يناير 2026، أعد `"odd"` لأن رقم اليوم (`27`) عدد فردي.

ملاحظة: الطابع الزمني بوحدة المللي ثانية ويجب استخدام التاريخ في توقيت UTC، وليس في توقيتك المحلي.

# --hints--

`odd_or_even_day(1769472000000)` يجب أن يعيد `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` يجب أن يعيد `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` يجب أن يعيد `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` يجب أن يعيد `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` يجب أن يعيد `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```
