---
id: 696655d24b614176d4c9b78b
title: "تحدٍّ 168: صورة بمقياس"
challengeType: 29
dashedName: challenge-168
---

# --description--

معطى سلسلة نصية تمثّل عرض وارتفاع صورة، ورقم لتكبير الصورة، أرجع العرض والارتفاع بعد التكبير.

- السلسلة النصية المُدخلة تكون بصيغة `"WxH"`. على سبيل المثال، `"800x600"`.
- المقياس هو رقم يُضرب به العرض والارتفاع.

أرجع الأبعاد بعد التكبير بنفس صيغة `"WxH"`.

# --hints--

`scale_image("800x600", 2)` يجب أن تُرجع `"1600x1200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("800x600", 2), "1600x1200")`)
}})
```

`scale_image("100x100", 10)` يجب أن تُرجع `"1000x1000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("100x100", 10), "1000x1000")`)
}})
```

`scale_image("1024x768", 0.5)` يجب أن تُرجع `"512x384"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("1024x768", 0.5), "512x384")`)
}})
```

`scale_image("300x200", 1.5)` يجب أن تُرجع `"450x300"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("300x200", 1.5), "450x300")`)
}})
```

# --seed--

## --seed-contents--

```py
def scale_image(size, scale):

    return size
```

# --solutions--

```py
def scale_image(size, scale):
    width, height = map(int, size.split("x"))

    new_width = round(width * scale)
    new_height = round(height * scale)

    return f"{new_width}x{new_height}"
```
