---
id: 696655d24b614176d4c9b78c
title: "تحدٍّ 169: فيزباز ميني"
challengeType: 29
dashedName: challenge-169
---

# --description--

أعطِ عددًا صحيحًا، وأرجع سلسلة نصية بناءً على القواعد التالية:

- إذا كان العدد يقبل القسمة على 3، أرجع `"Fizz"`.
- إذا كان العدد يقبل القسمة على 5، أرجع `"Buzz"`.
- إذا كان العدد يقبل القسمة على كل من 3 و5، أرجع `"FizzBuzz"`.
- خلاف ذلك، أرجع العدد المعطى كسلسلة نصية.

# --hints--

يجب أن تُرجع الدالة `fizz_buzz_mini(3)` القيمة `"Fizz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(3), "Fizz")`)
}})
```

يجب أن تُرجع الدالة `fizz_buzz_mini(4)` القيمة `"4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(4), "4")`)
}})
```

يجب أن تُرجع الدالة `fizz_buzz_mini(35)` القيمة `"Buzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(35), "Buzz")`)
}})
```

يجب أن تُرجع الدالة `fizz_buzz_mini(75)` القيمة `"FizzBuzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(75), "FizzBuzz")`)
}})
```

يجب أن تُرجع الدالة `fizz_buzz_mini(98)` القيمة `"98"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(98), "98")`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_mini(n):

    return n
```

# --solutions--

```py
def fizz_buzz_mini(n):
    if n % 3 == 0 and n % 5 == 0:
        return "FizzBuzz"
    if n % 3 == 0:
        return "Fizz"
    if n % 5 == 0:
        return "Buzz"
    return str(n)
```
