---
id: 696655d24b614176d4c9b78a
title: "تحدٍّ 167: حرف بينغو!"
challengeType: 29
dashedName: challenge-167
---

# --description--

بالنظر إلى رقم، أعد الحرف الخاص بلعبة بينغو المرتبط به (بحروف كبيرة). تُجمّع أرقام بينغو كما يلي:

| الحرف | نطاق الأرقام |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

# --hints--

يجب أن تُعيد الدالة `get_bingo_letter(75)` القيمة `"O"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(75), "O")`)
}})
```

يجب أن تُعيد الدالة `get_bingo_letter(54)` القيمة `"G"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(54), "G")`)
}})
```

يجب أن تُعيد الدالة `get_bingo_letter(25)` القيمة `"I"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(25), "I")`)
}})
```

يجب أن تُعيد الدالة `get_bingo_letter(38)` القيمة `"N"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(38), "N")`)
}})
```

يجب أن تُعيد الدالة `get_bingo_letter(11)` القيمة `"B"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(11), "B")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_bingo_letter(n):

    return n
```

# --solutions--

```py
def get_bingo_letter(n):
    if 1 <= n <= 15:
        return "B"
    if 16 <= n <= 30:
        return "I"
    if 31 <= n <= 45:
        return "N"
    if 46 <= n <= 60:
        return "G"
    if 61 <= n <= 75:
        return "O"
    return None
```
