---
id: 69738771fb5a7b8b24cca2a2
title: "تحدٍّ 176: يوم الغراب"
challengeType: 29
dashedName: challenge-176
---

# --description--

اليوم هو يوم الغراب، حيث يتنبأ الغراب بالطقس بناءً على ما إذا كان يرى ظله أم لا.

اعطِ قيمة تمثل ظهور الغراب، وأرجع التنبؤ الصحيح:

- إذا كانت القيمة المعطاة هي القيمة المنطقية `true` (رأى القندس ظله)، أعد `"Looks like we'll have six more weeks of winter."`.
- إذا كانت القيمة هي القيمة المنطقية `false` (لم ير القندس ظله)، أعد `"It's going to be an early spring."`.
- إذا كانت القيمة أي شيء آخر (لم يظهر القندس)، أعد `"No prediction this year."`.

# --hints--

`groundhog_day_prediction(True)` يجب أن تُعيد `"Looks like we'll have six more weeks of winter."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(True), "Looks like we'll have six more weeks of winter.")`)
}})
```

يجب أن تُعيد الدالة `groundhog_day_prediction(False)` القيمة `"It's going to be an early spring."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(False), "It's going to be an early spring.")`)
}})
```

يجب أن تُعيد الدالة `groundhog_day_prediction(None)` القيمة `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(None), "No prediction this year.")`)
}})
```

`groundhog_day_prediction(" ")` يجب أن تُعيد القيمة `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(" "), "No prediction this year.")`)
}})
```

يجب أن تُعيد الدالة `groundhog_day_prediction("True")` القيمة `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction("True"), "No prediction this year.")`)
}})
```

# --seed--

## --seed-contents--

```py
def groundhog_day_prediction(appearance):

    return appearance
```

# --solutions--

```py
def groundhog_day_prediction(appearance):
    if appearance is True:
        return "Looks like we'll have six more weeks of winter."
    if appearance is False:
        return "It's going to be an early spring."
    return "No prediction this year."
```
