---
id: 69738771fb5a7b8b24cca2a0
title: "تحدٍّ 174: محدد البرج"
challengeType: 29
dashedName: challenge-174
---

# --description--

بالنظر إلى سلسلة نصية للتاريخ بصيغة `"YYYY-MM-DD"`، أرجع البرج الفلكي لذلك التاريخ باستخدام المخطط التالي:

| نطاق التاريخ | البرج الفلكي |
| - | - |
| 21 مارس - 19 أبريل | `"Aries"` |
| 20 أبريل - 20 مايو | `"Taurus"` |
| 21 مايو - 20 يونيو | `"Gemini"` |
| 21 يونيو - 22 يوليو | `"Cancer"` |
| 23 يوليو - 22 أغسطس | `"Leo"` |
| 23 أغسطس - 22 سبتمبر | `"Virgo"` |
| 23 سبتمبر - 22 أكتوبر | `"Libra"` |
| 23 أكتوبر - 21 نوفمبر | `"Scorpio"` |
| 22 نوفمبر - 21 ديسمبر | `"Sagittarius"` |
| 22 ديسمبر - 19 يناير | `"Capricorn"` |
| 20 يناير - 18 فبراير | `"Aquarius"` |
| 19 فبراير - 20 مارس | `"Pisces"` |

- الأبراج الفلكية تعتمد فقط على الشهر واليوم، يمكنك تجاهل السنة.

# --hints--

`get_sign("2026-01-31")` يجب أن تُرجع `"Aquarius"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2026-01-31"), "Aquarius")`)
}})
```

`get_sign("2001-06-10")` يجب أن تُرجع `"Gemini"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2001-06-10"), "Gemini")`)
}})
```

`get_sign("1985-09-07")` يجب أن تُرجع `"Virgo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1985-09-07"), "Virgo")`)
}})
```

`get_sign("2023-03-19")` يجب أن تُرجع `"Pisces"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2023-03-19"), "Pisces")`)
}})
```

`get_sign("2045-11-05")` يجب أن تُرجع `"Scorpio"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2045-11-05"), "Scorpio")`)
}})
```


`get_sign("1985-12-06")` يجب أن تُرجع `"Sagittarius"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1985-12-06"), "Sagittarius")`)
}})
```

`get_sign("2025-12-30")` يجب أن تُرجع `"Capricorn"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2025-12-30"), "Capricorn")`)
}})
```

`get_sign("2018-10-08")` يجب أن تُرجع `"Libra"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2018-10-08"), "Libra")`)
}})
```

`get_sign("1958-05-04")` يجب أن تُرجع `"Taurus"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1958-05-04"), "Taurus")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_sign(date_str):

    return date_str
```

# --solutions--

```py
def get_sign(date_str):
    _, month_str, day_str = date_str.split("-")
    month = int(month_str)
    day = int(day_str)

    if (month == 3 and day >= 21) or (month == 4 and day <= 19):
        return "Aries"
    if (month == 4 and day >= 20) or (month == 5 and day <= 20):
        return "Taurus"
    if (month == 5 and day >= 21) or (month == 6 and day <= 20):
        return "Gemini"
    if (month == 6 and day >= 21) or (month == 7 and day <= 22):
        return "Cancer"
    if (month == 7 and day >= 23) or (month == 8 and day <= 22):
        return "Leo"
    if (month == 8 and day >= 23) or (month == 9 and day <= 22):
        return "Virgo"
    if (month == 9 and day >= 23) or (month == 10 and day <= 22):
        return "Libra"
    if (month == 10 and day >= 23) or (month == 11 and day <= 21):
        return "Scorpio"
    if (month == 11 and day >= 22) or (month == 12 and day <= 21):
        return "Sagittarius"
    if (month == 12 and day >= 22) or (month == 1 and day <= 19):
        return "Capricorn"
    if (month == 1 and day >= 20) or (month == 2 and day <= 18):
        return "Aquarius"
    if (month == 2 and day >= 19) or (month == 3 and day <= 20):
        return "Pisces"
```
