---
id: 69738771fb5a7b8b24cca2a5
title: "تحدٍّ 179: النقود في الجيب"
challengeType: 29
dashedName: challenge-179
---

# --description--

معطاة مصفوفة من الأعداد الصحيحة تمثل العملات المعدنية في جيبك، حيث كل عدد صحيح هو قيمة عملة بالسنت، أرجع المبلغ الإجمالي بالتنسيق `"$D.CC"`.

- 100 سنت تساوي دولارًا واحدًا.
- في قيمة الإرجاع، أدرج صفرًا بادئًا للمبالغ الأقل من دولار واحد ودائمًا رقمين بالضبط للسنتات.

# --hints--

`count_change([25, 10, 5, 1])` يجب أن تُرجع `"$0.41"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 10, 5, 1]), "$0.41")`)
}})
```

`count_change([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` يجب أن تُرجع `"$1.43"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43")`)
}})
```

`count_change([100, 25, 100, 1000, 5, 500, 2000, 25])` يجب أن تُرجع `"$37.55"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55")`)
}})
```

`count_change([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` يجب أن تُرجع `"$0.70"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70")`)
}})
```

`count_change([1])` يجب أن تُرجع `"$0.01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([1]), "$0.01")`)
}})
```

`count_change([25, 25, 25, 25])` يجب أن تُرجع `"$1.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 25, 25, 25]), "$1.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def count_change(change):

    return change
```

# --solutions--

```py
def count_change(change):
    total_cents = sum(change)
    dollars = total_cents // 100
    cents = total_cents % 100
    return f"${dollars}.{cents:02d}"
```
