---
id: 69738771fb5a7b8b24cca2a3
title: "تحدٍّ 177: مرآة السلسلة النصية"
challengeType: 29
dashedName: challenge-177
---

# --description--

بالنظر إلى سلسلة نصية، أعد سلسلة نصية جديدة تتكون من السلسلة النصية المعطاة مع نسخة معكوسة منها ملحقة في نهايتها.

# --hints--

`mirror("freeCodeCamp")` يجب أن تُعيد `"freeCodeCamppmaCedoCeerf"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf")`)
}})
```

`mirror("RaceCar")` يجب أن تُعيد `"RaceCarraCecaR"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("RaceCar"), "RaceCarraCecaR")`)
}})
```

`mirror("helloworld")` يجب أن تُعيد `"helloworlddlrowolleh"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("helloworld"), "helloworlddlrowolleh")`)
}})
```

`mirror("The quick brown fox...")` يجب أن تُعيد `"The quick brown fox......xof nworb kciuq ehT"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT")`)
}})
```

# --seed--

## --seed-contents--

```py
def mirror(s):

    return s
```

# --solutions--

```py
def mirror(s):
    reversed_s = s[::-1]
    return s + reversed_s
```
