---
id: 69738771fb5a7b8b24cca2a4
title: "تحدٍّ 178: اقتطاع النص"
challengeType: 29
dashedName: challenge-178
---

# --description--

إذا كانت السلسلة النصية 20 حرفًا أو أقل، أعدها كما هي. إذا كانت أطول من 20 حرفًا، اقتطع أول 17 حرفًا وألحق بها `"..."` في النهاية (ليصبح المجموع 20 حرفًا) ثم أعد النتيجة.

# --hints--

`truncate_text("Hello, world!")` يجب أن تُعيد `"Hello, world!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Hello, world!"), "Hello, world!")`)
}})
```

`truncate_text("This string should get truncated.")` يجب أن تُعيد `"This string shoul..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("This string should get truncated."), "This string shoul...")`)
}})
```

`truncate_text("Exactly twenty chars")` يجب أن تُعيد `"Exactly twenty chars"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Exactly twenty chars"), "Exactly twenty chars")`)
}})
```

`truncate_text(".....................")` يجب أن تُعيد `"...................."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("....................."), "....................")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(text):

    return text
```

# --solutions--

```py
def truncate_text(text):
    if len(text) <= 20:
        return text
    return text[:17] + "..."
```
