---
id: 697a49e6ff50d756c9b6935f
title: "تحدٍّ 182: ألعاب الشتاء 2026 اليوم 3: البياثلون"
challengeType: 29
dashedName: challenge-182
---

# --description--

معطاة مصفوفة من الأعداد الصحيحة، حيث تمثل كل قيمة عدد الأهداف التي تم إصابتها في جولة واحدة من البياثلون، أرجع المسافة الإجمالية للعقوبة التي يجب على الرياضي التزلج عليها.

- تتكون كل جولة من 5 أهداف.
- كل هدف مفقود يؤدي إلى دورة عقوبة بطول 150 مترًا.

# --hints--

`calculate_penalty_distance([4, 4])` يجب أن تُرجع `300`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 4]), 300)`)
}})
```

`calculate_penalty_distance([5, 5])` يجب أن تُرجع `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 5]), 0)`)
}})
```

`calculate_penalty_distance([4, 5, 3, 5])` يجب أن تُرجع `450`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 5, 3, 5]), 450)`)
}})
```

`calculate_penalty_distance([5, 4, 5, 5])` يجب أن تُرجع `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 4, 5, 5]), 150)`)
}})
```

`calculate_penalty_distance([4, 3, 0, 3])` يجب أن تُرجع `1500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 3, 0, 3]), 1500)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_penalty_distance(rounds):

    return rounds
```

# --solutions--

```py
def calculate_penalty_distance(rounds):
    total_penalty = 0

    for hits in rounds:
        misses = 5 - hits
        total_penalty += misses * 150

    return total_penalty
```
