---
id: 697a49e6ff50d756c9b69362
title: "تحدٍّ 185: ألعاب الشتاء 2026 اليوم 6: التزلج الفني"
challengeType: 29
dashedName: challenge-185
---

# --description--

معطى مصفوفة من درجات الحكام وعقوبات اختيارية، احسب الدرجة النهائية لروتين التزلج الفني.

المعلمة الأولى هي مصفوفة من 10 درجات حكام، كل منها رقم من 0 إلى 10. احذف أعلى وأدنى درجات الحكام واجمع الدرجات الثمانية المتبقية لتحصل على الدرجة الأساسية.

أي معلمات إضافية تُمرر إلى الدالة هي عقوبات. اطرح كل العقوبات من الدرجة الأساسية لتحصل على الدرجة النهائية.

# --hints--

`compute_score([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1)` يجب أن تُرجع `64`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1), 64)`)
}})
```

`compute_score([10, 10, 10, 10, 10, 10, 10, 10, 10, 10])` يجب أن تُرجع `80`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 10, 10, 10, 10, 10, 10, 10, 10, 10]), 80)`)
}})
```

`compute_score([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1)` يجب أن تُرجع `67`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1), 67)`)
}})
```

`compute_score([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0)` يجب أن تُرجع `67.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0), 67.5)`)
}})
```

`compute_score([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5)` يجب أن تُرجع `59`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5), 59)`)
}})
```

# --seed--

## --seed-contents--

```py
def compute_score(judge_scores, *penalties):

    return judge_scores
```

# --solutions--

```py
def compute_score(judge_scores, *penalties):
    sorted_scores = sorted(judge_scores)
    base_score = sum(sorted_scores[1:9])
    total_penalty = sum(penalties)

    return base_score - total_penalty
```
