---
id: 697a49e6ff50d756c9b69364
title: "تحدٍّ 187: ألعاب الشتاء 2026 اليوم 8: التزحلق على الزلاجات"
challengeType: 29
dashedName: challenge-187
---

# --description--

معطاة مصفوفة من خمسة أرقام، كل رقم يمثل الوقت (بالثواني) الذي استغرقه المتزحلق لإكمال جزء من المسار، حدّد أي جزء كان أسرع من حيث السرعة وما هي تلك السرعة.

المسار مقسم إلى الأجزاء التالية:

- الجزء 1: 320 مترًا
- الجزء 2: 280 مترًا
- الجزء 3: 350 مترًا
- الجزء 4: 300 مترًا
- الجزء 5: 250 مترًا

القيمة الأولى في المصفوفة المعطاة تمثل الوقت للجزء 1، والقيمة الثانية للجزء 2، وهكذا.

لحساب السرعة (بالمتر لكل ثانية) لجزء معين، اقسم المسافة على الوقت.

أرجع `"The luger's fastest speed was X m/s on segment Y."`. حيث `X` هي أسرع سرعة، مقربة إلى منزلتين عشريتين، و`Y` هو رقم الجزء الذي حدثت فيه أسرع سرعة.

# --hints--

`get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128])` يجب أن تُرجع `"The luger's fastest speed was 35.07 m/s on segment 5."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.")`)
}})
```

`get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284])` يجب أن تُرجع `"The luger's fastest speed was 37.75 m/s on segment 2."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.")`)
}})
```

`get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912])` يجب أن تُرجع `"The luger's fastest speed was 38.49 m/s on segment 3."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.")`)
}})
```

`get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840])` يجب أن تُرجع `"The luger's fastest speed was 37.69 m/s on segment 1."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.")`)
}})
```

`get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508])` يجب أن تُرجع `"The luger's fastest speed was 39.24 m/s on segment 4."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_fastest_speed(times):

    return times
```

# --solutions--

```py
def get_fastest_speed(times):
    distances = [320, 280, 350, 300, 250]
    fastest_speed = 0
    fastest_segment = 0

    for i, t in enumerate(times):
        speed = distances[i] / t
        if speed > fastest_speed:
            fastest_speed = speed
            fastest_segment = i + 1  # 1-indexed

    return f"The luger's fastest speed was {fastest_speed:.2f} m/s on segment {fastest_segment}."
```
