---
id: 697a49e6ff50d756c9b6935d
title: "تحدٍّ 180: ألعاب الشتاء 2026 اليوم 1: يوم الافتتاح"
challengeType: 29
dashedName: challenge-180
---

# --description--

اليوم يبدأ ألعاب الشتاء 2026. ستتلقى خلال الأيام السبعة عشر القادمة تحديات شفرة مستوحاة منها.

في التحدٍّ الأول، يُعطى لك رمز دولة مكون من حرفين وتحتاج إلى إرجاع رمز العلم الخاص بتلك الدولة.

استخدم هذه القائمة:

| الدولة | الرمز | العلم |
| - | -: | :--: |
| ألبانيا                    |            `"AL"` | `"🇦🇱"` |
| أندورا                    |            `"AD"` | `"🇦🇩"` |
| الأرجنتين                  |            `"AR"` | `"🇦🇷"` |
| أرمينيا                    |            `"AM"` | `"🇦🇲"` |
| أستراليا                  |            `"AU"` | `"🇦🇺"` |
| النمسا                    |            `"AT"` | `"🇦🇹"` |
| أذربيجان                 |            `"AZ"` | `"🇦🇿"` |
| بلجيكا                    |            `"BE"` | `"🇧🇪"` |
| بنين                      |            `"BJ"` | `"🇧🇯"` |
| بوليفيا                    |            `"BO"` | `"🇧🇴"` |
| البوسنة والهرسك     |            `"BA"` | `"🇧🇦"` |
| البرازيل                     |            `"BR"` | `"🇧🇷"` |
| بلغاريا                   |            `"BG"` | `"🇧🇬"` |
| كندا                     |            `"CA"` | `"🇨🇦"` |
| تشيلي                      |            `"CL"` | `"🇨🇱"` |
| الصين                      |            `"CN"` | `"🇨🇳"` |
| كولومبيا                   |            `"CO"` | `"🇨🇴"` |
| كرواتيا                    |            `"HR"` | `"🇭🇷"` |
| قبرص                     |            `"CY"` | `"🇨🇾"` |
| جمهورية التشيك             |            `"CZ"` | `"🇨🇿"` |
| الدنمارك                    |            `"DK"` | `"🇩🇰"` |
| الإكوادور                    |            `"EC"` | `"🇪🇨"` |
| إريتريا                    |            `"ER"` | `"🇪🇷"` |
| إستونيا                    |            `"EE"` | `"🇪🇪"` |
| فنلندا                    |            `"FI"` | `"🇫🇮"` |
| فرنسا                     |            `"FR"` | `"🇫🇷"` |
| جورجيا                    |            `"GE"` | `"🇬🇪"` |
| ألمانيا                    |            `"DE"` | `"🇩🇪"` |
| بريطانيا العظمى              |            `"GB"` | `"🇬🇧"` |
| اليونان                     |            `"GR"` | `"🇬🇷"` |
| غينيا بيساو              |            `"GW"` | `"🇬🇼"` |
| هايتي                      |            `"HT"` | `"🇭🇹"` |
| هونغ كونغ                  |            `"HK"` | `"🇭🇰"` |
| المجر                    |            `"HU"` | `"🇭🇺"` |
| أيسلندا                    |            `"IS"` | `"🇮🇸"` |
| الهند                      |            `"IN"` | `"🇮🇳"` |
| إيران                       |            `"IR"` | `"🇮🇷"` |
| أيرلندا                    |            `"IE"` | `"🇮🇪"` |
| إسرائيل                     |            `"IL"` | `"🇮🇱"` |
| إيطاليا                      |            `"IT"` | `"🇮🇹"` |
| جامايكا                    |            `"JM"` | `"🇯🇲"` |
| اليابان                      |            `"JP"` | `"🇯🇵"` |
| كازاخستان                 |            `"KZ"` | `"🇰🇿"` |
| كينيا                      |            `"KE"` | `"🇰🇪"` |
| كوسوفو                     |            `"XK"` | `"🇽🇰"` |
| قيرغيزستان                 |            `"KG"` | `"🇰🇬"` |
| لاتفيا                     |            `"LV"` | `"🇱🇻"` |
| لبنان                    |            `"LB"` | `"🇱🇧"` |
| ليختنشتاين              |            `"LI"` | `"🇱🇮"` |
| ليتوانيا                  |            `"LT"` | `"🇱🇹"` |
| لوكسمبورغ                 |            `"LU"` | `"🇱🇺"` |
| مدغشقر                 |            `"MG"` | `"🇲🇬"` |
| ماليزيا                   |            `"MY"` | `"🇲🇾"` |
| مالطا                      |            `"MT"` | `"🇲🇹"` |
| المكسيك                     |            `"MX"` | `"🇲🇽"` |
| مولدوفا                    |            `"MD"` | `"🇲🇩"` |
| موناكو                     |            `"MC"` | `"🇲🇨"` |
| منغوليا                   |            `"MN"` | `"🇲🇳"` |
| الجبل الأسود                 |            `"ME"` | `"🇲🇪"` |
| المغرب                    |            `"MA"` | `"🇲🇦"` |
| هولندا                |            `"NL"` | `"🇳🇱"` |
| نيوزيلندا                |            `"NZ"` | `"🇳🇿"` |
| نيجيريا                    |            `"NG"` | `"🇳🇬"` |
| مقدونيا الشمالية            |            `"MK"` | `"🇲🇰"` |
| النرويج                     |            `"NO"` | `"🇳🇴"` |
| باكستان                   |            `"PK"` | `"🇵🇰"` |
| الفلبين                |            `"PH"` | `"🇵🇭"` |
| بولندا                     |            `"PL"` | `"🇵🇱"` |
| البرتغال                   |            `"PT"` | `"🇵🇹"` |
| بورتو ريكو                |            `"PR"` | `"🇵🇷"` |
| رومانيا                    |            `"RO"` | `"🇷🇴"` |
| سان مارينو                 |            `"SM"` | `"🇸🇲"` |
| السعودية               |            `"SA"` | `"🇸🇦"` |
| صربيا                     |            `"RS"` | `"🇷🇸"` |
| سنغافورة                  |            `"SG"` | `"🇸🇬"` |
| سلوفاكيا                   |            `"SK"` | `"🇸🇰"` |
| سلوفينيا                   |            `"SI"` | `"🇸🇮"` |
| جنوب أفريقيا               |            `"ZA"` | `"🇿🇦"` |
| كوريا الجنوبية                |            `"KR"` | `"🇰🇷"` |
| إسبانيا                      |            `"ES"` | `"🇪🇸"` |
| السويد                     |            `"SE"` | `"🇸🇪"` |
| سويسرا                |            `"CH"` | `"🇨🇭"` |
| تايلاند                   |            `"TH"` | `"🇹🇭"` |
| ترينيداد وتوباغو          |            `"TT"` | `"🇹🇹"` |
| تركيا                     |            `"TR"` | `"🇹🇷"` |
| أوكرانيا                    |            `"UA"` | `"🇺🇦"` |
| الإمارات العربية المتحدة       |            `"AE"` | `"🇦🇪"` |
| الولايات المتحدة              |            `"US"` | `"🇺🇸"` |
| أوروغواي                    |            `"UY"` | `"🇺🇾"` |
| أوزبكستان                 |            `"UZ"` | `"🇺🇿"` |
| فنزويلا                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`get_flag("AL")` يجب أن تُرجع `"🇦🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AL"), "🇦🇱")`)
}})
```

`get_flag("AD")` يجب أن تُرجع `"🇦🇩"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AD"), "🇦🇩")`)
}})
```

`get_flag("AR")` يجب أن تُرجع `"🇦🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AR"), "🇦🇷")`)
}})
```

`get_flag("AM")` يجب أن تُرجع `"🇦🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AM"), "🇦🇲")`)
}})
```

`get_flag("AU")` يجب أن تُرجع `"🇦🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AU"), "🇦🇺")`)
}})
```

`get_flag("AT")` يجب أن تُرجع `"🇦🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AT"), "🇦🇹")`)
}})
```

`get_flag("AZ")` يجب أن تُرجع `"🇦🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AZ"), "🇦🇿")`)
}})
```

`get_flag("BE")` يجب أن تُرجع `"🇧🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BE"), "🇧🇪")`)
}})
```

`get_flag("BJ")` يجب أن تُرجع `"🇧🇯"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BJ"), "🇧🇯")`)
}})
```

`get_flag("BO")` يجب أن تُرجع `"🇧🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BO"), "🇧🇴")`)
}})
```

`get_flag("BA")` يجب أن تُرجع `"🇧🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BA"), "🇧🇦")`)
}})
```

`get_flag("BR")` يجب أن تُرجع `"🇧🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BR"), "🇧🇷")`)
}})
```

`get_flag("BG")` يجب أن تُرجع `"🇧🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BG"), "🇧🇬")`)
}})
```

`get_flag("CA")` يجب أن تُرجع `"🇨🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CA"), "🇨🇦")`)
}})
```

`get_flag("CL")` يجب أن تُرجع `"🇨🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CL"), "🇨🇱")`)
}})
```

`get_flag("CN")` يجب أن تُرجع `"🇨🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CN"), "🇨🇳")`)
}})
```

`get_flag("CO")` يجب أن تُرجع `"🇨🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CO"), "🇨🇴")`)
}})
```

`get_flag("HR")` يجب أن تُرجع `"🇭🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HR"), "🇭🇷")`)
}})
```

`get_flag("CY")` يجب أن تُرجع `"🇨🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CY"), "🇨🇾")`)
}})
```

`get_flag("CZ")` يجب أن تُرجع `"🇨🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CZ"), "🇨🇿")`)
}})
```

`get_flag("DK")` يجب أن تُرجع `"🇩🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DK"), "🇩🇰")`)
}})
```

`get_flag("EC")` يجب أن تُرجع `"🇪🇨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EC"), "🇪🇨")`)
}})
```

`get_flag("ER")` يجب أن تُرجع `"🇪🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ER"), "🇪🇷")`)
}})
```

`get_flag("EE")` يجب أن تُرجع `"🇪🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EE"), "🇪🇪")`)
}})
```

`get_flag("FI")` يجب أن تُرجع `"🇫🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FI"), "🇫🇮")`)
}})
```

`get_flag("FR")` يجب أن تُرجع `"🇫🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FR"), "🇫🇷")`)
}})
```

`get_flag("GE")` يجب أن تُرجع `"🇬🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GE"), "🇬🇪")`)
}})
```

`get_flag("DE")` يجب أن تُرجع `"🇩🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DE"), "🇩🇪")`)
}})
```

`get_flag("GB")` يجب أن تُرجع `"🇬🇧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GB"), "🇬🇧")`)
}})
```

`get_flag("GR")` يجب أن تُرجع `"🇬🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GR"), "🇬🇷")`)
}})
```

`get_flag("GW")` يجب أن تُرجع `"🇬🇼"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GW"), "🇬🇼")`)
}})
```

`get_flag("HT")` يجب أن تُرجع `"🇭🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HT"), "🇭🇹")`)
}})
```

`get_flag("HK")` يجب أن تُرجع `"🇭🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HK"), "🇭🇰")`)
}})
```

`get_flag("HU")` يجب أن تُرجع `"🇭🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HU"), "🇭🇺")`)
}})
```

`get_flag("IS")` يجب أن تُرجع `"🇮🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IS"), "🇮🇸")`)
}})
```

`get_flag("IN")` يجب أن تُرجع `"🇮🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IN"), "🇮🇳")`)
}})
```

`get_flag("IR")` يجب أن تُرجع `"🇮🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IR"), "🇮🇷")`)
}})
```

`get_flag("IE")` يجب أن تُرجع `"🇮🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IE"), "🇮🇪")`)
}})
```

`get_flag("IL")` يجب أن تُرجع `"🇮🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IL"), "🇮🇱")`)
}})
```

`get_flag("IT")` يجب أن تُرجع `"🇮🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IT"), "🇮🇹")`)
}})
```

`get_flag("JM")` يجب أن تُرجع `"🇯🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JM"), "🇯🇲")`)
}})
```

`get_flag("JP")` يجب أن تُرجع `"🇯🇵"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JP"), "🇯🇵")`)
}})
```

`get_flag("KZ")` يجب أن تُرجع `"🇰🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KZ"), "🇰🇿")`)
}})
```

`get_flag("KE")` يجب أن تُرجع `"🇰🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KE"), "🇰🇪")`)
}})
```

`get_flag("XK")` يجب أن تُرجع `"🇽🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("XK"), "🇽🇰")`)
}})
```

`get_flag("KG")` يجب أن تُرجع `"🇰🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KG"), "🇰🇬")`)
}})
```

`get_flag("LV")` يجب أن تُرجع `"🇱🇻"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LV"), "🇱🇻")`)
}})
```

`get_flag("LB")` يجب أن تُرجع `"🇱🇧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LB"), "🇱🇧")`)
}})
```

`get_flag("LI")` يجب أن تُرجع `"🇱🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LI"), "🇱🇮")`)
}})
```

`get_flag("LT")` يجب أن تُرجع `"🇱🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LT"), "🇱🇹")`)
}})
```

`get_flag("LU")` يجب أن تُرجع `"🇱🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LU"), "🇱🇺")`)
}})
```

`get_flag("MG")` يجب أن تُرجع `"🇲🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MG"), "🇲🇬")`)
}})
```

`get_flag("MY")` يجب أن تُرجع `"🇲🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MY"), "🇲🇾")`)
}})
```

`get_flag("MT")` يجب أن تُرجع `"🇲🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MT"), "🇲🇹")`)
}})
```

`get_flag("MX")` يجب أن تُرجع `"🇲🇽"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MX"), "🇲🇽")`)
}})
```

`get_flag("MD")` يجب أن تُرجع `"🇲🇩"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MD"), "🇲🇩")`)
}})
```

`get_flag("MC")` يجب أن تُرجع `"🇲🇨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MC"), "🇲🇨")`)
}})
```

`get_flag("MN")` يجب أن تُرجع `"🇲🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MN"), "🇲🇳")`)
}})
```

`get_flag("ME")` يجب أن تُرجع `"🇲🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ME"), "🇲🇪")`)
}})
```

`get_flag("MA")` يجب أن تُرجع `"🇲🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MA"), "🇲🇦")`)
}})
```

`get_flag("NL")` يجب أن تُرجع `"🇳🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NL"), "🇳🇱")`)
}})
```

`get_flag("NZ")` يجب أن تُرجع `"🇳🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NZ"), "🇳🇿")`)
}})
```

`get_flag("NG")` يجب أن تُرجع `"🇳🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NG"), "🇳🇬")`)
}})
```

`get_flag("MK")` يجب أن تُرجع `"🇲🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MK"), "🇲🇰")`)
}})
```

`get_flag("NO")` يجب أن تُرجع `"🇳🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NO"), "🇳🇴")`)
}})
```

`get_flag("PK")` يجب أن تُرجع `"🇵🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PK"), "🇵🇰")`)
}})
```

`get_flag("PH")` يجب أن تُرجع `"🇵🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PH"), "🇵🇭")`)
}})
```

`get_flag("PL")` يجب أن تُرجع `"🇵🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PL"), "🇵🇱")`)
}})
```

`get_flag("PT")` يجب أن تُرجع `"🇵🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PT"), "🇵🇹")`)
}})
```

`get_flag("PR")` يجب أن تُرجع `"🇵🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PR"), "🇵🇷")`)
}})
```

`get_flag("RO")` يجب أن تُرجع `"🇷🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RO"), "🇷🇴")`)
}})
```

`get_flag("SM")` يجب أن تُرجع `"🇸🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SM"), "🇸🇲")`)
}})
```

`get_flag("SA")` يجب أن تُرجع `"🇸🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SA"), "🇸🇦")`)
}})
```

`get_flag("RS")` يجب أن تُرجع `"🇷🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RS"), "🇷🇸")`)
}})
```

`get_flag("SG")` يجب أن تُرجع `"🇸🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SG"), "🇸🇬")`)
}})
```

`get_flag("SK")` يجب أن تُرجع `"🇸🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SK"), "🇸🇰")`)
}})
```

`get_flag("SI")` يجب أن تُرجع `"🇸🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SI"), "🇸🇮")`)
}})
```

`get_flag("ZA")` يجب أن تُرجع `"🇿🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ZA"), "🇿🇦")`)
}})
```

`get_flag("KR")` يجب أن تُرجع `"🇰🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KR"), "🇰🇷")`)
}})
```

`get_flag("ES")` يجب أن تُرجع `"🇪🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ES"), "🇪🇸")`)
}})
```

`get_flag("SE")` يجب أن تُرجع `"🇸🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SE"), "🇸🇪")`)
}})
```

`get_flag("CH")` يجب أن تُرجع `"🇨🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CH"), "🇨🇭")`)
}})
```

`get_flag("TH")` يجب أن تُرجع `"🇹🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TH"), "🇹🇭")`)
}})
```

`get_flag("TT")` يجب أن تُرجع `"🇹🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TT"), "🇹🇹")`)
}})
```

`get_flag("TR")` يجب أن تُرجع `"🇹🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TR"), "🇹🇷")`)
}})
```

`get_flag("UA")` يجب أن تُرجع `"🇺🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UA"), "🇺🇦")`)
}})
```

`get_flag("AE")` يجب أن تُرجع `"🇦🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AE"), "🇦🇪")`)
}})
```

`get_flag("US")` يجب أن تُرجع `"🇺🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("US"), "🇺🇸")`)
}})
```

`get_flag("UY")` يجب أن تُرجع `"🇺🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UY"), "🇺🇾")`)
}})
```

`get_flag("UZ")` يجب أن تُرجع `"🇺🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UZ"), "🇺🇿")`)
}})
```

`get_flag("VE")` يجب أن تُرجع `"🇻🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("VE"), "🇻🇪")`)
}})
```



# --seed--

## --seed-contents--

```py
def get_flag(code):

    return code
```

# --solutions--

```py
def get_flag(code):
    return "".join(chr(127397 + ord(char)) for char in code)
```
