---
id: 697a49e6ff50d756c9b69366
title: "تحدٍّ 189: ألعاب الشتاء 2026 اليوم 10: التزلج على الجبال"
challengeType: 29
dashedName: challenge-189
---

# --description--

معطى الانخفاض الرأسي، والمسافة الأفقية، ونوع تل التزلج، حدّد تصنيف صعوبة التل.

لتحديد التصنيف:

- احسب انحدار التل بقسمة الانخفاض على المسافة.
- ثم احسب الانحدار المعدل بناءً على نوع التل:
  - `"Downhill"`: اضرب الانحدار في 1.2
  - `"Slalom"`: اضرب الانحدار في 0.9
  - `"Giant Slalom"`: اضرب الانحدار في 1.0

أرجع:

- `"Green"` إذا كان الانحدار المعدل أقل من أو يساوي 0.1
- `"Blue"` إذا كان الانحدار المعدل أكبر من 0.1 وأقل من أو يساوي 0.25
- `"Black"` إذا كان الانحدار المعدل أكبر من 0.25


# --hints--

`get_hill_rating(95, 900, "Slalom")` يجب أن تُرجع `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(95, 900, "Slalom"), "Green")`)
}})
```

`get_hill_rating(620, 2800, "Downhill")` يجب أن تُرجع `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(620, 2800, "Downhill"), "Black")`)
}})
```

`get_hill_rating(420, 1680, "Giant Slalom")` يجب أن تُرجع `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(420, 1680, "Giant Slalom"), "Blue")`)
}})
```

`get_hill_rating(250, 3000, "Downhill")` يجب أن تُرجع `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(250, 3000, "Downhill"), "Green")`)
}})
```

`get_hill_rating(110, 900, "Slalom")` يجب أن تُرجع `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(110, 900, "Slalom"), "Blue")`)
}})
```

`get_hill_rating(380, 1500, "Giant Slalom")` يجب أن تُرجع `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(380, 1500, "Giant Slalom"), "Black")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_hill_rating(drop, distance, hill_type):

    return drop
```

# --solutions--

```py
def get_hill_rating(drop, distance, hill_type):
    steepness = drop / distance

    if hill_type == "Downhill":
        steepness *= 1.2
    elif hill_type == "Slalom":
        steepness *= 0.9
    elif hill_type == "Giant Slalom":
        steepness *= 1.0

    if steepness <= 0.1:
        return "Green"
    elif steepness <= 0.25:
        return "Blue"
    else:
        return "Black"
```
