---
id: 697a49e6ff50d756c9b6935e
title: "تحدٍّ 181: ألعاب الشتاء 2026 اليوم 2: التزلج على اللوح"
challengeType: 29
dashedName: challenge-181
---

# --description--

بما أن وضعية المتزلج على اللوح الابتدائية معروفة والدوران بالدرجات، حدّد وضعية الهبوط الخاصة به.

- وضعية المتزلج على اللوح تكون إما `"Regular"` أو `"Goofy"`.
- دوران الحركات هو مضاعفات 90 درجة. القيم الموجبة تعني دورانًا مع عقارب الساعة، والقيم السالبة تعني دورانًا عكس عقارب الساعة.
- تتغير وضعية الهبوط كل 180 درجة من الدوران.

على سبيل المثال، إذا كانت البداية `"Regular"` والدوران `90`، فارجع `"Regular"`. وإذا كانت البداية `"Regular"` والدوران `180` درجة، فارجع `"Goofy"`.

# --hints--

`get_landing_stance("Regular", 90)` يجب أن تُرجع `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` يجب أن تُرجع `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` يجب أن تُرجع `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` يجب أن تُرجع `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` يجب أن تُرجع `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` يجب أن تُرجع `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

`get_landing_stance("Goofy", 90)` يجب أن تُرجع `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 90), "Goofy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
