---
id: 697a49e9860d24853adef67b
title: "تحدٍّ 190: ألعاب الشتاء 2026 اليوم 11: هوكي الجليد"
challengeType: 29
dashedName: challenge-190
---

# --description--

معطاة مصفوفة من 6 فرق هوكي جليد وسجلاتها بعد مباريات الدوري المجزأ، حدّد المواجهات لجولة نصف النهائي.

- كل عنصر في المصفوفة يحتوي على فريق وسجله بالشكل `"TEAM: W-OTW-OTL-L"`. حيث:
  - `"W"` عدد الانتصارات في الوقت النظامي، كل منها بقيمة 3 نقاط
  - `"OTW"` عدد الانتصارات في الوقت الإضافي، كل منها بقيمة نقطتين
  - `"OTL"` عدد الخسائر في الوقت الإضافي، كل منها بقيمة نقطة واحدة
  - `"L"` عدد الخسائر، كل منها بقيمة 0 نقاط
  
على سبيل المثال، `"FIN: 2-2-1-0"` سيكون لديه 11 نقطة بعد جمع سجله.
  
احسب إجمالي النقاط لكل فريق وأرجع `"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."`. على سبيل المثال، `"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`get_semifinal_matchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])` يجب أن تُرجع `"The semi-final games will be FIN vs SWE and CAN vs USA."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.")`)
}})
```

`get_semifinal_matchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])` يجب أن تُرجع `"The semi-final games will be USA vs CZE and CAN vs FIN."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.")`)
}})
```

`get_semifinal_matchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])` يجب أن تُرجع `"The semi-final games will be CAN vs ITA and USA vs CZE."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.")`)
}})
```

`get_semifinal_matchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])` يجب أن تُرجع `"The semi-final games will be JPN vs ITA and AUT vs KOR."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_semifinal_matchups(teams):

    return teams
```

# --solutions--

```py
def get_semifinal_matchups(teams):
    team_points = []
    for team_str in teams:
        name, record = team_str.split(": ")
        W, OTW, OTL, L = map(int, record.split("-"))
        points = W * 3 + OTW * 2 + OTL * 1 + L * 0
        team_points.append((name, points))

    team_points.sort(key=lambda x: x[1], reverse=True)
    first, second, third, fourth = team_points[:4]

    return f"The semi-final games will be {first[0]} vs {fourth[0]} and {second[0]} vs {third[0]}."
```
