---
id: 697a49e9860d24853adef680
title: "تحدٍّ 195: ألعاب الشتاء 2026 اليوم 16: الكيرلنغ"
challengeType: 29
dashedName: challenge-195
---

# --description--

معطى مصفوفة 5x5 تمثل "البيت" في نهاية جولة الكيرلنغ، حدّد الفريق الذي يسجل وعدد النقاط التي يسجلها.

التخطيط:

- الخلية المركزية (الفهرس `[2, 2]`) هي "الزر".
- الثماني خلايا المحيطة مباشرة بالزر تمثل الحلقة 1.
- والـ 16 خلية على الحافة الخارجية للبيت تمثل الحلقة 2.

في المصفوفة المعطاة:

- `"."` تمثل مساحة فارغة.
- `"R"` تمثل مساحة بها حجر أحمر.
- `"Y"` تمثل مساحة بها حجر أصفر.

قواعد التسجيل:

- يمكن لفريق واحد فقط التسجيل في كل جولة.
- الفريق الذي يملك الحجر الأقرب إلى الزر يسجل.
- الفريق الذي يسجل يحصل على نقطة واحدة عن كل حجر له أقرب إلى الزر من أقرب حجر للفريق المنافس.
- كلما كان رقم الحلقة أقل، كان الحجر أقرب إلى المركز.
- إذا كان أقرب حجر لكل فريق على نفس المسافة من المركز، فلا يسجل أي فريق.

الإرجاع:

- سلسلة نصية بصيغة `"team: number_of_points"`، مثل: `"R: 2"`.
- أو `"No points awarded"` إذا لم يسجل أي فريق نقاطًا.

على سبيل المثال، معطى:

```js
[
  [".", ".", "R", ".", "."],
  [".", "R", ".", ".", "."],
  ["Y", ".", ".", ".", "."],
  [".", "R", ".", ".", "."],
  [".", ".", ".", ".", "."]
]
```

أرجع `"R: 2"`. الحجرتان الحمراء في الحلقة 1 متساويتان في القرب وهما الحجرتان الوحيدتان الأقرب من أقرب حجر أصفر.

# --hints--

`score_curling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]])` يجب أن تُرجع `"R:2"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]]), "R: 2")`)
}})
```

`score_curling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]])` يجب أن تُرجع `"Y: 3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]]), "Y: 3")`)
}})
```

`score_curling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]])` يجب أن تُرجع `"No points awarded"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]]), "No points awarded")`)
}})
```

`score_curling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]])` يجب أن تُرجع `"R: 4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]]), "R: 4")`)
}})
```

`score_curling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]])` يجب أن تُرجع `"Y: 1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]]), "Y: 1")`)
}})
```

`score_curling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]])` يجب أن تُرجع `"No points awarded"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]]), "No points awarded")`)
}})
```

# --seed--

## --seed-contents--

```py
def score_curling(house):

    return house
```

# --solutions--

```py
def score_curling(house):
  stones = {'R': [], 'Y': []}
  
  for i in range(5):
    for j in range(5):
      cell = house[i][j]
      if cell == 'R' or cell == 'Y':
        distance = max(abs(i - 2), abs(j - 2))
        stones[cell].append(distance)
  
  red_closest = min(stones['R']) if stones['R'] else float('inf')
  yellow_closest = min(stones['Y']) if stones['Y'] else float('inf')
  
  if red_closest == float('inf') and yellow_closest == float('inf'):
    return "No points awarded"
  
  if red_closest == yellow_closest:
    return "No points awarded"
  
  if red_closest < yellow_closest:
    winner = 'R'
    opponent_closest = yellow_closest
    winning_stones = stones['R']
  else:
    winner = 'Y'
    opponent_closest = red_closest
    winning_stones = stones['Y']
  
  points = sum(1 for d in winning_stones if d < opponent_closest)
  
  return "No points awarded" if points == 0 else f"{winner}: {points}"
```
