---
id: 697a49e9860d24853adef67e
title: "تحدٍّ 193: ألعاب الشتاء 2026 اليوم 14: تسلق الجبال على الثلج"
challengeType: 29
dashedName: challenge-193
---

# --description--

بناءً على عمق الثلج وانحدار الجبل، حدّد إذا كان هناك خطر انزلاق جليدي.

- قيم عمق الثلج هي `"Shallow"`، `"Moderate"`، أو `"Deep"`.
- قيم الانحدار هي `"Gentle"`، `"Steep"`، أو `"Very Steep"`.

أرجع `"Safe"` أو `"Risky"` بناءً على هذا الجدول:

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalanche_risk("Shallow", "Gentle")` يجب أن تُرجع `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Shallow", "Steep")` يجب أن تُرجع `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Steep"), "Safe")`)
}})
```

`avalanche_risk("Shallow", "Very Steep")` يجب أن تُرجع `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Very Steep"), "Safe")`)
}})
```

`avalanche_risk("Moderate", "Gentle")` يجب أن تُرجع `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Moderate", "Steep")` يجب أن تُرجع `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Steep"), "Risky")`)
}})
```

`avalanche_risk("Moderate", "Very Steep")` يجب أن تُرجع `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Very Steep"), "Risky")`)
}})
```

`avalanche_risk("Deep", "Gentle")` يجب أن تُرجع `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Deep", "Steep")` يجب أن تُرجع `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Steep"), "Risky")`)
}})
```

`avalanche_risk("Deep", "Very Steep")` يجب أن تُرجع `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Very Steep"), "Risky")`)
}})
```

# --seed--

## --seed-contents--

```py
def avalanche_risk(snow_depth, slope):

    return snow_depth
```

# --solutions--

```py
def avalanche_risk(snow_depth, slope):
    risk_table = {
        "Gentle": {
            "Shallow": "Safe",
            "Moderate": "Safe",
            "Deep": "Safe",
        },
        "Steep": {
            "Shallow": "Safe",
            "Moderate": "Risky",
            "Deep": "Risky",
        },
        "Very Steep": {
            "Shallow": "Safe",
            "Moderate": "Risky",
            "Deep": "Risky",
        },
    }

    return risk_table[slope][snow_depth]
```
