---
id: 697a49e6ff50d756c9b69365
title: "تحدٍّ 188: ألعاب الشتاء 2026 اليوم 9: التزلج الهيكلي"
challengeType: 29
dashedName: challenge-188
---

# --description--

أعطيت سلسلة نصية تمثل المنحنيات على مسار التزلج الهيكلي، حدّد صعوبة المسار.

- تتكون السلسلة النصية المعطاة فقط من الحروف:
  - `"L"` للدوران إلى اليسار
  - `"R"` للدوران إلى اليمين
  - `"S"` لجزء مستقيم

- كل تغيير في الاتجاه يضيف 15 نقطة (عندما يتبع `"L"` حرف `"R"` أو العكس).
- كل المنحنيات الأخرى تضيف 5 نقاط لكل منها (كل حروف `"L"` أو `"R"` الأخرى).
- الأجزاء المستقيمة تضيف 0 نقاط.

تعتمد صعوبة المسار على المجموع الكلي للنقاط. أرجع:

- `"Easy"` إذا كان المجموع من 0 إلى 100
- `"Medium"` إذا كان المجموع من 101 إلى 200
- `"Hard"` إذا كان المجموع أكثر من 200

# --hints--

`get_difficulty("SLSLLSRRLSRLRL")` يجب أن تُرجع `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLSLLSRRLSRLRL"), "Easy")`)
}})
```

`get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` يجب أن تُرجع `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard")`)
}})
```

`get_difficulty("SRRRRLSLLRLRSSRLSRL")` يجب أن تُرجع `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRRRRLSLLRLRSSRLSRL"), "Medium")`)
}})
```

`get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` يجب أن تُرجع `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard")`)
}})
```

`get_difficulty("SLLSSLRLSLSLRSLSSLRL")` يجب أن تُرجع `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium")`)
}})
```

`get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` يجب أن تُرجع `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_difficulty(track):

    return track
```

# --solutions--

```py
def get_difficulty(track):
  score = 0

  for i in range(len(track)):
    current = track[i]
    previous = track[i - 1] if i > 0 else None

    if current == 'S':
      score += 0
    elif current == 'L' or current == 'R':
      if previous and previous != 'S' and previous != current:
        score += 15
      else:
        score += 5

  if score <= 100:
    return 'Easy'
  elif score <= 200:
    return 'Medium'
  else:
    return 'Hard'
```
