---
id: 697a49e9860d24853adef67c
title: "تحدٍّ 191: ألعاب الشتاء 2026 اليوم 12: التزحلق على الزلاجات"
challengeType: 29
dashedName: challenge-191
---

# --description--

معطى مصفوفة تمثل أوزان الرياضيين في فريق التزحلق على الزلاجات ورقم يمثل وزن الزلاجة، حدّد ما إذا كان الفريق مؤهلاً للمشاركة في السباق.

- طول المصفوفة يحدد حجم الفريق: فرق من 1 أو 2 أو 4 أشخاص.
- جميع قيم الأوزان المعطاة بالكيلوغرام (كغ).

يجب أن يكون للزلاجة (الزلاجة وحدها) وزن الحد الأدنى التالي:

- 162 كغ لفريق من شخص واحد
- 170 كغ لفريق من شخصين
- 210 كغ لفريق من أربعة أشخاص

يجب ألا يتجاوز الوزن الكلي للزلاجة (الرياضيون بالإضافة إلى الزلاجة):

- 247 كغ لفريق من شخص واحد
- 390 كغ لفريق من شخصين
- 630 كغ لفريق من أربعة أشخاص

أرجع `"Eligible"` إذا استوفى الفريق جميع المتطلبات، أو `"Not Eligible"` إذا فشل الفريق في استيفاء واحد أو أكثر من المتطلبات.

# --hints--

`check_eligibility([78], 165)` يجب أن تُرجع `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([78], 165), "Eligible")`)
}})
```

`check_eligibility([80], 160)` يجب أن تُرجع `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 160), "Not Eligible")`)
}})
```

`check_eligibility([80], 170)` يجب أن تُرجع `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 170), "Not Eligible")`)
}})
```

`check_eligibility([85, 90], 170)` يجب أن تُرجع `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 90], 170), "Eligible")`)
}})
```

`check_eligibility([85, 95], 168)` يجب أن تُرجع `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 95], 168), "Not Eligible")`)
}})
```

`check_eligibility([112, 97], 185)` يجب أن تُرجع `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([112, 97], 185), "Not Eligible")`)
}})
```

`check_eligibility([110, 102, 90, 106], 222)` يجب أن تُرجع `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([110, 102, 90, 106], 222), "Eligible")`)
}})
```

`check_eligibility([106, 99, 90, 88], 205)` يجب أن تُرجع `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 90, 88], 205), "Not Eligible")`)
}})
```

`check_eligibility([106, 99, 103, 96], 227)` يجب أن تُرجع `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 103, 96], 227), "Not Eligible")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_eligibility(athlete_weights, sled_weight):

    return athlete_weights
```

# --solutions--

```py
def check_eligibility(athlete_weights, sled_weight):
    team_size = len(athlete_weights)

    rules = {
        1: {"min_sled": 162, "max_total": 247},
        2: {"min_sled": 170, "max_total": 390},
        4: {"min_sled": 210, "max_total": 630}
    }

    min_sled = rules[team_size]["min_sled"]
    max_total = rules[team_size]["max_total"]

    total_weight = sum(athlete_weights) + sled_weight

    if sled_weight < min_sled:
        return "Not Eligible"

    if total_weight > max_total:
        return "Not Eligible"

    return "Eligible"
```
