---
id: 698a1a73ade5ac0e19180fa2
title: "تحدٍّ 199: الفرق التتابعي"
challengeType: 29
dashedName: challenge-199
---

# --description--

معطاة مصفوفة من الأرقام، أرجع مصفوفة جديدة تحتوي على القيمة اللازمة للانتقال من كل رقم إلى الرقم التالي.

- للرقم الأخير، استخدم `0` لأنه لا يوجد رقم تالي.

على سبيل المثال، إذا كانت المصفوفة `[1, 2, 4, 7]`، أرجع `[1, 2, 3, 0]`.

# --hints--

يجب أن تُرجع الدالة `find_differences([1, 2, 4, 7])` القيمة `[1, 2, 3, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([1, 2, 4, 7]), [1, 2, 3, 0])`)
}})
```

يجب أن تُرجع الدالة `find_differences([10, 15, 19, 22, 24, 25])` القيمة `[5, 4, 3, 2, 1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([10, 15, 19, 22, 24, 25]), [5, 4, 3, 2, 1, 0])`)
}})
```

يجب أن تُرجع الدالة `find_differences([25, 20, 16, 13, 11, 10])` القيمة `[-5, -4, -3, -2, -1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([25, 20, 16, 13, 11, 10]), [-5, -4, -3, -2, -1, 0])`)
}})
```

يجب أن تُرجع الدالة `find_differences([0, 1, 2, 2, 3, 3, 4, 5])` القيمة `[1, 1, 0, 1, 0, 1, 1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([0, 1, 2, 2, 3, 3, 4, 5]), [1, 1, 0, 1, 0, 1, 1, 0])`)
}})
```

يجب أن تُرجع الدالة `find_differences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1])` القيمة `[1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1]), [1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_differences(arr):

    return arr
```

# --solutions--

```py
def find_differences(arr):
    result = []

    for i in range(len(arr)):
        if i == len(arr) - 1:
            result.append(0)
        else:
            result.append(arr[i + 1] - arr[i])

    return result
```
