---
id: 698a1a73ade5ac0e19180fa9
title: "تحدٍّ 206: قيم أوراق اللعب"
challengeType: 29
dashedName: challenge-206
---

# --description--

معطاة مصفوفة من أوراق اللعب، أعد مصفوفة جديدة تحتوي على القيمة الرقمية لكل ورقة.

قيم الأوراق:

- الآس (`"A"`) قيمته 1.
- الأوراق المرقمة (`"2"` - `"10"`) قيمتها تساوي الرقم الظاهر عليها: من 2 إلى 10 على التوالي.
- أوراق الوجه: الولد (`"J"`)، الملكة (`"Q"`)، والملك (`"K"`) كل منها قيمته 10.

الرموز:

- كل ورقة لها رمز: البستوني (`"S"`)، الكبة (`"C"`)، الماس (`"D"`)، أو القلوب (`"H"`).

صيغة الورقة:

- تمثّل كل ورقة كسلسلة نصية: `"valueSuit"`. على سبيل المثال: `"AS"` هو آس البستوني، `"10H"` هو عشرة القلوب، و`"QC"` هي ملكة الكبة.

# --hints--

`card_values(["3H", "4D", "5S"])` يجب أن تُعيد `[3, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["3H", "4D", "5S"]), [3, 4, 5])`)
}})
```

`card_values(["AS", "10S", "10H", "6D", "7D"])` يجب أن تُعيد `[1, 10, 10, 6, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7])`)
}})
```

`card_values(["8D", "QS", "2H", "JC", "9C"])` يجب أن تُعيد `[8, 10, 2, 10, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9])`)
}})
```

`card_values(["AS", "KS"])` يجب أن تُعيد `[1, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "KS"]), [1, 10])`)
}})
```

`card_values(["10H", "JH", "QH", "KH", "AH"])` يجب أن تُعيد `[10, 10, 10, 10, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def card_values(cards):

    return cards
```

# --solutions--

```py
def card_values(cards):
    values = []
    for card in cards:
        value_str = card[:-1]

        if value_str == "A":
            value = 1
        elif value_str in ["J", "Q", "K"]:
            value = 10
        else:
            value = int(value_str)

        values.append(value)
    return values
```
