---
id: 698a1a73ade5ac0e19180fa8
title: "تحدٍّ 205: عد المكعبات المثالية"
challengeType: 29
dashedName: challenge-205
---

# --description--

معطى عددان صحيحان، حدّد كم عدد المكعبات المثالية الموجودة في النطاق بين العددين وشاملهما.

- لا يضمن أن يكون العدد الأصغر هو المعلمة الأولى.
- العدد مكعب مثالي إذا وُجد عدد صحيح (`n`) بحيث `n * n * n = number`. على سبيل المثال، 27 هو مكعب مثالي لأن `3 * 3 * 3 = 27`.

# --hints--

يجب أن تُرجع الدالة `count_perfect_cubes(3, 30)` القيمة `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(3, 30), 2)`)
}})
```

يجب أن تُرجع الدالة `count_perfect_cubes(1, 30)` القيمة `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(1, 30), 3)`)
}})
```

يجب أن تُرجع الدالة `count_perfect_cubes(30, 0)` القيمة `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(30, 0), 4)`)
}})
```

يجب أن تُرجع الدالة `count_perfect_cubes(-64, 64)` القيمة `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(-64, 64), 9)`)
}})
```

يجب أن تُرجع الدالة `count_perfect_cubes(9214, -8127)` القيمة `41`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(9214, -8127), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_perfect_cubes(a, b):

    return a
```

# --solutions--

```py
def count_perfect_cubes(a, b):
    start, end = min(a, b), max(a, b)
    count = 0

    n = 0
    while n ** 3 <= end:
        if n ** 3 >= start:
            count += 1
        n += 1

    n = -1
    while n ** 3 >= start:
        if n ** 3 <= end:
            count += 1
        n -= 1

    return count
```
