---
id: 698a1a863194f1f4e63f645e
title: "تحدٍّ 207: أصغر فجوة"
challengeType: 29
dashedName: challenge-207
---

# --description--

معطاة سلسلة نصية، أرجع الجزء الفرعي بين الحرفين المتطابقين اللذين يفصل بينهما أصغر عدد من الأحرف (أصغر فجوة).

- سيكون هناك دائمًا زوج واحد على الأقل من الأحرف المتطابقة.
- يجب أن يستبعد الجزء الفرعي المرجع الأحرف المتطابقة.
- إذا تساوت فجوتان أو أكثر في الطول، أرجع الأحرف من الفجوة الأولى.

على سبيل المثال، معطى `"ABCDAC"`، أرجع `"DA"` لأن:

- فقط `"A"` و`"C"` يتكرران في السلسلة النصية.
- عدد الأحرف بين الحرفين `"A"` هو 3، وبين الحرفين `"C"` هو 2.
- لذلك أرجع السلسلة النصية بين الحرفين `"C"`.

# --hints--

`smallest_gap("ABCDAC")` يجب أن تُرجع `"DA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("ABCDAC"), "DA")`)
}})
```

`smallest_gap("racecar")` يجب أن تُرجع `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("racecar"), "e")`)
}})
```

`smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` يجب أن تُرجع `"#q6e&rkf(p"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p")`)
}})
```

`smallest_gap("Hello World")` يجب أن تُرجع `""`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("Hello World"), "")`)
}})
```

`smallest_gap("The quick brown fox jumps over the lazy dog.")` يجب أن تُرجع `"fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("The quick brown fox jumps over the lazy dog."), "fox")`)
}})
```

# --seed--

## --seed-contents--

```py
def smallest_gap(s):

    return s
```

# --solutions--

```py
def smallest_gap(s):
    min_gap = float("inf")
    result = ""

    for i in range(len(s)):
        for j in range(i + 1, len(s)):
            if s[i] == s[j]:
                gap = j - i - 1

                if gap < min_gap:
                    min_gap = gap
                    result = s[i + 1:j]

                break

    return result
```
