---
id: 698a1a863194f1f4e63f6460
title: "تحدٍّ 209: حجم العنصر"
challengeType: 29
dashedName: challenge-209
---

# --description--

معطى حجم النافذة، عرض عنصر بوحدات عرض الإطار `"vw"`، وارتفاع عنصر بوحدات ارتفاع الإطار `"vh"`، حدّد حجم العنصر بالبكسل.

- حجم النافذة المعطى وحجم العنصر المرجع عبارة عن سلاسل نصية بصيغة `"width x height"`، مثلًا `"1200 x 800"`.

- وحدات `"vw"` تمثل نسبة مئوية من عرض النافذة. على سبيل المثال، `"50vw"` تعني 50% من عرض النافذة.
- وحدات `"vh"` تمثل نسبة مئوية من ارتفاع النافذة. على سبيل المثال، `"50vh"` تعني 50% من ارتفاع النافذة.

# --hints--

`get_element_size("1200 x 800", "50vw", "50vh")` يجب أن تُرجع `"600 x 400"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "50vw", "50vh"), "600 x 400")`)
}})
```

`get_element_size("320 x 480", "25vw", "50vh")` يجب أن تُرجع `"80 x 240"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("320 x 480", "25vw", "50vh"), "80 x 240")`)
}})
```

`get_element_size("1000 x 500", "7vw", "3vh")` يجب أن تُرجع `"70 x 15"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1000 x 500", "7vw", "3vh"), "70 x 15")`)
}})
```

`get_element_size("1920 x 1080", "95vw", "100vh")` يجب أن تُرجع `"1824 x 1080"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1920 x 1080", "95vw", "100vh"), "1824 x 1080")`)
}})
```

`get_element_size("1200 x 800", "0vw", "0vh")` يجب أن تُرجع `"0 x 0"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "0vw", "0vh"), "0 x 0")`)
}})
```

`get_element_size("1440 x 900", "100vw", "114vh")` يجب أن تُرجع `"1440 x 1026"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1440 x 900", "100vw", "114vh"), "1440 x 1026")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_element_size(window_size, element_vw, element_vh):

    return window_size
```

# --solutions--

```py
def get_element_size(window_size, element_vw, element_vh):
    window_width, window_height = map(int, window_size.split(" x "))

    vw = int(element_vw.replace("vw", ""))
    vh = int(element_vh.replace("vh", ""))

    width_px = round((vw / 100) * window_width)
    height_px = round((vh / 100) * window_height)

    return f"{width_px} x {height_px}"
```
