---
id: 698a1a73ade5ac0e19180fa6
title: "تحدٍّ 203: مصفوفة مسطحة"
challengeType: 29
dashedName: challenge-203
---

# --description--

معطاة مصفوفة، حدّد إذا كانت مسطحة.

- المصفوفة مسطحة إذا لم يكن أي من عناصرها مصفوفات.

# --hints--

`is_flat([1, 2, 3, 4])` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, 2, 3, 4]), True)`)
}})
```

`is_flat([1, [2, 3], 4])` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, [2, 3], 4]), False)`)
}})
```

`is_flat([1, 0, False, True, "a", "b"])` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, 0, False, True, "a", "b"]), True)`)
}})
```

`is_flat(["a", [0], "b", True])` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat(["a", [0], "b", True]), False)`)
}})
```

`is_flat([1, [2, [3, [4, [5]]]], 6])` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, [2, [3, [4, [5]]]], 6]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_flat(arr):

    return arr
```

# --solutions--

```py
def is_flat(arr):
    for el in arr:
        if isinstance(el, list):
            return False
    return True
```
