---
id: 698a1a73ade5ac0e19180fa5
title: "تحدٍّ 202: إضافة علامات الترقيم"
challengeType: 29
dashedName: challenge-202
---

# --description--

أعطِ سلسلة نصية من الجمل التي تفتقد النقاط، أضف نقطة (`"."`) في الأماكن التالية:

- قبل كل فراغ يأتي مباشرةً قبل حرف كبير
- وفي نهاية السلسلة النصية

أرجع السلسلة النصية الناتجة.

# --hints--

`add_punctuation("Hello world")` يجب أن تُرجع `"Hello world."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Hello world"), "Hello world.")`)
}})
```

`add_punctuation("Hello world It's nice today")` يجب أن تُرجع `"Hello world. It's nice today."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Hello world It's nice today"), "Hello world. It's nice today.")`)
}})
```

`add_punctuation("JavaScript is great Sometimes")` يجب أن تُرجع `"JavaScript is great. Sometimes."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("JavaScript is great Sometimes"), "JavaScript is great. Sometimes.")`)
}})
```

`add_punctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z")` يجب أن تُرجع `"A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z"), "A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z.")`)
}})
```

`add_punctuation("Wait.. For it")` يجب أن تُرجع `"Wait... For it."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Wait.. For it"), "Wait... For it.")`)
}})
```

# --seed--

## --seed-contents--

```py
def add_punctuation(sentences):

    return sentences
```

# --solutions--

```py
def add_punctuation(sentences):
    result = ""
    length = len(sentences)

    for i, c in enumerate(sentences):
        if c == " " and i + 1 < length and sentences[i + 1].isupper():
            result += "."
        result += c

    if not result.endswith("."):
        result += "."

    return result
```
