---
id: 698a1a73ade5ac0e19180fa0
title: "تحدٍّ 197: توافق فصيلة الدم"
challengeType: 29
dashedName: challenge-197
---

# --description--

معطى فصيلة دم المتبرع وفصيلة دم المتلقي، حدّد ما إذا كان المتبرع يمكنه التبرع بالدم للمتلقي.

تتكون كل فصيلة دم من:

- حرف: `"A"`، `"B"`، `"AB"`، أو `"O"`
- وعامل Rh: `"+"` أو `"-"`

فصائل الدم ستكون واحدة من الحروف الصحيحة متبوعة بعامل Rh. على سبيل المثال، `"AB+"` و`"O-"` هما فصائل دم صحيحة.

قواعد الحروف:

- `"O"` يمكنه التبرع لأنواع الحروف الأخرى.
- `"A"` يمكنه التبرع لـ `"A"` و`"AB"`.
- `"B"` يمكنه التبرع لـ `"B"` و`"AB"`.
- `"AB"` يمكنه التبرع فقط لـ `"AB"`.

قواعد عامل Rh:

- السالب (`"-"`) يمكنه التبرع لكل من `"-"` و`"+"`.
- الموجب (`"+"`) يمكنه التبرع فقط لـ `"+"`.

يجب أن تنطبق كل من قاعدة الحرف وقاعدة عامل Rh ليتمكن المتبرع من التبرع للمتلقي.

# --hints--

`can_donate("B+", "B+")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B+", "B+"), True)`)
}})
```

`can_donate("O-", "AB-")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "AB-"), True)`)
}})
```

`can_donate("O+", "A-")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "A-"), False)`)
}})
```

`can_donate("A+", "AB+")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A+", "AB+"), True)`)
}})
```

`can_donate("A-", "B-")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A-", "B-"), False)`)
}})
```

`can_donate("B-", "AB+")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "AB+"), True)`)
}})
```

`can_donate("B-", "A+")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "A+"), False)`)
}})
```

`can_donate("O-", "O+")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "O+"), True)`)
}})
```

`can_donate("O+", "O-")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "O-"), False)`)
}})
```

`can_donate("AB+", "AB-")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("AB+", "AB-"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_donate(donor, recipient):

    return donor
```

# --solutions--

```py
def can_donate(donor, recipient):
    donor_type = donor[:-1]
    donor_rh = donor[-1]

    recipient_type = recipient[:-1]
    recipient_rh = recipient[-1]

    abo_compatibility = {
        "O": ["A", "B", "AB", "O"],
        "A": ["A", "AB"],
        "B": ["B", "AB"],
        "AB": ["AB"]
    }

    abo_match = recipient_type in abo_compatibility[donor_type]
    rh_match = (
        donor_rh == "-" or
        (donor_rh == "+" and recipient_rh == "+")
    )

    return abo_match and rh_match
```
