---
id: 698a1a73ade5ac0e19180fa1
title: "تحدٍّ 198: حساب أيام العمل"
challengeType: 29
dashedName: challenge-198
---

# --description--

بما أن لديك تاريخ بداية وتاريخ نهاية، أرجع عدد أيام العمل بينهما.

- التواريخ المعطاة بصيغة `"YYYY-MM-DD"`.
- أيام الأسبوع هي أيام العمل (من الاثنين إلى الجمعة).
- عطلات نهاية الأسبوع ليست أيام عمل (السبت والأحد).
- احسب مع تضمين تاريخ البداية وتاريخ النهاية.

# --hints--

`count_business_days("2026-02-24", "2026-02-26")` يجب أن تُرجع `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2026-02-26"), 3)`)
}})
```

`count_business_days("2026-02-24", "2026-02-28")` يجب أن تُرجع `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2026-02-28"), 4)`)
}})
```

`count_business_days("2026-02-21", "2026-03-01")` يجب أن تُرجع `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-21", "2026-03-01"), 5)`)
}})
```

`count_business_days("2026-03-08", "2026-03-17")` يجب أن تُرجع `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-03-08", "2026-03-17"), 7)`)
}})
```

`count_business_days("2026-02-24", "2027-02-24")` يجب أن تُرجع `262`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2027-02-24"), 262)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_business_days(start, end):

    return start
```

# --solutions--

```py
from datetime import datetime, timedelta

def count_business_days(start, end):
    start_date = datetime.strptime(start, "%Y-%m-%d")
    end_date = datetime.strptime(end, "%Y-%m-%d")
    
    count = 0
    current = start_date
    
    while current <= end_date:
        if current.weekday() < 5:
            count += 1
        current += timedelta(days=1)
    
    return count
```
