---
id: 698a1a863194f1f4e63f645f
title: "تحدٍّ 208: تتبع المسار"
challengeType: 29
dashedName: challenge-208
---

# --description--

معطاة مصفوفة من السلاسل النصية تمثل خريطة مسارك، أرجع سلسلة نصية بالحركات اللازمة للوصول إلى هدفك.

السلاسل النصية المعطاة ستحتوي على القيم التالية:

- `"C"`: موقعك الحالي
- `"G"`: هدفك
- `"T"`: أجزاء من المسار يمكن عبورها
- `"-"`: أجزاء من الخريطة لا يمكن عبورها

أرجع سلسلة نصية بالحركات اللازمة لاتباع المسار من موقعك إلى هدفك حيث:

- `"R"` تعني حركة إلى اليمين
- `"D"` تعني حركة إلى الأسفل
- `"L"` تعني حركة إلى اليسار
- `"U"` تعني حركة إلى الأعلى

- سيكون هناك دائمًا مسار واحد مستمر، بدون أي تفرعات، من موقعك الحالي إلى هدفك.
- كل موقع على المسار سيتصل بحد أقصى موقعين يمكن عبورهما.

على سبيل المثال، معطى:

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

أرجع `"RDDRDD"`.

# --hints--

`navigate_trail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` يجب أن تُرجع `"RDDRDD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD")`)
}})
```

`navigate_trail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` يجب أن تُرجع `"RRUUURR"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR")`)
}})
```

`navigate_trail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` يجب أن تُرجع `"DLDDRRRRD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD")`)
}})
```

`navigate_trail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` يجب أن تُرجع `"RRRDDL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL")`)
}})
```

`navigate_trail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` يجب أن تُرجع `"ULLLUUURRRRRRDDDR"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate_trail(map):

    return map
```

# --solutions--

```py
def navigate_trail(map):
    rows = len(map)
    cols = len(map[0])

    directions = [
        (0, 1, "R"),
        (1, 0, "D"),
        (0, -1, "L"),
        (-1, 0, "U"),
    ]

    for r in range(rows):
        for c in range(cols):
            if map[r][c] == "C":
                row, col = r, c

    result = ""
    prev = None

    while map[row][col] != "G":
        for dr, dc, move in directions:
            new_row = row + dr
            new_col = col + dc

            if (
                0 <= new_row < rows and
                0 <= new_col < cols and
                (map[new_row][new_col] == "T" or
                 map[new_row][new_col] == "G") and
                (not prev or (new_row, new_col) != prev)
            ):
                result += move
                prev = (row, col)
                row, col = new_row, new_col
                break

    return result
```
