---
id: 698a1a73ade5ac0e19180fa3
title: "تحدٍّ 200: عد الحروف والأرقام"
challengeType: 29
dashedName: challenge-200
---

# --description--

بالنظر إلى سلسلة نصية، أرجع رسالة تحتوي على عدد الحروف والأرقام التي تحتويها.

- الحروف هي `A-Z` و `a-z`.
- الأرقام هي `0-9`.
- تجاهل كل الأحرف الأخرى.

أرجع `"The string has X letters and Y numbers."`، حيث `"X"` هو عدد الحروف و`"Y"` هو عدد الأرقام. إذا كان أي من العددين 1، استخدم صيغة المفرد لذلك العنصر. مثلاً: `"1 letter"` بدلًا من `"1 letters"` و`"1 number"` بدلًا من `"1 numbers"`.

# --hints--

`count_letters_and_numbers("helloworld123")` يجب أن تُرجع `"The string has 10 letters and 3 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("helloworld123"), "The string has 10 letters and 3 numbers.")`)
}})
```

`count_letters_and_numbers("Catch 22")` يجب أن تُرجع `"The string has 5 letters and 2 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("Catch 22"), "The string has 5 letters and 2 numbers.")`)
}})
```

`count_letters_and_numbers("A1!")` يجب أن تُرجع `"The string has 1 letter and 1 number."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("A1!"), "The string has 1 letter and 1 number.")`)
}})
```

`count_letters_and_numbers("12345")` يجب أن تُرجع `"The string has 0 letters and 5 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("12345"), "The string has 0 letters and 5 numbers.")`)
}})
```

`count_letters_and_numbers("password")` يجب أن تُرجع `"The string has 8 letters and 0 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("password"), "The string has 8 letters and 0 numbers.")`)
}})
```

# --seed--

## --seed-contents--

```py
def count_letters_and_numbers(s):

    return s
```

# --solutions--

```py
def count_letters_and_numbers(s):
    letter_count = 0
    number_count = 0

    for char in s:
        if ('A' <= char <= 'Z') or ('a' <= char <= 'z'):
            letter_count += 1
        elif '0' <= char <= '9':
            number_count += 1

    letter_word = "letter" if letter_count == 1 else "letters"
    number_word = "number" if number_count == 1 else "numbers"

    return f"The string has {letter_count} {letter_word} and {number_count} {number_word}."
```
