---
id: 6994cff2290543b3aec9f50f
title: "تحدٍّ 210: مدقق HSL"
challengeType: 29
dashedName: challenge-210
---

# --description--

معطاة سلسلة نصية، حدّد ما إذا كانت قيمة لون CSS `hsl()` صحيحة.

- يجب أن تكون قيمة HSL صحيحة بالتنسيق `"hsl(h, s%, l%)"`، حيث:
  - يجب أن يكون `h` (الصبغة) رقمًا بين 0 و360 (شامل).
  - يجب أن تكون `s` (التشبع) نسبة مئوية بين 0% و100%.
  - يجب أن تكون `l` (فاتح اللون) نسبة مئوية بين 0% و100%.

- يُسمح بالمسافات فقط:
  - بعد القوس المفتوح
  - قبل و/أو بعد الفواصل
  - قبل و/أو بعد القوس المغلق

- يمكن أن تنتهي القيمة اختياريًا بفاصلة منقوطة (`";"`).

على سبيل المثال، `"hsl(240, 50%, 50%)"` هي قيمة HSL صحيحة.

# --hints--

يجب أن تُعيد الدالة `is_valid_hsl("hsl(240, 50%, 50%)")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(240, 50%, 50%)"), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_hsl("hsl( 200 , 50% , 75% )")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl( 200 , 50% , 75% )"), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_hsl("hsl(99,60%,80%);")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(99,60%,80%);"), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_hsl("hsl(0, 0%, 0%) ;")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(0, 0%, 0%) ;"), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_hsl("hsl(  10  ,  20%   ,  30%   )    ;")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(  10  ,  20%   ,  30%   )    ;"), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_hsl("hsl(361, 50%, 80%)")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(361, 50%, 80%)"), False)`)
}})
```

يجب أن تُعيد الدالة `is_valid_hsl("hsl(300, 101%, 70%)")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(300, 101%, 70%)"), False)`)
}})
```

يجب أن تُعيد الدالة `is_valid_hsl("hsl(200, 55%, 75)")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(200, 55%, 75)"), False)`)
}})
```

يجب أن تُعيد الدالة `is_valid_hsl("hsl (80, 20%, 10%)")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl (80, 20%, 10%)"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_hsl(hsl):

    return hsl
```

# --solutions--

```py
import re
def is_valid_hsl(hsl):
    pattern = r"^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$"
    
    match = re.match(pattern, hsl)
    if not match:
        return False

    hue = int(match.group(1))
    sat = int(match.group(2))
    light = int(match.group(3))

    if not (0 <= hue <= 360):
        return False
    if not (0 <= sat <= 100):
        return False
    if not (0 <= light <= 100):
        return False

    return True
```
