---
id: 6994cff2290543b3aec9f510
title: "تحدٍّ 211: مجموع المصفوفة"
challengeType: 29
dashedName: challenge-211
---

# --description--

أعطِ مصفوفة من الأرقام، أرجع مجموع كل الأرقام.

# --hints--

يجب أن تُرجع الدالة `sum_array([1, 2, 3, 4, 5])` القيمة `15`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([1, 2, 3, 4, 5]), 15)`)
}})
```

يجب أن تُرجع الدالة `sum_array([42])` القيمة `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([42]), 42)`)
}})
```

يجب أن تُرجع الدالة `sum_array([5, -2, 7, -3])` القيمة `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([5, -2, 7, -3]), 7)`)
}})
```

يجب أن تُرجع الدالة `sum_array([203, 145, -129, 6293, 523, -919, 845, 2434])` القيمة `9395`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395)`)
}})
```

يجب أن تُرجع الدالة `sum_array([0, 0])` القيمة `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([0, 0]), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_array(numbers):

    return numbers
```

# --solutions--

```py
def sum_array(numbers):

    return sum(numbers)
```
