---
id: 699c8e045ee7cb94ed2322d9
title: "تحدٍّ 218: القسمة المتساوية"
challengeType: 29
dashedName: challenge-218
---

# --description--

بما أن لديك عددين صحيحين، حدّد إذا كان يمكنك قسمة الأول بالتساوي على الثاني.

# --hints--

`is_evenly_divisible(4, 2)` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(4, 2), True)`)
}})
```

`is_evenly_divisible(7, 3)` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(7, 3), False)`)
}})
```

`is_evenly_divisible(5, 10)` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(5, 10), False)`)
}})
```

`is_evenly_divisible(48, 6)` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(48, 6), True)`)
}})
```

`is_evenly_divisible(3186, 9)` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(3186, 9), True)`)
}})
```

`is_evenly_divisible(4192, 11)` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(4192, 11), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_evenly_divisible(a, b):

    return a
```

# --solutions--

```py
def is_evenly_divisible(a, b):
    return a % b == 0
```
