---
id: 699c8e045ee7cb94ed2322d4
title: "تحدٍّ 213: محول طول الكلمات"
challengeType: 29
dashedName: challenge-213
---

# --description--

معطاة سلسلة نصية من الكلمات، أعد سلسلة نصية جديدة حيث تُستبدَل كل كلمة بطولها.

- الكلمات في السلسلة النصية المعطاة ستكون مفصولة بمسافة واحدة
- احتفظ بالمسافات في السلسلة النصية المعادة.

على سبيل المثال، إذا كانت السلسلة النصية `"hello world"`، فأعد `"5 5"`.

# --hints--

يجب أن تُعيد الدالة `convert_words("hello world")` القيمة `"5 5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("hello world"), "5 5")`)
}})
```

يجب أن تُعيد الدالة `convert_words("Thanks and happy coding")` القيمة `"6 3 5 6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("Thanks and happy coding"), "6 3 5 6")`)
}})
```

يجب أن تُعيد الدالة `convert_words("The quick brown fox jumps over the lazy dog")` القيمة `"3 5 5 3 5 4 3 4 3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3")`)
}})
```

يجب أن تُعيد الدالة `convert_words("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` القيمة `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_words(s):

    return s
```

# --solutions--

```py
def convert_words(s):

    return " ".join(str(len(word)) for word in s.split(" "))
```
