---
id: 699c8e045ee7cb94ed2322d6
title: "تحدٍّ 215: حاسبة رسوم الموقف"
challengeType: 29
dashedName: challenge-215
---

# --description--

معطى سلسلتا نص تمثلان وقت ركن سيارتك ووقت استلامها، احسب رسوم الموقف.

- ستكون السلسلتان المعطيتان بصيغة `"HH:MM"` باستخدام ساعة 24 ساعة. فمثلاً `"14:00"` تعني الساعة 2 ظهرًا.
- السلسلة الأولى تمثل وقت ركن السيارة، والثانية تمثل وقت استلامها.
- إذا كان وقت الاستلام قبل وقت الدخول، فهذا يعني أن الموقف امتد بعد منتصف الليل إلى اليوم التالي.

قواعد الرسوم:

- كل ساعة ركن تكلف 3 دولارات.
- يتم تقريب الساعات الجزئية إلى الساعة الكاملة التالية.
- إذا امتد الموقف طوال الليل (بعد منتصف الليل)، تُضاف رسوم إضافية قدرها 10 دولارات.
- هناك حد أدنى للرسوم بقيمة 5 دولارات (يُستخدم فقط إذا كان المجموع أقل من 5 دولارات).

أرجع التكلفة الإجمالية بصيغة `"$cost"`، مثلًا `"$5"`.

# --hints--

`calculate_parking_fee("09:00", "11:00")` يجب أن تُرجع `"$6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("09:00", "11:00"), "$6")`)
}})
```

`calculate_parking_fee("10:00", "10:30")` يجب أن تُرجع `"$5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("10:00", "10:30"), "$5")`)
}})
```

`calculate_parking_fee("08:10", "10:45")` يجب أن تُرجع `"$9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("08:10", "10:45"), "$9")`)
}})
```

`calculate_parking_fee("14:40", "23:10")` يجب أن تُرجع `"$27"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("14:40", "23:10"), "$27")`)
}})
```

`calculate_parking_fee("18:15", "01:30")` يجب أن تُرجع `"$34"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("18:15", "01:30"), "$34")`)
}})
```

`calculate_parking_fee("11:11", "11:10")` يجب أن تُرجع `"$82"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("11:11", "11:10"), "$82")`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_parking_fee(park_time, pickup_time):

    return park_time
```

# --solutions--

```py
import math

def calculate_parking_fee(park_time, pickup_time):
    def to_minutes(time_str):
        hours, minutes = map(int, time_str.split(":"))
        return hours * 60 + minutes

    entry_minutes = to_minutes(park_time)
    exit_minutes = to_minutes(pickup_time)

    overnight = False

    if exit_minutes < entry_minutes:
        overnight = True
        total_minutes = (24 * 60 - entry_minutes) + exit_minutes
    else:
        total_minutes = exit_minutes - entry_minutes

    hours = math.ceil(total_minutes / 60)

    cost = hours * 3

    if overnight:
        cost += 10

    if cost < 5:
        cost = 5

    return f"${cost}"
```
