---
id: 699c8e045ee7cb94ed2322d5
title: "تحدٍّ 214: مدقق سلسلة الدومينو"
challengeType: 29
dashedName: challenge-214
---

# --description--

معطى مصفوفة ثنائية الأبعاد تمثل متتالية من قطع الدومينو، حدّد ما إذا كانت تشكل سلسلة صحيحة.

- كل عنصر في المصفوفة يمثل قطعة دومينو ويكون مصفوفة من رقمين من 1 إلى 6، شاملة.
- لكي تكون السلسلة صحيحة، يجب أن يطابق الرقم الثاني من كل قطعة الرقم الأول من القطعة التالية.
- لا يشترط أن يطابق الرقم الأول من أول قطعة أو الرقم الأخير من آخر قطعة أي شيء.

# --hints--

يجب أن تُرجع `is_valid_domino_chain([[1, 3], [3, 6], [6, 5]])` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[1, 3], [3, 6], [6, 5]]), True)`)
}})
```

يجب أن تُرجع `is_valid_domino_chain([[6, 2], [3, 4], [4, 1]])` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[6, 2], [3, 4], [4, 1]]), False)`)
}})
```

يجب أن تُرجع `is_valid_domino_chain([[2, 5], [5, 6], [5, 1]])` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 5], [5, 6], [5, 1]]), False)`)
}})
```

يجب أن تُرجع `is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]), True)`)
}})
```

يجب أن تُرجع `is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_domino_chain(dominoes):

    return dominoes
```

# --solutions--

```py
def is_valid_domino_chain(dominoes):
    for i in range(len(dominoes) - 1):
        if dominoes[i][1] != dominoes[i + 1][0]:
            return False

    return True
```
