---
id: 699c8e045ee7cb94ed2322d8
title: "تحدٍّ 217: قطع الشطرنج المأسورة"
challengeType: 29
dashedName: challenge-217
---

# --description--

معطاة مصفوفة من السلاسل النصية تمثل قطع الشطرنج التي لا تزال لديك على الرقعة، احسب قيمة القطع التي أسرها خصمك.

في الشطرنج، تبدأ بـ 16 قطعة:

| القطعة  | الاختصار   | الكمية | القيمة |
| ------ | -------------- | -------- | ----- |
| جندي   | `"P"`          | 8        | 1     |
| رخ     | `"R"`          | 2        | 5     |
| حصان   | `"N"`          | 2        | 3     |
| فيل    | `"B"`          | 2        | 3     |
| وزير   | `"Q"`          | 1        | 9     |
| ملك    | `"K"`          | 1        | 0     |

- المصفوفة المعطاة ستحتوي فقط على الاختصارات أعلاه.
- أي من القطع الـ 16 غير المدرجة في المصفوفة المعطاة قد أُسرت.
- أعد القيمة الإجمالية لكل القطع المأسورة، إلا إذا...
- إذا أُسر الملك، أعد `"Checkmate"`.

# --hints--

`get_captured_value(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` يجب أن تُعيد `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8)`)
}})
```

`get_captured_value(["P", "P", "P", "P", "P", "R", "B", "K"])` يجب أن تُعيد `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "P", "P", "P", "P", "R", "B", "K"]), 26)`)
}})
```

`get_captured_value(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` يجب أن تُعيد `16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16)`)
}})
```

`get_captured_value(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` يجب أن تُعيد `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4)`)
}})
```

`get_captured_value(["P", "K"])` يجب أن تُعيد `38`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "K"]), 38)`)
}})
```

`get_captured_value(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` يجب أن تُعيد `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0)`)
}})
```

`get_captured_value(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` يجب أن تُعيد `"Checkmate"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_captured_value(pieces):

    return pieces
```

# --solutions--

```py
def get_captured_value(pieces):
    if "K" not in pieces:
        return "Checkmate"

    values = {"P":1, "R":5, "N":3, "B":3, "Q":9, "K":0}
    remaining = "PPPPPPPPRRNNBBQK"

    for p in pieces:
        remaining = remaining.replace(p, "", 1)

    return sum(values[p] for p in remaining)
```
