---
id: 699c8e045ee7cb94ed2322dd
title: "تحدٍّ 222: ظلال الاعتدال"
challengeType: 29
dashedName: challenge-222
---

# --description--

اليوم هو يوم الاعتدال، حيث تكون الشمس مباشرة فوق خط الاستواء وفوق الرأس تمامًا عند الظهر. معطى وقت، حدّد الظل الذي يلقيه عمود عمودي طوله 4 أقدام.

- سيكون الوقت سلسلة نصية بصيغة `"HH:MM"` بنظام 24 ساعة (على سبيل المثال، `"15:00"` تعني الساعة 3 مساءً).
- سيُعطى الوقت فقط بفواصل 30 دقيقة.

القواعد:

- تشرق الشمس عند الساعة 6 صباحًا مباشرةً من جهة `"east"`، وتغرب عند الساعة 6 مساءً مباشرةً من جهة `"west"`.
- يشير الظل دائمًا في الاتجاه المعاكس للشمس.
- طول الظل (بالأقدام) هو عدد الساعات بعيدًا عن الظهر مرفوعًا للقوة الثالثة.
- لا يوجد ظل قبل شروق الشمس (قبل 6 صباحًا)، أو بعد غروب الشمس (6 مساءً أو بعده)، أو عند الظهر.

الإرجاع:

- إذا وُجد ظل، أرجع `"(length)ft (direction)"`. على سبيل المثال، `"8ft west"`.
- وإلا، أرجع `"No shadow"`.

على سبيل المثال، معطى `"10:00"`، أرجع `"8ft west"` لأن الساعة 10 صباحًا تبعد ساعتين عن الظهر، إذًا 2<sup>3</sup> = 8 أقدام، والظل يشير إلى الغرب لأن الشمس في الشرق عند الساعة 10 صباحًا.

# --hints--

`get_shadow("10:00")` يجب أن تُرجع `"8ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` يجب أن تُرجع `"27ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` يجب أن تُرجع `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` يجب أن تُرجع `"166.375ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` يجب أن تُرجع `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` يجب أن تُرجع `"216ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` يجب أن تُرجع `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` يجب أن تُرجع `"91.125ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` يجب أن تُرجع `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
