---
id: 699c8e045ee7cb94ed2322da
title: "تحدٍّ 219: معالم الذكرى السنوية"
challengeType: 29
dashedName: challenge-219
---

# --description--

معطى عدد صحيح يمثل عدد سنوات زواج الزوجين، أرجع أحدث معلم ذكرى سنوية لهم وفقًا لهذا المخطط:

| سنوات الزواج | المعلم       |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- إذا لم يصلوا إلى أول معلم، أرجع `"Newlyweds"`.

# --hints--

`get_milestone(0)` يجب أن تُرجع `"Newlyweds"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(0), "Newlyweds")`)
}})
```

`get_milestone(1)` يجب أن تُرجع `"Paper"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(1), "Paper")`)
}})
```

`get_milestone(8)` يجب أن تُرجع `"Wood"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(8), "Wood")`)
}})
```

`get_milestone(10)` يجب أن تُرجع `"Tin"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(10), "Tin")`)
}})
```

`get_milestone(26)` يجب أن تُرجع `"Silver"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(26), "Silver")`)
}})
```

`get_milestone(45)` يجب أن تُرجع `"Ruby"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(45), "Ruby")`)
}})
```

`get_milestone(50)` يجب أن تُرجع `"Gold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(50), "Gold")`)
}})
```

`get_milestone(64)` يجب أن تُرجع `"Diamond"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(64), "Diamond")`)
}})
```

`get_milestone(71)` يجب أن تُرجع `"Platinum"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(71), "Platinum")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_milestone(years):

    return years
```

# --solutions--

```py
def get_milestone(years):
    if years < 1:
        return "Newlyweds"

    milestones = [
        (1, "Paper"),
        (5, "Wood"),
        (10, "Tin"),
        (25, "Silver"),
        (40, "Ruby"),
        (50, "Gold"),
        (60, "Diamond"),
        (70, "Platinum")
    ]

    for year, milestone in reversed(milestones):
        if years >= year:
            return milestone
```
