---
id: 69a890af247de743333bd4d0
title: "تحدٍّ 227: وقت التهدئة"
challengeType: 29
dashedName: challenge-227
---

# --description--

معطى طابعان زمنيّان، الأول يمثل وقت انتهاء المستخدم من اختبار، والثاني يمثل الوقت الحالي، حدّد ما إذا كان يمكن للمستخدم أن يخوض الاختبار مرة أخرى.

- كلا الطابعين الزمنيّين سيكونان بالتنسيق: `"YYYY-MM-DDTHH:MM:SS"`، مثلًا `"2026-03-25T14:00:00"`. لاحظ أن الوقت بنظام 24 ساعة.
- يجب على المستخدم الانتظار 48 ساعة على الأقل قبل إعادة الاختبار.

# --hints--

`can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00"), True)`)
}})
```

`can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00"), False)`)
}})
```

`can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00"), True)`)
}})
```

`can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_retake(finish_time, current_time):

    return finish_time
```

# --solutions--

```py
from datetime import datetime, timedelta

def can_retake(finish_time, current_time):
    fmt = "%Y-%m-%dT%H:%M:%S"
    finished_dt = datetime.strptime(finish_time, fmt)
    current_dt = datetime.strptime(current_time, fmt)
    diff = current_dt - finished_dt
    cooldown = timedelta(hours=48)

    return diff >= cooldown
```
