---
id: 69a890af247de743333bd4d1
title: "تحدٍّ 228: ليلة الفيلم"
challengeType: 29
dashedName: challenge-228
---

# --description--

معطى سلسلة نصية لليوم من الأسبوع، وسلسلة نصية أخرى لوقت العرض، وعدد صحيح من التذاكر، أرجع التكلفة الإجمالية لتذاكر الفيلم لذلك العرض.

اليوم المعطى سيكون واحدًا من:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

سيُعطى وقت العرض بالتنسيق `"H:MMam"` أو `"H:MMpm"`. على سبيل المثال `"10:00am"` أو `"10:00pm"`.

أرجع التكلفة الإجمالية بالتنسيق `"$D.CC"` باستخدام القواعد التالية:

- عطلة نهاية الأسبوع (الجمعة - الأحد): 12.00$ لكل تذكرة.
- أيام الأسبوع (الاثنين - الخميس): 10.00$ لكل تذكرة.
- العروض الصباحية (قبل 5:00 مساءً): اخصم 2.00$ لكل تذكرة (باستثناء أيام الثلاثاء).
- أيام الثلاثاء: كل التذاكر بسعر 5.00$ لكل تذكرة.

# --hints--

`get_movie_night_cost("Saturday", "10:00pm", 1)` يجب أن تُرجع `"$12.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Saturday", "10:00pm", 1), "$12.00")`)
}})
```

`get_movie_night_cost("Sunday", "10:00am", 1)` يجب أن تُرجع `"$10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Sunday", "10:00am", 1), "$10.00")`)
}})
```

`get_movie_night_cost("Tuesday", "7:20pm", 2)` يجب أن تُرجع `"$10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "7:20pm", 2), "$10.00")`)
}})
```

`get_movie_night_cost("Wednesday", "5:40pm", 3)` يجب أن تُرجع `"$30.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Wednesday", "5:40pm", 3), "$30.00")`)
}})
```

`get_movie_night_cost("Monday", "11:50am", 4)` يجب أن تُرجع `"$32.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Monday", "11:50am", 4), "$32.00")`)
}})
```

`get_movie_night_cost("Friday", "4:30pm", 5)` يجب أن تُرجع `"$50.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Friday", "4:30pm", 5), "$50.00")`)
}})
```

`get_movie_night_cost("Tuesday", "11:30am", 1)` يجب أن تُرجع `"$5.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "11:30am", 1), "$5.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):

    return day
```

# --solutions--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):
    if day == "Tuesday":
        price = 5
    else:
        if day in ["Friday", "Saturday", "Sunday"]:
            price = 12
        else:
            price = 10

        hour = int(showtime.split(":")[0])
        period = showtime[-2:]

        hour24 = hour

        if period == "pm" and hour != 12:
            hour24 += 12
        if period == "am" and hour == 12:
            hour24 = 0

        if hour24 < 17:
            price -= 2

    total = price * number_of_tickets

    return f"${total:.2f}"
```
