---
id: 69a890af247de743333bd4cc
title: "تحدٍّ 223: فك شفرة QR"
challengeType: 29
dashedName: challenge-223
---

# --description--

معطى مصفوفة 6x6 (مصفوفة من سلاسل نصية) تمثل رمز QR، أرجع السلسلة النصية للبيانات الثنائية في الرمز.

- قد يُعطى رمز QR بأي دوران بزيادات 90 درجة.
- الرمز الموجه بشكل صحيح يحتوي على مجموعة 2x2 من `1` (علامات التوجيه) في الزوايا السفلى اليسرى، العلوية اليسرى، والعلوية اليمنى.
- علامات التوجيه الثلاث 2x2 ليست جزءًا من البيانات الثنائية.
- تُقرأ البيانات الثنائية من اليسار إلى اليمين، ومن الأعلى إلى الأسفل (مثل الكتاب) عندما يكون رمز QR موجهًا بشكل صحيح.
- سيكون للرمز دائمًا اتجاه واحد صحيح فقط.

على سبيل المثال، معطى:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

أو معطى نفس الرمز مع اتجاه مختلف:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

أرجع `"000000000000000000000001"`، كل البيانات الثنائية باستثناء علامات التوجيه الثلاث 2x2.

# --hints--

`decode_qr(["110011", "110011", "000000", "000000", "110000", "110001"])` يجب أن تُرجع `"000000000000000000000001"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001")`)
}})
```

`decode_qr(["100011", "000011", "000000", "000000", "110011", "110011"])` يجب أن تُرجع `"000000000000000000000001"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001")`)
}})
```

`decode_qr(["110011", "111111", "010000", "110000", "110011", "110100"])` يجب أن تُرجع `"001101000011000000110100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100")`)
}})
```

`decode_qr(["011011", "101011", "101000", "100010", "110011", "111011"])` يجب أن تُرجع `"010001000100010101010110"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110")`)
}})
```

`decode_qr(["111100", "110001", "100011", "001101", "110011", "110011"])` يجب أن تُرجع `"010000100100100101001110"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode_qr(qr_code):

    return qr_code
```

# --solutions--

```py
def is_oriented(code):
    marker_indices = [
        (0, 0),
        (0, 1),
        (0, 4),
        (0, 5),
        (1, 0),
        (1, 1),
        (1, 4),
        (1, 5),
        (4, 0),
        (4, 1),
        (5, 0),
        (5, 1)
    ]

    for row, col in marker_indices:
        if code[row][col] != '1':
            return False

    return True


def rotate(code):
    new_code = ['', '', '', '', '', '']
    for i in range(6):
        for j in range(5, -1, -1):
            new_code[i] += code[j][i]
    return new_code


def get_data(code):
    data = ''
    data += code[0][2:4]
    data += code[1][2:4]
    data += code[2]
    data += code[3]
    data += code[4][2:]
    data += code[5][2:]
    return data


def decode_qr(qr_code):
    code = qr_code[:]

    while not is_oriented(code):
        code = rotate(code)

    return get_data(code)
```
