---
id: 69a890af247de743333bd4d2
title: "تحدٍّ 229: اقتطاع النص 2"
challengeType: 29
dashedName: challenge-229
---

# --description--

بالنظر إلى سلسلة نصية، أعد سلسلة جديدة مقتطعة بحيث لا يتجاوز العرض الكلي للأحرف 50 وحدة.

لكل حرف عرض محدد:

| الحروف | العرض |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

الجدول أعلاه يشمل جميع الحروف الكبيرة والصغيرة. بالإضافة إلى ذلك:

- الفراغات (`" "`) عرضها 2
- النقاط (`"."`) عرضها 1

- إذا كانت السلسلة النصية المعطاة 50 وحدة أو أقل، أعد السلسلة كما هي، وإلا
- اقتطع السلسلة وأضف ثلاث نقاط في النهاية (`"..."`) بحيث يكون عرضها الكلي، بما في ذلك النقاط الثلاث، قريبًا قدر الإمكان من 50 وحدة دون تجاوزها.

# --hints--

`truncate_text("The quick brown fox")` يجب أن تعيد `"The quick brown f..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The quick brown fox"), "The quick brown f...")`)
}})
```

`truncate_text("The silky smooth sloth")` يجب أن تعيد `"The silky smooth s..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The silky smooth sloth"), "The silky smooth s...")`)
}})
```

`truncate_text("THE LOUD BRIGHT BIRD")` يجب أن تعيد `"THE LOUD BRIG..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...")`)
}})
```

`truncate_text("The fast striped zebra")` يجب أن تعيد `"The fast striped z..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The fast striped zebra"), "The fast striped z...")`)
}})
```

`truncate_text("The big black bear")` يجب أن تعيد `"The big black bear"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The big black bear"), "The big black bear")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(s):

    return s
```

# --solutions--

```py
def truncate_text(s):
    MAX_WIDTH = 50
    ELLIPSIS = "..."
    ELLIPSIS_WIDTH = 3
    TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH

    char_groups = {
        "ilI.": 1,
        "fjrt ": 2,
        "abcdeghkmnopqrstuvwxyzJL": 3,
        "ABCDEFGHKMNOPQRSTUVWXYZ": 4
    }

    def get_char_width(char):
        for key, width in char_groups.items():
            if char in key:
                return width
        return 3

    def string_width(string):
        total = 0
        for char in string:
            total += get_char_width(char)
        return total

    if string_width(s) <= MAX_WIDTH:
        return s

    result = ""
    total_width = 0

    for char in s:
        char_width = get_char_width(char)

        if total_width + char_width > TRUNCATE_LIMIT:
            break

        result += char
        total_width += char_width

    return result + ELLIPSIS
```
