---
id: 69a890af247de743333bd4ce
title: "تحدٍّ 225: لا تكرار متتالي"
challengeType: 29
dashedName: challenge-225
---

# --description--

أعطِ سلسلة نصية، حدّد إذا لم تحتوي على أحرف مكررة.

- لا تحتوي السلسلة النصية على تكرارات إذا لم يظهر نفس الحرف مرتين أو أكثر على التوالي.

# --hints--

`has_no_repeats("hi world")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("hi world"), True)`)
}})
```

`has_no_repeats("hello world")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("hello world"), False)`)
}})
```

`has_no_repeats("abcdefghijklmnopqrstuvwxyz")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("abcdefghijklmnopqrstuvwxyz"), True)`)
}})
```

`has_no_repeats("freeCodeCamp")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("freeCodeCamp"), False)`)
}})
```

`has_no_repeats("The quick brown fox jumped over the lazy dog.")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("The quick brown fox jumped over the lazy dog."), True)`)
}})
```

`has_no_repeats("Mississippi")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("Mississippi"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_no_repeats(s):

    return s
```

# --solutions--

```py
def has_no_repeats(s):
    for i in range(1, len(s)):
        if s[i] == s[i-1]:
            return False
    return True
```
