---
id: 69a890af247de743333bd4cd
title: "تحدٍّ 224: كاشف تحميص القهوة"
challengeType: 29
dashedName: challenge-224
---

# --description--

أعطيت سلسلة نصية تمثل الحبوب المستخدمة لصنع فنجان قهوة، حدّد مستوى تحميص الفنجان.

- ستحتوي السلسلة النصية على الأحرف التالية، كل منها يمثل نوعًا من الحبوب:

- الفاصلة العليا (`'`) تمثل حبة محمصة خفيفة بقيمة نقطة واحدة لكل منها.
  - الشرطة (`-`) تمثل حبة محمصة متوسطة بقيمة نقطتين لكل منها.
  - النقطة (`.`) تمثل حبة محمصة داكنة بقيمة ثلاث نقاط لكل منها.

- يُحدد مستوى التحميص بمتوسط جميع الحبوب.

أرجع:

- `"Light"` إذا كان المتوسط أقل من 1.75.
- `"Medium"` إذا كان المتوسط بين 1.75 و2.5.
- `"Dark"` إذا كان المتوسط أكبر من 2.5.

# --hints--

`detect_roast("''-''''''-'-''--''''")` يجب أن تُرجع `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("''-''''''-'-''--''''"), "Light")`)
}})
```

`detect_roast(".'-''-''..'''.-.-''-")` يجب أن تُرجع `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".'-''-''..'''.-.-''-"), "Medium")`)
}})
```

`detect_roast("--.''--'-''.--..-.--")` يجب أن تُرجع `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("--.''--'-''.--..-.--"), "Medium")`)
}})
```

`detect_roast("-...'-......-..-...-")` يجب أن تُرجع `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-...'-......-..-...-"), "Dark")`)
}})
```

`detect_roast(".--.-..-......----.'")` يجب أن تُرجع `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".--.-..-......----.'"), "Medium")`)
}})
```

`detect_roast("..-..-..-..-....-.-.")` يجب أن تُرجع `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("..-..-..-..-....-.-."), "Dark")`)
}})
```

`detect_roast("-'-''''''..-'.''-'.'")` يجب أن تُرجع `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-'-''''''..-'.''-'.'"), "Light")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_roast(beans):

    return beans
```

# --solutions--

```py
def detect_roast(beans):
    total = 0

    for bean in beans:
        if bean == "'":
            total += 1
        elif bean == "-":
            total += 2
        elif bean == ".":
            total += 3

    avg = total / len(beans)

    if avg < 1.75:
        return "Light"
    elif avg <= 2.5:
        return "Medium"
    else:
        return "Dark"
```
