---
id: 69b1028d6e265413d0198a2a
title: "تحدٍّ 231: مدقق ISBN-10"
challengeType: 29
dashedName: challenge-231
---

# --description--

أعطِ سلسلة نصية، حدّد إذا كانت رقم ISBN-10 صالحًا.

يتكوّن ISBN-10 من شرطات (`"-"`) و10 أحرف أخرى. بعد إزالة الشرطات (`"-"`):

- يجب أن تكون الأحرف التسعة الأولى أرقامًا، و
- قد يكون الحرف الأخير رقمًا أو الحرف `"X"`، الذي يُمثّل الرقم 10.

للتحقق من صحته:

- اضرب كل رقم (أو قيمة) في موقعه (اضرب الرقم الأول في 1، والثاني في 2، وهكذا).
- اجمع كل النتائج معًا.
- إذا كان المجموع قابلًا للقسمة على 11، فهو صالح.

# --hints--

`is_valid_isbn10("0-306-40615-2")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-306-40615-2"), True)`)
}})
```

`is_valid_isbn10("0-306-40615-1")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-306-40615-1"), False)`)
}})
```

`is_valid_isbn10("0-8044-2957-X")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-8044-2957-X"), True)`)
}})
```

`is_valid_isbn10("X-306-40615-2")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("X-306-40615-2"), False)`)
}})
```

`is_valid_isbn10("0-6822-2589-4")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-6822-2589-4"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_isbn10(s):

    return s
```

# --solutions--

```py
def is_valid_isbn10(s):
    isbn = s.replace("-", "")

    if len(isbn) != 10:
        return False

    total = 0

    for i in range(9):
        if not isbn[i].isdigit():
            return False
        total += int(isbn[i]) * (i + 1)

    last = isbn[9]

    if last == "X":
        total += 10 * 10
    elif last.isdigit():
        total += int(last) * 10
    else:
        return False

    return total % 11 == 0
```
