---
id: 69b1028d6e265413d0198a29
title: "تحدٍّ 230: صف مثلث باسكال"
challengeType: 29
dashedName: challenge-230
---

# --description--

معطى عدد صحيح `n`، أعد الصف `n` من مثلث باسكال كمصفوفة.

في مثلث باسكال، يبدأ وينتهي كل صف بالعدد 1، وكل قيمة داخلية هي مجموع القيمتين الموجودتين فوقها مباشرة.

إليك أول 5 صفوف من المثلث:

```js
    1
   1 1
  1 2 1
 1 3 3 1
1 4 6 4 1
```

# --hints--

`pascal_row(5)` يجب أن تُعيد `[1, 4, 6, 4, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pascal_row(5), [1, 4, 6, 4, 1])`)
}})
```

`pascal_row(3)` يجب أن تُعيد `[1, 2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pascal_row(3), [1, 2, 1])`)
}})
```

`pascal_row(1)` يجب أن تُعيد `[1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pascal_row(1), [1])`)
}})
```

`pascal_row(10)` يجب أن تُعيد `[1, 9, 36, 84, 126, 126, 84, 36, 9, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pascal_row(10), [1, 9, 36, 84, 126, 126, 84, 36, 9, 1])`)
}})
```

`pascal_row(15)` يجب أن تُعيد `[1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pascal_row(15), [1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def pascal_row(n):

    return n
```

# --solutions--

```py
def pascal_row(n):
    if n == 1:
        return [1]

    row = [1]

    for _ in range(2, n + 1):
        next_row = [1]

        for i in range(1, len(row)):
            next_row.append(row[i - 1] + row[i])

        next_row.append(1)
        row = next_row

    return row
```
