---
id: 69b1028d6e265413d0198a2c
title: "تحدٍّ 233: منبّه الاستيقاظ"
challengeType: 29
dashedName: challenge-233
---

# --description--

معطاة سلسلة نصية تمثّل الوقت الذي ضبطت فيه المنبّه وسلسلة نصية تمثّل الوقت الذي استيقظت فيه فعليًا، حدّد إذا استيقظت مبكرًا، في الوقت المحدد، أو متأخرًا.

- يُعطى كلا الوقتين بصيغة 24 ساعة `"HH:MM"`.

أرجع:

- `"early"` إذا استيقظت قبل وقت المنبّه.
- `"on time"` إذا استيقظت في وقت المنبّه، أو خلال نافذة الغفوة التي تستمر 10 دقائق بعد وقت المنبّه.
- `"late"` إذا استيقظت بعد أكثر من 10 دقائق من وقت المنبّه.

كلا الوقتين في نفس اليوم.

# --hints--

`alarm_check("07:00", "06:45")` يجب أن تُرجع `"early"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("07:00", "06:45"), "early")`)
}})
```

`alarm_check("06:30", "06:30")` يجب أن تُرجع `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("06:30", "06:30"), "on time")`)
}})
```

`alarm_check("08:10", "08:15")` يجب أن تُرجع `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("08:10", "08:15"), "on time")`)
}})
```

`alarm_check("09:30", "09:45")` يجب أن تُرجع `"late"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("09:30", "09:45"), "late")`)
}})
```

`alarm_check("08:15", "08:25")` يجب أن تُرجع `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("08:15", "08:25"), "on time")`)
}})
```

`alarm_check("05:45", "05:56")` يجب أن تُرجع `"late"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("05:45", "05:56"), "late")`)
}})
```

`alarm_check("04:30", "04:00")` يجب أن تُرجع `"early"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("04:30", "04:00"), "early")`)
}})
```

# --seed--

## --seed-contents--

```py
def alarm_check(alarm_time, wake_time):

    return alarm_time
```

# --solutions--

```py
def alarm_check(alarm_time, wake_time):
    alarm_h, alarm_m = map(int, alarm_time.split(":"))
    wake_h, wake_m = map(int, wake_time.split(":"))

    alarm_minutes = alarm_h * 60 + alarm_m
    wake_minutes = wake_h * 60 + wake_m
    diff = wake_minutes - alarm_minutes

    if diff < 0:
        return "early"
    if diff <= 10:
        return "on time"
    return "late"
```
