---
id: 69b58ce40693f140c84c855b
title: "تحدٍّ 242: الرقم التالي في لعبة البنغو"
challengeType: 29
dashedName: challenge-242
---

# --description--

عند إعطائك رقم بنغو، أرجع الرقم التالي في التسلسل.

رقم البنغو هو حرف واحد يتبعه رقم ضمن نطاقه وفقًا لهذا المخطط:

| الحرف | نطاق الرقم |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

على سبيل المثال، عند إعطائك `"B10"`، أرجع `"B11"`، الرقم التالي في البنغو. إذا كان الرقم المعطى هو آخر رقم في البنغو، أرجع `"B1"`.

# --hints--

يجب أن تُرجع الدالة `get_next_bingo_number("B10")` القيمة `"B11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("B10"), "B11")`)
}})
```

يجب أن تُرجع الدالة `get_next_bingo_number("N33")` القيمة `"N34"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("N33"), "N34")`)
}})
```

يجب أن تُرجع الدالة `get_next_bingo_number("I30")` القيمة `"N31"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("I30"), "N31")`)
}})
```

يجب أن تُرجع الدالة `get_next_bingo_number("G60")` القيمة `"O61"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("G60"), "O61")`)
}})
```

يجب أن تُرجع الدالة `get_next_bingo_number("O75")` القيمة `"B1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("O75"), "B1")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_next_bingo_number(n):

    return n
```

# --solutions--

```py
def get_next_bingo_number(n):
    num = int(n[1:])
    next_num = 1 if num == 75 else num + 1

    if next_num <= 15:
        return "B" + str(next_num)
    if next_num <= 30:
        return "I" + str(next_num)
    if next_num <= 45:
        return "N" + str(next_num)
    if next_num <= 60:
        return "G" + str(next_num)
    return "O" + str(next_num)
```
