---
id: 69b58ce40693f140c84c8559
title: "تحدٍّ 240: أحرف الكلمة المتماثلة"
challengeType: 29
dashedName: challenge-240
---

# --description--

معطاة سلسلة نصية، حدّد إذا كانت كلمة متماثلة وأرجع الحرف الأوسط (إذا كان طولها فرديًا) أو الحرفين الأوسطين (إذا كان طولها زوجيًا).

- الكلمة المتماثلة هي سلسلة نصية تتطابق عند قراءتها من الأمام أو الخلف.
- إذا لم تكن كلمة متماثلة، أرجع `"none"`.

# --hints--

`palindrome_locator("racecar")` يجب أن تُرجع `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("racecar"), "e")`)
}})
```

`palindrome_locator("level")` يجب أن تُرجع `"v"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("level"), "v")`)
}})
```

`palindrome_locator("freecodecamp")` يجب أن تُرجع `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("freecodecamp"), "none")`)
}})
```

`palindrome_locator("noon")` يجب أن تُرجع `"oo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("noon"), "oo")`)
}})
```

`palindrome_locator("11100111")` يجب أن تُرجع `"00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("11100111"), "00")`)
}})
```

# --seed--

## --seed-contents--

```py
def palindrome_locator(s):

    return s
```

# --solutions--

```py
def palindrome_locator(s):
    if s != s[::-1]:
        return "none"
    mid = len(s) // 2
    return s[mid] if len(s) % 2 == 1 else s[mid - 1] + s[mid]
```
