---
id: 69b559d2903b9e4afe9075f8
title: "تحدٍّ 238: هروب تدوير الأرقام"
challengeType: 29
dashedName: challenge-238
---

# --description--

أعطِ عددًا صحيحًا موجبًا، حدّد إذا كان هو أو أي من تدويراته يقبل القسمة بالتساوي على عدد أرقامه.

التدوير يعني نقل الرقم الأول إلى النهاية. على سبيل المثال، بعد تدوير واحد، يصبح 123 هو 231.

- افحص التدوير `0` (العدد المعطى) أولًا.
- الأعداد المعطاة لا تحتوي على أصفار.
- أرجع رقم التدوير الأول إذا وُجد، أو `"none"` إذا لم يوجد.

# --hints--

`get_rotation(123)` يجب أن تُرجع `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(123), 0)`)
}})
```

`get_rotation(13579)` يجب أن تُرجع `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(13579), 3)`)
}})
```

`get_rotation(24681)` يجب أن تُرجع `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(24681), "none")`)
}})
```

`get_rotation(84138789345)` يجب أن تُرجع `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(84138789345), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rotation(n):

    return n
```

# --solutions--

```py
def get_rotation(n):
    s = str(n)
    digit_count = len(s)
    current = s

    for i in range(digit_count):
        if int(current) % digit_count == 0:
            return i
        current = current[1:] + current[0]

    return "none"
```
