---
id: 69b58ce40693f140c84c855c
title: "تحدٍّ 243: هجوم الرخ"
challengeType: 29
dashedName: challenge-243
---

# --description--

معطاة سلسلتا نصية لموقع رخين على رقعة شطرنج، حدّد إذا كان بإمكانهما مهاجمة بعضهما البعض.

رقعة الشطرنج القياسية 8×8، مع أعمدة معنونة من `A` إلى `H` (من اليسار إلى اليمين) وصفوف معنونة من `1` إلى `8` (من الأسفل إلى الأعلى). تبدو هكذا:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

يمكن للرُّخ التحرك بعدد مربعات غير محدود أفقيًا أو عموديًا. لذلك إذا كانا في نفس الصف أو العمود، يمكنهما مهاجمة بعضهما البعض.

# --hints--

`rook_attack("A1", "A8")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(rook_attack("A1", "A8"), True)`)
}})
```

`rook_attack("B4", "F4")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(rook_attack("B4", "F4"), True)`)
}})
```

`rook_attack("E3", "D4")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(rook_attack("E3", "D4"), False)`)
}})
```

`rook_attack("H7", "F6")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(rook_attack("H7", "F6"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def rook_attack(rook1, rook2):

    return rook1
```

# --solutions--

```py
def rook_attack(rook1, rook2):
    return rook1[0] == rook2[0] or rook1[1] == rook2[1]
```
