---
id: 69b559d2903b9e4afe9075f9
title: "تحدٍّ 239: ما هو اليوم؟"
challengeType: 29
dashedName: challenge-239
---

# --description--

معطى طابع زمني من نوع Unix بوحدة المللي ثانية، أرجع يوم الأسبوع.

أيام الإرجاع الصحيحة هي:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

تأكد من تجاهل المناطق الزمنية.

# --hints--

`get_day_of_week(1775492249000)` يجب أن تُرجع `"Monday"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_day_of_week(1775492249000), "Monday")`)
}})
```

`get_day_of_week(1766246400000)` يجب أن تُرجع `"Saturday"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_day_of_week(1766246400000), "Saturday")`)
}})
```

`get_day_of_week(33791256000000)` يجب أن تُرجع `"Tuesday"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_day_of_week(33791256000000), "Tuesday")`)
}})
```

`get_day_of_week(1773576000000)` يجب أن تُرجع `"Sunday"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_day_of_week(1773576000000), "Sunday")`)
}})
```

`get_day_of_week(0)` يجب أن تُرجع `"Thursday"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_day_of_week(0), "Thursday")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_day_of_week(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime, timezone

def get_day_of_week(timestamp):
    days = ["Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday", "Sunday"]
    dt = datetime.fromtimestamp(timestamp / 1000, tz=timezone.utc)
    return days[dt.weekday()]
```
