---
id: 69bc6cb30c1d112a2e110a03
title: "تحدٍّ 246: الأحرف الأولى من الاسم"
challengeType: 29
dashedName: challenge-246
---

# --description--

بالنظر إلى اسم كامل كسلسلة نصية، أرجع الأحرف الأولى منه.

- الأسماء التي تُؤخذ منها الأحرف الأولى مفصولة بمسافة.
- يجب أن تكون الأحرف الأولى بحروف كبيرة.
- يجب أن تُفصل الأحرف الأولى بنقاط.

على سبيل المثال، `"Tommy Millwood"` تُرجع `"T.M."`.

# --hints--

`get_initials("Tommy Millwood")` يجب أن تُرجع `"T.M."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Tommy Millwood"), "T.M.")`)
}})
```

`get_initials("Savanna Puddlesplash")` يجب أن تُرجع `"S.P."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Savanna Puddlesplash"), "S.P.")`)
}})
```

`get_initials("Frances Cowell Conrad")` يجب أن تُرجع `"F.C.C."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Frances Cowell Conrad"), "F.C.C.")`)
}})
```

`get_initials("Dragon")` يجب أن تُرجع `"D."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Dragon"), "D.")`)
}})
```

`get_initials("Dorothy Vera Clump Haverstock Norris")` يجب أن تُرجع `"D.V.C.H.N."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_initials(name):

    return name
```

# --solutions--

```py
def get_initials(name):
    return ''.join(word[0].upper() + '.' for word in name.split())
```
