---
id: 69bc6cb30c1d112a2e110a08
title: "تحدٍّ 251: مكتشف مجموع المصفوفة"
challengeType: 29
dashedName: challenge-251
---

# --description--

معطاة مصفوفة من الأرقام ورقم هدف، أعد أول مجموعة جزئية من رقمين أو أكثر مجموعها يساوي الهدف.

- المجموعة الجزئية "الأولى" هي التي تحتوي عناصرها على أدنى فهارس ممكنة، مع إعطاء الأولوية لأدنى فهرس أولًا.
- يمكن استخدام كل رقم في المصفوفة مرة واحدة فقط.
- إذا لم توجد مجموعة جزئية صالحة، أعد `"Sum not found"`.

أعد الأرقام المطابقة كمصفوفة بالترتيب الذي تظهر به في المصفوفة الأصلية.

# --hints--

`find_sum([1, 3, 5, 7], 6)` يجب أن تُعيد `[1, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([1, 3, 5, 7], 6), [1, 5])`)
}})
```

`find_sum([1, 2, 3, 4, 5], 5)` يجب أن تُعيد `[1, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([1, 2, 3, 4, 5], 5), [1, 4])`)
}})
```

`find_sum([1, 2, 3, 4, 5], 6)` يجب أن تُعيد `[1, 2, 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([1, 2, 3, 4, 5], 6), [1, 2, 3])`)
}})
```

`find_sum([-1, -2, 3, 4], 1)` يجب أن تُعيد `[-1, -2, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([-1, -2, 3, 4], 1), [-1, -2, 4])`)
}})
```

`find_sum([3, 1, 4, 1, 5, 9, 2, 6], 10)` يجب أن تُعيد `[3, 1, 4, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2])`)
}})
```

`find_sum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` يجب أن تُعيد `[1, 2, 3, 5, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9])`)
}})
```

`find_sum([7, 9, 4, 2, 5], 10)` يجب أن تُعيد `"Sum not found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([7, 9, 4, 2, 5], 10), "Sum not found")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_sum(arr, target):

    return arr
```

# --solutions--

```py
def find_sum(arr, target):
    def search(start, remaining, current):
        if remaining == 0 and len(current) >= 2:
            return current
        if start >= len(arr):
            return None
        for i in range(start, len(arr)):
            result = search(i + 1, remaining - arr[i], current + [arr[i]])
            if result is not None:
                return result
        return None

    return search(0, target, []) or "Sum not found"
```
