---
id: 69bc6cb30c1d112a2e110a06
title: "تحدٍّ 249: حساب السلاسل النصية"
challengeType: 29
dashedName: challenge-249
---

# --description--

معطى سلسلة نصية تحتوي أرقامًا وحروفًا أخرى، نفذ عمليات حسابية على الأرقام بناءً على عدد الأحرف غير الرقمية بين الأرقام.

- إذا كان عدد الأحرف التي تفصل بين رقمين زوجيًا، استخدم الجمع.
- إذا كان فرديًا، استخدم الطرح.
- الأرقام المتتالية تشكل رقمًا واحدًا.
- تُطبق العمليات من اليسار إلى اليمين.
- تجاهل الأحرف غير الرقمية في البداية والنهاية.

على سبيل المثال، إذا كانت السلسلة النصية `"3ab10c8"`، أرجع القيمة `5`. اجمع 3 و10 لتحصل على 13 لأن عدد الأحرف بينهما زوجي. ثم اطرح 8 من 13 لأن عدد الأحرف بين النتيجة و8 فردي.

# --hints--

`do_math("3ab10c8")` يجب أن تُرجع `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("3ab10c8"), 5)`)
}})
```

`do_math("6MINUS4")` يجب أن تُرجع `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("6MINUS4"), 2)`)
}})
```

`do_math("9plus3")` يجب أن تُرجع `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("9plus3"), 12)`)
}})
```

`do_math("5fkwo#10i#%.<>15P=@20!#B/25")` يجب أن تُرجع `15`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("5fkwo#10i#%.<>15P=@20!#B/25"), 15)`)
}})
```

`do_math("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` يجب أن تُرجع `67`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67)`)
}})
```

# --seed--

## --seed-contents--

```py
def do_math(s):

    return s
```

# --solutions--

```py
import re

def do_math(s):
    tokens = re.findall(r'\d+|\D+', s)
    result = None
    pending_op = None
    for token in tokens:
        if re.match(r'^\d+$', token):
            if result is None:
                result = int(token)
            elif pending_op == 'add':
                result += int(token)
            else:
                result -= int(token)
        else:
            pending_op = 'add' if len(token) % 2 == 0 else 'subtract'
    return result
```
