---
id: 69bc6cb30c1d112a2e110a04
title: "تحدٍّ 247: الحرف الأخير"
challengeType: 29
dashedName: challenge-247
---

# --description--

معطاة سلسلة نصية، أعد الحرف من السلسلة الذي يظهر آخرًا في الأبجدية.

- إذا تعادلت حرفان أو أكثر في الترتيب الأخير في الأبجدية، أعد الأول منها.
- تجاهل كل الأحرف غير الحروف.

# --hints--

`get_last_letter("world")` يجب أن تُعيد `"w"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_last_letter("world"), "w")`)
}})
```

`get_last_letter("Hello World")` يجب أن تُعيد `"W"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_last_letter("Hello World"), "W")`)
}})
```

`get_last_letter("The quick brown fox jumped over the lazy dog.")` يجب أن تُعيد `"z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_last_letter("The quick brown fox jumped over the lazy dog."), "z")`)
}})
```

`get_last_letter("HeLl0")` يجب أن تُعيد `"L"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_last_letter("HeLl0"), "L")`)
}})
```

`get_last_letter("!#$ er@R asd fT.,> 2t0e9")` يجب أن تُعيد `"T"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_last_letter("!#$ er@R asd fT.,> 2t0e9"), "T")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_last_letter(s):

    return s
```

# --solutions--

```py
def get_last_letter(s):
    letters = [c for c in s if c.isalpha()]
    return max(letters, key=lambda c: c.lower())
```
