---
id: 69bc6cb30c1d112a2e110a07
title: "تحدٍّ 250: المفتاح المخفي"
challengeType: 29
dashedName: challenge-250
---

# --description--

مرحبًا بك في تحدٍّ اليومي رقم 250!

معطاة سلسلة نصية مشفرة، فك شفرتها باستخدام مفتاح تشفير وأعدها.

للعثور على المفتاح:

- انظر إلى جميع تحديات اليوميات حتى اليوم التي يكون رقم التحدي فيها من مضاعفات 25 (بما في ذلك هذا التحدي).
- خذ الحرف الأول من كل عنوان من تلك التحديات وادمجها في سلسلة نصية. إذا بدأ العنوان بحرف غير حرف أبجدي، ابحث عن أول حرف أبجدي فيه.

لفك تشفير الرسالة، استعرض كل حرف في الرسالة المشفرة و:

1. انظر إلى الحرف المقابل في المفتاح (كرر المفتاح إذا كانت الرسالة أطول من المفتاح).
  2. حوّل حرف المفتاح إلى رقمه المقابل: `"A"` = 1، `"B"` = 2، ...، `"Z"` = 26.
  3. حرك الحرف المشفر إلى الوراء في الأبجدية بعدد تلك القيمة.
  4. إذا تجاوزت الحركة الحرف `"A"`، عد إلى `"Z"`.

على سبيل المثال، إذا بدأت الرسالة المشفرة بـ `"Y"` وكان أول حرف في المفتاح هو `"V"` (22)، حرك `"Y"` إلى الوراء 22 خانة لتحصل على `"C"`. كرر هذه العملية لكل حرف لفك تشفير الرسالة كاملة.

- تُحرك الحروف فقط، أما الفراغات فتُعاد كما هي.
- كل الحروف المعطاة والمعادة تكون بأحرف كبيرة.

# --hints--

`decode("YAVJYNXE")` يجب أن تُعيد `"CONGRATS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")` يجب أن تُعيد `"CODING LEGEND"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")` يجب أن تُعيد `"YOU ARE AWESOME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")` يجب أن تُعيد `"KEEP GOING"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` يجب أن تُعيد `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```
