---
id: 69c5f3d787b1725d5f00c8bd
title: "تحدٍّ 259: انفجار FizzBuzz"
challengeType: 29
dashedName: challenge-259
---

# --description--

معطى عدد صحيح، أرجع عدد الخطوات اللازمة لتحويل الكلمة `"fizzbuzz"` إلى سلسلة نصية تحتوي على عدد `"z"` على الأقل كما هو معطى، باستخدام القواعد التالية:

- ابدأ بالسلسلة النصية `"fizzbuzz"`.
- في كل خطوة، طبق قواعد FizzBuzz القياسية باستخدام موقع الحرف في السلسلة النصية (الحرف الأول `"f"` في الموقع 1).
  - إذا كان موقع الحرف قابلًا للقسمة على 3، استبدل الحرف بـ `"fizz"`
  - إذا كان قابلًا للقسمة على 5، استبدل الحرف بـ `"buzz"`
  - إذا كان قابلًا للقسمة على 3 و5، استبدل الحرف بـ `"fizzbuzz"`

بعد خطوة واحدة، تتحول `"fizzbuzz"` إلى `"fifizzzbuzzfizzzz"`، والتي تحتوي على 9 أحرف `"z"`.

# --hints--

يجب أن تُرجع الدالة `explode_fizzbuzz(9)` القيمة `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

يجب أن تُرجع الدالة `explode_fizzbuzz(15)` القيمة `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

يجب أن تُرجع الدالة `explode_fizzbuzz(51)` القيمة `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

يجب أن تُرجع الدالة `explode_fizzbuzz(52)` القيمة `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

يجب أن تُرجع الدالة `explode_fizzbuzz(359)` القيمة `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

يجب أن تُرجع الدالة `explode_fizzbuzz(789)` القيمة `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

يجب أن تُرجع الدالة `explode_fizzbuzz(54482)` القيمة `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

يجب أن تُرجع الدالة `explode_fizzbuzz(1000000)` القيمة `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```
