---
id: 69c5f3d787b1725d5f00c8b8
title: "تحدٍّ 254: كلمات ذات عدد فردي من الحروف"
challengeType: 29
dashedName: challenge-254
---

# --description--

بالنظر إلى سلسلة نصية من الكلمات، أعد فقط الكلمات التي تحتوي على عدد فردي من الحروف.

- الكلمات في السلسلة النصية المعطاة ستكون مفصولة بمسافة واحدة.
- أعد الكلمات مفصولة بمسافة واحدة.

# --hints--

`get_odd_words("This is a super good test")` يجب أن تُعيد `"a super"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odd_words("This is a super good test"), "a super")`)
}})
```

`get_odd_words("one two three four")` يجب أن تُعيد `"one two three"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odd_words("one two three four"), "one two three")`)
}})
```

`get_odd_words("banana split sundae with rainbow sprinkles on top")` يجب أن تُعيد `"split rainbow sprinkles top"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odd_words("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top")`)
}})
```

`get_odd_words("The quick brown fox jumped over the lazy river")` يجب أن تُعيد `"The quick brown fox the river"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odd_words("The quick brown fox jumped over the lazy river"), "The quick brown fox the river")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_odd_words(s):

    return s
```

# --solutions--

```py
def get_odd_words(s):
    return ' '.join(word for word in s.split(' ') if len(word) % 2 == 1)
```
