---
id: 69c5f3d787b1725d5f00c8ba
title: "تحدٍّ 256: أقرب اتجاه زمني"
challengeType: 29
dashedName: challenge-256
---

# --description--

معطى وقتان، حدّد ما إذا كان بإمكانك الانتقال من الأول إلى الثاني أسرع بالحركة إلى الأمام أو إلى الخلف.

- الأوقات معطاة بصيغة 24 ساعة (`"HH:MM"`)
- الساعة تدور بشكل دائري (23:59 تتحول إلى 00:00 عند الحركة إلى الأمام، و00:00 تتحول إلى 23:59 عند الحركة إلى الخلف)

أرجع:

- `"forward"` إذا كانت الحركة إلى الأمام أقصر
- `"backward"` إذا كانت الحركة إلى الخلف أقصر
- `"equal"` إذا استغرق الاتجاهان نفس الوقت

# --hints--

`get_direction("10:00", "12:00")` يجب أن تُرجع `"forward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("10:00", "12:00"), "forward")`)
}})
```

`get_direction("11:00", "05:00")` يجب أن تُرجع `"backward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("11:00", "05:00"), "backward")`)
}})
```

`get_direction("00:00", "12:00")` يجب أن تُرجع `"equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("00:00", "12:00"), "equal")`)
}})
```

`get_direction("15:45", "01:10")` يجب أن تُرجع `"forward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("15:45", "01:10"), "forward")`)
}})
```

`get_direction("03:30", "19:50")` يجب أن تُرجع `"backward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("03:30", "19:50"), "backward")`)
}})
```

`get_direction("06:30", "18:30")` يجب أن تُرجع `"equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("06:30", "18:30"), "equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_direction(time1, time2):

    return time1
```

# --solutions--

```py
def get_direction(time1, time2):
    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    start = to_minutes(time1)
    end = to_minutes(time2)
    forward = (end - start + 1440) % 1440
    backward = (start - end + 1440) % 1440

    if forward < backward:
        return 'forward'
    if backward < forward:
        return 'backward'
    return 'equal'
```
