---
id: 69c5f3d787b1725d5f00c8b7
title: "تحدٍّ 253: مكتشف الاختصارات"
challengeType: 29
dashedName: challenge-253
---

# --description--

أعطِ سلسلة نصية تمثل اختصارًا، وأرجع الاسم الكامل للمنظمة التي ينتمي إليها من القائمة أدناه:

- `"National Avocado Storage Authority"`
- `"Cats Infiltration Agency"`
- `"Fluffy Beanbag Inspectors"`
- `"Department Of Jelly"`
- `"Wild Honey Organization"`
- `"Eating Pancakes Administration"`

يجب أن تتطابق كل حرف في الاختصار المعطى مع الحرف الأول من كل كلمة في المنظمة التي ينتمي إليها، بنفس الترتيب.

# --hints--

`find_org("NASA")` يجب أن تُرجع `"National Avocado Storage Authority"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("NASA"), "National Avocado Storage Authority")`)
}})
```

`find_org("CIA")` يجب أن تُرجع `"Cats Infiltration Agency"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("CIA"), "Cats Infiltration Agency")`)
}})
```

`find_org("FBI")` يجب أن تُرجع `"Fluffy Beanbag Inspectors"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("FBI"), "Fluffy Beanbag Inspectors")`)
}})
```

`find_org("DOJ")` يجب أن تُرجع `"Department Of Jelly"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("DOJ"), "Department Of Jelly")`)
}})
```

`find_org("WHO")` يجب أن تُرجع `"Wild Honey Organization"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("WHO"), "Wild Honey Organization")`)
}})
```

`find_org("EPA")` يجب أن تُرجع `"Eating Pancakes Administration"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("EPA"), "Eating Pancakes Administration")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_org(acronym):

    return acronym
```

# --solutions--

```py
def find_org(acronym):
    orgs = [
        "National Avocado Storage Authority",
        "Cats Infiltration Agency",
        "Fluffy Beanbag Inspectors",
        "Department Of Jelly",
        "Wild Honey Organization",
        "Eating Pancakes Administration",
    ]
    for org in orgs:
        initials = ''.join(word[0] for word in org.split())
        if initials.upper() == acronym.upper():
            return org
    return "not found"
```
