---
id: 69c5f3d787b1725d5f00c8b9
title: "تحدٍّ 255: فريق تنظيف يوم الأرض"
challengeType: 29
dashedName: challenge-255
---

# --description--

اليوم هو يوم الأرض. مع وجود مصفوفة من العناصر التي نظفتها، أعد مجموع نقاط التنظيف الكلي بناءً على القواعد التالية.

العناصر المعطاة ستكون واحدة من:

| العنصر | القيمة الأساسية |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

العنصر النادر يُمثّل كـ `["rare", value]`. على سبيل المثال، `["rare", 80]`. العناصر النادرة لا تحصل على مكافأة التتابع.

- مكافأة التتابع: إذا ظهر نفس العنصر بشكل متتالي، يحصل على نقاط مكافأة متزايدة.
  - الظهور المتتالي الأول: القيمة الأساسية
  - الثاني: القيمة الأساسية + 1
  - الثالث: القيمة الأساسية + 2
  - وهكذا.
- مضاعف العنصر الخامس: كل عنصر خامس يتم جمعه يحصل على مضاعف.
  - العنصر الخامس: *2
  - العنصر العاشر: *3
  - وهكذا.

- طبق المضاعف بعد حساب أي مكافآت.

# --hints--

`get_cleanup_score(["bottle", "straw", "shoe", "battery"])` يجب أن تُعيد `44`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["bottle", "straw", "shoe", "battery"]), 44)`)
}})
```

`get_cleanup_score(["electronics", "straw", "newspaper", "bottle", "bag"])` يجب أن تُعيد `58`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["electronics", "straw", "newspaper", "bottle", "bag"]), 58)`)
}})
```

`get_cleanup_score(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` يجب أن تُعيد `79`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79)`)
}})
```

`get_cleanup_score(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` يجب أن تُعيد `358`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358)`)
}})
```

`get_cleanup_score(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` يجب أن تُعيد `383`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_cleanup_score(items):

    return items
```

# --solutions--

```py
def get_cleanup_score(items):
    values = {
        'bottle': 10, 'can': 6, 'bag': 8, 'tire': 35,
        'straw': 4, 'cardboard': 3, 'newspaper': 3,
        'shoe': 12, 'electronics': 25, 'battery': 18, 'mattress': 38
    }
    total = 0
    prev_item = None
    streak = 0
    for i, item in enumerate(items):
        is_rare = isinstance(item, list) and item[0] == 'rare'
        if is_rare:
            base = item[1]
            prev_item = None
            streak = 0
        else:
            base = values[item]
            streak = streak + 1 if item == prev_item else 0
            prev_item = item
        points = base if is_rare else base + streak
        if (i + 1) % 5 == 0:
            multiplier = (i + 1) // 5 + 1
            points *= multiplier
        total += points
    return total
```
