---
id: 69c6ff713a52713463aa7926
title: "تحدٍّ 260: نقاط الكلمة"
challengeType: 29
dashedName: challenge-260
---

# --description--

أعطِ كلمة، وأرجع نقاطها باستخدام جدول قياسي لقيمة الحروف:

| الحرف | القيمة |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- الحروف الكبيرة والأحرف الصغيرة لها نفس القيمة.

# --hints--

`get_word_score("hi")` يجب أن تُرجع `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("hi"), 17)`)
}})
```

`get_word_score("hello")` يجب أن تُرجع `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("hello"), 52)`)
}})
```

`get_word_score("hippopotamus")` يجب أن تُرجع `169`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("hippopotamus"), 169)`)
}})
```

`get_word_score("freeCodeCamp")` يجب أن تُرجع `94`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("freeCodeCamp"), 94)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_word_score(word):

    return word
```

# --solutions--

```py
def get_word_score(word):
    return sum(ord(c) - 96 for c in word.lower())
```
