---
id: 69c6ff713a52713463aa7929
title: "تحدٍّ 263: الكلمات المتقاطعة الثنائية"
challengeType: 29
dashedName: challenge-263
---

# --description--

بالنظر إلى حرف، حدّد ما إذا كان تمثيله الثنائي المكوّن من 8 بت يمكن إيجاده في الشبكة التالية، أفقيًا أو عموديًا في أي اتجاه:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

على سبيل المثال، `"A"` له التمثيل الثنائي `01000001`، الذي يظهر في الصف الأول من اليسار إلى اليمين.

# --hints--

`is_in_crossword("I")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("I"), True)`)
}})
```

`is_in_crossword("D")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("D"), True)`)
}})
```

`is_in_crossword("0")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("0"), True)`)
}})
```

`is_in_crossword("u")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("u"), True)`)
}})
```

`is_in_crossword("Y")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("Y"), False)`)
}})
```

`is_in_crossword("p")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("p"), False)`)
}})
```

`is_in_crossword("1")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("1"), False)`)
}})
```

`is_in_crossword("Q")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("Q"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_in_crossword(char):

    return char
```

# --solutions--

```py
def is_in_crossword(char):
    grid = [
        "01000001",
        "01101111",
        "01000100",
        "01100101",
        "01010010",
        "01010100",
        "01101000",
        "10101110",
    ]
    target = format(ord(char), '08b')
    cols = [''.join(grid[r][c] for r in range(8)) for c in range(8)]
    return any(target in seq or target in seq[::-1] for seq in grid + cols)
```
