---
id: 69c6ff713a52713463aa7928
title: "تحدٍّ 262: محلل استعلام URL"
challengeType: 29
dashedName: challenge-262
---

# --description--

معطى عنوان URL يحتوي على سلسلة استعلام، حلل سلسلة الاستعلام إلى كائن (أو قاموس) من أزواج المفتاح-القيمة.

- تبدأ سلسلة الاستعلام بعد `"?"`،
- تُفصل كل معلمة بـ `"&"`،
- يُفصل كل زوج مفتاح/قيمة بـ `"="`

على سبيل المثال، معطى `"https://example.com/search?name=Alice&age=30"`، أرجع:

```json
{
  "name": "Alice",
  "age": "30"
}
```

يجب أن تُعاد كل القيم كسلاسل نصية.

# --hints--

`parse_url_query("https://example.com/search?name=Alice&age=30")` يجب أن تُعيد `{"name": "Alice", "age": "30"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"})`)
}})
```

`parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python")` يجب أن تُعيد `{"skill": "programming", "language": "python"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"})`)
}})
```

`parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2")` يجب أن تُعيد `{"category": "books", "sort": "asc", "page": "2"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"})`)
}})
```

`parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now")` يجب أن تُعيد `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_url_query(url):

    return url
```

# --solutions--

```py
def parse_url_query(url):
    query_string = url.split('?')[1]
    return dict(param.split('=') for param in query_string.split('&'))
```
