---
id: 69c6ff713a52713463aa7927
title: "تحدٍّ 261: كلمات الأعداد"
challengeType: 29
dashedName: challenge-261
---

# --description--

أعطِ عددًا صحيحًا من 0 إلى 99، وأرجع تمثيله بالكلمات الإنجليزية.

- `0` تُرجع `"zero"`.
- الأعداد من 1 إلى 19 لها أسماء فريدة (`"one"`، `"two"`، ...، `"ten"`، `"eleven"`، ...، `"eighteen"`، `"nineteen"`).
- مضاعفات 10 من 20 إلى 90 لها أسماء خاصة بها (`"twenty"`، `"thirty"`، ...، `"eighty"`، `"ninety"`).
- الأعداد من 21 إلى 99 التي ليست مضاعفات 10 تُكتب ككلمتين مرتبطتين بشرطة. على سبيل المثال `"forty-two"` و `"fifty-three"`.

# --hints--

`get_number_words(0)` يجب أن تُرجع `"zero"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(0), "zero")`)
}})
```

`get_number_words(10)` يجب أن تُرجع `"ten"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(10), "ten")`)
}})
```

`get_number_words(19)` يجب أن تُرجع `"nineteen"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(19), "nineteen")`)
}})
```

`get_number_words(30)` يجب أن تُرجع `"thirty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(30), "thirty")`)
}})
```

`get_number_words(53)` يجب أن تُرجع `"fifty-three"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(53), "fifty-three")`)
}})
```

`get_number_words(7)` يجب أن تُرجع `"seven"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(7), "seven")`)
}})
```

`get_number_words(12)` يجب أن تُرجع `"twelve"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(12), "twelve")`)
}})
```

`get_number_words(60)` يجب أن تُرجع `"sixty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(60), "sixty")`)
}})
```

`get_number_words(67)` يجب أن تُرجع `"sixty-seven"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(67), "sixty-seven")`)
}})
```

`get_number_words(98)` يجب أن تُرجع `"ninety-eight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(98), "ninety-eight")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_words(n):

    return n
```

# --solutions--

```py
def get_number_words(n):
    ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
            "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
            "seventeen","eighteen","nineteen"]
    tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"]

    if n < 20:
        return ones[n]
    if n % 10 == 0:
        return tens[n // 10]
    return tens[n // 10] + "-" + ones[n % 10]
```
