---
id: 69cfca90e8a0a6d4d6871c4e
title: "تحدٍّ 264: مجموعات الكلمات المتجانسة"
challengeType: 29
dashedName: challenge-264
---

# --description--

بالنظر إلى مصفوفة من الكلمات، أعد مصفوفة ثنائية الأبعاد تحتوي الكلمات مجمعة في مجموعات من الكلمات المتجانسة.

- الكلمات متجانسة إذا احتوت على نفس الحروف بأي ترتيب.
- تنتمي كل كلمة إلى مجموعة واحدة فقط.
- ترتيب الإرجاع لا يهم.

على سبيل المثال، إذا كانت المصفوفة `["listen", "silent", "hello", "enlist", "world"]`، فأعد `[["listen", "silent", "enlist"], ["hello"], ["world"]]`.

# --hints--

`group_anagrams(["listen", "silent", "hello", "enlist", "world"])` يجب أن تُعيد `[["listen", "silent", "enlist"], ["hello"], ["world"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["listen", "silent", "hello", "enlist", "world"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["enlist", "listen", "silent"], ["hello"], ["world"]])`)
}})
```

`group_anagrams(["eat", "tea", "tan", "ate", "nat", "bat"])` يجب أن تُعيد `[["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["eat", "tea", "tan", "ate", "nat", "bat"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["ate", "eat", "tea"], ["bat"], ["nat", "tan"]])`)
}})
```

`group_anagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])` يجب أن تُعيد `[["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]])`)
}})
```

`group_anagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])` يجب أن تُعيد `[["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def group_anagrams(words):

    return words
```

# --solutions--

```py
def group_anagrams(words):
    groups = {}
    for word in words:
        key = ''.join(sorted(word))
        if key not in groups:
            groups[key] = []
        groups[key].append(word)
    return list(groups.values())
```
