---
id: 69cfca90e8a0a6d4d6871c4f
title: "تحدٍّ 265: أعمق الأقواس"
challengeType: 29
dashedName: challenge-265
---

# --description--

أعطِ سلسلة نصية تحتوي على أقواس متوازنة، أرجع محتوى أعمق الأقواس المسكنة.

- يمكن أن تكون الأقواس من أي نوع من الأنواع الثلاثة: `()`, `[]`, و `{}`.
- سيكون المُدخَل دائمًا يحتوي على مجموعة واحدة أعمق.

على سبيل المثال، إذا أعطيت `"(hello (world))"`، أرجع `"world"`.

# --hints--

`get_deepest_brackets("(hello (world))")` يجب أن تُرجع `"world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_deepest_brackets("(hello (world))"), "world")`)
}})
```

`get_deepest_brackets("[outer [inner] outer]")` يجب أن تُرجع `"inner"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_deepest_brackets("[outer [inner] outer]"), "inner")`)
}})
```

`get_deepest_brackets("{a{b}c{d{e}f}g}")` يجب أن تُرجع `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_deepest_brackets("{a{b}c{d{e}f}g}"), "e")`)
}})
```

`get_deepest_brackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` يجب أن تُرجع `"fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_deepest_brackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox")`)
}})
```

`get_deepest_brackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` يجب أن تُرجع `"C"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_deepest_brackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_deepest_brackets(s):

    return s
```

# --solutions--

```py
def get_deepest_brackets(s):
    max_depth = 0
    depth = 0
    content = ''
    start = 0

    for i, ch in enumerate(s):
        if ch in '([{':
            depth += 1
            if depth > max_depth:
                max_depth = depth
                start = i + 1
        elif ch in ')]}':
            if depth == max_depth:
                content = s[start:i]
            depth -= 1

    return content
```
