---
id: 69cfca90e8a0a6d4d6871c52
title: "تحدٍّ 268: الرقم النرجسي"
challengeType: 29
dashedName: challenge-268
---

# --description--

أعطِ عددًا صحيحًا موجبًا، حدّد ما إذا كان رقمًا نرجسيًا.

- الرقم نرجسي إذا كان مجموع كل رقم من أرقامه مرفوعًا لقوة عدد الأرقام يساوي الرقم نفسه.

على سبيل المثال، الرقم 153 يحتوي على 3 أرقام، و1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153، إذًا هو نرجسي.

# --hints--

يجب أن تُعيد الدالة `is_narcissistic(153)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(153), True)`)
}})
```

يجب أن تُعيد الدالة `is_narcissistic(154)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(154), False)`)
}})
```

يجب أن تُعيد الدالة `is_narcissistic(371)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(371), True)`)
}})
```

يجب أن تُعيد الدالة `is_narcissistic(512)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(512), False)`)
}})
```

يجب أن تُعيد الدالة `is_narcissistic(9)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9), True)`)
}})
```

يجب أن تُعيد الدالة `is_narcissistic(11)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(11), False)`)
}})
```

يجب أن تُعيد الدالة `is_narcissistic(9474)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9474), True)`)
}})
```

يجب أن تُعيد الدالة `is_narcissistic(6549)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(6549), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_narcissistic(n):

    return n
```

# --solutions--

```py
def is_narcissistic(n):
    digits = str(n)
    power = len(digits)
    return sum(int(d) ** power for d in digits) == n
```
