---
id: 69cfca90e8a0a6d4d6871c50
title: "تحدٍّ 266: يوم جيد"
challengeType: 29
dashedName: challenge-266
---

# --description--

أعطِ سلسلة نصية للوقت بصيغة `"HH:MM"` (ساعة 24)، أرجع:

- `"Good morning"` للأوقات من `05:00` إلى `11:59`
- `"Good afternoon"` للأوقات من `12:00` إلى `17:59`
- `"Good evening"` للأوقات من `18:00` إلى `21:59`
- `"Good night"` للأوقات من `22:00` إلى `04:59`

# --hints--

`get_greeting("06:30")` يجب أن تُرجع `"Good morning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("06:30"), "Good morning")`)
}})
```

`get_greeting("12:00")` يجب أن تُرجع `"Good afternoon"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("12:00"), "Good afternoon")`)
}})
```

`get_greeting("21:59")` يجب أن تُرجع `"Good evening"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("21:59"), "Good evening")`)
}})
```

`get_greeting("00:01")` يجب أن تُرجع `"Good night"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("00:01"), "Good night")`)
}})
```

`get_greeting("11:30")` يجب أن تُرجع `"Good morning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("11:30"), "Good morning")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_greeting(s):

    return s
```

# --solutions--

```py
def get_greeting(s):
    hours, minutes = map(int, s.split(':'))
    total = hours * 60 + minutes

    if 300 <= total < 720:
        return 'Good morning'
    if 720 <= total < 1080:
        return 'Good afternoon'
    if 1080 <= total < 1320:
        return 'Good evening'
    return 'Good night'
```
