---
id: 69cfca90e8a0a6d4d6871c51
title: "تحدٍّ 267: محول Parsec"
challengeType: 29
dashedName: challenge-267
---

# --description--

في مجرة بعيدة، تُستخدم الـ parsecs لقياس كل من الزمن والمسافة. معطى عدد صحيح من الـ parsecs، أرجع ما يعادله من الزمن أو المسافة.

- إذا كان العدد الصحيح المعطى فرديًا، فهو يمثل الزمن. وإذا كان زوجيًا، فهو يمثل المسافة.

استخدم معدلات التحويل هذه:

| Parsecs | الزمن/المسافة |
| - | - |
| 1 | ساعتان |
| 2 | 6 سنوات ضوئية |

أرجع القيمة المحولة كعدد صحيح.

# --hints--

`convert_parsecs(1)` يجب أن تُرجع `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(1), 2)`)
}})
```

`convert_parsecs(2)` يجب أن تُرجع `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(2), 6)`)
}})
```

`convert_parsecs(31)` يجب أن تُرجع `62`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(31), 62)`)
}})
```

`convert_parsecs(88)` يجب أن تُرجع `264`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(88), 264)`)
}})
```

`convert_parsecs(17)` يجب أن تُرجع `34`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(17), 34)`)
}})
```

`convert_parsecs(14)` يجب أن تُرجع `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(14), 42)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_parsecs(parsecs):

    return parsecs
```

# --solutions--

```py
def convert_parsecs(parsecs):
    if parsecs % 2 != 0:
        return parsecs * 2
    else:
        return parsecs * 3
```
