---
id: 69cfca90e8a0a6d4d6871c54
title: "تحدٍّ 270: أطول سلسلة نصية مشتركة"
challengeType: 29
dashedName: challenge-270
---

# --description--

معطاة سلسلة نصية، أعد أطول سلسلة نصية فرعية تظهر أكثر من مرة.

- يمكن أن تتداخل السلاسل النصية الفرعية.

# --hints--

`get_longest_substring("abracadabra")` يجب أن تُعيد `"abra"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("abracadabra"), "abra")`)
}})
```

`get_longest_substring("hello world hello")` يجب أن تُعيد `"hello"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("hello world hello"), "hello")`)
}})
```

`get_longest_substring("mississippi")` يجب أن تُعيد `"issi"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("mississippi"), "issi")`)
}})
```

`get_longest_substring("ha ha ha ha ha ha ha")` يجب أن تُعيد `"ha ha ha ha ha ha"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("ha ha ha ha ha ha ha"), "ha ha ha ha ha ha")`)
}})
```

`get_longest_substring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over")` يجب أن تُعيد `"the quick brown fox jumped over"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over"), "the quick brown fox jumped over")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_longest_substring(s):

    return s
```

# --solutions--

```py
def get_longest_substring(s):
    longest = ''

    for length in range(len(s) - 1, 0, -1):
        for i in range(len(s) - length + 1):
            sub = s[i:i + length]
            if s.find(sub) != s.rfind(sub):
                return sub

    return longest
```
