---
id: 69d03e549613fbdbe21d11e0
title: "تحدٍّ 271: تذكير الدواء"
challengeType: 29
dashedName: challenge-271
---

# --description--

معطاة مصفوفة من الأدوية وسلسلة نصية تمثل الوقت الحالي، أعد الدواء التالي الذي تحتاج إلى تناوله وكم من الوقت تبقى حتى تحتاج لتناوله.

- كل دواء يكون بالشكل `[name, lastTaken]`، حيث `name` هو اسم الدواء و`lastTaken` هو الوقت الذي تم تناوله فيه آخر مرة.
- كل الأوقات المعطاة ستكون بصيغة `"HH:MM"` (24 ساعة).

استخدم جدول الأدوية التالي:

| الاسم | الجدول |
| - | - |
| Deployxitrin | 08:00، 16:00 |
| Debuggamanizole | 07:00، 13:00، 21:00 |
| Mergeflictamine | كل 4 ساعات |

أعد سلسلة نصية بصيغة `"{name} in Hh Mm"`. على سبيل المثال، `"Debuggamanizole in 2h 0m"` أو `"Deployxitrin in 1h 5m"`.

# --hints--

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00")` يجب أن تعيد `"Debuggamanizole in 2h 0m"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00"), "Debuggamanizole in 2h 0m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55")` يجب أن تعيد `"Deployxitrin in 1h 5m"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55"), "Deployxitrin in 1h 5m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15")` يجب أن تعيد `"Mergeflictamine in 0h 45m"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15"), "Mergeflictamine in 0h 45m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59")` يجب أن تعيد `"Debuggamanizole in 0h 1m"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59"), "Debuggamanizole in 0h 1m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55")` يجب أن تعيد `"Debuggamanizole in 0h 5m"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55"), "Debuggamanizole in 0h 5m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00")` يجب أن تعيد `"Mergeflictamine in 3h 30m"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00"), "Mergeflictamine in 3h 30m")`)
}})
```

# --seed--

## --seed-contents--

```py
def medication_reminder(medications, current_time):

    return medications
```

# --solutions--

```py
def medication_reminder(medications, current_time):
    schedule = {
        'Deployxitrin': {'type': 'fixed', 'times': ['08:00', '16:00']},
        'Debuggamanizole': {'type': 'fixed', 'times': ['07:00', '13:00', '21:00']},
        'Mergeflictamine': {'type': 'interval', 'hours': 4},
    }

    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    current = to_minutes(current_time)
    earliest = float('inf')
    next_med = ''

    for name, last_taken in medications:
        med = schedule[name]
        if med['type'] == 'fixed':
            future_times = [to_minutes(t) for t in med['times'] if to_minutes(t) > current]
            if not future_times:
                continue
            next_due = min(future_times)
        else:
            next_due = to_minutes(last_taken) + med['hours'] * 60

        diff = next_due - current
        if 0 < diff < earliest:
            earliest = diff
            next_med = name

    h = earliest // 60
    m = earliest % 60
    return f"{next_med} in {h}h {m}m"
```
