---
id: 69e2383af7832c8032603b8f
title: "تحدٍّ 273: مدقق ISBN-13"
challengeType: 29
dashedName: challenge-273
---

# --description--

معطاة سلسلة نصية، حدّد إذا كانت رقم ISBN-13 صالحًا.

رقم ISBN-13 صالح:

- يحتوي فقط على أرقام وواصلات
- يحتوي على 13 رقمًا بالضبط بعد إزالة الواصلات
- يجتاز الفحص التالي:
  1. اضرب كل رقم في 1 أو 3 بالتناوب (اضرب الرقم الأول في 1، والثاني في 3، والثالث في 1، وهكذا).
  2. يجب أن يكون مجموع النتائج قابلاً للقسمة على 10.

# --hints--

`is_valid_isbn_13("9780306406157")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9780306406157"), True)`)
}})
```

`is_valid_isbn_13("97803064061570")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("97803064061570"), False)`)
}})
```

`is_valid_isbn_13("978-0-13-595705-9")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-0-13-595705-9"), True)`)
}})
```

`is_valid_isbn_13("978-030-64061A-4")` يجب أن تُرجع `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-030-64061A-4"), False)`)
}})
```

`is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9")` يجب أن تُرجع `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_isbn_13(s):

    return s
```

# --solutions--

```py
def is_valid_isbn_13(s):
    digits = s.replace("-", "")
    if not digits.isdigit() or len(digits) != 13:
        return False
    total = sum(int(d) * (1 if i % 2 == 0 else 3) for i, d in enumerate(digits))
    return total % 10 == 0
```
