---
id: 69e2383af7832c8032603b90
title: "تحدٍّ 274: أقدم شخص"
challengeType: 29
dashedName: challenge-274
---

# --description--

معطاة مصفوفة من الكائنات، كل منها يحتوي على خاصيتي `"name"` و`"age"`، أعد مصفوفة تحتوي على أسماء أقدم الأشخاص.

إذا شارك عدة أشخاص في أكبر عمر، أعد جميع أسمائهم بالترتيب الذي يظهرون به في المُدخَل.

# --hints--

`get_oldest([{ "name": "Brenda", "age": 40 }])` يجب أن تُعيد `["Brenda"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "Brenda", "age": 40}]), ["Brenda"])`)
}})
```

`get_oldest([{ "name": "Alice", "age": 30 }, { "name": "Bob", "age": 25 }])` يجب أن تُعيد `["Alice"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "Alice", "age": 30}, {"name": "Bob", "age": 25}]), ["Alice"])`)
}})
```

`get_oldest([{ "name": "Allison", "age": 25 }, { "name": "Bill", "age": 30 }, { "name": "Carol", "age": 30 }])` يجب أن تُعيد `["Bill", "Carol"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "Allison", "age": 25}, {"name": "Bill", "age": 30}, {"name": "Carol", "age": 30}]), ["Bill", "Carol"])`)
}})
```

`get_oldest([{ "name": "George", "age": 50 }, { "name": "Shirley", "age": 42 }, { "name": "Beth", "age": 48 }, { "name": "Holly", "age": 50 }, { "name": "Kevin", "age": 44 }, { "name": "Frank", "age": 47 }, { "name": "Zach", "age": 50 }, { "name": "Jennifer", "age": 43 }])` يجب أن تُعيد `["George", "Holly", "Zach"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "George", "age": 50}, {"name": "Shirley", "age": 42}, {"name": "Beth", "age": 48}, {"name": "Holly", "age": 50}, {"name": "Kevin", "age": 44}, {"name": "Frank", "age": 47}, {"name": "Zach", "age": 50}, {"name": "Jennifer", "age": 43}]), ["George", "Holly", "Zach"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_oldest(people):

    return people
```

# --solutions--

```py
def get_oldest(people):
    max_age = max(p["age"] for p in people)
    return [p["name"] for p in people if p["age"] == max_age]
```
