---
id: 69e2383af7832c8032603b92
title: "تحدٍّ 276: العنصر المخالف"
challengeType: 29
dashedName: challenge-276
---

# --description--

معطاة مصفوفة من الأعداد الصحيحة مرتبة ترتيبًا تصاعديًا ما عدا عنصر واحد في غير موضعه، أعد فهرس ذلك العنصر.

- إذا كان يمكن اعتبار أكثر من عنصر في غير موضعه، أعد فهرس أول عنصر منهم.

# --hints--

`find_offender([1, 6, 2, 3, 4, 5])` يجب أن تُعيد `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([1, 6, 2, 3, 4, 5]), 1)`)
}})
```

`find_offender([1, 2, 3, 5, 4, 5])` يجب أن تُعيد `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([1, 2, 3, 5, 4, 5]), 3)`)
}})
```

`find_offender([2, 1])` يجب أن تُعيد `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([2, 1]), 0)`)
}})
```

`find_offender([2, 4, 1, 6, 8])` يجب أن تُعيد `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([2, 4, 1, 6, 8]), 2)`)
}})
```

`find_offender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])` يجب أن تُعيد `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_offender(arr):

    return arr
```

# --solutions--

```py
def find_offender(arr):
    for i in range(len(arr)):
        without = arr[:i] + arr[i+1:]
        if all(without[j-1] <= without[j] for j in range(1, len(without))):
            return i
```
