---
id: 69f35a5bb823ed620fcb7cbd
title: "تحدٍّ 284: I قبل E"
challengeType: 29
dashedName: challenge-284
---

# --description--

عند إعطائك كلمة أو جملة، أعد نسخة مصححة حيث تتبع كل كلمة قاعدة "I قبل E إلا بعد C".

- إذا احتوت الكلمة على `"ei"` ولم تسبقها `"c"`، استبدلها بـ `"ie"`.
- إذا احتوت الكلمة على `"ie"` وسبقها `"c"`، استبدلها بـ `"ei"`.
- تُترك كل الكلمات الأخرى دون تغيير.

# --hints--

`i_before_e("beleive")` يجب أن تُعيد `"believe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("beleive"), "believe")`)
}})
```

`i_before_e("recieve")` يجب أن تُعيد `"receive"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("recieve"), "receive")`)
}})
```

`i_before_e("we recieved a breif")` يجب أن تُعيد `"we received a brief"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("we recieved a breif"), "we received a brief")`)
}})
```

`i_before_e("she beleived the friendly niece could percieve the greif")` يجب أن تُعيد `"she believed the friendly niece could perceive the grief"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief")`)
}})
```

`i_before_e("we recieved relief after the theif gave us a breif piece of feirce deceit")` يجب أن تُعيد `"we received relief after the thief gave us a brief piece of fierce deceit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit")`)
}})
```

# --seed--

## --seed-contents--

```py
def i_before_e(sentence):

    return sentence
```

# --solutions--

```py
import re

def i_before_e(sentence):
    words = []
    for word in sentence.split(" "):
        word = re.sub(r"([^c])ei", r"\1ie", word)
        word = re.sub(r"cie", "cei", word)
        words.append(word)
    return " ".join(words)
```
