---
id: 69f35a5bb823ed620fcb7cb9
title: "تحدٍّ 280: تاريخ معرف Mongo"
challengeType: 29
dashedName: challenge-280
---

# --description--

بالنظر إلى سلسلة نصية لمعرف MongoDB، أعد وقت إنشائه كسلسلة نصية بتنسيق ISO 8601.

- معرف MongoDB هو سلسلة سداسية عشرية مكونة من 24 حرفًا. تمثل الأحرف الثمانية الأولى طابعًا زمنيًا لنظام Unix (بالثواني) مشفرًا كعدد صحيح بالقاعدة 16.

على سبيل المثال، `"6a094b50bcf86cd799439011"` يحتوي على طابع زمني `"6a094b50"` بالسداسي العشري، وهو `1778994000` بالعشري، ويمثل وقت إنشاء `"2026-05-17T05:00:00.000Z"`.

# --hints--

`mongo_id_to_date("6a094b50bcf86cd799439011")` يجب أن يعيد `"2026-05-17T05:00:00.000Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("6a094b50bcf86cd799439011"), "2026-05-17T05:00:00.000Z")`)
}})
```

`mongo_id_to_date("695344eb1f4a4c1123042128")` يجب أن يعيد `"2025-12-30T03:20:11.000Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("695344eb1f4a4c1123042128"), "2025-12-30T03:20:11.000Z")`)
}})
```

`mongo_id_to_date("386da62df34123ac54617e56")` يجب أن يعيد `"2000-01-01T07:01:01.000Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("386da62df34123ac54617e56"), "2000-01-01T07:01:01.000Z")`)
}})
```

`mongo_id_to_date("69f571c3d7711807afd3dd55")` يجب أن يعيد `"2026-05-02T03:38:43.000Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("69f571c3d7711807afd3dd55"), "2026-05-02T03:38:43.000Z")`)
}})
```

`mongo_id_to_date("68adce01c0e1144d0a90295a")` يجب أن يعيد `"2025-08-26T15:08:49.000Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("68adce01c0e1144d0a90295a"), "2025-08-26T15:08:49.000Z")`)
}})
```

# --seed--

## --seed-contents--

```py
def mongo_id_to_date(s):

    return s
```

# --solutions--

```py
from datetime import datetime, timezone

def mongo_id_to_date(s):
    timestamp = int(s[:8], 16)
    dt = datetime.fromtimestamp(timestamp, tz=timezone.utc)
    return dt.strftime("%Y-%m-%dT%H:%M:%S.000Z")
```
