---
id: 69f35a5bb823ed620fcb7cbe
title: "تحدٍّ 285: وقت الاجتماع"
challengeType: 29
dashedName: challenge-285
---

# --description--

معطى مصفوفة ثلاثية الأبعاد تمثل نوافذ التوفر لعدة أشخاص، أعد أقرب وقت يكون فيه الجميع متاحًا لمدة ساعة واحدة. إذا لم يوجد مثل هذا الوقت، أعد `"None"`.

- توفر كل شخص هو مصفوفة من أزواج صحيحة `[start, end]` تمثل الوقت بنظام 24 ساعة. على سبيل المثال، `[10, 12]` تعني أن الشخص متاح من الساعة 10 إلى 12. تبدأ أوقات البداية من 0 إلى 23، وأوقات النهاية من 1 إلى 24.

على سبيل المثال، معطى:

```json
[
  [[10, 12], [15, 16]], // person 1
  [[11, 14], [15, 16]]  // person 2
]
```

أعد `11`، بداية أول ساعة مشتركة متاحة لهم.

# --hints--

`get_meeting_time([[[10, 12], [15, 16]], [[11, 14], [15, 16]]])` يجب أن تعيد `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[10, 12], [15, 16]], [[11, 14], [15, 16]]]), 11)`)
}})
```

`get_meeting_time([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]])` يجب أن تعيد `13`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]]), 13)`)
}})
```

`get_meeting_time([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]])` يجب أن تعيد `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]]), 9)`)
}})
```

`get_meeting_time([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]])` يجب أن تعيد `None`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]]), "None")`)
}})
```

`get_meeting_time([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]])` يجب أن تعيد `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]]), 21)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_meeting_time(availability):

    return availability
```

# --solutions--

```py
def get_meeting_time(availability):
    for hour in range(24):
        if all(
            any(start <= hour and hour + 1 <= end for start, end in person)
            for person in availability
        ):
            return hour
    return "None"
```
