---
id: 69f35a5bb823ed620fcb7cbb
title: "تحدٍّ 282: دين النوم"
challengeType: 29
dashedName: challenge-282
---

# --description--

معطاة مصفوفة من ساعات النوم كل ليلة حتى اليوم، وعدد ساعات مستهدف لكل ليلة، أرجع عدد ساعات النوم التي تحتاجها الليلة لتسديد دين النوم.

- أدرج ساعات الليلة ضمن إجمالي الوقت اللازم للتعويض.
- إذا نمت بما يكفي لتغطية الهدف الليلي أو أكثر، أرجع `0`.

# --hints--

`sleep_debt([6, 6, 6, 6, 6, 6], 8)` يجب أن تُرجع `20`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([6, 6, 6, 6, 6, 6], 8), 20)`)
}})
```

`sleep_debt([6, 7, 8, 4, 8, 6], 7)` يجب أن تُرجع `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([6, 7, 8, 4, 8, 6], 7), 10)`)
}})
```

`sleep_debt([10, 10, 9, 10, 9, 11], 9)` يجب أن تُرجع `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([10, 10, 9, 10, 9, 11], 9), 4)`)
}})
```

`sleep_debt([8, 7, 6, 7, 6, 8], 6)` يجب أن تُرجع `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([8, 7, 6, 7, 6, 8], 6), 0)`)
}})
```

`sleep_debt([8, 9, 10, 9, 10, 7], 7)` يجب أن تُرجع `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([8, 9, 10, 9, 10, 7], 7), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sleep_debt(hours_slept, target_hours):

    return hours_slept
```

# --solutions--

```py
def sleep_debt(hours_slept, target_hours):
    debt = (len(hours_slept) + 1) * target_hours - sum(hours_slept)
    return max(0, debt)
```
