---
id: 69f35a5bb823ed620fcb7cbc
title: "تحدٍّ 283: دمج سلسلتي نص"
challengeType: 29
dashedName: challenge-283
---

# --description--

بالنظر إلى سلسلتي نص، أعد سلسلة نصية جديدة تدمج حروفهما واحدًا تلو الآخر. إذا كانت إحدى السلسلتين أطول، أضف الحروف المتبقية في النهاية.

ابدأ بالحرف الأول من السلسلة النصية الأولى.

# --hints--

`zip_strings("abc", "123")` يجب أن تُعيد `"a1b2c3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("abc", "123"), "a1b2c3")`)
}})
```

`zip_strings("acegikmoqsuwy", "bdfhjlnprtvxz")` يجب أن تُعيد `"abcdefghijklmnopqrstuvwxyz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz")`)
}})
```

`zip_strings("day", "night")` يجب أن تُعيد `"dnaiyght"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("day", "night"), "dnaiyght")`)
}})
```

`zip_strings("python", "javascript")` يجب أن تُعيد `"pjyatvhaosncript"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("python", "javascript"), "pjyatvhaosncript")`)
}})
```

`zip_strings("feCdCm", "reoeap")` يجب أن تُعيد `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("feCdCm", "reoeap"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def zip_strings(a, b):

    return a
```

# --solutions--

```py
def zip_strings(a, b):
    result = ""
    for i in range(max(len(a), len(b))):
        if i < len(a):
            result += a[i]
        if i < len(b):
            result += b[i]
    return result
```
