---
id: 69f8c998d78ad3171a0713bb
title: "تحدٍّ 288: الرقم السري"
challengeType: 29
dashedName: challenge-288
---

# --description--

معطى رقم سري وتخمين، حدّد إذا كان التخمين صحيحًا.

أرجع:

- `"higher"` إذا كان الرقم السري أعلى من التخمين.
- `"lower"` إذا كان الرقم السري أقل من التخمين.
- `"you got it!"` إذا كان التخمين صحيحًا.

# --hints--

`guess_number(50, 30)` يجب أن تُرجع `"higher"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(50, 30), "higher")`)
}})
```

`guess_number(85, 99)` يجب أن تُرجع `"lower"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(85, 99), "lower")`)
}})
```

`guess_number(2026, 2026)` يجب أن تُرجع `"you got it!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(2026, 2026), "you got it!")`)
}})
```

`guess_number(92904, 11283)` يجب أن تُرجع `"higher"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(92904, 11283), "higher")`)
}})
```

`guess_number(230495, 423920)` يجب أن تُرجع `"lower"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(230495, 423920), "lower")`)
}})
```

`guess_number(120349, 120349)` يجب أن تُرجع `"you got it!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(120349, 120349), "you got it!")`)
}})
```

# --seed--

## --seed-contents--

```py
def guess_number(secret, guess):

    return secret
```

# --solutions--

```py
def guess_number(secret, guess):
    if guess == secret:
        return "you got it!"
    return "higher" if guess < secret else "lower"
```
