---
id: 69f8c998d78ad3171a0713bc
title: "تحدٍّ 289: مجموع الفروقات"
challengeType: 29
dashedName: challenge-289
---

# --description--

معطاة مصفوفة من الأرقام، أرجع مجموع الفروقات بين كل رقم والذي يليه.

على سبيل المثال، معطاة `[1, 3, 4]`، أرجع `3` (2 + 1).

# --hints--

`sum_of_differences([1, 3, 4])` يجب أن تُرجع `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([1, 3, 4]), 3)`)
}})
```

`sum_of_differences([5, -3, 3, 9, 10])` يجب أن تُرجع `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([5, -3, 3, 9, 10]), 5)`)
}})
```

`sum_of_differences([9, 6, 15, -20, 33, 14, 25, 16, -7])` يجب أن تُرجع `-16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16)`)
}})
```

`sum_of_differences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])` يجب أن تُرجع `25`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_of_differences(arr):

    return arr
```

# --solutions--

```py
def sum_of_differences(arr):
    return sum(arr[i + 1] - arr[i] for i in range(len(arr) - 1))
```
