---
id: 69f8c998d78ad3171a0713ba
title: "تحدٍّ 287: مصحح الأعداد الرومانية"
challengeType: 29
dashedName: challenge-287
---

# --description--

بالنظر إلى سلسلة نصية من الأعداد الرومانية المشوهة، أعد القيمة بصيغة الأعداد الرومانية القياسية.

سيستخدم المُدخَل فقط التدوين الجمعي، بحيث يضيف كل رمز قيمته إلى المجموع. للتذكير، هذه هي الرموز والقيم:

| الرمز | القيمة |
| ------ | ----- |
| `"I"`  | 1     |
| `"V"`  | 5     |
| `"X"`  | 10    |
| `"L"`  | 50    |
| `"C"`  | 100   |
| `"D"`  | 500   |
| `"M"`  | 1000  |

عند إعادة الترميز، استخدم أكبر رمز ممكن في كل خطوة، مع استخدام الأزواج الطارحة (`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`) عند الحاجة.

# --hints--

`fix_numerals("XIIIII")` يجب أن تُعيد `"XV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("XIIIII"), "XV")`)
}})
```

`fix_numerals("IIIILX")` يجب أن تُعيد `"LXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("IIIILX"), "LXIV")`)
}})
```

`fix_numerals("XXVVVIIIII")` يجب أن تُعيد `"XL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("XXVVVIIIII"), "XL")`)
}})
```

`fix_numerals("MDCCLXXXXVIIII")` يجب أن تُعيد `"MDCCXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("MDCCLXXXXVIIII"), "MDCCXCIX")`)
}})
```

`fix_numerals("IIIIVVVVXXXXLLLLCCDD")` يجب أن تُعيد `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV")`)
}})
```

`fix_numerals("ILCDMIVDIIXLCVCXDL")` يجب أن تُعيد `"MMCMLXXXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV")`)
}})
```

# --seed--

## --seed-contents--

```py
def fix_numerals(s):

    return s
```

# --solutions--

```py
def fix_numerals(s):
    values = {"I": 1, "V": 5, "X": 10, "L": 50, "C": 100, "D": 500, "M": 1000}
    encoding = [
        (1000, "M"), (900, "CM"), (500, "D"), (400, "CD"),
        (100, "C"), (90, "XC"), (50, "L"), (40, "XL"),
        (10, "X"), (9, "IX"), (5, "V"), (4, "IV"), (1, "I")
    ]

    total = sum(values[c] for c in s)

    result = ""
    for value, symbol in encoding:
        while total >= value:
            result += symbol
            total -= value
    return result
```
