---
id: 69f8c998d78ad3171a0713b9
title: "تحدٍّ 286: القضايا المفتوحة"
challengeType: 29
dashedName: challenge-286
---

# --description--

معطاة مصفوفة بأرقام القضايا ومصفوفة أخرى بأرقام طلبات السحب (PR)، أعد مصفوفة بالقضايا التي تبقى مفتوحة بعد دمج جميع طلبات السحب.

- يغلق طلب السحب قضية إذا كانت أرقامهما تدويرًا لبعضها البعض. على سبيل المثال، تُغلق القضية 123 بواسطة طلب السحب 231 أو 312.
- لا يغلق طلب السحب قضية تحمل نفس الرقم بالضبط. على سبيل المثال، طلب السحب 123 لا يغلق القضية 123. لذلك، لا يمكن إغلاق قضية أرقامها كلها متطابقة.
- يمكن أن يُحذف الصفر الرائد من أي رقم. على سبيل المثال، يغلق طلب السحب 201 القضية 12 (012، تدوير لـ 201). وبالمثل، تُغلق القضية 201 بواسطة طلب السحب 12.

أعد القضايا المفتوحة المتبقية بنفس الترتيب الذي أُعطيت به.

# --hints--

`get_open_issues([123, 234], [231])` يجب أن تُعيد `[234]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 234], [231]), [234])`)
}})
```

`get_open_issues([123, 345, 16], [345, 231])` يجب أن تُعيد `[345, 16]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 345, 16], [345, 231]), [345, 16])`)
}})
```

`get_open_issues([456, 332, 12, 15], [201, 945, 180])` يجب أن تُعيد `[456, 332, 15]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15])`)
}})
```

`get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` يجب أن تُعيد `[115, 296, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24])`)
}})
```

`get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` يجب أن تُعيد `[95, 395, 754, 296, 709, 237, 1802]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_open_issues(issues, prs):

    return issues
```

# --solutions--

```py
def get_open_issues(issues, prs):
    def get_rotations(n):
        s = str(n)
        rotations = set()
        for i in range(len(s)):
            rotated = s[i:] + s[:i]
            rotations.add(int(rotated))
        return rotations

    def are_rotations(a, b):
        if a == b:
            return False
        return b in get_rotations(a)

    return [issue for issue in issues if not any(are_rotations(pr, issue) or are_rotations(issue, pr) for pr in prs)]
```
