---
id: 69f8c998d78ad3171a0713bf
title: "تحدٍّ 292: نسبة العرض إلى الارتفاع الأوسع"
challengeType: 29
dashedName: challenge-292
---

# --description--

معطى سلسلتا نص لأبعاد صورة مختلفة، أرجع نسبة العرض إلى الارتفاع للصورة التي لها نسبة عرض إلى ارتفاع أكبر.

- ستكون السلاسل النصية المعطاة بصيغة `"WxH"`، مثلًا `"1920x1080"`.
- نسبة العرض إلى الارتفاع هي نسبة العرض إلى الارتفاع، مختزلة إلى أصغر أعداد صحيحة. مثلًا، `"1920x1080"` تُختزل إلى `"16:9"`.
- أرجع سلسلة نصية بصيغة `"W:H"`، مثلًا `"16:9"`.

# --hints--

`get_wider_aspect_ratio("1920x1080", "800x600")` يجب أن تُرجع `"16:9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1920x1080", "800x600"), "16:9")`)
}})
```

`get_wider_aspect_ratio("1080x1350", "2048x1536")` يجب أن تُرجع `"4:3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1080x1350", "2048x1536"), "4:3")`)
}})
```

`get_wider_aspect_ratio("640x480", "2440x1220")` يجب أن تُرجع `"2:1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("640x480", "2440x1220"), "2:1")`)
}})
```

`get_wider_aspect_ratio("360x640", "1080x1920")` يجب أن تُرجع `"9:16"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("360x640", "1080x1920"), "9:16")`)
}})
```

`get_wider_aspect_ratio("3440x1440", "2048x858")` يجب أن تُرجع `"43:18"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("3440x1440", "2048x858"), "43:18")`)
}})
```

`get_wider_aspect_ratio("12345x61234", "12534x51234")` يجب أن تُرجع `"2089:8539"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("12345x61234", "12534x51234"), "2089:8539")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_wider_aspect_ratio(a, b):

    return a
```

# --solutions--

```py
from math import gcd

def get_wider_aspect_ratio(a, b):
    def parse(s):
        w, h = map(int, s.split("x"))
        g = gcd(w, h)
        return w / h, f"{w // g}:{h // g}"

    ratio_a, simplified_a = parse(a)
    ratio_b, simplified_b = parse(b)

    return simplified_a if ratio_a >= ratio_b else simplified_b
```
