---
id: 69f8c998d78ad3171a0713be
title: "تحدٍّ 291: عد FizzBuzz"
challengeType: 29
dashedName: challenge-291
---

# --description--

معطى رقم بداية ورقم نهاية، احسب عدد مرات ظهور fizz و buzz في النطاق (شامل).

- الأرقام القابلة للقسمة على 3 تُحتسب كـ fizz.
- الأرقام القابلة للقسمة على 5 تُحتسب كـ buzz.
- الأرقام القابلة للقسمة على كل من 3 و 5 تُحتسب كـ fizz و buzz معًا.

أرجع كائنًا أو قاموسًا يحتوي على العدادات بالشكل: `{ fizz, buzz }`.

# --hints--

`fizz_buzz_count(1, 11)` يجب أن تُرجع `{"fizz": 3, "buzz": 2}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(1, 11), {"fizz": 3, "buzz": 2})`)
}})
```

`fizz_buzz_count(14, 41)` يجب أن تُرجع `{"fizz": 9, "buzz": 6}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(14, 41), {"fizz": 9, "buzz": 6})`)
}})
```

`fizz_buzz_count(24, 100)` يجب أن تُرجع `{"fizz": 26, "buzz": 16}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(24, 100), {"fizz": 26, "buzz": 16})`)
}})
```

`fizz_buzz_count(-635, -14)` يجب أن تُرجع `{"fizz": 207, "buzz": 125}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-635, -14), {"fizz": 207, "buzz": 125})`)
}})
```

`fizz_buzz_count(-5432, 6789)` يجب أن تُرجع `{"fizz": 4074, "buzz": 2444}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-5432, 6789), {"fizz": 4074, "buzz": 2444})`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_count(start, end):

    return start
```

# --solutions--

```py
def fizz_buzz_count(start, end):
    fizz = sum(1 for i in range(start, end + 1) if i % 3 == 0)
    buzz = sum(1 for i in range(start, end + 1) if i % 5 == 0)
    return {"fizz": fizz, "buzz": buzz}
```
