---
id: 69f90c1329a94b37e2a2086d
title: "تحدٍّ 294: تراكيب الأقواس العادية"
challengeType: 29
dashedName: challenge-294
---

# --description--

معطى عدد صحيح `n`، أعد عدد التراكيب الصالحة لـ `n` أزواج من الأقواس العادية.

- التركيبة الصالحة هي سلسلة نصية يكون لكل قوس عادي مفتوح قوس عادي مغلق مطابق، ولا يظهر أي قوس عادي مغلق قبل قوسه العادي المفتوح المقابل.

على سبيل المثال، عند إعطاء `2`، هناك تراكيب صالحة عددها 2:

```json
(())
()()
```

# --hints--

يجب أن تُعيد الدالة `get_combinations(2)` القيمة `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(2), 2)`)
}})
```

يجب أن تُعيد الدالة `get_combinations(3)` القيمة `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(3), 5)`)
}})
```

يجب أن تُعيد الدالة `get_combinations(5)` القيمة `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(5), 42)`)
}})
```

يجب أن تُعيد الدالة `get_combinations(8)` القيمة `1430`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(8), 1430)`)
}})
```

يجب أن تُعيد الدالة `get_combinations(13)` القيمة `742900`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(13), 742900)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_combinations(n):

    return n
```

# --solutions--

```py
def get_combinations(n):
    count = 0

    def generate(op, cl):
        nonlocal count
        if op == 0 and cl == 0:
            count += 1
            return
        if op > 0:
            generate(op - 1, cl)
        if cl > op:
            generate(op, cl - 1)

    generate(n, n)
    return count
```
