---
id: 6a0dcd03ee4e68698080ef68
title: "تحدٍّ 304: ترتيبات خط سير الرحلة"
challengeType: 29
dashedName: challenge-304
---

# --description--

معطاة مصفوفة تحتوي على توقفين اختياريين على الأقل لرحلة يومية، أرجع عدد ترتيبات خط السير الصالحة.

يتضمن خط السير دائمًا `"breakfast"` و`"lunch"` و`"dinner"`، ولن تُمرر هذه كمعلمات. يمكن وضع التوقفات الاختيارية في أي مكان ضمن خط السير، مع مراعاة القواعد التالية:

- يكون `"breakfast"` دائمًا في البداية، مع وجود توقف واحد على الأقل قبل `"lunch"`.
- يجب أن يظهر `"lunch"` قبل `"dinner"`، مع وجود توقف واحد على الأقل بينهما.
- لا يجوز أن يظهر أكثر من توقف اختياري واحد بعد `"dinner"`.

أرجع عدد الترتيبات الصالحة.

# --hints--

يجب أن تُرجع الدالة `get_itinerary_count(["library", "park"])` القيمة `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

يجب أن تُرجع الدالة `get_itinerary_count(["library", "park", "arcade"])` القيمة `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

يجب أن تُرجع الدالة `get_itinerary_count(["library", "park", "arcade", "store"])` القيمة `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

يجب أن تُرجع الدالة `get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` القيمة `840`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

يجب أن تُرجع الدالة `get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` القيمة `55440`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
