---
id: 6a0dcc730cb92a616f86f0c5
title: "تحدٍّ 301: آخر تحميل"
challengeType: 29
dashedName: challenge-301
---

# --description--

بالنظر إلى عدد مغارف مسحوق الغسيل المتبقية لديك ومصفوفة بعدد المغارف التي استخدمتها في الأيام السابقة، أعد عدد الأيام الكاملة التي يمكنك استخدام المسحوق فيها.

احسب متوسط استهلاكك اليومي من سجل الاستخدام وافترض هذا المقدار من الاستهلاك لكل يوم مستقبلاً.

# --hints--

`last_load_date(10, [2, 2, 2, 2, 2, 2, 2])` يجب أن تُعيد `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(10, [2, 2, 2, 2, 2, 2, 2]), 5)`)
}})
```

`last_load_date(16, [2, 3, 0, 3, 4, 2, 1])` يجب أن تُعيد `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(16, [2, 3, 0, 3, 4, 2, 1]), 7)`)
}})
```

`last_load_date(33, [5, 0, 4, 3, 3, 2])` يجب أن تُعيد `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(33, [5, 0, 4, 3, 3, 2]), 11)`)
}})
```

`last_load_date(50, [2, 0, 2, 9, 12, 0, 2])` يجب أن تُعيد `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(50, [2, 0, 2, 9, 12, 0, 2]), 12)`)
}})
```

`last_load_date(20, [13, 9, 12, 10, 8])` يجب أن تُعيد `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(20, [13, 9, 12, 10, 8]), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def last_load_date(scoops, usage):

    return scoops
```

# --solutions--

```py
def last_load_date(scoops, usage):
    avg = sum(usage) / len(usage)
    return int(scoops / avg)
```
