---
id: 6a0dcd03ee4e68698080ef6b
title: "تحدٍّ 307: تنظيمات التقسيم"
challengeType: 29
dashedName: challenge-307
---

# --description--

معطى شبكة ثنائية الأبعاد (مصفوفة من مصفوفات) تمثل تخطيط مباني المدينة، أرجع إحداثيات كل المباني التي تنتهك قواعد التقسيم.

كل خلية في الشبكة تحتوي على أحد المسميات من الجدول أدناه. المبنى يكون مخالفًا إذا كان أي من جيرانه (حتى 4 جيران) أفقيًا أو عموديًا من نوع لا يمكن أن يكون مجاورًا له.

| مسمى | النوع | لا يمكن أن يكون مجاورًا لـ |
| - | - | - |
| `"i"` | صناعي | `"R"`، `"I"` |
| `"A"` | زراعي | `"C"` |
| `"R"` | سكني | `"i"`، `"C"` |
| `"I"` | مؤسسي | `"i"` |
| `"C"` | تجاري | `"R"`، `"A"` |
| `""` (سلسلة نصية فارغة) | غير مطور | لا قيود |

أرجع إحداثيات كل الخلايا المخالفة كمصفوفة من أزواج `[row, col]` بأي ترتيب. إذا لم توجد مخالفات، أرجع مصفوفة فارغة.

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` يجب أن تُرجع `[[0, 0], [0, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` يجب أن تُرجع `[[0, 1], [1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` يجب أن تُرجع `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` يجب أن تُرجع `[[0, 1], [0, 2], [0, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` يجب أن تُرجع `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
