---
id: 6a0dcd03ee4e68698080ef66
title: "تحدٍّ 302: تعب السفر بسبب اختلاف التوقيت"
challengeType: 29
dashedName: challenge-302
---

# --description--

معطى مدينة المغادرة، ومدينة الوصول، ومدة الرحلة بالساعات، واتجاه السفر، أرجع عدد ساعات تعب السفر التي يعاني منها المسافر.

المدن المعطاة ستكون من القائمة التالية التي تتضمن إزاحة التوقيت العالمي المنسق (UTC):

| المدينة | الإزاحة |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

لحساب ساعات تعب السفر:

1. احسب فرق التوقيت بالساعات بين المدينتين.
2. حدّد معامل الاتجاه. إذا كان السفر نحو `"east"`، يكون 1.5، وإلا يكون 1.0.
3. احصل على ساعات تعب السفر بالمعادلة: فرق التوقيت + (مدة الرحلة * 0.1) * معامل الاتجاه

أرجع ساعات تعب السفر مقربة إلى منزلة عشرية واحدة.

# --hints--

`get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east")` يجب أن تُرجع `6.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east"), 6.5)`)
}})
```

`get_jet_lag_hours("London", "New York", 8, "west")` يجب أن تُرجع `5.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("London", "New York", 8, "west"), 5.8)`)
}})
```

`get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east")` يجب أن تُرجع `1.6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east"), 1.6)`)
}})
```

`get_jet_lag_hours("Dubai", "London", 7, "west")` يجب أن تُرجع `4.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Dubai", "London", 7, "west"), 4.7)`)
}})
```

`get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west")` يجب أن تُرجع `17.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west"), 17.5)`)
}})
```

`get_jet_lag_hours("Tokyo", "Dubai", 9, "west")` يجب أن تُرجع `5.9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Tokyo", "Dubai", 9, "west"), 5.9)`)
}})
```

`get_jet_lag_hours("New York", "Istanbul", 10, "east")` يجب أن تُرجع `9.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("New York", "Istanbul", 10, "east"), 9.5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):

    return departure_city
```

# --solutions--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):
    offsets = {
        "Los Angeles": -8,
        "New York": -5,
        "London": 0,
        "Istanbul": 3,
        "Dubai": 4,
        "Hong Kong": 8,
        "Tokyo": 9
    }
    timezone_diff = abs(offsets[arrival_city] - offsets[departure_city])
    multiplier = 1.5 if direction == "east" else 1.0
    score = timezone_diff + (flight_duration * 0.1) * multiplier
    return round(score, 1)
```
