---
id: 6a0dcd03ee4e68698080ef6c
title: "تحدٍّ 308: مدقق بطاقة الائتمان"
challengeType: 29
dashedName: challenge-308
---

# --description--

معطاة سلسلة نصية من الأرقام لرقم بطاقة ائتمان، حدّد إذا كان رقم البطاقة صحيحًا باستخدام الطريقة التالية:

- ابدأ من الرقم قبل الأخير، وضاعف كل رقم ثاني متجهًا نحو اليسار.
- إذا أدى الضرب إلى رقم أكبر من 9، اطرح 9.
- اجمع كل الأرقام (المضاعفة وغير المضاعفة).
- إذا كان المجموع قابلاً للقسمة على 10، فالرقم صحيح.

# --hints--

يجب أن تُعيد الدالة `is_valid_card("4532015112830366")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830366"), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_card("5425233430109903")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("5425233430109903"), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_card("371449635398431")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("371449635398431"), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_card("6011111111111117")` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("6011111111111117"), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_card("4532015112830367")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830367"), False)`)
}})
```

يجب أن تُعيد الدالة `is_valid_card("1234567890123456")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("1234567890123456"), False)`)
}})
```

يجب أن تُعيد الدالة `is_valid_card("4532015112830368")` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830368"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_card(number):

    return number
```

# --solutions--

```py
def is_valid_card(number):
    total = 0
    for i, digit in enumerate(reversed(number)):
        d = int(digit)
        if i % 2 == 1:
            d *= 2
            if d > 9:
                d -= 9
        total += d
    return total % 10 == 0
```
