---
id: 6a0dcc730cb92a616f86f0bf
title: "تحدٍّ 295: مدقق المخطط الجزء 1"
challengeType: 29
dashedName: challenge-295
---

# --description--

معطى كائن (JavaScript) أو قاموس (Python)، حدّد ما إذا كان يتطابق مع المخطط التالي:

```json
{
  username: string
}
```

- يُسمح بالمفاتيح الإضافية

# --hints--

يجب أن تُعيد الدالة `is_valid_schema({"username": "bob"})` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bob"}), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_schema({"username": "jen", "posts": 30})` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "jen", "posts": 30}), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_schema({"username": ""})` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": ""}), True)`)
}})
```

يجب أن تُعيد الدالة `is_valid_schema({"username": 7})` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 7}), False)`)
}})
```

يجب أن تُعيد الدالة `is_valid_schema({"posts": 25})` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"posts": 25}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    return isinstance(obj.get("username"), str)
```
