---
id: 6a15cadf5f240d05a2649557
title: "تحدٍّ 313: تكلفة الإيجار"
challengeType: 29
dashedName: challenge-313
---

# --description--

معطى طابع زمني للإيجار، وطابع زمني للإرجاع، ودرجة الإيجار، أرجع التكلفة الإجمالية للإيجار بما في ذلك أي رسوم تأخير.

- الطوابع الزمنية المعطاة هي سلاسل نصية بصيغة UTC ISO، على سبيل المثال: `"2026-06-18T18:30:00Z"`.
- درجة الإيجار هي عدد الأيام قبل موعد إرجاع الإيجار: `1`، `3`، أو `7`.
- يجب إرجاع الإيجارات بحلول الساعة 12:00 ظهرًا بتوقيت UTC أو قبلها في آخر يوم من فترة الإيجار. على سبيل المثال، إيجار ليوم واحد تم استعارته في أي وقت من 15 مارس يجب إرجاعه بحلول الساعة 12:00 ظهرًا بتوقيت UTC في 16 مارس.
- كل يوم يتجاوز موعد الإرجاع والوقت يترتب عليه رسوم تأخير.

التسعير كما يلي:

| الدرجة | التكلفة الأساسية | رسوم التأخير لكل يوم |
| - | - | - |
| يوم واحد | $4.99 | $3.99 |
| 3 أيام | $3.99 | $2.99 |
| 7 أيام | $2.99 | $0.99 |

أرجع التكلفة الإجمالية مقربة إلى منزلتين عشريتين بالشكل `"$D.CC"`.

# --hints--

`get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` يجب أن تُرجع `"$4.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99")`)
}})
```

`get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` يجب أن تُرجع `"$12.97"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97")`)
}})
```

`get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` يجب أن تُرجع `"$3.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99")`)
}})
```

`get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` يجب أن تُرجع `"$9.97"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97")`)
}})
```

`get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` يجب أن تُرجع `"$2.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99")`)
}})
```

`get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` يجب أن تُرجع `"$358.40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rental_cost(rented, returned, tier):

    return rented
```

# --solutions--

```py
from datetime import datetime, timedelta
import math

def get_rental_cost(rented, returned, tier):
    pricing = {
        1: {"base": 4.99, "late": 3.99},
        3: {"base": 3.99, "late": 2.99},
        7: {"base": 2.99, "late": 0.99},
    }
    rented_dt = datetime.fromisoformat(rented.replace("Z", "+00:00"))
    returned_dt = datetime.fromisoformat(returned.replace("Z", "+00:00"))
    due = rented_dt + timedelta(days=tier)
    due = due.replace(hour=12, minute=0, second=0, microsecond=0)
    base = pricing[tier]["base"]
    late = pricing[tier]["late"]
    if returned_dt <= due:
        return f"${base:.2f}"
    days_late = math.ceil((returned_dt - due).total_seconds() / 86400)
    return f"${base + days_late * late:.2f}"
```
