---
id: 6a15cadf5f240d05a264955e
title: "تحدٍّ 320: بنك الدم"
challengeType: 29
dashedName: challenge-320
---

# --description--

معطى مصفوفة لمخزون بنك الدم ومصفوفة لطلبات نوع دم المرضى، أرجع سلسلة نصية بالشكل `"X of Y patients served"`. حيث `X` هو أقصى عدد من المرضى يمكنهم الحصول على الدم من مخزون البنك، و`Y` هو العدد الإجمالي للمرضى.

كل مدخل في المصفوفتين هو أحد أنواع الدم التالية: `"AB"`، `"A"`، `"B"`، أو `"O"`.

قواعد التوافق:

- يمكن لـ `"AB"` أن يستقبل من أي نوع دم.
- يمكن لـ `"A"` أن يستقبل من `"A"` و`"O"`.
- يمكن لـ `"B"` أن يستقبل من `"B"` و`"O"`.
- يمكن لـ `"O"` أن يستقبل فقط من `"O"`.

المدخلات المكررة في المصفوفات تمثل الكمية.

# --hints--

`triage_blood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` يجب أن تُرجع `"4 of 4 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served")`)
}})
```

`triage_blood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` يجب أن تُرجع `"3 of 5 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served")`)
}})
```

`triage_blood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` يجب أن تُرجع `"4 of 5 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served")`)
}})
```

`triage_blood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` يجب أن تُرجع `"4 of 4 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served")`)
}})
```

`triage_blood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` يجب أن تُرجع `"8 of 11 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served")`)
}})
```

`triage_blood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` يجب أن تُرجع `"13 of 14 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_blood(bank, patients):

    return bank
```

# --solutions--

```py
def triage_blood(bank, patients):
    b = {}
    p = {}
    for t in bank:
        b[t] = b.get(t, 0) + 1
    for t in patients:
        p[t] = p.get(t, 0) + 1

    saved = 0

    def serve(patient, donors):
        nonlocal saved
        for donor in donors:
            n = min(p.get(patient, 0), b.get(donor, 0))
            p[patient] = p.get(patient, 0) - n
            b[donor] = b.get(donor, 0) - n
            saved += n

    serve("O", ["O"])
    serve("A", ["A", "O"])
    serve("B", ["B", "O"])
    serve("AB", ["AB", "A", "B", "O"])
    return f"{saved} of {len(patients)} patients served"
```
