---
id: 6a15cadf5f240d05a2649558
title: "تحدٍّ 314: التحليل إلى عوامل أولية"
challengeType: 29
dashedName: challenge-314
---

# --description--

أعطِ عددًا صحيحًا أكبر من 1، أرجع تحليله إلى عوامل أولية كمصفوفة أعداد مرتبة تصاعديًا.

التحليل إلى عوامل أولية هو مجموعة الأعداد الأولية التي تُضرب معًا لإنتاج العدد المعطى. لكل عدد مجموعة واحدة فقط. على سبيل المثال، التحليل إلى عوامل أولية للعدد 20 هو `[2, 2, 5]` لأن 2 * 2 * 5 = 20.

إذا كان العدد المعطى نفسه عددًا أوليًا، أرجعه في مصفوفة تحتوي على عنصر واحد فقط.

# --hints--

`prime_factorization(20)` يجب أن تُرجع `[2, 2, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(20), [2, 2, 5])`)
}})
```

`prime_factorization(17)` يجب أن تُرجع `[17]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(17), [17])`)
}})
```

`prime_factorization(15)` يجب أن تُرجع `[3, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(15), [3, 5])`)
}})
```

`prime_factorization(35)` يجب أن تُرجع `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(35), [5, 7])`)
}})
```

`prime_factorization(999)` يجب أن تُرجع `[3, 3, 3, 37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(999), [3, 3, 3, 37])`)
}})
```

`prime_factorization(360)` يجب أن تُرجع `[2, 2, 2, 3, 3, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(360), [2, 2, 2, 3, 3, 5])`)
}})
```

`prime_factorization(510510)` يجب أن تُرجع `[2, 3, 5, 7, 11, 13, 17]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(510510), [2, 3, 5, 7, 11, 13, 17])`)
}})
```

# --seed--

## --seed-contents--

```py
def prime_factorization(n):

    return n
```

# --solutions--

```py
def prime_factorization(n):
    factors = []
    divisor = 2
    while divisor * divisor <= n:
        while n % divisor == 0:
            factors.append(divisor)
            n //= divisor
        divisor += 1
    if n > 1:
        factors.append(n)
    return factors
```
