---
id: 6a151d32d772271dd2448c2d
title: "تحدٍّ 311: Spellcaster"
challengeType: 29
dashedName: challenge-311
---

# --description--

معطاة سلسلة من رموز التعويذات التي تلقيها، احسب المجموع الكلي للنقاط.

كل حرف في السلسلة يمثل تعويذة:

| Code | Spell| Category | Base Score |
| - | - | - | - |
| `"f"` | نار | التدمير | 3 |
| `"l"` | برق | التدمير | 3 |
| `"i"` | جليد | السيطرة | 2 |
| `"w"` | ريح | السيطرة | 2 |
| `"h"` | شفاء | الاستعادة | 1 |
| `"s"` | درع | الاستعادة | 1 |

يُطبق مضاعف التتابع بناءً على عدد التعويذات المتتالية التي أُلقيت من فئات مختلفة:

- التعويذة الأولى دائمًا تحصل على النقاط الأساسية.
- كل تعويذة متتالية من فئة مختلفة عن السابقة تزيد المضاعف بمقدار 1.
- إلقاء تعويذة من نفس الفئة السابقة يعيد المضاعف إلى 1.
- النقاط لكل تعويذة تساوي نقاطها الأساسية مضروبة في المضاعف الحالي.

أرجع المجموع الكلي للنقاط من متتالية التعويذات.

# --hints--

`cast("fihwl")` يجب أن تُرجع `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` يجب أن تُرجع `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` يجب أن تُرجع `37`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` يجب أن تُرجع `50`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` يجب أن تُرجع `101`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
