---
id: 6a15cadf5f240d05a264955f
title: "تحدٍّ 321: التهجئة الدورية"
challengeType: 29
dashedName: challenge-321
---

# --description--

أعطِ كلمة، حدّد إذا كان يمكن تهجئتها باستخدام رموز العناصر من الجدول الدوري.

- تجاهل حالة الأحرف عند تهجئة كلمة. يمكن تهجئة `"neon"` باستخدام الرموز `"Ne"`، `"O"`، و`"N"`.

إليك قائمة كاملة برموز العناصر:

```json
["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];
```

أرجع مصفوفة بالعناصر المستخدمة لتهجئة الكلمة، بحالتها الأصلية وبالترتيب الذي تهجأ به الكلمة. أو، مصفوفة فارغة إذا لم يمكن تهجئتها.

# --hints--

يجب أن تُرجع الدالة `get_periodic_spelling("neon")` القيمة `["Ne", "O", "N"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("neon"), ["Ne", "O", "N"])`)
}})
```

يجب أن تُرجع الدالة `get_periodic_spelling("rational")` القيمة `["Ra", "Ti", "O", "N", "Al"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("rational"), ["Ra", "Ti", "O", "N", "Al"])`)
}})
```

يجب أن تُرجع الدالة `get_periodic_spelling("yarn")` القيمة `["Y", "Ar", "N"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("yarn"), ["Y", "Ar", "N"])`)
}})
```

يجب أن تُرجع الدالة `get_periodic_spelling("carbon")` القيمة `["C", "Ar", "B", "O", "N"]` أو `["Ca", "Rb", "O", "N"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("carbon"), [["C", "Ar", "B", "O", "N"], ["Ca", "Rb", "O", "N"]])`)
}})
```

يجب أن تُرجع الدالة `get_periodic_spelling("noisy")` القيمة `["N", "O", "I", "S", "Y"]` أو `["No", "I", "S", "Y"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("noisy"), [["N", "O", "I", "S", "Y"], ["No", "I", "S", "Y"]])`)
}})
```

يجب أن تُرجع الدالة `get_periodic_spelling("bicycles")` القيمة `["B", "I", "C", "Y", "Cl", "Es"]` أو `["Bi", "C", "Y", "Cl", "Es"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("bicycles"), [["B", "I", "C", "Y", "Cl", "Es"], ["Bi", "C", "Y", "Cl", "Es"]])`)
}})
```

يجب أن تُرجع الدالة `get_periodic_spelling("optics")` القيمة `["O", "P", "Ti", "C", "S"]`، `["O", "P", "Ti", "Cs"]`، `["O", "Pt", "I", "C", "S"]`، أو `["O", "Pt", "I", "Cs"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("optics"), [["O", "P", "Ti", "C", "S"], ["O", "P", "Ti", "Cs"], ["O", "Pt", "I", "C", "S"], ["O", "Pt", "I", "Cs"]])`)
}})
```

يجب أن تُرجع الدالة `get_periodic_spelling("value")` القيمة `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("value"), [])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_periodic_spelling(word):

    return word
```

# --solutions--

```py
def get_periodic_spelling(word):
    elements = ["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"]
    lower = {e.lower(): e for e in elements}

    def spell(w, i):
        if i == len(w):
            return []
        two = w[i:i + 2]
        one = w[i:i + 1]
        if two in lower:
            rest = spell(w, i + 2)
            if rest is not None:
                return [lower[two]] + rest
        if one in lower:
            rest = spell(w, i + 1)
            if rest is not None:
                return [lower[one]] + rest
        return None

    result = spell(word.lower(), 0)
    return result if result is not None else []
```
