---
id: 6a15cadf5f240d05a264955b
title: "تحدٍّ 317: حاسبة مؤشر كتلة الجسم"
challengeType: 29
dashedName: challenge-317
---

# --description--

بالنظر إلى وزن بوحدة الباوند وطول بوحدة الإنش، أرجع مؤشر كتلة الجسم (BMI) مقربًا إلى منزلة عشرية واحدة.

لحساب مؤشر كتلة الجسم: اقسم الوزن على مربع الطول، ثم اضرب الناتج في 703.

# --hints--

يجب أن تُرجع الدالة `calculate_bmi(180, 70)` القيمة `25.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(180, 70), 25.8)`)
}})
```

يجب أن تُرجع الدالة `calculate_bmi(140, 64)` القيمة `24.0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(140, 64), 24.0)`)
}})
```

يجب أن تُرجع الدالة `calculate_bmi(160, 76)` القيمة `19.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(160, 76), 19.5)`)
}})
```

يجب أن تُرجع الدالة `calculate_bmi(200, 60)` القيمة `39.1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(200, 60), 39.1)`)
}})
```

يجب أن تُرجع الدالة `calculate_bmi(150, 68)` القيمة `22.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(150, 68), 22.8)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_bmi(weight, height):

    return weight
```

# --solutions--

```py
def calculate_bmi(weight, height):
    return round((weight / (height * height)) * 703, 1)
```
