---
id: 6a15cadf5f240d05a264955d
title: "تحدٍّ 319: محلل الـ Frontmatter"
challengeType: 29
dashedName: challenge-319
---

# --description--

أعطِ سلسلة نصية تمثل كتلة frontmatter، حللها وأرجع كائنًا (JavaScript) أو قاموسًا (Python) يحتوي على المفاتيح والقيم.

الـ frontmatter محاط بفواصل `---` ويحتوي على أزواج `key: value` بداخلها، واحد في كل سطر. على سبيل المثال:

```md
---
title: My Post
draft: false
views: 100
---
```

يجب أن تُرجع:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- يجب إرجاع الأعداد والقيم المنطقية والسلاسل النصية كلها بنوعها المناسب.
- السلسلة النصية المعطاة ستحتوي على أسطر جديدة مفصولة بحرف السطر الجديد (`"\n"`). المثال أعلاه سيُعطى كالتالي: `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`.

# --hints--

`parse_frontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` يجب أن تُرجع `{ title: "My Post", draft: False, views: 100 }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\ntitle: My Post\\ndraft: false\\nviews: 100\\n---"), {"title": "My Post", "draft": False, "views": 100})`)
}})
```

`parse_frontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` يجب أن تُرجع `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: False }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nid: 6a174db57256a112f932195c\\ntitle: My Book\\nlocale: en\\nwordCount: 10000\\npublished: false\\n---"), {"id": "6a174db57256a112f932195c", "title": "My Book", "locale": "en", "wordCount": 10000, "published": False})`)
}})
```

`parse_frontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` يجب أن تُرجع `{ version: "1.0.0", url: "https://example.com", private: True }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nversion: 1.0.0\\nurl: https://example.com\\nprivate: true\\n---"), {"version": "1.0.0", "url": "https://example.com", "private": True})`)
}})
```

`parse_frontmatter("---\nrating: 4.5\nprice: 9.99\n---")` يجب أن تُرجع `{ rating: 4.5, price: 9.99 }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nrating: 4.5\\nprice: 9.99\\n---"), {"rating": 4.5, "price": 9.99})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_frontmatter(s):

    return s
```

# --solutions--

```py
def parse_frontmatter(s):
    lines = [line for line in s.split("\n") if line != "---" and line != ""]
    result = {}
    for line in lines:
        key, *rest = line.split(":")
        value = ":".join(rest).strip()
        if value == "true":
            result[key.strip()] = True
        elif value == "false":
            result[key.strip()] = False
        else:
            try:
                n = float(value)
                result[key.strip()] = int(n) if n == int(n) else n
            except ValueError:
                result[key.strip()] = value
    return result
```
