---
id: 6a15cadf5f240d05a264955a
title: "تحدٍّ 316: لغة 1337"
challengeType: 29
dashedName: challenge-316
---

# --description--

بالنظر إلى سلسلة نصية من أحرف صغيرة، أعدها مترجمة إلى لغة leet باستبدال الحروف أدناه بنظائرها في لغة leet:

| الحرف | leet |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- تُترك الحروف التي لا يوجد لها بديل كما هي دون تغيير.

# --hints--

`make_leet("cool")` يجب أن تُعيد `"c001"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("cool"), "c001")`)
}})
```

`make_leet("leet")` يجب أن تُعيد `"1337"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("leet"), "1337")`)
}})
```

`make_leet("hacker")` يجب أن تُعيد `"h4ck3r"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("hacker"), "h4ck3r")`)
}})
```

`make_leet("satellite")` يجب أن تُعيد `"547311173"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("satellite"), "547311173")`)
}})
```

`make_leet("abcdefghijklmnopqrstuvwxyz")` يجب أن تُعيد `"4bcd3f9h1jk1mn0pqr57uvwxyz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz")`)
}})
```

# --seed--

## --seed-contents--

```py
def make_leet(s):

    return s
```

# --solutions--

```py
def make_leet(s):
    leet_map = {"a": "4", "e": "3", "g": "9", "i": "1", "l": "1", "o": "0", "t": "7", "s": "5"}
    return "".join(leet_map.get(c, c) for c in s)
```
