---
id: 6a15cadf5f240d05a2649559
title: "تحدٍّ 315: الانقلاب الصيفي"
challengeType: 29
dashedName: challenge-315
---

# --description--

اليوم هو الانقلاب الصيفي، أطول يوم في السنة في نصف الكرة الشمالي وأقصر يوم في نصف الكرة الجنوبي. عند إعطاء دائرة عرض، أرجع سلسلة نصية تمثّل ساعات النهار والليل.

- ستكون دائرة العرض بين 90 (القُطب الشمالي) و-90 (القُطب الجنوبي)، شاملة
- عدد ساعات النهار = 12 + (دائرة العرض / 90) * 12
- قرب عدد ساعات النهار لأقرب عدد زوجي

أرجع سلسلة نصية مكونة من 24 حرفًا تستخدم `"☀️"` لساعات النهار و`"🌑"` لساعات الليل، حيث:

- كل حرف يمثل ساعة واحدة، تبدأ من منتصف الليل (الساعة 0)
- شروق الشمس وغروبها يقعان بشكل متماثل حول الظهر

على سبيل المثال، دائرة عرض `0` (خط الاستواء) لها 12 ساعة من ضوء النهار، لذا يكون شروق الشمس عند الساعة 6:00 صباحًا وغروبها عند الساعة 6:00 مساءً. أرجع: `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`.

# --hints--

`get_daytime_hours(0)` يجب أن تُرجع `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(0), "🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(90)` يجب أن تُرجع `"☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(90), "☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️")`)
}})
```

`get_daytime_hours(-90)` يجب أن تُرجع `"🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-90), "🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(-33)` يجب أن تُرجع `"🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-33), "🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(66.5)` يجب أن تُرجع `"🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(66.5), "🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑")`)
}})
```

`get_daytime_hours(40)` يجب أن تُرجع `"🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(40), "🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑")`)
}})
```

`get_daytime_hours(-50)` يجب أن تُرجع `"🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-50), "🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_daytime_hours(latitude):

    return latitude
```

# --solutions--

```py
def get_daytime_hours(latitude):
    daylight_hours = round((12 + (latitude / 90) * 12) / 2) * 2
    sunrise = 12 - daylight_hours // 2
    sunset = 12 + daylight_hours // 2
    result = ""
    for hour in range(24):
        result += "☀️" if sunrise <= hour < sunset else "🌑"
    return result
```
