---
id: 6a1d9f98e819ed70a0e994e1
title: "تحدٍّ 335: خمسة نرد"
challengeType: 29
dashedName: challenge-335
---

# --description--

معطاة مصفوفة من خمسة نرد بقيم من 1 إلى 6، أرجع أفضل يد ممكنة.

هذه هي الأيدي مرتبة من الأدنى إلى الأعلى:

| اليد | الوصف |
|------|-------------|
| `"no pair"` | لا يوجد زوج أو أفضل |
| `"pair"` | نردان لهما نفس القيمة |
| `"two pair"` | زوجان مختلفان |
| `"three of a kind"` | ثلاثة نردات لها نفس القيمة |
| `"small straight"` | أربع قيم متتالية |
| `"large straight"` | خمس قيم متتالية |
| `"full house"` | ثلاثة من نفس النوع وزوج |
| `"four of a kind"` | أربعة نردات لها نفس القيمة |
| `"five of a kind"` | جميع النردات الخمسة لها نفس القيمة |

# --hints--

`five_dice([1, 1, 1, 1, 1])` يجب أن تُرجع `"five of a kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 1, 1, 1, 1]), "five of a kind")`)
}})
```

`five_dice([5, 5, 5, 6, 5])` يجب أن تُرجع `"four of a kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([5, 5, 5, 6, 5]), "four of a kind")`)
}})
```

`five_dice([2, 5, 6, 4, 3])` يجب أن تُرجع `"large straight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([2, 5, 6, 4, 3]), "large straight")`)
}})
```

`five_dice([4, 3, 3, 3, 1])` يجب أن تُرجع `"three of a kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([4, 3, 3, 3, 1]), "three of a kind")`)
}})
```

`five_dice([4, 6, 2, 6, 5])` يجب أن تُرجع `"pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([4, 6, 2, 6, 5]), "pair")`)
}})
```

`five_dice([1, 4, 5, 6, 2])` يجب أن تُرجع `"no pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 4, 5, 6, 2]), "no pair")`)
}})
```

`five_dice([1, 3, 4, 6, 2])` يجب أن تُرجع `"small straight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 3, 4, 6, 2]), "small straight")`)
}})
```

`five_dice([2, 2, 5, 2, 5])` يجب أن تُرجع `"full house"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([2, 2, 5, 2, 5]), "full house")`)
}})
```

`five_dice([6, 4, 5, 6, 4])` يجب أن تُرجع `"two pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([6, 4, 5, 6, 4]), "two pair")`)
}})
```

# --seed--

## --seed-contents--

```py
def five_dice(dice):

    return dice
```

# --solutions--

```py
def five_dice(dice):
    from collections import Counter
    counts = Counter(dice)
    values = sorted(counts.values(), reverse=True)
    keys = sorted(counts.keys())

    is_large_straight = len(keys) == 5 and keys[-1] - keys[0] == 4
    is_small_straight = not is_large_straight and any(
        keys[i + 3] - keys[i] == 3 and len(set(keys[i:i + 4])) == 4
        for i in range(len(keys) - 3)
    )

    if values[0] == 5: return "five of a kind"
    if values[0] == 4: return "four of a kind"
    if values[0] == 3 and values[1] == 2: return "full house"
    if is_large_straight: return "large straight"
    if is_small_straight: return "small straight"
    if values[0] == 3: return "three of a kind"
    if values[0] == 2 and values[1] == 2: return "two pair"
    if values[0] == 2: return "pair"
    return "no pair"
```
