---
id: 6a1d9f98e819ed70a0e994e2
title: "تحدٍّ 336: توافق الأبراج"
challengeType: 29
dashedName: challenge-336
---

# --description--

معطاة سلسلتا نصية لبرجيْن، أرجع نسبة توافقهما.

البرج مرتّب في عجلة مكوّنة من 12 موقعًا بهذا الترتيب: `"Aries"`، `"Taurus"`، `"Gemini"`، `"Cancer"`، `"Leo"`، `"Virgo"`، `"Libra"`، `"Scorpio"`، `"Sagittarius"`، `"Capricorn"`، `"Aquarius"`، `"Pisces"`، ثم تعود إلى `"Aries"` بعد `"Pisces"`. أوجد أقصر مسافة بين البرجيْن وأرجع نسبة التوافق:

| المسافة | التوافق |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscope_match("Libra", "Sagittarius")` يجب أن تُرجع `"80%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Libra", "Sagittarius"), "80%")`)
}})
```

`horoscope_match("Gemini", "Scorpio")` يجب أن تُرجع `"20%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Gemini", "Scorpio"), "20%")`)
}})
```

`horoscope_match("Pisces", "Aries")` يجب أن تُرجع `"40%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Pisces", "Aries"), "40%")`)
}})
```

`horoscope_match("Capricorn", "Cancer")` يجب أن تُرجع `"50%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Capricorn", "Cancer"), "50%")`)
}})
```

`horoscope_match("Aquarius", "Aquarius")` يجب أن تُرجع `"100%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Aquarius", "Aquarius"), "100%")`)
}})
```

`horoscope_match("Virgo", "Taurus")` يجب أن تُرجع `"90%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Virgo", "Taurus"), "90%")`)
}})
```

`horoscope_match("Leo", "Scorpio")` يجب أن تُرجع `"30%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Leo", "Scorpio"), "30%")`)
}})
```

# --seed--

## --seed-contents--

```py
def horoscope_match(sign1, sign2):

    return sign1
```

# --solutions--

```py
def horoscope_match(sign1, sign2):
    signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"]
    a = signs.index(sign1)
    b = signs.index(sign2)
    diff = abs(a - b)
    distance = min(diff, 12 - diff)
    compatibility = {0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%"}
    return compatibility[distance]
```
