---
id: 6a1d9f98e819ed70a0e994da
title: "تحدٍّ 328: روتين كابريكار"
challengeType: 29
dashedName: challenge-328
---

# --description--

بالنظر إلى رقم مكون من 4 أرقام، أرجع عدد المرات التي تحتاج فيها لتطبيق روتين كابريكار حتى تصل إلى 6174.

يعمل روتين كابريكار كما يلي:

- رتب الأرقام ترتيبًا تنازليًا لتكوين أكبر رقم
- رتب الأرقام ترتيبًا تصاعديًا لتكوين أصغر رقم (أضف أصفارًا بادئة إذا لزم الأمر)
- اطرح الرقم الأصغر من الأكبر
- كرر العملية مع الرقم الجديد

# --hints--

يجب أن تُرجع `kaprekar(1234)` القيمة `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(1234), 3)`)
}})
```

يجب أن تُرجع `kaprekar(2025)` القيمة `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(2025), 6)`)
}})
```

يجب أن تُرجع `kaprekar(7173)` القيمة `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(7173), 4)`)
}})
```

يجب أن تُرجع `kaprekar(3164)` القيمة `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(3164), 7)`)
}})
```

يجب أن تُرجع `kaprekar(8082)` القيمة `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(8082), 2)`)
}})
```

# --seed--

## --seed-contents--

```py
def kaprekar(n):

    return n
```

# --solutions--

```py
def kaprekar(n):
    steps = 0
    current = n
    while current != 6174:
        digits = list(str(current).zfill(4))
        desc = int("".join(sorted(digits, reverse=True)))
        asc = int("".join(sorted(digits)))
        current = desc - asc
        steps += 1
    return steps
```
