---
id: 6a1d9f98e819ed70a0e994db
title: "تحدٍّ 329: ملء الدلو"
challengeType: 29
dashedName: challenge-329
---

# --description--

مع شبكة ثنائية الأبعاد، وموقع بداية (`[row, col]`)، وقيمة جديدة، استبدل القيمة في موقع البداية وكل الخلايا المتصلة التي تحمل نفس القيمة بالقيمة الجديدة.

- الخلايا متصلة إذا كانت متجاورة أفقيًا أو عموديًا (وليس قطريًا).

أرجع الشبكة المحدثة.

# --hints--

`bucket_fill([["R", "G"], ["R", "G"]], [0, 1], "B")` يجب أن تُرجع `[["R", "B"], ["R", "B"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["R", "G"], ["R", "G"]], [0, 1], "B"), [["R", "B"], ["R", "B"]])`)
}})
```

`bucket_fill([["Y", "G", "G"], ["Y", "Y", "Y"], ["B", "Y", "R"]], [1, 2], "B")` يجب أن تُرجع `[["B", "G", "G"], ["B", "B", "B"], ["B", "B", "R"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["Y", "G", "G"], ["Y", "Y", "Y"], ["B", "Y", "R"]], [1, 2], "B"), [["B", "G", "G"], ["B", "B", "B"], ["B", "B", "R"]])`)
}})
```

`bucket_fill([["O", "O", "P"], ["P", "O", "O"], ["P", "P", "O"]], [2, 0], "R")` يجب أن تُرجع `[["O", "O", "P"], ["R", "O", "O"], ["R", "R", "O"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["O", "O", "P"], ["P", "O", "O"], ["P", "P", "O"]], [2, 0], "R"), [["O", "O", "P"], ["R", "O", "O"], ["R", "R", "O"]])`)
}})
```

`bucket_fill([["T", "T", "R", "T"], ["R", "T", "R", "T"], ["R", "T", "R", "T"], ["T", "T", "T", "T"]], [0, 3], "Y")` يجب أن تُرجع `[["Y", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["Y", "Y", "Y", "Y"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["T", "T", "R", "T"], ["R", "T", "R", "T"], ["R", "T", "R", "T"], ["T", "T", "T", "T"]], [0, 3], "Y"), [["Y", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["Y", "Y", "Y", "Y"]])`)
}})
```

`bucket_fill([["G", "B", "G", "B"], ["R", "B", "B", "G"], ["B", "G", "B", "R"], ["B", "G", "G", "B"]], [2, 2], "G")` يجب أن تُرجع `[["G", "G", "G", "B"], ["R", "G", "G", "G"], ["B", "G", "G", "R"], ["B", "G", "G", "B"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "B", "G", "B"], ["R", "B", "B", "G"], ["B", "G", "B", "R"], ["B", "G", "G", "B"]], [2, 2], "G"), [["G", "G", "G", "B"], ["R", "G", "G", "G"], ["B", "G", "G", "R"], ["B", "G", "G", "B"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def bucket_fill(grid, pos, new_value):

    return grid
```

# --solutions--

```py
def bucket_fill(grid, pos, new_value):
    row, col = pos
    target = grid[row][col]
    if target == new_value:
        return grid

    def fill(r, c):
        if r < 0 or r >= len(grid):
            return
        if c < 0 or c >= len(grid[0]):
            return
        if grid[r][c] != target:
            return
        grid[r][c] = new_value
        fill(r - 1, c)
        fill(r + 1, c)
        fill(r, c - 1)
        fill(r, c + 1)

    fill(row, col)
    return grid
```
