---
id: 6a1d9f98e819ed70a0e994de
title: "تحدٍّ 332: فرز القضايا"
challengeType: 29
dashedName: challenge-332
---

# --description--

معطى عدد من المللي ثانية منذ آخر منشور على قضية، وآخر رسالة نُشرت على القضية، حدّد ما يجب فعله مع القضية وفقًا لهذه القواعد:

- إذا كانت آخر رسالة منذ أقل من 7 أيام، أرجع `"leave it"`
- إذا كانت آخر رسالة منذ 7 أيام أو أكثر ومحتواها يحتوي على `"bump"` (بغض النظر عن حالة الأحرف)، أرجع `"close it"`
- خلاف ذلك، أرجع `"bump it"`

# --hints--

`triage_issue(86400000, "Lets fix it")` يجب أن تُرجع `"leave it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(86400000, "Lets fix it"), "leave it")`)
}})
```

`triage_issue(1209600000, "still waiting")` يجب أن تُرجع `"bump it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(1209600000, "still waiting"), "bump it")`)
}})
```

`triage_issue(864000000, "bump")` يجب أن تُرجع `"close it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(864000000, "bump"), "close it")`)
}})
```

`triage_issue(604800000, "Do we still want this?")` يجب أن تُرجع `"bump it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Do we still want this?"), "bump it")`)
}})
```

`triage_issue(604800000, "Bumping this")` يجب أن تُرجع `"close it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Bumping this"), "close it")`)
}})
```

`triage_issue(345600000, "I'll make a PR")` يجب أن تُرجع `"leave it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(345600000, "I'll make a PR"), "leave it")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(ms, message):

    return ms
```

# --solutions--

```py
def triage_issue(ms, message):
    seven_days = 7 * 24 * 60 * 60 * 1000
    if ms < seven_days:
        return "leave it"
    if "bump" in message.lower():
        return "close it"
    return "bump it"
```
