---
id: 6a1d9f98e819ed70a0e994dd
title: "تحدٍّ 331: أقرب مضاعف"
challengeType: 29
dashedName: challenge-331
---

# --description--

بالنظر إلى عددين صحيحين، قرب الأول إلى أقرب مضاعف للثاني.

# --hints--

`round_to_nearest_multiple(5, 3)` يجب أن تُرجع `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(5, 3), 6)`)
}})
```

`round_to_nearest_multiple(17, 4)` يجب أن تُرجع `16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(17, 4), 16)`)
}})
```

`round_to_nearest_multiple(43, 5)` يجب أن تُرجع `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(43, 5), 45)`)
}})
```

`round_to_nearest_multiple(38, 11)` يجب أن تُرجع `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(38, 11), 33)`)
}})
```

`round_to_nearest_multiple(93, 12)` يجب أن تُرجع `96`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(93, 12), 96)`)
}})
```

# --seed--

## --seed-contents--

```py
def round_to_nearest_multiple(num, multiple):

    return num
```

# --solutions--

```py
def round_to_nearest_multiple(num, multiple):
    return round(num / multiple) * multiple
```
