---
id: 6a1d9f98e819ed70a0e994d9
title: "تحدٍّ 327: ترحيل قاعدة بيانات"
challengeType: 29
dashedName: challenge-327
---

# --description--

بالنظر إلى كائنين من قاعدة بيانات، أعد الكائن الثاني مع ملء أي خاصيات مفقودة من الأول.

- لا يجب استبدال الحقول الموجودة بالفعل في السجل.

# --hints--

`migrate_record({ "username": "", "posts": 0 }, { "verified": True })` يجب أن تُعيد `{ "username": "", "posts": 0, "verified": True }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "verified": True }), { "username": "", "posts": 0, "verified": True })`)
}})
```

`migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "posts": 5 })` يجب أن تُعيد `{ "username": "camper", "posts": 5 }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "posts": 5 }), { "username": "camper", "posts": 5 })`)
}})
```

`migrate_record({ "username": "", "posts": 0, "verified": False }, { "username": "camper" })` يجب أن تُعيد `{ "username": "camper", "posts": 0, "verified": False }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0, "verified": False }, { "username": "camper" }), { "username": "camper", "posts": 0, "verified": False })`)
}})
```

`migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "role": "admin" })` يجب أن تُعيد `{ "username": "camper", "role": "admin", "posts": 0 }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "role": "admin" }), { "username": "camper", "role": "admin", "posts": 0 })`)
}})
```

`migrate_record({ "username": "", "email": "", "posts": 0, "verified": False, "role": "user", "banned": False }, { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin" })` يجب أن تُعيد `{ "username": "camper", "email": "camper@freecodecamp.org", "role": "admin", "posts": 0, "verified": False, "banned": False }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "email": "", "posts": 0, "verified": False, "role": "user", "banned": False }, { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin" }), { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin", "posts": 0, "verified": False, "banned": False })`)
}})
```

# --seed--

## --seed-contents--

```py
def migrate_record(schema, record):

    return schema
```

# --solutions--

```py
def migrate_record(schema, record):
    return {**schema, **record}
```
