---
id: 6a1d9f98e819ed70a0e994df
title: "تحدٍّ 333: تصنيف القضايا 2"
challengeType: 29
dashedName: challenge-333
---

# --description--

بالنظر إلى عنوان قضية ومصفوفة من المسميات الحالية، أرجع مصفوفة محدثة من المسميات بناءً على القواعد التالية:

إذا لم يكن للقضية أي مسميات، أضف:

- `"bug"` و`"needs triage"` إذا احتوى العنوان على `"error"` أو `"bug"`
- `"enhancement"` و`"discussing"` إذا احتوى العنوان على `"feature"` أو `"add"`

أما إذا كانت المسميات المعطاة تحتوي على:

- `"needs triage"` وكان العنوان يحتوي على `"simple"` أو `"easy"`، احذف `"needs triage"` وأضف `"good first issue"`
- `"discussing"` وكان العنوان يحتوي على `"planned"` أو `"next"`، احذف `"discussing"` وأضف `"on the roadmap"`
- وإلا، إذا كان `"needs triage"` أو `"discussing"` موجودًا، احذفه وأضف `"help wanted"`

إذا احتوى العنوان على:

- `"security"`، أضف مسمى `"critical"`

# --hints--

`triage_issue("app crashes with error", [])` يجب أن تُرجع `["bug", "needs triage"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", [])), sorted(["bug", "needs triage"]))`)
}})
```

`triage_issue("app crashes with error", ["bug", "needs triage"])` يجب أن تُرجع `["bug", "help wanted"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", ["bug", "needs triage"])), sorted(["bug", "help wanted"]))`)
}})
```

`triage_issue("add dark mode", [])` يجب أن تُرجع `["enhancement", "discussing"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", [])), sorted(["enhancement", "discussing"]))`)
}})
```

`triage_issue("add dark mode", ["enhancement", "discussing"])` يجب أن تُرجع `["enhancement", "help wanted"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted"]))`)
}})
```

`triage_issue("xss security bug", [])` يجب أن تُرجع `["bug", "needs triage", "critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("xss security bug", [])), sorted(["bug", "needs triage", "critical"]))`)
}})
```

`triage_issue("security vulnerability in auth", [])` يجب أن تُرجع `["critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("security vulnerability in auth", [])), sorted(["critical"]))`)
}})
```

`triage_issue("easy a11y fix", ["bug", "needs triage"])` يجب أن تُرجع `["bug", "good first issue"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("easy a11y fix", ["bug", "needs triage"])), sorted(["bug", "good first issue"]))`)
}})
```

`triage_issue("planned api migration", ["enhancement", "discussing"])` يجب أن تُرجع `["enhancement", "on the roadmap"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("planned api migration", ["enhancement", "discussing"])), sorted(["enhancement", "on the roadmap"]))`)
}})
```

`triage_issue("improve security", ["enhancement", "discussing"])` يجب أن تُرجع `["enhancement", "help wanted", "critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("improve security", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted", "critical"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(title, labels):

    return title
```

# --solutions--

```py
def triage_issue(title, labels):
    result = list(labels)
    t = title.lower()

    def has(label):
        return label in result

    def add(label):
        if not has(label):
            result.append(label)

    def remove(label):
        if label in result:
            result.remove(label)

    if len(result) == 0:
        if "error" in t or "bug" in t:
            add("bug")
            add("needs triage")
        elif "feature" in t or "add" in t:
            add("enhancement")
            add("discussing")
    else:
        if has("needs triage") and ("simple" in t or "easy" in t):
            remove("needs triage")
            add("good first issue")
        elif has("discussing") and ("planned" in t or "next" in t):
            remove("discussing")
            add("on the roadmap")
        elif has("needs triage") or has("discussing"):
            remove("needs triage")
            remove("discussing")
            add("help wanted")

    if "security" in t:
        add("critical")

    return result
```
