---
id: 6a1d9f98e819ed70a0e994e0
title: "تحدٍّ 334: التغيير الدقيق"
challengeType: 29
dashedName: challenge-334
---

# --description--

معطى مبلغ صحيح بالسنتات، أعد عدد الطرق المختلفة لصنع التغيير الدقيق باستخدام البنسات (1 سنت)، النيكل (5 سنتات)، الدايم (10 سنتات)، والكوارتير (25 سنتًا).

# --hints--

`exact_change(3)` يجب أن تُعيد `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(3), 1)`)
}})
```

`exact_change(9)` يجب أن تُعيد `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(9), 2)`)
}})
```

`exact_change(17)` يجب أن تُعيد `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(17), 6)`)
}})
```

`exact_change(39)` يجب أن تُعيد `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(39), 24)`)
}})
```

`exact_change(61)` يجب أن تُعيد `73`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(61), 73)`)
}})
```

`exact_change(99)` يجب أن تُعيد `213`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(99), 213)`)
}})
```

# --seed--

## --seed-contents--

```py
def exact_change(amount):

    return amount
```

# --solutions--

```py
def exact_change(amount):
    coins = [1, 5, 10, 25]
    dp = [0] * (amount + 1)
    dp[0] = 1
    for coin in coins:
        for i in range(coin, amount + 1):
            dp[i] += dp[i - coin]
    return dp[amount]
```
