---
id: 6a2037a68a0bc2aef0006003
title: "تحدٍّ 341: العد التنازلي لعيد الميلاد"
challengeType: 29
dashedName: challenge-341
---

# --description--

بما أن لديك تاريخ اليوم وتاريخ عيد ميلاد، أرجع عدد الأيام حتى عيد ميلاد الشخص التالي.

- يُعطى تاريخ اليوم كسلسلة نصية بصيغة `"YYYY-MM-DD"` مع أصفار بادئة، مثل: `"2026-07-16"`.
- يُعطى عيد الميلاد كسلسلة نصية بصيغة `"M/D"` بدون أصفار بادئة، مثل: `"9/7"`.
- إذا كان اليوم هو عيد ميلادهم، أرجع عدد الأيام حتى عيد ميلادهم التالي (ليس `0`).
- يجب احتساب سنوات الكبيسة.

# --hints--

`days_until_birthday("2026-07-16", "9/7")` يجب أن تُرجع `53`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "9/7"), 53)`)
}})
```

`days_until_birthday("2026-07-16", "3/22")` يجب أن تُرجع `249`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "3/22"), 249)`)
}})
```

`days_until_birthday("2026-07-16", "7/16")` يجب أن تُرجع `365`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "7/16"), 365)`)
}})
```

`days_until_birthday("2024-02-28", "3/1")` يجب أن تُرجع `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2024-02-28", "3/1"), 2)`)
}})
```

`days_until_birthday("2023-04-24", "12/30")` يجب أن تُرجع `250`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2023-04-24", "12/30"), 250)`)
}})
```

`days_until_birthday("2024-03-01", "2/29")` يجب أن تُرجع `1460`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2024-03-01", "2/29"), 1460)`)
}})
```

`days_until_birthday("2096-03-01", "2/29")` يجب أن تُرجع `2920`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2096-03-01", "2/29"), 2920)`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_birthday(today, birthday):

    return today
```

# --solutions--

```py
def days_until_birthday(today, birthday):
    from datetime import date

    year, month, day = [int(x) for x in today.split("-")]
    b_month, b_day = [int(x) for x in birthday.split("/")]

    today_date = date(year, month, day)

    def is_leap_year(y):
        return (y % 4 == 0 and y % 100 != 0) or y % 400 == 0

    def next_valid_birthday(from_year):
        for y in range(from_year, from_year + 9):
            if b_month == 2 and b_day == 29 and not is_leap_year(y):
                continue
            d = date(y, b_month, b_day)
            if d > today_date:
                return d

    return (next_valid_birthday(year) - today_date).days
```
