---
id: 6a2037a68a0bc2aef0006002
title: "تحدٍّ 340: محول Pig Latin"
challengeType: 29
dashedName: challenge-340
---

# --description--

معطاة سلسلة نصية، حوّلها إلى Pig Latin باستخدام القواعد التالية:

- إذا بدأت الكلمة بحرف صوتي (`"a"`، `"e"`، `"i"`، `"o"`، أو `"u"`)، أضف `"way"` في النهاية. على سبيل المثال، تتحول كلمة `"universe"` إلى `"universeway"`.
- إذا بدأت الكلمة بحرف أو أكثر من الحروف الساكنة، انقلها إلى النهاية وأضف `"ay"`. على سبيل المثال، تتحول كلمة `"hello"` إلى `"ellohay"`.
- حافظ على حالة الحرف الأول. على سبيل المثال، تتحول كلمة `"Hello"` إلى `"Ellohay"`.

# --hints--

`pig_latin("universe")` يجب أن تُرجع `"universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("universe"), "universeway")`)
}})
```

`pig_latin("hello")` يجب أن تُرجع `"ellohay"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello"), "ellohay")`)
}})
```

`pig_latin("hello universe")` يجب أن تُرجع `"ellohay universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello universe"), "ellohay universeway")`)
}})
```

`pig_latin("Hello universe")` يجب أن تُرجع `"Ellohay universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Hello universe"), "Ellohay universeway")`)
}})
```

`pig_latin("Pig Latin is fun")` يجب أن تُرجع `"Igpay Atinlay isway unfay"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Pig Latin is fun"), "Igpay Atinlay isway unfay")`)
}})
```

`pig_latin("The quick brown fox jumped over the lazy dog")` يجب أن تُرجع `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday")`)
}})
```

# --seed--

## --seed-contents--

```py
def pig_latin(s):

    return s
```

# --solutions--

```py
def pig_latin(s):
    vowels = "aeiou"

    def convert_word(word):
        lower_word = word.lower()
        if lower_word[0] in vowels:
            return word + "way"
        i = 0
        while i < len(word) and lower_word[i] not in vowels:
            i += 1
        consonants = word[:i]
        rest = word[i:]
        result = rest + consonants + "ay"
        if word[0].isupper():
            return result[0].upper() + result[1:].lower()
        return result

    return " ".join(convert_word(word) for word in s.split(" "))
```
