---
id: 6a2037a68a0bc2aef0006007
title: "تحدٍّ 345: مزج الكلمات"
challengeType: 29
dashedName: challenge-345
---

# --description--

بالنظر إلى كلمتين، أعد كلمة جديدة بدمج النصف الأول من الكلمة الأولى مع النصف الثاني من الكلمة الثانية.

- بالنسبة للكلمات ذات الطول الفردي، يكون النصف الأول هو النصف الأقصر.

# --hints--

`blend_words("turtle", "toucan")` يجب أن تُعيد `"turcan"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("turtle", "toucan"), "turcan")`)
}})
```

`blend_words("chipmunk", "flamingo")` يجب أن تُعيد `"chipingo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("chipmunk", "flamingo"), "chipingo")`)
}})
```

`blend_words("falcon", "pelican")` يجب أن تُعيد `"falican"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("falcon", "pelican"), "falican")`)
}})
```

`blend_words("hyena", "iguana")` يجب أن تُعيد `"hyana"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("hyena", "iguana"), "hyana")`)
}})
```

`blend_words("scorpion", "gorilla")` يجب أن تُعيد `"scorilla"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("scorpion", "gorilla"), "scorilla")`)
}})
```

`blend_words("platypus", "wolverine")` يجب أن تُعيد `"platerine"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("platypus", "wolverine"), "platerine")`)
}})
```

# --seed--

## --seed-contents--

```py
def blend_words(word1, word2):

    return word1
```

# --solutions--

```py
def blend_words(word1, word2):
    first_half = word1[:len(word1) // 2]
    second_half = word2[len(word2) // 2:]
    return first_half + second_half
```
