---
id: 6a2037a68a0bc2aef0005fff
title: "تحدٍّ 337: عداد التعداد"
challengeType: 29
dashedName: challenge-337
---

# --description--

بالنظر إلى سلسلة نصية من علامات التعداد، أعد العدد الإجمالي المُمثل.

- كل أنبوب `"|"` يُمثل عدداً واحداً.
- كل علامة خامسة تُمثل بشرطة مائلة للأمام `"/"`، لتكمل مجموعة من خمسة (`"||||/"`).
- تُفصل المجموعات بمسافة.

# --hints--

`get_tally_count("||||")` يجب أن تُعيد `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||"), 4)`)
}})
```

`get_tally_count("||||/")` يجب أن تُعيد `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/"), 5)`)
}})
```

`get_tally_count("||||/ |||")` يجب أن تُعيد `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ |||"), 8)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||")` يجب أن تُعيد `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||"), 17)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` يجب أن تُعيد `41`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_tally_count(s):

    return s
```

# --solutions--

```py
def get_tally_count(s):
    return len(s.replace(" ", ""))
```
