---
id: 6a2037a68a0bc2aef0006006
title: "تحدٍّ 344: النسبة الذهبية"
challengeType: 29
dashedName: challenge-344
---

# --description--

بالنظر إلى رقمين، حدّد إذا كانت نسبتُهما تقارب النسبة الذهبية.

- استخدم النسبة الذهبية `1.618`
- اقبل هامش خطأ `0.01`

# --hints--

يجب أن تُرجع الدالة `is_golden_ratio(21, 34)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(21, 34), True)`)
}})
```

يجب أن تُرجع الدالة `is_golden_ratio(15, 20)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(15, 20), False)`)
}})
```

يجب أن تُرجع الدالة `is_golden_ratio(8, 13)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(8, 13), True)`)
}})
```

يجب أن تُرجع الدالة `is_golden_ratio(10, 16)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(10, 16), False)`)
}})
```

يجب أن تُرجع الدالة `is_golden_ratio(1618, 1000)` القيمة `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(1618, 1000), True)`)
}})
```

يجب أن تُرجع الدالة `is_golden_ratio(88, 55)` القيمة `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(88, 55), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_golden_ratio(a, b):

    return a
```

# --solutions--

```py
def is_golden_ratio(a, b):
    golden_ratio = 1.618
    ratio = max(a, b) / min(a, b)
    return abs(ratio - golden_ratio) <= 0.01
```
