---
id: 6a2037a68a0bc2aef0006000
title: "تحدٍّ 338: حاسبة عمر الحيوانات الأليفة"
challengeType: 29
dashedName: challenge-338
---

# --description--

بما أن لديك نوع الحيوان الأليف وعمره بالسنوات البشرية، أرجع العمر المكافئ له بسنوات الحيوان الأليف باستخدام جدول التحويل التالي:

| الحيوان الأليف | المضاعف |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

يجب أن تُرجع الدالة `pet_years("dog", 5)` القيمة `35`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("dog", 5), 35)`)
}})
```

يجب أن تُرجع الدالة `pet_years("cat", 9)` القيمة `54`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("cat", 9), 54)`)
}})
```

يجب أن تُرجع الدالة `pet_years("rabbit", 3)` القيمة `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("rabbit", 3), 24)`)
}})
```

يجب أن تُرجع الدالة `pet_years("hamster", 4)` القيمة `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("hamster", 4), 120)`)
}})
```

يجب أن تُرجع الدالة `pet_years("guinea pig", 5)` القيمة `60`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("guinea pig", 5), 60)`)
}})
```

يجب أن تُرجع الدالة `pet_years("goldfish", 2)` القيمة `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("goldfish", 2), 12)`)
}})
```

يجب أن تُرجع الدالة `pet_years("bird", 1)` القيمة `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("bird", 1), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def pet_years(pet, age):

    return pet
```

# --solutions--

```py
def pet_years(pet, age):
    table = {
        "dog": 7,
        "cat": 6,
        "rabbit": 8,
        "hamster": 30,
        "guinea pig": 12,
        "goldfish": 6,
        "bird": 5,
    }
    return age * table[pet]
```
