---
id: 6a22d77ddf034bc4e35b1d5a
title: "تحدٍّ 353: تصنيف التباين 2"
challengeType: 29
dashedName: challenge-353
---

# --description--

معطى قيمتان للسطوع النسبي وقيمة منطقية تشير إلى ما إذا كان النص كبيرًا، أرجع تصنيف التباين وفقًا لمعايير WCAG باستخدام الطريقة التالية:

احسب نسبة التباين بإضافة 0.05 إلى كل قيمة سطوع، ثم اقسم القيمة الأعلى على الأدنى. ستكون القيمة الأعلى دائمًا هي المعلمة الأولى.

أرجع التصنيف بناءً على نسبة التباين باستخدام الجدول التالي:

| التصنيف | النص العادي | النص الكبير |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | أقل من 4.5 | أقل من 3.0 |

# --hints--

`get_contrast_rating(1.0, 0.0, False)` يجب أن تُرجع `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(1.0, 0.0, False), "AAA")`)
}})
```

`get_contrast_rating(0.9015, 0.1364, False)` يجب أن تُرجع `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.9015, 0.1364, False), "AA")`)
}})
```

`get_contrast_rating(0.8965, 0.1628, False)` يجب أن تُرجع `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.8965, 0.1628, False), "Fail")`)
}})
```

`get_contrast_rating(0.7469, 0.0957, True)` يجب أن تُرجع `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7469, 0.0957, True), "AAA")`)
}})
```

`get_contrast_rating(0.7489, 0.2018, True)` يجب أن تُرجع `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7489, 0.2018, True), "AA")`)
}})
```

`get_contrast_rating(0.6571, 0.1974, True)` يجب أن تُرجع `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.6571, 0.1974, True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(l1, l2, is_large_text):

    return l1
```

# --solutions--

```py
def get_contrast_rating(l1, l2, is_large_text):
    ratio = (l1 + 0.05) / (l2 + 0.05)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
