---
id: 6a22d77ddf034bc4e35b1d56
title: "تحدٍّ 349: إشارة الخلية"
challengeType: 29
dashedName: challenge-349
---

# --description--

معطى شبكة تحتوي على قراءات من ثلاثة أبراج خلوية، حدّد موقع الهاتف.

- كل خلية في الشبكة إما `0` (لا يوجد برج) أو عدد صحيح موجب يمثل عدد الخلايا إلى الهاتف، مقاسًا في خط مستقيم: أفقي، عمودي، أو قطري.
- أعد `[row, col]` للخلية التي تبعد العدد الصحيح من الخلايا عن جميع الأبراج الثلاثة.
- يوجد دائمًا حل واحد فقط.

# --hints--

يجب أن تُعيد الدالة `find_signal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` القيمة `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2])`)
}})
```

يجب أن تُعيد الدالة `find_signal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` القيمة `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1])`)
}})
```

يجب أن تُعيد الدالة `find_signal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` القيمة `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2])`)
}})
```

يجب أن تُعيد الدالة `find_signal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` القيمة `[3, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4])`)
}})
```

يجب أن تُعيد الدالة `find_signal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` القيمة `[3, 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_signal(grid):

    return grid
```

# --solutions--

```py
def find_signal(grid):
    rows = len(grid)
    cols = len(grid[0])

    towers = []
    for r in range(rows):
        for c in range(cols):
            if grid[r][c] != 0:
                towers.append((r, c, grid[r][c]))

    def is_valid(r, c, tr, tc, dist):
        dr = abs(r - tr)
        dc = abs(c - tc)
        return (dr == dist or dc == dist) and (dr == 0 or dc == 0 or dr == dc)

    for r in range(rows):
        for c in range(cols):
            if grid[r][c] != 0:
                continue
            if all(is_valid(r, c, tr, tc, dist) for tr, tc, dist in towers):
                return [r, c]
```
