---
id: 6a22d77ddf034bc4e35b1d5b
title: "تحدٍّ 354: تصنيف نسبة التباين 3"
challengeType: 29
dashedName: challenge-354
---

# --description--

معطى مصفوفتان تمثلان قيم RGB وقيمة منطقية تشير إلى ما إذا كان النص كبيرًا، أرجع تصنيف نسبة التباين وفقًا لمعايير WCAG باستخدام الطريقة التالية:

أولًا، حوّل كل قيمة RGB إلى اللمعان النسبي:

- اقسم كل قناة `[R, G, B]` على 255 لتحصل على قيمة بين `0` و `1`
- طبق صيغة تصحيح جاما على كل قناة:
  - إذا كانت قيمة القناة أقل من أو تساوي `0.04045`: `channel / 12.92`
  - وإلا: `((channel + 0.055) / 1.055) ^ 2.4`
- احسب اللمعان: `0.2126 * R + 0.7152 * G + 0.0722 * B`

بعدها، احسب نسبة التباين بإضافة 0.05 إلى كل قيمة لمعان، ثم اقسم القيمة الأعلى على الأدنى. القيمة الأعلى ستكون دائمًا المعلمة الأولى.

أرجع التصنيف بناءً على نسبة التباين باستخدام الجدول التالي:

| التصنيف | النص العادي | النص الكبير |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | أقل من 4.5 | أقل من 3.0 |

# --hints--

يجب أن تُرجع الدالة `get_contrast_rating([255, 255, 255], [0, 0, 0], False)` القيمة `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([255, 255, 255], [0, 0, 0], False), "AAA")`)
}})
```

يجب أن تُرجع الدالة `get_contrast_rating([215, 188, 188], [55, 55, 55], False)` القيمة `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([215, 188, 188], [55, 55, 55], False), "AA")`)
}})
```

يجب أن تُرجع الدالة `get_contrast_rating([143, 144, 210], [46, 47, 61], False)` القيمة `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([143, 144, 210], [46, 47, 61], False), "Fail")`)
}})
```

يجب أن تُرجع الدالة `get_contrast_rating([167, 167, 210], [53, 10, 53], True)` القيمة `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([167, 167, 210], [53, 10, 53], True), "AAA")`)
}})
```

يجب أن تُرجع الدالة `get_contrast_rating([135, 147, 155], [60, 70, 90], True)` القيمة `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([135, 147, 155], [60, 70, 90], True), "AA")`)
}})
```

يجب أن تُرجع الدالة `get_contrast_rating([125, 210, 195], [105, 130, 90], True)` القيمة `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([125, 210, 195], [105, 130, 90], True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):

    return rgb1
```

# --solutions--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):
    def to_luminance(rgb):
        channels = []
        for c in rgb:
            c = c / 255
            if c <= 0.04045:
                channels.append(c / 12.92)
            else:
                channels.append(((c + 0.055) / 1.055) ** 2.4)
        return 0.2126 * channels[0] + 0.7152 * channels[1] + 0.0722 * channels[2]

    l1 = to_luminance(rgb1)
    l2 = to_luminance(rgb2)
    ratio = (l1 + 0.05) / (l2 + 0.05)

    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
