---
id: 6a22d77ddf034bc4e35b1d57
title: "تحدٍّ 350: مسافة الحروف"
challengeType: 29
dashedName: challenge-350
---

# --description--

بالنظر إلى سلسلتين نصيتين متساويتين في الطول، أرجع مجموع أقصر المسافات بين كل زوج من الحروف.

- المُدخَل سيحتوي فقط على أحرف صغيرة
- يُعامل الأبجد كدائرة، لذا المسافة بين `a` و `z` هي 1.

# --hints--

`letter_distance("abc", "bcd")` يجب أن تُرجع `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("abc", "bcd"), 3)`)
}})
```

`letter_distance("abc", "xyz")` يجب أن تُرجع `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("abc", "xyz"), 9)`)
}})
```

`letter_distance("encrypt", "decrypt")` يجب أن تُرجع `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("encrypt", "decrypt"), 10)`)
}})
```

`letter_distance("algorithm", "codeblock")` يجب أن تُرجع `43`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("algorithm", "codeblock"), 43)`)
}})
```

`letter_distance("lobster", "penguin")` يجب أن تُرجع `47`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("lobster", "penguin"), 47)`)
}})
```

`letter_distance("alligator", "crocodile")` يجب أن تُرجع `55`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("alligator", "crocodile"), 55)`)
}})
```

# --seed--

## --seed-contents--

```py
def letter_distance(str1, str2):

    return str1
```

# --solutions--

```py
def letter_distance(str1, str2):
    total = 0
    for a, b in zip(str1, str2):
        diff = abs(ord(a) - ord(b))
        total += min(diff, 26 - diff)
    return total
```
