---
id: 6a22d77ddf034bc4e35b1d54
title: "تحدٍّ 347: نظرية الألعاب"
challengeType: 29
dashedName: challenge-347
---

# --description--

معطيتان سلسلتا نص متساويتان في الطول تمثلان استراتيجيات لاعبين في لعبة، أرجع النقاط كمصفوفة `[player1, player2]`.

- السلسلتان المعطيتان تحتويان فقط على أحد الحرفين: `"C"` (تعاون) أو `"D"` (خيانة).
- كل حرف يمثل جولة واحدة، تُحتسب النقاط كما يلي:
  - إذا تعاون اللاعبان، يحصل كل منهما على 3 نقاط.
  - إذا خانهما اللاعبان، يحصل كل منهما على نقطة واحدة.
  - إذا خانه لاعب وتعاون الآخر، يحصل الخائن على 5 نقاط والمتعاون على 0.

# --hints--

`play_game("CCCC", "CCCC")` يجب أن تُرجع `[12, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCC", "CCCC"), [12, 12])`)
}})
```

`play_game("DDDD", "DDDD")` يجب أن تُرجع `[4, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDDD", "DDDD"), [4, 4])`)
}})
```

`play_game("CCDD", "CDDD")` يجب أن تُرجع `[5, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCDD", "CDDD"), [5, 10])`)
}})
```

`play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD")` يجب أن تُرجع `[24, 34]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34])`)
}})
```

`play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` يجب أن تُرجع `[66, 21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21])`)
}})
```

# --seed--

## --seed-contents--

```py
def play_game(p1, p2):

    return p1
```

# --solutions--

```py
def play_game(p1, p2):
    scores = [0, 0]
    for i in range(len(p1)):
        if p1[i] == 'C' and p2[i] == 'C':
            scores[0] += 3
            scores[1] += 3
        elif p1[i] == 'D' and p2[i] == 'D':
            scores[0] += 1
            scores[1] += 1
        elif p1[i] == 'D' and p2[i] == 'C':
            scores[0] += 5
        else:
            scores[1] += 5
    return scores
```
