---
id: 6a26df3e2988bcdded204897
title: "تحدٍّ 361: الوقت المنطوق"
challengeType: 29
dashedName: challenge-361
---

# --description--

بالنظر إلى زوايا عقارب الساعة والدقائق في ساعة تناظرية بالدرجات (باتجاه عقارب الساعة من الرقم 12)، أرجع الوقت باللغة الإنجليزية المنطوقة.

حوّل زاوية عقرب الدقائق إلى دقائق (360° = 60 دقيقة)، ثم استخدم القواعد التالية:

| الدقائق | المنطوق |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (باستثناء 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (باستثناء 45) | `"X minutes to Z"` (حيث `X` هو 60 - الدقائق) |

حيث `Y` هو الساعة الحالية و`Z` هي الساعة التالية، كلاهما مستخلص من زاوية عقرب الساعة (360° = 12 ساعة).

ملاحظة: قد لا تقع زوايا العقارب بالضبط على رقم، فكر في تقريبها بطريقة ما.

# --hints--

`get_spoken_time(90, 0)` يجب أن تُرجع `"3 o'clock"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(90, 0), "3 o'clock")`)
}})
```

`get_spoken_time(160, 120)` يجب أن تُرجع `"20 minutes past 5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(160, 120), "20 minutes past 5")`)
}})
```

`get_spoken_time(255, 180)` يجب أن تُرجع `"half past 8"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(255, 180), "half past 8")`)
}})
```

`get_spoken_time(67.5, 92)` يجب أن تُرجع `"quarter past 2"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(67.5, 92), "quarter past 2")`)
}})
```

`get_spoken_time(200, 240)` يجب أن تُرجع `"20 minutes to 7"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(200, 240), "20 minutes to 7")`)
}})
```

`get_spoken_time(322.5, 273)` يجب أن تُرجع `"quarter to 11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(322.5, 273), "quarter to 11")`)
}})
```

`get_spoken_time(117.5, 335)` يجب أن تُرجع `"5 minutes to 4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(117.5, 335), "5 minutes to 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_time(hour_angle, minute_angle):

    return hour_angle
```

# --solutions--

```py
def get_spoken_time(hour_angle, minute_angle):
    minutes = int(minute_angle / 6)
    hour = int(hour_angle / 30) or 12
    next_hour = (hour % 12) + 1

    if minutes == 0:
        return f"{hour} o'clock"
    if minutes == 15:
        return f"quarter past {hour}"
    if minutes == 30:
        return f"half past {hour}"
    if minutes == 45:
        return f"quarter to {next_hour}"
    if minutes < 30:
        return f"{minutes} minutes past {hour}"
    return f"{60 - minutes} minutes to {next_hour}"
```
