---
id: 6a26df3e2988bcdded204893
title: "تحدٍّ 357: السلسلة الغذائية"
challengeType: 29
dashedName: challenge-357
---

# --description--

معطاة مصفوفة من أزواج `[predator, prey]`، أعد السلسلة الغذائية من المفترس الأعلى حتى القاع.

- المفترس الأعلى هو الحيوان الذي لا يكون فريسة لحيوان آخر.
- أعد السلسلة كمصفوفة من سلاسل نصية.

# --hints--

`get_food_chain([["cat", "mouse"]])` يجب أن تُعيد `["cat", "mouse"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["cat", "mouse"]]), ["cat", "mouse"])`)
}})
```

`get_food_chain([["wolf", "deer"], ["deer", "grass"]])` يجب أن تُعيد `["wolf", "deer", "grass"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"])`)
}})
```

`get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` يجب أن تُعيد `["hawk", "snake", "frog", "fly"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"])`)
}})
```

`get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` يجب أن تُعيد `["eagle", "fox", "rabbit", "grass"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"])`)
}})
```

`get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` يجب أن تُعيد `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_food_chain(pairs):

    return pairs
```

# --solutions--

```py
def get_food_chain(pairs):
    prey_set = {p for _, p in pairs}
    chain_map = dict(pairs)
    apex = next(predator for predator, _ in pairs if predator not in prey_set)
    chain = [apex]
    while chain[-1] in chain_map:
        chain.append(chain_map[chain[-1]])
    return chain
```
