---
id: 6a26df3e2988bcdded204894
title: "تحدٍّ 358: مترجم الرموز التعبيرية"
challengeType: 29
dashedName: challenge-358
---

# --description--

بالنظر إلى سلسلة نصية من الرموز التعبيرية، أعد العبارة باستخدام الجدول التالي:

| رمز تعبيري | كلمة |
|-------|------|
| 👶 | `"baby"` |
| 🐱 | `"cat"` |
| 🐕 | `"dog"` |
| 🐟 | `"fish"` |
| 🥵 | `"hot"` |
| 🧊 | `"ice"` |
| 🪨 | `"rock"` |
| 🦈 | `"shark"` |
| 🍲 | `"soup"` |
| ⭐ | `"star"` |

أعد الكلمات مفصولة بمسافات.

# --hints--

`get_emoji_phrase("🪨⭐")` يجب أن تُعيد `"rock star"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🪨⭐"), "rock star")`)
}})
```

`get_emoji_phrase("🥵🐕")` يجب أن تُعيد `"hot dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🥵🐕"), "hot dog")`)
}})
```

`get_emoji_phrase("👶🦈")` يجب أن تُعيد `"baby shark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("👶🦈"), "baby shark")`)
}})
```

`get_emoji_phrase("⭐🐟")` يجب أن تُعيد `"star fish"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("⭐🐟"), "star fish")`)
}})
```

`get_emoji_phrase("🧊🧊👶")` يجب أن تُعيد `"ice ice baby"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🧊🧊👶"), "ice ice baby")`)
}})
```

`get_emoji_phrase("🐱🐟🍲")` يجب أن تُعيد `"cat fish soup"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🐱🐟🍲"), "cat fish soup")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_emoji_phrase(s):

    return s
```

# --solutions--

```py
def get_emoji_phrase(s):
    table = {
        '👶': 'baby',
        '🐱': 'cat',
        '🐕': 'dog',
        '🐟': 'fish',
        '🥵': 'hot',
        '🧊': 'ice',
        '🪨': 'rock',
        '🦈': 'shark',
        '🍲': 'soup',
        '⭐': 'star',
    }
    return ' '.join(table[emoji] for emoji in s)
```
