---
id: 6a26df3e2988bcdded204895
title: "تحدٍّ 359: حاسبة مؤشر الهانديكاب في الغولف"
challengeType: 29
dashedName: challenge-359
---

# --description--

معطاة مصفوفة من نتائج الغولف ومصفوفة مقابلة لقيم البار لكل ملعب، أرجع مؤشر الهانديكاب للاعب الغولف باستخدام الطريقة التالية:

- احسب الفرق لكل جولة بطرح قيمة البار من النتيجة، ثم أرجع متوسط كل الفروق بعد تقريبه إلى منزلة عشرية واحدة.

# --hints--

`calculate_handicap([72, 72, 72], [72, 72, 72])` يجب أن تُرجع `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([72, 72, 72], [72, 72, 72]), 0)`)
}})
```

`calculate_handicap([80, 76, 78, 78], [72, 72, 72, 72])` يجب أن تُرجع `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([80, 76, 78, 78], [72, 72, 72, 72]), 6)`)
}})
```

`calculate_handicap([42, 45, 46, 44], [36, 36, 36, 36])` يجب أن تُرجع `8.3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3)`)
}})
```

`calculate_handicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])` يجب أن تُرجع `8.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8)`)
}})
```

`calculate_handicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])` يجب أن تُرجع `11.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_handicap(scores, pars):

    return scores
```

# --solutions--

```py
def calculate_handicap(scores, pars):
    differentials = [s - p for s, p in zip(scores, pars)]
    average = sum(differentials) / len(differentials)
    return int(average * 10 + 0.5) / 10
```
