---
id: 67116c175a77cb64c5c10d49
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在之前的課程中，你學會了如何使用像 `map`、`filter` 和 `reduce` 這類高階函式（程式）。

在這個工作坊中，你將建置一個函式庫管理應用程式，讓你有機會練習使用不同的高階函式（程式）。

要開始，創建一個名為 `library` 的變數並指定一個空陣列。

# --hints--

你應該創建一個 `library` 陣列。

```js
assert.isArray(library)
```

你的 `library` 陣列應該是空的。

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
