---
id: 67116d7584d0b469b14579c0
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

在 `library` 陣列中創建另一個物件，並包含以下屬性和值：

| Property | Value   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


你應該在你的 `library` 陣列的第二個條款中有一個物件。

```js
assert.isObject(library[1]);
```

你的物件應該有一個 `title` 屬性。

```js
assert.property(library[1], "title")
```

你的 `title` 屬性應該有 `"Atomic Habits"` 作為其值。

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

你的物件應該有一個 `author` 屬性。

```js
assert.property(library[1], "author")
```

你的 `author` 屬性應該有 `"James Clear"` 作為其值。

```js
assert.propertyVal(library[1], "author", "James Clear");
```

你的物件應該有一個 `about` 屬性。

```js
assert.property(library[1], "about")
```

你的 `about` 屬性應該具有值 `"A practical book about discarding bad habits and building good ones"`。

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

你的物件應該有一個 `pages` 屬性。

```js
assert.property(library[1], "pages")
```

你的 `pages` 屬性的值應該是一個數字。

```js
assert.isNumber(library[1]?.pages)
```

你的 `pages` 屬性應該設為 `320`。

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--
  
--fcc-editable-region--
];
```
