---
id: 67116d7584d0b469b14579bf
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

在 `library` 陣列中，創建一個具有以下屬性和值的物件：

| Property | Value   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

你應該在 `library` 陣列中有一個物件。

```js
assert.isObject(library[0])
```

你的物件應該有一個 `title` 屬性。

```js
assert.property(library[0], "title")
```

你的 `title` 屬性的值應該是 `"Your Next Five Moves: Master the Art of Business Strategy"`。

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

你的物件應該有一個 `author` 屬性。

```js
assert.property(library[0], "author")
```

你的 `author` 屬性應該具有值 `"Patrick Bet-David and Greg Dinkin"`。

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

你的物件應該有一個 `about` 屬性。

```js
assert.property(library[0], "about")
```

你的 `about` 屬性應該有值 `"A book on how to plan ahead"`。

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

你的物件應該有一個 `pages` 屬性。

```js
assert.property(library[0], "pages")
```

你的 `pages` 屬性的值應該是一個數字。

```js
assert.isNumber(library[0]?.pages)
```

你的 `pages` 屬性應該設為 `320`。

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
