---
id: 6827deed7fc01f074c90fc0d
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

用另一位作者測試你的 `getBooksByAuthor` 函式會很不錯。

首先將訊息 `"\nList of books by James Clear:\n"` 記錄到主控台。 

在那個 `console.log()` 下面，添加另一個 `console.log()`。在那個 `console.log()` 裡，使用 `library` 和 `"James Clear"` 作為引數呼叫 `getBooksByAuthor()` 函式。

現在，你應該會在主控台看到該特定作者的所有書籍。

# --hints--

你應該將訊息 `"\nList of books by James Clear:\n"` 記錄到主控台。

```js
assert.match(code, /console\.log\(("|')\\nList\s+of\s+books\s+by\s+James\s+Clear:\\n\1\);?/);
```

你應該有一個 `console.log(getBooksByAuthor(library, "James Clear"))`。

```js
assert.match(code, /console\.log\(\s*getBooksByAuthor\s*\(\s*library\s*,\s*(['"])\s*James\s+Clear\s*\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`).join("\n");
}

console.log(getBookInformation(library));

console.log("\nList of book summaries:\n");

function getBookSummaries(catalog) {
  return catalog.map((book) => book.about).join("\n");
}

console.log(getBookSummaries(library));

console.log("\nList of books by Arvid Kahl:\n");

function getBooksByAuthor(catalog, author) {
  return catalog.filter((book) => book.author === author);
}

console.log(getBooksByAuthor(library, "Arvid Kahl"));

--fcc-editable-region--

--fcc-editable-region--
```
