---
id: 669e2f60e83c011754f711f9
title: "Crea una pagina per un'agenzia di viaggi"
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una dichiarazione `DOCTYPE`.
1. Dovresti avere un elemento `html` con `lang` impostato su `en`.
1. Dovresti avere un elemento `head` contenente un elemento void `meta` con `charset` impostato su `utf-8` e un `title` con il testo `Travel Agency Page`.
1. Dovresti avere un tag `meta` nel tuo elemento `head` che contiene una breve descrizione del tuo sito web per la SEO.
1. Dovresti avere un elemento `h1` per presentare le tue destinazioni di viaggio. Dovrebbe essere l'unico elemento `h1` nella tua pagina.
1. Dovresti avere un paragrafo sotto l'elemento `h1` che introduce le opportunità di viaggio.
1. Dovresti avere un elemento `h2` con il testo `Packages`.
1. Dovresti avere un elemento `p` che introduce brevemente i vari pacchetti.
1. Dovresti avere un elemento lista non ordinata con due elementi lista. I due elementi lista dovrebbero avere rispettivamente il testo `Group Travels` e `Private Tours`. Il testo di ogni elemento lista dovrebbe essere racchiuso da un elemento àncora.
1. Dovresti avere un elemento `h2` con il testo `Top Itineraries`.
1. Dovresti avere almeno tre elementi `figure`, ognuno contenente un elemento àncora e un elemento `figcaption` che descrive l'immagine della figura.
1. I tre elementi àncora dovrebbero avere un elemento `img` con un appropriato attributo `alt` e un attributo `src` impostato su un'immagine valida come loro contenuto. Puoi usare `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` e `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` se vuoi.
1. Tutti i tuoi cinque elementi àncora dovrebbero avere un attributo `href` con il valore `https://www.freecodecamp.org/learn` e un attributo `target` con il valore `_blank`.

# --hints--

La tua pagina per l'agenzia di viaggi dovrebbe avere una dichiarazione `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Dovresti avere un elemento `html` con `lang` impostato su `en`.

```js
assert.match(code, /<\s*html(?=[\s>])[^>]*\slang\s*=\s*('|")en\1[^>]*>/i);
```

Il tuo elemento `html` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/html\s*>/i);
```

Dovresti avere un elemento `head` all'interno dell'elemento `html`.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>/i
);
```

Il tuo elemento `head` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/head\s*>/i);
```

Dovresti avere un elemento `title` all'interno del tuo elemento `head`.

```js
assert.isNotNull(document.querySelector('head title'));
```

Il tuo elemento `title` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/title\s*>/i);
```

Dovresti avere due elementi `meta` all'interno del tuo elemento `head`.

```js
assert.isAtLeast(document.querySelectorAll('head meta').length, 2);
```

Un elemento `meta` dovrebbe avere un attributo `name` con valore `description` e un attributo `content` non vuoto.

```js
assert(
  Array.from(document.querySelectorAll('head meta')).some(
    (meta) =>
      meta.getAttribute('name')?.toLowerCase() === 'description' &&
      meta.getAttribute('content')?.trim()
  )
);
```

Un elemento `meta` dovrebbe avere il suo attributo `charset` impostato su `utf-8`.

```js
const charsetMeta = document.querySelector('head meta[charset]');
assert.equal(charsetMeta?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Il tuo elemento `title` dovrebbe avere il nome della tua agenzia di viaggi.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Dovresti avere un elemento `body` all'interno del tuo elemento `html`.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>(?:(?!<\/html\s*>)[\s\S])*<\s*body\s*>/i
);
```

Il tuo elemento `body` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/body\s*>/i);
```

Dovresti avere un elemento `h1` per presentare le tue destinazioni di viaggio.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Dovresti avere solo un elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Dovresti avere un elemento `p` sotto l'elemento `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

Il tuo primo elemento `p` dovrebbe introdurre le opportunità di viaggio.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Il tuo primo elemento `h2` dovrebbe contenere il testo `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

Dovresti avere un elemento `p` sotto il tuo primo elemento `h2`.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Il tuo secondo elemento `p` dovrebbe introdurre brevemente i vari pacchetti.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

Dovresti avere un elemento lista non ordinata sotto il tuo secondo elemento `p`.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

Dovresti avere due elementi nella tua lista non ordinata.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Entrambi i tuoi elementi della lista dovrebbero contenere un elemento àncora.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

L'elemento àncora del tuo primo elemento della lista dovrebbe racchiudere il testo `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

L'elemento àncora del tuo secondo elemento della lista dovrebbe racchiudere il testo `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

Dovresti avere un elemento `h2` dopo la tua lista non ordinata.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

Il tuo secondo elemento `h2` dovrebbe contenere il testo `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

Dovresti avere almeno tre elementi `figure`.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Ogni elemento `figure` dovrebbe contenere un elemento àncora come primo elemento annidato.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Ogni elemento `figure` dovrebbe contenere un elemento `figcaption` come secondo elemento annidato.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Ogni `figcaption` dovrebbe contenere del testo.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Ognuno degli elementi `a` annidati nei tuoi elementi `figure` dovrebbe contenere un'immagine.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Ogni elemento `img` dovrebbe avere un attributo `src` valido.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Ogni elemento `img` dovrebbe avere un attributo `alt` con un valore appropriato.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Ogni elemento `a` dovrebbe avere un attributo `href` con il valore `https://www.freecodecamp.org/learn`. Non dimenticare i collegamenti negli elementi della lista.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Ogni elemento `a` dovrebbe avere un attributo `target` con il valore `_blank`. Non dimenticare i collegamenti negli elementi della lista.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
