---
id: 6690e10ebe2181212abc9652
title: 名刺をデザインする
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `div` に `class` 属性が `business-card` の値で設定されており、他のすべての要素を含む必要があります。  
1. `.business-card` 要素の中に、`img` 属性が `class` の値を持つ `profile-image` 要素がある必要があります。画像のソースは `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg` に設定してもかまいません。意味のある説明を持つ `alt` が必要です。  
1. `.business-card` 要素の中で、`img` 要素の後に、`p` 属性がそれぞれ `class`、`full-name`、`designation` の値を持つ3つの `company` 要素がある必要があります。  
1. 最初の `p` 要素にはあなたの名前を含めてください。  
1. 2番目の `p` 要素にはあなたの役職を含めてください。  
1. 3番目の `p` 要素にはあなたの会社名を含めてください。  
1. 3番目の `hr` 要素の下に `p` 要素がある必要があります。  
1. `hr` 要素の後に2つの `p` 要素がある必要があります。最初の `p` 要素にはメールアドレスのテキストを、2番目の `p` 要素には電話番号のテキストを含めてください。  
1. 2つの `p` 要素の後に、クラスが `a` の `portfolio-link` 要素があり、テキストは `Portfolio` で、有効なURLにリンクしている必要があります。  
1. `hr` 要素の後に、ポートフォリオリンクを含む `a` 要素がある必要があります。  
1. `div` 属性が `class` の値を持つ別の `social-media` 要素が必要です。この要素の中に、テキストが `h2` の `Connect with me` 要素がある必要があります。  
1. `.social-media` 要素の中に3つの `a` 要素がある必要があります。`a` 要素はそれぞれ `Twitter`、`LinkedIn`、`GitHub` のテキストを持ち、有効なウェブサイトへのリンクが設定されている必要があります。  
1. `styles.css` セクションの HTML ファイル内で `link` タグを使い、スタイルシート `head` をリンクしてください。  
1. ページの背景は `rosybrown` にしてください。全体のフォントは `Arial` で、フォールバックは `sans-serif` にしてください。  
1. `.business-card` セレクターには、要素の幅を `300px` に設定し、背景色を任意の色に設定するプロパティが必要です。また、全体に `padding` が `20px` で、上部に `margin` が `100px` あるように設定してください。テキストは中央揃えで、フォントサイズは `16px` にしてください。左右のマージンは `auto` に設定してください。  
1. `.profile-image` セレクターには `max-width` プロパティがあり、その値は `100%` である必要があります。  
1. すべての `p` 要素には上下に `margin` が `5px` で設定されている必要があります。  
1. ページ上のすべてのリンクは下線がないようにしてください。

# --hints--

メインの `div` があり、`class` 属性が `business-card` の値で設定されている必要があります。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

`.business-card` 要素の中に、`img` 要素があり、`class` 属性が `profile-image` の値である必要があります。

```js
assert.exists(document.querySelector('img.profile-image'));
```

画像のソースは有効な画像に設定してください。

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

画像の `alt` 属性は意味のある説明に設定してください。

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

`.business-card` 要素の中で、`img` 要素の後に、`p` 要素があり、`class` 属性が `full-name` の値である必要があります。

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

`.full-name` 要素の後に、`p` 要素があり、`class` 属性が `designation` の値である必要があります。

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

`.designation` 要素の後に、`p` 要素があり、`class` 属性が `company` の値である必要があります。

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

最初の `p` 要素はクラスが `full-name` で、あなたの名前を含めてください。

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

2番目の `p` 要素は `class` 属性が `designation` の値で、あなたの役職を含めてください。

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

3番目の `p` 要素は `class` 属性が `company` の値で、あなたの会社名を含めてください。

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

3番目の `p` 要素の後に、`class` 属性が `company` の値である `hr` 要素がある必要があります。

```js
assert.exists(document.querySelector('p.company + hr'));
```

最初の `hr` 要素の後に、メールアドレスのテキストを含む `p` 要素がある必要があります。

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

メールの `p` 要素の後に、電話番号のテキストを含む別の `p` 要素がある必要があります。

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

電話番号の `p` 要素の後に、クラスが `a` の `portfolio-link` 要素がある必要があります。

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

`.portfolio-link` 要素には `Portfolio` のテキストがある必要があります。

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

`.portfolio-link` 要素には有効な `href` がある必要があります。

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

ポートフォリオを含む `a` 要素の後に、区切りの `hr` 要素がある必要があります。

```js
assert.exists(document.querySelector('a + hr'));
```

2番目の `hr` 要素の後に、`div` 属性が `class` の値である `social-media` 要素がある必要があります。

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

`.social-media` 要素の中に、テキストが `h2` の `Connect with me` 要素がある必要があります。

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

`.social-media` 要素の中に3つの `a` 要素がある必要があります。

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

最初の `a` 要素には `Twitter` のテキストがある必要があります。

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

最初の `href` 要素の `a` は有効なリンクを指している必要があります。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

2番目の `a` 要素には `LinkedIn` のテキストがある必要があります。

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

2番目の `href` 要素の `a` は有効なリンクを指している必要があります。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

3番目の `a` 要素には `GitHub` のテキストがある必要があります。

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

3番目の `href` 要素の `a` は有効なリンクを指している必要があります。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

`link` 要素には `rel` 属性があり、その値は `stylesheet` である必要があります。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

`link` 要素には `href` 属性があり、その値は `styles.css` である必要があります。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

`body` セレクターでページの背景色を `rosybrown` に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

`Arial` 要素でページのフォントを `body` に設定し、フォールバックは `sans-serif` にしてください。

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

`.business-card` セレクターには `width` プロパティがあり、その値は `300px` である必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

`.business-card` セレクターには有効な色の値を持つ `background-color` プロパティがある必要があります。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

`.business-card` セレクターには `padding` プロパティがあり、その値は `20px` である必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

`.business-card` セレクターには上部マージンが `100px` に設定されている必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

`.business-card` 要素内のテキストは中央揃えにしてください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

`.business-card` 要素内のテキストのフォントサイズは `16px` にしてください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

`.business-card` 要素の左右のマージンは `auto` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

`.profile-image` セレクターには `max-width` プロパティがあり、その値は `100%` である必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

`p` 要素には上下に `margin` が `5px` で設定されている必要があります。

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

ページ上のリンクは下線なしにしてください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
