---
id: 68e97fe79367ad7b5dd6c9cd
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

이 단계에서는 텍스트가 `h1`인 `XYZ Bookstore` 요소를 추가하세요.

# --hints--

`h1` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h1'));
```

`h1` 요소의 텍스트는 `XYZ Bookstore`여야 합니다. 철자나 대소문자 오류가 없는지 다시 확인하세요.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

`h1` 요소는 `body` 요소 안에 있어야 합니다.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
