---
id: 68ebdbacdd3fa474132cc975
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`h1` 요소 아래에 `p` 요소를 추가하고, 이 텍스트를 넣으세요: `Browse our collection of amazing books!`.

# --hints--

`p` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('p'));
```

`p` 요소의 텍스트는 `Browse our collection of amazing books!` 여야 합니다.

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

`p` 요소는 `h1` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
