---
id: 68ec6cd5b7e8f5a8f7319e32
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

`div` 요소는 다른 HTML 요소를 그룹화하는 컨테이너로 사용됩니다. 주로 CSS 스타일을 공유할 HTML 요소를 그룹화할 때 `div` 요소를 사용합니다.

`p` 요소 아래에 `div` 요소를 추가하세요. 이 `div`는 책 카드의 컨테이너가 될 것입니다.

**참고**: 이 워크숍에서는 CSS를 적용하지 않습니다. 클래스와 그룹화된 요소는 CSS 스타일링에 유용하지만, 이 워크숍에서는 콘텐츠를 구조화하고 그룹화하는 데만 사용합니다. 스타일링 방법은 이후 모듈에서 배우게 됩니다.

# --hints--

`div` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('div'));
```

`div` 요소는 `p` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
