---
id: 68eca5412b5dd9d06b3d6404
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

두 번째 카드의 `h2` 요소 아래에 다음 텍스트가 들어간 `p` 요소를 추가하세요:

```md
This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.
```

# --hints--

두 번째 요소 안에 `p` 요소가 중첩되어 있어야 하며, 이 요소는 `card` 클래스를 가져야 합니다.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('p'));
```

두 번째 카드의 `p` 요소 텍스트는 `This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.` 여야 합니다.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('p')?.innerText.trim(), "This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.");
```

두 번째 카드의 `p` 요소는 `h2` 요소 아래에 있어야 합니다.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
