---
id: 68ec9332a9b5b2b32487bd00
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

`id` 속성은 HTML 요소에 고유 식별자를 추가합니다. 각 `id`는 페이지 내에서 고유해야 하며 한 번만 사용해야 합니다.

`id` 값에는 공백이 포함될 수 없으며 문자, 숫자, 밑줄, 대시만 포함해야 합니다.

예시는 다음과 같습니다:

```html
<p id="para">example paragraph</p>
```

`id` 클래스를 가진 요소에 `card` 속성을 추가하고 그 값을 `sally-adventure-book`로 설정하세요.

# --hints--

`card` 클래스를 가진 요소에 `id` 속성이 있어야 합니다.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

`card` 클래스를 가진 요소에 `id` 속성이 있고 그 값이 `sally-adventure-book`이어야 합니다.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">
  --fcc-editable-region--
      
    </div>
  </div>
</body>

</html>
```
