---
id: 68eca5c0065c82d256d29ca3
title: 17단계
challengeType: 0
dashedName: step-17
---

# --description--

두 번째 카드 안에 `button` 요소를 추가하고 `class` 속성을 `btn`로 설정한 뒤 텍스트를 `Buy Now`로 입력하세요.

두 `button` 요소는 이제 같은 `class`를 공유하므로 함께 일관되게 스타일이 적용될 수 있습니다.

# --hints--

두 번째 요소 안에 `button` 클래스를 가진 `card` 요소가 중첩되어 있어야 합니다.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

두 번째 카드의 `button` 요소에는 `class` 속성이 있어야 합니다.

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

두 번째 카드의 `button` 요소에는 값이 `class`인 `btn`이 있어야 합니다.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

두 번째 카드의 `button` 요소의 텍스트는 `Buy Now`여야 합니다.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
