---
id: 68eca6e5a6759ed4ea0034dc
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

`btn-container` 클래스를 가진 요소 안에 `button` 요소 두 개를 추가하세요:

첫 번째 버튼:

- Id: `view-cart-btn`
- 클래스: `btn`
- 텍스트: `View Cart`

두 번째 버튼:

- Id: `checkout-btn`
- 클래스: `btn`
- 텍스트: `Checkout`


축하합니다! div, 클래스, id를 사용해 서점 페이지 구조를 성공적으로 만들었습니다.

# --hints--

`button` 요소 안에 `.btn-container` 요소 두 개가 중첩되어 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

`btn-container` 클래스를 가진 요소 안의 첫 번째 버튼에는 `id` 속성이 `view-cart-btn` 값으로 있어야 합니다.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

`btn-container` 클래스를 가진 요소 안의 첫 번째 버튼에는 `class` 속성이 `btn` 값으로 있어야 합니다.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

첫 번째 버튼의 텍스트는 `View Cart` 이어야 합니다.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

`btn-container` 클래스를 가진 요소 안의 두 번째 버튼에는 `id` 속성이 `checkout-btn` 값으로 있어야 합니다.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

`class="btn-container"` 요소 안의 두 번째 버튼에는 `class` 속성이 `btn` 값으로 있어야 합니다.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

두 번째 버튼의 텍스트는 `Checkout` 이어야 합니다.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
  <p>Review your selections and continue to checkout.</p>
  
  <div class="btn-container">
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
