---
id: 68ec99e478211dc578699944
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

`h2` 클래스를 가진 첫 번째 요소 안에 있는 `card` 요소 아래에 다음 텍스트가 포함된 `p` 요소를 추가하세요:

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

`p` 클래스를 가진 요소 안에 `card` 요소가 중첩되어 있어야 합니다.

```js
assert.exists(document.querySelector('.card p'));
```

`p` 요소의 텍스트는 `This is an epic story of Sally and her dog Rex as they navigate through other worlds.` 여야 합니다.

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

`p` 요소는 `h2` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
