---
id: 68ec6d9a315221aa31e54816
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

`class` 속성은 하나 이상의 요소를 스타일링하기 위해 식별하는 데 사용됩니다. `id` 속성과 달리, 클래스 이름은 고유할 필요가 없으며 여러 요소가 같은 클래스를 공유할 수 있습니다.

예시는 다음과 같습니다:

```html
<p class="example">example paragraph</p>
```

`class` 요소에 `div` 속성을 추가하고 그 값을 `card-container`로 설정하세요.

# --hints--

`div` 요소에 `class` 속성이 있어야 합니다.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

`div` 요소에 `class` 속성이 있고 그 값이 `card-container`이어야 합니다.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
