---
id: 68eca2a795b333ca5fee30a8
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

이제 두 번째 책 카드를 만드세요. `div` 속성이 `class`로 설정된 또 다른 `card` 요소를 추가하세요. 동일한 클래스 이름을 여러 요소에 재사용하여 일관된 스타일을 적용할 수 있다는 점을 확인하세요.

# --hints--

`div` 속성 값이 `class`인 `card` 요소가 두 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

`card` 클래스를 가진 두 번째 요소는 `card-container` 클래스를 가진 요소 안에 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

`card` 클래스를 가진 두 요소는 형제 요소여야 합니다.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```
