---
id: 68ec98b38d83a3c28dd30efe
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

`card` 클래스를 가진 첫 번째 요소 안에 `h2` 텍스트를 가진 `Sally's SciFi Adventure` 요소를 추가하세요.

# --hints--

`h2` 클래스를 가진 요소 안에 `card` 요소가 중첩되어 있어야 합니다.

```js
assert.exists(document.querySelector('.card h2'));
```

`h2` 요소의 텍스트는 `Sally's SciFi Adventure`이어야 합니다.

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
