---
id: 68eca6364c5616d393389a30
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

HTML 요소는 다음과 같이 생겼다는 점을 기억하세요:

```html
<element attribute="value">
    inner text
</element>
```

클래스가 `card-container`인 요소 아래에, 이 텍스트를 가진 새로운 `p` 요소를 추가하세요:

```md
Review your selections and continue to checkout.
```

`p` 요소 아래에 `div` 속성이 `class`로 설정된 `btn-container` 요소를 만드세요. 이 컨테이너는 내비게이션 버튼 요소들을 그룹화할 것입니다.

# --hints--

새로 추가한 `p` 요소는 클래스가 `card-container`인 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('.card-container + p'));
```

새로운 `p` 요소의 텍스트는 `Review your selections and continue to checkout.`이어야 합니다.

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

`div`이 `class`인 `btn-container` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('div.btn-container'));
```

`div` 요소가 `class`가 `btn-container`인 상태로 새로 추가한 `p` 요소 뒤에 와야 합니다.

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
