---
id: 68eca4c4b0f952cf09fabe09
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

`card` 클래스를 가진 두 번째 요소 안에 `h2` 텍스트를 가진 `Dave's Cooking Adventure` 요소를 추가하세요.

# --hints--

`h2` 클래스를 가진 두 번째 요소 안에 `card` 요소가 중첩되어 있어야 합니다.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

`h2` 클래스를 가진 두 번째 요소 안의 `card` 요소는 `Dave's Cooking Adventure` 텍스트를 가져야 합니다.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText.trim(), "Dave's Cooking Adventure");
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
