---
id: 68eca11e99e3c5c894ca9d69
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

`button` 요소는 웹페이지에서 클릭 가능한 버튼을 만드는 데 사용됩니다. 버튼은 사용자가 클릭하여 동작을 수행할 수 있는 상호작용 요소입니다.

`button`가 `class`인 요소 안에 `card` 요소를 추가하고, 버튼에 `class` 속성을 `btn`로 설정하며, 텍스트는 `Buy Now`로 하세요.

# --hints--

`button` 클래스를 가진 요소 안에 `card` 요소가 중첩되어 있어야 합니다.

```js
assert.exists(document.querySelector('.card button'));
```

`button` 요소에 `class` 속성이 있어야 합니다.

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

`button` 요소에 `class`가 있고 그 값이 `btn`이어야 합니다.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

`button` 요소의 텍스트는 `Buy Now`여야 합니다.

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
