---
id: 69134b6cb43aa5f254950a10
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 다양한 책에 대한 정보를 보여주는 책 카드들을 만들어 서점 페이지를 구성할 것입니다. `div` 요소, 클래스, ID를 사용해 콘텐츠를 정리하는 연습을 하세요.

HTML 보일러플레이트를 만들어 서점 페이지를 시작하세요.

`<!DOCTYPE html>` 선언과 `html`, `head` 요소를 추가하세요.

`html` 요소에 `lang` 속성을 추가하고 값을 `"en"`으로 설정하세요.

# --hints--

코드에 `<!DOCTYPE html>` 선언이 포함되어야 합니다.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

코드에 여는 `<html>` 태그와 닫는 `<html>` 태그가 모두 포함되어야 합니다.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

`html` 요소에 `lang` 속성이 `"en"`으로 설정되어야 합니다.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

코드에 여는 `<head>` 태그와 닫는 `<head>` 태그가 모두 포함되어야 합니다.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

`head` 요소는 `html` 요소 안에 중첩되어야 합니다.

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
