---
id: 69134d14150d51f8066a3246
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

이제 HTML 문서 구조를 개선하여 페이지가 올바르게 인코딩되도록 하세요.

`head` 요소 안에 `<meta charset="UTF-8">` 요소를 추가하세요.

마지막으로, `head` 섹션 아래에 `body` 요소를 추가하세요. 이곳에 페이지에 보이는 모든 내용이 들어갑니다.

# --hints--

코드에 `<meta charset="UTF-8">` 요소가 포함되어야 합니다.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

코드에 여는 `<body>` 태그와 닫는 `<body>` 태그가 포함되어야 합니다.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

`body` 요소는 `head` 요소 아래에 있어야 합니다.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
  <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region--
</html>
```
