---
id: 691f85085787a8c4ca589fbb
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`title` 요소 안에 `head` 요소를 추가하세요.

페이지 제목을 `XYZ Bookstore Page`로 설정하세요.

# --hints--

코드에는 `<title>` 태그의 여는 태그와 닫는 태그가 모두 포함되어야 합니다.

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

제목은 `XYZ Bookstore Page`이어야 합니다.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

`title` 요소는 `head` 요소 안에 중첩되어야 합니다.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
--fcc-editable-region-- 
</head>

</html>
```
