---
id: 687d8e05a529a4838f1c961c
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Ifuatayo, ondoa neno kuu la `pass` ndani ya darasa la `Node`, kisha mpe darasa la `Node` njia ya `__init__` yenye vigezo viwili, `self` na `element`.

Ndani ya njia hiyo, weka thamani ya `element` kwa `self.element`. Hii inahifadhi thamani inayopitishwa wakati wa kuunda nodi.

# --hints--

Unapaswa kufafanua njia ya `__init__` ndani ya darasa lako la `Node`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

Njia yako ya `__init__` inapaswa kuwa na vigezo viwili, `self` na `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

Unapaswa kuweka thamani ya `element` kwa `self.element` ndani ya njia ya `__init__`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
