---
id: 687d7deea6a2a22fa8816a73
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika mafundisho yaliyopita, ulijifunza kuhusu baadhi ya miundo ya data inayojulikana. Katika warsha hii, utatekeleza orodha iliyounganishwa.

Kwanza, tengeneza darasa la vitu linaloitwa `LinkedList` ili wakilishe muundo utakaoshikilia nodi zako za orodha. Ongeza neno kuu la `pass` ndani ya darasa ili jaribio lipite.

# --hints--

Unapaswa kutengeneza darasa la vitu linaloitwa `LinkedList`. Kumbuka kuongeza neno kuu la `pass`.

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
