---
id: 687d8c6f066877768f7ee441
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Ndani ya darasa la vitu la `LinkedList`, badilisha neno kuu la `pass` na darasa la vitu la ndani linaloitwa `Node`. Hii itawakilisha nodi binafsi zinazohifadhi data katika orodha.

# --hints--

Unapaswa kuunda darasa la vitu la `Node` ndani ya darasa la vitu la `LinkedList`. Kumbuka kuongeza neno kuu la `pass` ndani ya darasa la vitu la `Node`.

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    pass
--fcc-editable-region--
```
