---
id: 687e366b3004a9831e4cbd41
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Ongeza njia iitwayo `is_empty` inayorejesha `self.length == 0`.

Utaitumia hii kufuatilia ni nodi ngapi ziko katika orodha. Hivyo, kwa `self.length == 0`, inarejesha `True`, ikimaanisha kwamba orodha haina nodi yoyote.

Ikiwa `self.length` ni kubwa kuliko `0`, inarejesha `False`, ikimaanisha kwamba orodha ina angalau nodi moja.

# --hints--

Unapaswa kuunda njia ya `is_empty` inayorejesha `self.length == 0`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
