---
id: 687e344cc1993f7316af030c
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Ndani ya darasa la `LinkedList`, ongeza njia ya `__init__` inayoweka `self.length` kuwa `0`. Hii huanzisha kielekezi kufuatilia ni nodi ngapi (vipengele) zipo sasa kwenye orodha.

# --hints--

Unapaswa kuweka `self.length` kuwa `0` ndani ya njia ya `__init__` ya darasa la `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
