---
id: 688003f589b875af60273dd9
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Bado ndani ya njia ya `__init__`, weka thamani ya `None` kwa `self.next`. Hii inaweka kidokezo cha `next` kwa `None`, ikimaanisha kuwa nodi mpya bado haionyeshi nodi nyingine yoyote katika orodha.

Baadaye, `next` inaweza kusahihishwa kuonyesha nodi nyingine ili kuunda mnyororo wa orodha iliyounganishwa.

# --hints--

Ndani ya njia ya `__init__`, unapaswa kuanzisha sifa ya kipengele ya `next` kwa `None`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
