---
id: 68800a64b06e80d9a31c131b
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Bado katika njia ya `__init__` ya darasa la vitu la `LinkedList`, weka `self.head` kuwa `None`. Hii hufanya kama rejea kwa nodi ya kwanza katika orodha. Orodha bado ni tupu kwa sasa, kwa hivyo, lazima uiweke kuwa `None`.


# --hints--

Unapaswa kuweka `self.head` kuwa `None` ndani ya njia ya `__init__` ya darasa la vitu la `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

    def __init__(self):
        self.length = 0
--fcc-editable-region--
        
--fcc-editable-region--
```
