---
id: 68857e0a49baf4751f737a2a
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Sasa, baada ya njia ya `is_empty`, tengeneza njia iitwayo `add` ndani ya darasa la vitu la `LinkedList`. Njia mpya inapaswa kuwa na vigezo viwili, `self` na `element`.

Ndani ya njia yako mpya, tengeneza kigezo kinachoitwa `node` na weka thamani yake kuwa mfano mpya wa `Node` kwa kutumia `self.Node(element)`. Hii itaunda nodi mpya ya kuongezwa kwenye orodha iliyounganishwa.

# --hints--

Unapaswa kutengeneza njia ya `add` ndani ya darasa la vitu la `LinkedList`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

Njia yako ya `add` inapaswa kuwa na vigezo viwili, `self` na `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

Unapaswa kuweka thamani ya `self.Node(element)` kwa `node` katika njia ya `add`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
