---
id: 688649c8673f842f49228a5e
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Ikiwa orodha si tupu, njia inapaswa kutafuta nodi ya mwisho na kuunganisha nodi mpya na kiashiria chake cha `next`.

Ili kufanya hivyo, tengeneza kipande cha `else` baada ya kipande cha `if`. Ndani ya kipande cha `else`, tengeneza kigezo cha `current_node` na weka thamani ya `self.head` ndani yake.

# --hints--

Unapaswa kuweka thamani ya `self.head` kwa `current_node` ndani ya kipande cha `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
