---
id: 68867c3898b09115d8c5140f
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Mzunguko wa `while` uliouunda sasa hivi unakuwezesha kupita nodi kila moja katika orodha hadi nodi ya mwisho ifikike. Mara hiyo itakapotokea, unaweza kuongeza nodi hiyo kwenye orodha.

Katika blokki ya `else`, mara moja baada ya blokki ya `while`, weka thamani ya `node` kwa `current_node.next`.

# --hints--

Unapaswa kuweka thamani ya `node` kwa `current_node.next` katika blokki ya `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].has_stmt("current_node.next = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
