---
id: 6888d05ba074461b64f8eaca
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Ifuatayo, utaunda njia nyingine ambayo itatumika kuondoa kipengele fulani kutoka kwenye orodha iliyounganishwa.

Unda njia iitwayo `remove` yenye vigezo viwili, `self` na `element`. Ndani yake, unda kigezo kinachoitwa `previous_node` na weka thamani ya `None` ndani yake.

# --hints--

Unapaswa kuunda njia iitwayo `remove`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_class("LinkedList").has_function("remove")
    `)
  )
})
```

Njia yako ya `remove` inapaswa kuwa na vigezo viwili, `self` na `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("remove").has_args("self, element")`) 
})
```

Unapaswa kuweka thamani ya `None` kwa `previous_node` katika njia ya `remove`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").find_body().is_equivalent("previous_node = None")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
