---
id: 6888fe6c745b6852ed240229
title: Hatua 18
challengeType: 20
dashedName: step-18
---

# --description--

Ifuatayo, utapitia orodha iliyounganishwa ili kupata nodi inayoshikilia kipengele kilichotajwa cha kuondoa.

Tengeneza mzunguko wa `while` unaokagua wakati `current_node` si `None` na wakati `current_node.element` si sawa na `element`. Ndani ya mzunguko wa `while`, weka thamani ya `current_node` kwa `previous_node`.

# --hints--

Unapaswa kuwa na mzunguko wa `while` ndani ya njia ya `remove`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
`)})
```

Mzunguko wako wa `while` unapaswa kukagua wakati `current_node` si `None` na wakati `current_node.element` si sawa na `element`.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0].find_conditions()[0]
  assert _cond.is_equivalent("current_node is not None and current_node.element != element")
`)})
```

Unapaswa kuweka thamani ya `current_node` kwa `previous_node` ndani ya mzunguko wa `while`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
