---
id: 688acb2053ac7978a06345f1
title: Hatua 21
challengeType: 20
dashedName: step-21
---

# --description--

Ifuatayo, tengeneza usemi wa `elif` unaothibitisha kuwa `previous_node` si `None`. Weka thamani ya `current_node.next` kwa `previous_node.next` ndani yake.

Kwa hiyo, ikiwa kipengele kinachotakiwa kuondolewa kimepatikana na si katika nodi ya kichwa, kidokezo cha `next` cha nodi iliyotangulia kitawekwa upya ili kuruka nodi ya sasa (kipengele kinachotakiwa kuondolewa), na kuiondoa kutoka katika orodha iliyounganishwa.

# --hints--

Unapaswa kutengeneza usemi wa `elif` unaothibitisha kuwa `previous_node` si `None`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("previous_node is not None")
`) })
```

Unapaswa kuweka thamani ya `current_node.next` kwa `previous_node.next` ndani ya usemi wa `elif`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return\\nelif previous_node is not None:\\n previous_node.next = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
        if current_node is None:
            return
--fcc-editable-region--
        
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
