---
id: 68b84df2648b4c96b25022f3
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Ni wakati wa kujaribu orodha yako iliyounganishwa hadi sasa! Mwishoni mwa msimbo wako, tengeneza mfano wa `LinkedList` uitwae `my_list`. Kisha, ukitumia njia ya `is_empty` uliyotengeneza katika hatua iliyopita, chapisha `my_list.is_empty()`.

# --hints--

Unapaswa kutengeneza mfano wa `LinkedList` uitwae `my_list`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

Unapaswa kuchapisha `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
