---
id: 68b9e951e4892721632acc26
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Nzuri! Sasa ni wakati wa kujaribu kuongeza vipengele kwenye orodha yako iliyounganishwa. Mwishoni mwa msimbo wako, baada ya kuunda mfano wa `LinkedList` uitwao `my_list` na kuchapisha `my_list.is_empty()`, ongeza `1` na `2` kwenye orodha kwa kutumia njia ya `add`.

Kisha, chapisha tena `my_list.is_empty()` kuthibitisha kuwa orodha haijabaki tupu. Mwishowe, chapisha `my_list.length` kuona urefu wa sasa wa orodha iliyounganishwa.

# --hints--

Unapaswa kuongeza `1` kwenye orodha kwa kutumia njia ya `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

Unapaswa kuongeza `2` kwenye orodha kwa kutumia njia ya `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

Unapaswa kuchapisha `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

Unapaswa kuchapisha `my_list.length`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
