---
id: 65ef1be2b958cabf712296cd
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Ora restringerai ripetutamente l'intervallo trovando il punto medio dell'intervallo corrente e confrontando il quadrato del punto medio con il valore obiettivo.

Per farlo, all'interno del blocco `else` crea un loop `for` che venga eseguito fino a `max_iterations` volte.

Per il tuo loop, usa la funzione `range`, che genera una sequenza di numeri su cui puoi iterare. La sintassi è `range(start, stop, step)`, dove `start` è l'intero di partenza (incluso), `stop` è l'ultimo intero (non incluso) e `step` è la differenza tra un numero e quello precedente nella sequenza.

Usa anche `_` come variabile del loop. `_` funge da segnaposto ed è utile quando hai bisogno di usare una variabile ma non ti serve effettivamente il suo valore.

# --hints--

Dovresti creare un loop `for` per iterare su `range(max_iterations)`. Usa `_` come variabile del loop.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
