---
id: 65ef1a720d6e65bbb0d22463
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--
Ora lavorerai sui casi in cui il `square_target` è un numero positivo diverso da `1` o `0`.

Crea una clausola `else` per gestire questi casi.

# --hints--

Crea una clausola `else` e non dimenticare di aggiungere una parola chiave `pass`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[2] == _Node()`))
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')
--fcc-editable-region--


--fcc-editable-region--
```
