---
id: 65ef190c6b51e9b5a5f7ed29
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Dai alla funzione `square_root_bisection` i seguenti parametri:

- `square_target`: Il numero di cui vuoi trovare la radice quadrata.
- `tolerance` (opzionale): La differenza accettabile tra il quadrato del valore approssimato della radice e il valore reale obiettivo (il valore predefinito è `1e-7`). La tolleranza `1e-7` implica che la soluzione sarà precisa entro `0.0000001` del valore reale ed è una buona scelta predefinita che bilancia precisione e prestazioni.
- `max_iterations` (opzionale): Il numero massimo di iterazioni da eseguire (il valore predefinito è `100`). Se il metodo non converge entro questo limite, si assume che la soluzione non sia stata trovata.
  
# --hints--

La tua funzione dovrebbe avere questi parametri: `square_target`, `tolerance = 1e-7` e `max_iterations = 100`. L’ordine è importante.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```
