---
id: 65ef19425d1b27b6c930bae6
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Se il numero di cui vuoi trovare la radice quadrata è negativo, il codice dovrebbe generare un errore perché la radice quadrata di un numero negativo non è definita nei numeri reali.

Rimuovi l'istruzione `pass` e crea un'istruzione `if` per verificare se `square_target` è minore di `0`.

# --hints--

Dovresti rimuovere la parola chiave `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

Dovresti creare un'istruzione `if` per verificare se `square_target < 0`.

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
