---
id: 65ef1a50049cf9bada13266f
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Crea un'istruzione `elif` per verificare se `square_target` è uguale a `0`. Se lo è, assegna il valore `0` alla variabile `root`. Inoltre, stampa il messaggio `'The square root of {square_target} is 0'`. Ricorda di formattare il messaggio usando una f-string.

# --hints--

Dovresti avere un'istruzione `elif` per controllare la condizione `square_target == 0`.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

Dovresti assegnare il valore `0` alla variabile `root` e passare l'argomento `f'The square root of {square_target} is 0'` a una chiamata `print`.


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
