---
id: 65ef181dad3d96b14a21a03e
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo progetto troverai la radice quadrata approssimativa di un numero dato usando il metodo della bisezione.

Il metodo della bisezione è una tecnica per trovare le radici di una funzione a valori reali. Funziona restringendo un intervallo in cui si trova la radice quadrata finché non converge a un valore entro una tolleranza specificata.

Inizia creando una funzione chiamata `square_root_bisection`. Lascia per ora i parametri vuoti.

# --hints--

Dovresti definire la funzione `square_root_bisection` senza parametri.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
