---
id: 65ef198fde24dfb7ff675b42
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

L'istruzione `raise` ti permette di forzare il verificarsi di un'eccezione specifica. Consiste nella parola chiave `raise` seguita dal tipo di eccezione, e ti consente di fornire un messaggio di errore personalizzato:

```py
raise ValueError("Invalid value")
```

Quando il codice sopra viene eseguito, viene sollevata una `ValueError` e all'utente viene mostrato il messaggio `"Invalid value"`.

Se `square_target` è minore di `0`, non è possibile calcolare una radice quadrata con valore reale. Perciò, solleva una `ValueError` con il messaggio `'Square root of negative number is not defined in real numbers'`. Non dimenticare di rimuovere la parola chiave `pass`.

# --hints--

Dovresti rimuovere la parola chiave `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

Dovresti sollevare una `ValueError` con il messaggio `'Square root of negative number is not defined in real numbers'` all'interno del corpo `if`.

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
