---
id: 68e6698b29e562d21b1d260c
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, vas a construir una lista de los principales navegadores web. La estructura básica de HTML ha sido proporcionada para ti.

Comienza agregando un encabezado a tu página que diga `List of Major Web Browsers` usando un elemento `h1` dentro del elemento `body`.

# --hints--

Debes tener un elemento `h1` dentro del elemento `body`.

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

Tu elemento `h1` debe contener el texto `List of Major Web Browsers`.

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
