---
id: 68ecfa1003d88990bb96a33d
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Las listas de descripción están compuestas por una lista de términos y detalles. El elemento `dt` representa el término de la descripción, mientras que el elemento `dd` representa el detalle de la descripción. 

El primer navegador que agregarás a tu lista de descripción será para Google Chrome.

Dentro de tu elemento `dl`, agrega un elemento `dt` que contenga el texto `Google Chrome`.

Debajo de tu elemento `dt`, agrega un elemento `dd` con el siguiente texto:

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

Debes tener un elemento `dt` dentro de tu elemento `dl`.

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

Tu elemento `dt` debe tener el texto `Google Chrome`.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

Debes tener un elemento `dd` debajo de tu elemento `dt`.

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

Tu elemento `dd` debe contener el texto `This is a free web browser developed by Google and first released in 2008.`

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```

