---
id: 68ecfa1003d88990bb96a33e
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Ahora es momento de agregar tu segundo navegador a la lista.

Debajo de tu entrada `Google Chrome`, agrega otro elemento `dt` que contenga el texto `Firefox`.

Debajo de tu elemento `dt`, agrega un elemento `dd` con el siguiente texto:

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

Debes tener dos elementos `dt` en tu código.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

Tu segundo elemento `dt` debe tener el texto `Firefox`.

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

Debes tener dos elementos `dd` en tu código.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

Tu segundo elemento `dd` debe contener el texto `This is a free web browser developed by the Mozilla Corporation and first created in 2004.`

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```

