---
id: 68ecfa0f03d88990bb96a33c
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Como recuerdas de una lección anterior, las listas de descripción se usan para presentar términos y definiciones en un formato organizado y fácil de leer. 

Aquí hay un ejemplo:

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

Debajo del elemento `h1`, crea un elemento `dl`. Este contendrá la lista de navegadores.

# --hints--

Debes tener un elemento `dl`.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

Tu elemento `dl` debe estar debajo del elemento `h1`.

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        --fcc-editable-region--
        <h1>List of Major Web Browsers</h1> 

        --fcc-editable-region--
    </body> 
</html>
```

