---
id: 69038f5efad6f722def4a214
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Ahora es momento de agregar tu cuarto navegador a la lista.

Agrega otro elemento `dt` que contenga el texto `Brave`.

Debajo de tu elemento `dt`, agrega un elemento `dd` con el siguiente texto:

```md
This is a free web browser first released in 2016 that is based on the Chromium web browser.
```

# --hints--

Debes tener cuatro elementos `dt` en tu código.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 4);
```

Tu cuarto elemento `dt` debe tener el texto `Brave`.

```js
const dt = document.querySelectorAll('dt')[3];
assert.strictEqual(dt?.textContent.trim(), "Brave");
```

Debes tener cuatro elementos `dd` en tu código.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 4);
```

Tu cuarto elemento `dd` debe contener el texto `This is a free web browser first released in 2016 that is based on the Chromium web browser.`.

```js
const ddElement = document.querySelectorAll('dd')[3];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser first released in 2016 that is based on the Chromium web browser.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            --fcc-editable-region--

            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
