---
id: 69038dd51a41ae215fcc030a
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

El tercer navegador que agregar치s a la lista ser치 para el navegador web Safari.

Agrega otro elemento `dt` que contenga el texto `Safari`.

Debajo de tu elemento `dt`, agrega un elemento `dd` con el siguiente texto:

```md
This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.
```

# --hints--

Debes tener tres elementos `dt` en tu c칩digo.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 3);
```

Tu tercer elemento `dt` debe tener el texto `Safari`.

```js
const dt = document.querySelectorAll('dt')[2];
assert.strictEqual(dt?.textContent.trim(), "Safari");
```

Debes tener tres elementos `dd` en tu c칩digo.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 3);
```

Tu tercer elemento `dd` debe contener el texto `This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.`.

```js
const ddElement = document.querySelectorAll('dd')[2];
assert.strictEqual(ddElement?.textContent.trim(), "This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            --fcc-editable-region--

            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
