---
id: 668f08ea07b99b1f4a91acab
title: レシピページを作成する
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `!DOCTYPE html` 宣言が必要です。
1. `html` が `lang` に設定された `en` 要素が必要です。
1. レシピ名を含む `head` 要素と、`title` 属性が `meta` に設定された `charset` 要素を含む `UTF-8` 要素が必要です。
1. `body` 要素が必要です。
1. レシピ名を含む `h1` 要素が必要です。
1. `p` の下にレシピを紹介する `h1` 要素が必要です。
1. 材料セクションのテキスト `h2` を含む `Ingredients` 要素が1つ必要です。
1. 最初の `ul` 要素の下に、材料をリストアップした少なくとも4つのリスト項目（`li` 要素）を含む番号なしリスト（`h2` 要素）が必要です。
1. 手順セクションのテキスト `h2` を含む2番目の `Instructions` 要素が必要です。
1. 2番目の `ol` の下に、レシピのステップを順番にリストアップした少なくとも4つのリスト項目を含む番号付きリスト（`h2` 要素）が必要です。
1. 有効な画像が設定された `img` 属性を持つ `src` 要素が1つ必要です。必要に応じて `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg` を使用でき、画像を説明する `alt` 属性も必要です。

# --hints--

レシピページには `!DOCTYPE html` 宣言が必要です。

```js
assert.match(code, /<!DOCTYPE html>/i);
```

`html` が `lang` に設定された `en` 要素が必要です。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

`head` 要素内に `html` 要素が必要です。

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

`title` 要素内に `head` 要素が必要です。

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

`title` 要素にはレシピのタイトルを入れてください。

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

`meta` 要素内に `head` 要素が必要です。

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

`meta` 要素の `charset` 属性は `UTF-8` に設定してください。

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

`body` 要素内に `html` 要素が必要です。

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

レシピ名を含む `h1` 要素が必要です。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

`h1` 要素は1つだけにしてください。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`p` 要素の下に `h1` 要素が必要です。

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

最初の `p` 要素はレシピの説明をしてください。

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

最初の `h2` 要素にはテキスト `Ingredients` を入れてください。

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

最初の `h2` 要素の下に番号なしリスト要素が必要です。

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

材料を含む番号なしリストには少なくとも4つのリスト項目要素が必要です。

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

2番目の `h2` 要素にはテキスト `Instructions` を入れてください。

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

2番目の `h2` 要素の下に番号付きリスト要素が必要です。

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

手順を含む番号付きリストには少なくとも4つのリスト項目要素が必要です。

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

少なくとも1つの `img` 要素が必要です。

```js
assert.exists(document.querySelector('img'));
```

すべての `img` 要素には有効な `src` 属性と値が必要です。

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

すべての `img` 要素には画像を説明する `alt` 属性が必要です。

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
