---
id: 660ee6e3a242da6bd579de69
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Um dos conceitos mais importantes em programação são as variáveis.Uma <dfn>variável</dfn> aponta para um endereço de memória específico que armazena um valor.Variáveis recebem um nome que pode ser usado em todo o seu código para acessar esse valor.

Declarar uma variável significa dar a ela um nome.Em JavaScript, isso é frequentemente feito com a palavra-chave `let`.Por exemplo, aqui está como você declararia uma variável `hello`:

```js
let hello;
```

A nomeação de variáveis segue regras específicas: os nomes podem incluir letras, números, sinais de dólar e underscores, mas não podem conter espaços e não devem começar com um número.

Use a palavra-chave `let` para declarar uma variável chamada `character`.

_Note_: It is common practice to end statements in JavaScript with a semicolon. `;`

# --hints--

Você deve usar `let` no seu código.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Você deve usar `character` no seu código.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

Você deve usar `let` para declarar uma variável `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Sua declaração deve terminar com um ponto e vírgula.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
