---
id: 660ef5105b8ba095307a0e50
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Agora registre sua variável `character` no console novamente. Você deve ver a string `"Hello"` e depois a string `"World"` no console.

# --hints--

Você deve usar `console.log` uma segunda vez.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Você deve registrar `character` no console uma segunda vez.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


Seu novo `console.log()` deve vir após sua reatribuição.

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
