---
id: 660eebd83100d37862268781
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Sua variável `character` atualmente não tem um valor.Você pode atribuir um valor usando o operador <dfn>assignment</dfn> `=`.Por exemplo:

```js
let hello = "Hello";
```

Atribuir um valor a uma variável no momento de sua declaração é conhecido como <dfn>inicialização</dfn>.

Inicialize sua variável `character` atribuindo a ela o valor `"Hello"` durante sua declaração.

# --hints--

Você deve usar o operador de atribuição `=`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

Você deve usar a string `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

Você deve usar aspas duplas para sua string `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

Você deve atribuir `"Hello"` à sua variável `character`.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
