---
id: 660ef6355e8f5a9e67fe5f46
title: Passo 16
challengeType: 1
dashedName: step-16
---

# --description--

Com o tipo de dado `number`, você pode realizar operações matemáticas como adição.Tente imprimir `count + 1` no console.

# --hints--

Você deve acessar o `console` no seu código.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Você deve acessar o `.log()` do `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Você deve registrar `count + 1` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
