---
id: 660ef55dd468079679ee0092
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

Você está pronto para declarar sua próxima variável.Remova ambas as declarações `console.log` e a reatribuição de `character`.

Também remova sua variável `secondCharacter`, mas deixe a inicialização de `character` inalterada.

# --hints--

Você não deve ter nenhuma declaração `console.log()` no seu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Você não deve reatribuir o valor de `character`.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

Sua variável `character` deve ter o valor `"Hello"`.

```js
assert.equal(character, "Hello");
```

Você não deve ter uma variável `secondCharacter`.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
