---
id: 660ef857f2806aa626d29d17
title: Passo 17
challengeType: 1
dashedName: step-17
---

# --description--

Você também pode realizar subtração (`-`), multiplicação (`*`) e divisão (`/`). Sinta-se à vontade para experimentar com os operadores e números no seu `console.log`. Quando estiver pronto para continuar, remova o `console.log`.

# --hints--

Você não deve ter um `console.log()` no seu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Sua variável `character` ainda deve ter o valor `"Hello"`.

```js
assert.equal(character, "Hello");
```

Sua variável `count` ainda deve ter o valor `8`.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
