---
id: 660ef0f7c4b8e68ccd1f0786
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

O JavaScript tem sete tipos de dados primitivos. Um deles é a String. Em JavaScript, uma <dfn>string</dfn> representa uma sequência de caracteres e pode ser delimitada por aspas simples (`'`) ou duplas (`"`).

Observe que strings são <dfn>imutáveis</dfn>, o que significa que, uma vez criadas, não podem ser alteradas. A variável ainda pode ser reatribuída a outro valor.

Altere sua string `"Hello"` para usar aspas simples.

# --hints--

Você não deve ter aspas duplas no seu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

Você deve usar aspas simples para sua string `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

Você ainda deve usar `let` no seu código.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Você ainda deve declarar uma variável `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Sua variável `character` ainda deve ter a string `"Hello"` como valor.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
