---
id: 660ef5c1904955978a986a5c
title: Passo 15
challengeType: 1
dashedName: step-15
---

# --description--

Use `let` para declarar uma variável `count`.Atribua a ele o <dfn>número</dfn> `8`.Ao usar um valor numérico, você não deve usar aspas.Por exemplo:

```js
let money = 100;
```

# --hints--

Você deve ter uma segunda palavra-chave `let` no seu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Você deve usar `let` para declarar uma variável `count`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

Você deve atribuir o número `8` à sua variável `count`.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
