---
id: 660ef19b95d3308e7dd31bb6
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

O <dfn>console</dfn> permite que você imprima e visualize a saída do JavaScript. Você pode enviar informações para o console usando `console.log()`. Por exemplo, este código irá imprimir `"Naomi"` no console:

```js
let developer = "Naomi";
console.log(developer);
```

O código acima acessa a variável `developer` pelo seu nome no `console.log()`. Observe que o valor entre os parênteses é o valor que será impresso.

Imprima o valor da variável `character` no console. Então, clique no botão "Console" para visualizar o console do JavaScript.

# --hints--

A inicialização da variável `character` não deve ser alterada.

```js
assert.strictEqual(character, 'Hello');
```

Você deve acessar o `console` no seu código.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Você deve acessar a propriedade `log` do `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Você deve usar parênteses para chamar o método `.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Você deve imprimir a variável `character` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
