---
id: 660ef31a5be625914a0102cd
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Quando uma variável é declarada com a palavra-chave `let`, você pode <dfn>reatribuir</dfn> (ou alterar o valor de) essa variável posteriormente. Neste exemplo, o valor de `programmer` é alterado de `"Naomi"` para `"CamperChan"`.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Observe que ao reatribuir uma variável, você **não** usa a palavra-chave `let` novamente.

Após seu `console.log`, atribua o valor `"World"` à sua variável `character`.

# --hints--

Você deve usar `character` duas vezes no seu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

Você deve usar o operador de atribuição para reatribuir `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

Você deve usar a string `"World"` no seu código.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

Sua variável `character` deve ter o valor `"World"`.

```js
assert.equal(character, "World");
```

Sua reatribuição não deve usar `let`.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
