---
id: 660f38c34a4de6141c0c369f
title: Passo 71
challengeType: 1
dashedName: step-71
---

# --description--

Seu gerador de pirâmide agora funciona como esperado.Mas esta é uma excelente oportunidade para explorar ainda mais o código que você escreveu.

O operador de adição não é a única forma de adicionar valores a uma variável.O operador <dfn>addition assignment</dfn> pode ser usado como uma forma abreviada para significar "pegar o valor original da variável, adicionar este valor e atribuir o resultado de volta à variável." Por exemplo, essas duas declarações gerariam o mesmo resultado:

```js
test = test + 1;
test += 1;
```

Atualize sua declaração de iteração no loop `for` para usar atribuição de adição.

# --hints--

Seu loop `for` não deve usar `i = i + 1`;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Seu loop `for` deve usar atribuição de adição com `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
