---
id: 660f0c34aad72dc712b97624
title: Passo 27
challengeType: 1
dashedName: step-27
---

# --description--

Outro método essencial para este projeto é o método `.pop()`.Remove o último elemento de um array e <dfn>retorna</dfn> esse elemento.

Quando um método retorna um valor, você pode pensar nisso como devolver o valor para você, tornando-o disponível para uso em outras partes do seu código.

Crie uma nova variável chamada `popped` e atribua a ela o resultado de `rows.pop()`.Então, registre `popped` no console.

# --hints--

Você deve declarar uma variável chamada `popped`.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

Você deve usar `let` para declarar sua variável chamada `popped`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

Você deve chamar o método `.pop()`.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

Você deve chamar o método `.pop()` no seu array `rows`. 

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

Você deve registrar sua variável `popped`. 

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
