---
id: 660f4990b1caa03b9dc97a43
title: Passo 97
challengeType: 1
dashedName: step-97
---

# --description--

Agora você não precisa mais da sua variável `done`.Remova a operação de incremento do seu loop e a declaração da variável para `done`.

# --hints--

Você não deve incrementar a variável `done`.

```js
assert.notMatch(__helpers.removeJSComments(code), /done\+\+/);
```

Você não deve mais ter uma variável `done`.

```js
assert.notMatch(__helpers.removeJSComments(code), /done/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (rows.length < count) {
  done++;
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
