---
id: 660f1cedf3676fe26122ebf6
title: Passo 38
challengeType: 1
dashedName: step-38
---

# --description--

Você deve ver os números zero a sete impressos no seu console, um por linha.Isto servirá como a base para gerar sua pirâmide.

Substitua sua instrução de log por uma instrução para adicionar `i` ao seu array `rows`.

# --hints--

Você não deve ter uma chamada `console.log`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Você deve chamar `.push()` no seu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Você deve adicionar `i` ao seu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Seu `.push()` deve acontecer no seu `for` loop. 

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
