---
id: 660f374d532dc41189cc9cc2
title: Passo 69
challengeType: 1
dashedName: step-69
---

# --description--

Agora é hora de um pouco de matemática.Considere uma pirâmide de três linhas.Se quisermos que fique centralizado, ficaria mais ou menos assim:

```md
··#··
·###·
#####
```

Espaços vazios foram substituídos por interpuncts, ou pontos médios, para melhor legibilidade.Se você extrapolar o padrão, pode ver que os espaços no início e no final de uma linha seguem um padrão.

Atualize suas strings de espaço em branco para serem repetidas `rowCount - rowNumber` vezes.

Abra o console para ver o resultado.

# --hints--

Você deve chamar `.repeat()` nas suas strings `" "` para repeti-las `rowCount - rowNumber` vezes.

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
