---
id: 660f09a2694b59c3a10ee304
title: Passo 23
challengeType: 1
dashedName: step-23
---

# --description--

Por enquanto, remova seu primeiro console log e sua atribuição `rows[rows.length - 1]`. Deixe a segunda declaração de log `rows` para depois.

# --hints--

Você deve remover sua instrução `console.log(rows[0])`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

Você deve remover a reatribuição de `rows[rows.length - 1]`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

Você não deve remover sua declaração `console.log(rows)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
