---
id: 660f415b76859a2736771607
title: Passo 78
challengeType: 1
dashedName: step-78
---

# --description--

Sua pirâmide desapareceu novamente.Tudo bem - isso é esperado.

Antes de criar seu novo loop, você precisa aprender sobre as declarações `if`.Uma <dfn>`if` statement</dfn> permite que você execute um bloco de código somente quando uma condição for atendida.Eles usam a seguinte sintaxe:

```js
if (condition) {
  logic
}
```

Crie uma declaração `if` com o booleano `true` como condição.No corpo, imprima a string `"Condition is true"`.

# --hints--

Você deve criar uma declaração `if`.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

Sua declaração `if` deve ter `true` como condição.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

O corpo do seu `if` deve registrar `"Condition is true"`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
