---
id: 660f540c2176ea5dec01306d
title: Passo 118
challengeType: 1
dashedName: step-118
---

# --description--

Bom trabalho! Experimente diferentes valores para suas variáveis `character`, `count` e `inverted`.

Quando estiver pronto para passar para o seu próximo projeto, defina `character` como `"!"`, `count` como `10` e `inverted` como `false` para continuar.

Parabéns por completar seu primeiro projeto em JavaScript!

# --hints--

Você deve definir `character` como `"!"`.

```js
assert.equal(character, "!");
```

Você deve definir `count` para `10`.

```js
assert.equal(count, 10);
```

Você deve definir `inverted` como `false`.

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
