---
id: 660f46460f9c36330ebc07d8
title: Passo 88
challengeType: 1
dashedName: step-88
---

# --description--

O operador de igualdade pode levar a alguns comportamentos estranhos em JavaScript. Por exemplo, `"0" == 0` é verdadeiro, mesmo que um seja uma string e o outro seja um número.

O operador <dfn>strict equality</dfn> `===` é usado para verificar se dois valores são iguais e compartilham o mesmo tipo. Como regra geral, este é o operador de igualdade que você deve sempre usar. Com o operador de igualdade estrita, `"0" === 0` se torna falso, porque embora possam ter o mesmo valor zero, eles não são do mesmo tipo.

Atualize sua condição `done == count` para usar o operador de igualdade estrita.

# --hints--

Sua condição `if` deve usar igualdade estrita.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
