---
id: 660f1bf673487ae0bb25b900
title: Passo 37
challengeType: 1
dashedName: step-37
---

# --description--

Seu loop agora deve executar oito vezes.Dentro do corpo do loop, imprima o valor do iterador `i` e veja o que acontece.

# --hints--

Você deve usar `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

Você deve registrar o valor de `i`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

Você deve registrar o valor de `i` no seu loop `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
