---
id: 660f39b444fd6f16d1e49c1f
title: Passo 73
challengeType: 1
dashedName: step-73
---

# --description--

Ao invés de precisar passar `i + 1` para sua chamada de `padRow`, você pode começar seu loop em `1`. Isso permitiria que você criasse um loop indexado a partir de 1.

Atualize seu iterador para começar em `1` ao invés de `0`.

# --hints--

Seu loop `for` deve inicializar `i` em `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
