---
id: 660f23b53db70af0f2620e78
title: Passo 43
challengeType: 1
dashedName: step-43
---

# --description--

Agora todos os seus números estão aparecendo na mesma linha.Isso não funcionará para criar uma pirâmide.

Você precisará adicionar uma nova linha a cada linha.No entanto, pressionar a tecla return para inserir uma quebra de linha entre aspas em JavaScript resultará em um erro de análise.Em vez disso, você precisa usar a <dfn>sequência de escape</dfn> especial `\n`, que é interpretada como uma nova linha quando a string é registrada.Por exemplo:

```js
lineOne = lineOne + "\n" + lineTwo;
```

Use um segundo operador de adição para adicionar uma nova linha após os valores `result` e `row`.

# --hints--

Você deve usar a sequência de escape `\n`.Lembre-se de que precisa ser uma string, então está envolvida em aspas.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

Você deve concatenar sua variável `row` à sua variável `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

Você deve concatenar a sequência de escape `\n` à sua variável `row`.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

Você deve atribuir toda a concatenação de volta à sua variável `result`.

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
