---
id: 660f4455f457ef2e3ec6920f
title: Passo 80
challengeType: 1
dashedName: step-80
---

# --description--

Agora a string não está mais sendo impressa, porque `false` não é `true`. Mas e quanto a outros valores?

Tente alterar a condição para a string `"false"`.

# --hints--

Sua declaração `if` deve ter a string `"false"` como condição.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*('|")false\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (false) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
