---
id: 660f061d259bbebc37461080
title: Passo 20
challengeType: 1
dashedName: step-20
---

# --description--

Você pode acessar os valores dentro de um array usando o <dfn>index</dfn> do valor. Um índice é um número que representa a posição do valor no array, começando em `0` para o primeiro valor.

Você pode acessar o valor usando <dfn>notação de colchetes</dfn>, como `array[0]`.

Use `console.log` e a notação de colchetes para imprimir o primeiro valor no seu array `rows`.

# --hints--

Você deve ter uma instrução de `console.log()` no código.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Você deve acessar seu array `rows`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

Você deve usar a notação de colchetes com seu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

Você deve usar a notação de colchetes para acessar o primeiro elemento do seu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

Você deve registrar o primeiro elemento do seu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
