---
id: 660f530d6e33d159e1bf4947
title: Passo 115
challengeType: 1
dashedName: step-115
---

# --description--

Se sua pirâmide não estiver invertida, então você vai querer ter um bloco `else` que construa a pirâmide na ordem normal.

Nos passos anteriores, você aprendeu como trabalhar com a declaração `else` assim:

```js
if (condition) {
  // if condition is true, run this code
} else {
  // if condition is false, run this code
}
```

Adicione um bloco `else` ao seu bloco `if`.

# --hints--

Você deve adicionar um bloco `else`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)\s*\{\s*rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)\s*\);\s*\}\s*else\s*\{/);
```

Seu bloco `else` deve estar vazio.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)\s*\{\s*rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)\s*\);\s*\}\s*else\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  }

}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
