---
id: 660f48a419b40238e2b8b4d5
title: Passo 93
challengeType: 1
dashedName: step-93
---

# --description--

Seu loop não depende mais da variável `continueLoop`.Isso torna a variável uma <dfn>declaração não utilizada</dfn>.Geralmente, você quer evitar declarações não utilizadas para prevenir confusão futura.

Remova sua variável `continueLoop`.

# --hints--

Você não deve mais ter uma variável `continueLoop`.

```js
assert.notMatch(__helpers.removeJSComments(code), /continueLoop/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
