---
id: 660f165270622fd4ec0da3f7
title: Passo 31
challengeType: 1
dashedName: step-31
---

# --description--

Declarar uma variável com a palavra-chave `let` permite que ela seja reatribuída. Isso significa que você poderia alterar `character` depois para ser um valor completamente diferente.

Para este projeto, você não vai querer alterar esses valores de variável. Então, em vez disso, você deve usar `const` para declará-las. Variáveis `const` são especiais.

Primeiro, uma variável `const` não pode ser reatribuída como uma variável `let`. Este código geraria um erro:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

Uma variável `const` também não pode ficar sem inicialização. Este código geraria um erro:

```js
const firstName;
```

Substitua suas palavras-chave `let` por `const`.

# --hints--

Você deve usar `const` para declarar sua variável `character`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

Você deve usar `const` para declarar sua variável `count`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

Você deve usar `const` para declarar sua variável `rows`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

Você não deve usar `let` no seu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
