---
id: 660f1b6e60bd9edf902c81fd
title: Passo 36
challengeType: 1
dashedName: step-36
---

# --description--

Sua declaração <dfn>iteration</dfn> dirá ao seu loop o que fazer com o iterator após cada execução.

Quando você reatribui uma variável, pode usar a variável para referenciar o valor anterior antes da reatribuição.Isso permite que você faça coisas como adicionar três a um número existente.Por exemplo, `bees = bees + 3` aumentaria o valor de `bees` em três.

Use essa sintaxe para substituir sua string `"iteration"` por uma instrução de reatribuição que aumenta `i` em um.

# --hints--

Você deve adicionar um ao seu `i` variável.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

Você deve atribuir `i + 1` de volta à sua variável `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Seu loop `for` deve aumentar `i` em `1` a cada iteração.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
