---
id: 660f4cde8dd305450514a1cb
title: Passo 105
challengeType: 1
dashedName: step-105
---

# --description--

Use um comentário de múltiplas linhas para comentar este loop também, para preparar para a próxima abordagem.

# --hints--

Seu loop `for` deve ser comentado usando comentário multilinha.

```js
assert.lengthOf(code.match(/\/\*\s*for\s*\(/g), 2);
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 1);
```

O corpo do seu loop `for` deve ser comentado como parte de um comentário multilinha.

```js
assert.notMatch(code, /\/\/\s*rows\.push/);
assert.lengthOf(code.match(/rows\.push/g), 3)
assert.notMatch(__helpers.removeJSComments(code), /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
