---
id: 660f4f79e2a82a4e92290f44
title: Passo 110
challengeType: 1
dashedName: step-110
---

# --description--

Às vezes você pode querer recuperar código anterior que você comentou.Você pode fazer isso removendo o `/*` e o `*/` ao redor desse código.Isto é chamado de <dfn>uncommenting</dfn>.

Descomente apenas seu primeiro loop `for`.Deixe o comentário de linha única e os outros dois comentários de múltiplas linhas no lugar.

# --hints--

Você deve descomentar seu primeiro loop `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

Você não deve remover seu comentário de linha única.

```js
assert.match(code, /\/\//);
```

Você não deve descomentar seu loop `while`.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Você não deve descomentar seu segundo loop `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
