---
id: 660f20473aef47e9b8c9afc6
title: Passo 40
challengeType: 1
dashedName: step-40
---

# --description--

Adicione uma instrução de log para imprimir o valor de `result`.Dependendo do console que você usar, pode ser que nada seja exibido.

# --hints--

Você deve adicionar uma chamada `console.log`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Você deve registrar sua variável `result`. 

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
