---
id: 660f2b6fd54ac1fc142804dd
title: Passo 47
challengeType: 1
dashedName: step-47
---

# --description--

A lógica para formatar esta pirâmide provavelmente vai ficar complicada, o que significa que é um ótimo momento para extrair esse código para uma função.

Uma <dfn>function</dfn> é um bloco de código que pode ser reutilizado em toda a sua aplicação.Funções são declaradas com a seguinte sintaxe:

```js
function name(parameter) {

}
```

A palavra-chave `function` informa ao JavaScript que a variável `name` será uma função. `parameter` é uma variável que representa um valor que é passado para a função quando ela é usada.Uma função pode ter tantos, ou tão poucos, <dfn>parameters</dfn> quanto você quiser.Como em um loop `for`, o espaço entre as chaves é o <dfn>corpo da função</dfn>.

Declare uma função `padRow`.Não crie nenhuma variável de parâmetro ainda.O corpo da função deve estar vazio.Lembre-se de que você precisa usar camel case para sua convenção de nomenclatura.

# --hints--

Você deve usar a palavra-chave `function`.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

Você deve declarar uma função `padRow`.

```js
assert.isFunction(padRow);
```

Sua função `padRow()` não deve ter nenhum parâmetro.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

Sua função `padRow()` deve ter um corpo vazio.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
