---
id: 660f383d4c772c12ff59904b
title: Passo 70
challengeType: 1
dashedName: step-70
---

# --description--

Você pode passar expressões completas como argumento.A função receberá o resultado da avaliação dessa expressão.Por exemplo, essas duas chamadas de função abaixo gerariam o mesmo resultado. (Observe que a regra da ordem das operações PEMDAS—Parenthesis, Exponents, Multiplication, Division, Addition, Subtraction—se aplica):

```js
test(2 * 3 + 1);
test(7);
```

Olhando para o padrão novamente:

```md
··#··
·###·
#####
```

Atualize o valor de `character` para ser repetido `2 * rowNumber - 1` vezes.

Abra o console novamente para ver o resultado atualizado.

# --hints--

Sua função deve passar `2 * rowNumber - 1` para sua chamada `.repeat()`.

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
