---
id: 660f229d2dbe09ef2954a4a1
title: Passo 42
challengeType: 1
dashedName: step-42
---

# --description--

Lembre-se no seu loop anterior que você usou o operador de adição `+` para aumentar o valor de `i` em `1`.

Você pode fazer algo semelhante com um valor string, adicionando uma nova string a uma string existente.Por exemplo, `hello = hello + " World";` adicionaria a string `" World"` à string existente armazenada na variável `hello`.Isto é chamado <dfn>concatenation</dfn>.

No seu loop `for...of`, use o operador de adição para concatenar o valor `row` ao valor `result`.

# --hints--

Você deve usar o operador de concatenação na sua variável `result`.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

Você deve concatenar `row` à sua variável `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

Você deve atribuir o resultado da sua concatenação de volta à variável `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
