---
id: 660f53ad3d39175c5d4335ac
title: Passo 117
challengeType: 1
dashedName: step-117
---

# --description--

Seu gerador de pirâmide agora está em um estado finalizado, com mais funcionalidades do que você planejou originalmente! O próximo passo é limpar seu código.

Remova todos os comentários, tanto de linha única quanto multilinha, do seu código.

# --hints--

Você não deve ter nenhum comentário de linha única no seu código.

```js
assert.notMatch(code, /\/\//);
```

Você não deve ter nenhum comentário multilinha no seu código.

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

Você não deve ter nenhum comentário no seu código.

```js
assert.equal(code, __helpers.removeJSComments(code));
```

Você deve remover o código que foi comentado por comentários multilinha.

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
