---
id: 660f0da9bf1035c9097af20a
title: Passo 28
challengeType: 1
dashedName: step-28
---

# --description--

Você deveria ter visto `"freeCodeCamp"` impresso no console. Isso acontece porque `.pop()` retorna o valor que foi removido do array e você adicionou `"freeCodeCamp"` ao final do array anteriormente.

Mas o que `.push()` retorna? Atribua seu `rows.push()` existente a uma nova variável `pushed` e registre-a.

# --hints--

Você deve declarar uma variável `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

Você deve usar `let` para declarar sua variável `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

Você deve atribuir `rows.push("freeCodeCamp")` à sua variável `pushed`.

```js
assert.equal(pushed, 4);
```

Você deve registrar sua variável `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
