---
id: 660f3915b41a441537ec9f5e
title: Passo 72
challengeType: 1
dashedName: step-72
---

# --description--

Porque você está apenas aumentando `i` em `1`, você pode usar o <dfn>operador de incremento</dfn> `++`. Este operador aumenta o valor de uma variável em 1, atualizando a atribuição para essa variável. Por exemplo, `test` se tornaria `8` aqui:

```js
let test = 7;
test++;
```

Substitua sua atribuição de adição pelo operador de incremento para a iteração do seu loop.

# --hints--

Seu loop `for` não deve usar atribuição de adição com `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

Seu loop `for` deve usar o operador de incremento em `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
