---
id: 660f17d4e9f227d86e834abd
title: Passo 33
challengeType: 1
dashedName: step-33
---

# --description--

Para gerar uma pirâmide, você precisará criar múltiplas linhas.Quando você precisa executar uma tarefa repetidamente até que uma condição seja atendida, você usará um <dfn>loop</dfn>.Existem muitas maneiras de escrever um loop.

Você vai começar com um `for` loop básico. `for` loops usam a seguinte sintaxe:

```js
for (iterator; condition; iteration) {
  logic;
}
```

Nos próximos passos, você explorará cada componente de um loop em detalhes.Por enquanto, construa um loop `for` que inclua os termos `"iterator"`, `"condition"` e `"iteration"` para os três componentes.Mantenha o <dfn>body</dfn> do loop, a seção dentro das chaves `{}`, vazia.


# --hints--

Você deve ter um loop `for`.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

O primeiro componente do seu loop `for` deve ser a string `"iterator"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

O segundo componente do seu loop `for` deve ser a string `"condition"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

O terceiro componente do seu loop `for` deve ser a string `"iteration"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

O corpo do seu loop `for` deve estar vazio.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
