---
id: 660f255022991ef34ed0ee88
title: Passo 44
challengeType: 1
dashedName: step-44
---

# --description--

Imprimir números não resultará em uma pirâmide visualmente atraente.Agora que você está exibindo o conteúdo formatado do seu array `rows`, é hora de atualizar seu loop original.

Em vez de adicionar `i` ao array, adicione o valor da sua variável `character`.

# --hints--

Você não deve mais incrementar sua variável `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Você deve atribuir valor à sua variável `character`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
