---
id: 660f2eccfe3f820304af1b39
title: Passo 64
challengeType: 1
dashedName: step-64
---

# --description--

Para saber como formatar uma linha, sua função `padRow` precisará saber em qual número de linha você está e quantas linhas no total estão sendo geradas.

A melhor maneira de fazer isso é criando parâmetros de função para eles. Dê à sua função `padRow` um parâmetro `rowNumber` e `rowCount`. Múltiplos parâmetros são separados por vírgula:

```js
function name(first, second) {

}
```

# --hints--

Sua função `padRow` deve ter um parâmetro `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

Você deve adicionar uma vírgula após o seu parâmetro `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

Sua função `padRow` deve ter um parâmetro `rowCount`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
