---
id: 660f447efc0e722f016c1be0
title: Passo 81
challengeType: 1
dashedName: step-81
---

# --description--

O texto reapareceu!Isso acontece porque `"false"` é uma string que, quando avaliada como um booleano, se torna `true`.Isso significa que `"false"` é um valor truthy.

Um <dfn>truthy value</dfn> é um valor que é considerado verdadeiro quando avaliado como booleano.A maioria dos valores que você encontrará em JavaScript será truthy.

Um <dfn>falsy value</dfn> é o oposto - um valor considerado falso quando avaliado como booleano.JavaScript tem uma lista definida de valores falsy.Alguns deles incluem `false`, `0`, `""`, `null`, `undefined` e `NaN`.

Tente alterar sua condição `if` para uma string vazia `""`, que é um valor falsy.

# --hints--

Sua declaração `if` deve ter `""` como condição.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
