---
id: 660f3dd626be3a1ffe27e5d1
title: Passo 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript também tem suporte para comentários de múltiplas linhas. Um comentário de múltiplas linhas começa com `/*` e termina com `*/`.

Ao contrário de um comentário de linha única, um comentário de múltiplas linhas encapsulará várias linhas.

Use `/*` e `*/` para transformar seu `for` atual, incluindo o corpo, em um comentário de múltiplas linhas.

# --hints--

Você deve iniciar um comentário de múltiplas linhas com `/*`.

```js
assert.match(code, /\/\*/);
```

Você deve terminar um comentário de múltiplas linhas com `*/`.

```js
assert.match(code, /\*\//);
```

Todo o seu loop `for` deve estar comentado.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
