---
id: 660f48e1d3682f39e81843c4
title: Passo 94
challengeType: 1
dashedName: step-94
---

# --description--

Seu gerador de pirâmide ainda está funcionando. No entanto, pode ser possível acabar com um loop infinito novamente.

Porque você está apenas verificando se `done` não é igual a `count`, se `done` fosse **maior** que `count` seu loop continuaria para sempre.

Atualize a condição do seu loop para verificar se `done` é menor ou igual a `count`.

# --hints--

Seu loop `while` deve verificar se `done` é menor ou igual a `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
