---
id: 660f4a83373de83ea101685f
title: Passo 100
challengeType: 1
dashedName: step-100
---

# --description--

Porque você vai fazer o loop na direção oposta, seu loop precisa rodar enquanto `i` for maior que `0`. Você pode usar o operador <dfn>greater than</dfn> `>` para isso.

Defina a condição do seu loop para executar quando `i` for maior que `0`.

# --hints--

Seu loop `for` deve executar quando `i` for maior que `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; false; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
