---
id: 660f4cffb1459d45e34902d1
title: Passo 106
challengeType: 1
dashedName: step-106
---

# --description--

Você pode realmente construir a pirâmide invertida sem precisar fazer um loop "para trás" como você fez.

Para fazer isso, você precisará aprender alguns novos métodos de array.Comece usando `const` para declarar uma variável `numbers`.Atribua a ele um array com os elementos `1`, `2` e `3`.Então registre o array `numbers`.

# --hints--

Você deve usar `const` para declarar uma variável `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

Sua variável `numbers` deve ser um array.

```js
assert.isArray(numbers);
```

Seu array `numbers` deve conter os elementos `1`, `2` e `3` nessa ordem.

```js
assert.deepEqual(numbers, [1,2,3]);
```

Você deve registrar seu array `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
