---
id: 660f44f10ea40f300b896a5e
title: Etapa 84
challengeType: 1
dashedName: step-84
---

# --description--

Agora que você praticou trabalhar com declarações `if...else if...else`, você pode removê-las do seu código.

Depois de concluir isso, use `let` para declarar uma variável `continueLoop` e atribuir a ela o booleano `false`.Então use `let` para declarar uma variável `done` e atribuir a ela o valor `0`.

# --hints--

Você não deve ter uma declaração `if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

Você não deve ter uma declaração `else if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Você não deve ter uma declaração `else`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

Você deve usar `let` para declarar uma variável `continueLoop`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

Sua variável `continueLoop` deve ter o valor `false`.

```js
assert.isFalse(continueLoop);
```

Você deve usar `let` para declarar uma variável `done`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

Sua variável `done` deve ter o valor `0`.

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
