---
id: 660f207334fabaeac3269c38
title: Passo 41
challengeType: 1
dashedName: step-41
---

# --description--

Para manipular a string `result`, você usará um tipo diferente de loop. Especificamente, um loop `for...of`, que itera sobre cada item em um objeto iterável e o atribui temporariamente a uma variável.

A sintaxe para um loop `for...of` é:

```js
for (const value of iterable) {

}
```

Observe que você pode usar `const` porque a variável existe apenas para uma única iteração, não durante todo o loop.

Crie um loop `for...of` para iterar pelo seu array `rows`, atribuindo cada valor a uma variável `row`.

# --hints--

Você deve usar outra palavra-chave `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

Você deve declarar uma variável `row`.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

Você deve usar `const` para declarar sua variável `row`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

Seu loop `for...of` deve declarar sua variável `row`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

Sua variável `row` deve ser extraída de `rows` usando a palavra-chave `of`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

O corpo do seu loop `for...of` deve estar vazio.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
