---
id: 660f3ce51f70571e1c5227c8
title: Passo 76
challengeType: 1
dashedName: step-76
---

# --description--

Comentários podem ser úteis para explicar por que seu código adota uma determinada abordagem ou para deixar notas de tarefas para seu eu futuro.

Em JavaScript, você pode usar `//` para deixar um comentário de uma única linha no seu código.

Adicione um comentário de linha única acima do seu `for` para lembrar-se de alterar o código para um tipo diferente de loop.

# --hints--

Você deve iniciar um comentário de linha única com `//`.

```js
assert.match(code, /\/\//);
```

Seu comentário de linha única deve ter pelo menos cinco caracteres.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
