---
id: 660f17294346b7d69e79db3d
title: Passo 32
challengeType: 1
dashedName: step-32
---

# --description--

Você está pronto para começar a construir seu gerador de pirâmide.Sua variável `character` servirá como o bloco de construção para a pirâmide.

"Hello" pode não funcionar muito bem para isso.Altere o valor de `character` para ser o caractere hash `"#"`.

# --hints--

Sua variável `character` deve ser um símbolo de hash.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
