---
id: 660f34626216270c682e2f7b
title: Passo 66
challengeType: 1
dashedName: step-66
---

# --description--

Uma <dfn>chamada de função</dfn> permite que você realmente use uma função. Você pode não ter percebido, mas os métodos como `.push()` que você tem usado são chamadas de função.

Uma função é chamada referenciando o nome da função e adicionando `()`. Aqui está como chamar uma função `test`:

```js
test();
```

Substitua o `character.repeat(i + 1)` na sua chamada `.push()` por uma chamada de função para sua função `padRow`. Ainda não adicione nenhum argumento a ele.

# --hints--

Você não deve usar `i + 1` na sua chamada `push`.

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

Você não deve usar `character.repeat` na sua chamada `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

Você deve chamar `padRow` na sua chamada `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

Você não deve ter argumentos na sua chamada `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
