---
id: 660f4de78f775e480ba2e451
title: Passo 107
challengeType: 1
dashedName: step-107
---

# --description--

O método `.unshift()` de um array permite adicionar um valor ao **início** do array, diferente do `.push()` que adiciona o valor no final do array. `.unshift()` retorna o novo tamanho do array no qual foi chamado.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

Use `const` para declarar uma variável `unshifted` e atribua a ela o resultado de chamar `.unshift()` no seu array `numbers`.Passe `5` como argumento.Então imprima sua variável `unshifted`.

# --hints--

Você deve declarar sua variável `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

Você deve chamar o método `.unshift()` no seu array `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

Você deve passar `5` como argumento para sua chamada `.unshift()`.

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

Você deve atribuir o resultado da sua chamada `.unshift()` à sua variável `unshifted`.

```js
assert.equal(unshifted, 4);
```

Você deve registrar sua variável `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
