---
id: 660f3ba3cceef11b6ba08b59
title: Passo 75
challengeType: 1
dashedName: step-75
---

# --description--

Infelizmente, agora a base da pirâmide desapareceu. Isso ocorre porque você criou outro <dfn>off-by-one error</dfn>.

Seu loop original percorreu valores de `i` de `0` a `7`, porque `count` é `8` e sua condição exige que `i` seja menor que `count`. Seu loop agora está executando para valores de `i` de `1` a `7`.

Seu loop precisa ser atualizado para rodar quando `i` for `8`, também. Olhando para sua lógica, isso significa que seu loop deve rodar quando `i` for <dfn>menor ou igual a</dfn> `count`. Você pode usar o operador menor ou igual a `<=` para isso.

Atualize a condição do seu loop para executar enquanto `i` for menor ou igual a `count`.

# --hints--

Seu loop `for` não deve verificar se `i` é menor que `count`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Seu loop `for` deve verificar se `i` é menor ou igual a `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
