---
id: 660f4e74f7fd3f4a99ac2e50
title: Passo 108
challengeType: 1
dashedName: step-108
---

# --description--

Arrays também têm um método `.shift()`.Isso removerá o **primeiro** elemento do array, diferente de `.pop()` que remove o último elemento.Aqui está um exemplo do método `.shift()`:

```js
const numbers = [1, 2, 3];
numbers.shift();
```

O array `numbers` seria `[2, 3]`.

Diretamente abaixo do seu array `numbers`, declare uma variável `shifted` e atribua a ela o resultado de chamar `.shift()` no array `numbers`.Na próxima linha, registre a variável `shifted` no console.

# --hints--

Você deve usar `const` para declarar uma variável `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

Você deve chamar `.shift()` no seu array `numbers`.

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

Você deve atribuir o resultado da sua chamada `.shift()` à sua variável `shifted`.

```js
assert.equal(shifted, 1);
```

Você deve registrar sua variável `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
