---
id: 660f49e32001983c90b75850
title: Passo 98
challengeType: 1
dashedName: step-98
---

# --description--

Esse é um loop muito limpo e funcional.  Nice work! Mas ainda há mais para explorar.

Use um comentário de múltiplas linhas para comentar seu loop `while`.

# --hints--

Seu loop `while` deve estar comentado.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

O corpo do seu loop `while` deve estar comentado.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
