---
id: 660f45ccf4ca5c31f253005a
title: Passo 86
challengeType: 1
dashedName: step-86
---

# --description--

Neste momento, se você alterar `continueLoop` para true, seu loop `while` irá rodar para sempre. Isso é chamado de <dfn>infinite loop</dfn> e você deve ter cuidado para evitar esses casos. Um loop infinito pode travar seu sistema, exigindo uma reinicialização completa para escapar.

Para evitar isso, comece usando o operador de incremento para aumentar o valor da variável `done` dentro do seu loop.

# --hints--

Seu loop `while` deve incrementar a variável `done`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
