---
id: 660f4b641290da41b2cf0dd9
title: Passo 103
challengeType: 1
dashedName: step-103
---

# --description--

Assim como na adição, existem diferentes operadores que você pode usar para subtração.O operador <dfn>subtraction assignment</dfn> `-=` subtrai o valor fornecido do valor atual da variável, e então atribui o resultado de volta à variável.

Substitua sua instrução de iteração pela instrução correta usando o operador de atribuição de subtração.

# --hints--

Seu loop `for` não deve usar `i = i - 1`.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /for\(leti=count;i>0;i=i-1\)\{rows\.push\(padRow\(i,count\)\);/);
```

Seu loop `for` deve usar atribuição de subtração para reduzir `i` em `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*-=\s*1\s*\)\s*\{\s*rows\.push\(\s*padRow\s*\(\s*i\s*,\s*count\s*\)\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i = i - 1) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
