---
id: 660f0f980e98e8cf77f1ce31
title: Passo 30
challengeType: 1
dashedName: step-30
---

# --description--

Altere sua declaração de `rows` para ser atribuída a um array vazio novamente.

Além disso, altere sua string `'Hello'` para usar aspas duplas novamente.Geralmente, não importa qual dos dois você prefira, mas você vai querer ser consistente nessa escolha ao longo do seu projeto.

# --hints--

Seu array `rows` deve estar vazio.

```js
assert.empty(rows);
```

Sua string `"Hello"` deve usar aspas duplas.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
