---
id: 660f0ee51d7460ce88cd248d
title: Passo 29
challengeType: 1
dashedName: step-29
---

# --description--

Você esperava ver `4` no console? `.push()` retorna o novo tamanho do array depois de adicionar o valor que você fornece.

É importante estar ciente dos valores que um método retorna.Reserve um tempo para experimentar com `.push()` e `.pop()`.Quando estiver pronto, remova todas as suas chamadas `.push()` e `.pop()` e suas declarações `console.log`.

# --hints--

Você não deve ter uma chamada `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

Você não deve ter uma chamada `.pop()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

Você não deve ter nenhuma instrução de log.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

Você não deve ter uma variável `popped`.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

Você não deve ter uma variável `pushed`.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
