---
id: 660f280dda5040f707c76b4a
title: Passo 45
challengeType: 1
dashedName: step-45
---

# --description--

Agora você tem uma série de caracteres `#`, mas a forma da pirâmide ainda está faltando.Felizmente, a variável `i` representa o número da "linha" atual no seu loop, permitindo que você a use para criar uma estrutura em forma de pirâmide.

Para conseguir isso, você usará o método `.repeat()` disponível para strings.Este método aceita um número como argumento, especificando o número de vezes para repetir a string alvo.Por exemplo, usar `.repeat()` para gerar a string `"Code!Code!Código!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

Use o método `.repeat()` no seu `character` e passe `i` como número.

# --hints--

Você deve usar o método `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

Você deve usar o método `.repeat()` na sua variável `character`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

Você deve passar `i` para o seu método `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

Você deve usar o método `.repeat()` no método `.push()`

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
