---
id: 660f0a55847d6cc485f29ba5
title: Passo 26
challengeType: 1
dashedName: step-26
---

# --description--

Um <dfn>method</dfn> em JavaScript é uma função que está associada a certos valores ou objetos.Um exemplo que você já encontrou é o método `.log()`, que faz parte do objeto `console`.

Arrays têm seus próprios métodos e o primeiro que você vai explorar é o método `.push()`.Isso permite que você "push" um valor para o final de um array.Aqui está um exemplo para adicionar o número `12` ao final de um array:

```js
array.push(12);
```

Use `.push()` para adicionar a string `"freeCodeCamp"` ao final do seu array `rows`.Adicione este código antes do seu `console.log` para que você possa ver a alteração que fez no seu array.

# --hints--

Você deve usar `.push()` no seu código.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

Você deve usar o método `.push()` do seu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Você deve passar a string `"freeCodeCamp"` para o seu método `.push()`. 

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
