---
id: 660f1e3f047bf4e403268713
title: Passo 39
challengeType: 1
dashedName: step-39
---

# --description--

Infelizmente, agora você não pode ver o que seu loop está fazendo.

Use `let` para declarar uma variável `result` e atribuir a ela uma string vazia.Uma string vazia é representada por aspas sem nada entre elas, como `""`.

# --hints--

Você deve declarar uma variável `result`.

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

Você deve usar `let` para declarar sua variável `result`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

Sua variável `result` deve ser uma string vazia.

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
