---
id: 660f2a70ad6225fa503e71c3
title: Passo 46
challengeType: 1
dashedName: step-46
---

# --description--

Você está chegando mais perto! Neste ponto, você está encontrando o que é conhecido como um <dfn>off-by-one error</dfn>, um problema frequente em linguagens de indexação baseada em zero como JavaScript.

O primeiro índice do seu array `rows` é `0`, por isso você começa seu loop `for` com `i = 0`. Mas repetir uma string zero vezes resulta em nada para imprimir.

Para corrigir isso, adicione `1` ao valor de `i` na sua chamada `.repeat()`. Não atribua de volta a `i` como você fez nas condições do seu loop.

# --hints--

Você deve adicionar `1` a `i` no seu método `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
