---
id: 660f3b664421471aa595170f
title: Passo 74
challengeType: 1
dashedName: step-74
---

# --description--

A pirâmide parece um pouco estranha agora. Porque você está começando o loop em `1` ao invés de `0`, você não precisa adicionar um a `i` quando passá-lo para `padRow`.

Atualize o primeiro argumento da sua chamada `padRow` para ser `i`.

# --hints--

Você não deve passar `i + 1` para `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

Você deve passar `i` para `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
