---
id: 660f455b044d3230ed971e98
title: Passo 85
challengeType: 1
dashedName: step-85
---

# --description--

Um loop <dfn>`while`</dfn> será executado repetidamente até que a `condition` especificada não seja mais verdadeira.Possui a seguinte sintaxe:

```js
while (condition) {
  logic;
}
```

Use essa sintaxe para declarar um loop `while` com `continueLoop` como condição.O corpo deve estar vazio.

# --hints--

Você deve usar um laço `while`.

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

Seu loop `while` deve usar `continueLoop` como condição.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
