---
id: 660f50a21fe7645252804f2b
title: Passo 112
challengeType: 1
dashedName: step-112
---

# --description--

E se você tivesse uma maneira de alternar entre uma pirâmide invertida e uma pirâmide padrão?

Comece declarando uma variável `inverted` e atribuindo a ela o valor `true`. Você não está alterando essa variável no seu código, mas precisará usar `let` para que nossos testes possam modificá-la depois.

# --hints--

Você deve declarar uma variável `inverted` com `let`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

Você deve inicializar `inverted` com o valor `true`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
