---
id: 660f2fbd45b520046cac68e8
title: Passo 65
challengeType: 1
dashedName: step-65
---

# --description--

Lembre-se que em uma etapa anterior, você aprendeu sobre valores de retorno.Uma função pode <dfn>return</dfn> um valor para sua aplicação consumir separadamente.

Em uma função, a palavra-chave `return` é usada para especificar um valor de retorno.Por exemplo, esta função retornaria o valor dado ao primeiro parâmetro:

```js
function name(parameter) {
  return parameter;
}
```

Use a palavra-chave `return` para retornar o valor da variável `character`, repetida `rowNumber` vezes.

# --hints--

Você deve usar o método `.repeat()`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

Você deve usar o método `.repeat()` na sua variável `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

Você deve passar `rowNumber` para sua chamada `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

Você deve usar a palavra-chave `return`.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

Você deve retornar o resultado da sua chamada `.repeat()`.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
