---
id: 660f34e99571070d56d2f231
title: Passo 67
challengeType: 1
dashedName: step-67
---

# --description--

Sua função `padRow` tem dois parâmetros que você definiu. Os valores são fornecidos para esses parâmetros quando uma função é chamada.

Os valores que você fornece para uma chamada de função são chamados de <dfn>arguments</dfn> e você <dfn>pass</dfn> arguments para uma chamada de função. Aqui está uma chamada de função com `"Hello"` passada como argumento:

```js
test("Hello");
```

Passe `i + 1` e `count` como argumentos para sua chamada `padRow`. Como os parâmetros, os argumentos são separados por vírgula.

# --hints--

Você deve passar `i + 1` para sua chamada `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

Você deve ter uma vírgula após seu argumento `i + 1`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

Você deve passar `count` como seu segundo argumento.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
