---
id: 660f033cf051ebb50ea3bf48
title: Passo 18
challengeType: 1
dashedName: step-18
---

# --description--

Em programação, você frequentemente precisará trabalhar com muitos dados. Existem muitas estruturas de dados que podem ajudar você a organizar e gerenciar seus dados. Uma das estruturas de dados mais básicas é um <dfn>array</dfn>.

Um <dfn>array</dfn> é um tipo de dado não primitivo que pode conter uma série de valores. Tipos de dados não primitivos diferem dos tipos de dados primitivos porque podem conter dados mais complexos. Tipos de dados primitivos como strings e números só podem conter um valor por vez.

Arrays são denotados usando colchetes (`[]`). Aqui está um exemplo de uma variável com o valor de um array vazio:

```js
let array = [];
```

Declare uma variável `rows` e atribua a ela um array vazio.

# --hints--

Você deve ter uma variável `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

Você deve usar `let` para declarar sua variável `rows`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

Você deve atribuir um array vazio à sua variável `rows`.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
