---
id: 660f47afe4c98536715d5fa4
title: Passo 91
challengeType: 1
dashedName: step-91
---

# --description--

O operador <dfn>strict inequality</dfn> `!==` permite que você verifique se dois valores não são iguais ou não têm o mesmo tipo.A sintaxe é semelhante ao operador de igualdade: `value !== 4`.

Atualmente, o loop `while` executa apenas se `continueLoop` for verdadeiro.Atualize a condição do loop `while` para verificar se `done` não é igual a `count`.

# --hints--

Seu loop `while` deve verificar se `done` e `count` não são iguais.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
