---
id: 660f039ff313dbb696b007ca
title: Passo 19
challengeType: 1
dashedName: step-19
---

# --description--

Quando um array contém valores, ou <dfn>elementos</dfn>, esses valores são separados por vírgulas.Aqui está um array que contém duas strings:

```js
let array = ["first", "second"];
```

Altere sua declaração de `rows` para ser um array com as strings `"Naomi"`, `"Quincy"` e `"CamperChan"`.A ordem dos valores em um array é importante, então siga essa ordem.Lembre-se de que strings diferenciam maiúsculas de minúsculas.

# --hints--

O primeiro elemento no seu array deve ser a string `"Naomi"`.

```js
assert.equal(rows[0], "Naomi");
```

O segundo elemento no seu array deve ser a string `"Quincy"`.

```js
assert.equal(rows[1], "Quincy");
```

O terceiro elemento no seu array deve ser a string `"CamperChan"`.

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
