---
id: 660f359af3e32e0f1a6880b7
title: Passo 68
challengeType: 1
dashedName: step-68
---

# --description--

Você agora deve ver o mesmo conjunto de caracteres no seu console. Sua função `padRow` está fazendo exatamente a mesma coisa que você fazia antes, mas agora está em uma seção reutilizável própria.

Use o operador de adição para concatenar um espaço único `" "` no início e no final da sua string repetida `character`.

Lembre-se de que você pode usar o operador `+` para concatenar strings assim:

```js
" " + "string"
```

# --hints--

Você deve concatenar um único espaço ao início do valor retornado.

```js
assert.match(padRow(1, 1), /^\s/);
```

Você deve concatenar um espaço único ao final do valor retornado.

```js
assert.match(padRow(1, 1), /\s$/);
```

Sua função `padRow()` deve retornar a série repetida de `character` com um espaço antes e depois da série.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
