---
id: 660f07d231941bc11719f664
title: Passo 21
challengeType: 1
dashedName: step-21
---

# --description--

Arrays são especiais porque são considerados <dfn>mutáveis</dfn>.Isso significa que você pode alterar o valor em um índice diretamente.

Por exemplo, este código atribuiria o número `25` ao segundo elemento do array:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

Atualize o **terceiro** elemento do seu array `rows` para ser o número `10`.Então imprima o array `rows` no seu console.

# --hints--

Você deve usar a notação de colchetes no array `rows` novamente.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

Você deve acessar o terceiro elemento do array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

Você deve usar o operador de atribuição no terceiro elemento do array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

Você deve atribuir o valor `10` ao terceiro elemento do seu array `rows`.

```js
assert.equal(rows[2], 10);
```

Você deve ter uma segunda declaração `console.log` no seu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Você deve registrar o array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
