---
id: 660f4a1472f8e63d76162ce5
title: Passo 99
challengeType: 1
dashedName: step-99
---

# --description--

E se você fizesse sua pirâmide de cabeça para baixo, ou <dfn>invertida</dfn>? Hora de experimentar!

Comece criando um novo loop `for`. Declare seu iterador `i` e atribua a ele o valor de `count`, depois use o booleano `false` para suas declarações de condição e iteração.

# --hints--

Seu código deve ter um loop `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

Seu loop `for` deve inicializar `i` com o valor de `count`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count/);
```

Seu loop `for` deve usar `false` como condição.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false/);
```


Seu loop `for` deve usar `false` como a iteração.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
