---
id: 660f4934fb48f63abd5ae371
title: Passo 95
challengeType: 1
dashedName: step-95
---

# --description--

Usar `done` para rastrear o número de linhas que foram geradas é funcional, mas você pode realmente simplificar a lógica um pouco mais.

Arrays têm uma propriedade especial `length` que permite ver quantos valores ou <dfn>elementos</dfn> estão no array. Você acessaria essa propriedade usando uma sintaxe como `myArray.length`.

Observe que `rows.length` na chamada de `padRow` resultaria em um erro de um a mais, porque `done` é incrementado *antes* da chamada.

Atualize sua condição para verificar se `rows.length` é menor que `count`.

# --hints--

Seu loop `while` deve verificar se `rows.length` é menor que `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
