---
id: 660f18f059fe0fda192ce394
title: Passo 34
challengeType: 1
dashedName: step-34
---

# --description--

Seu loop agora precisa de um iterador adequado. O <dfn>iterator</dfn> é uma variável que você pode declarar especificamente no seu loop `for` para controlar como o loop itera ou percorre sua lógica.

É uma convenção comum usar `i` como sua variável iteradora em um loop. Um loop `for` permite que você declare isso entre parênteses `()`. Por exemplo, aqui está um loop `for` que declara uma variável `index` e atribui a ela o valor `100`.

```js
for (let index = 100; "second"; "third") {

}
```

Substitua a string `"iterator"` por uma declaração `let` para a variável `i`. Atribua a ele o valor `0` para começar. Isso dará à variável `i` o valor `0` na **primeira vez** que seu loop for executado.

# --hints--

Você deve usar `let` para declarar uma variável `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

Você deve atribuir `0` à sua variável `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

Seu loop `for` deve iniciar um iterador `i` em `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
