---
id: 660f1a00ac619ddc1e259a66
title: Passo 35
challengeType: 1
dashedName: step-35
---

# --description--

A <dfn>condição</dfn> de um loop `for` informa ao loop quantas vezes ele deve iterar. Quando a `condition` se tornar falsa, o loop irá parar.

Em JavaScript, um valor Booleano pode ser `true` ou `false`. Estes não são strings - você aprenderá mais sobre a diferença mais adiante.

Por enquanto, você usará o operador <dfn>less than</dfn> (`<`). Isso permite que você verifique se o valor à esquerda é menor que o valor à direita. Por exemplo, `count < 3` avaliaria para `true` se `count` for `2` e `false` se `count` for `4`.

Substitua sua string `"condition"` por uma condição para verificar se `i` é menor que `count`.

# --hints--

Você deve usar o operador less than.

```js
assert.match(__helpers.removeJSComments(code), /</);
```

Você deve usar o operador menor que para verificar se `i` é menor que `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Seu loop `for` deve usar `i < count` como condição.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
