---
id: 6610c105bbdacc114d6cdc44
title: Passo 48
challengeType: 1
dashedName: step-48
---

# --description--

Para usar uma função, você precisa chamá-la.Uma <dfn>function call</dfn> instrui sua aplicação a executar o código da função onde quer que você escolha chamá-la.A sintaxe para uma chamada de função é o nome da função seguido por parênteses.Por exemplo, este código define e chama uma função `test`.

```js
function test() {

}
test();
```

Chame sua função `padRow`.

# --hints--

Você deve chamar a função `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

O corpo da sua função `padRow` deve estar vazio.

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
