---
id: 6610c77d50636722e5b6be17
title: Passo 60
challengeType: 1
dashedName: step-60
---

# --description--

Abaixo da instrução `return`, registre a string `"This works!"` no console.

Depois de fazer isso, você verá que a string `"This works!"` não é exibida no console e a linha `console.log("This works!")` está esmaecida.

Copie o log do console e cole-o acima da instrução `return`. Agora, a string `"This works!"` deve aparecer no console.

Uma coisa importante a saber sobre a palavra-chave `return` é que ela não apenas define um valor a ser retornado da sua função, mas também interrompe a execução do seu código dentro de uma função ou de um bloco de instruções. Isso significa que qualquer código após uma instrução `return` não será executado.

# --hints--

Sua função `padRow` deve retornar a variável `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

Seu primeiro `console.log` deve vir depois da sua palavra-chave `return`.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

Seu segundo `console.log` deve vir antes da sua palavra-chave `return`.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  
  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
