---
id: 6610c2d8d67563174fcf96dc
title: Passo 52
challengeType: 1
dashedName: step-52
---

# --description--

Quando você tem um valor que está explicitamente escrito no seu código, como a string `"Hello!"` na sua função, ele é considerado <dfn>hard-coded</dfn>. Hard-coding um valor dentro de uma função pode não torná-la tão reutilizável quanto você gostaria.

Em vez disso, você pode definir <dfn>parameters</dfn> para a função. Parâmetros são variáveis especiais que recebem um valor quando você chama a função e podem ser usados na sua função para alterar dinamicamente o resultado do código da função.

Para adicionar um parâmetro à sua função, você precisa adicionar um nome de variável dentro dos parênteses. Por exemplo, esta função `demo` tem um parâmetro `name`:

```js
function demo(name) {

}
```

`name` parece um parâmetro útil, então vá em frente e adicione-o à sua função `padRow`.

# --hints--

Sua função `padRow` deve ter um parâmetro `name`.

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
