---
id: 6610bf6fa14d700beed1b109
title: Passo 87
challengeType: 1
dashedName: step-87
---

# --description--

O operador <dfn>equality</dfn> `==` é usado para verificar se dois valores são iguais. Para comparar dois valores, você usaria uma declaração como `value == 8`.

Abaixo de `done++` dentro do seu loop, adicione uma declaração `if`. A instrução deve verificar se `done` é igual a `count` usando o operador de igualdade.

# --hints--

Você deve usar uma declaração `if` no seu loop. Deve ser adicionado após `done++`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

Sua declaração `if` deve usar o operador de igualdade para comparar `done` e `count` na condição.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
