---
id: 6610b8f6a98d25f4d485a94d
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

O valor padrão de uma variável não inicializada é `undefined`. Este é um tipo de dado especial que representa um valor que ainda não tem uma definição.

Você ainda pode atribuir um valor a uma variável não inicializada. Aqui está um exemplo:

```js
let uninitialized;
uninitialized = "assigned";
```

Atribua a string `"Test"` à sua variável `secondCharacter` abaixo da sua declaração. Abra o console para ver como seu log mudou.

# --hints--

Você não deve inicializar `secondCharacter`. Lembre-se de que inicialização significa atribuir um valor quando você declara a variável.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

Você deve usar o operador de atribuição em `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

Você deve atribuir a string `"Test"` à sua variável `secondCharacter`.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
