---
id: 6610c16c4fa0df12c0e30675
title: Passo 49
challengeType: 1
dashedName: step-49
---

# --description--

Você está chamando sua função `padRow`, mas não está fazendo nada com essa chamada de função. Todas as funções em JavaScript <dfn>return</dfn> um valor, o que significa que elas fornecem o resultado definido da chamada para você usar em outro lugar.

Para ver o resultado de chamar sua função `padRow`, declare uma variável `call` e atribua a essa variável sua chamada existente de `padRow`.

# --hints--

Você deve declarar uma variável `call`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

Você deve usar `const` para declarar sua variável `call`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

Você deve atribuir a `call` o resultado da sua chamada `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
