---
id: 6610c1d97b1671140f95cfbb
title: Passo 50
challengeType: 1
dashedName: step-50
---

# --description--

Agora adicione uma instrução de log para imprimir o valor da sua variável `call`.

# --hints--

Você deve adicionar uma chamada `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2)
```

Você deve registrar sua variável `call`. Não esqueça o ponto e vírgula.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*call\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
