---
id: 6610bbed59bc2a0194d85533
title: Passo 22
challengeType: 1
dashedName: step-22
---

# --description--

Perceba como o valor dentro do seu array `rows` foi alterado diretamente? Isto é chamado <dfn>mutation</dfn>. À medida que você aprende mais sobre arrays, você aprenderá quando deve mutar um array e quando não deve.

Antes de continuar, esta é uma ótima oportunidade para aprender um uso comum de array. Atualmente, seu código acessa o último elemento no array com `rows[2]`. Mas você pode não saber quantos elementos há em um array quando quiser o último.

Você pode usar a propriedade `.length` de um array - isso retorna o número de elementos no array. Para obter o último elemento de qualquer array, você pode usar a seguinte sintaxe:

```js
array[array.length - 1]
```

`array.length` retorna o número de elementos no array. Subtraindo `1`, você obtém o índice do último elemento no array. Você pode aplicar esse mesmo conceito ao seu array `rows`.

Atualize seu `rows[2]` para acessar dinamicamente o último elemento no array `rows`. Consulte o exemplo acima para ajudá-lo.

Você não deve ver nenhuma alteração no seu console.

# --hints--

Você deve usar a propriedade `.length` do seu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

Você deve usar `rows.length` na sua notação de colchetes.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

Você deve subtrair `1` do comprimento na sua notação de colchetes.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
