---
id: 6610c48c4ea0891afa7c4696
title: Passo 54
challengeType: 1
dashedName: step-54
---

# --description--

Se você abrir seu console novamente, verá que sua função `padRow` está retornando `undefined`, mesmo que você tenha definido um valor de retorno!Isso ocorre porque os parâmetros precisam receber um valor quando você **chama** a função.

Quando você passa um valor para uma chamada de função, esse valor é referido como um <dfn>argumento</dfn>.Aqui está um exemplo de chamada da função `demo` passando `"Naomi"` como argumento para o parâmetro `name`.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

Passe seu próprio nome como argumento para o parâmetro `name` na sua chamada `padRow`.Lembre-se de que seu nome é uma string, então você precisará usar aspas.

# --hints--

Você deve passar uma string para sua chamada `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
