---
id: 6610c71600966a2191d3a64a
title: Passo 59
challengeType: 1
dashedName: step-59
---

# --description--

Valores retornados de uma função são usados ao chamar a função. Você pode usar a chamada da função diretamente como o valor que ela retorna ou capturar o valor retornado em uma variável. Dessa forma, você pode usar o valor atribuído a uma variável com escopo local, fora da função em que foi criada.

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

Para usar seu valor `"Testing"`, retorne-o da função `padRow` atualizando sua declaração de `return` para retornar apenas a variável `test`.


# --hints--

Sua função `padRow` deve retornar a variável `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
