---
id: 6610b8017d1671f2814e8c77
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Quando você declara uma variável sem inicializá-la, ela é considerada <dfn>uninitialized</dfn>.Atualmente, sua variável `secondCharacter` está não inicializada.

Adicione um `console.log()` para ver qual é o valor da sua variável `secondCharacter`.

# --hints--

Você deve adicionar um segundo `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Você deve registrar sua variável `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
