---
id: 6610c6541c82551f95e765ab
title: Passo 58
challengeType: 1
dashedName: step-58
---

# --description--

Variáveis também podem ser declaradas dentro de uma função. Essas variáveis são consideradas no <dfn>local scope</dfn> ou <dfn>block scope</dfn>. Uma variável declarada dentro de uma função só pode ser usada dentro dessa função. Se você tentar acessá-lo fora da função, receberá um erro de referência.

Para ver isso em ação, use `const` para declarar uma variável `test` na sua função `padRow`. Inicialize-o com o valor `"Testing"`.

Então, abaixo da sua função, tente registrar `test` no console. Você verá um erro porque não está definido fora do escopo local da função. Remova esse `console.log` para passar nos testes e continuar.

# --hints--

Sua função deve declarar uma variável `test`.

```js
assert.match(padRow.toString(), /var\s+test/);
```

Você deve inicializar `test` com o valor `"Testing"`. Não esqueça o ponto e vírgula.

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

Sua variável `test` deve vir antes da palavra-chave `return`.

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
