---
id: 6610c8cfe4cf4d278e35c156
title: Passo 63
challengeType: 1
dashedName: step-63
---

# --description--

Como esperado, sua função agora retorna `undefined` novamente.Sua variável `call` não é mais necessária, então remova a declaração `call` e o `console.log` para a variável `call`.

# --hints--

Você não deve ter uma declaração `call`.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

Você não deve registrar sua variável `call`.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
