---
id: 6610c424b7119919b62932f4
title: Passo 53
challengeType: 1
dashedName: step-53
---

# --description--

Uma função não precisa retornar um valor codificado diretamente.Ele pode retornar o valor armazenado em uma variável.Parâmetros são variáveis especiais para uma função, então eles também podem ser retornados.

Altere sua função `padRow` para `return` o parâmetro `name` diretamente.

# --hints--

Sua função `padRow` deve retornar o valor do parâmetro `name`.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
