---
id: 6610b741b54b90f0c0fb3d58
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Quando os nomes das variáveis têm mais de uma palavra, existem convenções específicas de nomenclatura para como você capitaliza as palavras.Em JavaScript, a convenção a ser usada é <dfn>camel case</dfn>.

Camel case significa que a primeira palavra no nome está inteiramente em minúsculas, mas as palavras seguintes estão todas com a primeira letra maiúscula.Aqui estão alguns exemplos de camel case:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

Use camel case para declarar uma nova variável `secondCharacter`.

# --hints--

Você deve declarar uma variável `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

Você deve usar `let` para declarar sua variável `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

Você não deve atribuir um valor à sua variável `secondCharacter`. Não esqueça o ponto e vírgula.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
