---
id: 6610c21b3ef82015573ffbbe
title: Passo 51
challengeType: 1
dashedName: step-51
---

# --description--

Sua variável `call` tem um valor `undefined`, mesmo que você a tenha definido! Isso ocorre porque sua função `padRow` atualmente não retorna um valor. Por padrão, funções retornam `undefined` como seu valor.

Para retornar outra coisa, você precisa usar a palavra-chave `return`. Aqui está um exemplo de uma função que retorna a string `"Functions are cool!"`:

```js
function demo() {
  return "Functions are cool!";
}
```

Use a palavra-chave `return` para fazer sua função retornar a string `"Hello!"`.

# --hints--

Sua função `padRow` deve usar a palavra-chave `return`.

```js
assert.match(padRow.toString(), /return/);
```

Sua função `padRow` deve retornar a string `"Hello!"`.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
