---
id: 6610c87eac0f0b256d7b037e
title: Passo 62
challengeType: 1
dashedName: step-62
---

# --description--

Porque sua função não estava mais usando o parâmetro, alterar o argumento não a afetou.

Vá em frente e remova a declaração `test` e a instrução `return` da sua função `padRow`, para que a função fique vazia novamente.

# --hints--

Sua função `padRow` não deve ter uma variável `test`.

```js
assert.notMatch(padRow.toString(), /test/);
```

Sua função `padRow` não deve retornar um valor.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
