---
id: 6610c538372aa61cc0f5b122
title: Passo 57
challengeType: 1
dashedName: step-57
---

# --description--

Variáveis em JavaScript estão disponíveis em um <dfn>escopo</dfn> específico. Em outras palavras, onde uma variável é declarada determina onde no seu código ela pode ser usada.

O primeiro escopo é o escopo global. Variáveis que são declaradas fora de qualquer "bloco" como uma função ou `for` loop estão no <dfn>escopo global</dfn>. Suas variáveis `character`, `count` e `rows` estão todas no escopo global.

Quando uma variável está no escopo global, uma função pode acessá-la em sua definição. Aqui está um exemplo de uma função usando uma variável global `title`:

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

Este exemplo retornaria `"Professor Naomi"`. Atualize sua função `padRow` para retornar o valor da concatenação da sua variável `character` ao início do parâmetro `name`.

# --hints--

Sua função `padRow` deve concatenar `character` ao início de `name`.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

Sua função `padRow` deve retornar o resultado de `character + name`.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
