---
id: 6610b9f7619764fad5fd516d
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Você também pode atribuir o valor de uma variável a outra variável.Por exemplo:

```js
let first = "One";
let second = "Two";
second = first;
```

A variável `second` agora teria o valor `"One"`.

Para ver isso em ação, altere sua atribuição de `secondCharacter` de `"Test"` para sua variável `character`.

Então abra o console para ver o que é registrado.

# --hints--

Você não deve atribuir o valor `"Test"` à sua variável `secondCharacter`.

```js
assert.notEqual(secondCharacter, "Test");
```

Você deve atribuir o valor da variável `character` à sua variável `secondCharacter`.Não esqueça seu ponto e vírgula.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

Sua variável `secondCharacter` agora deve ter o valor `"World"`.


```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
