---
id: 6610c83b52583e245a079217
title: Passo 61
challengeType: 1
dashedName: step-61
---

# --description--

Agora sua variável `call` tem o valor `"Testing"`. Mas sua função não está mais usando o parâmetro `name`.

Remova o parâmetro `name` da declaração da sua função, depois remova a string `"CamperChan"` da chamada `padRow`.

Além disso, remova ambos os `console.log` da função `padRow`.
# --hints--

Sua função `padRow` não deve ter um parâmetro `name`.

```js
assert.notMatch(padRow.toString(), /name/);
```

Você não deve passar `"CamperChan"` para sua chamada `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

Você ainda deve chamar sua função `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Você não deve ter um `console.log` antes da sua palavra-chave `return`.

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

Você não deve ter um `console.log` após sua palavra-chave `return`.

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
