---
id: 66458f0a05df478aa627629e
title: Passo 82
challengeType: 1
dashedName: step-82
---

# --description--

O texto sumiu de novo! Strings vazias são avaliadas como `false`, tornando-as um <dfn>valor falso</dfn>. Você aprenderá mais sobre valores verdadeiros e falsos em projetos futuros.

Além das declarações `if`, o JavaScript também possui declarações <dfn>else if</dfn>. As declarações `else if` permitem que você verifique múltiplas condições em um único bloco de código.

Aqui está a sintaxe para uma instrução `else if`:

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

Se a primeira condição for `false`, o JavaScript verificará a próxima condição na cadeia. Se a segunda condição for `false`, o JavaScript verificará a terceira condição, e assim por diante.

Abaixo da sua instrução `if`, adicione uma instrução `else if` que verifica se `5` é menor que `10`. Em seguida, dentro do corpo da instrução `else if`, registre a string `"5 is less than 10"` no console.

Verifique o console para ver os resultados.

# --hints--

Você deve ter uma instrução `else if`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Sua instrução `else if` deve verificar se `5` é menor que `10`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

Você deve registrar a string `"5 is less than 10"` para o console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

