---
id: 664599653fcd6e97104f9261
title: Passo 83
challengeType: 1
dashedName: step-83
---

# --description--

Às vezes você vai querer executar um código diferente quando todas as condições `if...else if` forem `false`.Você pode fazer isso adicionando um bloco `else`. 

Um bloco `else` será avaliado somente se as condições nos blocos `if` e `else if` não forem atendidas.

Aqui o bloco `else` é adicionado ao bloco `else if`. 

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

Adicione um bloco `else` ao bloco `else if`.Dentro do bloco `else`, registre a string `"This is the else block"` no console.

Para ver os resultados no console, você pode alterar manualmente o `<` na declaração `else if` para `>`.Isso fará com que a condição seja `false` e o bloco `else` será executado.

# --hints--

Você deve ter um bloco `else`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

O seu bloco `else` deve registrar a string `"This is the else block"` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

