---
id: 664ed97c55a99e5ffab759dc
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

Antes de seguir em frente, você deve reservar um momento para revisar os conceitos que aprendeu.

Use a palavra-chave `let` para declarar uma variável `profession` e uma variável `age`. Inicialize `profession` com a string `"teacher"`, mas não inicialize `age` com nenhum valor.

Registre ambas as suas variáveis no console para ver os resultados.

# --hints--

Você deve declarar uma variável `profession`.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

Você deve usar `let` para declarar a variável `profession`.

```js
assert.match(code, /let\s+profession/);
```

Você deve atribuir a string `"teacher"` à variável `profession`.

```js
assert.strictEqual(profession, "teacher");
```

Você deve declarar uma variável `age`.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

Você deve usar `let` para declarar a variável `age`.

```js
assert.match(code, /let\s+age/);
```

Você não deve atribuir um valor à variável `age`.

```js
assert.isUndefined(age);
```

Você deve registrar `age` no console.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

Você deve registrar `profession` no console.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
