---
id: 664ed9344ba3495e1f67f8ea
title: Passo 14
challengeType: 1
dashedName: step-14
---

# --description--

Agora que você revisou declaração e inicialização, remova o código que você escreveu para sua revisão.

Não remova sua variável `character`.

# --hints--

Você não deve ter uma variável `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

Você não deve ter uma variável `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

Você não deve ter uma declaração `console.log` para `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

Você não deve ter uma declaração `console.log` para `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

Você ainda deve ter sua variável `character`.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
