---
id: 665732e62eb8b30c249b5c2c
title: Passo 25
challengeType: 1
dashedName: step-25
---

# --description--

Agora você está pronto para passar para o próximo conjunto de lições sobre arrays.

Remova todo o seu código da etapa anterior.

# --hints--

Você deve remover todas as referências à sua variável `cities`.

```js
assert.notMatch(__helpers.removeJSComments(code), /cities/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let cities = ["London", "New York", "Mumbai"];
console.log(cities);
cities[cities.length - 1] = "Mexico City";
console.log(cities);
--fcc-editable-region--
console.log(rows);
```
