---
id: 665732da4815b70bb083915e
title: Passo 24
challengeType: 1
dashedName: step-24
---

# --description--

Nos últimos passos, você aprendeu tudo sobre como trabalhar com arrays. Reserve um momento para revisar o que você aprendeu.

Comece declarando uma variável `cities` e inicializando-a como um array das strings `"London"`, `"New York"` e `"Mumbai"`. Então registre essa variável no console.

Após o registro, altere o último elemento de `cities` para a string `"Mexico City"`, depois registre a variável `cities` novamente.

Quando feito corretamente, você deve ver essa saída no console.

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

Você deve usar `let` para declarar uma variável `cities`.

```js
assert.match(code, /let\s+cities/);
```

Você deve atribuir um array das strings `"London"`, `"New York"` e `"Mumbai"` à variável `cities`.

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

Você deve usar `console.log()` para registrar todo o array `cities` no console.

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

Você deve atualizar o último elemento do array `cities` para a string `"Mexico City"`. Lembre-se de que você pode acessar o último elemento de um array usando `array[array.length - 1]`.

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

Você deve ter duas declarações `console.log(cities)` no seu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
