---
id: 66643c93e05093c728abdbe9
title: Passo 55
challengeType: 1
dashedName: step-55
---

# --description--

Antes de continuar, reserve um momento para revisar como as funções funcionam.

Declare uma função chamada `addTwoNumbers`.Esta função deve receber dois argumentos e retornar a soma desses dois argumentos.

Sua função não deve usar valores codificados diretamente.Um exemplo de uma função codificada diretamente pode ser:

```js
function sayName(firstName, lastName) {
  return "John Doe";
}

sayName("Camper", "Cat");
```

Esta função retornaria `"John Doe"` independentemente dos argumentos passados para os parâmetros `firstName` e `lastName`, então `"John Doe"` é considerado um valor codificado diretamente.

Declare uma variável `sum` e atribua a ela o valor da chamada da sua função `addTwoNumbers` com `5` e `10` como argumentos.Registre a variável `sum` no console.

# --hints--

Você deve ter uma função chamada `addTwoNumbers`.

```js
assert.isFunction(addTwoNumbers);
```

Sua função `addTwoNumbers` deve ter dois parâmetros.

```js
assert.lengthOf(addTwoNumbers, 2);
```

Sua função deve retornar a soma dos dois parâmetros.

```js
assert.strictEqual(addTwoNumbers(5,10), 15);
```

Sua função não deve retornar um valor codificado diretamente.Ou seja, deve funcionar com quaisquer dois argumentos numéricos.

```js
assert.strictEqual(addTwoNumbers(3, 5), 8);
```

Você deve declarar uma variável `sum`.

```js
assert.isDefined(sum);
```

Sua variável `sum` deve ter o valor `15`.

```js
assert.strictEqual(sum, 15);
```

Você deve atribuir a `sum` o valor da chamada da função `addTwoNumbers` com `5` e `10` como argumentos.

```js
assert.match(code, /sum\s*=\s*addTwoNumbers\s*\(/);
```

Você deve registrar sua variável `sum`.

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
