---
id: 669fcb06c3034a39f5431a38
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare creating some actual content, add a `body` element below the `head` element.

# --hints--

Dovresti avere un tag di apertura `<body>`.

```js
assert.match(code, /<body>/i);
```

Dovresti avere un tag di chiusura `</body>`.

```js
assert.match(code, /<\/body>/i);
```

Non dovresti cambiare l'elemento `head`. Assicurati di non aver eliminato il tag di chiusura.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

L'elemento `body` dovrebbe essere all'interno dell'elemento `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>

  --fcc-editable-region--
</html>
```
