---
id: 6998b99a8e195cd0c1f211df
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, você vai construir um rastreador simples de vida selvagem usando objetos JavaScript.

Aqui está um exemplo de um objeto:

```js
const animal = {}; // empty object
```

Isso cria um objeto vazio.

Agora crie uma variável chamada `tiger` e atribua a ela um objeto vazio.

# --hints--

Você deve criar uma variável chamada `tiger`.

```js
assert.exists(tiger);
```

A variável `tiger` deve ser atribuída a um objeto vazio.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
