---
id: 6999b07c34c2cde21ee6e1d8
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Neste passo, você vai criar uma função que verifica se um objeto possui uma propriedade específica.

Aqui está um exemplo de uso do `hasOwnProperty`:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

Neste exemplo, a função `hasColor` usa `hasOwnProperty` para verificar se a propriedade `color` existe em um objeto.

Agora crie uma função chamada `hasHabitat`. A função deve receber um parâmetro chamado `animal`.

Use o método `hasOwnProperty` para retornar `true` se o objeto `animal` tiver uma propriedade chamada `"habitat"` e `false` caso contrário.

Depois de criar a função, use `console.log` para chamar `hasHabitat(tiger)` e `hasHabitat(elephant)` para que você possa ver ambos os resultados no console.

# --hints--

Você deve criar uma função chamada `hasHabitat`.

```js
assert.isFunction(hasHabitat);
```

A função `hasHabitat` deve ter um único parâmetro, `animal`.

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` deve usar o método `hasOwnProperty` para verificar a propriedade `habitat`.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

Você deve registrar `hasHabitat(tiger)` no console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

Você deve registrar `hasHabitat(elephant)` no console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

Chamar `hasHabitat(tiger)` deve retornar `true` após adicionar um habitat.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Chamar `hasHabitat(elephant)` deve retornar `false` se nenhum habitat existir.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` deve usar os parâmetros da função e funcionar com qualquer objeto.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
