---
id: 6999b0d79ea938edcb032237
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

Neste passo final da oficina Wildlife Tracker, você vai criar uma função que acessa uma propriedade de um objeto usando notação por colchetes.

Aqui está um exemplo de acesso a uma propriedade usando notação por colchetes:

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

Neste exemplo, a notação por colchetes permite acessar uma propriedade usando uma variável.

Agora crie uma função chamada `getProperty`.

A função deve receber dois parâmetros: `animal` e `propertyName`.

Retorne o valor da propriedade usando notação por colchetes.

Depois de criar a função, use `console.log` para chamar `getProperty(tiger, "species")` e `getProperty(elephant, "age")` para que você possa ver os valores retornados no console.

Com isso, a oficina Wildlife Tracker está completa!

# --hints--

Você deve criar uma função chamada `getProperty`.

```js
assert.isFunction(getProperty);
```

A função `getProperty` deve ter dois parâmetros: `animal` e `propertyName`.

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty` deve usar notação por colchetes para acessar a propriedade.

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

Você deve registrar `getProperty(tiger, "species")` no console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

Você deve registrar `getProperty(elephant, "age")` no console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

Chamar `getProperty(tiger, "species")` deve retornar `"Tiger"`.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

Chamar `getProperty(elephant, "age")` deve retornar `12`.

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty` deve usar os parâmetros da função e funcionar com qualquer objeto.

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
