---
id: 6999a0c1698b93d5d41156ed
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora adicione outra propriedade ao objeto `tiger` chamada `isEndangered`.

Defina a propriedade `isEndangered` como `true`.

# --hints--

`tiger` deve ter uma propriedade `isEndangered`.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` deve ser igual a `true`.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
