---
id: 6999adc72c02812a28c01e31
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Neste passo, vocĂȘ vai criar uma funĂ§ĂŁo que atualiza uma propriedade de um objeto.

Aqui estĂĄ um exemplo de atualizaĂ§ĂŁo de uma propriedade dentro de uma funĂ§ĂŁo:

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

Neste exemplo, a propriedade `age` Ă© atualizada para um novo valor.

Agora crie uma funĂ§ĂŁo chamada `updateAge`. A funĂ§ĂŁo deve receber dois parĂąmetros: `animal` e `newAge`.

Dentro da funĂ§ĂŁo, atualize a propriedade `age` do objeto `animal` para `newAge`. Retorne o objeto `animal` atualizado.

Depois de criar a funĂ§ĂŁo, use `console.log` para chamar `updateAge(elephant, 12)` para que vocĂȘ possa ver o objeto `elephant` atualizado no console.

# --hints--

VocĂȘ deve criar uma funĂ§ĂŁo chamada `updateAge`.

```js
assert.isFunction(updateAge);
```

A funĂ§ĂŁo `updateAge` deve ter dois parĂąmetros: `animal` e `newAge`.

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

`updateAge` deve usar notaĂ§ĂŁo de ponto para atualizar a propriedade `age`.

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

A funĂ§ĂŁo `updateAge` deve retornar o objeto `animal` atualizado.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

VocĂȘ deve registrar `updateAge(elephant, 12)` no console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

Chamar `updateAge(elephant, 12)` deve atualizar a propriedade age.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` deve usar os parĂąmetros da funĂ§ĂŁo e funcionar com qualquer objeto.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
