---
id: 6999b0198d18d402c36cc5fe
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Neste passo, você vai criar uma função que remove uma propriedade de um objeto.

Aqui está um exemplo de remoção de uma propriedade usando `delete`:

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

Neste exemplo, a propriedade `canFly` é removida do objeto `bird`.

Agora crie uma função chamada `removeEndangeredStatus`. A função deve receber um parâmetro chamado `animal`.

Dentro da função, remova a propriedade `isEndangered` do objeto `animal` usando a palavra-chave `delete`. Retorne o objeto `animal` atualizado.

Depois de criar a função, use `console.log` para chamar `removeEndangeredStatus(tiger)` para que você possa ver o objeto atualizado no console.

# --hints--

Você deve criar uma função chamada `removeEndangeredStatus`.

```js
assert.isFunction(removeEndangeredStatus);
```

A função `removeEndangeredStatus` deve ter um único parâmetro, `animal`.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

A função `removeEndangeredStatus` deve retornar o objeto `animal` atualizado.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` deve usar a palavra-chave `delete` para remover a propriedade `isEndangered`.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

Você deve registrar `removeEndangeredStatus(tiger)` no console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

Chamar `removeEndangeredStatus(tiger)` deve remover a propriedade `isEndangered`.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` deve usar os parâmetros da função e funcionar com qualquer objeto.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
