---
id: 6999faf886634b72be7a819b
title: passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Para este rastreador de vida selvagem, você vai atualizar um objeto com novas propriedades.

Aqui está um exemplo de criação de um objeto com uma propriedade:

```js
const animal = {
  name: "Lion" // property: value
};
```

Agora atualize o objeto `tiger` para que ele inclua uma propriedade chamada `species`.

Defina seu valor como `"Tiger"`.

# --hints--

`tiger` deve ser um objeto.

```js
assert.isObject(tiger);
```

`tiger` deve ter uma propriedade `species`.

```js
assert.property(tiger, "species");
```

`tiger.species` deve ser igual a `"Tiger"`.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
