---
id: 69999ff7f7862fdfb885ea5d
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Agora atualize seu objeto `tiger` adicionando uma nova propriedade chamada `age`.

Defina a propriedade `age` com o valor `5`.

# --hints--

`tiger` deve ter uma propriedade `age`.

```js
assert.property(tiger, "age");
```

`tiger.age` deve ser igual a `5`.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
