---
id: 6999ac171928136168654774
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Neste passo, você vai criar uma função para acessar a propriedade `species` de um objeto.

Aqui está um pequeno exemplo de como acessar uma propriedade de um objeto:

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

Neste exemplo, a função recebe um objeto como parâmetro e retorna sua propriedade `species`.

Agora crie uma função chamada `getSpecies`.

A função deve receber um parâmetro chamado `animal`.

Retorne a propriedade `species` do `animal` usando notação de ponto.

Depois de criar a função, use `console.log` para chamar `getSpecies(tiger)` para que você possa ver o resultado no console.

# --hints--

Você deve criar uma função chamada `getSpecies`.

```js
assert.isFunction(getSpecies);
```

A função `getSpecies` deve ter um único parâmetro, `animal`.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Você deve registrar `getSpecies(tiger)` no console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` deve acessar a propriedade species usando notação de ponto.

```js
assert.match(code, /animal\.species/);
```

Chamar `getSpecies(tiger)` deve retornar `"Tiger"`.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` deve usar o parâmetro da função e funcionar com qualquer objeto.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
