---
id: 699a00de9f564bb0effaa14d
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Neste passo, você vai criar uma função para acessar a propriedade `age` de um objeto.

Aqui está um exemplo:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

Neste exemplo, a função recebe um objeto como parâmetro e retorna sua propriedade `age`.

Agora crie uma função chamada `getAge`.

A função deve receber um parâmetro chamado `animal`.

Retorne a propriedade `age` do `animal` usando notação de ponto.

Depois de criar a função, use `console.log` para chamar `getAge(tiger)` para que você possa ver o resultado no console.

# --hints--

Você deve criar uma função chamada `getAge`.

```js
assert.isFunction(getAge);
```

A função `getAge` deve ter um único parâmetro, `animal`.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` deve acessar a propriedade `age` usando notação de ponto.

```js
assert.match(code, /animal\.age/);
```

Você deve registrar `getAge(tiger)` no console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

Chamar `getAge(tiger)` deve retornar `5`.

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` deve usar o parâmetro da função e funcionar com qualquer objeto.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
