---
id: 69ba8da6b296391d913096a0
title: passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Agora crie um segundo objeto chamado `elephant`.

Adicione as seguintes propriedades:

- `species` com o valor `"Elephant"`
- `age` com o valor `10`
- `isEndangered` com o valor `true`

# --hints--

Você deve criar uma variável chamada `elephant`.

```js
assert.exists(elephant);
```

`elephant` deve ter uma propriedade `species` igual a `"Elephant"`.

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` deve ter uma propriedade `age` igual a `10`.

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` deve ter uma propriedade `isEndangered` igual a `true`.

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
