---
id: 68e4c81577243e4db1676491
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop üben Sie den Umgang mit semantischem HTML, indem Sie eine Webseite erstellen, die einige von Quincy Larsons Tipps für den Einstieg als Entwickler enthält. Der grundlegende Boilerplate-Code für HTML wurde für Sie vorbereitet.

Beginnen Sie damit, ein `h1`-Element mit dem Text `Quincy's Tips for Getting a Developer Job` zu erstellen.

# --hints--

Sie sollten ein `h1`-Element haben.

```js
assert.exists(document.querySelector('h1'));
```

Ihr `h1`-Element sollte den Text `Quincy's Tips for Getting a Developer Job` enthalten.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
