---
id: 68e65ec5348e40a739903e6d
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Erstellen Sie nun ein Absatz-Element unterhalb des `h1`-Elements. Fügen Sie in diesen Absatz den Text `Learning to code is hard, but as Quincy Larson says, You can become a developer.` ein.

# --hints--

Sie sollten ein `p`-Element haben.

```js
assert.exists(document.querySelector('p'));
```

Ihr Absatz sollte den Text `Learning to code is hard, but as Quincy Larson says, You can become a developer.` enthalten. Überprüfen Sie Ihre Rechtschreibung sorgfältig.

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

Ihr Absatz sollte unterhalb des `h1`-Elements sein.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
