---
id: 68e90335bd53a901c41dfc13
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

In einer vorherigen Lektion haben Sie gelernt, dass das `q`-Element verwendet wird, um zitierten Text vom umgebenden Inhalt zu unterscheiden.

Hier ist ein Beispiel:

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

Die meisten modernen Browser fügen automatisch Anführungszeichen um ein inline-Zitat hinzu, wenn Sie das `q`-Element verwenden.

Umgeben Sie im Absatz `You can become a developer.` mit einem inline-Zitat-Element, ohne den restlichen Absatz zu verändern.

# --hints--

Sie sollten ein `q`-Element haben.

```js
assert.exists(document.querySelector('q'));
```

Ihr `q`-Element sollte innerhalb des Absatzes verschachtelt sein, der unter dem `h1`-Element steht.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

Ihr `q`-Element sollte den Text `You can become a developer.` enthalten.

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

Innerhalb Ihres `q`-Elements darf der Text keine führenden oder nachfolgenden Leerzeichen enthalten (keine Leerzeichen, Tabs oder Zeilenumbrüche vor dem ersten oder nach dem letzten Zeichen).

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

Sie sollten nur `q`-Tags um den passenden Teil des Satzes hinzufügen, ohne den Originaltext zu verändern. Das `p`-Element sollte weiterhin lauten: `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
