---
id: 68e964b84f954c322376e592
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Wenn die Quelle eines Zitats eine Website ist, können Sie sie mit dem `cite`-Attribut zitieren. Der Wert dieses Attributs sollte eine gültige URL sein. Obwohl dieses Attribut die Darstellung des Blockzitats nicht verändert, ist es sehr hilfreich, um Screenreadern und Suchmaschinen mehr Informationen über das Zitat zu geben.

Hier ist ein Beispiel für ein inline-Zitat-Element mit einem `cite`-Attribut:

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

Fügen Sie dem inline-Zitat-Element das `cite`-Attribut mit dieser URL hinzu:

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

Ihr `q`-Element sollte ein `cite`-Attribut haben.

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

Das `cite`-Attribut Ihres `q`-Elements sollte den Wert `https://www.freecodecamp.org/news/learn-to-code-book/` haben.

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
