---
id: 68ea7e7b916e8270483689b6
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Unter dem `h2`-Element fügen Sie ein weiteres Zitat von Quincy hinzu. Dieses Mal ist das Zitat kein Teil eines größeren Absatzes. Stattdessen ist der gesamte Absatz ein Zitat. Um solchen zitierten Text zu kennzeichnen, sollten Sie das Blockzitat-Element verwenden: `blockquote`. Im Browser sehen Sie, dass der Text leicht eingerückt ist.

Hier ist ein Beispiel für ein Blockzitat-Element mit zitiertem Text:

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Fügen Sie nun im ersten Abschnitt unter dem `h2`-Element ein Blockzitat-Element mit dem Text `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?` hinzu.

# --hints--

Innerhalb des ersten `section`-Elements sollte sich unter dem `blockquote`-Element ein `h2`-Element befinden.

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

Ihr `blockquote`-Element sollte den Text `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?` enthalten.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
