---
id: 68eab6a2b37d2f13602d0c3f
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Sie können zitierten Text direkt innerhalb des Blockzitat-Elements schreiben, so wie Sie es im ersten `section` Element gemacht haben. Alternativ können Sie ihn in einem oder mehreren Absatz-Elementen einschließen. Das ist hilfreich, wenn der Text mehrere Absätze hat, Sie diese aber im selben Blockzitat-Element behalten möchten.

Hier ist ein Beispiel mit zwei Absätzen:

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  <p>So many people ask me each week: is college still worth it? In this 1-hour video I answer this question and other commonly asked questions about university.</p>
  <p>I've been in adult education for two decades at this point, and even though I'm not a labor market economist, I do feel confident enough to answer these questions.</p>
</blockquote>
```

Fügen Sie im zweiten Abschnitt, innerhalb des bestehenden Blockzitat-Elements, vier `p` Elemente mit den folgenden Texten in der angegebenen Reihenfolge hinzu:

- `So much of getting a job is who you know.`
- `It's OK to be an introvert, but you do need to push your boundaries.`
- `Create GitHub, Twitter, LinkedIn, and Discord accounts.`
- `Go to tech meetups and conferences. Travel if you have to.`

# --hints--

Sie sollten vier `p` Elemente verschachtelt innerhalb des `blockquote` Elements im zweiten Abschnitt haben.

```js
const pElements = document.querySelectorAll('main > section:nth-of-type(2) > blockquote > p');
assert.lengthOf(pElements, 4);
```

Ihr erstes `p` Element sollte den Text `So much of getting a job is who you know.` enthalten

```js
const firstPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(1)');
assert.equal(firstPEl?.innerText.trim(), 'So much of getting a job is who you know.');
```

Ihr zweites `p` Element sollte den Text `It's OK to be an introvert, but you do need to push your boundaries.` enthalten

```js
const secondPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(2)');
assert.equal(secondPEl?.innerText.trim(), "It's OK to be an introvert, but you do need to push your boundaries.");
```

Ihr drittes `p` Element sollte den Text `Create GitHub, Twitter, LinkedIn, and Discord accounts.` enthalten

```js
const thirdPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(3)');
assert.equal(thirdPEl?.innerText.trim(), 'Create GitHub, Twitter, LinkedIn, and Discord accounts.');
```

Ihr viertes `p`-Element sollte den Text `Go to tech meetups and conferences. Travel if you have to.` enthalten.

```js
const fourthPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(4)');
assert.equal(fourthPEl?.innerText.trim(), 'Go to tech meetups and conferences. Travel if you have to.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>
      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
        --fcc-editable-region--
          
        --fcc-editable-region--
        </blockquote>
      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
