---
id: 68ea9d23318b6481ac3de777
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Genau wie beim inline-Zitat-Element können Sie auch einem Blockzitat-Element ein `cite`-Attribut hinzufügen.

Hier ist ein Beispiel für ein Blockzitat-Element mit einem `cite`-Attribut:

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Fügen Sie nun dem Blockzitat-Element ein `cite`-Attribut mit der URL `https://www.freecodecamp.org/news/learn-to-code-book/` hinzu.

# --hints--

Ihr `blockquote`-Element sollte ein `cite`-Attribut haben.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.exists(firstBlockquoteEl?.getAttribute('cite'));
```

Das `cite`-Attribut sollte den Wert `https://www.freecodecamp.org/news/learn-to-code-book/` haben.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
      --fcc-editable-region--
        <blockquote>
      --fcc-editable-region--
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
