---
id: 68ebe9cd7523d340c0b343d3
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Im dritten und letzten `section`-Element erstellen Sie die gleiche HTML-Struktur wie im vorherigen Abschnitt.

Fügen Sie zuerst innerhalb des letzten Abschnitts ein `h2`-Element mit dem Text `Importance of Building a Reputation` hinzu.

Dann fügen Sie unter dem `h2`-Element ein Blockzitat-Element mit einem `cite`-Attribut hinzu, das auf `https://www.freecodecamp.org/news/learn-to-code-book/` gesetzt ist.

Verschachteln Sie zuletzt drei Absatz-Elemente innerhalb des Blockzitat-Elements mit den folgenden Texten in der Reihenfolge:

1. `Share short video demos of your projects.`
2. `Keep applying to speak at bigger and bigger conferences.`
3. `Hang out at hackerspaces and help people who are even newer to coding than you.`

Herzlichen Glückwunsch! Damit haben Sie diesen Workshop abgeschlossen.

# --hints--

Sie sollten ein `h2`-Element verschachtelt im dritten Abschnitt haben.

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > h2'));
```

Ihr `h2`-Element sollte den Text `Importance of Building a Reputation` enthalten.

```js
const h2El = document.querySelector('main > section:nth-of-type(3) > h2');
assert.equal(h2El?.innerText.trim(), 'Importance of Building a Reputation');
```

Sie sollten ein `blockquote`-Element verschachtelt im dritten Abschnitt haben.

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > blockquote'));
```

Ihr `blockquote`-Element sollte unter Ihrem `h2`-Element sein.

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > h2 + blockquote'));
```

Ihr `blockquote`-Element sollte ein `cite`-Attribut besitzen.

```js
const blockquoteEl = document.querySelector('main > section:nth-of-type(3) > blockquote');
assert.exists(blockquoteEl?.getAttribute('cite'));
```

Ihr `cite`-Attribut des `blockquote`-Elements sollte den Wert `https://www.freecodecamp.org/news/learn-to-code-book/` haben.

```js
const blockquoteEl = document.querySelector('main > section:nth-of-type(3) > blockquote');
assert.equal(blockquoteEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

Sie sollten drei `p`-Elemente verschachtelt im `blockquote`-Element im dritten Abschnitt haben.

```js
const pElements = document.querySelectorAll('main > section:nth-of-type(3) > blockquote > p');
assert.lengthOf(pElements, 3);
```

Ihr erstes `p`-Element sollte den Text `Share short video demos of your projects.` enthalten.

```js
const firstPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(1)');
assert.equal(firstPEl?.innerText.trim(), 'Share short video demos of your projects.');
```

Ihr zweites `p`-Element sollte den Text `Keep applying to speak at bigger and bigger conferences.` enthalten.

```js
const secondPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(2)');
assert.equal(secondPEl?.innerText.trim(), "Keep applying to speak at bigger and bigger conferences.");
```

Ihr drittes `p`-Element sollte den Text `Hang out at hackerspaces and help people who are even newer to coding than you.` enthalten.

```js
const thirdPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(3)');
assert.equal(thirdPEl?.innerText.trim(), 'Hang out at hackerspaces and help people who are even newer to coding than you.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>
      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>So much of getting a job is who you know.</p>
          <p>It's OK to be an introvert, but you do need to push your boundaries.</p>
          <p>Create GitHub, Twitter, LinkedIn, and Discord accounts.</p>
          <p>Go to tech meetups and conferences. Travel if you have to.</p>
        </blockquote>
      </section>
      <section>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>

      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>So much of getting a job is who you know.</p>
          <p>It's OK to be an introvert, but you do need to push your boundaries.</p>
          <p>Create GitHub, Twitter, LinkedIn, and Discord accounts.</p>
          <p>Go to tech meetups and conferences. Travel if you have to.</p>
        </blockquote>
      </section>

      <section>
        <h2>Importance of Building a Reputation</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>Share short video demos of your projects.</p>
          <p>Keep applying to speak at bigger and bigger conferences.</p>
          <p>Hang out at hackerspaces and help people who are even newer to coding than you.</p>
        </blockquote>
      </section>
    </main>
  </body>
</html>
```
