---
id: 68552464ef378e1051a55c28
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Dentro del segundo `div`, crea dos elementos anchor. El primero debe tener el texto `Learn more`, y el segundo debe tener el texto `Start listening`.

Para el atributo `href` de los dos elementos anchor, configúralos ambos en `#`.

# --hints--

Debes crear dos elementos `a` dentro del segundo elemento `div`

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

Tu primer elemento `a` debe tener el texto `Learn more` y un atributo `href` con valor `#` dentro del segundo elemento `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

Tu segundo elemento `a` debe tener el texto `Start listening` y un atributo `href` con valor `#` dentro del segundo elemento `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
