---
id: 685523c05fb41d0848c4b52f
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro del primer elemento `div`, crea un elemento `span` con el texto `Soundflow`, un elemento `h1` con el texto `Discover New Music` y un elemento `p` con el texto `Stream your favorite tracks and discover new artists.`

# --hints--

Debes crear un elemento `span` con el texto `Soundflow` dentro del primer elemento `div`.

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

Debes crear un elemento `h1` con el texto `Discover New Music` dentro del primer elemento `div`.

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

Debes crear un elemento `p` con el texto `Stream your favorite tracks and discover new artists.` dentro del primer elemento `div`.

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
