---
id: 68551b677a5f6898af364b85
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, aprenderás más sobre Tailwind CSS usando sus clases utilitarias de flexbox para construir un componente de llamada a la acción (CTA) responsivo.

La estructura básica de HTML ha sido proporcionada para ti, así que primero necesitas construir los elementos HTML.

Comienza creando un elemento `div`. Dentro de ese `div`, crea dos elementos `div` más separados.

# --hints--

Debes crear un elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Debes tener dos elementos `div` más dentro del elemento `div`.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
