---
id: 6858f3a67aa8580ff37016c7
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Si ajustas el ancho de la pestaña de vista previa al ancho idéntico al de un dispositivo móvil nuevamente, notarás que el contenido siempre está uno al lado del otro. Esto no es lo que quieres en dispositivos móviles.

Para arreglar eso, agrega la clase `flex-col` al `div`. Esto cambiará la dirección del flex a columna, por lo que los elementos se apilarán uno encima del otro.

Aquí tienes cómo usar la clase de utilidad `flex-col`:

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

Tu elemento `div` de apertura debe tener la clase de utilidad `flex-col`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
