---
id: 6858e7e23bf9e475e14cb931
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Ahora que la estructura HTML está construida, es momento de pasar al estilizado adecuado con Tailwind.

Comienza con el elemento `div` que envuelve todo el contenido. Establece su color de fondo asignándole una `class` con el valor `bg-indigo-600`.

# --hints--

Tu elemento `div` externo debe tener un atributo `class` establecido a `bg-indigo-600`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("bg-indigo-600"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
   --fcc-editable-region--
  </body>
</html>
```
