---
id: 6858e6408a0999603b2a4fee
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Para el `h1`, añade las clases `font-bold` y `text-4xl`. Esto hará que el texto sea más negrita, más grande y también le dará más espacio entre líneas.

Además, agrega la clase `my-4` para aumentar el margen superior e inferior del `h1` en 16 píxeles, o `1rem`.

# --hints--

Tu elemento `h1` debe tener la clase de utilidad `font-bold`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("font-bold"));
```

Tu elemento `h1` debe tener la clase de utilidad `text-4xl`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("text-4xl"));
```

Tu elemento `h1` debe tener la clase de utilidad `my-4`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("my-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
         --fcc-editable-region--
        <h1>Discover New Music</h1>
         --fcc-editable-region--
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
