---
id: 6858ee73ed8106ccb214942a
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Ahora que todo está alineado a la izquierda, necesitas centrar todo el contenido agregando un margen automático en el eje x.

Agrega la clase Tailwind que corresponda a ese elemento `div` externo. Para ver ese efecto, ajusta la pestaña de vista previa hasta que todo el contenido esté centrado.

# --hints--

El elemento externo `div` debe tener la clase de utilidad `mx-auto`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--    
  </body>
</html>
```
