---
id: 6858ef974c8ad5dcd7abb9f1
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

El contenido dentro del `div` puede estar mejor espaciado y organizado. Para el resto del espaciado y esa organización, necesitas usar flexbox.

Agrega la clase `flex` al `div` para que los dos elementos hijos `div` estén uno al lado del otro.

Además, agrega `justify-around` para que el espacio alrededor y entre el contenido se distribuya de manera uniforme, y `items-center` para que todo el contenido finalmente esté centrado en el `div`. 

# --hints--

El elemento externo `div` debe tener la clase de utilidad `flex`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

El elemento externo `div` debe tener la clase de utilidad `justify-around`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

Tu elemento externo `div` debe tener la clase de utilidad `items-center`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
