---
id: 6858eadd3144c89d45e02bc7
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Si te das cuenta, el `div` externo está demasiado pegado arriba y el contenido dentro de él está demasiado amontonado a la izquierda. Así que, ahora deberías trabajar en algo de espacio.

Agrega la clase `mt-8` al `div` para empujarlo un poco hacia abajo y `p-4` para separar el texto del lado izquierdo.

# --hints--

El elemento externo `div` debe tener la clase de utilidad `mt-8`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mt-8"));
```

El elemento externo `div` debe tener la clase de utilidad `p-4`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
