---
id: 6858f90c54647a48ddb8822d
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

En este punto, el contenido, es decir, los dos elementos `div` dentro del `div` exterior, estarán apilados uno encima del otro sin importar el tamaño de la pantalla.

Eso es porque tienes la clase `flex-col` en el elemento `div` que envuelve. Para solucionarlo, usa el prefijo `lg:` para agregar la clase `flex-row` al `div`. 

Esto significa que en dispositivos con pantallas más grandes, la dirección del flex será `row`. Para ver eso en efecto nuevamente, ajusta el navegador a un ancho mayor.

# --hints--

El elemento externo `div` debe tener la clase de utilidad `lg:flex-row`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
