---
id: 6858ebf2552046ab3de9f17c
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

En este punto, notarás que hay demasiado espacio a la derecha del contenido en el `div`. Puedes minimizar esto agregando un ancho.

Agrega la clase `md:w-1/2` al `div` para reducir el ancho al 50% en pantallas medianas y más grandes. Es decir, pantallas de 768 píxeles en adelante.

Para ver eso en efecto, expande la pestaña de vista previa.

# --hints--

El elemento externo `div` debe tener la clase de utilidad `md:w-1/2`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
