---
id: 68597114730d7c6b8838d204
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

El párrafo bajo el `h1` se ve bien, así que no necesita más estilos. Eso significa que puedes pasar a los elementos anchor dentro del segundo elemento `div`.

Dale al primer elemento anchor las clases `bg-white`, `hover:bg-gray-100` y `text-indigo-600`. Esto le da un fondo blanco, un color de texto índigo y un fondo ligeramente más claro cuando pasas el cursor sobre él.

Para el segundo elemento anchor, asígnale las clases `bg-pink-500`, `hover:bg-pink-600` y `text-white`.

# --hints--

Tu primer elemento anchor debe tener la clase de utilidad `bg-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

Tu primer elemento anchor debe tener la clase de utilidad `hover:bg-gray-100`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

Tu primer elemento anchor debe tener la clase de utilidad `text-indigo-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

Tu segundo elemento anchor debe tener la clase de utilidad `bg-pink-500`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

Tu segundo elemento anchor debe tener la clase de utilidad `hover:bg-pink-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

Tu segundo elemento ancla debe tener la clase de utilidad `text-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
