---
id: a5deed1811a43193f9f1c841
title: Implementa un Saltador de Elementos
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

En este laboratorio crearás una función que omite elementos en un arreglo hasta que encuentra uno aceptable basado en una función de prueba específica.

Por ejemplo, para un arreglo como `[1, 1, 1, 2, 1, 1, 1]` y una función de prueba `function(n) {return n === 2}`, el primer elemento que cumple con esto es el que está en el índice 3, por lo que todos los elementos antes de ese deben ser descartados, y la salida debe ser los elementos restantes `[2, 1, 1, 1]`.

**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de Usuario**

1. Debes tener una función `dropElements` que acepte un arreglo (`arr`) y una función (`func`) como argumentos.
1. La función `dropElements` debe iterar a través del array y eliminar elementos comenzando desde el primero hasta que `func` devuelva `true` para un elemento.
1. La función `dropElements` debería devolver los elementos restantes en el array si se cumple la condición.
1. Si la condición nunca se cumple, debe devolver un arreglo vacío.

# --hints--

Debes tener una función `dropElements`.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` debería devolver `[3, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` debe devolver `[1, 0, 1]`.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` debería devolver `[1, 2, 3]`.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` debería devolver `[]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` debería devolver `[7, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` debería devolver `[3, 9, 2]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
