---
id: 66c63bfa4685e4e3843fa0dc
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個雜貨購物列表來繼續學習陣列。

首先為你的主控台添加一個 `console.log`，將字串 `"Grocery shopping list"` 記錄到主控台。

# --hints--

你的程式碼中應該有一個 `console.log`。

```js
assert.match(code, /console\.log(.*)/);
```

你的 `console.log` 應該輸出訊息 `"Grocery shopping list"`。請仔細檢查拼字和大小寫錯誤。

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
