---
id: 66c6491fe4c8e0f16845425f
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

現在是為你的列表添加另一種水果的時候了。

使用先前相同的陣列方法，為你的 `shoppingList` 陣列添加字串 `"Grapes"` 到結尾。

然後，為你的程式碼添加一個 `console.log`，並在 `console.log` 裡呼叫 `getShoppingListMsg` 函式，以查看更新後的列表已記錄到主控台。

# --hints--

你的 `shoppingList` 陣列現在應該包含兩個條款。

```js
assert.lengthOf(shoppingList, 2);
```

你的 `shoppingList` 陣列應該包含字串 `"Grapes"`。

```js
assert.equal(shoppingList[1], "Grapes");
```

你應該在 `console.log` 裡呼叫 `getShoppingListMsg(shoppingList)` 函式（程式）。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
