---
id: 66c645b345a39bf04864dc50
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

要查看更新後的購物列表，你需要將目前的購物列表和一則簡短的訊息記錄到主控台。

由於這個訊息將在整個工作坊中反覆使用，最好創建一個可重複使用的函式（程式）。

創建一個名為 `getShoppingListMsg` 的函式（程式），該函式（程式）接受一個陣列作為參數，並傳回字串 `"Current Shopping List: "`，後面接提供的陣列內容。

你可以使用模板字面值或使用 `+` 運算子將字串與傳遞給函式的陣列串接起來。

# --hints--

你應該在你的程式碼中有一個名為 `getShoppingListMsg` 的函式。

```js
assert.isFunction(getShoppingListMsg);
```

你的 `getShoppingListMsg` 函式應該傳回一個字串。

```js
assert.isString(getShoppingListMsg(shoppingList));
```

你的 `getShoppingListMsg` 函式應該接受一個陣列並傳回訊息 `"Current Shopping List: "`，後面接著該陣列的內容。請仔細檢查空格和拼寫錯誤。

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
