---
id: 66c644c08ebcd2ef90c750bd
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

在之前的課程中，你學會了如何使用 `push` 方法將項目添加到陣列的末端，如下所示：

```js
fruits.push('pear');
```

使用 `push` 方法，為你的 `shoppingList` 陣列添加字串 `"Apples"`。

# --hints--

你的 `shoppingList` 陣列中應該有一個項目。

```js
assert.lengthOf(shoppingList, 1);
```

你的 `shoppingList` 陣列中應該包含字串 `"Apples"`。請仔細檢查拼寫和大小寫錯誤。

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
