---
id: 66c64ad16796c7f2419b45c5
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

現在是時候開始將項目添加到購物列表的頂部。

首先為你的程式碼添加一個 `console.log`，將訊息 `"It looks like we need to get some cooking oil."` 記錄到主控台。

# --hints--

你應該將訊息 `"It looks like we need to get some cooking oil."` 記錄到主控台。

```js
assert.match(code, /console\.log\((['"])(It\s+looks\s+like\s+we\s+need\s+to\s+get\s+some\s+cooking\s+oil\.)\1\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
