---
id: 66c645b345a39bf04864dc50
title: 步驟5
challengeType: 1
dashedName: step-5
---

# --description--

要查看更新後的購物清單，你需要將當前購物清單和一條簡短信息記錄到控制檯。

由於這條信息將在整個工作坊中反覆使用，因此最好創建一個可重複使用的函數。

創建一個名爲 `getShoppingListMsg` 的函數，將數組作爲參數並返回字符串 `"Current Shopping List: "` 和所提供數組的內容。

你可以使用模板字面量或使用 `+` 操作符進行字符串連接，將字符串與傳遞給函數的數組結合起來。

# --hints--

你的代碼中應該有一個名爲 `getShoppingListMsg` 的函數。

```js
assert.isFunction(getShoppingListMsg);
```

你的 `getShoppingListMsg` 函數應返回一個字符串。

```js
assert.isString(getShoppingListMsg(shoppingList));
```

你的 `getShoppingListMsg` 函數應接受一個數組作爲參數。


```js
assert.strictEqual(getShoppingListMsg(["Anchovies", "Olives"]), "Current Shopping List: Anchovies,Olives");
```

你的 `getShoppingListMsg` 函數應接受一個數組並返回信息 `"Current Shopping List: "` 和該數組的內容。 仔細檢查空格和拼寫錯誤。

```js
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
