---
id: 66c64095f61166ede6e39a84
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

對於這個購物列表，你將使用陣列來表述你需要購買的項目。

在之前的課程中，你學會了如何像這樣創建陣列：

```js
const fruits = ['apple', 'banana', 'orange'];
```

在此步驟中，創建一個名為 `shoppingList` 的變數，並將其指定為空陣列。

# --hints--

你應該有一個名為 `shoppingList` 的變數。

```js
assert.isNotNull(shoppingList);
```

你應該將一個陣列指定給 `shoppingList` 變數。

```js
assert.isArray(shoppingList);
```

你的 `array` 應該是空的。

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
