---
id: 66c64210413532ee9d3bd342
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

在接下來的工作坊部分，你將練習為你的雜貨列表添加食物項目。

首先使用 `console.log` 記錄訊息 `"It will be nice to have some fruit to eat."`。 

# --hints--

你應該將訊息 `"It will be nice to have some fruit to eat."` 記錄到主控台。

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
