---
id: 66c64ad16796c7f2419b45c5
title: 步驟 8
challengeType: 1
dashedName: step-8
---

# --description--

現在是時候開始在雜貨清單前面添加物品了。

首先添加一個 `console.log`，將信息 `"It looks like we need to get some cooking oil."` 記錄到控制檯。

# --hints--

你應該在控制檯中記錄消息 `"It looks like we need to get some cooking oil."`。

```js
assert.match(code, /console\.log\((['"])(It\s+looks\s+like\s+we\s+need\s+to\s+get\s+some\s+cooking\s+oil\.)\1\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
