---
id: 66c73a7798f6f62b2ae58f22
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

現在是為你的雜貨列表開頭添加更多項目的時候了。

首先為你的程式碼添加一個 `console.log` 述語，該述語會記錄訊息 `"It might be nice to get a dessert."`

在該 `console` 述語下方，使用正確的陣列方法，將字串 `"Chocolate Cake"` 添加到 `shoppingList` 陣列的開頭。

最後，為你的 `console.log` 添加一個 `console.log`，並在 `console.log` 裡呼叫 `getShoppingListMsg` 函式，以查看更新後的列表已記錄到主控台。

# --hints--

你應該有一個 `console.log`，用來記錄訊息 `"It might be nice to get a dessert."`。

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

你的 `shoppingList` 陣列應該有總共 `6` 個條款。

```js
assert.lengthOf(shoppingList, 6);
```

你應該使用 `unshift` 方法，將字串 `"Chocolate Cake"` 添加到 `shoppingList` 陣列的開頭。

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

你應該在 `console.log` 裡呼叫 `getShoppingListMsg(shoppingList)`。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
