---
id: 66c742d045c9fc2e09fa64b1
title: 第 19 步
challengeType: 1
dashedName: step-19
---

# --description--

對購物列表做的最後一項更改是更新列表中的第一個項目。

在之前的課程中，你學會了如何使用括號表示法和你想更新的項目的索引來更新項目。

這裡是如何更新陣列中一個項目的提醒：

```js
const array = [1, 2, 3, 4, 5];

array[0] = 10;
// Result: [10, 2, 3, 4, 5]
console.log(array); 
```

將 `shoppingList` 陣列中的第一個項目更新為 `"Canola Oil"`。

# --hints--

你應該將 `shoppingList` 陣列中的第一個項目更新為 `"Canola Oil"`。

```js
assert.equal(shoppingList[0], "Canola Oil");
assert.notEqual(shoppingList[1], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

shoppingList.shift();

--fcc-editable-region--

--fcc-editable-region--
```
