---
id: 66c730ee6ae076281721d0b9
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

在之前的課程中，你學會了如何使用 `pop` 方法從陣列的末端移除項目。 

這裡是如何使用 `pop` 方法的提醒：

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

使用 `pop` 方法從 `shoppingList` 陣列中移除最後一個項目。

# --hints--

你的 `shoppingList` 陣列現在應該有 `5` 個條款。

```js
assert.lengthOf(shoppingList, 5);
```

你的 `shoppingList` 陣列中不應再包含字串 `"Potato Chips"`。

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
