---
id: 66c726c34ecf1e238aa9d7d5
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

在之前的課程中，你學會了如何使用 `unshift()` 方法將元素添加到陣列的開頭。 

這裡是如何使用 `unshift()` 方法的提醒：

```js
array.unshift(item1, item2, ..., itemX);
```

使用 `unshift()` 方法將字串 `"Vegetable Oil"` 添加到 `shoppingList` 陣列的開頭。

# --hints--

你的 `shoppingList` 陣列中應該有三個條款。

```js
assert.lengthOf(shoppingList, 3);
```

你應該使用 `unshift()` 方法，將 `"Vegetable Oil"` 添加到 `shoppingList` 陣列的開頭。

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
