---
id: 66c72b0ffbc5522525768558
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

在先前的步驟中，你已經檢視過如何使用 `push` 方法將一個項目添加到陣列的末端。 

但是 `push` 方法接受多個引數，所以你可以像這樣將多個項目添加到陣列的末端：

```js
array.push(item1, item2, item3);
```

在此步驟中，使用 `push` 方法將字串 `"Popcorn"、"Beef Jerky"、"Potato Chips"` 添加到 `shoppingList` 陣列。

順序很重要，所以請確保依照列表中所列的順序添加項目。

# --hints--

你的 `shoppingList` 陣列應該有總共 `6` 個字串。

```js
assert.lengthOf(shoppingList, 6);
```

你的 `shoppingList` 應該包含字串 `"Popcorn"`。

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

你的 `shoppingList` 應該包含字串 `"Beef Jerky"`。

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

你的 `shoppingList` 應該包含字串 `"Potato Chips"`。

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
