---
id: 66c72f4d0528bd268a82107b
title: 步驟 12
challengeType: 1
dashedName: step-12
---

# --description--

現在是時候將更新後的 `shoppingList` 陣列記錄到主控台。

為你的程式碼添加另一個 `console.log`，並在 `console.log` 裡呼叫 `getShoppingListMsg` 函式，以查看更新後的列表記錄在主控台中。

# --hints--

你應該在 `console.log` 裡呼叫 `getShoppingListMsg(shoppingList)`。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 4);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");

--fcc-editable-region--

--fcc-editable-region--
```
