---
id: 66c72b0ffbc5522525768558
title: 步驟11
challengeType: 1
dashedName: step-11
---

# --description--

在前面的步驟中，你回顧瞭如何使用 `push` 方法在數組末尾添加一個項目。

但是 `push` 方法可接受多個參數，因此你可以像這樣在數組末尾添加多個項目：

```js
array.push(item1, item2, item3);
```

在此步驟中，使用 `push` 方法在 `shoppingList` 數組中添加字符串 `"Popcorn", "Beef Jerky", "Potato Chips"`。

順序很重要，因此請確保按照列出的順序添加項目。

# --hints--

你的 `shoppingList` 數組總共應有 `6` 個字符串。

```js
assert.lengthOf(shoppingList, 6);
```

你的 `shoppingList` 應包含字符串 `"Popcorn"`。

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

你的 `shoppingList` 應包含字符串 `"Beef Jerky"`。

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

你的 `shoppingList` 應包含字符串 `"Potato Chips"`。

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
