---
id: 66c72a55418cc9247b710827
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

接著，為你的程式碼添加一個 `console.log`，並在 `console.log` 裡呼叫 `getShoppingListMsg` 函式，以查看更新後的列表記錄在主控台中。

# --hints--

你應該在 `console.log` 裡呼叫 `getShoppingListMsg(shoppingList)`。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
