---
id: 66c73a7798f6f62b2ae58f22
title: 步驟 16
challengeType: 1
dashedName: step-16
---

# --description--

現在是時候在雜貨清單的開頭添加更多物品了。

首先添加 `console.log` 語句，記錄信息 `"It might be nice to get a dessert."`

在 `console` 語句下方，使用正確的數組方法將字符串 `"Chocolate Cake"` 添加到 `shoppingList` 數組的開頭。

最後，添加 `console.log`，並在 `console.log` 中調用 `getShoppingListMsg` 函數，查看控制檯記錄的更新列表。

# --hints--

你應該有一個 `console.log` 來記錄信息 `"It might be nice to get a dessert."`。

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

你的 `shoppingList` 數組總共應有 `6` 個項目。

```js
assert.lengthOf(shoppingList, 6);
```

你應該使用 `unshift` 方法將字符串 `"Chocolate Cake"` 添加到 `shoppingList` 數組的開頭。

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

你應該在 `console.log` 中調用 `getShoppingListMsg(shoppingList)`。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
