---
id: 66c72a55418cc9247b710827
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

接著，為你的 `console.log` 添加一個呼叫 `getShoppingListMsg` 函式，並將 `shoppingList` 陣列作為引數放入 `console.log` 中，以便在主控台看到更新後的列表記錄。

# --hints--

你應該在 `console.log` 裡呼叫 `getShoppingListMsg(shoppingList)`。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
