---
id: 66c74079c30b1c2d166cb9a4
title: 步驟 18
challengeType: 1
dashedName: step-18
---

# --description--

在之前的課程中，你學會了如何使用 `shift` 方法從陣列的開頭移除一個項目。 

這裡是如何使用 `shift` 方法的提醒：

```js
const array = [1, 2, 3, 4, 5];
array.shift();

// Result: [2, 3, 4, 5]
console.log(array); 
```

使用 `shift` 方法從 `shoppingList` 陣列中移除第一個項目。

# --hints--

你的 `shoppingList` 陣列應該有總共 `5` 個條款。

```js
assert.lengthOf(shoppingList, 5);
```

你應該使用 `shift` 方法來移除 `shoppingList` 陣列中的第一個項目。

```js
assert.strictEqual(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

--fcc-editable-region--

--fcc-editable-region--
```
