---
id: 66cbe2319d3845545a293a0b
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

現在是時候查看記錄到主控台的訊息。

為你的 `console.log` 添加一個 `console.log`，並在 `console.log` 內呼叫 `getShoppingListMsg` 函式（程式），以查看訊息是否已記錄到主控台。

# --hints--

你應該在 `console.log` 裡呼叫 `getShoppingListMsg(shoppingList)` 函式（程式）。

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
