---
id: 587d8259367417b2b2512c85
title: Implementieren Sie den Selection-Sort-Algorithmus
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

In diesem Labor implementieren Sie den Selection-Sort-Algorithmus. Selection Sort funktioniert, indem er den kleinsten Wert in einer Liste auswählt und ihn mit dem ersten Wert der Liste vertauscht. Dann beginnt er an der zweiten Position, wählt den kleinsten Wert in der verbleibenden Liste aus und vertauscht ihn mit dem zweiten Element. Er fährt fort, die Liste zu durchlaufen und Elemente zu vertauschen, bis er das Ende der Liste erreicht hat. Nun ist die Liste sortiert. Selection Sort hat in allen Fällen eine quadratische Zeitkomplexität.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und lassen Sie alle Tests erfolgreich durchlaufen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `selectionSort`-Funktion haben.
1. Die `selectionSort`-Funktion sollte ein Array entgegennehmen und ein Array mit denselben Elementen zurückgeben, das aber von klein nach groß sortiert ist.

# --hints--

`selectionSort` sollte eine Funktion sein.

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` sollte ein Array zurückgeben, das bis auf die Reihenfolge unverändert ist.

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` sollte ein sortiertes Array zurückgeben (von klein nach groß).

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` sollte die eingebaute `.sort()`-Methode nicht verwenden.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
