---
id: 612e8279827a28352ce83a72
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

الآن، انسخ مجموعة العناصر السبعة `.key` والصق مجموعتين إضافيتين داخل عنصر `.keys`.

# --hints--

يجب أن يكون لديك 21 عنصرًا من `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

يجب أن يكون لديك 15 عنصرًا من `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

يجب أن تكون كل عناصر `.key` داخل عنصر `.keys` الخاص بك.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

يجب أن تكون كل عناصر `.black--key` داخل عنصر `.keys` الخاص بك.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css

```
