---
id: 612e813b3ba67633222cbe54
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

تذكر، يمكن لخاصية `class` أن تحمل عدة قيم. لتمييز مفاتيحك البيضاء عن السوداء، أضف قيمة ثانية لـ `class` وهي `black--key` إلى عناصر `.key` الثانية والثالثة والخامسة والسادسة والسابعة.

# --hints--

يجب أن يحتوي عنصر `.key` الثاني أيضًا على خاصية `class` بقيمة `black--key`.

```js
const key = document.querySelectorAll('.key')[1];
assert.isTrue(key?.className.includes('black--key'));
```

يجب أن يحتوي عنصر `.key` الثالث على `black--key` في الخاصية `class`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert.isTrue(third?.classList?.contains('black--key'));
```

يجب أن يحتوي عنصر `.key` الخامس على `black--key` في الخاصية `class`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert.isTrue(fifth?.classList?.contains('black--key'));
```

يجب أن يحتوي عنصر `.key` السادس على `black--key` في الخاصية `class`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert.isTrue(sixth?.classList.contains('black--key'));
```

يجب أن يحتوي عنصر `.key` السابع على `black--key` في الخاصية `class`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert.isTrue(seventh?.classList?.contains('black--key'));
```

يجب أن يكون لديك خمسة عناصر `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 5);
```

يجب أن يكون لديك سبعة عناصر `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
--fcc-editable-region--
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css

```
