---
id: 612e89d254fe5d3df7d6693d
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

الآن بعد أن أعدت ضبط نموذج الصندوق في html، تحتاج إلى تطبيق ذلك على العناصر الموجودة بداخله أيضًا. للقيام بذلك، يجب تعيين خاصية `box-sizing` لجميع العناصر الأخرى لتكون موروثة، مما يضمن أن العناصر المستهدفة تعتمد نفس القيمة التي لدى عنصر الأب.

ستحتاج أيضًا إلى استهداف العناصر المزيفة، وهي كلمات مفتاحية خاصة تتبع المحدد. العنصران المزيفان اللذان ستستخدمهما هما `::before` و`::after`.

محدد `::before` ينشئ عنصرًا مزيفًا يكون الطفل الأول للعنصر المحدد، بينما محدد `::after` ينشئ عنصرًا مزيفًا يكون الطفل الأخير للعنصر المحدد. تُستخدم هذه العناصر المزيفة غالبًا لإنشاء محتوى تجميلي، سترى ذلك لاحقًا في هذا المشروع.

في الوقت الحالي، أنشئ محدد CSS لاستهداف جميع العناصر باستخدام `*`، وضمّن العناصر المزيفة `::before` و`::after`. عيّن خاصية `box-sizing` إلى `inherit`.

# --hints--

يجب أن يكون لديك محدد `html`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

يجب أن يحتوي محدد `html` على خاصية `box-sizing` مضبوطة على `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

يجب أن يكون لديك محدد `*, ::before, ::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

يجب أن يحتوي محدد `*, ::before, ::after` على خاصية `box-sizing` مضبوطة على `inherit`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing, 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
