---
id: 612e7d1c29fb872d6384379c
title: خطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

قم بالتسكين لعنصر `div` ثاني داخل عنصر `div` الموجود لديك، واضبط `class` ليكون `keys`.

# --hints--

يجب إنشاء عنصر `div` ثاني.

```js
const divDiv = document.querySelectorAll('div');
assert.lengthOf(divDiv, 2);
```

يجب أن يكون عنصر `div` الجديد داخل عنصر `div` الموجود لديك.

```js
const div = document.querySelector('div');
assert.lengthOf(div?.children, 1);
assert.equal(div?.children?.[0]?.localName, 'div');
```

يجب أن يكون لعنصر `div` الجديد `class` مضبوطًا على `keys`.

```js
const div = document.querySelector('div');
assert.equal(div?.children?.[0]?.className, 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
