---
id: 612e804c54d5e7308d7ebe56
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

داخل عنصر `.keys`، أضف سبعة عناصر `div`، كل منها يحمل الصنف `key`.

# --hints--

يجب إنشاء سبعة عناصر `div` جديدة.

```js
const divDivDiv = document.querySelectorAll('div');
assert.lengthOf(divDivDiv, 9);
```

يجب أن تكون عناصر `div` السبعة الجديدة داخل عنصر `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf([...keys?.children], 7);
[...keys?.children].forEach(
  child => assert.equal(child.tagName, 'DIV')
)
```

يجب أن تحتوي جميع عناصر `div` السبعة الجديدة على الخاصية `class` مضبوطة على `key`.

```js
const keys = document.querySelector('.keys');
[...keys?.children].forEach(child => assert.isTrue(child?.classList.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css

```
