---
id: 587d7b84367417b2b2512b37
title: 작은따옴표와 큰따옴표 혼용 사용 잡기
challengeType: 1
forumTopicId: 301188
dashedName: catch-mixed-usage-of-single-and-double-quotes
---

# --description--

JavaScript는 문자열을 선언할 때 작은따옴표(`'`)와 큰따옴표(`"`)를 모두 사용할 수 있습니다. 어느 쪽을 사용할지는 일반적으로 개인 취향에 달려 있지만, 몇 가지 예외가 있습니다.

문자열에 축약형이나 따옴표로 감싼 다른 텍스트가 있을 때 두 가지 선택지가 있으면 좋습니다. 단, 문자열을 너무 일찍 닫아 구문 오류가 발생하지 않도록 주의하세요.

다음은 따옴표를 혼용한 예시입니다:

```js
const grouchoContraction = "I've had a perfectly wonderful evening, but this wasn't it.";
const quoteInString = "Groucho Marx once said 'Quote me as saying I was mis-quoted.'";
const uhOhGroucho = 'I've had a perfectly wonderful evening, but this wasn't it.';
```

처음 두 개는 올바르지만, 세 번째는 올바르지 않습니다.

물론, 한 가지 스타일의 따옴표만 사용하는 것도 괜찮습니다. 문자열 안의 따옴표는 역슬래시(`\`) 이스케이프 문자를 사용해 이스케이프할 수 있습니다:

```js
const allSameQuotes = 'I\'ve had a perfectly wonderful evening, but this wasn\'t it.';
```

# --instructions--

`href` 값에 대해 서로 다른 따옴표를 사용하거나 기존 따옴표를 이스케이프하도록 문자열을 고치세요. 전체 문자열은 큰따옴표로 감싸세요.

# --hints--

`href` 값 `#Home` 주위의 따옴표를 변경하거나 이스케이프해서 고치세요.

```js
assert(__helpers.removeJSComments(code).match(/<a href=\s*?('|\\")#Home\1\s*?>/g));
```

전체 문자열은 큰따옴표로 감싸야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/"<p>.*?<\/p>";/g));
```

# --seed--

## --seed-contents--

```js
let innerHtml = "<p>Click here to <a href="#Home">return home</a></p>";
console.log(innerHtml);
```

# --solutions--

```js
let innerHtml = "<p>Click here to <a href=\"#Home\">return home</a></p>";
console.log(innerHtml);
```
