---
id: 587d7b86367417b2b2512b3d
title: 유효한 종료 조건으로 무한 루프 방지
challengeType: 1
forumTopicId: 301192
dashedName: prevent-infinite-loops-with-a-valid-terminal-condition
---

# --description--

마지막 주제는 무서운 무한 루프입니다. 루프는 프로그램이 특정 횟수만큼 또는 조건이 충족될 때까지 코드 블록을 실행해야 할 때 매우 유용하지만, 루프를 종료하는 종료 조건이 필요합니다. 무한 루프는 브라우저를 멈추거나 충돌시키고, 프로그램 실행에 혼란을 일으켜 아무도 원하지 않습니다.

이 섹션 소개에서 무한 루프 예제가 있었습니다 — `while` 내부의 `loopy()` 루프에서 종료 조건이 없었습니다. 이 함수를 절대 호출하지 마세요!

```js
function loopy() {
  while(true) {
    console.log("Hello, world!");
  }
}
```

종료 조건이 루프 코드를 언제 벗어날지 프로그램에 알려주는 역할을 하므로, 프로그래머가 이 조건이 결국 도달하도록 보장해야 합니다. 한 가지 오류는 종료 조건과 반대 방향으로 카운터 변수를 증가시키거나 감소시키는 것입니다. 또 다른 오류는 루프 코드 내에서 카운터나 인덱스 변수를 증가시키거나 감소시키는 대신 실수로 재설정하는 것입니다.

# --instructions--

`myFunc()` 함수에는 종료 조건 `i != 4`이 `false`로 평가되어 루프를 종료하지 못하기 때문에 무한 루프가 있습니다 — `i`는 매 반복마다 2씩 증가하여 `i`가 홀수로 시작하므로 4를 건너뛰게 됩니다. 종료 조건의 비교 연산자를 수정하여 루프가 `i`가 4 이하일 때만 실행되도록 하세요.

# --hints--

코드는 `for` 루프의 종료 조건(중간 부분)에 있는 비교 연산자를 변경해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/i\s*?<=\s*?4;/g).length == 1);
```

코드는 루프 종료 조건의 비교 연산자를 수정해야 합니다.

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?!=\s*?4;/g));
```

# --seed--

## --seed-contents--

```js
function myFunc() {
  for (let i = 1; i != 4; i += 2) {
    console.log("Still going!");
  }
}
```

# --solutions--

```js
function myFunc() {
 for (let i = 1; i <= 4; i += 2) {
   console.log("Still going!");
 }
}
```
