---
id: 587d7b86367417b2b2512b3b
title: 인덱싱할 때 발생하는 오프 바이 원 오류 잡기
challengeType: 1
forumTopicId: 301189
dashedName: catch-off-by-one-errors-when-using-indexing
---

# --description--

<dfn>오프 바이 원 오류</dfn>(때때로 OBOE라고도 함)는 문자열이나 배열의 특정 인덱스를 대상으로 하려고 할 때(일부를 자르거나 접근할 때) 또는 인덱스를 반복할 때 발생합니다. JavaScript 인덱싱은 1이 아니라 0부터 시작하므로 마지막 인덱스는 항상 항목 길이보다 1 작습니다. 길이와 같은 인덱스에 접근하려 하면 프로그램이 "인덱스 범위 초과" 참조 오류를 발생시키거나 `undefined`를 출력할 수 있습니다.

인덱스 범위를 인수로 받는 문자열이나 배열 메서드를 사용할 때는 문서를 읽고 해당 범위가 포함 범위인지(주어진 인덱스의 항목이 반환되는 부분에 포함되는지) 이해하는 것이 도움이 됩니다. 다음은 오프 바이 원 오류의 몇 가지 예입니다:

```js
let alphabet = "abcdefghijklmnopqrstuvwxyz";
let len = alphabet.length;
for (let i = 0; i <= len; i++) {
  console.log(alphabet[i]);
}
for (let j = 1; j < len; j++) {
  console.log(alphabet[j]);
}
for (let k = 0; k < len; k++) {
  console.log(alphabet[k]);
}
```

여기 첫 번째 예는 반복 횟수가 한 번 더 많고, 두 번째 예는 한 번 적게 반복하여 첫 번째 인덱스 0을 놓칩니다. 세 번째 예가 올바릅니다.

# --instructions--

다음 함수에서 두 개의 인덱싱 오류를 고쳐 1부터 5까지 모든 숫자가 콘솔에 출력되도록 하세요.

# --hints--

코드는 반복문이 첫 번째 인덱스에서 시작하도록 초기 조건을 설정해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/i\s*?=\s*?0\s*?;/g).length == 1);
```

코드는 반복문이 인덱스를 0부터 시작하도록 초기 조건을 고쳐야 합니다.

```js
assert(!__helpers.removeJSComments(code).match(/i\s?=\s*?1\s*?;/g));
```

코드는 반복문이 마지막 인덱스에서 멈추도록 종료 조건을 설정해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/i\s*<\s*len\s*;|i\s*<=\s*len\s*-\s*1\s*;/g).length == 1);
```

코드는 반복문이 길이보다 1 작은 값에서 멈추도록 종료 조건을 고쳐야 합니다.

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?<=\s*?len;/g));
```

# --seed--

## --seed-contents--

```js
function countToFive() {
  let firstFive = "12345";
  let len = firstFive.length;
  // Only change code below this line
  for (let i = 1; i <= len; i++) {
  // Only change code above this line
    console.log(firstFive[i]);
  }
}

countToFive();
```

# --solutions--

```js
function countToFive() {
 let firstFive = "12345";
 let len = firstFive.length;
 // Only change code below this line
 for (let i = 0; i < len; i++) {
 // Only change code above this line
   console.log(firstFive[i]);
 }
}

countToFive();
```
