---
id: 587d7b83367417b2b2512b37
title: freeCodeCamp 콘솔과 브라우저 콘솔의 차이점 이해하기
challengeType: 1
forumTopicId: 301193
dashedName: understanding-the-differences-between-the-freecodecamp-and-browser-console
---

# --description--

일부 freeCodeCamp 도전 과제에 자체 콘솔이 포함된 것을 보셨을 겁니다. 이 콘솔은 브라우저 콘솔과 약간 다르게 작동합니다.

메시지를 출력하기 위해 `console`와 함께 사용할 수 있는 여러 메서드가 있습니다. 예를 들어 `log`, `warn`, `clear` 등이 있습니다. freeCodeCamp 콘솔은 `log` 메시지만 출력하는 반면, 브라우저 콘솔은 모든 메시지를 출력합니다. 코드를 변경하면 자동으로 실행되어 로그를 보여줍니다. freeCodeCamp 콘솔은 코드가 실행될 때마다 지워집니다.

# --instructions--

먼저, 로그를 볼 수 있도록 브라우저 콘솔을 여세요. 대부분의 브라우저에서 freeCodeCamp 상단 네비게이션 바를 오른쪽 클릭하고 `inspect`를 클릭하면 됩니다. 그런 다음 열리는 창에서 `console` 탭을 찾으세요.

그다음 `console.log`를 사용해 `output` 변수를 로그하세요. 두 콘솔을 보면서 로그를 확인하세요. 마지막으로 로그 뒤에 `console.clear`를 사용해 브라우저 콘솔을 지우세요. 두 콘솔의 차이를 확인하세요.

# --hints--

`console.log()`를 사용해 `output` 변수를 출력하세요.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/console\.log\(output\)/));
```

브라우저 콘솔을 지우려면 `console.clear()`를 사용하세요.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/console.clear\(\)/)
);
```

로그 뒤에 콘솔을 지우세요.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/console\.log\(output\)[\s\S]*console.clear\(\)/)
);
```

# --seed--

## --seed-contents--

```js
let output = "Get this to show once in the freeCodeCamp console and not at all in the browser console";

```

# --solutions--

```js
let output = "Get this to show once in the freeCodeCamp console and not at all in the browser console";

console.log(output);
console.clear();
```
