---
id: 587d7b84367417b2b2512b36
title: '닫히지 않은 소괄호, 괄호, 중괄호, 따옴표 잡기'
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--

알아야 할 또 다른 구문 오류는 모든 여는 소괄호, 괄호, 중괄호, 따옴표가 닫히는 쌍이 있어야 한다는 점입니다. 쌍 중 하나를 빼먹는 실수는 기존 코드를 편집하면서 쌍 중 하나를 삽입할 때 자주 발생합니다. 또한, 콜백 함수를 메서드의 인수로 추가하는 등 코드 블록을 다른 블록 안에 중첩할 때 주의하세요.

이 실수를 피하는 한 가지 방법은 여는 문자를 입력하자마자 바로 닫는 쌍을 포함시키고, 커서를 그 사이로 옮겨 계속 코딩하는 것입니다. 다행히도 대부분 최신 코드 편집기는 쌍의 두 번째 절반을 자동으로 생성해 줍니다.

# --instructions--

코드에서 두 쌍의 오류를 고치세요.

# --hints--

코드에서 배열의 빠진 부분을 고쳐야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

코드에서 `.reduce()` 메서드의 빠진 부분을 고쳐야 합니다. 콘솔 출력에는 `Sum of array values is: 6`가 보여야 합니다.

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```
