---
id: 587d7b85367417b2b2512b38
title: 대입 연산자 대신 동등 연산자 사용 오류 잡기
challengeType: 1
forumTopicId: 301191
dashedName: catch-use-of-assignment-operator-instead-of-equality-operator
---

# --description--

분기 프로그램, 즉 특정 조건이 충족되면 다른 작업을 수행하는 프로그램은 JavaScript에서 `if`, `else if`, `else` 문에 의존합니다. 조건은 때때로 결과가 어떤 값과 같은지 테스트하는 형태를 취합니다.

이 논리는 (적어도 영어로는) "if x equals y, then ..."라고 말하며, 이는 문자 그대로 `=` 또는 대입 연산자를 사용해 코드로 번역될 수 있습니다. 이로 인해 프로그램의 제어 흐름이 예상치 못하게 됩니다.

이전 도전 과제에서 다뤘듯이, JavaScript에서 대입 연산자(`=`)는 변수 이름에 값을 할당합니다. 그리고 `==`와 `===` 연산자는 동등성을 검사합니다(삼중 `===`는 값과 타입이 모두 같은 엄격한 동등성을 테스트합니다).

아래 코드는 `x`에 2를 할당하는데, 이는 `true`로 평가됩니다. JavaScript에서 거의 모든 값은 자체적으로 `true`로 평가되지만, "falsy" 값으로 알려진 `false`, `0`, `""` (빈 문자열), `NaN`, `undefined`, 그리고 `null`는 예외입니다.

```js
let x = 1;
let y = 2;
if (x = y) {

} else {

}
```

이 예제에서 `if` 문 내의 코드 블록은 `y`가 falsy가 아닌 한 `y`의 어떤 값에 대해서도 실행됩니다. 여기서 실행되길 기대하는 `else` 블록은 실제로 실행되지 않습니다.

# --instructions--

조건을 수정해 프로그램이 올바른 분기를 실행하고, `result`에 적절한 값이 할당되도록 하세요.

# --hints--

조건이 대입 대신 동등성을 검사하도록 코드를 수정하세요.

```js
assert(result == 'Not equal!');
```

조건은 동등성을 테스트하기 위해 `==` 또는 `===`를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/x\s*?===?\s*?y/g));
```

# --seed--

## --seed-contents--

```js
let x = 7;
let y = 9;
let result = "to come";

if(x = y) {
  result = "Equal!";
} else {
  result = "Not equal!";
}

console.log(result);
```

# --solutions--

```js
let x = 7;
let y = 9;
let result = "to come";

if(x === y) {
 result = "Equal!";
} else {
 result = "Not equal!";
}

console.log(result);
```
