---
id: 587d7b86367417b2b2512b3c
title: 루프 안에서 변수를 다시 초기화할 때 주의하세요
challengeType: 1
forumTopicId: 301194
dashedName: use-caution-when-reinitializing-variables-inside-a-loop
---

# --description--

때때로 루프 안에서 정보를 저장하거나, 카운터를 증가시키거나, 변수를 다시 설정해야 할 때가 있습니다. 문제가 될 수 있는 상황은 변수를 다시 초기화해야 하는데 하지 않거나, 반대로 초기화하지 말아야 하는데 초기화하는 경우입니다. 특히 종료 조건에 사용되는 변수를 실수로 다시 설정하면 무한 루프가 발생할 수 있어 매우 위험합니다.

루프의 각 사이클마다 `console.log()`를 사용해 변수 값을 출력하면 변수를 다시 설정하거나 초기화하지 않아 발생하는 버그를 발견할 수 있습니다.

# --instructions--

다음 함수는 `m` 행과 `n` 열로 이루어진 2차원 배열을 0으로 채워 생성하려고 합니다. 하지만 외부 루프에서 `row` 변수를 다시 초기화(빈 배열로 설정)하지 않아 예상한 결과가 나오지 않습니다. 코드를 수정해 3x2 크기의 0으로 채워진 배열을 반환하도록 하세요. 배열은 `[[0, 0], [0, 0], [0, 0]]`와 같아야 합니다.

# --hints--

코드는 `matrix` 변수를 3개의 행과 각 행마다 2개의 열로 이루어진 0으로 채워진 배열로 설정해야 합니다.

```js
assert(JSON.stringify(matrix) == '[[0,0],[0,0],[0,0]]');
```

`matrix` 변수는 3개의 행을 가져야 합니다.

```js
assert(matrix.length == 3);
```

`matrix` 변수는 각 행마다 2개의 열을 가져야 합니다.

```js
assert(
  matrix[0].length == 2 && matrix[1].length === 2 && matrix[2].length === 2
);
```

`zeroArray(4,3)`는 4개의 행과 3개의 열로 이루어진 0으로 채워진 배열을 반환해야 합니다.

```js
assert(JSON.stringify(zeroArray(4,3)) == '[[0,0,0],[0,0,0],[0,0,0],[0,0,0]]');
```

# --seed--

## --seed-contents--

```js
function zeroArray(m, n) {
  // Creates a 2-D array with m rows and n columns of zeroes
  let newArray = [];
  let row = [];
  for (let i = 0; i < m; i++) {
    // Adds the m-th row into newArray

    for (let j = 0; j < n; j++) {
      // Pushes n zeroes into the current row to create the columns
      row.push(0);
    }
    // Pushes the current row, which now has n zeroes in it, to the array
    newArray.push(row);
  }
  return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```

# --solutions--

```js
function zeroArray(m, n) {
 // Creates a 2-D array with m rows and n columns of zeroes
 let newArray = [];
 for (let i = 0; i < m; i++) {
   let row = [];
   // Adds the m-th row into newArray

   for (let j = 0; j < n; j++) {
     // Pushes n zeroes into the current row to create the columns
     row.push(0);
   }
   // Pushes the current row, which now has n zeroes in it, to the array
   newArray.push(row);
 }
 return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```
