---
id: 587d7b84367417b2b2512b35
title: 잘못 쓴 변수명과 함수명 잡기
challengeType: 1
forumTopicId: 301186
dashedName: catch-misspelled-variable-and-function-names
---

# --description--

`console.log()`와 `typeof` 메서드는 프로그램 출력의 중간 값과 타입을 확인하는 두 가지 주요 방법입니다. 이제 버그가 흔히 나타나는 형태를 살펴볼 차례입니다. 빠르게 타이핑하는 사람이라면 공감할 수 있는 구문 수준 문제 중 하나는 바로 사소한 철자 오류입니다.

변수명이나 함수명에서 문자가 뒤바뀌거나 빠지거나 대소문자가 잘못된 경우, 브라우저는 존재하지 않는 객체를 찾으려 하며 참조 오류 형태로 불평합니다. JavaScript 변수명과 함수명은 대소문자를 구분합니다.

# --instructions--

코드에서 두 개의 철자 오류를 고쳐 `netWorkingCapital` 계산이 제대로 작동하게 하세요.

# --hints--

netWorkingCapital 계산에 사용된 두 변수의 철자를 확인하세요. 콘솔 출력에 "Net working capital is: 2"가 보여야 합니다.

```js
assert(netWorkingCapital === 2);
```

코드에 잘못 쓴 변수가 없어야 합니다.

```js
assert(!__helpers.removeJSComments(code).match(/recievables/g));
```

`receivables` 변수는 코드에서 올바르게 선언되고 사용되어야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/receivables/g).length == 2);
```

코드에 잘못 쓴 변수가 없어야 합니다.

```js
assert(!__helpers.removeJSComments(code).match(/payable;/g));
```

`payables` 변수는 코드에서 올바르게 선언되고 사용되어야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/payables/g).length == 2);
```

# --seed--

## --seed-contents--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = recievables - payable;
console.log(`Net working capital is: ${netWorkingCapital}`);
```

# --solutions--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = receivables - payables;
console.log(`Net working capital is: ${netWorkingCapital}`);
```
