---
id: 587d7b85367417b2b2512b39
title: 함수 호출 후 누락된 여는 소괄호와 닫는 소괄호 찾기
challengeType: 1
forumTopicId: 301185
dashedName: catch-missing-open-and-closing-parenthesis-after-a-function-call
---

# --description--

함수나 메서드가 인수를 받지 않을 때, 호출할 때 (빈) 여는 소괄호와 닫는 소괄호를 포함하는 것을 잊을 수 있습니다. 종종 함수 호출 결과는 코드 내 다른 용도로 변수에 저장됩니다. 이 오류는 변수 값(또는 타입)을 콘솔에 기록하여 함수가 반환하는 예상 값 대신 함수 참조가 변수에 할당된 것을 확인하면 감지할 수 있습니다.

다음 예제의 변수들은 서로 다릅니다:

```js
function myFunction() {
  return "You rock!";
}
let varOne = myFunction;
let varTwo = myFunction();
```

여기서 `varOne`는 함수 `myFunction`이고, `varTwo`는 문자열 `You rock!`입니다.

# --instructions--

변수 `result`가 함수 `getNine`를 호출하여 반환된 값으로 설정되도록 코드를 수정하세요.

# --hints--

코드는 변수 `result`가 함수 `getNine`가 반환하는 숫자로 설정되도록 고쳐야 합니다.

```js
assert(result == 9);
```

코드는 `getNine` 함수를 호출해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/getNine\(\)/g).length == 2);
```

# --seed--

## --seed-contents--

```js
function getNine() {
  let x = 6;
  let y = 3;
  return x + y;
}

let result = getNine;
console.log(result);
```

# --solutions--

```js
function getNine() {
 let x = 6;
 let y = 3;
 return x + y;
}

let result = getNine();
console.log(result);
```
