---
id: 587d7b83367417b2b2512b33
title: JavaScript 콘솔을 사용해 변수 값을 확인하기
challengeType: 1
forumTopicId: 18372
dashedName: use-the-javascript-console-to-check-the-value-of-a-variable
---

# --description--

Chrome과 Firefox 모두 JavaScript를 디버깅할 수 있는 훌륭한 JavaScript 콘솔, 즉 개발자 도구를 제공합니다.

Chrome 메뉴에서 개발자 도구를, Firefox 메뉴에서 웹 콘솔을 찾을 수 있습니다. 다른 브라우저나 모바일 기기를 사용 중이라면 데스크톱용 Firefox나 Chrome으로 전환하는 것을 강력히 권장합니다.

소괄호 안에 있는 내용을 콘솔에 "출력하는" `console.log()` 메서드는 가장 유용한 디버깅 도구일 것입니다. 코드의 적절한 위치에 배치하면 변수의 중간 값을 확인할 수 있습니다. 출력 결과가 어떻게 나올지 미리 예상하는 습관을 들이는 것이 좋습니다. 코드 전반에 걸쳐 계산 상태를 확인할 수 있는 점검 지점을 두면 문제 발생 위치를 좁히는 데 도움이 됩니다.

다음은 문자열 `Hello world!` 을 콘솔에 출력하는 예입니다:

```js
console.log('Hello world!');
```

# --instructions--

코드에서 표시된 위치에 변수 `console.log()`의 값을 출력하기 위해 `a` 메서드를 사용하세요.

# --hints--

코드에서 변수 `console.log()`의 값을 확인하기 위해 `a`를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\(a\)/g));
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 1;
a++;
// Only change code below this line


let sumAB = a + b;
console.log(sumAB);
```

# --solutions--

```js
var a = 5; console.log(a);
```
