---
id: 587d7b84367417b2b2512b34
title: 변수의 타입을 확인하기 위해 typeof 사용하기
challengeType: 1
forumTopicId: 18374
dashedName: use-typeof-to-check-the-type-of-a-variable
---

# --description--

변수의 자료 구조, 즉 타입을 확인하려면 `typeof`를 사용할 수 있습니다. 여러 데이터 타입을 다룰 때 디버깅에 유용합니다. 두 숫자를 더하는 중이라고 생각했는데 실제로는 하나가 문자열이라면 결과가 예상과 다를 수 있습니다. 타입 오류는 계산이나 함수 호출에서 숨어 있을 수 있으니 주의하세요. 특히 JavaScript Object Notation (JSON) 객체 형태의 외부 데이터를 접근하고 다룰 때 더욱 조심해야 합니다.

다음은 `typeof`를 사용한 예시입니다:

```js
console.log(typeof "");
console.log(typeof 0);
console.log(typeof []);
console.log(typeof {});
```

순서대로 콘솔에는 `string`, `number`, `object`, `object` 문자열이 출력됩니다.

JavaScript는 일곱 가지 원시(불변) 데이터 타입 `Boolean`, `Null`, `Undefined`, `Number`, `String`, `Symbol` (ES6부터 추가), `BigInt` (ES2020부터 추가)와 변경 가능한 항목을 위한 한 가지 타입 `Object`을 인식합니다. JavaScript에서 배열은 기술적으로 객체의 한 종류임을 참고하세요.

# --instructions--

코드에서 두 변수 `console.log()`와 `typeof` 각각의 `seven`를 확인하기 위해 `three` 문장 두 개를 추가하세요.

# --hints--

변수 타입을 확인하려면 두 개의 `typeof` 문장에 `console.log()`를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\s*\(typeof[\( ].*\)?\)/g).length == 2);
```

변수 `typeof`의 타입을 확인하려면 `seven`를 사용하세요.

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]seven\)?/g));
```

변수 `typeof`의 타입을 확인하려면 `three`를 사용하세요.

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]three\)?/g));
```

# --seed--

## --seed-contents--

```js
let seven = 7;
let three = "3";
console.log(seven + three);
// Only change code below this line
```

# --solutions--

```js
let seven = 7;let three = "3";console.log(typeof seven);
console.log(typeof three);
```
