---
id: 587d7b85367417b2b2512b3a
title: 함수 호출 시 잘못된 순서로 전달된 인수 잡기
challengeType: 1
forumTopicId: 301184
dashedName: catch-arguments-passed-in-the-wrong-order-when-calling-a-function
---

# --description--

함수 호출에 관한 논의를 이어가면서, 다음으로 주의해야 할 버그는 함수의 인수가 잘못된 순서로 전달되는 경우입니다. 만약 인수들이 서로 다른 타입이라면, 예를 들어 함수가 배열과 정수를 기대하는데 순서가 바뀌면, 이는 런타임 오류를 발생시킬 가능성이 큽니다. 인수들이 모두 같은 타입(예: 모두 정수)이라면, 코드의 논리가 맞지 않게 됩니다. 이러한 문제를 피하려면 모든 필수 인수를 올바른 순서로 전달하는지 꼭 확인하세요.

# --instructions--

함수 `raiseToPower`는 밑수를 지수로 거듭제곱합니다. 불행히도, 이 함수가 제대로 호출되지 않았습니다 - `power`의 값이 예상한 8이 되도록 코드를 고치세요.

# --hints--

변수 `power`가 3의 2제곱이 아니라 2의 3제곱이 되도록 코드를 고치세요.

```js
assert(power == 8);
```

`raiseToPower` 함수 호출에 올바른 인수 순서를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/raiseToPower\(\s*?base\s*?,\s*?exp\s*?\);/g));
```

# --seed--

## --seed-contents--

```js
function raiseToPower(b, e) {
  return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(exp, base);
console.log(power);
```

# --solutions--

```js
function raiseToPower(b, e) {
 return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(base, exp);
console.log(power);
```
