---
id: 66d6a7a3e1aa411e94bf2346
title: 집 페인팅 만들기
challengeType: 25
dashedName: build-a-house-painting
demoType: onClick
---

# --description--

이 실습에서는 HTML을 사용해 집의 구조를 만듭니다. 그런 다음 CSS 위치 지정으로 창문과 문 같은 집 요소들을 배치하세요.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `div`에 `id`가 `house`로 설정되어 있어야 합니다.
1. `#house`에 자식 요소들이 기준이 될 수 있도록 `position`가 `relative`로 설정되어 있어야 합니다.
1. `#house`는 너비가 `500px`이고 높이가 `400px`이어야 합니다. 최솟값이나 최댓값은 설정하지 마세요.
1. `#house`에 배경색과 테두리가 설정되어 있어야 합니다.
1. `div` 안에 `#house` 요소 다섯 개가 있어야 하며, 각각 `id` 값이 `chimney`, `roof`, `window-1`, `window-2`, `door`이어야 합니다.
1. `#house`의 직계 자식 모두에 `position`가 `absolute`로 설정되어 있어야 합니다.
1. `#roof`, `#chimney`, `#window-1`, `#window-2`, `#door`는 너비, 높이, 테두리, 배경색이 설정되어 있어야 합니다.
1. `#roof`에 top 값이 `0`로 설정되어 있어야 합니다.
1. `#door`은 집의 아래쪽에 배치되어야 합니다.
1. 창문은 `#roof` 아래에 배치되어야 하며 `#door` 높이의 3분의 1 이상 높아야 합니다.
1. 창문과 문 모두 `left` 또는 `right` 중 하나가 집 테두리 안에 배치되도록 설정되어야 합니다.
1. `#chimney`에 top 값이 `#house`의 맨 위에 오도록 설정되어야 합니다.
1. `#chimney`에 `z-index`가 설정되어 집 뒤에 배치되어야 합니다.

**참고:** CSS를 적용하려면 HTML에 스타일시트를 꼭 연결하세요.

# --hints--

`div`에 `id`가 `house`로 설정되어 있어야 합니다.

```js
assert.exists(document.querySelector("div#house"));
```

`#house`를 선택해 `position`를 `relative`로 설정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("position"), "relative");
```

`#house`를 선택해 `border` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("border"));
```

`#house`를 선택해 `background-color` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("background-color"));
```

`#house`에 `min-height` 속성이 없어야 합니다.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("min-height"));
```

`#house`에 `min-width` 속성이 없어야 합니다.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("min-width"));
```

`#house`에 `max-height` 속성이 없어야 합니다.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("max-height"));
```

`#house`에 `max-width` 속성이 없어야 합니다.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("max-width"));
```

`#house`를 선택해 너비를 `500px`로 설정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("width"), "500px")
```

`#house`를 선택해 높이를 `400px`로 설정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("height"), "400px")
```

`div` 안에 `id`가 `chimney`가 `#house`로 설정되어 있어야 합니다.

```js
assert.exists(document.querySelector("div#house div#chimney"));
```

`#chimney`에 `position` 속성이 `absolute`로 설정되어야 합니다.

```js
assert.equal(getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("position"), "absolute");
```

`#chimney`를 선택해 `width` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("width"));
```

`#chimney`를 선택해 `height` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("height"));
```

`#chimney`를 선택해 `border` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("border"));
```

`#chimney`를 선택해 `background-color` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("background-color"));
```

`div` 안에 `id`가 `roof`가 `#house`로 설정되어 있어야 합니다.

```js
assert.exists(document.querySelector("div#house  div#roof"));
```

`#roof`에 `position` 속성이 `absolute`로 설정되어야 합니다.

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("position"), "absolute");
```

`#roof`를 선택해 `width` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("width"));
```

`#roof`를 선택해 `height` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("height"));
```

`#roof`를 선택해 `border` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("border"));
```

`#roof`를 선택해 `background-color` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("background-color"));
```

`div` 안에 `id`가 `window-1`가 `#house`로 설정되어 있어야 합니다.

```js
assert.exists(document.querySelector("div#house  div#window-1"));
```

`#window-1`에 `position` 속성이 `absolute`로 설정되어야 합니다.

```js
assert.equal(getComputedStyle(document.querySelector("#window-1"))?.getPropertyValue("position"), "absolute");
```

`#window-1`를 선택해 `width` 속성을 설정하세요.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window1 || windows);
```

`#window-1`를 선택해 `height` 속성을 설정하세요.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window1 || windows);
```

`#window-1`를 선택해 `border` 속성을 설정하세요.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window1 || windows);
```

`#window-1`를 선택해 `background-color` 속성을 설정하세요.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window1 || windows);
```

`div` 안에 `id`가 `window-2`가 `#house`로 설정되어 있어야 합니다.

```js
assert.exists(document.querySelector("div#house  div#window-2"));
```

`#window-2`에 `position` 속성이 `absolute`로 설정되어야 합니다.

```js
assert.equal(getComputedStyle(document.querySelector("#window-2"))?.getPropertyValue("position"), "absolute");
```

`#window-2`를 선택해 `width` 속성을 설정하세요.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window2 || windows);
```

`#window-2`를 선택해 `height` 속성을 설정하세요.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window2 || windows);
```

`#window-2`를 선택해 `border` 속성을 설정하세요.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window2 || windows);
```

`#window-2`를 선택해 `background-color` 속성을 설정하세요.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window2 || windows);
```

`div` 안에 `id`가 `door`가 `#house`로 설정되어 있어야 합니다.

```js
assert.exists(document.querySelector("div#house  div#door"));
```

`#door`에 `position` 속성이 `absolute`로 설정되어야 합니다.

```js
assert.equal(getComputedStyle(document.querySelector("#door"))?.getPropertyValue("position"), "absolute");
```

`#door`를 선택해 `width` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("width"));
```

`#door`를 선택해 `height` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("height"));
```

`#door`를 선택해 `border` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("border"));
```

`#door`를 선택해 `background-color` 속성을 설정하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("background-color"));
```

`#roof`에 `top` 속성이 `0`로 설정되어야 합니다.

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("top"), "0px");
```

`#door`은 집 아래쪽에 배치되어야 합니다.

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorBottom = Number(door.getPropertyValue("bottom").replace("px", ""));
assert.strictEqual(doorBottom, 0);
```

`left`에 `right` 또는 `#door` 중 하나를 설정해 집 테두리 안에 배치하세요.

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorLeft = Number(door.getPropertyValue("left").replace("px", ""));
const doorRight = Number(door.getPropertyValue("right").replace("px", ""));
assert.isTrue(doorLeft > 0 || doorRight > 0);
```

`#chimney`에 `top` 값이 있어 `#house` 맨 위에 오도록 하세요.

```js
const chimney = getComputedStyle(document.querySelector("#chimney"));
assert.approximately(parseFloat(chimney?.getPropertyValue("top")), - parseFloat(chimney?.getPropertyValue("height")), 2);
```

`#chimney`에 `z-index`가 있어 집 뒤에 배치되도록 하세요.

```js
const houseZIndex = getComputedStyle(document.querySelector("#house"))?.getPropertyValue("z-index");
const chimneyZIndex = getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("z-index");
if (houseZIndex === "auto") {
    assert.isBelow(Number(chimneyZIndex), 0)
} else {
    assert.isBelow(Number(chimneyZIndex), Number(houseZIndex));
}
```

창문은 `#roof` 아래에 배치되어야 하며 `#door` 높이의 3분의 1 이상 높아야 합니다.

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const door = getComputedStyle(document.querySelector("#door"))
const roof = getComputedStyle(document.querySelector("#roof"))
const w1Top = Number(window1.getPropertyValue("top").replace("px", ""));
const w1Bottom = Number(window1.getPropertyValue("bottom").replace("px", ""));
const w2Top = Number(window2.getPropertyValue("top").replace("px", ""));
const w2Bottom = Number(window2.getPropertyValue("bottom").replace("px", ""));
const doorHeight = Number(door.getPropertyValue("height").replace("px", ""));
const roofHeight = Number(roof.getPropertyValue("height").replace("px", ""));
assert.isAbove(w1Top, roofHeight);
assert.isAbove(w1Bottom, doorHeight/3);
assert.isAbove(w2Top, roofHeight);
assert.isAbove(w2Bottom, doorHeight/3);
```

창문에 `left` 또는 `right` 중 하나를 설정해 집 테두리 안에 배치하세요.

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const w1Left = Number(window1.getPropertyValue("left").replace("px", ""));
const w1Right = Number(window1.getPropertyValue("right").replace("px", ""));
const w2Left = Number(window2.getPropertyValue("left").replace("px", ""));
const w2Right = Number(window2.getPropertyValue("right").replace("px", ""));
assert.isTrue(w1Left > 0 || w1Right > 0);
assert.isTrue(w2Left > 0 || w2Right > 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div id="house">
        <div id="chimney">
          <div id="smoke"></div>
        </div>
        <div id="roof"></div>
        <div id="window-1"></div>
        <div id="window-2"></div>
        <div id="door">
          <div id="door-knob"></div>
        </div>
        <div id="welcome">WELCOME</div>
    </div>
</body>
</html>
```

```css
* {
    box-sizing: border-box;
  }

  body {
    margin: 0;
    padding: 0;
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    background: linear-gradient(
      #b3e6ff,
      #b3e6ff 60%,
      green 60%
    );
  }

  #house {
    flex: 0 0 auto;
    width: 500px;
    height: 400px;
    position: relative;
    background-color: #ff9980;
    border: 6px solid #b35900;
  }

  #chimney {
    position: absolute;
    width: 90px;
    height: 100px;
    top: -100px;
    left: 65%;
    background: repeating-linear-gradient(
      #e6e6e6,
      #e6e6e6 10%,
      black 10%,
      black 12%
    );
    border: 2px solid black;
    z-index: -1;
  }

  #smoke {
    width: 40px;
    height: 40px;
    background-color: rgba(0,0,0,0.3);
    position: absolute;
    top: -20px;
    left: 20px;
    border-radius: 50%;
    box-shadow: 0 -10px 10px 10px rgba(0,0,0,0.3);
    border: none;
  }

  #roof {
    border: 2px;
    width: 100%;
    height: 110px;
    background-color: green;
    position: absolute;
    top: 0;
    background: repeating-linear-gradient(
      45deg,
      #b35900,
      #b35900 2%,
      transparent 2%,
      transparent 5%
    ), repeating-linear-gradient(
      -45deg,
      #b35900,
      #b35900 2%,
      #ff9980 2%,
      #ff9980 5%
    )
  }

  #window-1, #window-2 {
    width: 100px;
    height: 100px;
    background: linear-gradient(
      to right,
      #ffffb3,
      #ffffb3 48%,
      #b35900 48%,
      #b35900 52%,
      #ffffb3 52%,
      #ffffb3 100%
    );
    position: absolute;
    top: 40%;
    border: 6px solid #b35900;
  }

  #window-1 {
    left: 7.5%;
  }

  #window-2 {
    right: 7.5%;
  }

  #door {
    width: 130px;
    height: 180px;
    position: absolute;
    bottom: 0%;
    left: 185px;
    background-color: #e6e6e6;
    border: 6px solid #b35900;
    transform: translateY(6px);
  }

  #door-knob {
    width: 20px;
    height: 20px;
    background-color: #b35900;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    right: 5%;
  }

  #welcome {
    font-family: sans-serif;
    font-weight: bold;
    width: 130px;
    height: 40px;
    position: absolute;
    bottom: -47px;
    left: 178px;
    background-color: #85e085;
    border: 2px solid black;
    transform: skewX(-20deg);
    display: flex;
    align-items: center;
    justify-content: center;
  }
```
