---
id: 68e45a2b4ab9b2f48b1cc6df
title: Progetta una pagina di layout per i piani tariffari
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente**

1. La tua pagina dovrebbe avere un elemento `h1` con il testo `Pricing Plans`.
2. La tua pagina dovrebbe avere un elemento `div` con la classe `pricing-container` sotto l'elemento `h1`.
    - Il selettore `.pricing-container` dovrebbe avere una proprietà `display` con il valore `flex` e una proprietà `flex-wrap` con il valore `wrap`.
3. All'interno dell'elemento `.pricing-container` dovresti avere tre elementi `div` con la classe `pricing-card` per rappresentare i piani tariffari.
    - Uno degli elementi `.pricing-card` dovrebbe avere la classe `basic-plan` oltre alla classe `pricing-card`.
    - Uno degli elementi `.pricing-card` dovrebbe avere la classe `pro-plan` oltre alla classe `pricing-card`.
    - Uno degli elementi `.pricing-card` dovrebbe avere la classe `premium-plan` oltre alla classe `pricing-card`.
4. Il tuo elemento `.basic-plan` dovrebbe avere un elemento `h2` con il testo `Basic`.
5. Il tuo elemento `.basic-plan` dovrebbe avere un elemento `p` con il testo `$9/month`.
6. Il tuo elemento `.pro-plan` dovrebbe avere un elemento `h2` con il testo `Pro`.
7. Il tuo elemento `.pro-plan` dovrebbe avere un elemento `p` con il testo `$19/month`.
8. Il tuo elemento `.premium-plan` dovrebbe avere un elemento `h2` con il testo `Premium`.
9. Il tuo elemento `.premium-plan` dovrebbe avere un elemento `p` con il testo `$29/month`.
10. Ognuno dei tuoi elementi `.pricing-card` dovrebbe:
    - Usare Flexbox per impilare il suo contenuto in colonna e giustificare lo spazio tra i figli usando `space-between`.
    - Impostare la proprietà `flex` su `0 0 200px` per dargli una larghezza coerente e impedirne la crescita o la riduzione nel layout.
    - Impostare la proprietà `border` su `2px solid black` per vedere come le diverse schede occupano spazio.
11. L'elemento `.basic-plan` dovrebbe apparire per primo nel layout. Dovresti usare la proprietà `order` per questo.
12. L'elemento `.pro-plan` dovrebbe apparire secondo nel layout. Dovresti usare la proprietà `order` e impostare la sua proprietà `flex-grow` su `2` così occupa più spazio degli altri piani.
13. L'elemento `.premium-plan` dovrebbe venire per ultimo nel layout. Dovresti usare la proprietà `order` per questo.

**Nota:** Assicurati di collegare il tuo foglio di stile nel tuo HTML e applicare il CSS.

# --hints--

Dovresti avere un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

La tua pagina dovrebbe avere un elemento `h1` con il testo `Pricing Plans`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

Dovresti avere un elemento `div` sotto l'elemento `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Il tuo elemento `div` dovrebbe avere la classe `pricing-container`.

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

Dovresti avere un selettore `.pricing-container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

Il tuo selettore `pricing-container` dovrebbe avere la proprietà `display` con valore `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

Il tuo selettore `.pricing-container` dovrebbe avere la proprietà `flex-wrap` con valore `wrap`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

Dovresti avere tre elementi `div` con la classe `pricing-card` all'interno dell'elemento `.pricing-container`.

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

Uno dei tuoi elementi `.pricing-card` dovrebbe avere la classe `basic-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

Uno dei tuoi elementi `.pricing-card` dovrebbe avere la classe `premium-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

Uno dei tuoi elementi `.pricing-card` dovrebbe avere la classe `pro-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

Dovresti avere un selettore `.pricing-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

Il tuo selettore `.pricing-card` dovrebbe avere la proprietà `display` con valore `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

Il tuo selettore `.pricing-card` dovrebbe avere la proprietà `flex-direction` con valore `column`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

Il tuo selettore `.pricing-card` dovrebbe avere la proprietà `justify-content` con valore `space-between`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

Il tuo selettore `.pricing-card` dovrebbe avere la proprietà `flex` con valore `0 0 200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

Il tuo selettore `.pricing-card` dovrebbe avere la proprietà `border` con valore `2px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

Il tuo selettore `.basic-plan` dovrebbe avere la proprietà `order` con valore `0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

Il tuo elemento `.basic-plan` dovrebbe avere un elemento `h2` con il testo `Basic`.

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

Il tuo elemento `.basic-plan` dovrebbe avere un elemento `p` con il testo `$9/month`.

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

Il tuo elemento `.pro-plan` dovrebbe avere un elemento `h2` con il testo `Pro`.

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

Il tuo elemento `.pro-plan` dovrebbe avere un elemento `p` con il testo `$19/month`.

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

Il tuo elemento `.pro-plan` dovrebbe avere la proprietà `order` con valore `1`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

Il tuo elemento `.pro-plan` dovrebbe avere la proprietà `flex-grow` con valore `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

Il tuo elemento `.premium-plan` dovrebbe avere un elemento `h2` con il testo `Premium`.

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

Il tuo elemento `.premium-plan` dovrebbe avere un elemento `p` con il testo `$29/month`.

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

Il tuo elemento `premium-plan` dovrebbe essere l'ultimo elemento nel layout e avere la proprietà `order` con valore `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```
