---
id: 669fc938d38e6e38ace9251e
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

В межах елемента `head` вкладіть елемент `meta` з атрибутом `charset` зі значенням `"UTF-8"`.

Під цим елементом `meta` додайте елемент `title`.

Текстом елемента `title` має бути `Mr. Whiskers' Blog`.

# --hints--

Ви повинні мати елемент `meta`.

```js
assert.isNotNull(document.querySelector("meta"));
```

Елемент `meta` є порожнім елементом, він не повинен мати кінцевий тег `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Тег `meta` повинен мати атрибут `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Атрибут `charset` повинен мати значення `"UTF-8"`.

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

Елемент `meta` має бути вкладеним в елементі `head`.

```js
const meta = document.querySelector('head > meta');
assert.strictEqual(meta?.parentElement?.tagName, 'HEAD');
```

Ви повинні мати початковий тег `title`.

```js
assert.match(code, /<title>/i);
```

Ви повинні мати кінцевий тег `title`.

```js
assert.match(code, /<\/title>/i);
```

Елемент `title` має бути вкладеним в елементі `head`.

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

Елемент `title` повинен мати текст `Mr. Whiskers' Blog`. Можливо, доведеться перевірити написання.

```js
const titleText = document.querySelector('title')?.innerText
assert.strictEqual(titleText, "Mr. Whiskers' Blog");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
  --fcc-editable-region--
</html>
```
