---
id: 612e78af05201622d4bab8aa
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Час почати працювати з піаніно. Створіть елемент `div` всередині вашого елемента `body` з множиною `id`, встановленою на `piano`.

# --hints--

Вам слід створити новий елемент `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Ваш `div` має бути всередині вашого елемента `body`.

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

Ваш `div` має мати множину `id`, встановлену на `piano`.

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
