---
id: 612e77aba7ca691f598feb02
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте два теги `meta`: один для оптимізації вашої сторінки для мобільних пристроїв, а інший — щоб вказати прийнятий `charset` для сторінки.

# --hints--

Ви повинні додати два елементи `meta` на вашу сторінку.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

У вас має бути два елементи `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Один елемент `meta` повинен мати множину `name`, встановлену в `viewport`, а `content` — у `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Інший елемент `meta` повинен мати атрибут `charset`, встановлений у `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
