---
id: 612e6afc009b450a437940a1
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Почніть із базової структури HTML. Додайте `DOCTYPE` оголошення та елементи `html`, `head`, `body` і `title`.

Встановіть мову цієї сторінки на англійську. Встановіть `title` на `Piano`.

# --hints--

Ваш код має містити референс `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Після референса `DOCTYPE` має бути пробіл.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Визначте тип документа як `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Закрийте оголошення `DOCTYPE` кінцевим тегом `>` після типу.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Ваш елемент `html` має мати початковий тег з атрибутом `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш елемент `html` має мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Оголошення `DOCTYPE` має бути на початку вашого HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ви маєте мати початковий тег `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви маєте мати кінцевий тег `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви маєте мати початковий тег `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви маєте мати кінцевий тег `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` і `body` мають бути братськими.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Елемент `head` має бути всередині елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Елемент `body` має бути всередині елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Ваш код має містити елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

У вашого проєкту має бути заголовок `Piano`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

Пам’ятайте, що регістр і написання мають значення для заголовка.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
