---
id: 612e89d254fe5d3df7d6693d
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Тепер, коли ви скинули модель блоку `html`, потрібно передати це й елементам всередині. Для цього можна встановити властивість `box-sizing` у значення `inherit`, що скаже цільовим елементам використовувати те саме значення, що й батьківський елемент.

Також потрібно націлитися на псевдоелементи — це спеціальні ключові слова, які йдуть після селектора. Два псевдоелементи, які ви будете використовувати, — це псевдоелементи `::before` і `::after`.

Селектор `::before` створює псевдоелемент, який є першим дочірнім елементом вибраного елемента, а селектор `::after` створює псевдоелемент, який є останнім дочірнім елементом вибраного елемента. Ці псевдоелементи часто використовують для створення косметичного вмісту, що ви побачите пізніше в цьому проєкті.

Поки що створіть CSS-селектор, щоб націлитися на всі елементи з `*`, і включіть псевдоелементи `::before` і `::after`. Встановіть властивість `box-sizing` у значення `inherit`.

# --hints--

Ви маєте мати селектор `*, ::before, ::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Ваш селектор `*, ::before, ::after` має мати властивість `box-sizing`, встановлену в `inherit`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
