---
id: 612e804c54d5e7308d7ebe56
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Усередині вашого елемента `.keys` додайте сім елементів `div`. Усі вони повинні мати клас `key`.

# --hints--

Ви маєте створити сім нових елементів `div`.

```js
const divDivDiv = document.querySelectorAll('div');
assert(divDivDiv?.length === 9);
```

Ваші сім нових елементів `div` мають бути всередині елемента `.keys`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].length === 7);
assert([...keys?.children].every(child => child?.tagName === 'DIV'));
```

Усі сім нових елементів `div` повинні мати множину `class`, встановлену в `key`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].every(child => child?.classList?.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
