---
id: 662a6bc12cde72c32fb526f0
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Um <dfn>interface</dfn> é como um projeto para uma classe.Uma interface contém um conjunto de métodos e propriedades que uma classe deve implementar.  

Comece este projeto declarando uma classe vazia chamada `Equation`.Você usará esta classe para definir uma interface, um modelo para uma equação genérica.

# --hints--

Você deve definir uma nova classe chamada `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
