---
id: 662bd8260da84bdd5feae419
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Você quer que a classe `LinearEquation` implemente e não simplesmente herde todos os métodos definidos dentro da classe `Equation`, que deve agir como uma interface.

Atualmente, a classe `Equation` é simplesmente a classe pai de `LinearEquation`. Nos próximos passos você aprenderá como transformá-lo em uma interface formal.

Por enquanto, crie uma instância de `Equation` e atribua-a a uma variável `eq` e uma instância de `LinearEquation` e atribua-a a uma variável `lin_eq`.

# --hints--

Você deve declarar uma variável `eq` e atribuir a ela uma instância de `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

Você deve declarar uma variável `lin_eq` e atribuir a ela uma instância de `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
