---
id: 662bd456896f16d9bd03f1a6
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Dentro da classe `Equation`, defina dois novos métodos de instância chamados `solve` e `analyze`.

# --hints--

Você deve definir um método chamado `solve` dentro da classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

Seu método `solve` deve receber um parâmetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

Você deve definir um método chamado `analyze` dentro da classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

Seu método `analyze` deve receber um parâmetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
