---
id: 662bd552e1c1d2db1b88ba47
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Agora, defina outra classe chamada `LinearEquation` e faça com que ela herde de `Equation`.Você usará esta classe para representar equações lineares.

# --hints--

Você deve definir uma classe chamada `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

Sua classe `LinearEquation` deve herdar da classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
