---
id: 662bde88dc84f1e249801b1a
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` significa *Abstract Base Classes*. A classe `ABC` permite que você transforme uma classe comum em uma classe abstrata, que é uma classe que atua como um modelo para classes concretas.

Modifique sua declaração de `import` para importar apenas a classe `ABC` do módulo `abc`. Você pode importar um objeto específico `x` de um módulo `y` seguindo a construção de importação `from y import x`.

Então, transforme sua classe `Equation` em uma classe abstrata fazendo-a herdar de `ABC`.

# --hints--

Você deve importar `ABC` do módulo `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

Sua classe `Equation` deve herdar de `ABC`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
