---
id: 662bdd364bf2cde1487922a9
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Ao contrário de outras linguagens de programação, Python não implementa interfaces em sua linguagem principal, mas a biblioteca padrão do Python permite que você defina interfaces de forma simples.

Para este projeto, você usará utilitários do módulo `abc`. Portanto, importe este módulo no seu código.

# --hints--

Você deve importar o módulo `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
