---
id: 662fbcef5f05e1b84f541a0c
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Cada objeto de equação será instanciado passando tantos argumentos quantos os coeficientes da equação, começando do n-ésimo grau de \( x \) até o grau zero, incluindo o possível coeficiente com o valor `0`.

Por exemplo, `LinearEquation(4, 5)` representaria a equação \( 4x + 5 = 0 \), com `4` sendo o coeficiente do primeiro grau (o maior aqui) e `5` o coeficiente do grau zero.

Você precisa verificar se o número correto de argumentos foi passado para instanciar o objeto de equação.

Dentro do método `__init__`, crie uma instrução `if` para verificar se o comprimento de `args` é diferente do número de coeficientes que a equação deve ter (< code>grau + 1`).Se for, gere um `TypeError` e use a seguinte string para fornecer uma mensagem personalizada: `f"'{self.__class__.__name__}' objeto leva {self.degree + 1} argumentos posicionais mas {len(args)} foram dados"`.

Em seguida, corrija o erro passando `2` e `3` para instanciar `lin_eq`.

# --hints--

Você deve criar uma instrução `if` que verifique se o número de coeficientes usados ​​para instanciar a equação é diferente de `grau + 1`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

Você deve gerar um `Terroroso` na nova instrução `‘id’` e usar a ‘strings’ fornecida para retornar uma mensagem de erro personalizada.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

Você deve passar `2` e `3` para instanciar `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
