---
id: 662fc3eba556a6bf800d48c1
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

A função embutida `isinstance()` recebe dois argumentos e retorna um Booleano indicando se o objeto passado como o primeiro argumento é uma instância da classe passada como o segundo argumento.

```py
isinstance(7, int) # True
```

Outra coisa que você quer verificar é se todo argumento é um número.Após seu primeiro `if`, crie um loop `for` que itere sobre `args` e verifique se o argumento na iteração atual não é uma instância de `int` ou `float`.Use a função `isinstance()` e passe uma tupla contendo `int` e `float` como o segundo argumento.

Se o argumento não for um número, lance um `TypeError` dizendo `"Coefficients must be of type 'int' or 'float'"`.

# --hints--

Você deve criar um loop `for` que itere sobre `args` após sua declaração `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

Você deve criar uma declaração `if` que verifique se o coeficiente atual não é uma instância de `int` ou `float` dentro do loop `for`.

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

Você deve usar a string fornecida para gerar um `TypeError` dentro da declaração `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
