---
id: 662fa2e2cf27c09f21f4f5d0
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Em Python, os tipos de dados são reconhecidos durante a execução (quando o código é executado).Portanto, você não precisa especificar o tipo de dado de uma variável quando a declara.No entanto, você pode anotar uma variável para esclarecer que ela conterá um tipo de dado específico com `variable: <data type> = value` ou apenas `variable: <data type>`.Observe que o interpretador Python não impõe os tipos usados para anotar variáveis e normalmente você precisaria de ferramentas externas para isso.

Dentro da classe `Equation`, defina um atributo de classe `degree`.Não atribua um valor a ele.Em vez disso, use uma anotação de tipo `int` para mostrar que ele armazenará um número inteiro dentro das classes concretas.

Mais tarde, você usará este atributo de classe como parte do processo de validação dos argumentos passados para instanciar os objetos da equação.

# --hints--

Você deve definir um atributo de classe chamado `degree` e anotá-lo com `int` dentro da classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
