---
id: 662f6d7c92381a3049e4c987
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Uma interface não precisa definir apenas métodos abstratos, mas também pode implementar métodos para serem herdados pelas classes concretas.

Antes de cuidar da implementação real de `solve` e `analyze`, dentro da classe `Equation`, defina um método `__init__`.Não use nenhum decorator nele.

# --hints--

Você deve definir um método `__init__` na sua classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

Seu método `__init__` deve receber um parâmetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
