---
id: 662f96576ef178927de87975
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Para ser reconhecido como um método abstrato, um método deve ser decorado com o decorador `@abstractmethod`.

Um <dfn>decorator</dfn> é usado em Python para modificar o comportamento de uma função. Aqui está um exemplo de como usar um decorator chamado `spam`:

```py
@spam
def foo():
    pass
```

Modifique sua declaração de importação para importar o decorador `abstractmethod` e decorar os métodos `solve` e `analyze` da classe `Equation`. Isso gerará duas exceções.

Uma vez que uma classe que herda de `ABC` tenha um método abstrato, a classe não pode mais ser instanciada. Portanto, exclua a instância `Equation` para se livrar do erro.

O outro erro ocorre porque a classe `LinearEquation` deve implementar todos os métodos abstratos definidos na interface. Certifique-se de defini-los dentro da classe `LinearEquation`, também. Você não deve usar o decorator `abstractmethod` na classe concreta.

# --hints--

Você deve importar `abstractmethod` do módulo `abc`.

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

Você deve decorar com `@abstractmethod` o método `solve` dentro da classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

Você deve decorar com `@abstractmethod` o método `analyze` dentro da classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

Você deve definir um método chamado `solve` dentro da classe `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

Seu método `solve` deve receber um parâmetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

Você deve definir um método chamado `analyze` dentro da classe `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

Seu método `solve` deve receber um parâmetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
