---
id: 6639fdcc701833a54c364211
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Após validar os coeficientes, você precisa armazená-los em um atributo de instância. Use uma compreensão de dicionário para criar um dicionário no qual a chave seja o grau do coeficiente e o valor correspondente seja o coeficiente e atribua-o a um atributo chamado `coefficients`.

Por exemplo, um objeto `LinearEquation` instanciado com `2` e `4` deve ter o seguinte atributo `coefficients`: `{1: 2, 0: 4}`, porque `2` corresponde ao primeiro grau de `x` e `4` corresponde ao grau zero de `x`.

Crie os pares chave-valor no seu novo dicionário seguindo a mesma ordem de `args`.

# --hints--

Você deve declarar um atributo chamado `coefficients` dentro do seu método `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_variable("self.coefficients")`)) })
```

Você deve usar uma compreensão de dicionário para armazenar seus coeficientes.

```js
({ test: () => runPython(`
import ast
node = _Node(_code).find_class("Equation").find_function("__init__").find_variable("self.coefficients")
assert isinstance(node.tree.value, ast.DictComp)
`) })
```

Seu atributo `coefficients` deve ser um dicionário contendo pares chave-valor na forma grau-coeficiente. Lembre-se de seguir a mesma ordem na qual os coeficientes são armazenados dentro de `args`.

```js
({ test: () => runPython(`
actual1 = list(LinearEquation(1, 6).coefficients.items())
expected1 = list({1: 1, 0: 6}.items())
actual2 = list(LinearEquation(-3.5, 0).coefficients.items())
expected2 = list({1: -3.5, 0: 0}.items())
assert actual1 == expected1
assert actual2 == expected2
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
--fcc-editable-region--
        
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass

    
lin_eq = LinearEquation(2, 3)
```
