---
id: 6639f947d3a1818c9322c64a
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

A última etapa de validação dos coeficientes é verificar se o coeficiente de maior grau é diferente de zero. Lembre-se de que o coeficiente de maior grau deve ser passado como o primeiro argumento ao instanciar o objeto.

Adicione uma declaração `if` para isso e levante um `ValueError` usando a seguinte string para fornecer uma mensagem personalizada: `'Highest degree coefficient must be different from zero'`.

# --hints--

Você deve criar uma declaração `if` que verifique se o primeiro coeficiente passado para instanciar a equação é igual a zero.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_conditions()[0]
cond.is_equivalent("args[0] == 0") or cond.is_equivalent("0 == args[0]") or cond.is_equivalent("not args[0]")
`)) })
```

Você deve lançar um `ValueError` dentro da nova declaração `if` e usar a string fornecida para retornar uma mensagem de erro personalizada.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_bodies()[0].has_stmt("raise ValueError('Highest degree coefficient must be different from zero')")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
