---
id: 663b7fefd437bd984e091cbf
title: Passo 29
challengeType: 20
dashedName: step-29
---

# --description--

Em seguida, crie uma nova classe chamada `QuadraticEquation` e faça com que ela herde de `Equation`.Você usará esta nova classe para representar equações quadráticas, que são equações de segundo grau na forma $ax^2 + bx + c = 0$. 

Dentro da sua nova classe, defina um atributo de classe `degree` com o valor `2`, que é o grau de uma equação quadrática.Além disso, defina os métodos `solve` e `analyze`.Você cuidará da implementação nas seguintes etapas.

# --hints--

Você deve criar uma nova classe chamada `QuadraticEquation` e fazê-la herdar da classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").inherits_from("Equation")`)) })
```

Você deve definir um método `solve` dentro da classe `QuadraticEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("solve")`)) })
```

Seu método `solve` deve receber um único parâmetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("solve").has_args("self")`)) })
```

Você deve definir um método `analyze` dentro da classe `QuadraticEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("analyze")`)) })
```

Seu método `analyze` deve receber um único parâmetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("analyze").has_args("self")`)) })
```

Você deve definir um atributo de classe `degree` dentro da classe `QuadraticEquation` e atribuir a ele o valor `2`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_variable("degree").is_equivalent("degree = 2")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}
--fcc-editable-region--

--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
