---
id: 663b83a28943e6aa6275a514
title: Passo 19
challengeType: 20
dashedName: step-19
---

# --description--

Ainda dentro da classe `Equation`, defina um método `__str__` para fornecer uma representação em string adequada aos objetos de equação que você vai criar.

Por enquanto, dentro do método `__str__`, declare uma variável `terms` e atribua a ela uma lista vazia.Você usará essa variável para armazenar cada termo (coeficiente vezes \( x^n \)) da sua equação.

Então, declare uma variável `equation_string`, atribua a ela o resultado da junção dos elementos na lista `terms` com um espaço.Finalmente, retorne `equation_string`.

# --hints--

Você deve definir um método `__str__` dentro da classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

Seu método `__str__` deve receber um parâmetro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

Você deve declarar uma variável `terms` e atribuir a ela uma lista vazia dentro do método `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

Você deve declarar uma variável `equation_string` e atribuir a ela o resultado da junção dos elementos em `terms` com um espaço dentro do método `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

Você deve retornar `equation_string` do seu método `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
