---
id: 663c9f31306353460da54542
title: Passo 23
challengeType: 20
dashedName: step-23
---

# --description--

Como você pode ver, o sinal `+` está faltando na saída.O sinal de número é exibido por padrão somente se for negativo.Para alterar esse comportamento, você pode escrever dois pontos após a expressão a ser avaliada dentro das chaves da sua f-string e especificar a opção `+`.Isso permitirá que você exiba o sinal tanto para números positivos e negativos.

Modifique a string em suas duas cláusulas condicionais adicionando `:+` dentro das chaves após `coefficient`.

# --hints--

Você deve modificar a string para anexar à lista `terms` dentro da sua declaração `if` para `f'{coefficient:+}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[0].is_equivalent("terms.append(f'{coefficient:+}')")`)) })
```

Você deve modificar a string para inserir na lista `terms` dentro da sua cláusula `elif` para `f'{coefficient:+}x'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[1].is_equivalent("terms.append(f'{coefficient:+}x')")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient}')
            elif n == 1:
                terms.append(f'{coefficient}x')                
        equation_string = ' '.join(terms)        
        return equation_string
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
