---
id: 664f0389424a6f7aa15fd3e5
title: Passo 36
challengeType: 20
dashedName: step-36
---

# --description--

Outro tipo de asserção lookaround é o lookahead. Lookahead positivo e negativo são indicados por `(?=...)` e `(?!...)`, respectivamente. Eles são usados para corresponder a um padrão se seguido por uma certa sequência de caracteres, que não é consumida:

```py
spam = 'black back bat'
re.sub('a(?=t)', 'o', spam) == 'black back bot' # True
re.sub('a(?!t)', 'o', spam) == 'block bock bat' # True
```

No exemplo acima, o padrão `a(?=t)` contém um lookahead positivo, que é usado para corresponder ao caractere `a` somente quando seguido por um `t`. Na última linha do exemplo, o padrão `a(?!t)` contém um lookahead negativo, que é usado para corresponder ao caractere `a` somente se **não** for seguido por um `t`. Novamente, em ambos os casos, o caractere contido no lookahead não é consumido.

Adicione um lookahead positivo ao seu padrão regex para que o caractere `1` seja substituído somente se for seguido pelo caractere `x`.

# --hints--

Você deve modificar seu padrão regex usando um lookahead positivo para substituir o caractere `1` somente se for seguido por um `x`. Não remova o lookbehind negativo do seu padrão.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("re.sub(r'(?<!\\d)1(?=x)', '', equation_string.strip('+'))")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re


class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f"{coefficient:+}x**{n}")
        equation_string = ' '.join(terms) + ' = 0'
--fcc-editable-region--
        return re.sub(r'(?<!\d)1', '', equation_string.strip('+'))        
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2

    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c
   
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(11, -1, 1)
print(quadr_eq)

```
