---
id: 66759e32b88fb5459b1e0234
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

O método `__init_subclass__` é chamado sempre que a classe que o define é subclassificada e ele permite personalizar as classes filhas. O método recebe um parâmetro chamado por convenção `cls` (que significa "class"), que representa a nova subclasse.

Defina um método `__init_subclass__` na sua classe `Equation` e dê a ele um parâmetro `cls`.

# --hints--

Você deve definir um método `__init_subclass__` com um parâmetro `cls` na sua classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
