---
id: 6675aaf418b41157f6ccd692
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

É hora de voltar para o método `__init__`.Dependendo do tipo de equação, você precisará passar um número variável de argumentos durante a instanciação.

Adicione um segundo parâmetro `args` ao método e use o operador `*` para fazer com que ele aceite um número variável de argumentos.

# --hints--

Seu método `__init__` deve receber dois parâmetros, `self` e `*args`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
