---
id: 6675a38a8b535e4ff3274520
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

A função embutida `hasattr` recebe um objeto como seu primeiro argumento e uma string que representa o nome de um atributo como seu segundo argumento. Retorna um booleano indicando se o objeto possui o atributo especificado.

Agora você vai usar o método `__init_subclass__` para verificar se a classe filha possui o atributo `degree` no momento da instanciação.

Crie uma declaração `if` para verificar se `cls` não possui um atributo `degree`. Se for o caso, levante um `AttributeError` e use a string `f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"` para fornecer uma mensagem personalizada.

Depois disso, corrija o erro que apareceu no terminal declarando um atributo de classe `degree` dentro da classe `LinearEquation`.  This attribute should represent the degree of the equation, which is the exponent of the highest \\( x \\) term. Portanto, atribua o inteiro `1` ao atributo `degree`.

# --hints--

Você deve criar uma declaração `if` que verifique se `cls` não possui o atributo `degree` dentro do método `__init_subclass__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_conditions()[0].is_equivalent("not hasattr(cls, 'degree')")`)) })
```

Você deve levantar um `AttributeError` usando a string fornecida dentro do seu `if` statement.

```js
({ test: () => runPython(`
raise_stmt = 'raise AttributeError(f"Cannot create \\'{cls.__name__}\\' class: missing required attribute \\'degree\\'")'
node = _Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_bodies()[0]
assert node.has_stmt(raise_stmt) 
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        pass

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    
--fcc-editable-region--
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
