---
id: 66ebd4ae2812430bb883c787
title: Costruisci un hub per eventi
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

In questo laboratorio utilizzerai gli elementi HTML semantici per creare la struttura di una pagina web. Aggiungerai contenuti e immagini per farla sembrare un vero hub per eventi.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `header`.
1. All’interno dell’elemento `header` dovresti avere un elemento `h1` che contiene il testo `Event Hub` e un elemento `nav`.
1. All’interno dell’elemento `nav` dovresti avere una lista non ordinata di due elementi contenenti collegamenti a diverse sezioni della pagina. Il primo elemento dovrebbe avere il testo `Upcoming Events` e il secondo elemento dovrebbe avere il testo `Past Events`.
1. Ogni collegamento dovrebbe essere rappresentato da un elemento `a` con un attributo `href` che collega alla rispettiva sezione della pagina, `#upcoming-events` e `#past-events` rispettivamente.
1. Dovresti avere un elemento `main` che contiene le diverse sezioni della pagina.
1. All’interno dell’elemento `main` dovresti avere due elementi `section`.
1. Il primo elemento `section` dovrebbe avere un attributo `id` con il valore `upcoming-events`.
1. All’interno della sezione `#upcoming-events` dovresti avere:
   
   - Un elemento `h2` con il testo `Upcoming Events`.
   - Due elementi `article`. Ogni articolo dovrebbe rappresentare un evento e dovrebbe avere:
      - Un elemento `h3` per il titolo dell’evento.
      - Un elemento `p` per la descrizione dell’evento. Puoi aggiungere una data in fondo se vuoi.

1. Il secondo elemento `section` dovrebbe avere un attributo `id` con il valore `past-events`.
1. All’interno della sezione `#past-events` dovresti avere:

   - Un elemento `h2` con il testo `Past Events`.
   - Due elementi `article`. Ogni elemento articolo dovrebbe rappresentare un evento passato e dovrebbe avere:
     - Un elemento `h3` per il titolo dell’evento,
     - Un elemento `p` per la descrizione dell’evento. Puoi aggiungere una data in fondo se vuoi.
     - Un elemento immagine con l’attributo `src` che punta a un file immagine e l’attributo `alt` con una descrizione dell’immagine.

**Nota:** Puoi usare qualsiasi testo per le descrizioni e le date degli eventi. Se vuoi, puoi usare i seguenti URL per le immagini:

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

Dovresti avere un elemento `header`.

```js
assert.isNotNull(document.querySelector("header"));
```

Il tuo elemento `header` dovrebbe venire dopo il tag di apertura `body`.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

All’interno dell’elemento `header` dovresti avere un elemento `h1` che contiene il testo `Event Hub`.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

All’interno dell’elemento `header`, dopo l’elemento `h1`, dovresti avere un elemento `nav`.

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

Il tuo elemento `nav` dovrebbe contenere una lista non ordinata di due elementi.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

Il primo elemento nella lista non ordinata dovrebbe essere un collegamento.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

Il secondo elemento nella lista non ordinata dovrebbe essere un collegamento.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

Il testo del primo elemento nella lista non ordinata dovrebbe essere `Upcoming Events`.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

Il primo elemento nella lista non ordinata dovrebbe avere l’attributo `href` impostato su `#upcoming-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

Il testo del secondo elemento nella lista non ordinata dovrebbe essere `Past Events`.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

Il secondo elemento nella lista non ordinata dovrebbe avere l’attributo `href` impostato su `#past-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

Dovresti avere un elemento `main` dopo il tag di chiusura dell’elemento `header`.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

All’interno dell’elemento `main` dovresti avere due elementi `section`.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

Il tuo primo elemento `section` dovrebbe avere un attributo `id` con il valore `upcoming-events`.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

Il tuo secondo elemento `section` dovrebbe avere un attributo `id` con il valore `past-events`.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

All’interno della sezione `#upcoming-events` dovresti avere un elemento `h2` con il testo `Upcoming Events`.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

All’interno della sezione `#upcoming-events` dovresti avere due elementi `article` sotto l’elemento `h2`.

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Entrambi gli elementi `article` all’interno della sezione `#upcoming-events` dovrebbero avere un elemento `h3` per il titolo dell’evento.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Entrambi gli elementi `article` all’interno della sezione `#upcoming-events` dovrebbero avere un elemento paragrafo per la descrizione dell’evento.

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

All’interno della sezione `#past-events` dovresti avere un elemento `h2` con il testo `Past Events`.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

All’interno della sezione `#past-events` dovresti avere due elementi `article` sotto l’elemento `h2`.

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Entrambi gli elementi `article` all’interno della sezione `#past-events` dovrebbero avere un elemento `h3` per il titolo dell’evento.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Entrambi gli elementi `article` all’interno della sezione `#past-events` dovrebbero avere un elemento paragrafo per la descrizione dell’evento.

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Entrambi gli elementi `article` all’interno della sezione `#past-events` dovrebbero avere un elemento immagine.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

Entrambi gli elementi immagine all’interno della sezione `#past-events` dovrebbero avere l’attributo `src` che punta a un file immagine.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

Entrambi gli elementi immagine all’interno della sezione `#past-events` dovrebbero avere l’attributo `alt` con una descrizione dell’immagine.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

Ogni elemento `h3` dovrebbe avere il titolo dell’evento.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

Ogni elemento `p` dovrebbe avere la descrizione dell’evento.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```
