---
id: 587d824a367417b2b2512c44
title: Controllo del prezzo delle azioni
challengeType: 4
forumTopicId: 301572
dashedName: stock-price-checker
---

# --description--

Crea un'app JavaScript full-stack che sia funzionalmente simile a questa: <a href="https://stock-price-checker.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker.freecodecamp.rocks/</a>.

Poiché tutte le API affidabili per i prezzi delle azioni richiedono una chiave API, abbiamo creato una soluzione alternativa. Usa <a href="https://stock-price-checker-proxy.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker-proxy.freecodecamp.rocks/</a> per ottenere informazioni aggiornate sui prezzi delle azioni senza dover registrare una tua chiave.

Lavorare su questo progetto ti richiederà di scrivere il tuo codice usando uno dei seguenti metodi:

-   Clona <a href="https://github.com/freeCodeCamp/boilerplate-project-stockchecker/" target="_blank" rel="noopener noreferrer nofollow">questo repository GitHub</a> e completa il progetto localmente.
-   Usa un site builder a tua scelta per completare il progetto. Assicurati di includere tutti i file dal nostro repository GitHub.

# --instructions--

1.  Imposta la variabile d'ambiente `NODE_ENV` su `test`, senza virgolette
2.  Completa il progetto in `routes/api.js` o creando un gestore/controller
3.  Aggiungerai tutte le funzionalità di sicurezza in `server.js`
4.  Creerai tutti i test funzionali in `tests/2_functional-tests.js`

**Nota** Considerazioni sulla privacy: Poiché è richiesto che venga accettato solo 1 like per IP, dovrai salvare gli indirizzi IP. È importante rispettare le leggi sulla privacy dei dati come il Regolamento generale sulla protezione dei dati (GDPR). Un'opzione è ottenere il permesso per salvare i dati dell'utente, ma è molto più semplice anonimizzarli. Per questa sfida, ricorda di anonimizzare gli indirizzi IP prima di salvarli nel database. Se ti servono idee su come farlo, puoi scegliere di hashare i dati, troncarli o impostare parte dell'indirizzo IP a 0.

Scrivi i seguenti test in `tests/2_functional-tests.js`:

-   Visualizzare un'azione: richiesta GET a `/api/stock-prices/`
-   Visualizzare un'azione e metterle like: richiesta GET a `/api/stock-prices/`
-   Visualizzare la stessa azione e metterle like di nuovo: richiesta GET a `/api/stock-prices/`
-   Visualizzare due azioni: richiesta GET a `/api/stock-prices/`
-   Visualizzare due azioni e metterle like: richiesta GET a `/api/stock-prices/`

# --hints--

Puoi fornire il tuo progetto, non l'URL di esempio.

```js
  assert(
    !/.*\/stock-price-checker\.freecodecamp\.rocks/.test(code)
  );
```

Dovresti impostare le policy di sicurezza dei contenuti per permettere solo il caricamento di script e CSS dal tuo server.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("script-src 'self'")
  );
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("style-src 'self'")
  );
```

Puoi inviare una richiesta `GET` a `/api/stock-prices`, passando un simbolo azionario NASDAQ come parametro di query `stock`. L'oggetto restituito conterrà una proprietà chiamata `stockData`.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  assert.property(parsed, 'stockData');
```

La proprietà `stockData` include il simbolo `stock` come stringa, il `price` come numero e `likes` come numero.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker.price, 'number');
  assert.typeOf(ticker.likes, 'number');
  assert.typeOf(ticker.stock, 'string');
```

Puoi anche passare un campo `like` come `true` (booleano) per aggiungere il tuo like all'azione o alle azioni. Deve essere accettato solo 1 like per IP.

```js

```

Se passi 2 azioni, il valore restituito sarà un array con informazioni su entrambe. Invece di `likes`, mostrerà `rel_likes` (la differenza tra i like delle due azioni) per entrambi gli oggetti `stockData`.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG&stock=MSFT'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker, 'array');
  assert.property(ticker[0], 'rel_likes');
  assert.property(ticker[1], 'rel_likes');
```

Tutti e 5 i test funzionali sono completi e superati.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 5);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
  });
```
