---
id: 587d824a367417b2b2512c45
title: Bacheca messaggi anonima
challengeType: 4
forumTopicId: 301568
dashedName: anonymous-message-board
---

# --description--

Crea un'app JavaScript full-stack che sia funzionalmente simile a questa: <a href="https://anonymous-message-board.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://anonymous-message-board.freecodecamp.rocks/</a>.

Lavorare su questo progetto ti porterà a scrivere il tuo codice usando uno dei seguenti metodi:

-   Clona <a href="https://github.com/freeCodeCamp/boilerplate-project-messageboard/" target="_blank" rel="noopener noreferrer nofollow">questo repository GitHub</a> e completa il progetto localmente.
-   Usa un site builder a tua scelta per completare il progetto. Assicurati di includere tutti i file dal nostro repository GitHub.

# --instructions--

1.  Imposta `NODE_ENV` per testare senza virgolette quando sei pronto a scrivere i test e DB alla stringa di connessione del tuo database (in `.env`)
2.  Si consiglia di creare controller/gestori e gestire il routing in `routes/api.js`
3.  Aggiungerai tutte le funzionalità di sicurezza in `server.js`

Scrivi i seguenti test in `tests/2_functional-tests.js`:

-   Creare un nuovo thread: richiesta POST a `/api/threads/{board}`
-   Visualizzare i 10 thread più recenti con 3 risposte ciascuno: richiesta GET a `/api/threads/{board}`
-   Eliminare un thread con password errata: richiesta DELETE a `/api/threads/{board}` con un `delete_password` non valido
-   Eliminare un thread con password corretta: richiesta DELETE a `/api/threads/{board}` con un `delete_password` valido
-   Segnalare un thread: richiesta PUT a `/api/threads/{board}`
-   Creare una nuova risposta: richiesta POST a `/api/replies/{board}`
-   Visualizzare un singolo thread con tutte le risposte: richiesta GET a `/api/replies/{board}`
-   Eliminare una risposta con password errata: richiesta DELETE a `/api/replies/{board}` con un `delete_password` non valido
-   Eliminare una risposta con password corretta: richiesta DELETE a `/api/replies/{board}` con un `delete_password` valido
-   Segnalare una risposta: richiesta PUT a `/api/replies/{board}`

# --hints--

Puoi fornire il tuo progetto, non l'URL di esempio.

```js
  assert(
    !/.*\/anonymous-message-board\.freecodecamp\.rocks/.test(
      code
    )
  );
```

Consenti che il tuo sito venga caricato in un iFrame solo sulle tue pagine.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-frame-options']?.includes('SAMEORIGIN'));
```

Non consentire il prefetch DNS.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-dns-prefetch-control']?.includes('off'));
```

Consenti che il tuo sito invii il referrer solo per le tue pagine.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['referrer-policy']?.includes('same-origin'));
```

Puoi inviare una richiesta POST a `/api/threads/{board}` con dati di modulo che includono `text` e `delete_password`. Il record salvato nel database avrà almeno i campi `_id`, `text`, `created_on` (data e ora), `bumped_on` (data e ora, inizia uguale a `created_on`), `reported` (booleano), `delete_password` e `replies` (array).

```js
  const date = new Date();
  const text = `fcc_test_${date}`;
  const deletePassword = 'delete_me';
  const data = { text, delete_password: deletePassword };
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(url + '/api/threads/fcc_test');
    const parsed = await checkData.json();
    try {
      assert.equal(parsed[0].text, text);
      assert.isNotNull(parsed[0]._id);
      assert.equal(new Date(parsed[0].created_on).toDateString(), date.toDateString());
      assert.equal(parsed[0].bumped_on, parsed[0].created_on);
      assert.isArray(parsed[0].replies);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Puoi inviare una richiesta POST a `/api/replies/{board}` con dati di modulo che includono `text`, `delete_password` e `thread_id`. Questo aggiornerà la data `bumped_on` alla data del commento. Nell'array `replies` del thread verrà salvato un oggetto con almeno le proprietà `_id`, `text`, `created_on`, `delete_password` e `reported`.

```js
  const url = code;
  const body = await fetch(url + '/api/threads/fcc_test');
  const thread = await body.json();

  const date = new Date();
  const text = `fcc_test_reply_${date}`;
  const delete_password = 'delete_me';
  const thread_id = thread[0]._id;
  const replyCount = thread[0].replies.length;

  const data = { text, delete_password, thread_id };
  const res = await fetch(url + '/api/replies/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
    const parsed = await checkData.json();
    try {
      assert.equal(parsed.replies.length, replyCount + 1);
      assert.equal(parsed.replies[0].text, text);
      assert.equal(parsed._id, thread_id);
      assert.equal(parsed.bumped_on, parsed.replies[0].created_on);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Puoi inviare una richiesta GET a `/api/threads/{board}`. Verrà restituito un array dei 10 thread più recenti spinti in alto sulla bacheca con solo le 3 risposte più recenti per ciascuno. I campi `reported` e `delete_password` non saranno inviati al client.

```js
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test');

  if (res.ok) {
    const threads = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isAtMost(threads.length, 10);
      for (let i = 0; i < threads.length; i++) {
        assert.containsAllKeys(threads[i], ["_id", "text", "created_on", "bumped_on", "replies"]);
        assert.isAtMost(threads[i].replies.length, 3);
        assert.notExists(threads[i].delete_password);
        assert.notExists(threads[i].reported);
        for (let j = 0; j < threads[i].replies.length; j++) {
          assert.notExists(threads[i].replies[j].delete_password);
          assert.notExists(threads[i].replies[j].reported);
        }
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Puoi inviare una richiesta GET a `/api/replies/{board}?thread_id={thread_id}`. Verrà restituito l'intero thread con tutte le risposte, escludendo gli stessi campi dal client come nel test precedente.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);

  if (res.ok) {
    const thread = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isObject(thread);
      assert.containsAllKeys(thread, ["_id", "text", "created_on", "bumped_on", "replies"]);
      assert.isArray(thread.replies);
      assert.notExists(thread.delete_password);
      assert.notExists(thread.reported);
      for (let i = 0; i < thread.replies.length; i++) {
        assert.notExists(thread.replies[i].delete_password);
        assert.notExists(thread.replies[i].reported);
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Puoi inviare una richiesta DELETE a `/api/threads/{board}` e passare `thread_id` e `delete_password` per eliminare il thread. Verrà restituita la stringa `incorrect password` o `success`.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  let data = { thread_id, delete_password: "wrong_password" };
  const res_invalid = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  data = { thread_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    const not_deleted = await res_invalid.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      assert.equal(not_deleted, "incorrect password");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Puoi inviare una richiesta DELETE a `/api/replies/{board}` e passare `thread_id`, `reply_id` e `delete_password`. Verrà restituita la stringa `incorrect password` o `success`. In caso di successo, il testo di `reply_id` sarà cambiato in `[deleted]`.

```js
  const url = code;

  const thread_data = {
    text: "fcc_test_thread",
    delete_password: "delete_me",
  };
  await fetch(`${url}/api/threads/fcc_test`, {
    method: "POST",
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(thread_data)
  });
  let res = await fetch(`${url}/api/threads/fcc_test`);
  let threads = await res.json();
  const thread_id = threads[0]._id;
  
  const reply_data = { thread_id, text: "fcc_test_reply", delete_password: "delete_me" };
  await fetch(`${url}/api/replies/fcc_test`, {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(reply_data)
  });
  res = await fetch(`${url}/api/threads/fcc_test`);
  threads = await res.json();
  const reply_id = threads[0].replies[0]._id;

  const data = { thread_id, reply_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/replies/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
      const thread = await res.json();
      assert.equal(thread._id, thread_id);
      assert.equal(thread.replies[0]._id, reply_id);
      assert.equal(thread.replies[0].text, "[deleted]");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Puoi inviare una richiesta PUT a `/api/threads/{board}` e passare `thread_id`. Verrà restituita la stringa `reported`. Il valore `reported` di `thread_id` sarà cambiato in `true`.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const data = { thread_id };
  
  res = await fetch(`${url}/api/threads/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Puoi inviare una richiesta PUT a `/api/replies/{board}` e passare `thread_id` e `reply_id`. Verrà restituita la stringa `reported`. Il valore `reported` di `reply_id` sarà cambiato in `true`.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const reply_id = threads[0].replies[0]._id;
  const data = { thread_id, reply_id };
  
  res = await fetch(`${url}/api/replies/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Tutti e 10 i test funzionali sono completi e superati.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 10);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
    assert.isAtLeast(test.assertions.length, 1);
  });
```
