---
id: 698b42492a7f5f737895892e
title: 构建一个邻接列表到矩阵的转换器
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

在本实验中，你将构建一个函数，将图的邻接列表表现转换为邻接矩阵。邻接列表是一个对象，其中每个密钥代表一个节点，且对应的值是一个数组，表示该密钥节点连接的节点。邻接矩阵是一个二维数组，其中位置 `[i][j]` 的条目为 `1` 表示存在从节点 `i` 到节点 `j` 的边，否则为 `0`。

例如，给定邻接列表：

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

对应的邻接矩阵为：

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该定义一个名为 `adjacencyListToMatrix` 的函数，将邻接列表转换为邻接矩阵。
2. 该函数应接受一个表现无权（无向或有向）`Graph` 邻接列表的`object`作为其`参数`。
3. 该函数应该：
   - 将邻接列表转换为邻接矩阵。
   - 打印邻接矩阵中的每一行。
   - 返回邻接矩阵。

例如，`adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})` 应该打印：

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

并返回 `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`。

# --hints--

你应该定义一个名为 `adjacencyListToMatrix` 的函数。

```js
assert.isFunction(adjacencyListToMatrix);
```

`adjacencyListToMatrix` 函数应有一个参数。

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

该函数应正确地从邻接列表中确定节点数。

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

该函数应正确将矩阵值设置为 `1`，用于存在的边。

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

该函数应打印矩阵的每一行。

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

该函数应返回邻接矩阵。

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

当给定邻接列表 `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}` 时，函数应返回 `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`。

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

当给定邻接列表 `{0: [1], 1: [0]}` 时，函数应返回 `[[0, 1], [1, 0]]`。

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

当给定邻接列表 `{0: [], 1: [], 2: []}` 时，函数应返回 `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`。

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
