---
id: 5d5a813321b9e3db6c106a46
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript – це потужна мова, яка дозволяє створювати інтерактивні вебсайти.

*Примітка*: для всіх інших проєктів у цій навчальній програмі вам потрібно буде мати базові знання HTML та CSS. Якщо ви новачок у HTML та CSS, пройдіть [сертифікацію з адаптивного вебдизайну](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

To get started, create your standard HTML boilerplate with a `DOCTYPE`, `html`, `head`, and `body`, then add a `meta` tag for the `charset`. Add a `title` element and use the text `RPG - Dragon Repeller` for it. Include a `link` tag for your stylesheet to link the `styles.css` file.

Вкінці створіть елемент `div` з `id` зі значенням `game` в межах `body`.

# --hints--

Код повинен містити `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

You should include a space after the `DOCTYPE` reference.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Визначте тип документа як `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

You should close the `DOCTYPE` declaration with a `>` after the type.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Елемент `html` повинен мати початковий тег. Don't forget the `lang` attribute.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Елемент `html` повинен мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Your `DOCTYPE` declaration should be at the beginning of your HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ви повинні мати початковий тег `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви повинні мати кінцевий тег `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви повинні мати початковий тег `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви повинні мати кінцевий тег `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

The `head` and `body` elements should be siblings.

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

Елемент `head` має бути в межах елемента `html`.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

Елемент `body` має бути в межах елемента `html`.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

Код повинен мати елемент `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Елемент `meta` повинен мати атрибут `charset` зі значенням `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Код повинен мати елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ви повинні мати кінцевий тег `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Ваш код повинен мати елемент `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Ви повинні мати елемент `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Елемент `div` повинен мати атрибут `id` зі значенням `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Елемент `div` має бути в межах елемента `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
