---
id: 62a115879a6d51422652cbfc
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

У вашій рольовій грі користувачі зможуть відстежувати своїй стати, купувати зброю та битися з монстрами. Перш ніж ви зможете продовжити інтерактивну частину гри у JavaScript, вам потрібно спочатку створити елементи HTML, які відображатимуть інформацію про гру.

Створіть чотири елементи `div` в межах елемента `#game`. Give them the following respective `id` values, in order: `stats`, `controls`, `monsterStats`, and `text`.

# --hints--

Створіть чотири нові елементи `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

Надайте одному з нових елементів `div` властивість `id` зі значенням `stats`.

```js
assert.exists(document.querySelector('div#stats'));
```

Надайте одному з нових елементів `div` властивість `id` зі значенням `controls`.

```js
assert.exists(document.querySelector('div#controls'));
```

Надайте одному з нових елементів `div` властивість `id` зі значенням `monsterStats`.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

Надайте одному з нових елементів `div` властивість `id` зі значенням `text`.

```js
assert.exists(document.querySelector('div#text'));
```

You should place the new `div` elements in the correct order.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

You should place the new `div` elements within the `#game` element.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
