---
id: 62a3b41c9494f937560640ab
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Оголосіть ще дві змінні з іменами `monsterHealth` та `inventory`.

Для змінної `inventory` призначте до неї значення масиву, що містить рядок `"stick"`.

Пам’ятайте, що ви працювали з масивами у попередньому проєкті таким чином:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

Використайте `let`, щоб оголосити змінну `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

You should use camelCase to name your `monsterHealth` variable.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` should not have a value.

```js
assert.isUndefined(monsterHealth);
```

You should not assign a value to your `monsterHealth` variable.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

Використайте `let`, щоб оголосити змінну `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

Змінна `inventory` досі має бути масивом.

```js
assert.isArray(inventory);
```

Змінна `inventory` повинна мати лише одне значення.

```js
assert.lengthOf(inventory, 1);
```

Змінна `inventory` повинна містити рядок `"stick"`.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
