---
id: 62a3a0a3c0a4b32915d26a6e
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Вилучіть рядок `console.log("Hello World");`, щоб почати писати код свого проєкту.

У попередньому проєкті ви навчилися як працювати зі змінними та ключовим словом `let` таким чином:

```js
let camperbot;
```

Ви також вивчили, як ініціалізувати змінну значенням таким чином:

```js
let age = 32;
```

Використайте ключове слово `let`, щоб оголосити змінну під назвою `xp` та призначте до неї значення `0`, число.

# --hints--

You should not have a `console.log("Hello World");` line in your code.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Використайте ключове слово `let`, щоб оголосити змінну.

```js
assert.match(code, /let/);
```

Оголосіть змінну під назвою `xp`.

```js
assert.match(code, /let\s+xp/);
```

`xp` повинна мати значення `0`.

```js
assert.equal(xp, 0);
```

Ініціалізуйте змінну `xp` значенням `0`.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
