---
id: 655cd899f8de09431eabb40c
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Merge sort ni algoriti ya kupanga inayotumia kanuni ya kugawanya-na-kushinda kupanga makusanyo ya data. Hiyo ni, inagawa mkusanyo katika sehemu ndogo ndogo, na kushinda sehemu hizo kwa kuzipanga kwa uhuru, kisha kuunganisha sehemu zilizopangwa.

Anza kwa kutangaza kitendakazi kinachoitwa `merge_sort` kinachopokea kigezo `array`. Kitendakazi hiki kitashughulikia zoezi la kupanga orodha ya nambari. Tumia neno kuu la `pass` ndani ya mwili wa kitendakazi.

# --hints--

Unapaswa kutangaza kitendakazi chenye jina `merge_sort` chenye kigezo kimoja: `array`. Usisahau kutumia neno kuu la `pass` ndani ya mwili wa kitendakazi.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
