---
id: 6564aee9c077774ea49c3faf
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Algoriti ya kupanga kwa kutumia merge sort hufanya vitendo vitatu:

- Gawanya mfululizo usiopangwa wa vitu katika sehemu ndogo ndogo
- Panga vitu katika sehemu ndogo ndogo
- Unganisha sehemu ndogo ndogo zilizopangwa

Haya yote hufanyika kwa mfululizo wa kurudia hadi sehemu ndogo ndogo ziumanishwe kuwa mfululizo mzima uliopangwa. Anza kwa kugawanya mfululizo.

Kwanza, badilisha neno la `pass` na variable `middle_point`. Tumia operator ya kugawanya kwa nambari kamili (`//`) kugawanya urefu wa orodha ya `array` kwa 2 na weka thamani hiyo kwenye variable mpya ya `middle_point`. Kumbuka kuingiza msimbo wako kwa mpangilio sahihi.

# --hints--

Haupaswi kuwa na `pass` ndani ya mwili wa kitendakazi.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

Unapaswa kutangaza variable ya `middle_point` na kuwekea `len(array) // 2` thamani yake.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
