---
id: 656627b47bd2d2a4afbc945d
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Katika hatua iliyopita, ulipata nambari ya katikati. Unaweza kuitumia kugawanya `array` vipande viwili na kuweka kila sehemu katika vigezo vipya.

Tumia sintaksia ya slice kutoa nusu ya kushoto ya `array` na weka thamani yake katika kigezo kinachoitwa `left_part`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `left_part` katika kitendakazi chako cha `merge_sort`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

Unapaswa kutumia sintaksia ya slice kutoa nusu ya kushoto ya safu ya data na weka thamani yake katika kigezo cha `left_part`.

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
--fcc-editable-region--
    
--fcc-editable-region--
```
