---
id: 656639771fed09b5c0e8fe71
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Sasa baada ya kugawanya orodha ya `array` kuwa orodha mbili tofauti, utaendelea kugawanya kila orodha hadi kila kipengele kiwe katika orodha yake binafsi. Orodha yenye nambari moja tu daima imepangwa.

Ili kufanya hivyo, itumie wito wa `merge_sort` kwa mfululizo ndani ya kitendakazi chako na pita `left_part` kama hoja kwa wito huo.

# --hints--

Unapaswa kuita `merge_sort` ukiwa na hoja `left_part` chini kabisa ya mwili wa kitendakazi chako.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
--fcc-editable-region--
    
--fcc-editable-region--
```
