---
id: 656680b0fc79f2c38a34d90e
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Sasa ni wakati wa kupanga na kuunganisha orodha (`left_part` na `right_part`) katika orodha ya awali.

Unaweza kufanya hivi kwa kulinganisha vipengele vyote kwenye orodha zote mbili, na kuunganisha kipengele kidogo zaidi kwenye orodha kuu. Utafanya kulinganisha hii kwa kielezo zote katika `left_part` na `right_part`.

Tengeneza vigezo vitatu: `left_array_index`, `right_array_index`, na `sorted_index` na uweke thamani zao kuwa `0`. Vigezo hivi vitakusaidia kufuatilia kila kielezo wakati wa mchakato wa kupanga.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `left_array_index` ndani ya kitendakazi chako cha `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

Thamani ya `left_array_index` inapaswa kuwa `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

Unapaswa kuwa na kigezo kinachoitwa `right_array_index` ndani ya kitendakazi chako cha `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

Thamani ya `right_array_index` inapaswa kuwa `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

Unapaswa kuwa na kigezo kinachoitwa `sorted_index` ndani ya kitendakazi chako cha `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

Thamani ya `sorted_index` inapaswa kuwa `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)
--fcc-editable-region--
    
--fcc-editable-region--
```
