---
id: 656709e50ed928da35260276
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Katika hatua iliyopita, uliweka thamani ya kipengele katika `left_part` kwenye orodha ya `array` kwa sababu kilikuwa kidogo zaidi. Lakini hili halitakuwa hali kila wakati. Katika baadhi ya kesi za kulinganisha, kipengele kilichopo upande wa kulia kinaweza kuwa kidogo zaidi.

Tengeneza kifungu cha `else` kitakachotekelezwa wakati kipengele katika `left_part` si kidogo kuliko kipengele katika `right_part`.

Ndani ya blokki ya `else`, weka thamani ya `right_part[right_array_index]` kwa `array[sorted_index]`.

# --hints--

Unapaswa kuwa na `array[sorted_index] = right_part[right_array_index]` ndani ya mwili wa kifungu chako cha `else`.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_bodies()[1].find_body()[0].is_equivalent("array[sorted_index] = right_part[right_array_index]")
`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
