---
id: 656710d1e0ec62253426db24
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Mzunguko wa `while` uliouunda unalinganisha kipengele kimoja kutoka `left_part` na kingine kutoka `right_part`, kisha unaongeza kipengele kidogo zaidi kwenye orodha kuu ya `array`.

Utaendelea na operesheni hii hadi pasipokuwepo tena vipengele vya kulinganisha. Lakini `left_part` bado inaweza kuwa na vipengele wakati `right_part` haina, na kinyume chake.

Unda mzunguko mwingine wa `while` unaoendesha wakati `left_array_index` iko chini ya `len(left_part)`. Katika hatua zinazofuata, utaitumia kunakili vipengele vilivyobaki katika `left_part` kwenda kwenye orodha ya `array`.

# --hints--

Mzunguko wako wa `while` unapaswa kuwa na sharti hili: `left_array_index < len(left_part)`. Usisahau kutumia neno kuu la `pass` ndani ya mwili wa mzunguko.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

Unapaswa kuwa na neno kuu la `pass` ndani ya mwili wa mzunguko wako wa `while`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
    
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
