---
id: 656702f8b4cbd8cbf0a433c6
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Ifuatayo, unahitaji kuunda mzunguko unaoendelea mradi tu kuna vipengele vinavyobaki katika `left_part` na `right_part`.

Kwa ajili hiyo, unda mzunguko wa `while` wenye masharti mawili: moja linalokagua kama `left_array_index` ni ndogo kuliko urefu wa `left_part` na sharti lingine linalokagua kama `right_array_index` ni ndogo kuliko urefu wa `right_part`.

# --hints--

Mzunguko wako wa `while` unapaswa kukagua kama `left_array_index` ni ndogo kuliko urefu wa `left_part` na kama `right_array_index` ni ndogo kuliko urefu wa `right_part`.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

Tumia neno kuu la `pass` ndani ya mwili wa mzunguko.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
