---
id: 6569b4e0bd29d17d4b53b16c
title: Hatua 19
challengeType: 20
dashedName: step-19
---

# --description--

Sasa, utaiga mantiki ile ile ya mzunguko wa `while` kwa `right_part`.

Tengeneza mzunguko wa `while` unaoendelea wakati `right_array_index` ni mdogo kuliko `len(right_part)`. Ndani ya mzunguko wa `while`, weka thamani ya `right_part[right_array_index]` kwa `array[sorted_index]`. Mwishoni, ongeza `right_array_index` na `sorted_index` kwa `1`.

# --hints--

Mzunguko wako wa `while` unapaswa kukagua kama `right_array_index` ni mdogo kuliko `len(right_part)`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

Unapaswa kuweka thamani ya `right_part[right_array_index]` kwa `array[sorted_index]` ndani ya mwili wa mzunguko.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

Unapaswa kuongeza `right_array_index` kwa `1` ndani ya mzunguko wa `while`. Kumbuka kuwa ongezeko linapaswa kufanyika baada ya kuweka kipengele cha sehemu ya kulia kwenye safu ya data iliyopangwa.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

Unapaswa kuongeza thamani ya sasa ya `sorted_index` kwa `1` ndani ya mzunguko wa `while`. Kumbuka kuwa ongezeko linapaswa kufanyika baada ya kuweka kipengele cha sehemu ya kulia kwenye safu ya data iliyopangwa.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
    
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
    
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
