---
id: 6569b743630ee592a65a7e2f
title: Hatua 20
challengeType: 20
dashedName: step-20
---

# --description--

Kabla ya kupima kitendakazi cha `merge_sort()`, unahitaji kuunda kesi ya msingi inayosimamisha utekelezaji wa kitendakazi wakati urefu wa `array` ni chini au sawa na `1`.

Kesi hii ya msingi itasimamisha wito la mzunguko. Bila yake, operesheni ya kupanga kwa kutumia merge sort ingekuwa inaendelea hata wakati orodha imeshapangwa au haina vipengele ndani yake.

Mara moja baada ya tamko la kitendakazi, unda tamko la `if` lenye sharti hili: `len(array) <= 1`. Ndani ya mwili wa tamko la `if` ongeza tamko la `return` kusimamisha utekelezaji wa kitendakazi.

# --hints--

Unapaswa kuunda tamko la `if` ili kukagua kama `len(array) <= 1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

Unapaswa kuwa na tamko la `return` ndani ya mwili wa tamko la `if`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_ifs()[0].find_body().has_stmt("return")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    
--fcc-editable-region--
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
    
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
    
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
