---
id: 6569b19d31a09b65b4bc390a
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Bado ukiwa ndani ya mzunguko wa `while`, ongeza thamani ya `left_array_index` kwa 1.

# --hints--

Unapaswa kuongeza thamani ya sasa ya `left_array_index` kwa 1 ndani ya mzunguko wa `while`. Kumbuka kuwa ongezeko linapaswa kufanywa baada ya kuweka kipengele cha sehemu ya kushoto katika safu ya data iliyopangwa.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[1].find_body()
    
    cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "left_array_index += 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "left_array_index = left_array_index + 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "left_array_index = 1 + left_array_index") 
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
        
--fcc-editable-region--
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]

--fcc-editable-region--
```
