---
id: 6577254891048c8f2c19e961
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Chini kabisa ya mwili wa kitendakazi chako, fanya wito wa kitendakazi cha `merge_sort()` tena. Wakati huu, pitisha `right_part` kama hoja kwa wito wa kitendakazi.

# --hints--

Unapaswa kufanya wito wa `merge_sort()` ukiwa na hoja `right_part` chini kabisa ya mwili wa kitendakazi chako.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
