---
id: 657f3a661730891aa64f3568
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Tumia sintaksia ya slice kuchukua nusu ya kulia ya `array` na weka thamani yake katika kigezo kinachoitwa `right_part`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `right_part`

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

Unapaswa kutumia sintaksia ya slice kuchukua nusu ya kulia ya safu ya data na weka thamani yake katika kigezo cha `right_part`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
