---
id: 69b83e35f19ba26ba1fa517a
title: Construir um sequenciador de semáforo
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

Neste laboratório, você vai simular ciclos configuráveis de semáforo e registrar anomalias.

Você vai trabalhar com objetos de configuração que descrevem as fases de um semáforo. Cada objeto de configuração tem as seguintes propriedades:

- `fault`: uma flag booleano que dispara a terminação antecipada quando `true`.
- `phases`: um array de objetos de fase.

Cada objeto de fase dentro de `phases` tem as seguintes propriedades:

- `color`: uma string que representa a cor do sinal (`"green"`, `"yellow"` ou `"red"`).
- `duration`: um inteiro positivo que representa quanto tempo a fase dura em segundos.

Você pode consultar os objetos `config1`, `config2`, `config3` e `config4` fornecidos como exemplos de possíveis objetos de configuração.

**Nota**: não adicione declarações extras de `console.log()`, pois podem fazer os testes falharem.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

# user stories

1. Você deve ter uma função chamada `runSequence` com dois parâmetros: `config` e `cycles`, onde `config` representa um objeto de configuração e `cycles` representa o número máximo de vezes que a sequência pode ser executada antes de parar.

1. Você deve implementar `runSequence(config, cycles)` usando um loop `for` ou `while` para iterar através de cada fase ao longo do número dado de ciclos.

1. A função `runSequence` deve:
    - Registrar `No phases found` e retornar imediatamente se `config.phases` estiver vazio.
    - Registrar `Faulted phase!` e parar a simulação antecipadamente se `config.fault` estiver definido como `true`.
    - Registrar `Invalid phase detected` se `duration <= 0`.
    - Registrar `Switching to [color] for [duration] s` para cada fase válida. Substitua `[color]` e `[duration]` pelas propriedades correspondentes do objeto de fase.
    - Por exemplo, `runSequence(config1, 1)` deve registrar:

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. Você deve ter uma função chamada `generateTimeline` com dois parâmetros: `config` e `cycles`.

1. A função `generateTimeline` deve:
    - Registrar o tempo acumulado decorrido após cada fase ao longo dos ciclos em um array, adicionando a `duration` de cada fase ao total acumulado enquanto você itera.
    - Processar todas as fases com falha e inválidas sem validação, mesmo se `config.fault` for `true` ou `duration <= 0`.
    - Retornar o array de timestamps acumulados.
    - Por exemplo, `generateTimeline(config1, 1)` deve retornar o array `[5, 7, 11]`.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

Você não deve fazer alterações nos objetos `config1`, `config2`, `config3` e `config4`.

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

Você deve ter uma função chamada `runSequence`.

```js
assert.isFunction(runSequence);
```

Sua função `runSequence` deve aceitar dois parâmetros, `config` e `cycles`.

```js
const regex = __helpers.functionRegex('runSequence', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regex);
```

`runSequence(config1, 1)` deve registrar `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` nessa ordem.

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` deve registrar `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`, `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` nessa ordem.

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` deve registrar `Switching to red for 3 s`, `Invalid phase detected`, `Switching to green for 6 s` nessa ordem.

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` deve registrar `Faulted phase!`.

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` deve registrar `No phases found`.

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

Você deve ter uma função chamada `generateTimeline`.

```js
assert.isFunction(generateTimeline);
```

Sua função `generateTimeline` deve aceitar dois parâmetros, `config` e `cycles`.

```js
const regEx = __helpers.functionRegex('generateTimeline', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regEx);
```

`generateTimeline(config1, 1)` deve retornar `[5, 7, 11]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` deve retornar `[5, 7, 11, 16, 18, 22]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` deve retornar `[3, 1, 7, 10, 8, 14]`.

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` deve retornar `[5, 7, 13]`.

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` deve retornar `[]`.

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
