---
id: 69b83e35f19ba26ba1fa517a
title: Construir um sequenciador de semĂĄforo
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

Neste laboratĂłrio, vocĂȘ vai simular ciclos configurĂĄveis de semĂĄforo e registrar anomalias.

VocĂȘ vai trabalhar com objetos de configuraĂ§ĂŁo que descrevem as fases de um semĂĄforo. Cada objeto de configuraĂ§ĂŁo tem as seguintes propriedades:

- `fault`: uma flag booleano que dispara a terminaĂ§ĂŁo antecipada quando `true`.
- `phases`: um array de objetos de fase.

Cada objeto de fase dentro de `phases` tem as seguintes propriedades:

- `color`: uma string que representa a cor do sinal (`"green"`, `"yellow"` ou `"red"`).
- `duration`: um inteiro positivo que representa quanto tempo a fase dura em segundos.

VocĂȘ pode consultar os objetos `config1`, `config2`, `config3` e `config4` fornecidos como exemplos de possĂ­veis objetos de configuraĂ§ĂŁo.

**Nota**: nĂŁo adicione declaraĂ§Ă”es extras de `console.log()`, pois podem fazer os testes falharem.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratĂłrio.

# user stories

1. VocĂȘ deve ter uma funĂ§ĂŁo chamada `runSequence` com dois parĂąmetros: `config` e `cycles`, onde `config` representa um objeto de configuraĂ§ĂŁo e `cycles` representa o nĂșmero mĂĄximo de vezes que a sequĂȘncia pode ser executada antes de parar.

1. VocĂȘ deve implementar `runSequence(config, cycles)` usando um loop `for` ou `while` para iterar atravĂ©s de cada fase ao longo do nĂșmero dado de ciclos.

1. A funĂ§ĂŁo `runSequence` deve:
    - Registrar `No phases found` e retornar imediatamente se `config.phases` estiver vazio.
    - Registrar `Faulted phase!` e parar a simulaĂ§ĂŁo antecipadamente se `config.fault` estiver definido como `true`.
    - Registrar `Invalid phase detected` se `duration <= 0`.
    - Registrar `Switching to [color] for [duration] s` para cada fase vĂĄlida. Substitua `[color]` e `[duration]` pelas propriedades correspondentes do objeto de fase.
    - Por exemplo, `runSequence(config1, 1)` deve registrar:

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. VocĂȘ deve ter uma funĂ§ĂŁo chamada `generateTimeline` com dois parĂąmetros: `config` e `cycles`.

1. A funĂ§ĂŁo `generateTimeline` deve:
    - Registrar o tempo acumulado decorrido apĂłs cada fase ao longo dos ciclos em um array, adicionando a `duration` de cada fase ao total acumulado enquanto vocĂȘ itera.
    - Processar todas as fases com falha e invĂĄlidas sem validaĂ§ĂŁo, mesmo se `config.fault` for `true` ou `duration <= 0`.
    - Retornar o array de timestamps acumulados.
    - Por exemplo, `generateTimeline(config1, 1)` deve retornar o array `[5, 7, 11]`.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

VocĂȘ nĂŁo deve fazer alteraĂ§Ă”es nos objetos `config1`, `config2`, `config3` e `config4`.

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

VocĂȘ deve ter uma funĂ§ĂŁo chamada `runSequence`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.runSequence);
```

Sua funĂ§ĂŁo `runSequence` deve aceitar dois parĂąmetros, `config` e `cycles`.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.runSequence?.parameters;
assert.equal(parameters?.length, 2);
assert.equal(parameters?.[0]?.toString(), "config");
assert.equal(parameters?.[1]?.toString(), "cycles");
```

`runSequence(config1, 1)` deve registrar `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` nessa ordem.

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` deve registrar `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`, `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` nessa ordem.

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` deve registrar `Switching to red for 3 s`, `Invalid phase detected`, `Switching to green for 6 s` nessa ordem.

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` deve registrar `Faulted phase!`.

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` deve registrar `No phases found`.

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

VocĂȘ deve ter uma funĂ§ĂŁo chamada `generateTimeline`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.generateTimeline);
```

Sua funĂ§ĂŁo `generateTimeline` deve aceitar dois parĂąmetros, `config` e `cycles`.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.generateTimeline?.parameters;
assert.equal(parameters?.length, 2);
assert.equal(parameters?.[0]?.toString(), "config");
assert.equal(parameters?.[1]?.toString(), "cycles");
```

`generateTimeline(config1, 1)` deve retornar `[5, 7, 11]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` deve retornar `[5, 7, 11, 16, 18, 22]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` deve retornar `[3, 1, 7, 10, 8, 14]`.

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` deve retornar `[5, 7, 13]`.

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` deve retornar `[]`.

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
