---
id: 683ec95f2f8781355556ff81
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

アルゴリズムの流れをこの初期段階で理解するために、動作をテストする必要があります。

そのためには、まずループから抜け出す必要があります。現在の実装では1回の繰り返ししか許されないため、条件が満たされないと無限ループになってしまいます。

`if`文の直後に、`break`キーワードを使って`while`ループから抜け出してください。そして、`while`ループの後に、値が見つからなかったことを示すために空のリストを返してください。

# --hints--

`break`文の後に`while`ループから抜け出すために`if`キーワードを使うべきです。

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

`while`ループの後に空のリストを返すべきです。

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--
```
