---
id: 683ec95f2f8781355556ff7b
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

二分探索では、最も低いインデックスと最も高いインデックスがあります。

最も低いインデックスは、現在探索している値の範囲の左端の境界を表します。この場合は `search_list` です。一方、最も高いインデックスは、探索している値の右端の境界を表します。

これらの両方を考慮するために、探索しているリスト全体を考えなければなりません。そこで、`binary_search` 関数の中で、値が `low` の `0` 変数と、探索しているリストの最後のインデックスを表す値の `high` 変数を定義してください。

# --hints--

`low` 関数の中に `binary_search` という名前の変数があるはずです。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

`low` 変数は `0` に設定されているはずです。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

`high` 関数の中に `binary_search` という名前の変数があるはずです。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

`high` 変数は `len(search_list) - 1` に設定されているはずです。

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []
    
--fcc-editable-region--
```
