---
id: 683ec95f2f8781355556ff7e
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`mid` インデックスを計算したので、次はそのインデックスにある `search_list` から実際に値を取得する必要があります。この値が、検索対象である `value` パラメータと比較する値になります。

`value_at_middle` 変数を作成し、`search_list` の `mid` インデックスにある要素を代入してください。

# --hints--

`value_at_middle` 変数を作成し、`search_list[mid]` に設定してください。

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2

        
--fcc-editable-region--
```
