---
id: 683ec95f2f8781355556ff7d
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

ループを始めるには、現在の探索範囲の中央のインデックスを見つける必要があります。

中央のインデックスは、`low` と `high` の変数の平均を計算することで求められます。2つの値を足した後に床除算を使って平均を求める必要があるため、答えは常に切り捨てられた整数になります。インデックスは整数でなければならないからです。

`pass` のループから `while` キーワードを削除し、`mid` と `low` の平均を設定した `high` 変数を作成してください。

# --hints--

while ループから `pass` キーワードを削除し、`mid` を設定した `(low + high) // 2` 変数を作成してください。

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        pass
--fcc-editable-region--
```
