---
id: 683ec95f2f8781355556ff7f
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

今度は、ターゲットかどうかに関わらず、探索中に通ったステップを追跡するために `value_at_middle` を `path_to_target` リストに含めてください。

# --hints--

`append()` メソッドを使って `value_at_middle` を `path_to_target` リストに追加してください。

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid] \\npath_to_target.append(value_at_middle)")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        
--fcc-editable-region--
```
