---
id: 683ec95f2f8781355556ff82
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

関数は `search_list` というリストと、探している `value` を受け取ることを思い出してください。

では、`binary_search([1, 2, 3, 4, 5], 3)` と `binary_search([1, 2, 3, 4, 5, 9], 4)` を使って関数を呼び出し、呼び出し結果をすぐに出力してください。

# --hints--

関数を `binary_search([1, 2, 3, 4, 5], 3)` で呼び出し、すぐに出力してください。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

関数を `binary_search([1, 2, 3, 4, 5, 9], 4)` で呼び出し、すぐに出力してください。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        break
    return []
          
--fcc-editable-region--

--fcc-editable-region--
```
