---
id: 683ec95f2f8781355556ff80
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

次に行うのは、対象の `value` が中央にあるかどうかを確認する条件を作成することです。

`while` ループの中で、対象の `if` が `value` と等しいかどうかをチェックする `value_at_middle` 文を作成してください。もしそうなら、`path_to_target` 変数を返してください。

# --hints--

`if` が `value` と等しいかどうかをチェックする `value_at_middle` 文を作成してください。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

`path_to_target` 文の中で `if` 変数を返してください。

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
--fcc-editable-region--
        
--fcc-editable-region--
```
