---
id: 683ec95f2f8781355556ff7c
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

検索範囲を定義したので、次は二分探索アルゴリズムを実行するループを作成しましょう。

二分探索は、検索範囲を繰り返し狭めていくことで動作します。この処理は、チェックすべき有効な範囲がある限り続きます。

これは、`while` ポインターが `low` ポインター以下である限り続く `high` ループで表現できます。

もし `low` が `high` より大きくなった場合、検索範囲が空になったことを意味し、`value` はリストに存在しません。

`binary_search` 関数の中で、`while` が high 以下であるかを条件にした `low` ループを作成してください。ループの中には、今はプレースホルダーとして `pass` キーワードを追加しましょう。

# --hints--

`while` が `low` 以下の間、実行される `high` ループを作成してください。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
--fcc-editable-region--
    
--fcc-editable-region--
```
