---
id: 683ec95f2f8781355556ff7a
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

作成している二分探索アルゴリズムは、値が見つかったかどうかを示すために`True`や`False`だけを返すわけではありません。値を見つけるまでの経路も返します。そのために、経路を追跡する`path_to_target`変数を用意します。

`binary_search`関数の中で、まず`pass`キーワードを`path_to_target`変数に置き換え、空のリストで初期化してください。

# --hints--

`path_to_target`関数内に空のリストをセットした`binary_search`変数を作成してください。`pass`キーワードは忘れずに削除しましょう。

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    pass
--fcc-editable-region--
```
