---
id: 683ec95f2f8781355556ff83
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

最初の呼び出しが `[3]` を返すのがわかります。これは `3` が `[1, 2, 3, 4, 5]` の中央の値だからです。

一方、2回目の呼び出しは空のリストを返します。これは、中央の値を確認した後、ループが現在 `low` または `high` の変数を更新して検索範囲を狭めるロジックを持っていないために起こります。単に中央の値を1回チェックして、それ以上の検索をせずに終了してしまいます。

2分探索が検索範囲を狭め続けられるように、`elif` が `value` より大きいかどうかをチェックする `value_at_middle` ブロックを追加してください。今は `pass` キーワードを `elif` の中に追加してください。

# --hints--

`elif` が `value` より大きい場合のために `value_at_middle` ブロックを追加してください。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--          
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
