---
id: 683ec39565dbc7e3375263b1
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

前のレッスンで、二分探索と線形探索の違いについて学びました。

このワークショップでは、二分探索アルゴリズムを作成して実践します。

まず、`binary_search` 関数を `search_list` と `value` のパラメータで作成してください。

`search_list` は関数が探索する要素のリストで、`value` は関数が `search_list` の中から見つけようとするターゲットです。

テストを通過させるために、今は関数内に `pass` キーワードを追加してください。

# --hints--

`binary_search` 関数を作成してください。

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

`binary_search` 関数は `search_list` と `value` のパラメータを持つべきです。

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
