---
id: 683ec95f2f8781355556ff84
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

`elif` の条件が true の場合、`pass` キーワードを削除し、`low` を `1` 変数に加えて `mid` 変数の値を更新してください。

これにより、リスト内の現在の検索エリアの右半分まで検索範囲が拡張されます。なぜなら、`value` が `value_at_middle` より大きい場合、`value` は現在の検索エリアの右半分にあることを意味するからです。

# --hints--

`pass` キーワードを削除し、`low` 変数を `mid + 1` に更新してください。

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid + 1")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        elif value > value_at_middle:
            pass
--fcc-editable-region--
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
```
