---
id: 684169b00a073c9b9f0a703c
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

2回目の関数呼び出しでリストに `3, 5, 4` の値が入っているのに気づきましたか？これは、二分探索が意図した通りに動作していることを示しています。

こうして動作しました：アルゴリズムはまず初期リストの中央として `3` をチェックしました。`4` が `3` より大きいため、探索は右半分に移動しました。次に新しい中央として `5` を調べました。`4` が `5` より小さいので、探索は左に移動し、最終的に `4` を最終範囲の中央として特定しました。

関数をもう一度テストするには、`[1, 3, 5, 9, 14, 22], 10` を使って呼び出し、すぐに出力してください。この場合、値は見つかりません。

# --hints--

関数を `binary_search([1, 3, 5, 9, 14, 22], 10)` で呼び出し、すぐに出力してください。

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return []
    
print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
