---
id: a789b3483989747d63b0e427
title: بناء محدد أكبر رقم
challengeType: 26
dashedName: build-the-largest-number-finder
---

# --description--

في هذا المختبر، ستبني دالة تُرجع مصفوفة تتكون من أكبر رقم في كل مصفوفة فرعية مقدمة.

تذكر، يمكنك التكرار خلال مصفوفة باستخدام حلقة `for` بسيطة، والوصول إلى كل عنصر باستخدام صيغة المصفوفة `arr[i]`.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة `largestOfAll` تأخذ مصفوفة من المصفوفات كمعلمة.
2. يجب أن تُرجع الدالة مصفوفة تحتوي على أكبر رقم من كل مصفوفة فرعية.

# --hints--

يجب أن تنشئ دالة باسم `largestOfAll`.

```js
assert.isFunction(largestOfAll);
```

يجب أن تأخذ `largestOfAll` معلمة واحدة فقط.

```js
assert.lengthOf(largestOfAll, 1);
```

يجب أن تُرجع `largestOfAll([[4, 5, 1, 3], [13, 27, 18, 26], [32, 35, 37, 39], [1000, 1001, 857, 1]])` مصفوفة.

```js
assert.isArray(
  largestOfAll([
    [4, 5, 1, 3],
    [13, 27, 18, 26],
    [32, 35, 37, 39],
    [1000, 1001, 857, 1]
  ])
);
```

يجب أن تُرجع `largestOfAll([[13, 27, 18, 26], [4, 5, 1, 3], [32, 35, 37, 39], [1000, 1001, 857, 1]])` القيمة `[27, 5, 39, 1001]`.

```js
assert.deepEqual(
  largestOfAll([
    [13, 27, 18, 26],
    [4, 5, 1, 3],
    [32, 35, 37, 39],
    [1000, 1001, 857, 1]
  ]),
  [27, 5, 39, 1001]
);
```

يجب أن تُرجع `largestOfAll([[4, 9, 1, 3], [13, 35, 18, 26], [32, 35, 97, 39], [1000000, 1001, 857, 1]])` القيمة `[9, 35, 97, 1000000]`.

```js
assert.deepEqual(
  largestOfAll([
    [4, 9, 1, 3],
    [13, 35, 18, 26],
    [32, 35, 97, 39],
    [1000000, 1001, 857, 1]
  ]),
  [9, 35, 97, 1000000]
);
```

يجب أن تُرجع `largestOfAll([[17, 23, 25, 12], [25, 7, 34, 48], [4, -10, 18, 21], [-72, -3, -17, -10]])` القيمة `[25, 48, 21, -3]`.

```js
assert.deepEqual(
  largestOfAll([
    [17, 23, 25, 12],
    [25, 7, 34, 48],
    [4, -10, 18, 21],
    [-72, -3, -17, -10]
  ]),
  [25, 48, 21, -3]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function largestOfAll(arr) {
  return arr.map(subArr => Math.max.apply(null, subArr));
}

largestOfAll([
  [4, 5, 1, 3],
  [13, 27, 18, 26],
  [32, 35, 37, 39],
  [1000, 1001, 857, 1]
]);
```
