---
id: 699d20623e57e8ffee885451
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

La funzione `parseCard` prenderà una stringa di catalogo grezza e restituirà un oggetto strutturato. Prima di costruire la tua funzione, prepara un modo per vedere il suo output mentre lavori.

Chiama `parseCard` con `rawCatalogCards[2]` come argomento, assegna il risultato a una variabile chiamata `cardResult` e stampa `cardResult` nella console.

Al momento `parseCard` non ha un corpo, quindi restituisce `undefined`. Man mano che aggiungi codice alla funzione nei prossimi passaggi, l'output stampato ti mostrerà esattamente cosa la funzione restituisce in ogni fase.

# --hints--

Dovresti chiamare `parseCard` con `rawCatalogCards[2]` come argomento.

```js
assert.match(__helpers.removeJSComments(code), /parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)/);
```

Dovresti assegnare il risultato della chiamata a `parseCard` con `rawCatalogCards[2]` a una variabile chiamata `cardResult`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+cardResult\s*=\s*parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)\s*;?/);
```

Dovresti stampare `cardResult` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*cardResult\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {}

--fcc-editable-region--

--fcc-editable-region--
```
