---
id: 699d20623e57e8ffee885458
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Il campo year arriva come stringa, ma ti serve come numero per un calcolo successivo. Usa `parseInt()` per convertirlo.

L'anno può anche essere vuoto, quindi ti serve un fallback. L'operatore ternario `condition ? valueIfTrue : valueIfFalse` restituisce uno di due valori in base a una condizione.

Aggiorna la proprietà `year` nell'istruzione return in:

```js
year: year ? parseInt(year) : "Unknown"
```

# --hints--

Quando viene chiamato con una card che ha un anno valido, la proprietà `year` dell'oggetto restituito dovrebbe essere un numero, non una stringa.

```js
assert.isNumber(parseCard("The Shining | King, Stephen | 1977 | Shelf K1").year);
```

Quando viene chiamato con `"The Shining | King, Stephen | 1977 | Shelf K1"`, la proprietà `year` dovrebbe essere `1977`.

```js
assert.strictEqual(parseCard("The Shining | King, Stephen | 1977 | Shelf K1").year, 1977);
```

Quando viene chiamato con una stringa che ha un campo year vuoto, la proprietà `year` dovrebbe essere `"Unknown"`.

```js
assert.strictEqual(parseCard("X | A |  | Y").year, "Unknown");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year,
    location: location || "Unknown"
  };
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
