---
id: 699d1f8e51fcdc8230f779fb
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop costruirai un sistema di catalogo per una biblioteca del patrimonio. Analizzerai dati grezzi di libri, cercherai e raggrupperai gli elementi, verificherai la qualità dei dati ed esporterai i risultati in diversi formati.

Ogni scheda del catalogo è archiviata come una singola stringa nel formato `"Title | Author | Year | Location"`, con i campi separati dal carattere pipe (`|`).

Userai un array di queste stringhe durante tutto il workshop per costruire e testare le tue funzioni di catalogo. Crea un array chiamato `rawCatalogCards` e aggiungi a esso le seguenti due stringhe:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

Dovresti dichiarare una variabile chiamata `rawCatalogCards`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` dovrebbe essere un array.

```js
assert.isArray(rawCatalogCards);
```

Il primo elemento di `rawCatalogCards` dovrebbe essere `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

Il secondo elemento di `rawCatalogCards` dovrebbe essere `"The Shining | King, Stephen | 1977 | Shelf K1"`.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
