---
id: 699d20623e57e8ffee885457
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Alcune schede del catalogo hanno campi mancanti. Per esempio, la voce 16 ha un autore vuoto, quindi `parseCard` attualmente restituisce una stringa vuota per quel campo.

L'operatore `||` fornisce un fallback: `a || b` restituisce `a` quando `a` è truthy, e `b` altrimenti. Poiché una stringa vuota è falsy, `"" || "Unknown"` restituisce `"Unknown"`.

Aggiorna `title`, `author` e `location` nell'istruzione return per usare `|| "Unknown"` come fallback. Lascia `year` per il passaggio successivo.

# --hints--

Quando viene chiamato con una stringa che ha un campo autore vuoto, la proprietà `author` dell'oggetto restituito dovrebbe essere `"Unknown"`.

```js
assert.strictEqual(parseCard("X |  | 2000 | Y").author, "Unknown");
```

Quando viene chiamato con una stringa che ha un campo titolo vuoto, la proprietà `title` dell'oggetto restituito dovrebbe essere `"Unknown"`.

```js
assert.strictEqual(parseCard(" | A | 2000 | Y").title, "Unknown");
```

Quando viene chiamato con una stringa che ha un campo location vuoto, la proprietà `location` dell'oggetto restituito dovrebbe essere `"Unknown"`.

```js
assert.strictEqual(parseCard("X | A | 2000 | ").location, "Unknown");
```

Non dovresti ancora aggiungere un fallback per `year`.

```js
assert.strictEqual(parseCard("X | A |  | Y").year, "");
```

Quando tutti i campi sono presenti, `parseCard` dovrebbe restituire i loro valori reali e non `"Unknown"`.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--
  return {
    title: title,
    author: author,
    year: year,
    location: location
  };
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
