---
id: 69d5e6f7a8b9c0d1e2f3a4b5
title: Passaggio 22
challengeType: 1
dashedName: step-22
---

# --description--

Ora completa `renderEntry`. Dichiara una costante fallback per ogni campo: `const title = entry.title || "Unknown"`, e applica lo stesso pattern per `author`, `year` e `location`.

Restituisci un template literal multilinea che formatta la voce come una scheda biblioteca. Usa `"-".repeat(25)` per il bordo superiore e inferiore, e includi righe etichettate per `Title`, `Author`, `Year` e `Location`. Per esempio, una riga etichettata per il titolo sarebbe `Title: ${title}`.

# --hints--

`renderEntry` dovrebbe restituire una stringa.

```js
assert.isString(renderEntry(catalog[0]));
```

Quando viene chiamato con `catalog[0]`, la stringa restituita dovrebbe contenere `"From a Buick 8"`.

```js
assert.include(renderEntry(catalog[0]), "From a Buick 8");
```

Quando viene chiamato con `catalog[0]`, la stringa restituita dovrebbe contenere `"King, Stephen"`.

```js
assert.include(renderEntry(catalog[0]), "King, Stephen");
```

Quando viene chiamato con `catalog[0]`, la stringa restituita dovrebbe contenere l'anno `2002`.

```js
assert.include(renderEntry(catalog[0]), "2002");
```

Quando viene chiamato con `catalog[0]`, la stringa restituita dovrebbe contenere `"Shelf K7"`.

```js
assert.include(renderEntry(catalog[0]), "Shelf K7");
```

La stringa restituita dovrebbe includere un bordo superiore e inferiore usando `"-".repeat(25)`.

```js
assert.isAbove(renderEntry(catalog[0]).split("-".repeat(25)).length, 2);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

function groupByDecade(catalog) {
  const grouped = {};
  for (let i = 0; i < catalog.length; i++) {
    const book = catalog[i];
    if (book.year === "Unknown") {
      if (!grouped["Unknown"]) {
        grouped["Unknown"] = [];
      }
      grouped["Unknown"].push(book);
      continue;
    }
    const decade = Math.floor(book.year / 10) * 10;
    const decadeKey = `${decade}s`;
    if (!grouped[decadeKey]) {
      grouped[decadeKey] = [];
    }
    grouped[decadeKey].push(book);
  }
  return grouped;
}

const byDecade = groupByDecade(catalog);

function renderEntry(entry) {
--fcc-editable-region--

--fcc-editable-region--
}

console.log(renderEntry(catalog[0]));
```
