---
id: 69b279bae0f20eda63163cef
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

现在你已经将 `array` 列表分成了两个独立的列表，你需要继续将每个列表分割，直到每个元素都单独成为一个列表。只有包含单个数字的列表才是已排序的。

为此，在你的函数内部递归调用 `mergeSort`，并将 `leftPart` 作为调用的参数传入。

# --hints--

你应该在同一个 `mergeSort(leftPart)` 函数内部调用 `mergeSort`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
