---
id: 69b279bae0f20eda63163cf3
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

在循环内部，开始比较每一半中当前的元素，这样你可以先合并较小的值。

添加一个 `if` 块，用来查看 `leftPart` 中的当前元素是否小于 `rightPart` 中的当前元素。
# --hints--

`if` 条件应该查看 `leftPart[leftArrayIndex]` 是否小于 `rightPart[rightArrayIndex]`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
