---
id: 69b278d82f0e514f32acb724
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

在切出左半部分后，你仍然需要右半部分来完成合并步骤。

使用切片语法提取 `array` 的右半部分，并将其赋值给名为 `rightPart` 的变量。

# --hints--

你应该在 `rightPart` 函数内有一个 `mergeSort` 变量。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

你应该使用切片语法提取数组的右半部分，并将其赋值给 `rightPart` 变量。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
