---
id: 69b279bae0f20eda63163cf2
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

接下来，你需要创建一个循环，只要 `leftPart` 和 `rightPart` 中还有元素，就继续执行。

为此，创建一个 `while` 循环，包含两个条件：一个检查 `leftArrayIndex` 是否小于 `leftPart` 的长度，另一个检查 `rightArrayIndex` 是否小于 `rightPart` 的长度。

# --hints--

你应该创建一个 `while` 循环，检查 `leftArrayIndex` 是否小于 `leftPart` 的长度，以及 `rightArrayIndex` 是否小于 `rightPart` 的长度。

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
