---
id: 69b2713ff6c21fed2d3ad88b
title: 第一步
challengeType: 1
dashedName: step-1
---

# --description--

归并排序是一种使用分治原则对数据集合进行排序的排序算法。也就是说，它将一个集合“分割”成更小的子部分，然后通过独立排序这些子部分来“征服”它们，最后合并已排序的子部分。

首先创建一个名为 `mergeSort` 的函数，该函数接受参数 `array`。

# --hints--

你应该声明一个名为 `mergeSort` 的函数，该函数带有一个 `array` 参数。

```js
assert.isFunction(mergeSort);
const params = __helpers.getFunctionParams(mergeSort.toString());
assert.lengthOf(params, 1);
assert.deepEqual(params[0].name, "array");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
