---
id: 69b278519ef46b2fd8bf1546
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

你已经有了一个 `middlePoint` 值。使用它将数组分成两半，并准备对左半部分进行排序。

使用切片语法提取 `array` 的左半部分，并将其赋值给名为 `leftPart` 的变量。

# --hints--

你应该在 `leftPart` 函数内声明一个 `mergeSort` 变量。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

你应该将 `array.slice(0, middlePoint)` 赋值给 `leftPart`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
