---
id: 69b279bae0f20eda63163cfd
title: 第 19 步
challengeType: 1
dashedName: step-19
---

# --description--

最后，你应该将右半部分的任何剩余值拷贝到 `array` 中。

为此，添加一个 `while` 循环，只要 `rightArrayIndex` 小于 `rightPart.length` 就运行，将 `rightPart[rightArrayIndex]` 赋值给 `array[sortedIndex]`，并同时递增 `rightArrayIndex` 和 `sortedIndex`。

# --hints--

你应该创建一个条件为 `while` 的 `rightArrayIndex < rightPart.length` 循环。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

循环体应将 `rightPart[rightArrayIndex]` 赋值给 `array[sortedIndex]`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?/
);
```

赋值后应同时递增 `rightArrayIndex` 和 `sortedIndex`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*rightArrayIndex\s*(\+\+|\+=\s*1)\s*;?[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
