---
id: 69b279bae0f20eda63163cf9
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

你创建的 `while` 循环比较 `leftPart` 中的一个元素和 `rightPart` 中的另一个元素，然后将较小的元素添加到主 `array` 列表中。

它会继续执行这个操作，直到没有元素可以比较。但 `leftPart` 可能仍有元素剩余，而 `rightPart` 没有，反之亦然。

创建另一个 `while` 循环，当 `leftArrayIndex` 小于 `leftPart` 的长度时运行。在接下来的步骤中，你将用它将 `leftPart` 中剩余的元素拷贝到 `array` 列表中。

# --hints--

你应该创建一个条件为 `while` 的 `leftArrayIndex < leftPart.length` 循环。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
