---
id: 69b279bae0f20eda63163cf0
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

在你的函数体底部，再次调用 `mergeSort()` 函数。这一次，将 `rightPart` 作为参数传入函数调用。

# --hints--

你应该在函数体底部调用 `mergeSort()`，并传入参数 `rightPart`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
