---
id: 69b279bae0f20eda63163cf5
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

在将 `leftPart` 中的元素赋值给排序后的数组后，将 `leftArrayIndex` 增加 `1`。

# --hints--

在将元素插入排序后的数组后，你应该将 `leftArrayIndex` 的当前值增加。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftArrayIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
