---
id: 69b279bae0f20eda63163cf6
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

在之前的步骤中，你将 `leftPart` 中的元素赋值给了 `array` 列表，因为它更小。但情况并非总是如此。在某些比较情况下，右边的元素可能更小。

创建一个 `else` 分支，当 `leftPart` 中的元素不小于 `rightPart` 中的元素时执行。

在 `else` 块内，将 `rightPart[rightArrayIndex]` 赋值给 `array[sortedIndex]`。

# --hints--

你应该创建一个包含 `else` 的 `array[sortedIndex] = rightPart[rightArrayIndex]` 块。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
