---
id: 69b2713ff6c21fed2d3ad88b
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

归并排序是一种使用分治原理对数据集合进行排序的排序算法。也就是说，它将一个集合“分割”成更小的子部分，并通过独立排序这些子部分来“征服”它们，然后合并已排序的子部分。

首先创建一个名为 `mergeSort` 的函数，该函数接受参数 `array`。

# --hints--

你应该声明一个名为 `mergeSort` 的函数，该函数带有一个 `array` 参数。

```js
assert.isFunction(mergeSort);

const explorer = await __helpers.Explorer(code);
const mergeSortFn = explorer.allFunctions.mergeSort;
assert.exists(mergeSortFn, 'You should declare a mergeSort function');

const params = mergeSortFn.parameters;
assert.lengthOf(
  params,
  1,
  'mergeSort should take a single parameter'
);
assert.isTrue(
  params[0].matches('array'),
  'The parameter of mergeSort should be named array'
);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
