---
id: 69b279bae0f20eda63163cf1
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

现在是时候将列表 `leftPart` 和 `rightPart` 排序并合并到原始的 `array` 中了。

你可以通过比较两个列表中的元素，将较小的元素合并到主列表中。你将对 `leftPart` 和 `rightPart` 中的所有索引进行这种比较。

创建三个变量：`leftArrayIndex`、`rightArrayIndex` 和 `sortedIndex`，并将它们的值设置为 `0`。这些变量将帮助你在排序过程中跟踪每个索引。

# --hints--

你应该声明一个 `leftArrayIndex` 变量，并初始化为 `0`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

你应该声明一个 `rightArrayIndex` 变量，并初始化为 `0`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

你应该声明一个 `sortedIndex` 变量，并初始化为 `0`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);

--fcc-editable-region--
  
--fcc-editable-region--
}
```
