---
id: 69b2762b17d23fbee40a6c00
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

归并排序算法执行三个操作：

- 将一个未排序的序列分成若干子部分
- 对子部分中的项进行排序
- 合并已排序的子部分

上述过程递归进行，直到子部分合并成完整的已排序序列。我们先从分割序列开始。

首先，在 `mergeSort` 函数内部，使用 `middlePoint` 声明一个变量 `const`，并赋值为 `Math.floor(array.length / 2)`，这样你就可以在递归之前将序列分成两半。

# --hints--

你应该声明一个 `middlePoint` 变量，并将 `Math.floor(array.length / 2)` 赋值给它。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
