---
id: 69b279bae0f20eda63163cf8
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

你在前几步创建的 `if` 和 `else` 语句会将元素赋值给排序后的数组。

每个赋值给排序数组的元素都会占据列表中的一个索引。因此，每次赋值后你必须移动到排序数组中的下一个索引。

在 `if`/`else` 块的下面，但仍在 `while` 循环内，将 `sortedIndex` 增加 `1`。

# --hints--

你应该在 `sortedIndex` 子句之后，将 `1` 的当前值增加 `else`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
