---
id: 69b279bae0f20eda63163cf4
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

当 `if` 条件计算结果为 `true` 时，表示 `leftPart` 列表中的元素小于它与 `rightPart` 列表中元素比较的那个元素。

在这种情况下，你可以将 `leftArrayIndex` 处找到的元素放入排序数组中的 `leftPart` 位置。

在 `if` 块内，将 `leftPart[leftArrayIndex]` 赋值给 `array[sortedIndex]`。

# --hints--

你应该在 `leftPart[leftArrayIndex]` 内将 `array[sortedIndex]` 赋值给 `if`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
