---
id: 68cabc534d2e741f33a3fe1d
title: ゲームキャラクターのステータストラッカーを作成する
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

このラボでは、ゲームキャラクターのステータストラッカーを作成します。このプログラムでは、特定の属性を持つキャラクターを作成し、その属性を更新し、キャラクターの現在のステータスを取得できます。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. ゲームキャラクターを表し、キャラクターのステータスを管理する`GameCharacter`という名前のクラスを作成してください。
2. インスタンス化されたとき、新しい`GameCharacter`オブジェクトは以下の属性を持つ必要があります:
    - `_name`はインスタンス化時に与えられた文字列に設定されていること。
    - `_health`は`100`に設定されていること。
    - `_mana`は`50`に設定されていること。
    - `_level`は`1`に設定されていること。
3. キャラクターの名前への読み取り専用アクセスのために`name`プロパティを作成してください。
4. `health`プロパティについて:
     - 現在のヘルスを返すゲッターを定義してください。
     - ヘルスが`0`未満に設定されるのを防ぎ、`100`で上限を設定するセッターを定義してください。
5. `mana`プロパティについて:
     - 現在のマナを返すゲッターを定義してください。
     - マナが`0`未満に設定されるのを防ぎ、`50`で上限を設定するセッターを定義してください。
6. キャラクターの現在のレベルを返す`level`のゲッターを作成してください。
7. `level_up`という名前のメソッドを定義してください。このメソッドは:
    - キャラクターのレベルを1上げます。
    - ヘルスを`100`に、マナを`50`に対応するプロパティのセッターを使ってリセットします。
    - `<name> leveled up to <level>!`の形式でメッセージを出力します（`<name>`と`<level>`はキャラクターの名前と新しいレベルに置き換えられます）。
8. 以下を含むフォーマットされた文字列を返す`__str__`メソッドを定義してください:
    - キャラクターの名前
    - キャラクターのレベル
    - キャラクターの現在のヘルス
    - キャラクターの現在のマナ
   例えば、インスタンス化直後の`Kratos`という名前のキャラクターは以下のように表されます:

   ```md
   Name: Kratos
   Level: 1
   Health: 100
   Mana: 50
   ```

## 使用例

```python
hero = GameCharacter('Kratos') # Creates a new character named Kratos
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

`GameCharacter`クラスを作成してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

`__init__`クラスに`GameCharacter`メソッドを作成してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

`__init__`メソッドは2つのパラメータを持ち、最初のパラメータは`self`である必要があります。

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.__init__).parameters)
    assert params[0] == "self"
    assert len(params) == 2
    `) 
})
```

インスタンス化時に与えられた文字列を`self._name`メソッド内で`__init__`に設定してください。

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

`100`メソッド内で`self._health`に`__init__`を代入してください。

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

`50`メソッド内で`self._mana`に`__init__`を代入してください。

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

`1`メソッド内で`self._level`に`__init__`を代入してください。

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

`name`クラスに`GameCharacter`メソッドを作成してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

`name`メソッドはパラメータを1つだけ持ち、`self`である必要があります。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

`self._name`メソッドで`name`を返してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

`name`メソッドには`@property`デコレーターを付けてください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

`health`クラスに`GameCharacter`メソッドを作成してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

`health`メソッドはパラメータを1つだけ持ち、`self`である必要があります。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

`self._health`メソッドで`health`を返してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

`health`メソッドには`@property`デコレーターを付けてください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

`health`のセッターを作成してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

`health`のセッターには`@health.setter`デコレーターを付けてください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

`health`のセッターは2つのパラメータを持ち、最初のパラメータは`self`である必要があります。

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.health.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

`self._health`のセッターに与えられた値が`0`未満の場合、`health`を`0`に設定してください。

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = -1
    assert char.health == 0
    char.health = -10
    assert char.health == 0
`)})
```

`self._health`のセッターに与えられた値が`100`より大きい場合、`health`を`100`に設定してください。

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 101
    assert char.health == 100
    char.health = 200
    assert char.health == 100
`)})
```

`self._health`のセッターに与えられた値が`health`と`0`の間の場合、`100`をその値に設定してください。

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 95
    assert char.health == 95
    char.health = 1
    assert char.health == 1
    char.health = 0
    assert char.health == 0
    char.health = 100
    assert char.health == 100
`)})
```

`mana`クラスに`GameCharacter`メソッドを作成してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

`mana`メソッドはパラメータを1つだけ持ち、`self`である必要があります。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

`self._mana`メソッドで`mana`を返してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

`mana`メソッドには`@property`デコレーターを付けてください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

`mana`のセッターを作成してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1]`) 
})
```

`mana`のセッターには`@mana.setter`デコレーターを付けてください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1].has_decorators("mana.setter")`) 
})
```

`mana`のセッターは2つのパラメータを持ち、最初のパラメータは`self`である必要があります。

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.mana.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

`self._mana`のセッターに与えられた値が`0`未満の場合、`mana`を`0`に設定してください。

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = -1
    assert char.mana == 0
    char.mana = -10
    assert char.mana == 0
`)})
```

`self._mana`のセッターに与えられた値が`50`より大きい場合、`mana`を`50`に設定してください。

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 51
    assert char.mana == 50
    char.mana = 100
    assert char.mana == 50
`)})
```

`self._mana`のセッターに与えられた値が`mana`と`0`の間の場合、`50`をその値に設定してください。

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 25
    assert char.mana == 25
    char.mana = 0
    assert char.mana == 0
    char.mana = 50
    assert char.mana == 50
`)})
```

`level`クラスに`GameCharacter`メソッドを作成してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level")`) 
})
```

`level`メソッドはパラメータを1つだけ持ち、`self`である必要があります。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_args("self")`) 
})
```

`self._level`メソッドで`level`を返してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_return("self._level")`)
})
```

`level`メソッドには`@property`デコレーターを付けてください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_decorators("property")`) 
})
```

`level_up`クラスに`GameCharacter`メソッドを作成してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level_up")`) 
})
```

`level_up`メソッドはパラメータを1つだけ持ち、`self`である必要があります。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_args("self")`) 
})
```

`self._level`メソッドで`1`を`level_up`だけ増やしてください。

```js
({ 
  test: () => runPython(`
    node = _Node(_code).find_class("GameCharacter").find_function("level_up")
    assert (
        node.has_stmt("self._level += 1") or
        node.has_stmt("self._level = self._level + 1")
    )
`)
})
```

`self.health`メソッドで`100`を`level_up`に設定してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.health = 100")`) 
})
```

`self.mana`メソッドで`50`を`level_up`に設定してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.mana = 50")`) 
})
```

`level_up`メソッドは`<name> leveled up to <level>!`を出力してください（`<name>`と`<level>`はキャラクターの名前と新しいレベルに置き換えられます）。

```js
({
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.level_up()
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
            
    assert "test leveled up to 2!" in output
`)
})
```

`__str__`クラスに`GameCharacter`メソッドを作成してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__str__")`) 
})
```

`__str__`メソッドはパラメータを1つだけ持ち、`self`である必要があります。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__str__").has_args("self")`) 
})
```

`__str__`メソッドは指定されたフォーマットでキャラクターのステータスを含む文字列を返してください。

```js
({ 
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.mana = 15
    gc.health = 77
    print(gc)
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
    expected = """Name: test
    Level: 1
    Health: 77
    Mana: 15"""
    assert expected in output
`)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self._name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
