---
id: 68e2c945cc1d8e778152be31
title: プレイヤーインターフェースを作成する
challengeType: 27
dashedName: build-a-player-interface
---

# --description--

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `Player` クラスを継承した抽象クラス `abc.ABC` を定義してください。
1. `Player` クラスは、以下を設定する `__init__` メソッドを持つ必要があります：

   - `moves` 属性を空のリストに設定する。
   - `position` 属性を `(0, 0)` に設定する。
   - `path` 属性を初期位置を含むリストに設定する。

1. `Player` クラスは `make_move` という名前のメソッドを持ち、以下を行います：

   - `random.choice` を使って、具体的なクラスで定義された `moves` 属性からランダムな移動を取得する。
   - 選択された移動の値を現在の位置に加算し、`position` 属性を更新する。
   - 新しい `position` タプルを `path` 属性に追加する。
   - 新しい `position` を返す。

1. `Player` クラスは具体的なクラスで実装される抽象メソッド `level_up` を持つ必要があります。
1. `Pawn` クラスを継承した `Player` クラスを定義してください。
1. `Pawn` クラスは親の `super()` メソッドを `__init__` を使って呼び出し、その後 `moves` 属性を `x, y` 座標を表すタプルのリストに設定します。
1. 各座標タプルは、上、下、左、右の方向に `1` 単位移動を表す必要があります。
1. `Pawn` クラスは具体的な `level_up` メソッドを実装し、`moves` 属性にさらに移動を追加します。追加される移動は4つの斜め方向の移動を表します（例えば、`1` 単位下＋`1` 単位左）。

**注意:** 標準ライブラリのモジュールはエイリアスを使わずにインポートしてください。`Player` クラスがインスタンス化可能になるまでは、`Pawn` クラスに関連するテストは失敗します。

# --hints--

`Player` という名前のクラスが必要です。

```js
({ test: () => runPython(`assert _Node(_code).has_class("Player")`) })
```

`Player` クラスは `ABC` モジュールの `abc` クラスを継承する必要があります。

```js
({ test: () => runPython(`
from abc import ABC
import ast
cls = _Node(_code).find_class("Player")
assert cls.inherits_from("ABC") or cls.inherits_from("abc.ABC")
`) })
```

`Player` クラスは `__init__` メソッドを持つ必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("__init__")`) })
```

`Player` の `__init__` メソッドは単一のパラメータ `self` を持つ必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("__init__").has_args("self")`) })
```

`Player` の `__init__` メソッドは `self.moves` を空のリストに設定する必要があります。

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.moves = []") or init.has_stmt("self.moves = list()")
`) })
```

`Player` の `__init__` メソッドは `self.position` を `(0, 0)` に設定する必要があります。

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.position = (0, 0)")
`) })
```

`Player` の `__init__` メソッドは `self.path` を初期位置を含むリストに設定する必要があります。

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.path = [self.position]")
`) })
```

`Player` クラスは `make_move` メソッドを持つ必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("make_move")`) })
```

`Player` の `make_move` メソッドは単一のパラメータ `self` を持つ必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("make_move").has_args("self")`) })
```

`Player` の `make_move` メソッドは `random.choice` を使って `moves` 属性からランダムな移動を取得する必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").block_has_call("choice", "make_move")`) })
```

`Player` の `make_move` メソッドはランダムに選択された移動の座標を加算して `position` 属性を更新する必要があります。

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.position == (1, 2)
  p.make_move()
  assert p.position == (2, 4)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

`Player` の `make_move` メソッドは新しい `position` タプルを `path` 属性に追加する必要があります。

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.path == [(0, 0), (1, 2)]
  p.make_move()
  assert p.path == [(0, 0), (1, 2), (2, 4)]

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

`Player` の `make_move` メソッドは更新された `position` 属性を返す必要があります。

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  assert p.make_move() == (1, 2)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

`Player` クラスは `level_up` メソッドを持つ必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("level_up")`) })
```

`Player` の `level_up` メソッドは単一のパラメータ `self` を持つ必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("level_up").has_args("self")`) })
```

`Player` の `level_up` メソッドは抽象メソッドである必要があります。

```js
({ test: () => runPython(`
target = _Node(_code).find_class("Player").find_function("level_up")
assert target.has_decorators("abstractmethod") or target.has_decorators("abc.abstractmethod")
`) })
```

`Player` クラスは抽象クラスである必要があります。

```js
({ test: () => runPython(`
try:
  Player()
except TypeError as e:
  assert str(e) == "Can't instantiate abstract class Player with abstract method level_up"
else:
  assert False, "Player class should not be instantiable"
`) })
```

`Pawn` という名前のクラスが必要です。

```js
({ test: () => runPython(`assert _Node(_code).has_class("Pawn")`) })
```

`Pawn` クラスは `Player` クラスを継承する必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").inherits_from("Player")`) })
```

`Pawn` クラスは `__init__` メソッドを持つ必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("__init__")`) })
```

`Pawn` の `__init__` メソッドは単一のパラメータ `self` を持つ必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_args("self")`) })
```

`Pawn` の `__init__` メソッドは `__init__` 関数を使って親の `super` を呼び出す必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_stmt("super().__init__()")`) })
```

`Pawn` の `__init__` メソッドは `moves` 属性を `x, y` 座標を表すタプルのリストに設定する必要があります。各座標タプルは上、下、左、右の方向に `1` 単位移動を表します。

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]
assert isinstance(p.moves, list)
assert len(p.moves) == 4
assert all(move in p.moves for move in moves)
`) })
```

`Pawn` クラスは `level_up` メソッドを持つ必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("level_up")`) })
```

`Pawn` の `level_up` メソッドは `1` 単位の4つの斜め移動を `moves` 属性に追加する必要があります。

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1), (1, 1), (1, -1), (-1, 1), (-1, -1)]
p.level_up()
assert isinstance(p.moves, list)
assert len(p.moves) == 8
assert all(move in p.moves for move in moves)
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
from abc import ABC, abstractmethod
from random import choice


class Player(ABC):
    def __init__(self):
        self.moves = []
        self.position = (0, 0)
        self.path = [self.position]

    def make_move(self):
        x, y = self.position
        move = choice(self.moves)
        new_x = move[0] + x
        new_y = move[1] + y
        self.position = (new_x, new_y)
        self.path.append(self.position)
        return self.position

    @abstractmethod
    def level_up(self):
        pass


class Pawn(Player):
    def __init__(self):
        super().__init__()
        self.moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]

    def level_up(self):
        self.moves.extend([(1, 1), (1, -1), (-1, 1), (-1, -1)])
```
