---
id: 6852f01116912861c7d1b0e6
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Dentro del método `__init__`, necesitas asignar los valores de los parámetros a los atributos del objeto para que cada email pueda almacenar su información.

Dentro de tu método `__init__`, elimina la palabra clave `pass` y asigna los parámetros a `self.sender`, `self.receiver`, `self.subject` y `self.body`.

# --hints--

Debes asignar el parámetro `sender` a `self.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

Debes asignar el parámetro `receiver` a `self.receiver`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

Debes asignar el parámetro `subject` a `self.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

Debes asignar el parámetro `body` a `self.body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
