---
id: 6852ea0c3cadb54e897acdac
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, vas a construir un *Simulador de Email* que simula enviar, recibir y gestionar correos electr칩nicos entre diferentes usuarios. Aprender치s sobre clases, objetos y c칩mo organizar el c칩digo de manera orientada a objetos.

Comienza creando una clase llamada `Email` usando la palabra clave `class`.

# --hints--

Debes crear una clase llamada `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
