---
id: 6852f0c3fcedb964ab959921
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Ahora vamos a probar que el atributo `read` se agregó correctamente a tu clase `Email`. Como ya tienes un objeto email de los pasos anteriores, imprime el atributo `read` para ver que ahora es `False` por defecto.

# --hints--

Debes imprimir `email_obj.read` para probar el atributo read.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
