---
id: 6852ef95ccc27660275692cd
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Tu clase `Email` necesita almacenar información sobre cada correo electrónico cuando se crea.

Dentro de tu clase `Email`, elimina la palabra clave `pass` existente y crea el método `__init__` que tome `self`, `sender`, `receiver`, `subject` y `body` como parámetros.

# --hints--

Debes definir un método `__init__` dentro de la clase `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

Tu método `__init__` debe tener los parámetros: `self`, `sender`, `receiver`, `subject` y `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
