---
id: 6852f11d412c29665985a915
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Ahora que los usuarios pueden enviar correos electrónicos, necesitan una forma de almacenar los correos que reciben. Cada usuario debe tener una bandeja de entrada para recopilar sus correos.

Agrega un atributo `inbox` al método `__init__` de la clase `User` e inicialízalo como una lista vacía `[]`. Esto almacenará todos los correos electrónicos que el usuario reciba.

# --hints--

Debes agregar `self.inbox = []` al método `__init__` de la clase `User`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
