---
id: 6852f06fdde1b4634a263f03
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora vamos a probar la clase `Email` creando un objeto de correo electrónico y verificando que almacena la información correctamente. Imprimirás un par de atributos para comprobar que todo funciona.

Crea un objeto de correo electrónico llamado `email_obj` con `alice@example.com` como remitente, `bob@example.com` como destinatario, `Hello` como asunto y `Hi Bob!` como cuerpo. Luego imprime los atributos `sender` y `subject` en declaraciones de impresión separadas para verificar que estén almacenados correctamente.

# --hints--

Debes crear un objeto de correo electrónico llamado `email_obj` con los parámetros especificados.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

Debes imprimir `email_obj.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

Debes imprimir `email_obj.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--

--fcc-editable-region--
```
