---
id: 6853f2da9b8ed891e9fdc5fc
title: Pasos 23
challengeType: 20
dashedName: step-23
---

# --description--

Ahora vamos a mostrar la información del remitente y del receptor. Para mostrar quién envió y quién recibió el correo electrónico, agrega dos sentencias print al método `display_full_email` en este formato:

- `From: sender` donde `sender` es reemplazado con el nombre del remitente
- `To: receiver` donde `receiver` se reemplaza con el nombre del receptor

# --hints--

Debes usar f-strings para imprimir el nombre del remitente usando `From: {self.sender.name}`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender.name}')")`));
    }
})
```

Debes usar f-strings para imprimir el nombre del receptor usando `To: {self.receiver.name}`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver.name}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
