---
id: 6853f4ded9f84896ff182dd8
title: Paso 31
challengeType: 20
dashedName: step-31
---

# --description--

Después de verificar que la bandeja de entrada esté vacía, imprime el mensaje `\nYour Emails:`. Después de eso, itera sobre todos los correos en la bandeja de entrada usando un `for` loop con enumeración. Esta es la sintaxis para la enumeración con un número inicial: `enumerate(iterable, start=start_number)`.

Usa enumeración para numerar los correos electrónicos comenzando desde `1`. Usa `i` (el índice) y `email` (el objeto de correo electrónico) como las variables de iteración.

Dentro del ciclo, imprime una f-string con el índice de la iteración seguido de un `.`, luego un espacio y la representación en cadena del objeto email en este formato:

```md
i. string_representation
```

Donde `i` es el índice y `string_representation` es la representación del objeto email.

# --hints--

Debes imprimir el encabezado `\nYour Emails:` fuera de la sentencia `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

Debes tener un ciclo `for` que itere sobre `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

Tu ciclo `for` debe usar `i` y `email` para iterar sobre `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

Dentro del ciclo, debes imprimir cada resumen de correo electrónico con un prefijo numerado seguido de la cadena formateada del método `__str__`. No olvides agregar un `.` después del número.

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        
--fcc-editable-region--
```
