---
id: 6853f4584f508e95cc7a9f4d
title: Paso 30
challengeType: 20
dashedName: step-30
---

# --description--

Ahora crearás un método para listar todos los correos en la bandeja de entrada. Este método debe manejar el caso cuando la bandeja de entrada está vacía y mostrar una lista numerada de correos cuando haya correos presentes.

Agrega un método llamado `list_emails` a tu clase `Inbox` que tome `self` como parámetro. Primero, crea una estructura `if` para verificar si la bandeja de entrada está vacía probando la lista `emails`. Si está vacía, imprime el mensaje `Your inbox is empty.\n`. Agrega una estructura `return` para salir del método.

# --hints--

Debes definir un método llamado `list_emails` en la clase `Inbox` que tome `self` como parámetro.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

Debes verificar si `self.emails` está vacío usando `if not self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

Debes imprimir el mensaje `Your inbox is empty.\n` si no hay correos electrónicos.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

Debes tener una sentencia `return` al final de la sentencia `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
