---
id: 6859400f8f7f872e50e1238c
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

La clase `User` necesita un método para enviar correos electrónicos a otros usuarios. Al enviar un correo, crearás un nuevo objeto `Email` y lo agregarás a la bandeja de entrada del receptor.

Crea un método llamado `send_email` en tu clase `User`. Este método debe recibir parámetros para `receiver`, `subject` y `body`. 

Dentro del método, crea un nuevo objeto `Email` con los siguientes valores de parámetros y asígnalo a una variable llamada `email`:

- `sender`: `self` (el usuario actual)
- `receiver`: `receiver` (el usuario que recibe el correo electrónico)
- `subject`: `subject` (el asunto del correo electrónico)
- `body`: `body` (el cuerpo del correo electrónico)

# --hints--

Debes definir un método `send_email` en la clase `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

El método `send_email` debe tener parámetros: `self`, `receiver`, `subject` y `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

El método debe crear un objeto `Email` llamado `email` con los parámetros correctos.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
