---
id: 68593ea22eac7e228ed585ce
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Los correos electrónicos deben rastrear si han sido leídos o no. Agrega un atributo `read` al método `__init__` y configúralo en `False` por defecto, ya que los correos nuevos comienzan como no leídos.

# --hints--

Debes asignar `False` a `self.read` en el método `__init__`.

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
