---
id: 685c1eba966c6396850736ea
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Aunque los usuarios pueden enviar correos electrónicos y tener bandejas de entrada, necesitamos una clase dedicada para gestionar las operaciones de la bandeja de entrada de manera eficiente. 

La clase `Inbox` almacenará una lista de correos electrónicos y proporcionará métodos para agregar nuevos correos, listar todos los correos y gestionarlos.

Crea una nueva clase llamada `Inbox` con un método `__init__` que inicialice una lista vacía llamada `emails`.

# --hints--

Debes crear una clase `Inbox`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

La clase `Inbox` debe tener un método `__init__` que tome `self` como parámetro.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

Debes inicializar `self.emails = []` en el método `__init__` de la clase `Inbox`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
