---
id: 685c1e32254e5993a5c8823e
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Ahora que tienes una configuración inicial de la clase `Email`, es momento de crear usuarios que puedan enviar y recibir correos electrónicos. Cada usuario tendrá un nombre y podrá realizar operaciones de correo electrónico.

Crea una clase `User` con un método `__init__` que tome `self` y `name` como parámetros. Dentro de la clase, asigna el parámetro `name` a `self.name`.

# --hints--

Debes crear una clase `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

La clase `User` debe tener un método `__init__` que reciba `self` y `name` como parámetros.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

Debes asignar el parámetro `name` a `self.name`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
