---
id: 685cdfd1932d8ac6ad986813
title: Paso 28
challengeType: 20
dashedName: step-28
---

# --description--

Dentro del método `__str__`, mostrarás si un correo electrónico ha sido leído o no. 

Las expresiones condicionales te permiten asignar un valor basado en una condición en una sola línea.

Aquí hay un ejemplo:

```py
x = 10
y = 'Even' if x % 2 == 0 else 'Odd' # y will be even
```

Dentro del método, antes, usa una expresión condicional para asignar la cadena `Read` a una variable `status` si el correo electrónico está leído y `Unread` si no lo está. 

# --hints--

El método `__str__` debe crear una variable `status` que use una expresión condicional.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_stmt("status = 'Read' if self.read else 'Unread'")`));
    }
})
```

El método `__str__` debe devolver la información del correo electrónico con el estado incluido.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_variable("status")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    def __str__(self):
        pass
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
