---
id: 685d2fd1f309b9191f257c55
title: Paso 40
challengeType: 20
dashedName: step-40
---

# --description--

Antes de integrar marcas de tiempo en nuestro sistema de correo electrónico, practiquemos cómo trabajar con el formato de datetime. La función `datetime.datetime.now()` nos da la fecha y hora actuales, y podemos usar el método `strftime()` para formatearlo de diferentes maneras.

Aquí te explico cómo funciona `strftime()` con los códigos de formato:

```python
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))  # Output: 2024-03-15 (year-month-day with - separator)
```

Los códigos de formato como `%Y` (año), `%m` (mes), `%d` (día) indican a `strftime()` qué incluir, y puedes agregar separadores como `-` entre ellos.

Al final de tu código, crea una variable llamada `current_time` y asígnale `datetime.datetime.now()`. Luego usa `strftime()` para imprimir la hora en formato horas:minutos:segundos usando `:` como separador. 

Usa estos códigos de formato: `%H` para horas (formato de 24 horas), `%M` para minutos y `%S` para segundos.

# --hints--

Debes crear una variable llamada `current_time` y asignarle `datetime.datetime.now()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_stmt("current_time = datetime.datetime.now()")`));
    }
})
```

Debes usar `strftime()` para imprimir la hora en formato `%H:%M:%S`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(current_time.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

--fcc-editable-region--

--fcc-editable-region--
```
