---
id: 688754d6d59578df1313cb3c
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

¡Ahora es momento de conectar las clases User e Inbox para que los correos realmente puedan ser entregados! 

Los usuarios necesitan tener objetos Inbox adecuados en lugar de listas simples, y el método `send_email` debe entregar los correos electrónicos en la bandeja de entrada del receptor.

Actualiza el `User` class `__init__` de manera que cada usuario obtenga un objeto Inbox real en lugar de una lista vacía creada anteriormente. 

# --hints--

Debes actualizar la clase User para inicializar `self.inbox = Inbox()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
