---
id: 68875554c6c32de788afc64e
title: Paso 21
challengeType: 20
dashedName: step-21
---

# --description--

Los usuarios necesitan una manera de leer correctamente sus correos electrónicos.

Agrega un método `display_full_email` a la clase `Email` que solo tome `self` como parámetro. Dentro de este método, llama al método `mark_as_read()` para marcar el correo como leído cuando alguien lo vea.

Este método se usaría para mostrar el contenido completo del email de una manera bien formateada.

# --hints--

Debes definir un método `display_full_email` en la clase `Email`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

El método `display_full_email` debería tomar solo `self` como parámetro.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

El método debe llamar a `self.mark_as_read()` para marcar el correo como leído.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
