---
id: 6887550fe36f64e27074475a
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

Haz que Alice envíe un correo electrónico a Bob para ver si la entrega del correo funciona correctamente.

Usa el método `send_email` de Alice para enviarle un correo a Bob con el asunto `"Hello"` y el cuerpo `"Hi Bob, how are you?"`.

# --hints--

Debes llamar a `alice.send_email()` con los parámetros correctos.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call('alice.send_email(bob, "Hello", "Hi Bob, how are you?")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
alice = User("Alice")
bob = User("Bob")

--fcc-editable-region--
```
