---
id: 688754f3494b17e0b352d67d
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

¡Ahora es momento de probar nuestro sistema completo de correo electrónico! Vamos a crear algunos usuarios y ver la funcionalidad del correo en acción.

Crea dos objetos User: `alice` con nombre `"Alice"` y `bob` con nombre `"Bob"`. Esto demostrará cómo se pueden crear usuarios e interactuar entre ellos.

# --hints--

Debes crear un objeto User llamado `alice` con el nombre `"Alice"`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("alice").is_equivalent('alice = User("Alice")')`));
    }
})
```

Debes crear un objeto User llamado `bob` con el nombre `"Bob"`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("bob").is_equivalent('bob = User("Bob")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--
```
