---
id: 68b9366fae8b13c1278761b1
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Tu bandeja de entrada necesita una forma de recibir nuevos correos electrónicos. Cuando alguien envía un correo a un usuario, debe añadirse a su bandeja de entrada.

Agrega un método llamado `receive_email` a tu clase `Inbox` que reciba `self` y un objeto de correo `email` como parámetros. Dentro del cuerpo del método, agrega el `email` a la lista `emails` usando el método `append`.

# --hints--

Debes definir un método llamado `receive_email` en la clase `Inbox` que reciba `self` y `email` como parámetros.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

Debes agregar el `email` a `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
