---
id: 68b930ee0e7b8fb81895e59a
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Ahora vas a agregar un método para marcar un correo electrónico como leído. Cuando alguien abre un correo, quieres cambiar su estado de no leído a leído. Este método será simple: solo necesita establecer el atributo `read` en `True`.

Agrega un método llamado `mark_as_read` a tu clase `Email`. Este método debe tomar solo `self` como parámetro y establecer el atributo `read` en `True`.

# --hints--

Debes definir un método llamado `mark_as_read` en la clase `Email` que solo tome `self` como parámetro.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

El método debe establecer `self.read` en `True`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
