---
id: 68b9350183ff05bd106c79e4
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Ahora, prueba el método `mark_as_read` que creaste en la clase `Email`. Usa el método en el objeto `Email` existente `email_obj` para cambiar el estado. Después de marcar el correo como leído, imprime el atributo `read` de `email_obj` para confirmar el cambio.

# --hints--

Debes llamar al método `mark_as_read` en el `email_obj` para marcar el correo como leído.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

Debes imprimir el atributo `read` del `email_obj` después de llamar al método `mark_as_read`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)

--fcc-editable-region--
```
