---
id: 587d7b84367417b2b2512b36
title: 'Gundua Mabano ya Kawaida, Mabano, Mabano ya Mviringo na Nukuu Zisizofunguliwa'
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--

Kosa lingine la sintaksia la kuzingatia ni kwamba mabano yote ya kawaida, mabano, mabano ya mviringo, na nukuu lazima yawe na jozi ya kufunga. Kusahau kipande hufanyika zaidi wakati unarekebisha msimbo uliopo na kuingiza vitu vyenye aina moja ya jozi. Pia, kuwa makini unapopanga ndani mabano ya msimbo mengine, kama vile kuongeza kitendakazi cha mwitiko kama hoja kwa njia.

Njia moja ya kuepuka kosa hili ni mara tu unapoandika tabia ya kufungua, mara moja jumuisha jozi ya kufunga, kisha weka kidokezo cha kuandika katikati yao na endelea kuandika msimbo. Kwa bahati nzuri, wahariri wengi wa msimbo wa kisasa hutengeneza sehemu ya pili ya jozi moja kwa moja.

# --instructions--

Rekebisha makosa mawili ya jozi katika msimbo.

# --hints--

Msimbo wako unapaswa kurekebisha kipande kilichokosekana cha safu ya data.

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

Msimbo wako unapaswa kurekebisha kipande kilichokosekana cha njia ya `.reduce()`. Matokeo ya konsoli yanapaswa kuonyesha kwamba `Sum of array values is: 6`.

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```
