---
id: 587d7b83367417b2b2512b33
title: Tumia konsoli ya JavaScript kukagua thamani ya kigezo
challengeType: 1
forumTopicId: 18372
dashedName: use-the-javascript-console-to-check-the-value-of-a-variable
---

# --description--

Zote Chrome na Firefox zina konsoli nzuri za JavaScript, zinazojulikana pia kama zana za msanidi, kwa ajili ya kurekebisha makosa ya JavaScript yako.

Unaweza kupata zana za msanidi katika menyu ya Chrome yako au Konsoli ya Mtandao katika menyu ya Firefox. Ikiwa unatumia kivinjari tofauti, au simu ya mkononi, tunapendekeza sana kubadili kwenda Firefox au Chrome ya kompyuta ya mezani.

Njia ya `console.log()`, ambayo "inachapisha" matokeo ya kile kilicho ndani ya mabano yake ya kawaida kwenye konsoli, itakuwa chombo muhimu zaidi cha kurekebisha makosa. Kuipanga katika sehemu muhimu za msimbo wako kunaweza kukuonyesha thamani za kati za vigezo. Ni desturi nzuri kuwa na wazo la kile matokeo yanapaswa kuwa kabla ya kuangalia kile yalivyo. Kuwa na sehemu za kukagua hali ya mahesabu yako katika msimbo mzima kutasaidia kupunguza mahali hitilafu iko.

Hapa kuna mfano wa kuchapisha mfuatano wa herufi `Hello world!` kwenye konsoli:

```js
console.log('Hello world!');
```

# --instructions--

Tumia njia ya `console.log()` kuchapisha thamani ya kigezo `a` pale ilipoonyeshwa kwenye msimbo.

# --hints--

Msimbo wako unapaswa kutumia `console.log()` kukagua thamani ya kigezo `a`.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\(a\)/g));
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 1;
a++;
// Only change code below this line


let sumAB = a + b;
console.log(sumAB);
```

# --solutions--

```js
var a = 5; console.log(a);
```
