---
id: 587d7b84367417b2b2512b35
title: Gundua Majina ya Vigezo na Vitendakazi Vilivyoandikwa Vibaya
challengeType: 1
forumTopicId: 301186
dashedName: catch-misspelled-variable-and-function-names
---

# --description--

Njia za `console.log()` na `typeof` ndizo njia kuu mbili za kukagua thamani za kati na aina za matokeo ya programu. Sasa ni wakati wa kuangazia aina za kawaida za hitilafu. Tatizo moja la ngazi ya sintaksia ambalo waandishi wa msimbo wa kasi wanaweza kuelewana nalo ni hitilafu ya tahajia.

Herufi zilizobadilishwa nafasi, kukosekana, au kuandikwa kwa herufi kubwa isivyostahili katika jina la kigezo au kitendakazi zitafanya kivinjari kutafuta kitu ambacho hakipo - na kuliambia hilo kwa njia ya hitilafu ya rejea. Majina ya vigezo na vitendakazi katika JavaScript yanazingatia herufi kubwa na ndogo.

# --instructions--

Rekebisha makosa mawili ya tahajia katika msimbo ili hesabu ya `netWorkingCapital` ifanye kazi.

# --hints--

Kagua tahajia ya vigezo viwili vinavyotumika katika hesabu ya netWorkingCapital, matokeo ya konsoli yanapaswa kuonyesha "Net working capital is: 2".

```js
assert(netWorkingCapital === 2);
```

Hakuna vigezo vyovyote vilivyoandikwa vibaya katika msimbo.

```js
assert(!__helpers.removeJSComments(code).match(/recievables/g));
```

Kigezo cha `receivables` kinapaswa kutangazwa na kutumika ipasavyo katika msimbo.

```js
assert(__helpers.removeJSComments(code).match(/receivables/g).length == 2);
```

Hakuna vigezo vyovyote vilivyoandikwa vibaya katika msimbo.

```js
assert(!__helpers.removeJSComments(code).match(/payable;/g));
```

Kigezo cha `payables` kinapaswa kutangazwa na kutumika ipasavyo katika msimbo.

```js
assert(__helpers.removeJSComments(code).match(/payables/g).length == 2);
```

# --seed--

## --seed-contents--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = recievables - payable;
console.log(`Net working capital is: ${netWorkingCapital}`);
```

# --solutions--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = receivables - payables;
console.log(`Net working capital is: ${netWorkingCapital}`);
```
