---
id: 587d7b85367417b2b2512b3a
title: Kamata hoja zilizopitishwa kwa mpangilio usio sahihi wakati wa kuita kitendakazi
challengeType: 1
forumTopicId: 301184
dashedName: catch-arguments-passed-in-the-wrong-order-when-calling-a-function
---

# --description--

Kuendelea na mjadala kuhusu kuitwa kwa vitendakazi, hitilafu inayofuata ya kuzingatia ni pale hoja za kitendakazi zinapowasilishwa kwa mpangilio usio sahihi. Ikiwa hoja ni aina tofauti, kama vile kitendakazi kinachotarajia safu ya data na nambari kamili, hii huenda ikatupa kosa la wakati wa utekelezaji. Ikiwa hoja ni aina ile ile (kama nambari kamili zote, kwa mfano), basi mantiki ya msimbo haitakuwa na maana. Hakikisha unatoa hoja zote zinazohitajika, kwa mpangilio sahihi ili kuepuka matatizo haya.

# --instructions--

Kitendakazi `raiseToPower` huchukua nambari msingi na kuiongezea nguvu. Kwa bahati mbaya, hakuitwa kwa usahihi - rekebisha msimbo ili thamani ya `power` iwe 8 kama inavyotarajiwa.

# --hints--

Msimbo wako unapaswa kurekebisha kigezo `power` ili kiwe sawa na 2 iliyoinuliwa kwa nguvu ya 3, si 3 iliyoinuliwa kwa nguvu ya 2.

```js
assert(power == 8);
```

Msimbo wako unapaswa kutumia mpangilio sahihi wa hoja kwa wito wa kitendakazi `raiseToPower`.

```js
assert(__helpers.removeJSComments(code).match(/raiseToPower\(\s*?base\s*?,\s*?exp\s*?\);/g));
```

# --seed--

## --seed-contents--

```js
function raiseToPower(b, e) {
  return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(exp, base);
console.log(power);
```

# --solutions--

```js
function raiseToPower(b, e) {
 return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(base, exp);
console.log(power);
```
