---
id: 587d7b84367417b2b2512b34
title: Tumia typeof kukagua aina ya kigezo
challengeType: 1
forumTopicId: 18374
dashedName: use-typeof-to-check-the-type-of-a-variable
---

# --description--

Unaweza kutumia `typeof` kukagua muundo wa data, au aina, ya kigezo. Hii ni muhimu katika kurekebisha makosa unapotumia aina mbalimbali za data. Ikiwa unadhani unaongeza nambari mbili, lakini moja ni mfuatano wa herufi, matokeo yanaweza kuwa yasiyotegemewa. Makosa ya aina yanaweza kuwepo katika mahesabu au wito wa kitendakazi. Kuwa makini hasa unapotumia na kufanya kazi na data za nje katika mfumo wa JavaScript Object Notation (JSON).

Hapa kuna mifano inayotumia `typeof`:

```js
console.log(typeof "");
console.log(typeof 0);
console.log(typeof []);
console.log(typeof {});
```

Kwa mpangilio, konsoli itaonyesha mfuatano wa herufi `string`, `number`, `object`, na `object`.

JavaScript inatambua aina saba za data za msingi (zisizobadilika): `Boolean`, `Null`, `Undefined`, `Number`, `String`, `Symbol` (mpya na ES6), na `BigInt` (mpya na ES2020), na aina moja kwa vitu vinavyobadilika: `Object`. Kumbuka kuwa katika JavaScript, safu za data kwa kiufundi ni aina ya kitu.

# --instructions--

Ongeza tamko mbili za `console.log()` kukagua `typeof` za vigezo viwili `seven` na `three` katika msimbo.

# --hints--

Msimbo wako unapaswa kutumia `typeof` katika tamko mbili za `console.log()` kukagua aina ya vigezo.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\s*\(typeof[\( ].*\)?\)/g).length == 2);
```

Msimbo wako unapaswa kutumia `typeof` kukagua aina ya kigezo `seven`.

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]seven\)?/g));
```

Msimbo wako unapaswa kutumia `typeof` kukagua aina ya kigezo `three`.

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]three\)?/g));
```

# --seed--

## --seed-contents--

```js
let seven = 7;
let three = "3";
console.log(seven + three);
// Only change code below this line
```

# --solutions--

```js
let seven = 7;let three = "3";console.log(typeof seven);
console.log(typeof three);
```
