---
id: 587d7b85367417b2b2512b39
title: Gundua Kukosekana kwa Mabano ya Kawaida ya Kufungua na Kufunga Baada ya Wito wa Kitendakazi
challengeType: 1
forumTopicId: 301185
dashedName: catch-missing-open-and-closing-parenthesis-after-a-function-call
---

# --description--

Wakati kitendakazi au njia haichukui hoja yoyote, unaweza kusahau kujumuisha mabano ya kawaida ya kufungua na kufunga (tupu) unapotaita. Mara nyingi matokeo ya wito wa kitendakazi huhifadhiwa katika kigezo kwa matumizi mengine katika msimbo wako. Hitilafu hii inaweza kugunduliwa kwa kuandika thamani za vigezo (au aina zao) kwenye konsoli na kuona kuwa moja limewekwa kuwa rejea ya kitendakazi, badala ya thamani inayotarajiwa kitendakazi kurudisha.

Vigezo katika mfano ufuatao ni tofauti:

```js
function myFunction() {
  return "You rock!";
}
let varOne = myFunction;
let varTwo = myFunction();
```

Hapa `varOne` ni kitendakazi `myFunction`, na `varTwo` ni mfuatano wa herufi `You rock!`.

# --instructions--

Rekebisha msimbo ili kigezo `result` kiwe kimewekwa thamani inayorudishwa na wito wa kitendakazi `getNine`.

# --hints--

Msimbo wako unapaswa kurekebisha kigezo `result` ili kiwe kimewekwa nambari ambayo kitendakazi `getNine` hurudisha.

```js
assert(result == 9);
```

Msimbo wako unapaswa kuita kitendakazi `getNine`.

```js
assert(__helpers.removeJSComments(code).match(/getNine\(\)/g).length == 2);
```

# --seed--

## --seed-contents--

```js
function getNine() {
  let x = 6;
  let y = 3;
  return x + y;
}

let result = getNine;
console.log(result);
```

# --solutions--

```js
function getNine() {
 let x = 6;
 let y = 3;
 return x + y;
}

let result = getNine();
console.log(result);
```
