---
id: 587d7b86367417b2b2512b3c
title: Tumia tahadhari unaporudiaanzisha vigezo ndani ya mzunguko
challengeType: 1
forumTopicId: 301194
dashedName: use-caution-when-reinitializing-variables-inside-a-loop
---

# --description--

Wakati mwingine ni muhimu kuhifadhi taarifa, kuongeza kielekezi, au kuweka upya vigezo ndani ya mzunguko. Tatizo linaweza kutokea pale ambapo vigezo vinapaswa kuanzishwa upya, lakini havifanyi hivyo, au kinyume chake. Hii ni hatari hasa ikiwa kwa bahati mbaya unaweka upya kigezo kinachotumika kwa hali ya mwisho, na kusababisha mzunguko usioisha.

Kuchapisha thamani za vigezo kila mzunguko kwa kutumia `console.log()` kunaweza kugundua tabia yenye hitilafu inayohusiana na kuweka upya, au kushindwa kuweka upya kigezo.

# --instructions--

Kitendakazi kilicho hapa kinapaswa kuunda safu ya data ya vipimo viwili yenye mistari `m` na safu `n` za sifuri. Kwa bahati mbaya, hakitengenezi matokeo yanayotarajiwa kwa sababu kigezo `row` hakijaanzishwa upya (kurejeshwa kuwa safu tupu) katika mzunguko wa nje. Rekebisha msimbo ili urudishe safu ya data sahihi ya 3x2 ya sifuri, inayofanana na `[[0, 0], [0, 0], [0, 0]]`.

# --hints--

Msimbo wako unapaswa kuweka kigezo `matrix` kuwa safu ya data yenye mistari 3 na safu 2 za sifuri kila moja.

```js
assert(JSON.stringify(matrix) == '[[0,0],[0,0],[0,0]]');
```

Kigezo `matrix` kinapaswa kuwa na mistari 3.

```js
assert(matrix.length == 3);
```

Kigezo `matrix` kinapaswa kuwa na safu 2 katika kila mstari.

```js
assert(
  matrix[0].length == 2 && matrix[1].length === 2 && matrix[2].length === 2
);
```

`zeroArray(4,3)` inapaswa kurudisha safu ya data yenye mistari 4 na safu 3 za sifuri kila moja.

```js
assert(JSON.stringify(zeroArray(4,3)) == '[[0,0,0],[0,0,0],[0,0,0],[0,0,0]]');
```

# --seed--

## --seed-contents--

```js
function zeroArray(m, n) {
  // Creates a 2-D array with m rows and n columns of zeroes
  let newArray = [];
  let row = [];
  for (let i = 0; i < m; i++) {
    // Adds the m-th row into newArray

    for (let j = 0; j < n; j++) {
      // Pushes n zeroes into the current row to create the columns
      row.push(0);
    }
    // Pushes the current row, which now has n zeroes in it, to the array
    newArray.push(row);
  }
  return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```

# --solutions--

```js
function zeroArray(m, n) {
 // Creates a 2-D array with m rows and n columns of zeroes
 let newArray = [];
 for (let i = 0; i < m; i++) {
   let row = [];
   // Adds the m-th row into newArray

   for (let j = 0; j < n; j++) {
     // Pushes n zeroes into the current row to create the columns
     row.push(0);
   }
   // Pushes the current row, which now has n zeroes in it, to the array
   newArray.push(row);
 }
 return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```
