---
id: 587d7b84367417b2b2512b37
title: Gundua Matumizi Mchanganyiko ya Alama za Nukuu Moja na Maradufu
challengeType: 1
forumTopicId: 301188
dashedName: catch-mixed-usage-of-single-and-double-quotes
---

# --description--

JavaScript inaruhusu matumizi ya alama za nukuu moja (`'`) na maradufu (`"`) kutangaza mfuatano wa herufi. Kuamua ni ipi ya kutumia kwa kawaida hutegemea upendeleo binafsi, isipokuwa katika baadhi ya hali.

Kuwa na chaguo mbili ni nzuri wakati mfuatano wa herufi una vifupisho au kipande kingine cha maandishi kilicho ndani ya nukuu. Tafadhali kuwa makini usifunge mfuatano wa herufi mapema sana, kwani hilo linasababisha kosa la sintaksia.

Hapa kuna mifano ya mchanganyiko wa alama za nukuu:

```js
const grouchoContraction = "I've had a perfectly wonderful evening, but this wasn't it.";
const quoteInString = "Groucho Marx once said 'Quote me as saying I was mis-quoted.'";
const uhOhGroucho = 'I've had a perfectly wonderful evening, but this wasn't it.';
```

Mifano miwili ya kwanza ni sahihi, lakini ya tatu si sahihi.

Bila shaka, ni sawa kutumia mtindo mmoja tu wa alama za nukuu. Unaweza kuepuka alama za nukuu ndani ya mfuatano wa herufi kwa kutumia mkwaju wa nyuma (`\`):

```js
const allSameQuotes = 'I\'ve had a perfectly wonderful evening, but this wasn\'t it.';
```

# --instructions--

Rekebisha mfuatano wa herufi ili utumie alama tofauti za nukuu kwa thamani ya `href`, au kuepuka alama zilizopo. Hifadhi alama za nukuu maradufu kuzunguka mfuatano mzima.

# --hints--

Msimbo wako unapaswa kurekebisha alama za nukuu kuzunguka thamani ya `href` `#Home` kwa kubadilisha au kuziepuka.

```js
assert(__helpers.removeJSComments(code).match(/<a href=\s*?('|\\")#Home\1\s*?>/g));
```

Msimbo wako unapaswa kuhifadhi alama za nukuu maradufu kuzunguka mfuatano mzima.

```js
assert(__helpers.removeJSComments(code).match(/"<p>.*?<\/p>";/g));
```

# --seed--

## --seed-contents--

```js
let innerHtml = "<p>Click here to <a href="#Home">return home</a></p>";
console.log(innerHtml);
```

# --solutions--

```js
let innerHtml = "<p>Click here to <a href=\"#Home\">return home</a></p>";
console.log(innerHtml);
```
