---
id: 5d822fd413a79914d39e98c9
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the CSS Variables Skyline project! Start by adding the `!DOCTYPE html` declaration at the top of the document so the browser knows what type of document it's reading.

# --hints--

代码应包含 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

应该在 `DOCTYPE` 声明之后包含一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

应该将文档类型定义为 `html`。

```js
assert(code.match(/html/gi));
```

你应该在类型之后使用 `>` 关闭 `DOCTYPE` 声明。

```js
assert(code.match(/html\s*>/gi));
```

`DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
