---
id: 5f716bee5838c354c728a7c5
title: الخطوة 52
challengeType: 0
dashedName: step-52
---

# --description--

تحت الحلوى (dessert) التي أضفتها للتو، أضف بقية الحلوى والأسعار باستخدام ثلاثة عناصر `article`، كل منها مع يحتوي علي عنصرين `p`. يجب أن يحتوي كل عنصر على النص الصحيح للحلوى والسعر، ويجب أن يحتوي كل عنصر على classes الصحيحة.

```bash
Cherry Pie 2.75
Cheesecake 3.00
Cinnamon Roll 2.50
```

# --hints--

يجب أن يكون لديك أربع عناصر `article` في `section` الثاني الخاص بك.

```js
const secondSection = document.querySelectorAll('section')[1]
const articles = secondSection.querySelectorAll('article')
assert(articles.length === 4)
```

يجب أن يكون لديك أربع عناصر `.dessert`.

```js
assert.lengthOf(document.querySelectorAll('.dessert'),4);
```

يجب أن يكون لديك أربع عناصر `.price` جديدة.

```js
const section = [...document.querySelectorAll('section')].at(-1);
const prices = section.querySelectorAll(`:scope ${'.price'}`);
assert.lengthOf(prices, 4);
```

عنصر `section` الخاص بك يجب أن يحتوي على ثماني عناصر `p`.

```js
const section = [...document.querySelectorAll('section')].at(-1);
const paragraphs = section.querySelectorAll(`:scope ${'p'}`);

assert.lengthOf(paragraphs,8);
```

يجب أن تحتوي عناصر `.dessert` على النصوص `Donut`, و `Cherry Pie`, و `Cheesecake`, و `Cinnamon Roll`.

```js
const dessert = document.querySelectorAll('.dessert');
assert.match(dessert?.[0]?.innerText, /donut/i);
assert.match(dessert?.[1]?.innerText, /cherry pie/i);
assert.match(dessert?.[2]?.innerText, /cheesecake/i);
assert.match(dessert?.[3]?.innerText, /cinnamon roll/i);
```

يجب أن تحتوي عناصر `.price` الجديدة على النصوص `1.50`, و `2.75`, و `3.00`, و `2.50`.

```js
const section = [...document.querySelectorAll('section')].at(-1);
const prices = section?.querySelectorAll(`:scope ${'.price'}`);
assert.match(prices?.[0]?.innerText, /1\.50/);
assert.match(prices?.[1]?.innerText, /2\.75/);
assert.match(prices?.[2]?.innerText, /3\.00/);
assert.match(prices?.[3]?.innerText, /2\.50/);
```

يجب ألا يكون لديك أي مسافات بين عناصر `p` الخاصة بك.

```js
assert(!code.match(/<\/p>\s+<p/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
--fcc-editable-region--
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%;
}
```
