---
id: 587d7dbc367417b2b2512bb1
title: シンプルなJSX要素を作成する
challengeType: 6
forumTopicId: 301390
dashedName: create-a-simple-jsx-element
---

# --description--

ReactはFacebookが作成・管理しているオープンソースのビューライブラリです。モダンなウェブアプリケーションのユーザーインターフェース（UI）をレンダリングするのに最適なツールです。

ReactはJSXというJavaScriptの構文拡張を使っており、JavaScriptの中に直接HTMLを書くことができます。これにはいくつかの利点があります。HTMLの中でJavaScriptのプログラム的な力をフルに使えること、そしてコードの可読性を保ちやすいことです。基本的にJSXはすでに学んだHTMLに似ていますが、いくつかの重要な違いがあり、これらはチャレンジを通じて説明されます。

例えば、JSXはJavaScriptの構文拡張なので、JSXの中に直接JavaScriptを書くことができます。これを行うには、JavaScriptとして扱いたいコードを波括弧, 中括弧で囲みます：`{ 'this is treated as JavaScript code' }`。これは今後のチャレンジでも使うので覚えておいてください。

ただし、JSXは有効なJavaScriptではないため、JSXコードはJavaScriptにコンパイルされる必要があります。この処理にはトランスパイラのBabelがよく使われます。利便性のために、このチャレンジではすでに裏でBabelが追加されています。もし構文的に無効なJSXを書いた場合は、最初のテストが失敗します。

なお、内部的にはチャレンジで`ReactDOM.render(JSX, document.getElementById('root'))`が呼び出されています。この関数呼び出しによって、JSXがReact独自の軽量なDOM表現に変換されます。Reactは自身のDOMのスナップショットを使って、実際のDOMの特定の部分だけを効率的に更新します。

# --instructions--

現在のコードはJSXを使って`div`要素を定数`JSX`に代入しています。`div`を`h1`要素に置き換え、その中にテキスト`Hello JSX!`を追加してください。

# --hints--

定数`JSX`は`h1`要素を返すべきです。

```js
assert(JSX.type === 'h1');
```

`h1`タグにはテキスト`Hello JSX!`を含める必要があります。

```js
assert(Enzyme.shallow(JSX).contains('Hello JSX!'));
```

# --seed--

## --seed-contents--

```jsx
const JSX = <div></div>;
```

# --solutions--

```jsx
const JSX = <h1>Hello JSX!</h1>;
```
