---
id: 5a24c314108439a4d4036160
title: JSXでHTMLのクラスを定義する
challengeType: 6
forumTopicId: 301393
dashedName: define-an-html-class-in-jsx
---

# --description--

JSXの記述に慣れてきたところで、HTMLとどう違うのか気になるかもしれません。

これまでのところ、HTMLとJSXはまったく同じに見えるかもしれません。

JSXでの大きな違いの一つは、HTMLのクラスを定義するのに`class`という単語を使えなくなったことです。これは`class`がJavaScriptの予約語だからです。代わりにJSXでは`className`を使います。

実際、JSXではすべてのHTML属性やイベントの参照の命名慣例がキャメルケースになります。たとえば、JSXでのクリックイベントは`onClick`であり、`onclick`ではありません。同様に、`onchange`は`onChange`になります。これは微妙な違いですが、今後覚えておくべき重要なポイントです。

# --instructions--

JSXコード内の`myDiv`に`div`というクラスを適用してください。

# --hints--

定数`JSX`は`div`要素を返すはずです。

```js
assert.strictEqual(JSX.type, 'div');
```

`div`は`myDiv`というクラスを持つべきです。

```js
assert.strictEqual(JSX.props.className, 'myDiv');
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h1>Add a class to this div</h1>
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div className = 'myDiv'>
  <h1>Add a class to this div</h1>
</div>);
```
