---
id: 5a24c314108439a4d4036169
title: ステートレス関数コンポーネントにプロップスを渡す
challengeType: 6
forumTopicId: 301402
dashedName: pass-props-to-a-stateless-functional-component
---

# --description--

前のチャレンジでは、JSX要素の作成と合成、関数コンポーネント、ES6スタイルのクラスコンポーネントについて多く学びました。この基礎をもとに、Reactで非常に一般的なもう一つの機能、**プロップス**について見ていきましょう。Reactでは、子コンポーネントにプロップス（プロパティ）を渡すことができます。たとえば、`App` コンポーネントがステートレス関数コンポーネントである`Welcome`という子コンポーネントをレンダリングするとします。`Welcome` に`user` プロパティを渡すには、次のように書きます。

```jsx
<App>
  <Welcome user='Mark' />
</App>
```

ここでは、Reactがサポートする**カスタムHTML属性**を使ってコンポーネントに渡しています。この場合、作成したプロパティ`user`が`Welcome`コンポーネントに渡されます。`Welcome`はステートレス関数コンポーネントなので、この値には次のようにアクセスできます。

```jsx
const Welcome = (props) => <h1>Hello, {props.user}!</h1>
```

この値を`props`と呼ぶのが一般的で、ステートレス関数コンポーネントの場合は、JSXを返す関数の引数として考えます。関数の本体で引数の値にアクセスできます。クラスコンポーネントの場合は少し異なります。

# --instructions--

コードエディターには`Calendar`と`CurrentDate`のコンポーネントがあります。`CurrentDate`コンポーネントから`Calendar`をレンダリングするときに、JavaScriptの`date`オブジェクトから現在の日付を代入した`Date`プロパティを渡してください。そして、この`prop`を`CurrentDate`コンポーネントでアクセスし、その値を`p`タグ内に表示してください。`prop`の値をJavaScriptとして評価するには、中括弧で囲む必要があることに注意してください。たとえば`date={Date()}`のようにします。

# --hints--

`Calendar`コンポーネントは単一の`div`要素を返すべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`Calendar`コンポーネントの2番目の子は`CurrentDate`コンポーネントであるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).name() === 'CurrentDate';
  })()
);
```

`CurrentDate`コンポーネントは`date`というプロップを持つべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).props().date;
  })()
);
```

`date`プロップは`CurrentDate`の中に文字列のテキストを含むべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    const prop = mockedComponent.children().childAt(1).props().date;
    return typeof prop === 'string' && prop.length > 0;
  })()
);
```

`date`プロップは`Date()`を呼び出して生成されるべきです。

```js
assert(/<CurrentDatedate={Date\(\)}\/>/.test(__helpers.removeWhiteSpace(code)));
```

`CurrentDate`コンポーネントは`date`プロップの値を`p`タグ内でレンダリングするべきです。

```js
let date = 'dummy date';
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(CurrentDate, { date })
    );
    return mockedComponent.find('p').html().includes(date);
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: </p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: {props.date}</p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate date={Date()} />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
