---
id: 5a24c314108439a4d4036166
title: Reactコンポーネントの合成
challengeType: 6
forumTopicId: 301381
dashedName: compose-react-components
---

# --description--

ReactコンポーネントとJSXを使ったより複雑な合成が続く中で、重要なポイントがあります。ES6スタイルのクラスコンポーネントを他のコンポーネント内でレンダリングすることは、これまでのチャレンジで使ったシンプルなコンポーネントをレンダリングするのと何も変わりません。JSX要素、ステートレス関数コンポーネント、ES6クラスコンポーネントを他のコンポーネント内でレンダリングできます。

# --instructions--

コードエディターでは、`TypesOfFood`コンポーネントがすでに`Vegetables`というコンポーネントをレンダリングしています。また、前回のチャレンジで使った`Fruits`コンポーネントもあります。

`Fruits`の中に2つのコンポーネントをネストしてください。まず`NonCitrus`、次に`Citrus`です。これらのコンポーネントは裏で用意されています。次に、`Fruits`クラスコンポーネントを`TypesOfFood`コンポーネントの中に、`h1`見出し要素の下、`Vegetables`の上にネストしてください。結果として、2種類の異なるコンポーネントタイプを使った一連のネストされたコンポーネントができます。

# --before-all--

```js
globalThis.NonCitrus = function NonCitrus() {
  return React.createElement('div', null);
};
globalThis.Citrus = function Citrus() {
  return React.createElement('div', null);
};
globalThis.Vegetables = function Vegetables() {
  return React.createElement('div', null);
};
```

# --hints--

`TypesOfFood`コンポーネントは単一の`div`要素を返すべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`TypesOfFood`コンポーネントは`Fruits`コンポーネントを返すべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

`Fruits`コンポーネントは`NonCitrus`コンポーネントと`Citrus`コンポーネントを返すべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return (
      mockedComponent.find('Fruits').children().find('NonCitrus').length ===
        1 &&
      mockedComponent.find('Fruits').children().find('Citrus').length === 1
    );
  })()
);
```

`TypesOfFood`コンポーネントは`Vegetables`コンポーネントの下に`Fruits`コンポーネントを返すべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
        <Vegetables />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }
        <NonCitrus />
        <Citrus />
        { /* Change code above this line */ }
      </div>
    )
  }
}

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
    render() {
      return (
        <div>
        <h1>Types of Food:</h1>
          { /* Change code below this line */ }
          <Fruits />
          { /* Change code above this line */ }
          <Vegetables />
        </div>
      );
    }
};
```
