---
id: 5a24c314108439a4d403618b
title: 兄弟要素に一意のキー属性を付ける
challengeType: 6
forumTopicId: 301394
dashedName: give-sibling-elements-a-unique-key-attribute
---

# --description--

前回のチャレンジでは、`map` メソッドを使ってユーザーの入力に基づき複数の要素を動的にレンダリングする方法を示しました。しかし、その例には重要な部分が欠けていました。配列で要素を作成するときは、それぞれに一意の値を持つ `key` 属性を設定する必要があります。React はこれらのキーを使って、どのアイテムが追加、変更、または削除されたかを追跡します。これにより、リストが何らかの形で変更されたときの再レンダリング処理がより効率的になります。

**注意:** キーは兄弟要素間で一意であればよく、アプリケーション全体でグローバルに一意である必要はありません。

# --instructions--

コードエディターにはいくつかのフロントエンドフレームワークを含む配列と、ステートレス関数コンポーネント `Frameworks()` があります。`Frameworks()` は前回のチャレンジのように配列を番号なしリストにマッピングする必要があります。`map` コールバックの記述を完成させて、`li` 配列の各フレームワークに対して `frontEndFrameworks` 要素を返してください。今回は、各 `li` に一意の値を持つ `key` 属性を必ず付けてください。`li` 要素には `frontEndFrameworks` からのテキストも含める必要があります。

通常、キーはレンダリングされる要素を一意に識別するものにします。最終手段として配列のインデックスを使うこともできますが、通常は一意の識別子を使うことをお勧めします。

# --hints--

`Frameworks` コンポーネントは存在し、ページにレンダリングされるべきです。

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('Frameworks').length === 1
);
```

`Frameworks` は `h1` 要素をレンダリングするべきです。

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('h1').length === 1);
```

`Frameworks` は `ul` 要素をレンダリングするべきです。

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('ul').length === 1);
```

`ul` タグは 6 つの子 `li` 要素をレンダリングするべきです。

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('ul').children().length ===
    6 &&
    Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .childAt(0)
      .name() === 'li' &&
    Enzyme.mount(React.createElement(Frameworks)).find('li').length === 6
);
```

各リストアイテム要素には一意の `key` 属性が必要です。

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(Frameworks)).find('ul');
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key(),
      ul.childAt(4).key(),
      ul.childAt(5).key()
    ]);
    return keys.size === 6;
  })()
);
```

各リストアイテム要素には `frontEndFrameworks` からのテキストが含まれるべきです。

```js
assert(
  (() => {
    const li = Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .children();
    return [...Array(5)].every((_, i) =>
      frontEndFrameworks.includes(li.at(i).text())
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = null; // Change this line
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```

# --solutions--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = frontEndFrameworks.map((fw, i) => <li key={i}>{fw}</li>);
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```
