---
id: 5a24c314108439a4d4036185
title: より簡潔な条件分岐に && を使う
challengeType: 6
forumTopicId: 301413
dashedName: use--for-a-more-concise-conditional
---

# --description--

前回のチャレンジでは `if/else` 文が機能しましたが、同じ結果をより簡潔に実現する方法があります。コンポーネントで複数の条件を追跡していて、それぞれの条件に応じて異なる要素をレンダリングしたいと想像してください。少しずつ異なるUIを返すために多くの `else if` 文を書くと、コードが重複してミスの原因になることがあります。代わりに、`&&` 論理演算子を使って、より簡潔に条件ロジックを実行できます。これは、条件が `true` かどうかをチェックし、そうであればマークアップを返したい場合に可能です。例を示します。

```jsx
{condition && <p>markup</p>}
```

`condition` が `true` の場合、マークアップが返されます。条件が `false` の場合は、`false` を評価した後すぐに `condition` が返され、何も返されません。これらの文はJSXに直接含めることができ、各条件の後に `&&` を書くことで複数の条件を連結できます。これにより、多くのコードを繰り返すことなく、`render()` メソッド内でより複雑な条件ロジックを扱えます。

# --instructions--

前の例をもう一度解いて、`h1` が `display` が `true` の場合にのみレンダリングされるようにしてください。ただし、`&&` 文の代わりに `if/else` 論理演算子を使ってください。

# --hints--

`MyComponent` は存在し、レンダリングされるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length;
  })()
);
```

`display` が `true` に設定されているとき、`div`、`button`、および `h1` がレンダリングされるべきです。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 2 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 1
  );
```

`display` が `false` に設定されているとき、`div` と `button` のみがレンダリングされるべきです。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 1 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 0
  );
```

render メソッドは `&&` 論理演算子を使って `this.state.display` の条件をチェックするべきです。

```js
assert(code.includes('&&'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         {this.state.display && <h1>Displayed!</h1>}
       </div>
    );
  }
};
```
