---
id: 5a24c314108439a4d4036167
title: DOMにクラスコンポーネントをレンダリングする
challengeType: 6
forumTopicId: 301404
dashedName: render-a-class-component-to-the-dom
---

# --description--

以前のチャレンジでReactDOM APIを使ってJSX要素をDOMにレンダリングしたことを覚えているかもしれません。Reactコンポーネントをレンダリングする手順も非常に似ています。ここ数回のチャレンジはコンポーネントとコンポジションに焦点を当てていたため、レンダリングは裏で自動的に行われていました。しかし、ReactのコードはReactDOM APIを呼び出さなければDOMにレンダリングされません。

構文のおさらいです：`ReactDOM.render(componentToRender, targetNode)`。最初の引数はレンダリングしたいReactコンポーネントです。2番目の引数は、そのコンポーネントをレンダリングしたいDOMノードです。

ReactコンポーネントはJSX要素とは少し違う方法で`ReactDOM.render()`に渡されます。JSX要素の場合はレンダリングしたい要素名を渡しますが、Reactコンポーネントの場合はネストされたコンポーネントをレンダリングする時と同じ構文を使う必要があります。例えば`ReactDOM.render(<ComponentToRender />, targetNode)`のようにです。この構文はES6クラスコンポーネントと関数コンポーネントの両方で使います。

# --instructions--

`Fruits`と`Vegetables`の両方のコンポーネントは裏で定義されています。`TypesOfFood`コンポーネントの子として両方のコンポーネントをレンダリングし、その後`TypesOfFood`をDOMにレンダリングしてください。`div`と`id='challenge-node'`が利用可能です。

# --hints--

`TypesOfFood`コンポーネントは単一の`div`要素を返すべきです。

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`TypesOfFood`コンポーネントは`Fruits`要素の後に`h1`コンポーネントをレンダリングするべきです。

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

`TypesOfFood`コンポーネントは`Vegetables`の後に`Fruits`コンポーネントをレンダリングするべきです。

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

`TypesOfFood`コンポーネントはidが`div`の`challenge-node`内にDOMへレンダリングされるべきです。

```js
assert(
  (function () {
    const hasRenderCall = /ReactDOM\.render\(<TypesOfFood\/>,document\.getElementById\((['"`])challenge-node\1\)\);?/.test(
      __helpers.removeWhiteSpace(code)
    );
    if (!hasRenderCall) {
      return false;
    }

    const e = React.createElement;
    globalThis.Fruits = function Fruits() {
      return e(
        'div',
        null,
        e('h2', null, 'Fruits:'),
        e('h4', null, 'Non-Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Apples'),
          e('li', null, 'Blueberries'),
          e('li', null, 'Strawberries'),
          e('li', null, 'Bananas')
        ),
        e('h4', null, 'Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Lemon'),
          e('li', null, 'Lime'),
          e('li', null, 'Orange'),
          e('li', null, 'Grapefruit')
        )
      );
    };
    globalThis.Vegetables = function Vegetables() {
      return e(
        'div',
        null,
        e('h2', null, 'Vegetables:'),
        e(
          'ul',
          null,
          e('li', null, 'Brussel Sprouts'),
          e('li', null, 'Broccoli'),
          e('li', null, 'Squash')
        )
      );
    };

    ReactDOM.render(
      React.createElement(TypesOfFood),
      document.getElementById('challenge-node')
    );

    const root = document.getElementById('challenge-node').childNodes[0];
    if (!root) {
      return false;
    }
    const html = root.innerHTML;
    return (
      html.includes(
        '<div><h2>Fruits:</h2><h4>Non-Citrus:</h4><ul><li>Apples</li><li>Blueberries</li><li>Strawberries</li><li>Bananas</li></ul><h4>Citrus:</h4><ul><li>Lemon</li><li>Lime</li><li>Orange</li><li>Grapefruit</li></ul></div>'
      ) &&
      html.includes(
        '<div><h2>Vegetables:</h2><ul><li>Brussel Sprouts</li><li>Broccoli</li><li>Squash</li></ul></div>'
      )
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
```

# --solutions--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}
          <Fruits />
           <Vegetables />
         {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
ReactDOM.render(<TypesOfFood />, document.getElementById('challenge-node'));
```
