---
id: 5a24c314108439a4d4036187
title: 三項式を使った条件付きレンダリング
challengeType: 6
forumTopicId: 301414
dashedName: use-a-ternary-expression-for-conditional-rendering
---

# --description--

動的レンダリングの手法に進む前に、組み込みのJavaScript条件文を使って表示内容を制御する最後の方法をご紹介します。それが<dfn>三項演算子</dfn>です。三項演算子はJavaScriptの`if/else`文のショートカットとしてよく使われます。従来の`if/else`文ほど強力ではありませんが、React開発者の間で非常に人気があります。その理由の一つは、JSXがコンパイルされる仕組みによって、`if/else`文をJSXコード内に直接挿入できないためです。数チャレンジ前に気づいたかもしれませんが、`if/else`文が必要な場合は常に`return`文の*外側*にありました。JSX内で条件ロジックを実装したい場合、三項式は優れた代替手段となります。三項演算子は3つの部分から成りますが、複数の三項式を組み合わせることも可能です。基本的な構文は以下の通りです。

```jsx
condition ? expressionIfTrue : expressionIfFalse;
```

# --instructions--

コードエディターには`CheckUserAge`コンポーネントの`render()`メソッド内で3つの定数が定義されています。名前は`buttonOne`、`buttonTwo`、`buttonThree`です。これらはそれぞれボタン要素を表すシンプルなJSX式が代入されています。まず、`CheckUserAge`の状態を初期化し、`input`と`userAge`の両方を空文字列の値に設定してください。

コンポーネントがページに情報をレンダリングしたら、ユーザーが操作できるようにします。コンポーネントの`return`文内で、次のロジックを実装する三項式を設定してください。ページが最初に読み込まれたときは送信ボタン`buttonOne`をページにレンダリングします。次に、ユーザーが年齢を入力してボタンをクリックしたら、年齢に応じて異なるボタンをレンダリングします。ユーザーが`18`未満の数字を入力した場合は`buttonThree`をレンダリングし、`18`以上の数字を入力した場合は`buttonTwo`をレンダリングしてください。

# --hints--

`CheckUserAge`コンポーネントは1つの`input`要素と1つの`button`要素をレンダリングする必要があります。

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('input')
    .length === 1 &&
    Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('button')
      .length === 1
);
```

`CheckUserAge`コンポーネントの状態は、`userAge`プロパティと`input`プロパティの両方を空文字列の値で初期化してください。

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).state().input === '' &&
    Enzyme.mount(React.createElement(CheckUserAge)).state().userAge === ''
);
```

`CheckUserAge`コンポーネントが最初にDOMにレンダリングされたとき、`button`の内部テキストはSubmitである必要があります。

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('button').text() ===
    'Submit'
);
```

`input`要素に18未満の数字が入力され、`button`がクリックされたとき、`button`の内部テキストは`You Shall Not Pass`と表示される必要があります。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter3AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '3' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter17AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '17' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge3 = enter3AndClickButton();
  const userAge17 = enter17AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge3 === 'You Shall Not Pass' &&
      userAge17 === 'You Shall Not Pass'
  );
})();
```

`input`要素に18以上の数字が入力され、`button`がクリックされたとき、`button`の内部テキストは`You May Enter`と表示される必要があります。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter35AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '35' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const userAge35 = enter35AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge18 === 'You May Enter' &&
      userAge35 === 'You May Enter'
  );
})();
```

数字が提出された後、`input`の値が再び変更されると、`button`は再び`Submit`と表示される必要があります。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const changeInputDontClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '5' } });
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter10AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '10' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const changeInput1 = changeInputDontClickButton();
  const userAge10 = enter10AndClickButton();
  const changeInput2 = changeInputDontClickButton();
  assert(
    userAge18 === 'You May Enter' &&
      changeInput1 === 'Submit' &&
      userAge10 === 'You Shall Not Pass' &&
      changeInput2 === 'Submit'
  );
})();
```

コード内に`if/else`文を含めないでください。

```js
assert(
  new RegExp(/(\s|;)if(\s|\()/).test(
    code
  ) === false
);
```

# --seed--

## --seed-contents--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      userAge: '',
      input: ''
    };
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {this.state.userAge === ''
          ? buttonOne
          : this.state.userAge >= 18
          ? buttonTwo
          : buttonThree}
      </div>
    );
  }
}
```
