---
id: 5a24c314108439a4d4036178
title: 制御された入力を作成する
challengeType: 6
forumTopicId: 301385
dashedName: create-a-controlled-input
---

# --description--

アプリケーションでは、`state` とレンダリングされた UI の間でより複雑なやり取りがある場合があります。たとえば、`input` や `textarea` のようなテキスト入力用のフォームコントロール要素は、ユーザーが入力するたびに DOM 内で独自の状態を保持します。React を使うと、この変更可能な状態を React コンポーネントの `state` に移動できます。ユーザーの入力はアプリケーションの `state` の一部となり、React がその入力フィールドの値を制御します。通常、ユーザーが入力できる入力フィールドを持つ React コンポーネントがある場合、それは制御された入力フォームになります。

# --instructions--

コードエディタには、制御された `ControlledInput` 要素を作成するための `input` というコンポーネントの骨組みがあります。コンポーネントの `state` はすでに空の文字列を保持する `input` プロパティで初期化されています。この値は、ユーザーが `input` フィールドに入力するテキストを表します。

まず、`handleChange()` という名前のメソッドを作成し、`event` というパラメータを持たせてください。このメソッドが呼び出されると、`event` 要素からの文字列を含む `input` オブジェクトを受け取ります。メソッド内で `event.target.value` を使ってこの文字列にアクセスできます。コンポーネントの `input` の `state` プロパティをこの新しい文字列で更新してください。

`render` メソッド内で、`input` タグの上に `h4` 要素を作成してください。`value` 属性を追加し、コンポーネントの `input` の `state` プロパティと等しくしてください。次に、`onChange` プロパティを追加し、`handleChange()` イベントハンドラーメソッドを設定してください。

入力ボックスに文字を入力すると、そのテキストは `handleChange()` メソッドで処理され、ローカルの `input` にある `state` プロパティとして設定され、ページ上の `input` ボックスの値としてレンダリングされます。コンポーネントの `state` は入力データに関する唯一の真実の情報源です。

最後に、コンストラクタで必要なバインディングを追加するのを忘れないでください。

# --hints--

`ControlledInput` は `div` 要素を返し、その中に `input` と `p` タグを含める必要があります。

```js
assert(
  Enzyme.mount(React.createElement(ControlledInput))
    .find('div')
    .children()
    .find('input').length === 1 &&
    Enzyme.mount(React.createElement(ControlledInput))
      .find('div')
      .children()
      .find('p').length === 1
);
```

`ControlledInput` の状態は、空の文字列に設定された `input` プロパティで初期化する必要があります。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(ControlledInput)).state('input'),
  ''
);
```

入力要素に文字を入力すると、状態と入力の値が更新され、`p` 要素は入力中にこの状態をレンダリングします。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(ControlledInput));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return waitForIt(() => mockedComponent.state('input'));
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => ({
      state: mockedComponent.state('input'),
      text: mockedComponent.find('p').text(),
      inputVal: mockedComponent.find('input').props().value
    }));
  };
  const before = await _1();
  const after = await _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.text === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
```

# --seed--

## --seed-contents--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        { /* Change code below this line */}

        { /* Change code above this line */}
        <h4>Controlled Input:</h4>
        <p>{this.state.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  render() {
    return (
      <div>
        <input
          value={this.state.input}
          onChange={this.handleChange} />
        <h4>Controlled Input:</h4>

        <p>{this.state.input}</p>
      </div>
    );
  }
};
```
