---
id: 5a24c314108439a4d4036184
title: If-Else 条件でレンダリングする
challengeType: 6
forumTopicId: 301410
dashedName: render-with-an-if-else-condition
---

# --description--

JavaScriptを使ってレンダリングされるビューを制御するもう一つの方法は、レンダリングされる要素を条件に結びつけることです。条件が真のときはあるビューがレンダリングされ、偽のときは別のビューがレンダリングされます。これはReactコンポーネントの`if/else`メソッド内で標準の`render()`文を使って実現できます。

# --instructions--

MyComponentには、UIにある要素を表示するかどうかを追跡する`boolean`がstateにあります。`button`はこの値の状態を切り替えます。現在は毎回同じUIをレンダリングしています。`render()`メソッドを`if/else`文を使って書き換えてください。`display`が`true`の場合は現在のマークアップを返し、それ以外の場合は`h1`要素を含まないマークアップを返してください。

**注意:** テストに合格するには`if/else`を書く必要があります。三項演算子の使用はここでは合格しません。

# --hints--

`MyComponent`は存在し、レンダリングされるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length === 1;
  })()
);
```

`display`が`true`に設定されているとき、`div`、`button`、`h1`がレンダリングされるべきです。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 2 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 1
  );
```

`display`が`false`に設定されているときは、`div`と`button`だけがレンダリングされるべきです。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 0
  );
```

renderメソッドは`if/else`文を使って`this.state.display`の条件をチェックするべきです。

```js
assert(code.includes('if') && code.includes('else'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line

    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
 }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    if (this.state.display) {
      return (
         <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
           <h1>Displayed!</h1>
         </div>
      );
    } else {
      return (
        <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
         </div>
      );
    }
  }
};
```
