---
id: 5a24c314108439a4d403616e
title: this.props を使ってプロップスにアクセスする
challengeType: 6
forumTopicId: 301375
dashedName: access-props-using-this-props
---

# --description--

これまでのチャレンジでは、子コンポーネントにプロップスを渡す基本的な方法を扱いました。しかし、渡す先の子コンポーネントがステートレス関数コンポーネントではなく、ES6クラスコンポーネントだったらどうでしょうか？ES6クラスコンポーネントは、プロップスにアクセスするために少し異なる慣例を使います。

クラスコンポーネント内で自分自身を参照するときは、`this` キーワードを使います。クラスコンポーネント内でプロップスにアクセスするには、アクセスするコードの前に `this` を付けます。例えば、ES6クラスコンポーネントに `data` というプロップがある場合、JSX内では `{this.props.data}` と書きます。

# --instructions--

親コンポーネント `Welcome` 内に `App` コンポーネントのインスタンスをレンダリングしてください。ここで、`Welcome` に `name` というプロップを与え、その値に文字列を代入してください。子コンポーネント `Welcome` の中で、`name` タグ内に `strong` プロップにアクセスしてください。

# --hints--

`App` コンポーネントは単一の `div` 要素を返すべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`App` の子は `Welcome` コンポーネントであるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.children().childAt(0).name() === 'Welcome'
    );
  })()
);
```

`Welcome` コンポーネントは `name` という名前のプロップを持つべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.find('Welcome').props().name;
  })()
);
```

`Welcome` コンポーネントは `name` プロップとして渡された文字列を `strong` タグ内に表示するべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.find('strong').text() ===
      mockedComponent.find('Welcome').props().name
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome />
            { /* Change code above this line */ }
        </div>
    );
  }
};

class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong></strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};
```

# --solutions--

```jsx
class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong>{this.props.name}</strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};

class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome name="Quincy"/>
            { /* Change code above this line */ }
        </div>
    );
  }
};
```
