---
id: 5a24c314108439a4d403617d
title: ライフサイクルメソッド componentDidMount の使い方
challengeType: 6
forumTopicId: 301422
dashedName: use-the-lifecycle-method-componentdidmount
---

# --description--

ほとんどのウェブ開発者は、ある時点でAPIエンドポイントを呼び出してデータを取得する必要があります。Reactを使っている場合、この処理をどこで行うかを知ることが重要です。

Reactのベストプラクティスとしては、API呼び出しやサーバーへの呼び出しはライフサイクルメソッド`componentDidMount()`の中に置くことです。このメソッドはコンポーネントがDOMにマウントされた後に呼び出されます。ここで`setState()`を呼び出すと、コンポーネントの再レンダリングがトリガーされます。このメソッド内でAPIを呼び出し、APIから返されたデータで状態を設定すると、データを受け取った時点で自動的に更新が行われます。

# --instructions--

`componentDidMount()`にはモックのAPI呼び出しがあります。これはサーバーからデータを取得する動作をシミュレートするために2.5秒後に状態を設定します。この例ではサイトの現在のアクティブユーザー数の合計をリクエストしています。renderメソッド内で、テキスト`activeUsers`の後に`h1`の値を`Active Users:`にレンダリングしてください。プレビューで何が起こるかを確認し、タイムアウトの時間を変えて異なる効果を試してみてください。

# --hints--

`MyComponent`は`div`要素をレンダリングし、その中に`h1`タグを含めるべきです。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

コンポーネントの状態は`componentDidMount`の中のタイムアウト関数で更新する必要があります。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return new RegExp('setTimeout(.|\n)+setState(.|\n)+activeUsers').test(
      String(mockedComponent.instance().componentDidMount)
    );
  })()
);
```

`h1`タグは`activeUsers`の状態から`MyComponent`の値をレンダリングするべきです。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ activeUsers: 1237 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ activeUsers: 1000 });
    return mockedComponent.find('h1').text();
  };
  assert(new RegExp('1237').test(first()) && new RegExp('1000').test(second()));
})();
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        {/* Change code below this line */}
        <h1>Active Users: </h1>
        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        <h1>Active Users: {this.state.activeUsers}</h1>
      </div>
    );
  }
}
```
