---
id: 5a24c314108439a4d4036165
title: ネストされたコンポーネントをレンダリングするために React を使う
challengeType: 6
forumTopicId: 301420
dashedName: use-react-to-render-nested-components
---

# --description--

前回のチャレンジでは、2つのコンポーネントを組み合わせる簡単な方法を紹介しましたが、React ではコンポーネントを組み合わせる方法が他にもたくさんあります。

コンポーネントの組み合わせは React の強力な機能の一つです。React を使うときは、前回のチャレンジの App の例のように、ユーザーインターフェースをコンポーネントという単位で考え始めることが重要です。UI を基本的な構成要素に分解し、それらのパーツがコンポーネントになります。これにより、UI を担当するコードとアプリケーションのロジックを担当するコードを分けることができます。複雑なプロジェクトの開発やメンテナンスが大幅に簡単になります。

# --instructions--

コードエディターには `TypesOfFruit` と `Fruits` という2つの関数コンポーネントが定義されています。`TypesOfFruit` コンポーネントを取り、それを `Fruits` コンポーネントの中に組み合わせる、つまりネストしてください。次に `Fruits` コンポーネントを取り、それを `TypesOfFood` コンポーネントの中にネストしてください。結果として、子コンポーネントが親コンポーネントの中にネストされ、その親コンポーネントもさらに別の親コンポーネントの中にネストされる形になります！

# --hints--

`TypesOfFood` コンポーネントは単一の `div` 要素を返す必要があります。

```js
assert(Enzyme.shallow(React.createElement(TypesOfFood)).type() === 'div');
```

`TypesOfFood` コンポーネントは `Fruits` コンポーネントを返す必要があります。

```js
assert(
  Enzyme.shallow(React.createElement(TypesOfFood)).props().children[1].type
    .name === 'Fruits'
);
```

`Fruits` コンポーネントは `TypesOfFruit` コンポーネントを返す必要があります。

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('h2').html() ===
    '<h2>Fruits:</h2>'
);
```

`TypesOfFruit` コンポーネントは `h2` と `ul` の要素を返す必要があります。

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('ul').text() ===
    'ApplesBlueberriesStrawberriesBananas'
);
```

# --seed--

## --seed-contents--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }

      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }
        <TypesOfFruit />
      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }
        <Fruits />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
