---
id: 5a24c314108439a4d4036188
title: プロップスから条件付きレンダリング
challengeType: 6
forumTopicId: 301405
dashedName: render-conditionally-from-props
---

# --description--

これまでに、`if/else`、`&&`、および三項演算子（`condition ? expressionIfTrue : expressionIfFalse`）を使って、何をいつレンダリングするかを条件付きで決める方法を見てきました。しかし、これらの概念を強力なReactの機能であるプロップスと組み合わせることができる重要なトピックがまだあります。React開発者は、プロップの値を使って自動的にレンダリングの判断をすることが非常に一般的です。

このチャレンジでは、子コンポーネントを設定してプロップスに基づいてレンダリングの判断を行います。三項演算子も使いますが、ここで扱う他のいくつかの概念も同様に役立つことがわかるでしょう。

# --instructions--

コードエディターには、部分的に定義された2つのコンポーネントがあります。親コンポーネントは`GameOfChance`、子コンポーネントは`Results`です。これらは、ユーザーがボタンを押して勝ちか負けかを確認するシンプルなゲームを作成するために使われます。

まず、実行するたびに異なる値をランダムに返す簡単な式が必要です。`Math.random()`を使うことができます。このメソッドは呼び出されるたびに`0`（含む）から`1`（含まない）までの値を返します。したがって、50/50の確率にするには、式の中で`Math.random() >= .5`を使います。統計的に、この式は50％の確率で`true`を返し、残りの50％で`false`を返します。renderメソッド内で、`null`を上記の式に置き換えて変数宣言を完成させてください。

これでコード内でランダムな判断を行うための式ができました。次にこれを実装します。`Results`の子として`GameOfChance`コンポーネントをレンダリングし、`expression`を`fiftyFifty`という名前のプロップとして渡してください。`Results`コンポーネント内で、`h1`から渡される`You Win!`プロップに基づいて、テキストが`You Lose!`または`fiftyFifty`の`GameOfChance`要素をレンダリングする三項式を書いてください。最後に、`handleClick()`メソッドが各ターンを正しくカウントしていることを確認し、ユーザーが何回プレイしたかがわかるようにしてください。これにより、ユーザーが連続して勝ったり負けたりした場合でもコンポーネントが実際に更新されたことがわかります。

# --hints--

`GameOfChance`コンポーネントは存在し、ページにレンダリングされる必要があります。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('GameOfChance').length,
  1
);
```

`GameOfChance`は単一の`button`要素を返す必要があります。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('button').length,
  1
);
```

`GameOfChance`は`Results`というプロップを持つ`fiftyFifty`コンポーネントの単一インスタンスを返す必要があります。

```js
assert(
  Enzyme.mount(React.createElement(GameOfChance)).find('Results').length ===
    1 &&
    Enzyme.mount(React.createElement(GameOfChance))
      .find('Results')
      .props()
      .hasOwnProperty('fiftyFifty') === true
);
```

`GameOfChance`のstateは、`counter`というプロパティを`1`の値で初期化する必要があります。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).state().counter,
  1
);
```

`GameOfChance`コンポーネントが最初にDOMにレンダリングされるとき、内部テキストが`p`の`Turn: 1`要素が返される必要があります。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('p').text(),
  'Turn: 1'
);
```

ボタンがクリックされるたびに、カウンターのstateが1ずつ増加し、カウンターのstateの値`p`を含むテキスト`Turn: N`を持つ単一の`N`要素がDOMにレンダリングされる必要があります。

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    count: comp.state('counter'),
    text: comp.find('p').text()
  });
  const _1 = () => {
    simulate();
    return result();
  };
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _1_val = _1();
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  assert(
    _1_val.count === 2 &&
      _1_val.text === 'Turn: 2' &&
      _2_val.count === 3 &&
      _2_val.text === 'Turn: 3' &&
      _3_val.count === 4 &&
      _3_val.text === 'Turn: 4' &&
      _4_val.count === 5 &&
      _4_val.text === 'Turn: 5' &&
      _5_val.count === 6 &&
      _5_val.text === 'Turn: 6'
  );
})();
```

`GameOfChance`コンポーネントが最初にDOMにマウントされ、以降ボタンがクリックされるたびに、`h1`または`You Win!`のどちらかをランダムにレンダリングする単一の`You Lose!`要素が返される必要があります。注意：これはランダムに失敗する場合があります。その場合は、もう一度試してください。

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    h1: comp.find('h1').length,
    text: comp.find('h1').text()
  });
  const _1 = result();
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _6 = () => {
    simulate();
    return result();
  };
  const _7 = () => {
    simulate();
    return result();
  };
  const _8 = () => {
    simulate();
    return result();
  };
  const _9 = () => {
    simulate();
    return result();
  };
  const _10 = () => {
    simulate();
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  const _6_val = _6();
  const _7_val = _7();
  const _8_val = _8();
  const _9_val = _9();
  const _10_val = _10();
  const __text = new Set([
    _1.text,
    _2_val.text,
    _3_val.text,
    _4_val.text,
    _5_val.text,
    _6_val.text,
    _7_val.text,
    _8_val.text,
    _9_val.text,
    _10_val.text
  ]);
  const __h1 = new Set([
    _1.h1,
    _2_val.h1,
    _3_val.h1,
    _4_val.h1,
    _5_val.h1,
    _6_val.h1,
    _7_val.h1,
    _8_val.h1,
    _9_val.h1,
    _10_val.h1
  ]);
  assert(__text.size === 2 && __h1.size === 1);
})();
```

# --seed--

## --seed-contents--

```jsx
class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    {/* Change code below this line */}
    return <h1></h1>;
    {/* Change code above this line */}
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      // Complete the return statement:
      return {
        counter: prevState
      }
    });
  }
  render() {
    const expression = null; // Change this line
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        {/* Change code below this line */}

        {/* Change code above this line */}
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```

# --solutions--

```jsx
// We want this to be deterministic for testing purposes.
const randomSequence = [true, false, false, true, true, false, false, true, true, false];
let index = 0;
const fiftyFifty = () => randomSequence[index++ % randomSequence.length];

class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h1>{this.props.fiftyFifty ? 'You Win!' : 'You Lose!'}</h1>;
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      return {
        counter: prevState.counter + 1
      }
    });
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        <Results fiftyFifty={fiftyFifty()} />
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```
