---
id: 5a24c314108439a4d4036177
title: シンプルなカウンターを書く
challengeType: 6
forumTopicId: 301425
dashedName: write-a-simple-counter
---

# --description--

これまでに学んだ概念を組み合わせて、より複雑なステートフルコンポーネントを設計できます。これには、`state`の初期化、`state`を設定するメソッドの作成、そしてこれらのメソッドを起動するクリックハンドラの代入が含まれます。

# --instructions--

`Counter`コンポーネントは`count`に`state`の値を保持しています。2つのボタンがあり、それぞれ`increment()`と`decrement()`というメソッドを呼び出します。これらのメソッドを作成して、適切なボタンがクリックされたときにカウンターの値が1ずつ増減するようにしてください。また、リセットボタンがクリックされたときにカウントを0に設定する`reset()`メソッドも作成してください。

**注意:** ボタンの`className`は変更しないでください。また、新しく作成したメソッドにはコンストラクタで必要なバインディングを追加することを忘れないでください。

# --hints--

`Counter`は、テキスト内容がこの順番である3つのボタンを含む`div`要素を返すべきです：`Increment!`、`Decrement!`、`Reset`。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Counter));
    return (
      mockedComponent.find('.inc').text() === 'Increment!' &&
      mockedComponent.find('.dec').text() === 'Decrement!' &&
      mockedComponent.find('.reset').text() === 'Reset'
    );
  })()
);
```

`Counter`のステートは、`count`プロパティを`0`に設定して初期化するべきです。

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
assert(mockedComponent.find('h1').text() === 'Current Count: 0');
```

増加ボタンをクリックすると、カウントが`1`ずつ増加します。

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.inc').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: 1');
```

減少ボタンをクリックすると、カウントが`1`ずつ減少します。

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.dec').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: -1');
```

リセットボタンをクリックすると、カウントが`0`にリセットされます。

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.setState({ count: 5 });
const currentCountElement = mockedComponent.find('h1');
assert(currentCountElement.text() === 'Current Count: 5');
mockedComponent.find('.reset').simulate('click');
assert(currentCountElement.text() === 'Current Count: 0');
```

# --seed--

## --seed-contents--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
  this.increment = this.increment.bind(this);
 this.decrement = this.decrement.bind(this);
 this.reset = this.reset.bind(this);
 }
  reset() {
    this.setState({
      count: 0
    });
  }
  increment() {
    this.setState(state => ({
      count: state.count + 1
    }));
  }
  decrement() {
    this.setState(state => ({
      count: state.count - 1
    }));
  }
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```
