---
id: 5a24c314108439a4d403616a
title: 配列をプロップスとして渡す
challengeType: 6
forumTopicId: 301401
dashedName: pass-an-array-as-props
---

# --description--

前回のチャレンジでは、親コンポーネントから子コンポーネントへ情報を `props` またはプロパティとして渡す方法を示しました。このチャレンジでは、配列を `props` として渡す方法を見ていきます。配列をJSX要素に渡すには、JavaScriptとして扱い、波括弧で囲む必要があります。

```jsx
<ParentComponent>
  <ChildComponent colors={["green", "blue", "red"]} />
</ParentComponent>
```

子コンポーネントはその後、配列のプロパティ `colors` にアクセスできます。プロパティにアクセスするときに `join()` のような配列メソッドを使うことができます。

```jsx
const ChildComponent = (props) => <p>{props.colors.join(', ')}</p>
```

これにより、すべての `colors` 配列の項目がカンマ区切りの文字列に結合され、`<p>green, blue, red</p>` が生成されます。後で、Reactで配列のデータをレンダリングする他の一般的なメソッドについて学びます。

# --instructions--

コードエディターには `List` と `ToDo` のコンポーネントがあります。`List` コンポーネントから各 `ToDo` をレンダリングするときに、例えば `tasks` のようなやることタスクの配列を代入した `["walk dog", "workout"]` プロパティを渡してください。次に、`tasks` コンポーネントでこの `List` 配列にアクセスし、その値を `p` 要素内に表示してください。`join(", ")` を使って `props.tasks` 配列を `p` 要素内にカンマ区切りのリストとして表示してください。今日のリストには少なくとも2つのタスクがあり、明日のリストには少なくとも3つのタスクがあるべきです。

# --hints--

`ToDo` コンポーネントは単一の外側の `div` を返すべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().type() === 'div';
  })()
);
```

`ToDo` コンポーネントの3番目の子は `List` コンポーネントのインスタンスであるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(2).name() === 'List';
  })()
);
```

`ToDo` コンポーネントの5番目の子は `List` コンポーネントのインスタンスであるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(4).name() === 'List';
  })()
);
```

`List` コンポーネントの両方のインスタンスは `tasks` という名前のプロパティを持ち、`tasks` は配列の型であるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      Array.isArray(mockedComponent.find('List').get(0).props.tasks) &&
      Array.isArray(mockedComponent.find('List').get(1).props.tasks)
    );
  })()
);
```

今日のタスクを表す最初の `List` コンポーネントは2つ以上の項目を持つべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(0).props.tasks.length >= 2;
  })()
);
```

明日のタスクを表す2番目の `List` コンポーネントは3つ以上の項目を持つべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(1).props.tasks.length >= 3;
  })()
);
```

`List` コンポーネントは `tasks` プロップの値を `p` タグ内にレンダリングするべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      mockedComponent
        .find('p')
        .get(0)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(0)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',') &&
      mockedComponent
        .find('p')
        .get(1)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(1)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',')
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const List = (props) => {
  { /* Change code below this line */ }
  return <p>{}</p>
  { /* Change code above this line */ }
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        { /* Change code below this line */ }
        <List/>
        <h2>Tomorrow</h2>
        <List/>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const List= (props) => {
  return <p>{props.tasks.join(', ')}</p>
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        <List tasks={['study', 'exercise']} />
        <h2>Tomorrow</h2>
        <List tasks={['call Sam', 'grocery shopping', 'order tickets']} />
      </div>
    );
  }
};
```
