---
id: 5a24c314108439a4d4036172
title: ユーザーインターフェースでの状態のレンダリングの別の方法
challengeType: 6
forumTopicId: 301408
dashedName: render-state-in-the-user-interface-another-way
---

# --description--

コンポーネント内で`state`にアクセスする別の方法があります。`render()`メソッドの`return`文の前に、JavaScriptを直接書くことができます。たとえば、関数を宣言したり、`state`や`props`からデータにアクセスしたり、そのデータに対して計算を行ったりできます。そして、任意のデータを変数に代入し、`return`文でアクセスできます。

# --instructions--

`MyComponent`のrenderメソッド内で、`const`を`name`という名前で定義し、コンポーネントの`state`内のnameの値を代入してください。このコードの部分ではJavaScriptを直接書けるため、この参照を波括弧で囲む必要はありません。

次に、return文の中で、この値を`h1`タグ内に変数`name`を使ってレンダリングしてください。return文ではJSX構文（JavaScript用の波括弧）を使う必要があることを忘れないでください。

# --hints--

`MyComponent`は、状態に`name`というキーと`freeCodeCamp`という値を持っているはずです。

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

`MyComponent`は、単一の`h1`で囲まれた`div`の見出し要素をレンダリングするはずです。

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

レンダリングされる`h1`タグは、`{name}`への参照を含むべきです。

```js
assert(/<h1>\n*\s*\{\s*name\s*\}\s*\n*<\/h1>/.test(code));
```

レンダリングされる`h1`の見出し要素は、コンポーネントの状態からレンダリングされたテキストを含むべきです。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(firstValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line

    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line
    const name = this.state.name;
    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
