---
id: 5a24c314108439a4d4036179
title: 制御されたフォームを作成する
challengeType: 6
forumTopicId: 301384
dashedName: create-a-controlled-form
---

# --description--

前回のチャレンジでは、Reactが`input`や`textarea`のような特定の要素の内部状態を制御できることを示しました。これにより、それらは制御されたコンポーネントになります。この仕組みは、通常のHTMLの`form`要素を含む他のフォーム要素にも適用されます。

# --instructions--

`MyForm`コンポーネントは、空の`form`と提出ハンドラで設定されています。フォームが提出されると、提出ハンドラが呼び出されます。

フォームを提出するボタンを追加しました。このボタンには、フォームを制御するボタンであることを示す`type`が`submit`に設定されています。`input`内に`form`要素を追加し、前回のチャレンジのようにその`value`と`onChange()`属性を設定してください。その後、`handleSubmit`メソッドを完成させて、コンポーネントの状態プロパティ`submit`をローカルの`state`にある現在の入力値に設定してください。

**注意:** フォームの提出ハンドラ内で`event.preventDefault()`を呼び出して、デフォルトのフォーム提出動作（ウェブページの更新）を防ぐ必要があります。キャンパーの利便性のために、ここではデフォルト動作が無効化されており、更新によるチャレンジコードのリセットを防いでいます。

最後に、`h1`の後に`form`タグを作成し、コンポーネントの`submit`から`state`の値をレンダリングしてください。これでフォームに入力してボタンをクリック（またはEnterキーを押す）すると、入力内容がページに表示されるはずです。

# --hints--

`MyForm`は、`div`と`form`タグを含む`h1`要素を返すべきです。フォームには`input`と`button`が含まれている必要があります。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyForm));
    return (
      mockedComponent.find('div').children().find('form').length === 1 &&
      mockedComponent.find('div').children().find('h1').length === 1 &&
      mockedComponent.find('form').children().find('input').length === 1 &&
      mockedComponent.find('form').children().find('button').length === 1
    );
  })()
);
```

`MyForm`の状態は、両方とも空の文字列に設定された`input`と`submit`プロパティで初期化する必要があります。

```js
assert(
  Enzyme.mount(React.createElement(MyForm)).state('input') === '' &&
    Enzyme.mount(React.createElement(MyForm)).state('submit') === ''
);
```

`input`要素に入力すると、コンポーネントの状態の`input`プロパティが更新されるべきです。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return mockedComponent.state('input');
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return {
      state: mockedComponent.state('input'),
      inputVal: mockedComponent.find('input').props().value
    };
  };
  const before = _1();
  const after = _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
})();
```

フォームを提出すると、`handleSubmit`が実行され、状態の`submit`プロパティが現在の入力値に設定されるべきです。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'SubmitInput' } });
    return mockedComponent.state('submit');
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.state('submit');
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'SubmitInput');
})();
```

`handleSubmit`は`event.preventDefault`を呼び出すべきです。

```js
assert(
  __helpers.isCalledWithNoArgs(
    'event.preventDefault',
    MyForm.prototype.handleSubmit.toString()
  )
);
```

`h1`の見出し要素は、コンポーネントの状態の`submit`フィールドの値をレンダリングするべきです。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return mockedComponent.find('h1').text();
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.find('h1').text();
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'TestInput');
})();
```

# --seed--

## --seed-contents--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          {/* Change code below this line */}

          {/* Change code above this line */}
          <button type='submit'>Submit!</button>
        </form>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    event.preventDefault();
    this.setState(state => ({
      submit: state.input
    }));
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          <input value={this.state.input} onChange={this.handleChange} />
          <button type='submit'>Submit!</button>
        </form>
        <h1>{this.state.submit}</h1>
      </div>
    );
  }
}
```
