---
id: 5a24c314108439a4d4036164
title: コンポジションでコンポーネントを作成する
challengeType: 6
forumTopicId: 301383
dashedName: create-a-component-with-composition
---

# --description--

ここでは複数のReactコンポーネントをどのように組み合わせるかを見ていきます。アプリを作っていて、`Navbar`、`Dashboard`、`Footer`という3つのコンポーネントを作成したと想像してください。

これらのコンポーネントを組み合わせるには、`App`の*親*コンポーネントを作成し、この3つのコンポーネントを*子*としてレンダリングすればよいです。Reactコンポーネント内で子としてコンポーネントをレンダリングするには、JSX内にカスタムHTMLタグとしてコンポーネント名を書きます。例えば、`render`メソッド内で次のように書けます。

```jsx
return (
 <App>
  <Navbar />
  <Dashboard />
  <Footer />
 </App>
)
```

Reactが別のコンポーネントを参照するカスタムHTMLタグ（この例のように`< />`で囲まれたコンポーネント名）に出会うと、そのタグの場所にそのコンポーネントのマークアップをレンダリングします。これにより、`App`コンポーネントと`Navbar`、`Dashboard`、`Footer`の親子関係が示されます。

# --instructions--

コードエディターには`ChildComponent`というシンプルな関数コンポーネントと`ParentComponent`というクラスコンポーネントがあります。`ChildComponent`内で`ParentComponent`をレンダリングして、この2つを組み合わせてください。`ChildComponent`タグはスラッシュで閉じることを忘れないでください。

**注意:** `ChildComponent`はES6のアロー関数で定義されています。これはReactでよく使われる方法です。

# --hints--

Reactコンポーネントは単一の`div`要素を返す必要があります。

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.type() === 'div';
  })()
);
```

コンポーネントは2つのネストされた要素を返す必要があります。

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.children().length === 2;
  })()
);
```

コンポーネントは2番目の子として`ChildComponent`を返す必要があります。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ParentComponent));
    return (
      mockedComponent.find('ParentComponent').find('ChildComponent').length ===
      1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }


        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }
        <ChildComponent />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
