---
id: 5a24c314108439a4d403616c
title: デフォルトプロップスのオーバーライド
challengeType: 6
forumTopicId: 301399
dashedName: override-default-props
---

# --description--

デフォルトプロップスを設定する機能はReactで便利な機能です。デフォルトプロップスをオーバーライドする方法は、コンポーネントに対して明示的にプロップの値を設定することです。

# --instructions--

`ShoppingCart` コンポーネントは現在、子コンポーネント `Items` をレンダリングしています。この `Items` コンポーネントには、整数 `quantity` に設定されたデフォルトプロップ `0` があります。`10` に対して `quantity` の値を渡してデフォルトプロップをオーバーライドしてください。

**注意:** コンポーネントにプロップを追加する構文はHTMLの属性を追加する方法に似ています。ただし、`quantity` の値が整数なので、引用符で囲まずに波括弧で囲む必要があります。例えば `{100}` のようになります。この構文はJSXに対して、波括弧内の値を直接JavaScriptとして解釈するよう指示しています。

# --hints--

`ShoppingCart` コンポーネントがレンダリングされるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

`Items` コンポーネントがレンダリングされるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

`Items` コンポーネントには `{ quantity: 10 }` のプロップが `ShoppingCart` コンポーネントから渡されるべきです。

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
      return (
        mockedComponent.find('Items').props().quantity == 10 &&
        code.replace(/ /g, '')
          .includes('<Itemsquantity={10}/>')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items />
    { /* Change code above this line */ }
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items quantity = {10} />
    { /* Change code above this line */ }
  }
};
```
