---
id: 5a24c314108439a4d403617a
title: 子コンポーネントにステートをプロップスとして渡す
challengeType: 6
forumTopicId: 301403
dashedName: pass-state-as-props-to-child-components
---

# --description--

これまでのチャレンジで、子のJSX要素や子のReactコンポーネントにプロップスを渡す例をたくさん見てきました。これらのプロップスがどこから来るのか疑問に思うかもしれません。よくあるパターンは、アプリにとって重要な`state`を持つステートフルなコンポーネントがあり、それが子コンポーネントをレンダリングするというものです。これらの子コンポーネントには、その`state`の一部にアクセスできるようにし、それをプロップスとして渡します。

例えば、`App`コンポーネントが`Navbar`を含む他のコンポーネントをレンダリングしているとします。`App`には多くのユーザー情報を含む`state`がありますが、`Navbar`はユーザー名だけにアクセスして表示すればよい場合があります。その`state`の一部を`Navbar`コンポーネントにプロップとして渡します。

このパターンはReactの重要なパラダイムを示しています。まずは*単方向データフロー*です。ステートはアプリケーションのコンポーネントツリーの上位にあるステートフルな親コンポーネントから子コンポーネントへ一方向に流れます。子コンポーネントは必要なステートデータだけを受け取ります。次に、複雑なステートフルなアプリは、ほんの数個、あるいは単一のステートフルコンポーネントに分解できます。残りのコンポーネントは親からプロップスとしてステートを受け取り、そのステートからUIをレンダリングします。これにより、ステート管理はコードの一部で行い、UIレンダリングは別の部分で行うという分離が生まれます。このステートロジックとUIロジックの分離はReactの重要な原則の一つです。正しく使うことで、複雑なステートフルアプリの設計がずっと管理しやすくなります。

# --instructions--

`MyApp`コンポーネントはステートフルで、子として`Navbar`コンポーネントをレンダリングします。`name`プロパティをその`state`に渡して子コンポーネントに渡し、`name`のレンダーメソッドの一部である`h1`タグ内に`Navbar`を表示してください。`name`は`Hello, my name is:`のテキストの後に表示されるべきです。

# --hints--

`MyApp`コンポーネントは内部に`Navbar`コンポーネントをレンダリングするはずです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return (
      mockedComponent.find('MyApp').length === 1 &&
      mockedComponent.find('Navbar').length === 1
    );
  })()
);
```

`Navbar`コンポーネントは`MyApp`ステートプロパティ`name`をプロップスとして受け取るべきです。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').props());
  };
  const navProps = await setState();
  assert(navProps.name === 'TestName');
```

`h1`要素は`Navbar`内で`name`プロップをレンダリングするべきです。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const navH1Before = mockedComponent.find('Navbar').find('h1').text();
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').find('h1').text());
  };
  const navH1After = await setState();
  assert(new RegExp('TestName').test(navH1After) && navH1After !== navH1Before);
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         {/* Change code below this line */}
         <Navbar />
         {/* Change code above this line */}
       </div>
    );
  }
};

class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      {/* Change code below this line */}
      <h1>Hello, my name is: </h1>
      {/* Change code above this line */}
    </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         <Navbar name={this.state.name}/>
       </div>
    );
  }
};
class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      <h1>Hello, my name is: {this.props.name}</h1>
    </div>
    );
  }
};
```
