---
id: 5a24c314108439a4d4036176
title: 要素を切り替えるために State を使う
challengeType: 6
forumTopicId: 301421
dashedName: use-state-to-toggle-an-element
---

# --description--

状態を更新するときに、前の状態を知る必要がある場合があります。しかし、状態の更新は非同期で行われる場合があり、これは React が複数の`setState()`呼び出しをまとめて一度の更新にすることがあるということです。つまり、次の値を計算するときに `this.state` や `this.props` の前の値に頼ることはできません。したがって、次のようなコードは使わないでください。

```jsx
this.setState({
  counter: this.state.counter + this.props.increment
});
```

代わりに、`setState` に状態とプロップスにアクセスできる関数を渡すべきです。`setState` に関数を使うことで、状態とプロップスの最新の値を扱っていることが保証されます。つまり、上記は次のように書き換えるべきです。

```jsx
this.setState((state, props) => ({
  counter: state.counter + props.increment
}));
```

また、`props` を使わずに `state` のみが必要な場合は、次のような形も使えます。

```jsx
this.setState(state => ({
  counter: state.counter + 1
}));
```

オブジェクトリテラルは丸括弧で囲む必要があることに注意してください。そうしないと JavaScript はそれをコードブロックとみなします。

# --instructions--

`MyComponent` には `visibility` プロパティがあり、`false` で初期化されています。レンダーメソッドは `visibility` の値が true の場合はあるビューを返し、false の場合は別のビューを返します。

現在、コンポーネントの `visibility` にある `state` プロパティを更新する方法がありません。この値は true と false の間で切り替わる必要があります。ボタンにはクリックハンドラがあり、`toggleVisibility()` というクラスメソッドを呼び出します。このメソッドを定義するために `setState` に関数を渡してください。メソッドが呼び出されると `state` の `visibility` が反対の値に切り替わるようにします。もし `visibility` が `false` なら、メソッドはそれを `true` に設定し、その逆も同様です。

最後に、ボタンをクリックして `state` に基づくコンポーネントの条件付きレンダリングを確認してください。

**ヒント:** `this` でメソッドに `constructor` キーワードをバインドするのを忘れないでください！

# --hints--

`MyComponent` は `div` 要素を返し、その中に `button` を含むべきです。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('div').find('button')
    .length,
  1
);
```

`MyComponent` の状態は `visibility` プロパティを `false` に設定して初期化するべきです。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).state('visibility'),
  false
);
```

ボタン要素をクリックすると、状態の `visibility` プロパティが `true` と `false` の間で切り替わり、`h1` 要素が条件付きでレンダリングされます。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ visibility: false });
    return mockedComponent.state('visibility');
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const third = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(!firstValue && secondValue && !thirdValue);
})();
```

`setState` には無名関数を渡すべきです。

```js
const paramRegex = '[a-zA-Z$_]\\w*(,[a-zA-Z$_]\\w*)?';
assert(
  new RegExp(
    'this\\.setState\\((function\\(' +
      paramRegex +
      '\\){|([a-zA-Z$_]\\w*|\\(' +
      paramRegex +
      '\\))=>)'
  ).test(__helpers.removeWhiteSpace(code))
);
```

`this` は `setState` の中で使ってはいけません。

```js
assert(!/this\.setState\([^}]*this/.test(code));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    this.toggleVisibility = this.toggleVisibility.bind(this);
  }
  toggleVisibility() {
    this.setState(state => ({
      visibility: !state.visibility
    }));
  }
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```
