---
id: 5a24c314108439a4d403618a
title: Array.map() を使って要素を動的にレンダリングする
challengeType: 6
forumTopicId: 301417
dashedName: use-array-map-to-dynamically-render-elements
---

# --description--

条件付きレンダリングは便利ですが、コンポーネントが不特定多数の要素をレンダリングする必要がある場合もあります。リアクティブプログラミングでは、ユーザーの操作に大きく依存するため、アプリケーションの状態が実行時までわからないことがよくあります。プログラマーはその未知の状態に対応できるようにコードを書いておく必要があります。Reactでの`Array.map()`の使用はこの考え方を示しています。

たとえば、シンプルな「To Doリスト」アプリを作成するとします。プログラマーとして、ユーザーがリストに何個の項目を持つかはわかりません。プログラムを使う人が「今日は洗濯の日だ」と決めるずっと前に、コンポーネントが正しい数のリスト要素を動的にレンダリングできるように設定する必要があります。

# --instructions--

コードエディターには`MyToDoList`コンポーネントの大部分がセットアップされています。このコードは、コントロールされたフォームのチャレンジを完了していれば見覚えがあるはずです。`textarea`と`button`があり、それぞれの状態を追跡するメソッドもありますが、まだページには何もレンダリングされていません。

`constructor`の中で`this.state`オブジェクトを作成し、2つの状態を定義してください。`userInput`は空の文字列で初期化し、`toDoList`は空の配列で初期化します。次に、`null`メソッド内の`render()`変数の隣にある`items`の値を削除してください。その代わりに、コンポーネントの内部状態に保存されている`toDoList`配列をmapで処理し、各項目に対して`li`を動的にレンダリングしてください。`eat, code, sleep, repeat`に文字列`textarea`を入力し、ボタンをクリックして何が起こるか試してみましょう。

**注意:** このようなマッピング操作で作成される兄弟要素にはすべて一意の`key`属性を付ける必要があることはご存知かもしれません。これは次のチャレンジのテーマなので心配しなくて大丈夫です。

# --hints--

MyToDoListコンポーネントが存在し、ページにレンダリングされている必要があります。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return mockedComponent.find('MyToDoList').length === 1;
  })()
);
```

`MyToDoList`の最初の子要素は`textarea`要素であるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(0).type() ===
      'textarea'
    );
  })()
);
```

`MyToDoList`の2番目の子要素は`br`要素であるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(1).type() === 'br'
    );
  })()
);
```

`MyToDoList`の3番目の子要素は`button`要素であるべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(2).type() ===
      'button'
    );
  })()
);
```

`MyToDoList`の状態は空の配列`toDoList`で初期化されているべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      Array.isArray(initialState.toDoList) === true &&
      initialState.toDoList.length === 0
    );
  })()
);
```

`MyToDoList`の状態は空の文字列`userInput`で初期化されているべきです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      typeof initialState.userInput === 'string' &&
      initialState.userInput.length === 0
    );
  })()
);
```

`Create List`ボタンがクリックされると、`MyToDoList`コンポーネントは`textarea`要素に入力されたカンマ区切りリストの各項目に対応するリストアイテム要素を含む番号なしリストを動的に返すべきです。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  const state_1 = () => {
    return mockedComponent.find('ul').find('li');
  };
  const setInput = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      'testA, testB, testC'
    );
  };
  const click = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_2 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const setInput_2 = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      't1, t2, t3, t4, t5, t6'
    );
  };
  const click_1 = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_3 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const awaited_state_1 = state_1();
  const awaited_setInput = setInput();
  const awaited_click = click();
  const awaited_state_2 = state_2();
  const awaited_setInput_2 = setInput_2();
  const awaited_click_1 = click_1();
  const awaited_state_3 = state_3();
  assert(
    awaited_state_1.length === 0 &&
      awaited_state_2.nodes.length === 3 &&
      awaited_state_3.nodes.length === 6 &&
      awaited_state_2.text === 'testAtestBtestC' &&
      awaited_state_3.text === 't1t2t3t4t5t6'
  );
})();
```

# --seed--

## --seed-contents--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = null; // Change this line
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```

# --solutions--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      toDoList: [],
      userInput: ''
    };
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = this.state.toDoList.map((item, i) => {
      return <li key={i}>{item}</li>;
    });
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```
