---
id: 5a24c314108439a4d4036161
title: "自己終了JSXタグについて学ぶ"
challengeType: 6
forumTopicId: 301396
dashedName: learn-about-self-closing-jsx-tags
---

# --description--

これまでに、JSXがHTMLと異なる重要な点として、HTMLのクラスを定義する際に`className`と`class`を使い分けることを見てきました。

JSXがHTMLと異なるもう一つの重要な点は、自己終了タグの考え方です。

HTMLでは、ほとんどのタグに開始タグと終了タグの両方があります：`<div></div>`。終了タグには、閉じるタグ名の前に必ずスラッシュが付きます。ただし、HTMLには<dfn>void elements</dfn>と呼ばれる特別な要素があり、これらは別の要素が始まる前に開始タグと終了タグの両方を必要としません。

例えば、改行タグは`<br>`または`<br />`と書けますが、内容を含まないため`<br></br>`と書くべきではありません。

JSXではルールが少し異なります。どのJSX要素も自己終了タグで書くことができ、すべての要素は閉じる必要があります。例えば改行タグは、有効なJSXとしてトランスパイル可能にするために常に`<br />`と書かなければなりません。一方、`<div>`は`<div />`または`<div></div>`と書けます。最初の構文では`<div />`に何も含めることができません。この構文は後のチャレンジでReactコンポーネントをレンダリングする際に役立つことがわかります。

# --instructions--

コードエディターのエラーを修正して、有効なJSXとなり正常にトランスパイルされるようにしてください。内容は変更せず、タグを閉じる必要がある箇所だけ閉じてください。

# --hints--

定数`JSX`は`div`要素を返すべきです。

```js
assert.strictEqual(JSX.type, 'div');
```

`div`は`br`タグを含むべきです。

```js
assert(Enzyme.shallow(JSX).find('br').length === 1);
```

`div`は`hr`タグを含むべきです。

```js
assert(Enzyme.shallow(JSX).find('hr').length === 1);
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h2>Welcome to React!</h2> <br >
    <p>Be sure to close all tags!</p>
    <hr >
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div>
  <h2>Welcome to React!</h2> <br />
  <p>Be sure to close all tags!</p>
  <hr />
</div>
);
```
