---
id: 5a24c314108439a4d4036173
title: this.setState で状態を設定する
challengeType: 6
forumTopicId: 301412
dashedName: set-state-with-this-setstate
---

# --description--

前のチャレンジではコンポーネントの `state` と `constructor` での状態の初期化方法を扱いました。コンポーネントの `state` を変更する方法もあります。React はコンポーネントの `state` を更新するための `setState` というメソッドを提供しています。コンポーネントクラス内で `setState` メソッドを次のように呼び出します：`this.setState()`。このとき、キーと値のペアを持つオブジェクトを渡します。キーは状態のプロパティで、値は更新された状態のデータです。例えば、状態に `username` を保存していてそれを更新したい場合は、次のようになります。

```jsx
this.setState({
  username: 'Lewis'
});
```

React では `state` を直接変更することはせず、状態が変わるときは常に `this.setState()` を使うことが期待されています。また、React はパフォーマンス向上のために複数の状態更新をまとめて処理する場合があります。つまり、`setState` メソッドによる状態更新は非同期になる場合があります。`setState` メソッドにはこの問題を回避するための別の構文もあります。これはあまり使うことはありませんが、覚えておくと良いでしょう。詳しくは<a href="https://www.freecodecamp.org/news/what-is-state-in-react-explained-with-examples/" target="_blank" rel="noopener noreferrer nofollow">Reactの記事</a>をご覧ください。

# --instructions--

コードエディタには `button` 要素があり、`onClick()` ハンドラが設定されています。このハンドラはブラウザで `button` がクリックイベントを受け取ると発動し、`handleClick` に定義された `MyComponent` メソッドを実行します。`handleClick` メソッド内で、`state` を使ってコンポーネントの `this.setState()` を更新してください。`name` の `state` プロパティを文字列 `React Rocks!` に設定してください。

ボタンをクリックして、レンダリングされた状態の更新を確認しましょう。クリックハンドラのコードの仕組みがまだ完全に理解できなくても心配いりません。これは今後のチャレンジで扱います。

# --hints--

`MyComponent` の状態はキーと値のペア `{ name: Initial State }` で初期化されているはずです。

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'Initial State'
);
```

`MyComponent` は `h1` 見出し要素をレンダリングする必要があります。

```js
assert(Enzyme.mount(React.createElement(MyComponent)).find('h1').length === 1);
```

レンダリングされた `h1` 見出し要素にはコンポーネントの状態からレンダリングされたテキストが含まれている必要があります。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(/<h1>TestName<\/h1>/.test(firstValue));
```

`handleClick` の `MyComponent` メソッドを呼び出すと、状態の name プロパティが `React Rocks!` に設定されるはずです。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'Before' });
    return waitForIt(() => mockedComponent.state('name'));
  };
  const second = () => {
    mockedComponent.instance().handleClick();
    return waitForIt(() => mockedComponent.state('name'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Before' && secondValue === 'React Rocks!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
     // Change code below this line
    this.setState({
      name: 'React Rocks!'
    });
    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```
