---
id: 5a24c314108439a4d403617b
title: コールバックをプロップスとして渡す
challengeType: 6
forumTopicId: 301400
dashedName: pass-a-callback-as-props
---

# --description--

`state` を子コンポーネントにプロップスとして渡せますが、渡せるのはデータだけではありません。ハンドラ関数やReactコンポーネントで定義された任意のメソッドも子コンポーネントに渡せます。これにより、子コンポーネントが親コンポーネントとやり取りできるようになります。メソッドは通常のプロップと同じように子に渡します。名前を付けて割り当て、子コンポーネント内の `this.props` でそのメソッド名にアクセスできます。

# --instructions--

コードエディタには3つのコンポーネントが示されています。`MyApp` コンポーネントは親で、`GetInput` と `RenderInput` の子コンポーネントをレンダリングします。`GetInput` のレンダーメソッドに `MyApp` コンポーネントを追加し、`input` の `inputValue` から `MyApp` に代入された `state` という名前のプロップを渡してください。また、`handleChange` というプロップを作成し、入力ハンドラ `handleChange` を渡してください。

次に、`RenderInput` のレンダーメソッドに `MyApp` を追加し、`input` から `inputValue` を渡す `state` というプロップを作成してください。完了すると、`input` コンポーネントの `GetInput` フィールドに入力できるようになり、そこから親のハンドラメソッドがプロップス経由で呼び出されます。これにより親の `state` の入力が更新され、両方の子にプロップスとして渡されます。コンポーネント間のデータの流れと、単一の真実の源が親コンポーネントの `state` にあることを観察してください。正直なところ、この例はやや作為的ですが、Reactコンポーネント間でデータとコールバックを渡す方法を示すのに役立ちます。

# --hints--

`MyApp` コンポーネントがレンダリングされるはずです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('MyApp').length === 1;
  })()
);
```

`GetInput` コンポーネントがレンダリングされるはずです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('GetInput').length === 1;
  })()
);
```

`RenderInput` コンポーネントがレンダリングされるはずです。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('RenderInput').length === 1;
  })()
);
```

`GetInput` コンポーネントは `MyApp` の状態プロパティ `inputValue` をプロップスとして受け取り、`input` 状態を変更する `MyApp` 要素を含むはずです。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: '' });
    return waitForIt(() => mockedComponent.state());
  };
  const state_2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => mockedComponent.state());
  };
  const updated_1 = await state_1();
  const updated_2 = await state_2();
  assert(updated_1.inputValue === '' && updated_2.inputValue === 'TestInput');
```

`RenderInput` コンポーネントは `MyApp` の状態プロパティ `inputValue` をプロップスとして受け取るはずです。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: 'TestName' });
    return waitForIt(() => mockedComponent);
  };
  const updated_1 = await state_1();
  assert(updated_1.find('p').text().includes('TestName'));
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
  this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
         <GetInput
           input={this.state.inputValue}
           handleChange={this.handleChange}/>
         <RenderInput
           input={this.state.inputValue}/>
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```
