---
id: 5a24c314108439a4d4036189
title: コンポーネントの状態に基づいてインラインCSSを条件付きで変更する
challengeType: 6
forumTopicId: 301380
dashedName: change-inline-css-conditionally-based-on-component-state
---

# --description--

ここまでで、条件付きレンダリングとインラインスタイルのいくつかの応用を見てきました。ここではこれら両方のトピックを組み合わせたもう一つの例を紹介します。Reactコンポーネントの状態に基づいてCSSを条件付きでレンダリングすることもできます。これを行うには、条件をチェックし、その条件が満たされた場合に、レンダー メソッド内のJSX要素に代入されたスタイルオブジェクトを変更します。

このパラダイムを理解することは重要です。なぜなら、これは従来のDOM要素を直接変更してスタイルを適用する方法（例えばjQueryで非常に一般的な方法）から大きく変わるからです。その方法では、要素が変わるタイミングを追跡し、実際の操作も直接処理する必要があります。変更を追跡するのが難しくなり、UIが予測不可能になる可能性があります。条件に基づいてスタイルオブジェクトを設定すると、UIがアプリケーションの状態の関数としてどのように見えるべきかを記述します。情報の流れが一方向に明確に流れるため、Reactでアプリケーションを書く際に推奨される方法です。

# --instructions--

コードエディターには、スタイル付きのボーダーを持つシンプルな制御された入力コンポーネントがあります。ユーザーが入力ボックスに15文字以上入力した場合、このボーダーを赤くスタイル付けしたいです。この条件をチェックし、条件が有効な場合は入力のボーダースタイルを`3px solid red`に設定してください。入力にテキストを入力して試すことができます。

# --hints--

`GateKeeper`コンポーネントは`div`要素をレンダリングする必要があります。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.find('div').length === 1;
  })()
);
```

`GateKeeper`コンポーネントは、空の文字列に設定された状態キー`input`で初期化する必要があります。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.state().input === '';
  })()
);
```

`GateKeeper`コンポーネントは`h3`タグと`input`タグをレンダリングする必要があります。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('h3').length === 1 &&
      mockedComponent.find('input').length === 1
    );
  })()
);
```

`input`タグは最初、`1px solid black`プロパティに対して`border`のスタイルを持つ必要があります。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('input').props().style.border === '1px solid black'
    );
  })()
);
```

`input`タグは、状態の入力値が15文字を超える場合に`3px solid red`のボーダーでスタイル付けする必要があります。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  let initialStyle = mockedComponent.find('input').props().style.border;
  const state_1 = () => {
    mockedComponent.setState({ input: 'this is 15 char' });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const state_2 = () => {
    mockedComponent.setState({
      input: 'A very long string longer than 15 characters.'
    });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const style_1 = await state_1();
  const style_2 = await state_2();
  assert(
    initialStyle === '1px solid black' &&
      style_1 === '1px solid black' &&
      style_2 === '3px solid red'
  );
```

# --seed--

## --seed-contents--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    // Change code below this line

    // Change code above this line
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    if (this.state.input.length > 15) {
      inputStyle.border = '3px solid red';
    };
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```
