---
id: 5a24c314108439a4d4036180
title: shouldComponentUpdate で再レンダリングを最適化する
challengeType: 6
forumTopicId: 301398
dashedName: optimize-re-renders-with-shouldcomponentupdate
---

# --description--

これまで、どのコンポーネントでも新しい `state` または新しい `props` を受け取ると、自身とすべての子コンポーネントが再レンダリングされていました。これは通常問題ありません。しかし、React には子コンポーネントが新しい `state` または `props` を受け取ったときに呼び出せるライフサイクルメソッドがあり、そのコンポーネントが更新すべきかどうかを明確に宣言できます。そのメソッドが `shouldComponentUpdate()` で、`nextProps` と `nextState` をパラメータとして受け取ります。

このメソッドはパフォーマンスを最適化するのに便利です。例えば、デフォルトの動作では、コンポーネントは新しい `props` を受け取ると `props` が変わっていなくても再レンダリングされます。`shouldComponentUpdate()` を使って `props` を比較し、これを防ぐことができます。このメソッドはコンポーネントを更新するかどうかを React に伝える `boolean` の値を返す必要があります。現在のプロップス（`this.props`）と次のプロップス（`nextProps`）を比較して、更新が必要かどうかを判断し、それに応じて `true` または `false` を返してください。

# --instructions--

`shouldComponentUpdate()` メソッドは `OnlyEvens` というコンポーネントに追加されています。現在、このメソッドは `true` を返しているため、`OnlyEvens` は新しい `props` を受け取るたびに再レンダリングされます。このメソッドを修正して、新しいプロップスの `OnlyEvens` が偶数の場合にのみ `value` が更新されるようにしてください。`Add` ボタンをクリックし、ライフサイクルフックがトリガーされる際のブラウザのコンソールでのイベントの順序を確認してください。

# --hints--

`Controller` コンポーネントは子として `OnlyEvens` コンポーネントをレンダリングする必要があります。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Controller));
    return (
      mockedComponent.find('Controller').length === 1 &&
      mockedComponent.find('OnlyEvens').length === 1
    );
  })()
);
```

`shouldComponentUpdate` メソッドは `OnlyEvens` コンポーネントに定義されている必要があります。

```js
assert(
  (() => {
    const child = React.createElement(OnlyEvens)
      .type.prototype.shouldComponentUpdate.toString()
      .replace(/s/g, '');
    return child !== 'undefined';
  })()
);
```

`OnlyEvens` コンポーネントは `h1` の値をレンダリングする `this.props.value` タグを返す必要があります。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 1000 });
    return mockedComponent.find('h1').html();
  };
  const second = () => {
    mockedComponent.setState({ value: 10 });
    return mockedComponent.find('h1').html();
  };
  const firstValue = first();
  const secondValue = second();
  assert(firstValue === '<h1>1000</h1>' && secondValue === '<h1>10</h1>');
})();
```

`OnlyEvens` は `nextProps.value` が偶数のときのみ再レンダリングされるべきです。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 8 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ value: 7 });
    return mockedComponent.find('h1').text();
  };
  const third = () => {
    mockedComponent.setState({ value: 42 });
    return mockedComponent.find('h1').text();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(firstValue === '8' && secondValue === '8' && thirdValue === '42');
})();
```

# --seed--

## --seed-contents--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return true;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```

# --solutions--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return nextProps.value % 2 === 0;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```
