---
id: 5a24c314108439a4d403617e
title: イベントリスナーを追加する
challengeType: 6
forumTopicId: 301377
dashedName: add-event-listeners
---

# --description--

`componentDidMount()` メソッドは、特定の機能のために追加する必要があるイベントリスナーを取り付ける最適な場所でもあります。React はブラウザに存在するネイティブのイベントシステムをラップする合成イベントシステムを提供しています。これは、ネイティブイベントが異なるブラウザ間で異なる動作をする場合でも、合成イベントシステムはユーザーのブラウザに関係なくまったく同じ動作をすることを意味します。

すでに `onClick()` のような合成イベントハンドラを使っています。React の合成イベントシステムは、DOM 要素で管理するほとんどのインタラクションにとても便利です。ただし、document や window オブジェクトにイベントハンドラを取り付けたい場合は、直接行う必要があります。

# --instructions--

`componentDidMount()` メソッドで `keydown` イベントのイベントリスナーを取り付け、これらのイベントがコールバック `handleKeyPress()` をトリガーするようにしてください。イベント（文字列）を最初の引数に、コールバックを2番目の引数に取る `document.addEventListener()` を使えます。

次に、`componentWillUnmount()` で同じイベントリスナーを削除してください。`document.removeEventListener()` に同じ引数を渡せます。React コンポーネントがアンマウントされて破棄される前にクリーンアップを行うために、このライフサイクルメソッドを使うのが良い習慣です。イベントリスナーの削除はそのようなクリーンアップの一例です。

# --hints--

`MyComponent` は `div` 要素をレンダリングし、その中に `h1` タグをラップする必要があります。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').children().find('h1').length === 1;
  })()
);
```

`keydown` リスナーは `componentDidMount` で document に取り付ける必要があります。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const didMountString = mockedComponent
      .instance()
      .componentDidMount.toString();
    return new RegExp(
      'document.addEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(didMountString);
  })()
);
```

`keydown` リスナーは `componentWillUnmount` で document から削除する必要があります。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const willUnmountString = mockedComponent
      .instance()
      .componentWillUnmount.toString();
    return new RegExp(
      'document.removeEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(willUnmountString);
  })()
);
```

コンポーネントがマウントされたら、`enter` を押すと状態が更新され、レンダリングされた `h1` タグも更新されます。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const beforeState = mockedComponent.state('message');
  const beforeText = mockedComponent.find('h1').text();
  const pressEnterKey = () => {
    mockedComponent.instance().handleKeyPress({ keyCode: 13 });
    return waitForIt(() => {
      mockedComponent.update();
      return {
        state: mockedComponent.state('message'),
        text: mockedComponent.find('h1').text()
      };
    });
  };
  const afterKeyPress = await pressEnterKey();
  assert(
    beforeState !== afterKeyPress.state && beforeText !== afterKeyPress.text
  );
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleEnter = this.handleEnter.bind(this);
    this.handleKeyPress = this.handleKeyPress.bind(this);
  }
  // Change code below this line
  componentDidMount() {

  }
  componentWillUnmount() {

  }
  // Change code above this line
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleKeyPress = this.handleKeyPress.bind(this);
    this.handleEnter = this.handleEnter.bind(this);  }
  componentDidMount() {
    // Change code below this line
    document.addEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  componentWillUnmount() {
    // Change code below this line
    document.removeEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```
