---
id: bad87fee1348bd9aec908857
title: "Використовуйте коментарі, щоб пояснити код"
challengeType: 0
forumTopicId: 18347
dashedName: use-comments-to-clarify-code
---

# --description--

Коли ми почнемо використовувати jQuery, ми змінюватимемо HTML-елементи без необхідності фактично змінювати їх у HTML.

Давайте переконаємося, що всі знають, що не слід безпосередньо змінювати цей код.

Пам’ятайте, що ви можете почати коментар з `<!--` і закінчити коментар `-->`

Додайте коментар у верхній частині вашого HTML, який каже `Code below this line should not be changed`

# --hints--

Ви повинні почати коментар з `<!--` у верхній частині вашого HTML.

```js
assert.match(code,(/^\s*<!--/));
```

Ваш коментар має містити текст `Code below this line should not be changed`.

```js
assert.match(code,(/<!--(?!(>|->|.*-->.*this line))\s*.*this line.*\s*-->/gi));
```

Ви повинні закрити свій коментар `-->`.

```js
assert.match(code,(/-->.*\n+.+/g));
```

У вас має бути однакова кількість відкривачів і закривачів коментарів.

```js
assert.equal(code.match(/<!--/g).length,code.match(/-->/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<!-- Code below this line should not be changed -->
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
