---
id: bad87fee1348bd9aec908846
title: створити заголовок bootstrap
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

Тепер давайте створимо щось з нуля, щоб попрактикуватися у HTML, CSS та Bootstrap.

Ми створимо майданчик для jQuery, який незабаром використаємо у наших завданнях з jQuery.

Для початку створіть елемент `h3` з текстом `jQuery Playground`.

Забарвіть ваш елемент `h3` за допомогою класу Bootstrap `text-primary` і вирівняйте його по центру за допомогою класу Bootstrap `text-center`.

# --hints--

Ви повинні додати елемент `h3` на вашу сторінку.

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

Ваш елемент `h3` повинен мати кінцевий тег.

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

Ваш елемент `h3` має бути зафарбований, застосувавши клас `text-primary`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

Ваш елемент `h3` має бути вирівняний по центру, застосувавши клас `text-center`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

Ваш елемент `h3` повинен містити текст `jQuery Playground`.

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```
