---
id: bad87fee1348bd9aec908848
title: Створіть Bootstrap wells
challengeType: 0
forumTopicId: 16825
dashedName: create-bootstrap-wells
---

# --description--

У Bootstrap є клас `well`, який може створити візуальне відчуття глибини для ваших колонок.

Вкладіть один елемент `div` з класом `well` всередину кожного з ваших вкладених елементів `col-xs-6` `div`.

# --hints--

Ви повинні додати елемент `div` з класом `well` всередину кожного з ваших елементів `div` з класом `col-xs-6`.

```js
const wells = document.querySelectorAll('div.col-xs-6 > div.well');
assert.lengthOf( wells,2 ); 
```

Обидва ваші елементи `div` з класом `col-xs-6` мають бути вкладені всередину вашого елемента `div` з класом `row`.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

Усі ваші елементи `div` повинні мати кінцеві теги.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g)?.length , code.match(/<div/g)?.length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">

    </div>
    <div class="col-xs-6">

    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
  </div>
</div>
```
