---
id: bad87fee1348bd9aec908847
title: Розбийте свій рядок Bootstrap
challengeType: 0
forumTopicId: 18306
dashedName: split-your-bootstrap-row
---

# --description--

Тепер, коли у нас є рядок Bootstrap, давайте розіб’ємо його на два стовпці, щоб розмістити наші елементи.

Створіть два `div` елементи всередині вашого рядка, обидва з класом `col-xs-6`.

# --hints--

Два `div class="col-xs-6"` елементи мають бути вкладені всередину вашого `div class="row"` елемента.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

Усі ваші `div` елементи повинні мати кінцеві теги.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length , code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">


  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6"></div>
    <div class="col-xs-6"></div>
  </div>
</div>
```
