---
id: bad87fee1348bd9aec908746
title: Розмістіть нашу сторінку всередині контейнера Bootstrap container-fluid
challengeType: 0
forumTopicId: 18198
dashedName: house-our-page-within-a-bootstrap-container-fluid-div
---

# --description--

Тепер переконаймося, що весь вміст на вашій сторінці адаптивний для мобільних пристроїв.

Вкладіть ваш `h3` елемент всередину `div` елемента з класом `container-fluid`.

# --hints--

Ваш `div` елемент має мати клас `container-fluid`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container-fluid'));
```

Кожен з ваших `div` елементів має мати кінцеві теги.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);

assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

Ваш `h3` елемент має бути вкладений всередину `div` елемента.

```js
const divElement = document.querySelector('div');
const divChildren = divElement?.querySelectorAll(`:scope ${'h3'}`)
assert.lengthOf(divChildren,1);
```

# --seed--

## --seed-contents--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```

# --solutions--

```html
<div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
</div>
```
