---
id: 587d7fb0367417b2b2512bef
title: Hudumia faili la HTML
challengeType: 2
forumTopicId: 301516
dashedName: serve-an-html-file
---

# --description--

Unaweza kujibu maombi kwa kutumia faili kwa njia ya `res.sendFile(path)`. Unaweza kuiweka ndani ya mshughulikiaji wa njia ya data `app.get('/', ...)`. Chini ya pazia, njia hii itaweka vichwa vinavyofaa ili kuagiza kivinjari chako jinsi ya kushughulikia faili unayotaka kutuma, kulingana na aina yake. Kisha itasoma na kutuma faili. Njia hii inahitaji njia kamili ya faili. Tunapendekeza utumie kigezo cha kimataifa cha Node `__dirname` kuhesabu njia kama hii:

```js
absolutePath = __dirname + '/relativePath/file.ext'
```

# --instructions--

Tuma faili la `/views/index.html` kama jibu kwa maombi ya GET kwenye njia ya data `/`. Ukitazama programu iliyo hewani, unapaswa kuona kichwa cha HTML kikubwa (na fomu ambayo tutaitumia baadaye…), bila mtindo wowote kuwekewa.

**Kumbuka:** Unaweza kuhariri suluhisho la changamoto iliyopita au kuunda suluhisho jipya. Ukitengeneza suluhisho jipya, kumbuka kuwa Express hupitia njia za data kutoka juu hadi chini, na hufanya mshughulikiaji kwa mechi ya kwanza. Lazima utoe maoni kwenye msimbo wa suluhisho lililotangulia, au seva itaendelea kujibu kwa mfuatano wa herufi. 

# --hints--

Programu yako inapaswa kuhudumia faili views/index.html

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /<h1>.*<\/h1>/,
    'Your app does not serve the expected HTML'
  );
```
