---
id: 587d7fb2367417b2b2512bf5
title: Pata Ingizo la Kigezo cha Njia ya Data kutoka kwa Mteja
challengeType: 2
forumTopicId: 301513
dashedName: get-route-parameter-input-from-the-client
---

# --description--

Unapojenga API, tunapaswa kuruhusu watumizi kuwasiliana nasi kuhusu wanachotaka kupata kutoka kwa huduma yetu. Kwa mfano, ikiwa mteja anaomba taarifa kuhusu mtumizi aliyehifadhiwa katika hifadhidata, wanahitaji njia ya kutufahamisha ni mtumizi gani wanayevutiwa naye. Njia moja inayowezekana kufanikisha hili ni kwa kutumia vigezo vya njia ya data. Vigezo vya njia ya data ni sehemu za URL zilizopewa majina, zikitenganishwa na alama za slash (/). Kila sehemu hukamata thamani ya sehemu ya URL inayolingana na nafasi yake. Thamani zilizokamatwa zinaweza kupatikana katika kitu cha `req.params`.

<blockquote>route_path: '/user/:userId/book/:bookId'<br>actual_request_URL: '/user/546/book/6754' <br>req.params: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

Jenga seva ya kurudia, iliyowekwa kwenye njia ya data `GET /:word/echo`. Jibu kwa kitu cha JSON, ukichukua muundo wa `{echo: word}`. Unaweza kupata neno la kurudiwa katika `req.params.word`. Unaweza kujaribu njia yako ya data kutoka kwenye upau wa anwani wa kivinjari chako, ukitembelea baadhi ya njia zinazolingana, mfano `your-app-rootpath/freecodecamp/echo`.

# --hints--

Test 1 : Seva yako ya kurudia inapaswa kurudia maneno kwa usahihi

```js
  const response = await fetch(code + '/eChOtEsT/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'eChOtEsT',
    'Test 1: the echo server is not working as expected'
  );
```

Test 2 : Seva yako ya kurudia inapaswa kurudia maneno kwa usahihi

```js
  const response = await fetch(code + '/ech0-t3st/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'ech0-t3st',
    'Test 2: the echo server is not working as expected'
  );
```
