---
id: 646c59e942f35541923104bf
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Поверніться до свого HTML і створіть елемент `main`. Всередині цього елемента `main` додайте елемент `div` з класом `cat-head`.

# --hints--

У вас має бути елемент `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

У вас має бути елемент `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

Ваш елемент `div` має мати клас `cat-head`.

```js
assert.equal(document.querySelector('div')?.getAttribute('class'), 'cat-head');
```

Ваш елемент `div` має бути всередині тегу `main`.

```js
assert.lengthOf(document.querySelectorAll('main div'), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
