---
id: 646c5e727e56e743c9aed4a1
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Позиціонування в CSS дозволяє вам встановити, як ви хочете розмістити елемент у браузері. Воно має властивість `position`, яку можна призначити `static`, `absolute`, `relative`, `sticky` або `fixed`.

Після того, як ви призначите властивість `position` елементу, ви зможете переміщувати елемент, встановлюючи піксельне або відсоткове значення для однієї або кількох властивостей `top`, `right`, `left` або `bottom`.

`static` — це позиціонування за замовчуванням для всіх елементів. Якщо ви призначите його елементу, ви не зможете переміщувати його за допомогою `top`, `right`, `left` або `bottom`.

Дайте `.cat-head` властивість `position` зі значенням `static`, потім встановіть властивості `top` і `left` зі значенням `100px` кожній.


# --hints--

Ваш селектор `.cat-head` має мати властивість `position`, встановлену в `static`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'static')
```

Ваш селектор `.cat-head` має мати властивість `top`, встановлену в `100px`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px')
```

Ваш селектор `.cat-head` має мати властивість `left`, встановлену в `100px`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
