---
id: 646c48df8674cf2b91020ecc
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Вам надано базовий шаблонний код. Почніть із додавання `link` елемента всередину вашого `head` елемента. Для цього `link` елемента встановіть атрибут `rel` зі значенням `stylesheet` та `href` зі значенням `./styles.css`.

# --hints--

У вас має бути `link` елемент.

```js
assert.exists(document.querySelector("head>link"));
```

Ваш `link` елемент має мати атрибути `rel` та `href`, встановлені відповідно на `stylesheet` і `./styles.css`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
