---
id: 646c5ace05e4be4211407935
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Щоб побачити елемент `.cat-head`, надайте йому фон із лінійним градієнтом з `#5e5e5e` на `85%` та `#45454f` на `100%`.

Ви можете не помітити різниці між цими двома кольорами, але вона є.

# --hints--

Ваш селектор `.cat-head` має мати властивість `background`.

```js
assert.match(code, /background:/)
```

Ваша властивість `background` має використовувати функцію `linear-gradient`.

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

Ваша функція `linear-gradient` має встановити перший колір як `#5e5e5e` на `85%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

Ваша функція `linear-gradient` має встановити другий колір як `#45454f` на `100%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
