---
id: 646dec359bef3b7811fba5a6
title: Крок 50
challengeType: 0
dashedName: step-50
---

# --description--

Перемістіть праве внутрішнє око на позицію з `position` рівним `absolute`, `top` рівним `8px` і `left` рівним `18px`. Також надайте йому радіус кордону `60%` і поверніть на `-5deg`.

# --hints--

Ваш селектор `.cat-right-inner-eye` має мати властивість `position`, встановлену в `absolute`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-eye')?.position, 'absolute');
```

Ваш селектор `.cat-right-inner-eye` має мати властивість `top`, встановлену в ` 8px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-eye')?.top, '8px');
```

Ваш селектор `.cat-right-inner-eye` має мати властивість `left`, встановлену в `18px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-eye')?.left, '18px');
```

Ваш селектор `.cat-right-inner-eye` має мати властивість `border-radius`, встановлену в `60%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-eye')?.borderRadius, '60%');
```

Ваш селектор `.cat-right-inner-eye` має мати властивість `transform`, встановлену в `-5deg`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-right-inner-eye')?.transform, 'rotate(-5deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>

        <div class="cat-eyes">
          <div class="cat-left-eye">
            <div class="cat-left-inner-eye"></div>
          </div>
          <div class="cat-right-eye">
            <div class="cat-right-inner-eye"></div>
          </div>
        </div>        
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-left-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}

.cat-right-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}

.cat-left-eye {
  position: absolute;
  top: 54px;
  left: 39px;
  border-radius: 60%;
  transform: rotate(25deg);
  width: 30px;
  height: 40px;
  background-color: #000;
}

.cat-right-eye {
  position: absolute;
  top: 54px;
  left: 134px;
  border-radius: 60%;
  transform: rotate(-25deg);
  width: 30px;
  height: 40px;
  background-color: #000;
}

.cat-left-inner-eye {
  position: absolute;
  top: 8px;
  left: 2px;
  width: 10px;
  height: 20px;
  transform: rotate(10deg);
  background-color: #fff;
  border-radius: 60%;
}

.cat-right-inner-eye {
  --fcc-editable-region--
  
  --fcc-editable-region--
  width: 10px;
  height: 20px;
  background-color: #fff;
}
```
