---
id: 6476fd4213318f6ee211028a
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Тепер замініть значення властивості `position` вашого `.cat-head` на `fixed`. Залиште `top` і `left` без змін.

Після цього прокрутіть сторінку вгору і вниз, щоб побачити, як працює значення `fixed`.

# --hints--

Ваш селектор `.cat-head` має мати властивість `position`, встановлену в `fixed`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'fixed');
```

Ваш селектор `.cat-head` має мати властивість `top`, встановлену в `100px`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

Ваш селектор `.cat-head` має мати властивість `left`, встановлену в `100px`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
