---
id: 6476f5c17f99146071ee884c
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Ви могли помітити, що нічого не відбувається. Елемент `.cat-head` не змістився, незважаючи на встановлення `top` і `left` зі значенням `100px` кожен. Але це не так із позиціюванням `relative`.

Коли ви використовуєте значення `relative`, елемент все одно позиціонується відповідно до нормального потоку документа, але стають активними значення `top`, `left`, `bottom` і `right`.

Замість `static` задайте вашому `.cat-head` позицію `relative`, а властивості `top` і `left` залиште без змін.

# --hints--

Ваш селектор `.cat-head` повинен мати властивість `position`, встановлену в `relative`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'relative');
```

Ваш селектор `.cat-head` повинен мати властивість `top`, встановлену в `100px`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

Ваш селектор `.cat-head` повинен мати властивість `left`, встановлену в `100px`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
