---
id: 64770351e8586671ec0911f0
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

`sticky` позиціювання — це гібрид `relative` та `fixed` позиціювання. Воно дозволяє елементу **закріпитися** на певній позиції всередині його контейнера або вікна перегляду, залежно від положення прокрутки.

Змініть значення властивості `position` у `.cat-head` на `sticky`, встановіть `top` у `0`, а потім видаліть `left` та його значення.

**Примітка**: Щоб побачити, як працює `sticky`, потрібно розмістити кілька текстів перед і після вашого елемента `.cat-head`. Якщо ви прокрутите вниз після цього, побачите, що `.cat-head` закріплюється зверху і залишається там.

# --hints--

Ваш селектор `.cat-head` має мати властивість `position`, встановлену в `sticky`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'sticky');
```

Ваш селектор `.cat-head` має мати властивість `top`, встановлену в `0`. Переконайтеся, що ви додали крапку з комою.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

У вашому коді не повинно бути властивості `left` та її значення.

```js
assert.notMatch(code, /left:\s*100px;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
