---
id: 647d855ac12cd436059acd39
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Використайте універсальний селектор, щоб додати `box-sizing: border-box;` до CSS. Це гарантує, що вказані значення ширини та висоти елементів містять відступ та кордон.

# --hints--

Ви повинні мати універсальний селектор (`*`).

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Універсальний селектор повинен мати властивість `box-sizing` зі значенням `border-box`.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
