---
id: 6676a8b01e56ec1a1e07c202
title: Крок 21
challengeType: 0
dashedName: step-21
---

# --description--

Зверніть увагу, що тепер у вас є білий квадрат із товстими кордонами. Кордони мають діагональні краї, які можна використати для створення трикутників. You will need them to create your cat's ears.

У тому самому селекторі класу налаштуйте `height` та `width` на `0`. Встановіть лівий, правий та верхній кордони на `transparent`.

Видаліть властивість `background-color`, і ви побачите трикутник.

# --hints--


Ви повинні мати селектор `.cat-right-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Селектор `.cat-right-ear` повинен мати властивість `height` зі значенням `0`. Не забудьте про крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '0px')
```

Селектор `.cat-right-ear` повинен мати властивість `width` зі значенням `0`. Не забудьте про крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '0px')
```

Селектор `.cat-right-ear` повинен мати властивість `border-left` зі значенням `35px solid transparent`. Не забудьте про крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid transparent')
```

Селектор `.cat-right-ear` повинен мати властивість `border-right` зі значенням `35px solid transparent`. Не забудьте про крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid transparent')
```

Селектор `.cat-right-ear` повинен мати властивість `border-top` зі значенням `35px solid transparent`. Не забудьте про крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid transparent')
```


Селектор `.cat-right-ear` повинен мати властивість `border-bottom` зі значенням `35px solid red`. Не забудьте про крапку з комою.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

Селектор `.cat-right-ear` не повинен мати властивості `background-color`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
