---
id: 694c84106e9023f289c03b84
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

로스터는 단 한 명의 우주 비행사 지휘관으로 시작합니다. 다음 속성을 가진 `firstAstronaut`라는 이름의 객체를 만드세요:

| Key | Value |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

`firstAstronaut` 변수를 만드세요.

```js
assert.isDefined(firstAstronaut);
```

`firstAstronaut`는 객체여야 합니다.

```js
assert.isObject(firstAstronaut);
```

`firstAstronaut` 객체에 올바른 속성 값이 포함되어야 합니다.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
