---
id: 694c9c186c089f7bdff0f3c6
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

일반적으로 함수 내부에서 입력을 검증하는 것이 좋은 습관입니다. `for` 반복문을 사용해 `crew` 배열을 순차적으로 반복하며 입력된 `id`과 동일한 `astronaut`가 이미 있는지 확인하세요. 중복이 발견되면 오류 메시지 `console.log("Duplicate ID: " + astronaut.id)`를 기록하고 `return`를 호출해 함수를 조기에 종료하세요.

다음은 이 기법을 자주 사용하는 예시입니다:

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

함수는 `for` 반복문을 사용해 기존 `crew` 멤버를 검사해야 합니다 (유효한 `for (...)` 형태면 모두 괜찮습니다).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

`for` 반복문 안에서 중복된 `id`이 있는지 확인하세요. 중복이 발견되면 `"Duplicate ID: " + astronaut.id"`을 콘솔에 기록하고 `return`를 호출해 함수를 종료하세요.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

유효한 입력이 주어지면 `addCrewMember` 함수는 `console.log()` 오류 메시지나 조기 `return` 종료를 호출하지 않아야 합니다.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
