---
id: 694ce8609fb805ba21cebc02
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

미션 컨트롤에서 승무원 명단에 있는 두 명의 위치를 바꾸는 메서드가 필요합니다. 세 개의 매개변수를 가진 `swapCrewMembers`라는 빈 함수를 만드세요.

- `crew`: 우주비행사 객체들의 배열

- `fromIndex`: 위치를 바꿀 첫 번째 우주비행사의 인덱스

- `toIndex`: 위치를 바꿀 두 번째 우주비행사의 인덱스

# --hints--

`swapCrewMembers`라는 이름의 함수를 만드세요.

```js
assert.isFunction(swapCrewMembers);
```

`swapCrewMembers` 함수는 `crew`, `fromIndex`, `toIndex` 매개변수를 가져야 합니다.

```js
const regex = __helpers.functionRegex('swapCrewMembers', ['crew', 'fromIndex', 'toIndex']);
assert.match(__helpers.removeJSComments(swapCrewMembers.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

--fcc-editable-region--

--fcc-editable-region--
```
