---
id: 694cabb9f1afa9446f36ffc0
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

이제 명단을 만들기 시작할 수 있습니다. `addCrewMember()`를 사용해 `firstAstronaut`을 `squad`에 추가하세요.

# --hints--

`addCrewMember()`를 `squad`와 `firstAstronaut`로 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /addCrewMember\s*\(\s*squad\s*,\s*firstAstronaut\s*\)/);
assert.strictEqual(squad.length, 1);
assert.strictEqual(squad[0], firstAstronaut);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
  console.log(`Added ${astronaut.name} as ${astronaut.role}`);
}

--fcc-editable-region--

--fcc-editable-region--
```
