---
id: 694c84330cbe604821b9deef
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

미션 컨트롤에서는 명단에 새 멤버를 추가하는 메서드가 필요합니다. `addCrewMember`라는 이름의 빈 함수를 만들고, 두 개의 매개변수 `crew`과 `astronaut`를 받도록 하세요.

# --hints--

`addCrewMember`라는 이름의 함수를 만들어야 합니다.

```js
assert.isFunction(addCrewMember);
```

`addCrewMember` 함수는 `crew`와 `astronaut` 매개변수를 가져야 합니다.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
