---
id: 694c40d86fe6a78d20b6559c
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 JavaScript를 활용해 우주 임무를 위한 우주비행사 명단을 만듭니다.

먼저, `squad`라는 이름의 빈 배열을 만드세요.

# --hints--

`squad` 변수를 만들어야 합니다.

```js
assert.isDefined(squad);
```

`squad` 변수는 배열이어야 합니다.

```js
assert.isArray(squad);
```

`squad` 배열은 비어 있어야 합니다.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
