---
id: 694ca14c8ef00056ad146e51
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

`astronaut` 객체를 `crew` 배열에 추가하고 `Added ${astronaut.name} as ${astronaut.role}`를 콘솔에 출력하세요. 로그 메시지에는 템플릿 리터럴이나 문자열 연결 방식을 사용할 수 있습니다.

# --hints--

`astronaut`를 `crew` 배열에 추가해야 합니다.

```js
const crew = [];
const astronauts = [
  { id: 1, name: "Abraham", role: "Commander", isEVAEligible: true, priority: 1 },
  { id: 2, name: "Boris", role: "Pilot", isEVAEligible: false, priority: 2 },
  { id: 3, name: "Catarina", role: "Engineer", isEVAEligible: true, priority: 3 }
];

astronauts.forEach(astro => addCrewMember(crew, astro));

assert.strictEqual(crew.length, astronauts.length);
astronauts.forEach((astro, idx) => {
  assert.deepEqual(crew[idx], astro);
});
```

우주비행사의 `name`를 다음 형식인 `Added {astronaut.name} as {astronaut.role}`로 로그에 출력해야 합니다.

```js
function captureConsoleLogs(fn) {
  const originalLog = console.log;
  const logs = [];

  console.log = function (...args) {
    logs.push(args.join(" "));
  };

  try {
    fn();
  } finally {
    console.log = originalLog;
  }

  return logs;
}

const crewForLogging = [];
const astronautsToLog = [
  { id: 1, name: "Abraham", role: "Commander" },
  { id: 2, name: "Boris", role: "Pilot" },
  { id: 3, name: "Catarina", role: "Engineer" }
];

const logs = captureConsoleLogs(() => {
  astronautsToLog.forEach(astro => addCrewMember(crewForLogging, astro));
});

assert.strictEqual(logs.length, astronautsToLog.length);

astronautsToLog.forEach((astro, idx) => {
  assert.match(
    logs[idx],
    new RegExp(`Added\\s*${astro.name}\\s*as\\s*${astro.role}`, 'i'),
  );
});
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
