---
id: 694d8ddd106340237efb0c80
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

`splice()` 메서드는 중간을 포함해 배열의 어느 위치에서든 요소를 추가하거나 제거하여 배열을 수정할 수 있습니다. `splice()`은 제거된 요소를 포함하는 배열을 반환하므로, 원본을 변경하지 않고 배열 내 두 요소를 교환하는 데 사용할 수 있습니다. 다음 기법을 사용하면 한 줄로 이 작업을 할 수 있습니다:

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

예시는 다음과 같습니다:

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

세 번째 줄에 적용된 이 기법은 다음과 같이 작동합니다:

- `splice(3, 1, copyArray[1])`는 `3` 인덱스에 있는 요소(`55`)를 제거합니다

- `copyArray[1]`(`97`)에서 가져온 요소를 `3` 인덱스에 삽입합니다

- `splice()`은 제거된 요소를 포함하는 배열 `[55]`를 반환합니다

- `[0]`은 그 배열에서 `55`을 추출합니다

- 그 값(`55`)을 `copyArray[1]`에 다시 할당하여 교환을 완료합니다

이 기법을 `updatedCrew` 배열에 사용해 `fromIndex`와 `toIndex`에 있는 우주비행사를 교환하세요.

# --hints--

`splice()` 메서드를 호출하여 교환을 수행해야 합니다.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

교환을 수행하려면 `splice()` 배열과 함께 `updatedCrew`을 호출해야 합니다. 한 줄 기법을 사용하세요.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

`swapCrewMembers` 함수는 원본 입력 `crew` 배열을 변경하지 않아야 합니다.

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
