---
id: 694e0c6e2cf8b9b45755f0f5
title: 17단계
challengeType: 1
dashedName: step-17
---

# --description--

먼저, `eligible`라는 빈 배열을 만드세요. 그런 다음 입력 `crew` 배열에 있는 모든 우주비행사를 순차적으로 반복하면서, 그들이 EVA-자격이 있다면(즉, 그들의 `isEVAEligible` 속성이 `true`인 경우) `eligible` 배열에 추가하세요. 반복이 끝나면 `eligible` 배열을 반환하세요.

# --hints--

`eligible`라는 배열을 만들어야 합니다.

```js
const funcStr = __helpers.removeJSComments(getEVAReadyCrew.toString());
assert.match(
  funcStr,
  /(var|let|const)\s+eligible/,
  "You must have at least one space between the declaration keyword (let/const) and 'eligible'"
);
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(getEVAReadyCrew.toString()));
assert.match(cleaned, /(?:^|[^a-zA-Z])(var|let|const)eligible=\[\]/);
```

입력 `for` 배열을 순차적으로 반복하려면 `crew` 반복문을 사용하세요.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(getEVAReadyCrew.toString()));
assert.match(cleaned, /for\s*\(/);
```

`eligible` 배열을 반환해야 합니다.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(getEVAReadyCrew.toString()));
assert.match(cleaned, /returneligible;?}$/);

const sample2 = [{ isEVAEligible: true }];
assert.isArray(getEVAReadyCrew(sample2));
```

`getEVAReadyCrew()`을 호출한 후 반환된 배열에는 입력 배열 `crew`에서 EVA-자격이 있는 우주비행사만 포함되어야 합니다.

```js
const sample = [
  { name: "A", isEVAEligible: true },
  { name: "B", isEVAEligible: false },
  { name: "C", isEVAEligible: true }
];
const result = getEVAReadyCrew(sample);
assert.deepEqual(result.map(a => a.name), ["A", "C"]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];

  return updatedCrew; 
}

const updatedSquad = swapCrewMembers(squad, 2, 5);

function getEVAReadyCrew(crew) {
--fcc-editable-region--

--fcc-editable-region--
}
```
