---
id: 694e1931f6509ae16aa1c36f
title: 18단계
challengeType: 1
dashedName: step-18
---

# --description--

미션 컨트롤에서 EVA 자격이 있는 우주비행사 목록을 우선순위 내림차순으로도 정렬해야 한다고 알렸습니다. 배열을 정렬하는 방법은 여러 가지가 있지만, 가장 기본적인 방법 중 하나는 `bubble sort`입니다.

`Bubble sort`은 목록을 반복해서 순회하며 인접한 항목을 비교하고, 순서가 잘못된 경우 서로 교환하는 방식으로 작동합니다. 각 반복이 끝날 때마다 정렬 기준에 따라 가장 먼저 와야 할 항목이 배열 내 올바른 위치에 점점 더 가까워지거나(“버블”처럼) 이동합니다. 다음은 `crew`을 사용해 `bubble sort`를 우선순위 내림차순으로 정렬하는 방법입니다:

```js
// Outer loop: controls how many passes we make
for (let i = 0; i < crew.length - 1; i++) {
  // Inner loop: compares neighboring items
  for (let j = 0; j < crew.length - 1 - i; j++) {
    // If current member has lower priority than next, swap
    if (crew[j].priority < crew[j + 1].priority) {
      // Using a temp variable for the swap
        const temp = crew[j];
        crew[j] = crew[j + 1];
        crew[j + 1] = temp;
    }
  }
}
```

`sortByPriorityDescending` 매개변수를 받는 `crew`라는 새 헬퍼 함수를 만드세요. 이 함수는 입력받은 `crew` 배열을 우선순위 내림차순으로 직접 정렬해야 하며, 두 개의 중첩된 `for` 반복문을 사용하고, 아무것도 반환하지 않아야 합니다.

# --hints--

`sortByPriorityDescending` 매개변수를 가진 `crew`라는 함수가 있어야 합니다.

```js
assert.isFunction(sortByPriorityDescending);
const regex = __helpers.functionRegex('sortByPriorityDescending', ['crew']);
const cleaned = __helpers.removeJSComments(sortByPriorityDescending.toString());
assert.match(cleaned, regex);
```

`sortByPriorityDescending` 함수는 정렬을 수행하기 위해 두 개의 중첩된 `for` 반복문을 사용해야 합니다.

```js
const loopMatches = sortByPriorityDescending.toString().match(/for\s*\(/g) || [];
assert.isAtLeast(loopMatches.length, 2);
```

`sortByPriorityDescending` 함수는 입력받은 `crew`을 `priority` 기준으로 내림차순 정렬해야 합니다.

```js
const crewSample = [
  { name: "A", priority: 2 },
  { name: "B", priority: 5 },
  { name: "C", priority: 3 }
];
sortByPriorityDescending(crewSample);
assert.deepEqual(crewSample.map(a => a.name), ["B", "C", "A"]);
```

`sortByPriorityDescending` 함수는 입력받은 `crew` 배열을 직접 제자리에서 정렬해야 하며, 새 배열을 반환하지 않아야 합니다.

```js
const crewSample2 = [
  { name: "X", priority: 1 },
  { name: "Y", priority: 2 }
];
const result = sortByPriorityDescending(crewSample2);
assert.strictEqual(result, undefined);
assert.deepEqual(crewSample2.map(a => a.priority), [2, 1]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];

  return updatedCrew; 
}

const updatedSquad = swapCrewMembers(squad, 2, 5);

--fcc-editable-region--

--fcc-editable-region--

function getEVAReadyCrew(crew) {
  const eligible = [];
  for (const astronaut of crew) {
    if (astronaut.isEVAEligible) eligible.push(astronaut);
  }
  
  return eligible;
}
```
