---
id: 694f8510e8a0a3d89e146a01
title: 24단계
challengeType: 1
dashedName: step-24
---

# --description--

`n`을 변경하지 않고 `array` 크기의 청크로 나누려면 `for` 반복문과 `slice()` 메서드를 다음과 같이 결합할 수 있습니다:

```js
const result = [];
for (let i = 0; i < array.length; i += n) {
  result.push(array.slice(i, i + n));
}
```

위 예제는:

- `result`라는 빈 배열을 만듭니다

- 원본 `array`를 `n` 크기 단위로 반복합니다

- `array` 크기의 `n` 청크를 만들고 `result`을 사용해 `slice()`에 추가합니다

`chunkCrew` 함수를 완성하려면 다음을 해야 합니다:

- `chunks`라는 배열을 만듭니다

- `crew` 크기 단위로 `size`를 반복하면서 `crew`에서 청크를 만들어 `chunks`을 사용해 `slice()`에 추가합니다

- `chunks` 배열을 반환합니다

# --hints--

`chunks`라는 빈 배열을 만들어야 합니다.

```js
const funcStr = __helpers.removeJSComments(chunkCrew.toString());
assert.match(
  funcStr,
  /(var|let|const)\s+chunks/,
  "You must have at least one space between the declaration keyword (let/const) and 'chunks'"
);
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(chunkCrew.toString()));
assert.match(cleaned, /(?:^|[^a-zA-Z])(var|let|const)chunks=\[\]/);
```

청크를 만들려면 `for` 반복문을 사용하세요.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(chunkCrew.toString()));
assert.match(cleaned, /for\s*\(/);
```

`for` 반복문 안에서 `push()` 배열에 `chunks`를 호출해 잘라낸 청크(이는 `slice()` 배열에 `crew`를 호출해 만들어야 합니다)를 삽입하세요.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(chunkCrew.toString()));
assert.match(cleaned, /chunks\.push\(crew\.slice\(/);
```

반복문이 끝난 후 `chunks` 배열을 반환해야 합니다.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(chunkCrew.toString()));
assert.match(cleaned, /returnchunks;?}$/);
```

`chunkCrew` 함수는 입력 배열 `size`를 변경하지 않고 `crew` 길이의 청크를 반환해야 합니다.

```js
const sample = [1, 2, 3, 4, 5];
const copyBefore = [...sample];
const chunks = chunkCrew(sample, 2);
assert.isArray(chunks);
assert.deepEqual(chunks, [[1, 2], [3, 4], [5]]);
assert.deepEqual(sample, copyBefore, 'chunkCrew should not mutate the original crew');
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];

  return updatedCrew; 
}

const updatedSquad = swapCrewMembers(squad, 2, 5);

function sortByPriorityDescending(crew) {
  for (let i = 0; i < crew.length - 1; i++) {
    for (let j = 0; j < crew.length - 1 - i; j++) {
      if (crew[j].priority < crew[j + 1].priority) {
        const temp = crew[j];
        crew[j] = crew[j + 1];
        crew[j + 1] = temp;
      }
    }
  }
}

function getEVAReadyCrew(crew) {
  const eligible = [];
  for (const astronaut of crew) {
    if (astronaut.isEVAEligible) eligible.push(astronaut);
  }
  sortByPriorityDescending(eligible); 

  return eligible;
}

const EVAReadySquad = getEVAReadyCrew(updatedSquad);

function chunkCrew(crew, size) {
  if (size < 1) {
    console.log("Chunk size must be >= 1");
    return;
  }
  
--fcc-editable-region--

--fcc-editable-region--
}
```
