---
id: 69504c249850e861921dca82
title: 26단계
challengeType: 1
dashedName: step-26
---

# --description--

`EVAChunks`가 본질적으로 배열 안에 배열이 들어있는, 즉 2차원 배열이라는 것을 눈치챘을 수 있습니다. 중첩된 `for` 반복문을 사용해 2차원 배열의 데이터를 출력할 수 있습니다:

```js
for (let i = 0; i < rootArray.length; i++) {
  console.log(`Group ${i + 1}:`);
  for (let j = 0; j < rootArray[i].length; j++) {
    console.log(rootArray[i][j].property);
  }
}
```

위 예시에서 `rootArray`는 주 배열이고, `rootArray[i]`는 주 배열 안의 하위 배열이며, `rootArray[i][j]`는 그 하위 배열 안의 한 객체이고 `.property`는 그 객체 안의 값을 접근합니다.

중첩된 `for` 반복문을 사용해 `EVAChunks`를 순차적으로 반복하며 각 청크에 있는 모든 우주비행사의 `name`를 출력하세요. 다음을 꼭 지키세요:

- `EVAChunks`를 루트 배열로 사용하세요
- 청크 사이에 `Chunk ${i+1}:`를 출력하세요

# --hints--

중첩된 `for` 반복문을 사용해 `EVAChunks`를 순차적으로 반복하며 각 청크에 있는 모든 우주비행사의 `name`를 출력하세요. 청크 사이에 `Chunk ${i+1}:`를 출력하세요.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
assert.match(
  cleaned,
  /for\([^)]*EVAChunks[^)]*\)\{[^}]*console\.log[^;]*Chunk\$\{[^}]+\}:/,
  'First for loop should iterate EVAChunks and log Chunk ${i+1}:'
);
assert.match(
  cleaned,
  /for\([^)]*EVAChunks[^)]*\)\{[\s\S]*?for\([^)]*EVAChunks[^)]*\)[\s\S]*?console\.log[^;]*EVAChunks\[[^\]]+\]\[[^\]]+\]\.name/,
  'Nested for loop should iterate inner EVAChunks elements using double indexing and log .name'
);

function captureConsoleFromCode(code) {
  const logs = [];
  const originalLog = console.log;
  console.log = (...args) => logs.push(args.join(" "));
  try {
    const userScriptFn = new Function(code);
    userScriptFn();
  } finally {
    console.log = originalLog;
  }
  return logs;
}
const logs = captureConsoleFromCode(code)
const expectedLogs = [
  "Chunk 1:", "Elise", "Felix", "Irene",
  "Chunk 2:", "Caroline", "Andy", "Hank"
];
assert.strictEqual(
  logs.length,
  expectedLogs.length,
  "You must log each chunk and each astronaut name exactly once"
);
assert.deepEqual(
  logs,
  expectedLogs,
  "Console output must log 'Chunk #' followed by the names of astronauts in each chunk in order"
);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];

  return updatedCrew; 
}

const updatedSquad = swapCrewMembers(squad, 2, 5);

function sortByPriorityDescending(crew) {
  for (let i = 0; i < crew.length - 1; i++) {
    for (let j = 0; j < crew.length - 1 - i; j++) {
      if (crew[j].priority < crew[j + 1].priority) {
        const temp = crew[j];
        crew[j] = crew[j + 1];
        crew[j + 1] = temp;
      }
    }
  }
}

function getEVAReadyCrew(crew) {
  const eligible = [];
  for (const astronaut of crew) {
    if (astronaut.isEVAEligible) eligible.push(astronaut);
  }
  sortByPriorityDescending(eligible); 

  return eligible;
}

const EVAReadySquad = getEVAReadyCrew(updatedSquad);
function chunkCrew(crew, size) {
  if (size < 1) {
    console.log("Chunk size must be >= 1");
    return;
  }

  const chunks = [];
  for (let i = 0; i < crew.length; i += size) {
    chunks.push(crew.slice(i, i + size));
  }

  return chunks;
}

const EVAChunks = chunkCrew(EVAReadySquad, 3);

--fcc-editable-region--

--fcc-editable-region--
```
