---
id: 612e77aba7ca691f598feb02
title: 步驟2
challengeType: 0
dashedName: step-2
---

# --description--

Add two `meta` tags, one to optimize your page for mobile devices, and one to specify an accepted `charset` for the page.

# --hints--

應該在頁面中添加兩個 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

一個 `meta` 元素應將 `name` 設置爲 `viewport`，並將 `content` 設置爲 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一個 `meta` 元素應該將 `charset` 屬性設置爲 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
