---
id: a77dbc43c33f39daa4429b4f
title: crea una funzione per verificare un booleano
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

In questo laboratorio, creerai una funzione che verifica se un valore è classificato come un primitivo booleano.

I primitivi booleani sono `true` e `false`.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `booWho` che riceve un argomento.
1. Se l'argomento ricevuto è un primitivo booleano, la funzione dovrebbe restituire `true`.
1. Se l'argomento è qualsiasi altro valore, la funzione dovrebbe restituire `false`.

# --hints--

Dovresti avere una funzione `booWho`.

```js
assert.isFunction(booWho);
```

`booWho(true)` dovrebbe restituire `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` dovrebbe restituire `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` dovrebbe restituire `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` dovrebbe restituire `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` dovrebbe restituire `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` dovrebbe restituire `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` dovrebbe restituire `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` dovrebbe restituire `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` dovrebbe restituire `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` dovrebbe restituire `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
