---
id: 6a89c69c111694edc6d1c628
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Jaribu kitendakazi chako cha `countup` kwa kumpa `5` ndani ya wito wa `print()`. Matokeo yanapaswa kuwa `[1, 2, 3, 4, 5]`.

Kwa hivyo, warsha ya kitendakazi cha countup imekamilika.

# --hints--

Unapaswa kufanya wito la `countup(5)` ndani ya kitendakazi cha `print()`.

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
