---
id: 6a89c69c8dbfda74151fa0b6
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda kitendakazi cha `countup` kinachotumia urudiaji kuunda orodha ya nambari kutoka `1` hadi nambari iliyotolewa.

Anza kwa kufafanua kitendakazi kinachoitwa `countup` chenye kigezo kimoja kinachoitwa `number`. Ongeza neno la maneno `pass` ndani ya mwili wa kitendakazi ili kitendakazi kiwe halali wakati bado kiko tupu.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `countup` chenye kigezo kinachoitwa `number`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

Kitendakazi chako cha `countup` kinapaswa kuwa na neno la maneno `pass`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
