---
id: 6a89c69c6af682856558d78b
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Wito wa kurudia hurudisha orodha inayojumuisha nambari ndogo ndogo. Tumia `append()` kuongeza `number` kwenye `count_list`, kisha rudisha `count_list`.

# --hints--

Unapaswa kuongeza `number` kwenye `count_list`.

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

Unapaswa kurudisha `count_list` kutoka kwenye kitendakazi chako cha `countup`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

Kitendakazi chako cha `countup` kinapaswa kurudisha `[1, 2, 3]` kinapoitwa na `3`.

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
