---
id: 6a89c69c91bc7365fbd4b733
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Ongeza kesi ya msingi katika kitendakazi chako cha `countup`. Badilisha `pass` na usemi wa `if` unaokagua kama `number` ni chini ya `1`. Ndani ya usemi wa `if`, rudisha orodha tupu.

# --hints--

Kitendakazi chako cha `countup` kinapaswa kurudisha orodha tupu wakati `number` ni chini ya `1`.

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

Kitendakazi chako cha `countup` hakipaswi kurudisha orodha tupu wakati `number` ni kubwa au sawa na `1`.

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
