---
id: 6a89c69c99298438d6fbe60b
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Sasa ongeza kesi ya kurudia. Baada ya usemi wa `if`, tengeneza kigezo kinachoitwa `count_list`. Weka thamani yake kuwa matokeo ya wito wa `countup` ukiwa na `number - 1`.

Kila wito hutumia nambari ndogo zaidi, hivyo kurudia kunasogea kuelekea kesi ya msingi.

# --hints--

Unapaswa kuweka thamani ya `countup(number - 1)` katika kigezo kinachoitwa `count_list`.

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
