---
id: 6842a6cd9836f0114a5b7a8a
title: 建置一個數字樣式產生器
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

在這個實驗中，你將透過建置一個小型應用程式來練習 Python 的基礎，該應用程式會產生一個數字樣式。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該定義一個名為 `number_pattern` 的函式（程式），它接受一個單一參數 `n`（表述一個正整數）。
1. `number_pattern` 應該使用 `for` 迴圈。
1. `number_pattern(n)` 應該傳回一個字串，該字串包含從 1 到 `n`（包含）所有整數，並以空格分隔。 例如，`number_pattern(4)` 應該傳回字串 `1 2 3 4`。
1. 如果傳遞給函式的引數不是整數值，函式應該傳回 `Argument must be an integer value.`。
1. 如果傳遞給函式的引數小於 1，該函式應該傳回 `Argument must be an integer greater than 0.`。

# --hints--

你應該有一個 `number_pattern` 函式（程式）。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

`number_pattern` 函式（程式）應該有一個名為 `n` 的參數。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)` 應該傳回 `1 2 3 4`。

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)` 應該傳回 `1 2 3 4 5 6 7 8 9 10 11 12`。

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern` 應該傳回以空格分隔的正整數列表。

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2
    
    `)
})
```

`number_pattern` 應該在傳入非整數值時傳回 `Argument must be an integer value.`。

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern` 應該在傳入非正整數時傳回 `Argument must be an integer greater than 0.`。

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
