---
id: 61fd5a93fd62bb35968adeab
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Налаштуйте свій HTML із `DOCTYPE`, `html`, що вказують, що цей документ англійською, `head` та `body` елементами.

Додайте вашому `head` елементу відповідні `meta` елементи для `charset` та `viewport`, `title` елемент із відповідним заголовком і `link` елемент для вашого стилю.

# --hints--

Ваш код має містити оголошення `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Ваш код має містити елемент `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Ваш `html` елемент має мати початковий тег із атрибутом `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш `html` елемент має мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Ваш код має містити елемент `head` всередині елемента `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Ваш код має містити елемент `body` всередині елемента `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Ваш `head` елемент має йти перед вашим `body` елементом.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

У вас має бути два `meta` елементи.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Один `meta` елемент має мати множину `name` зі значенням `viewport` та `content` зі значенням `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Інший `meta` елемент має мати атрибут `charset` зі значенням `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Ваш код має містити елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ваш `title` має містити текст.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Ваш код має містити елемент `link`.

```js
assert(/<link/.test(code))
```

Ваш `link` елемент має бути всередині вашого `head` елемента.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

Ваш `link` елемент має мати атрибут `rel` зі значенням `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

Ваш `link` елемент має мати атрибут `href` зі значенням `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
