---
id: 61fd71d596e8f253b9408b39
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Елемент `tr` використовується для позначення рядка таблиці. Додайте елемент `tr` всередину вашого елемента `thead`. У новому елементі `tr` додайте елемент `td`, а потім три елементи `th`.

Елемент `td` позначає комірку з даними, а елемент `th` позначає комірку заголовка.

# --hints--

Ваш елемент `thead` має містити елемент `tr`.

```js
assert(document.querySelector('thead')?.children?.[0]?.localName === 'tr');
```

Ваш елемент `tr` має мати елемент `td` як перший дочірній елемент.

```js
assert(document.querySelector('tr')?.children?.[0]?.localName === 'td');
```

Ваш елемент `tr` має містити три елементи `th` після елемента `td`.

```js
assert(document.querySelector('tr')?.children?.[1]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[2]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[3]?.localName === 'th');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
