---
id: 61fd6343fb42013d99bcd7f3
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Всередині вашого елемента `body` вкладіть елемент `section` всередину елемента `main`.

# --hints--

Першим елементом всередині елемента `body` має бути ваш елемент `main`.

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

У вашому елементі `section` має бути елемент `main`.

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
