---
id: 61fd67a656743144844941cb
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Читачі екрана оголошують HTML-елементи відповідно до потоку документа. Зрештою ми хочемо, щоб у балансовому звіті був заголовок «Balance Sheet» і підзаголовок «AcmeWidgetCorp». Однак такий порядок не має сенсу, якщо його оголошує читач екрана.

Додайте вашому існуючому `span` атрибут `class` зі значенням `flex` і додайте всередину два елементи `span`. Першому надайте текст `AcmeWidgetCorp`. Другому — текст `Balance Sheet`. Ви використаєте CSS, щоб змінити порядок тексту на сторінці, але порядок у HTML буде логічнішим для читача екрана.

# --hints--

Ваш існуючий елемент `span` має мати атрибут `class` зі значенням `flex`.

```js
assert(document.querySelector('h1')?.children?.[0]?.classList?.contains('flex'));
```

Ваш існуючий елемент `span` має містити два нові елементи `span`.

```js
assert(document.querySelector('.flex')?.children?.[0]?.localName === 'span');
assert(document.querySelector('.flex')?.children?.[1]?.localName === 'span');
```

Нові елементи `span` не повинні мати атрибуту `class`.

```js
assert(!document.querySelector('.flex')?.children?.[0]?.classList?.length);
assert(!document.querySelector('.flex')?.children?.[1]?.classList?.length);
```

Перший новий елемент `span` має містити текст `AcmeWidgetCorp`.

```js
assert(document.querySelector('.flex')?.children?.[0]?.textContent === 'AcmeWidgetCorp');
```

Другий новий елемент `span` має містити текст `Balance Sheet`.

```js
assert(document.querySelector('.flex')?.children?.[1]?.textContent === 'Balance Sheet');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
