---
id: 5f332a88dc25a0fd25c7687a
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

The name of the cafe is `CAMPER CAFE`. So, add an `h1` element within your `main` element. Give it the name of the cafe in capitalized letters to make it stand out.

# --hints--

`<h1>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<h1>/i));
```

終了タグ `</h1>` が 1 つ必要です。

```js
assert(code.match(/<\/h1>/i));
```

既存の `main` 要素を変更しないでください。

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

`h1` 要素は `main` 要素の中にネストされている必要があります。

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

`h1` 要素は `CAMPER CAFE` というテキストをもつ必要があります。

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

