---
id: 5f33294a6af5e9188dbdb8f3
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

It's time to add some menu content. Add a `main` element within the existing `body` element. It will eventually contain pricing information about coffee and desserts offered by the cafe.

# --hints--

`<main>` の開始タグが 1 つ必要です。

```js
assert.match(code, /<main>/i);
```

終了タグ `</main>` が 1 つ必要です。

```js
assert.match(code, /<\/main>/i);
```

`body` 要素を変更しないでください。 誤って終了タグを削除していないか確認してください。

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

`main` タグは `body` タグの中に置く必要があります。

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
