---
id: 5f344fc1520b6719f2e35605
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

There will be two sections on the menu, one for coffees and one for desserts. Add a `section` element within the `main` element so you have a place to put all the coffees available.

# --hints--

`<section>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<section\s*>/i));
```

`</section>` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/section\s*>/i));
```

既存の `main` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

`section` 要素は `main` 要素の中にある必要があります。

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
