---
id: 6998b99a8e195cd0c1f211df
title: 第一步
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将使用 JavaScript 对象构建一个简单的野生动物追踪器。

这是一个对象的示例：

```js
const animal = {}; // empty object
```

这会创建一个空对象。

现在创建一个名为 `tiger` 的变量，并赋值为空对象。

# --hints--

你应该创建一个名为 `tiger` 的变量。

```js
assert.exists(tiger);
```

`tiger` 变量应赋值为空对象。

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
