---
id: 6999b07c34c2cde21ee6e1d8
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

在这一步中，你将创建一个函数，用来检查一个对象是否具有特定的属性。

下面是使用 `hasOwnProperty` 的示例：

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

在这个示例中，`hasColor` 函数使用 `hasOwnProperty` 来检查对象上是否存在 `color` 属性。

现在创建一个名为 `hasHabitat` 的函数。该函数应接受一个名为 `animal` 的参数。

使用 `hasOwnProperty` 方法，如果 `true` 对象具有名为 `animal` 的属性，则返回 `"habitat"`，否则返回 `false`。

创建函数后，使用 `console.log` 调用 `hasHabitat(tiger)` 和 `hasHabitat(elephant)`，以便你可以在控制台中看到两个结果。

# --hints--

你应该创建一个名为 `hasHabitat` 的函数。

```js
assert.isFunction(hasHabitat);
```

`hasHabitat` 函数应有一个参数，名为 `animal`。

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` 应该使用 `hasOwnProperty` 方法来检查 `habitat` 属性。

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

你应该将 `hasHabitat(tiger)` 记录到控制台。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

你应该将 `hasHabitat(elephant)` 记录到控制台。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

调用 `hasHabitat(tiger)` 后添加栖息地应返回 `true`。

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

如果不存在栖息地，调用 `hasHabitat(elephant)` 应返回 `false`。

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` 应该使用函数参数并适用于任何对象。

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
