---
id: 69999ff7f7862fdfb885ea5d
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

现在通过添加一个名为 `tiger` 的新属性来更新你的 `age` 对象。

将 `age` 属性设置为值 `5`。

# --hints--

`tiger` 应该有一个 `age` 属性。

```js
assert.property(tiger, "age");
```

`tiger.age` 应该相等于 `5`。

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
