---
id: 6999b0d79ea938edcb032237
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

在《野生动物追踪器》工作坊的最后一步中，你将创建一个函数，使用方括弧表示法访问对象的属性。

下面是使用方括弧表示法访问属性的示例：

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

在这个示例中，方括弧表示法允许你使用变量访问属性。

现在创建一个名为 `getProperty` 的函数。

该函数应接受两个参数：`animal` 和 `propertyName`。

使用方括弧表示法返回该属性的值。

创建函数后，使用 `console.log` 调用 `getProperty(tiger, "species")` 和 `getProperty(elephant, "age")`，以便你可以在控制台看到返回的值。

至此，《野生动物追踪器》工作坊完成！

# --hints--

你应该创建一个名为 `getProperty` 的函数。

```js
assert.isFunction(getProperty);
```

`getProperty` 函数应有两个参数：`animal` 和 `propertyName`。

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty` 应该使用方括弧表示法访问属性。

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

你应该将 `getProperty(tiger, "species")` 记录到控制台。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

你应该将 `getProperty(elephant, "age")` 记录到控制台。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

调用 `getProperty(tiger, "species")` 应该返回 `"Tiger"`。

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

调用 `getProperty(elephant, "age")` 应该返回 `12`。

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty` 应该使用函数参数，并且能处理任何对象。

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
