---
id: 6999adc72c02812a28c01e31
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

在这一步，你将创建一个函数，用于更新对象的属性。

下面是一个在函数内部更新属性的示例：

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

在这个示例中，`age` 属性被更新为一个新值。

现在创建一个名为 `updateAge` 的函数。该函数应接受两个参数：`animal` 和 `newAge`。

在函数内部，将 `age` 对象的 `animal` 属性更新为 `newAge`。返回更新后的 `animal` 对象。

创建函数后，使用 `console.log` 调用 `updateAge(elephant, 12)`，这样你就可以在控制台看到更新后的 `elephant` 对象。

# --hints--

你应该创建一个名为 `updateAge` 的函数。

```js
assert.isFunction(updateAge);
```

`updateAge` 函数应有两个参数：`animal` 和 `newAge`。

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

`updateAge` 应该使用点符号来更新 `age` 属性。

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

`updateAge` 函数应返回更新后的 `animal` 对象。

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

你应该将 `updateAge(elephant, 12)` 输出到控制台。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

调用 `updateAge(elephant, 12)` 应该更新 age 属性。

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` 应该使用函数参数，并能处理任何对象。

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
