---
id: 6999ac171928136168654774
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

在这一步，你将创建一个函数来访问对象的 `species` 属性。

下面是一个从对象访问属性的小例子：

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

在这个例子中，函数接受一个对象作为参数并返回它的 `species` 属性。

现在创建一个名为 `getSpecies` 的函数。

该函数应接受一个名为 `animal` 的参数。

使用点符号返回 `species` 的 `animal` 属性。

创建函数后，使用 `console.log` 调用 `getSpecies(tiger)`，以便你可以在控制台看到结果。

# --hints--

你应该创建一个名为 `getSpecies` 的函数。

```js
assert.isFunction(getSpecies);
```

`getSpecies` 函数应有一个参数，名为 `animal`。

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

你应该将 `getSpecies(tiger)` 记录到控制台。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` 应该使用点符号访问 species 属性。

```js
assert.match(code, /animal\.species/);
```

调用 `getSpecies(tiger)` 应该返回 `"Tiger"`。

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` 应该使用函数参数并能处理任何对象。

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
