---
id: 6999faf886634b72be7a819b
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

对于这个野生动物追踪器，你将更新一个对象，添加新的属性。

这是一个创建带有属性的对象的示例：

```js
const animal = {
  name: "Lion" // property: value
};
```

现在更新 `tiger` 对象，使其包含一个名为 `species` 的属性。

将它的值设置为 `"Tiger"`。

# --hints--

`tiger` 应该是一个对象。

```js
assert.isObject(tiger);
```

`tiger` 应该有一个 `species` 属性。

```js
assert.property(tiger, "species");
```

`tiger.species` 应该相等于 `"Tiger"`。

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
