---
id: 6999b0198d18d402c36cc5fe
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

在这一步，你将创建一个函数，用于从对象中删除一个属性。

下面是使用 `delete` 删除属性的示例：

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

在这个示例中，`canFly` 属性从 `bird` 对象中被删除。

现在创建一个名为 `removeEndangeredStatus` 的函数。该函数应接受一个名为 `animal` 的参数。

在函数内部，使用 `isEndangered` 关键字从 `animal` 对象中删除 `delete` 属性。返回更新后的 `animal` 对象。

创建函数后，使用 `console.log` 调用 `removeEndangeredStatus(tiger)`，以便你可以在控制台中查看更新后的对象。

# --hints--

你应该创建一个名为 `removeEndangeredStatus` 的函数。

```js
assert.isFunction(removeEndangeredStatus);
```

`removeEndangeredStatus` 函数应有一个参数，名为 `animal`。

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`removeEndangeredStatus` 函数应返回更新后的 `animal` 对象。

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` 应该使用 `delete` 关键字删除 `isEndangered` 属性。

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

你应该将 `removeEndangeredStatus(tiger)` 记录到控制台。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

调用 `removeEndangeredStatus(tiger)` 应该删除 `isEndangered` 属性。

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` 应该使用函数参数，并能处理任何对象。

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
