---
id: 6999a0c1698b93d5d41156ed
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

现在为 `tiger` 对象添加另一个名为 `isEndangered` 的属性。

将 `isEndangered` 属性设置为 `true`。

# --hints--

`tiger` 应该有一个 `isEndangered` 属性。

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` 应该相等于 `true`。

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
