---
id: 699a00de9f564bb0effaa14d
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

在这一步，你将创建一个函数来访问一个对象的 `age` 属性。

这是一个示例：

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

在这个示例中，函数接受一个对象作为参数并返回它的 `age` 属性。

现在创建一个名为 `getAge` 的函数。

该函数应接受一个名为 `animal` 的参数。

使用点符号返回 `age` 的 `animal` 属性。

创建函数后，使用 `console.log` 调用 `getAge(tiger)`，以便你可以在控制台看到结果。

# --hints--

你应该创建一个名为 `getAge` 的函数。

```js
assert.isFunction(getAge);
```

`getAge` 函数应有一个参数，名为 `animal`。

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` 应该使用点符号访问 `age` 属性。

```js
assert.match(code, /animal\.age/);
```

你应该将 `getAge(tiger)` 记录到控制台。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

调用 `getAge(tiger)` 应该返回 `5`。

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` 应该使用函数参数并适用于任何对象。

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
