---
id: 69ba8da6b296391d913096a0
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

现在创建第二个名为 `elephant` 的对象。

添加以下属性：

- `species`，值为 `"Elephant"`
- `age`，值为 `10`
- `isEndangered`，值为 `true`

# --hints--

你应该创建一个名为 `elephant` 的变量。

```js
assert.exists(elephant);
```

`elephant` 应该有一个 `species` 属性，值等于 `"Elephant"`。

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` 应该有一个 `age` 属性，值等于 `10`。

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` 应该有一个 `isEndangered` 属性，值等于 `true`。

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
