---
id: 69a5f35669099ed52f8563b1
title: 构建智能储藏室补货程序
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

在本实验中，你将使用基本的 JavaScript 概念，如数组、对象、循环和条件式，构建一个小型储藏室管理程序。

你将模拟接收一批储藏室物品，决定如何处理每一项物品，并将结果组织起来以便存储。

**目标：** 完成以下用户故事并通过所有测试以完成实验。

`rawData` 数组包含以管道符分隔的字串，格式为 `sku|name|qty|expires|zone`，其中 `zone` 是可选的。

**用户故事：**

1. 你应该实现一个 `parseShipment(rawData)` 函数，该函数接受一个字串数组并返回一个具有 `{ sku, name, qty, expires, zone }` 属性的对象数组。

- 应忽略货物中重复的 `sku`。  
- 当货物字串中未提供 `zone` 时，应默认赋值为 `"general"`。  
- `qty` 的值应转换为数字。

2. 你应该实现一个 `planRestock(pantry, shipment)` 函数，该函数比较当前储藏室与即将到货的货物，并以 `{ type, item }` 形式返回一个动作数组，其中 `type` 是 `"restock"`、`"discard"` 或 `"donate"` 之一，`item` 是解析后的货物对象。

   `pantry` 参数是一个对象数组，形状与解析后的货物项相同（`{ sku, name, qty, expires, zone }`）。

- 如果货物物品的 `qty` 为 `0` 或更少，动作 `type` 应为 `"discard"`，无论该物品是否存在于储藏室。
   - 否则，如果货物物品的 `sku` 已存在于储藏室，动作 `type` 应为 `"restock"`。
   - 否则（货物物品的 `sku` 不存在于储藏室），动作 `type` 应为 `"donate"`。

3. 你应该实现一个 `groupByZone(actions)` 函数，根据每个物品的 `zone` 属性将动作分组到存储区域。该函数应返回一个对象，其中每个键是区域名称，值是属于该区域的动作数组。例如，如果动作包含区域 `"fridge"` 和 `"pantry"` 的物品，结果应为 `{ fridge: [...], pantry: [...] }`。

4. 你应该实现一个 `clonePantry(pantry)` 函数，返回储藏室的深拷贝，以便计划更改不会影响原始列表。深拷贝意味着创建一个包含新对象的新数组，因此修改副本不会改变原始储藏室。

5. 你应该将所有函数结合使用来处理一批货物，并将最终分组的结果对象记录到控制台。

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该定义一个名为 `parseShipment` 的函数，该函数接受一个名为 `rawData` 的字串数组作为参数。

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

你的 `parseShipment` 函数应将数组中的货物字串转换为具有属性 `sku`、`name`、`qty`、`expires` 和 `zone` 的对象。

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A10",
  name: "Tomatoes",
  qty: 5,
  expires: "2027-01-01",
  zone: "fridge"
});
```

应忽略货物中重复的 SKU。

``` js
const rawData = [
  "C32|Eggs|3|2027-01-01",
  "C32|Eggs|3|2027-01-01"
];

const result = parseShipment(rawData);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C32");
```

当 `zone` 在货物字串中不存在时，应将其赋值为 `general`。

``` js
const rawData = ["B21|Bananas|10|2027-01-01"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].zone, "general");
```

`qty` 的值应转换为数字。

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

你应该定义一个名为 `planRestock` 的函数，该函数接受两个参数：`pantry` 和 `shipment`。

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

你的 `planRestock` 函数应返回一个具有 `type` 和 `item` 属性的动作数组。

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 6, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

货物中数量为 `0` 或更少的物品应创建一个类型为 `discard` 的动作，无论该物品是否在储藏室中。

``` js
const pantry = [{ sku: "D43", name: "Pineapples", qty: 2, expires: "2027-01-01", zone: "general" }];
const shipment = [{ sku: "D43", name: "Pineapples", qty: 0, expires: "2027-01-01", zone: "general" },
                  { sku: "E54", name: "Peppers", qty: -1, expires: "2027-01-01", zone: "fridge" }
                ];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
assert.strictEqual(result[1].type, "discard");
```

当货物物品已存在于储藏室（相同 `sku`）时，动作应为 `restock`。

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

如果货物物品不存在于储藏室，动作应为 `"donate"`。

``` js
const pantry = [];
const shipment = [{ sku: "C32", name: "Eggs", qty: 3, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

你应该定义一个名为 `groupByZone` 的函数，该函数接受一个名为 `actions` 的参数。

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

你的 `groupByZone` 函数应返回按每个对象的 `zone` 属性分组的动作。

``` js

const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }, 
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

你的 `groupByZone` 函数应正确地按内容和数量分组动作。

``` js
const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "donate",  item: { sku: "B1", zone: "fridge" } },
  { type: "discard", item: { sku: "C1", zone: "pantry" } },
];

const result = groupByZone(actions);

assert.lengthOf(result.fridge, 2);
assert.lengthOf(result.pantry, 1);
assert.strictEqual(result.fridge[0].item.sku, "A1");
assert.strictEqual(result.pantry[0].type, "discard");
```

你应该定义一个名为 `clonePantry` 的函数，该函数接受一个名为 `pantry` 的参数。

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

`clonePantry` 应返回一个新数组，而不是原始储藏室数组。

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.isArray(pantryClone);
assert.notStrictEqual(pantryClone, pantry);
```

克隆的储藏室中的对象也应是新对象（深拷贝）。

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.notStrictEqual(pantryClone[0], pantry[0]);
assert.deepEqual(pantryClone[0], pantry[0]);

pantryClone[0].qty = 100;
assert.strictEqual(pantry[0].qty, 4);
```

这些函数应协同工作以处理货物并分组生成的动作。

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];

const shipment = parseShipment(rawData);
const pantryClone = clonePantry(pantry);
const actions = planRestock(pantryClone, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.strictEqual(grouped.fridge[0].item.qty, 5);
```

你应将按区域分组的结果动作记录到日志。

``` js
const groupedLog = getLogs().find(log => {
  if (!log || typeof log !== "object" || Array.isArray(log)) return false;

  return Object.values(log).some(zoneActions => {
    if (!Array.isArray(zoneActions)) return false;

    return zoneActions.some(action =>
      action &&
      ["restock", "discard", "donate"].includes(action.type) &&
      action.item &&
      typeof action.item.sku === "string"
    );
  });
});

assert.isOk(groupedLog, "Log the grouped actions object to the console.");
```

# --seed--

## --seed-contents--

```js
const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01",
  "B21|Bananas|10|2027-01-01",
  "C32|Eggs|3|2027-01-01|fridge",
  "C32|Eggs|3|2027-01-01",
  "D43|Pineapples|0|2027-01-01",
  "E54|Peppers|-1|2027-01-01|fridge"
];
```

# --solutions--

```js

function parseShipment(rawData) {

  // Early return if rawData is not an Array
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let line of rawData) {

    const parts = line.split("|");

    const sku = parts[0].trim();

    if (seenSKUs.includes(sku)) continue; // Avoid duplicates

    const name = parts[1].trim();
    const qty = Number(parts[2].trim()); // Parse qty into a number
    const expires = parts[3].trim();
    const zone = ( parts[4] ) ? parts[4].trim() : "general";

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires,
      zone: zone
    });

    seenSKUs.push(sku);

  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let item of shipment) {

    var exists = false;

    for (let stock of pantry)
    {
      if (stock.sku === item.sku)
      {
        exists = true;
      }
    }

    if (item.qty > 0) {
      if (exists) {
        actions.push({ type: "restock", item: item });
      } else{
        actions.push({ type: "donate", item: item });
      }
    } else {
      actions.push({ type: "discard", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {

  const byZone = {};

  for (const action of actions) {

    const zone = action.item.zone;

    if (!byZone[zone]) {
      byZone[zone] = [];
    }

    byZone[zone].push(action);
  }

  return byZone;

}

function clonePantry(pantry) {
  const pantryClone = [];

  for (let item of pantry) {
    pantryClone.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return pantryClone;
}

const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01", // Restock existing item
  "B21|Bananas|10|2027-01-01", // Donate new item without zone
  "C32|Eggs|3|2027-01-01|fridge", // Donate to a defined zone
  "C32|Eggs|3|2027-01-01", // Duplicated SKU in shipment
  "D43|Pineapples|0|2027-01-01", // Discard with quantity of 0
  "E54|Peppers|-1|2027-01-01|fridge" // Discard even if it's not in pantry
];

const shipment = parseShipment(rawData);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

// Show items expanded in zones
console.log(grouped);

```
