---
id: 69b279bae0f20eda63163cfe
title: الخطوة 20
challengeType: 1
dashedName: step-20
---

# --description--

قبل اختبار الدالة `mergeSort()`، يجب إنشاء حالة أساسية توقف تنفيذ الدالة عندما يكون طول الـ `array` أقل من أو يساوي `1`.

هذه الحالة الأساسية توقف استدعاء التكرار. بدونها، ستستمر عملية الفرز بالدمج في العمل حتى عندما تكون القائمة مرتبة أو لا تحتوي على عناصر.

بعد إعلان الدالة مباشرة، أنشئ جملة `if` تتحقق مما إذا كان `array.length` أقل من أو يساوي 1. داخل جسم جملة `if`، أضف جملة `return` لإيقاف تنفيذ الدالة.

# --hints--

يجب أن تنشئ جملة `if` بالشرط `array.length <= 1`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

يجب أن تحتوي جسم جملة `if` على جملة `return`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }
  
  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
