---
id: 69b2762b17d23fbee40a6c00
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

تُنفذ خوارزمية دمج الترتيب ثلاث عمليات:

- قسّم متتالية غير مرتبة من العناصر إلى أجزاء فرعية
- رتب العناصر في الأجزاء الفرعية
- دمج الأجزاء الفرعية المرتبة

تحدث العمليات السابقة بشكل متكرر حتى تُدمج الأجزاء الفرعية في المتتالية المرتبة الكاملة. ابدأ بتقسيم المتتالية.

أولًا، داخل دالة `mergeSort`، أعلن عن متغير `middlePoint` باستخدام `const` وعيّن له القيمة `Math.floor(array.length / 2)` لتتمكن من تقسيم المتتالية إلى نصفين قبل التكرار.

# --hints--

يجب إعلان متغير `middlePoint` وتعيين القيمة `Math.floor(array.length / 2)` له.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
