---
id: 69b279bae0f20eda63163cef
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

الآن بعد أن قسمت قائمة `array` إلى قائمتين منفصلتين، استمر في تقسيم كل قائمة حتى يصبح كل عنصر بمفرده في قائمته الخاصة. القائمة التي تحتوي على رقم واحد دائمًا ما تكون مرتبة.

للقيام بذلك، استدعِ دالة `mergeSort` بشكل متكرر داخل دالتك ومرر `leftPart` كمعلمة للاستدعاء.

# --hints--

يجب أن تستدعي `mergeSort(leftPart)` داخل نفس دالة `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
