---
id: 69b278519ef46b2fd8bf1546
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

لديك بالفعل قيمة `middlePoint`. استخدمها لتقسيم المصفوفة إلى نصفين واحتفظ بالنصف الأيسر جاهزًا للفرز.

استخدم صيغة القطع لاستخراج النصف الأيسر من `array` وعيّنه في متغير باسم `leftPart`.

# --hints--

يجب إعلان متغير `leftPart` داخل الدالة `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

يجب تعيين `array.slice(0, middlePoint)` إلى `leftPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
