---
id: 69b2713ff6c21fed2d3ad88b
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

فرز الدمج هو خوارزمية فرز تستخدم مبدأ القسمة والتغلب لفرز مجموعات البيانات. أي أنها "تقسم" المجموعة إلى أجزاء فرعية أصغر، و"تتغلب" على الأجزاء الفرعية بفرزها بشكل مستقل، ثم تدمج الأجزاء الفرعية المرتبة.

ابدأ بإنشاء دالة تسمى `mergeSort` تأخذ المعلمة `array`.

# --hints--

يجب إعلان دالة باسم `mergeSort` بمعلمة واحدة هي `array`.

```js
assert.isFunction(mergeSort);
const regex = __helpers.functionRegex('mergeSort', ['array']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
