---
id: 69b279bae0f20eda63163cf3
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

داخل الحلقة، ابدأ بمقارنة العناصر الحالية من كل نصف حتى تتمكن من دمج القيمة الأصغر أولاً.

أضف كتلة `if` تتحقق مما إذا كان العنصر الحالي في `leftPart` أقل من العنصر الحالي في `rightPart`.

# --hints--

يجب أن تتحقق شرطية `if` مما إذا كان `leftPart[leftArrayIndex]` أقل من `rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
