---
id: 69b2713ff6c21fed2d3ad88b
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

فرز الدمج هو خوارزمية فرز تستخدم مبدأ القسمة والتغلب لفرز مجموعات البيانات. أي أنها "تقسم" المجموعة إلى أجزاء فرعية أصغر، و"تتغلب" على الأجزاء الفرعية بفرزها بشكل مستقل، ثم تدمج الأجزاء الفرعية المرتبة.

ابدأ بإنشاء دالة تسمى `mergeSort` تأخذ المعلمة `array`.

# --hints--

يجب إعلان دالة باسم `mergeSort` بمعلمة واحدة هي `array`.

```js
assert.isFunction(mergeSort);

const explorer = await __helpers.Explorer(code);
const mergeSortFn = explorer.allFunctions.mergeSort;
assert.exists(mergeSortFn, 'You should declare a mergeSort function');

const params = mergeSortFn.parameters;
assert.lengthOf(
  params,
  1,
  'mergeSort should take a single parameter'
);
assert.isTrue(
  params[0].matches('array'),
  'The parameter of mergeSort should be named array'
);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
