---
id: 69b278519ef46b2fd8bf1546
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

لديك بالفعل قيمة `middlePoint`. استخدمها لتقسيم المصفوفة إلى نصفين واحتفظ بالنصف الأيسر جاهزًا للفرز.

استخدم صيغة القطع لاستخراج النصف الأيسر من `array` وعيّنه في متغير باسم `leftPart`.

# --hints--

يجب إعلان متغير `leftPart` داخل الدالة `mergeSort`.

```js
const explorer = await __helpers.Explorer(code);
const { leftPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftPart,
  'You should declare a leftPart variable inside mergeSort'
);
```

يجب تعيين `array.slice(0, middlePoint)` إلى `leftPart`.

```js
const explorer = await __helpers.Explorer(code);
const { leftPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftPart,
  'You should declare a leftPart variable inside mergeSort'
);
assert.isTrue(
  leftPart.value.matches('array.slice(0, middlePoint)'),
  'You should assign array.slice(0, middlePoint) to leftPart'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
