---
id: 69b279bae0f20eda63163cf1
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

حان الوقت الآن لفرز ودمج القائمتين (`leftPart` و `rightPart`) في المصفوفة الأصلية `array`.

يمكنك فعل ذلك بمقارنة العناصر في كلتا القائمتين، ودمج العنصر الأصغر في القائمة الرئيسية. ستقوم بهذه المقارنة لكل الفهارس في `leftPart` و `rightPart`.

أنشئ ثلاث متغيرات: `leftArrayIndex` و `rightArrayIndex` و `sortedIndex`، وضع قيمها على `0`. ستساعدك هذه المتغيرات على تتبع كل فهرس أثناء عملية الفرز.

# --hints--

يجب إعلان متغير `leftArrayIndex` مهيأ بقيمة `0`.

```js
const explorer = await __helpers.Explorer(code);
const { leftArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftArrayIndex,
  'You should declare a leftArrayIndex variable inside mergeSort'
);
assert.isTrue(
  leftArrayIndex.value.matches('0'),
  'leftArrayIndex should be initialized to 0'
);
```

يجب إعلان متغير `rightArrayIndex` مهيأ بقيمة `0`.

```js
const explorer = await __helpers.Explorer(code);
const { rightArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightArrayIndex,
  'You should declare a rightArrayIndex variable inside mergeSort'
);
assert.isTrue(
  rightArrayIndex.value.matches('0'),
  'rightArrayIndex should be initialized to 0'
);
```

يجب إعلان متغير `sortedIndex` مهيأ بقيمة `0`.

```js
const explorer = await __helpers.Explorer(code);
const { sortedIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  sortedIndex,
  'You should declare a sortedIndex variable inside mergeSort'
);
assert.isTrue(
  sortedIndex.value.matches('0'),
  'sortedIndex should be initialized to 0'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
