---
id: 69b279bae0f20eda63163cf6
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

في خطوة سابقة، عيّنت العنصر في `leftPart` إلى قائمة `array` لأنه كان أصغر. لكن هذا لن يكون الحال دائمًا. في بعض حالات المقارنة، قد يكون العنصر في الجانب الأيمن أصغر.

أنشئ بند `else` للتنفيذ عندما لا يكون العنصر في `leftPart` أقل من العنصر في `rightPart`.

داخل كتلة `else`، عيّن `rightPart[rightArrayIndex]` إلى `array[sortedIndex]`.

# --hints--

يجب أن تنشئ كتلة `else` تحتوي على `array[sortedIndex] = rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
