---
id: 69b2762b17d23fbee40a6c00
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

تُنفذ خوارزمية دمج الترتيب ثلاث عمليات:

- قسّم متتالية غير مرتبة من العناصر إلى أجزاء فرعية
- رتب العناصر في الأجزاء الفرعية
- دمج الأجزاء الفرعية المرتبة

تحدث العمليات السابقة بشكل متكرر حتى تدمج الأجزاء الفرعية في المتتالية المرتبة الكاملة. ابدأ بتقسيم المتتالية.

أولًا، داخل دالة `mergeSort`، أعلن عن متغير `middlePoint` باستخدام `const` وعيّن له القيمة `Math.floor(array.length / 2)` لتتمكن من تقسيم المتتالية إلى نصفين قبل التكرار.

# --hints--

يجب إعلان متغير `middlePoint` وتعيين القيمة `Math.floor(array.length / 2)` له.

```js
const explorer = await __helpers.Explorer(code);

const { middlePoint } = explorer.allFunctions.mergeSort.variables;

assert.exists(
  middlePoint,
  'middlePoint should be declared inside mergeSort'
);

assert.isTrue(
  middlePoint.matches('const middlePoint = Math.floor(array.length / 2);'),
  'middlePoint should be declared with const and assigned Math.floor(array.length / 2)'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
