---
id: 69b279bae0f20eda63163cf8
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

عبارات `if` و `else` التي أنشأتها في الخطوات السابقة ستُعيّن العناصر إلى المصفوفة المرتبة.

كل عنصر مُعيَّن إلى المصفوفة المرتبة يشغل فهرسًا في القائمة. لذلك يجب الانتقال إلى الفهرس التالي في المصفوفة المرتبة بعد كل تعيين.

أسفل كتلة `if`/`else`، ولكن لا تزال داخل حلقة `while`، زد قيمة `sortedIndex` بمقدار `1`.

# --hints--

يجب زيادة القيمة الحالية لـ `sortedIndex` بمقدار `1` بعد بند `else`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
