---
id: 69b279bae0f20eda63163cf5
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

بعد تعيين العنصر في `leftPart` إلى المصفوفة المرتبة، زد قيمة `leftArrayIndex` بمقدار `1`.

# --hints--

يجب إضافة `1` إلى القيمة الحالية لـ `leftArrayIndex` بعد إدخال العنصر في المصفوفة المرتبة.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftArrayIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
