---
id: 69b279bae0f20eda63163cf0
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

في أسفل جسم الدالة، استدعِ الدالة `mergeSort()` مرة أخرى. هذه المرة، مرر `rightPart` كمعلمة لاستدعاء الدالة.

# --hints--

يجب أن تستدعي `mergeSort()` مع المعلمة `rightPart` في أسفل جسم الدالة.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
