---
id: 69b278d82f0e514f32acb724
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

بعد تقطيع النصف الأيسر، لا يزال عليك الحصول على النصف الأيمن لإكمال خطوة الدمج.

استخدم صيغة التقطيع لاستخراج النصف الأيمن من `array` وعيّنه في متغير باسم `rightPart`.

# --hints--

يجب أن يكون لديك متغير `rightPart` داخل دالة `mergeSort`.

```js
const explorer = await __helpers.Explorer(code);
const { rightPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightPart,
  'You should declare a rightPart variable inside mergeSort'
);
```

يجب أن تستخدم صيغة التقطيع لاستخراج النصف الأيمن من المصفوفة وتعيينه إلى المتغير `rightPart`.

```js
const explorer = await __helpers.Explorer(code);
const { rightPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightPart,
  'You should declare a rightPart variable inside mergeSort'
);
assert.isTrue(
  rightPart.value.matches('array.slice(middlePoint)'),
  'You should assign array.slice(middlePoint) to rightPart'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
