---
id: 69b279bae0f20eda63163cf1
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

حان الوقت الآن لفرز ودمج القائمتين (`leftPart` و `rightPart`) في المصفوفة الأصلية `array`.

يمكنك فعل ذلك بمقارنة العناصر في كلتا القائمتين، ودمج العنصر الأصغر في القائمة الرئيسية. ستقوم بهذه المقارنة لكل الفهارس في `leftPart` و `rightPart`.

أنشئ ثلاث متغيرات: `leftArrayIndex` و `rightArrayIndex` و `sortedIndex`، وضع قيمها على `0`. ستساعدك هذه المتغيرات على تتبع كل فهرس أثناء عملية الفرز.

# --hints--

يجب إعلان متغير `leftArrayIndex` مهيأ بقيمة `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

يجب إعلان متغير `rightArrayIndex` مهيأ بقيمة `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

يجب إعلان متغير `sortedIndex` مهيأ بقيمة `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
