---
id: 69b279bae0f20eda63163cf2
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

بعد ذلك، عليك إنشاء حلقة تستمر طالما توجد عناصر متبقية في كل من `leftPart` و `rightPart`.

لهذا، أنشئ حلقة `while` تحتوي على شرطين: الأول يتحقق مما إذا كان `leftArrayIndex` أقل من طول `leftPart`، والآخر يتحقق مما إذا كان `rightArrayIndex` أقل من طول `rightPart`.

# --hints--

يجب أن تنشئ حلقة `while` تتحقق مما إذا كان `leftArrayIndex` أقل من طول `leftPart` وإذا كان `rightArrayIndex` أقل من طول `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
