---
id: 69b279bae0f20eda63163cf9
title: الخطوة 15
challengeType: 1
dashedName: step-15
---

# --description--

تُقارن حلقة `while` التي أنشأتها عنصرًا واحدًا من `leftPart` مع عنصر آخر في `rightPart`، ثم تضيف العنصر الأصغر إلى القائمة الرئيسية `array`.

ستستمر هذه العملية حتى لا يبقى عناصر للمقارنة. لكن قد تبقى عناصر في `leftPart` بينما لا يوجد في `rightPart`، والعكس صحيح.

أنشئ حلقة `while` أخرى تعمل عندما يكون `leftArrayIndex` أقل من طول `leftPart`. في الخطوات التالية، ستستخدمها لنسخ العناصر المتبقية في `leftPart` إلى قائمة `array`.

# --hints--

يجب أن تنشئ حلقة `while` بشرط `leftArrayIndex < leftPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
