---
id: 69b279bae0f20eda63163cf4
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

عندما تُقيَّم شرط `if` إلى `true`، فهذا يعني أن العنصر في قائمة `leftPart` أصغر من العنصر الذي يُقارن به في قائمة `rightPart`.

في هذه الحالة، يمكنك وضع العنصر الموجود عند `leftArrayIndex` في `leftPart` داخل المصفوفة المرتبة.

داخل كتلة `if`، عيّن `leftPart[leftArrayIndex]` إلى `array[sortedIndex]`.

# --hints--

يجب أن تعيّن `leftPart[leftArrayIndex]` إلى `array[sortedIndex]` داخل `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
