---
id: 6814d8e1516e86b171929de4
title: "Sfida 1: equilibrio delle vocali"
challengeType: 29
dashedName: challenge-1
---

# --description--

Data una stringa, determina se il numero di vocali nella prima metà della stringa è uguale al numero di vocali nella seconda metà.

- La stringa può contenere qualsiasi carattere.
- Le lettere `a`, `e`, `i`, `o` e `u`, sia in maiuscolo sia in minuscolo, sono considerate vocali.
- Se la stringa ha un numero dispari di caratteri, ignora il carattere centrale.

# --hints--

`is_balanced("racecar")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("racecar"), True)`)
}})
```

`is_balanced("Lorem Ipsum")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("Lorem Ipsum"), True)`)
}})
```

`is_balanced("Kitty Ipsum")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("Kitty Ipsum"), False)`)
}})
```

`is_balanced("string")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("string"), False)`)
}})
```

`is_balanced(" ")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced(" "), True)`)
}})
```

`is_balanced("abcdefghijklmnopqrstuvwxyz")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("abcdefghijklmnopqrstuvwxyz"), False)`)
}})
```

`is_balanced("123A#b!E&*456-o.U")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("123A#b!E&*456-o.U"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_balanced(s):

    return s
```

# --solutions--

```py
def is_balanced(s):
    vowels = set("aeiouAEIOU")
    n = len(s)
    half = n // 2

    first_half = s[:half]
    second_half = s[-half:]

    def count_vowels(sub):
        return sum(1 for char in sub if char in vowels)

    return count_vowels(first_half) == count_vowels(second_half)
```
