---
id: 681cb1b0dab50c87ddb2e519
title: "Sfida 8: calcolo del fattoriale"
challengeType: 29
dashedName: challenge-8
---

# --description--

Dato un numero intero da zero a 20, restituisci il fattoriale di quel numero. Il fattoriale di un numero è il prodotto di tutti i numeri tra 1 e il numero dato.

- Il fattoriale di zero è 1.

# --hints--

`factorial(0)` dovrebbe restituire `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(0), 1)`)
}})
```

`factorial(5)` dovrebbe restituire `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(5), 120)`)
}})
```

`factorial(20)` dovrebbe restituire `2432902008176640000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(20), 2432902008176640000)`)
}})
```

# --seed--

## --seed-contents--

```py
def factorial(n):

    return n
```

# --solutions--

```py
def factorial(n):
    result = 1
    for i in range(1, n + 1):
        result *= i
    return result
```
