---
id: 681cb1afdab50c87ddb2e515
title: "Sfida 4: S  P  A  Z  I  O  J  A  M"
challengeType: 29
dashedName: challenge-4
---

# --description--

Data una stringa, rimuovi tutti gli spazi dalla stringa, inserisci due spazi tra ogni carattere, trasforma tutte le lettere alfabetiche in maiuscolo e restituisci il risultato.

- I caratteri non alfabetici devono rimanere invariati (eccetto gli spazi).

# --hints--

`space_jam("freeCodeCamp")` dovrebbe restituire `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("   free   Code   Camp   ")` dovrebbe restituire `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("Hello World?!")` dovrebbe restituire `"H  E  L  L  O  W  O  R  L  D  ?  !"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !")`)
}})
```

`space_jam("C@t$ & D0g$")` dovrebbe restituire `"C  @  T  $  &  D  0  G  $"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $")`)
}})
```

`space_jam("allyourbase")` dovrebbe restituire `"A  L  L  Y  O  U  R  B  A  S  E"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E")`)
}})
```

# --seed--

## --seed-contents--

```py
def space_jam(s):

    return s
```

# --solutions--

```py
def space_jam(s):
    s = s.replace(" ", "")
    s = "  ".join(s)
    return s.upper()
```
