---
id: 681cb05adab50c87ddb2e513
title: "Sfida 2: verifica della base"
challengeType: 29
dashedName: challenge-2
---

# --description--

Data una stringa che rappresenta un numero e una base intera da 2 a 36, determina se il numero è valido in quella base.

- La stringa può contenere numeri interi e caratteri maiuscoli o minuscoli.
- Il controllo deve essere case-insensitive.
- La base può essere qualsiasi numero da 2 a 36.
- Un numero è valido se ogni carattere è una cifra valida nella base data.
- Esempi di cifre valide per le basi:
  - Base 2: 0-1
  - Base 8: 0-7
  - Base 10: 0-9
  - Base 16: 0-9 e A-F
  - Base 36: 0-9 e A-Z

# --hints--

`is_valid_number("10101", 2)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("10101", 2), True)`);
}})
```

`is_valid_number("10201", 2)` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("10201", 2), False)`)
}})
```

`is_valid_number("76543210", 8)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("76543210", 8), True)`)
}})
```

`is_valid_number("9876543210", 8)` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("9876543210", 8), False)`)
}})
```

`is_valid_number("9876543210", 10)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("9876543210", 10), True)`)
}})
```

`is_valid_number("ABC", 10)` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 10), False)`)
}})
```

`is_valid_number("ABC", 16)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 16), True)`)
}})
```

`is_valid_number("Z", 36)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("Z", 36), True)`)
}})
```

`is_valid_number("ABC", 20)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 20), True)`)
}})
```

`is_valid_number("4B4BA9", 16)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("4B4BA9", 16), True)`)
}})
```

`is_valid_number("5G3F8F", 16)` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("5G3F8F", 16), False)`)
}})
```

`is_valid_number("5G3F8F", 17)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("5G3F8F", 17), True)`)
}})
```

`is_valid_number("abc", 10)` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("abc", 10), False)`)
}})
```

`is_valid_number("abc", 16)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("abc", 16), True)`)
}})
```

`is_valid_number("AbC", 16)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("AbC", 16), True)`)
}})
```

`is_valid_number("z", 36)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("z", 36), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_number(n, base):

    return n
```

# --solutions--

```py
def is_valid_number(n, base):
    allChars = "0123456789abcdefghijklmnopqrstuvwxyz"
    newN = n.lower()

    availableChars = allChars[0:base]

    for char in newN:
        if char not in availableChars:
            return False

    return True
```
