---
id: 6821ebe4237de8297eaee794
title: "Sfida 18: il secondo migliore"
challengeType: 29
dashedName: challenge-18
---

# --description--

Dato un array di numeri interi che rappresentano il prezzo di diversi laptop e un numero intero che rappresenta il tuo budget, restituisci:

1. Il secondo laptop più costoso se rientra nel tuo budget, oppure
2. Il laptop più costoso che rientra nel tuo budget, oppure
3. `0` se nessun laptop rientra nel tuo budget.

- I prezzi duplicati devono essere ignorati.

# --hints--

`get_laptop_cost([1500, 2000, 1800, 1400], 1900)` dovrebbe restituire `1800`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900)` dovrebbe restituire `1800`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 2200)` dovrebbe restituire `1899`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 2200), 1899)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 1000)` dovrebbe restituire `0`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 1000), 0)`)
}})
```

`get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` dovrebbe restituire `1400`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_laptop_cost(laptops, budget):

    return laptops
```

# --solutions--

```py
def get_laptop_cost(laptops, budget):
    laptops = sorted(set(laptops), reverse=True)

    if len(laptops) >= 2 and budget >= laptops[1]:
        return laptops[1]
    if not laptops or budget < laptops[-1]:
        return 0

    for price in laptops[2:]:
        if budget >= price:
            return price

    return 0
```
