---
id: 6821ebea237de8297eaee795
title: "Sfida 19: luce di candela"
challengeType: 29
dashedName: challenge-19
---

# --description--

Dato un numero intero che rappresenta il numero di candele con cui inizi, e un numero intero che indica quante candele bruciate servono per crearne una nuova, restituisci il numero di candele che avrai usato dopo aver creato e bruciato quante più candele possibile.

Per esempio, se hai 7 candele e servono 2 candele bruciate per farne una nuova:

1. Brucia 7 candele per ottenere 7 avanzi,
2. Ricicla 6 avanzi per ottenere 3 nuove candele (rimane 1 avanzo),
3. Brucia 3 candele per ottenere altri 3 avanzi (4 in totale),
4. Ricicla 4 avanzi per ottenere 2 nuove candele,
5. Brucia 2 candele per ottenere 2 avanzi,
6. Ricicla 2 avanzi per ottenere 1 nuova candela,
7. Brucia 1 candela.

Avrai bruciato in totale 13 candele nell'esempio.

# --hints--

`burn_candles(7, 2)` dovrebbe restituire `13`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(7, 2), 13)`)
}})
```

`burn_candles(10, 5)` dovrebbe restituire `12`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(10, 5), 12)`)
}})
```

`burn_candles(20, 3)` dovrebbe restituire `29`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(20, 3), 29)`)
}})
```

`burn_candles(17, 4)` dovrebbe restituire `22`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(17, 4), 22)`)
}})
```

`burn_candles(2345, 3)` dovrebbe restituire `3517`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(2345, 3), 3517)`)
}})
```

# --seed--

## --seed-contents--

```py
def burn_candles(candles, leftovers_needed):

    return candles
```

# --solutions--

```py
def burn_candles(candles, leftovers_needed):
    total_burned = 0
    unused_leftovers = 0

    while candles > 0:
        total_burned += candles
        leftovers = candles + unused_leftovers
        candles = leftovers // leftovers_needed
        unused_leftovers = leftovers % leftovers_needed

    return total_burned
```
