---
id: 6821ebd4237de8297eaee791
title: "Sfida 15: camelCase"
challengeType: 29
dashedName: challenge-15
---

# --description--

Data una stringa, restituisci la sua versione camelCase seguendo queste regole:

- Le parole nell'argomento stringa sono separate da uno o più caratteri appartenenti al seguente insieme: spazio (` `), trattino (`-`) o underscore (`_`). Considera qualsiasi sequenza di questi come una separazione tra parole.
- La prima parola deve essere tutta in minuscolo.
- Ogni parola successiva deve iniziare con una lettera maiuscola, con il resto in minuscolo.
- Tutti gli spazi e i separatori devono essere rimossi.

# --hints--

`to_camel_case("hello world")` dovrebbe restituire `"helloWorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("hello world"), "helloWorld")`)
}})
```

`to_camel_case("HELLO WORLD")` dovrebbe restituire `"helloWorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("HELLO WORLD"), "helloWorld")`)
}})
```

`to_camel_case("secret agent-X")` dovrebbe restituire `"secretAgentX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("secret agent-X"), "secretAgentX")`)
}})
```

`to_camel_case("FREE cODE cAMP")` dovrebbe restituire `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("FREE cODE cAMP"), "freeCodeCamp")`)
}})
```

`to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` dovrebbe restituire `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_camel_case(s):

    return s
```

# --solutions--

```py
import re
def to_camel_case(s):
    words = re.split(r'[_\- ]+', s)

    camel = [
        words[0].lower() if words else ''
    ] + [
        word.capitalize() for word in words[1:]
    ]

    return ''.join(camel)
```
