---
id: 6821ebce237de8297eaee790
title: "Sfida 14: battaglia tra caratteri"
challengeType: 29
dashedName: challenge-14
---

# --description--

Date due stringhe che rappresentano il tuo esercito e un esercito avversario, ogni carattere del tuo esercito combatte contro il carattere nella stessa posizione dell'esercito avversario seguendo queste regole:

- I caratteri `a-z` hanno una forza di `1-26`, rispettivamente.
- I caratteri `A-Z` hanno una forza di `27-52`, rispettivamente.
- Le cifre `0-9` hanno una forza pari al loro valore facciale.
- Tutti gli altri caratteri hanno valore zero.
- Ogni carattere può combattere solo una battaglia.

Per ogni battaglia, vince il carattere più forte. L'esercito con più vittorie vince la guerra. Restituisci i seguenti valori:

- `"Opponent retreated"` se il tuo esercito ha più caratteri dell'esercito avversario.
- `"We retreated"` se l'esercito avversario ha più caratteri del tuo.
- `"We won"` se il tuo esercito ha vinto più battaglie.
- `"We lost"` se l'esercito avversario ha vinto più battaglie.
- `"It was a tie"` se entrambi gli eserciti hanno vinto lo stesso numero di battaglie.

# --hints--

`battle("Hello", "World")` dovrebbe restituire `"We lost"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` dovrebbe restituire `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` dovrebbe restituire `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` dovrebbe restituire `"Opponent retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` dovrebbe restituire `"We retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` dovrebbe restituire `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` dovrebbe restituire `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```
