---
id: 6821ebdf237de8297eaee793
title: "Sfida 17: ordine non convenzionale delle operazioni"
challengeType: 29
dashedName: challenge-17
---

# --description--

Dato un array di numeri interi e un array di operatori stringa, applica le operazioni ai numeri in sequenza da sinistra a destra. Ripeti le operazioni secondo necessità finché tutti i numeri non sono stati usati. Restituisci il risultato finale.

Per esempio, dato `[1, 2, 3, 4, 5]` e `['+', '*']`, restituisci il risultato della valutazione di `1 + 2 * 3 + 4 * 5` da sinistra a destra ignorando il normale ordine delle operazioni.

- Gli operatori validi sono `+`, `-`, `*`, `/` e `%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` dovrebbe restituire `3`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3)`)
}})
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` dovrebbe restituire `38`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38)`)
}})
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` dovrebbe restituire `60`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60)`)
}})
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` dovrebbe restituire `30`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30)`)
}})
```

`evaluate([33, 11, 29, 13], ['/', '-'])` dovrebbe restituire `-2`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([33, 11, 29, 13], ['/', '-']), -2)`)
}})
```

# --seed--

## --seed-contents--

```py
def evaluate(numbers, operators):

    return numbers
```

# --solutions--

```py
def do_math(a, b, operator):
    if operator == '+':
        return a + b
    elif operator == '-':
        return a - b
    elif operator == '*':
        return a * b
    elif operator == '/':
        return a / b
    else:
        return a % b

def evaluate(numbers, operators):
    total = numbers[0]

    for i in range(1, len(numbers)):
        operator = operators[(i - 1) % len(operators)]
        total = do_math(total, numbers[i], operator)

    return total
```
