---
id: 6821ebf8237de8297eaee798
title: "Sfida 22: sequenza tribonacci"
challengeType: 29
dashedName: challenge-22
---

# --description--

La sequenza tribonacci è una serie di numeri in cui ogni numero è la somma dei tre precedenti. Partendo da `0`, `0` e `1`, i primi 10 numeri della sequenza sono `0`, `0`, `1`, `1`, `2`, `4`, `7`, `13`, `24`, `44`.

Dato un array contenente i primi tre numeri di una sequenza tribonacci e un intero che rappresenta la lunghezza della sequenza, restituisci un array contenente la sequenza della lunghezza indicata.

- La tua funzione dovrebbe gestire sequenze di qualsiasi lunghezza maggiore o uguale a zero.
- Se la lunghezza è zero, restituisci un array vuoto.
- Nota che i numeri iniziali fanno parte della sequenza.

# --hints--

`tribonacci_sequence([0, 0, 1], 20)` dovrebbe restituire `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513])`)
}})
```

`tribonacci_sequence([21, 32, 43], 1)` dovrebbe restituire `[21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([21, 32, 43], 1), [21])`)
}})
```

`tribonacci_sequence([0, 0, 1], 0)` dovrebbe restituire `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 0), [])`)
}})
```

`tribonacci_sequence([10, 20, 30], 2)` dovrebbe restituire `[10, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 2), [10, 20])`)
}})
```

`tribonacci_sequence([10, 20, 30], 3)` dovrebbe restituire `[10, 20, 30]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 3), [10, 20, 30])`)
}})
```

`tribonacci_sequence([123, 456, 789], 8)` dovrebbe restituire `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134])`)
}})
```

# --seed--

## --seed-contents--

```py
def tribonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def tribonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return [start_sequence[0], start_sequence[1]]
    if length == 3:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2] + sequence[-3]
        sequence.append(next_value)

    return sequence
```
