---
id: 6821ebee237de8297eaee796
title: "Sfida 20: duplicati in un array"
challengeType: 29
dashedName: challenge-20
---

# --description--

Dato un array di numeri interi, restituisci un array di numeri interi che compaiono più di una volta nell'array iniziale, ordinati in ordine crescente. Se nessun valore compare più di una volta, restituisci un array vuoto.

- Includi solo una istanza di ogni valore nell'array restituito.

# --hints--

`find_duplicates([1, 2, 3, 4, 5])` dovrebbe restituire `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([1, 2, 3, 4, 5]), [])`)
}})
```

`find_duplicates([1, 2, 3, 4, 1, 2])` dovrebbe restituire `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([1, 2, 3, 4, 1, 2]), [1, 2])`)
}})
```

`find_duplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` dovrebbe restituire `[-6, 0, 2, 4, 5, 23]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_duplicates(arr):

    return arr
```

# --solutions--

```py
def find_duplicates(arr):
    seen = set()
    duplicates = set()

    for num in arr:
        if num in seen:
            duplicates.add(num)
        else:
            seen.add(num)

    return sorted(duplicates)
```
