---
id: 68216ef80f957572e7c340c5
title: "Sfida 12: decodificatore di messaggi"
challengeType: 29
dashedName: challenge-12
---

# --description--

Data una stringa di messaggio segreto e un numero intero che rappresenta il numero di lettere con cui il messaggio è stato spostato per codificarlo, restituisci la stringa decodificata.

- Un numero positivo significa che il messaggio è stato spostato in avanti nell'alfabeto.
- Un numero negativo significa che il messaggio è stato spostato indietro nell'alfabeto.
- Il caso conta, i caratteri decodificati devono mantenere il caso dei loro corrispondenti codificati.
- I caratteri non alfabetici non devono essere decodificati.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` dovrebbe restituire `"This is a secret message."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.")`)
}})
```

`decode("Byffi Qilfx!", 20)` dovrebbe restituire `"Hello World!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Byffi Qilfx!", 20), "Hello World!")`)
}})
```

`decode("Zqd xnt njzx?", -1)` dovrebbe restituire `"Are you okay?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Zqd xnt njzx?", -1), "Are you okay?")`)
}})
```

`decode("oannLxmnLjvy", 9)` dovrebbe restituire `"freeCodeCamp"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("oannLxmnLjvy", 9), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message, shift):

    return message
```

# --solutions--

```py
def decode(message, shift):
    decoded_message = []
    for char in message:
        if char.isalpha():
            base = ord('a') if char.islower() else ord('A')
            new_char = chr((ord(char) - base - shift) % 26 + base)
            decoded_message.append(new_char)
        else:
            decoded_message.append(char)
    return ''.join(decoded_message)
```
