---
id: 6821ebda237de8297eaee792
title: "Sfida 16: inversione delle parentesi tonde"
challengeType: 29
dashedName: challenge-16
---

# --description--

Data una stringa che contiene parentesi tonde correttamente annidate, restituisci la versione decodificata della stringa seguendo queste regole:

- Tutti i caratteri all'interno di ogni coppia di parentesi tonde devono essere invertiti.
- Le parentesi tonde devono essere rimosse dal risultato finale.
- Se le parentesi tonde sono annidate, la coppia più interna deve essere invertita per prima, e poi il suo risultato deve essere incluso nell'inversione della coppia esterna.
- Presumi che tutte le parentesi tonde siano bilanciate in modo uniforme e correttamente annidate.

# --hints--

`decode("(f(b(dc)e)a)")` dovrebbe restituire `"abcdef"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` dovrebbe restituire `"Can you read this?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` dovrebbe restituire `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```
