---
id: 68adce01c0e1144d0a90295e
title: "Sfida 29: generatore di acronimi"
challengeType: 29
dashedName: challenge-29
---

# --description--

Dato una stringa contenente una o più parole, restituisci un acronimo delle parole usando le seguenti regole:

- L'acronimo deve consistere nella prima lettera di ogni parola in maiuscolo, salvo diversa indicazione.
- L'acronimo deve ignorare la prima lettera di queste parole a meno che non siano la prima parola della stringa data: `a`, `for`, `an`, `and`, `by` e `of`.
- Le lettere dell'acronimo devono essere restituite nell'ordine in cui appaiono.
- L'acronimo non deve contenere spazi.

# --hints--

`build_acronym("Search Engine Optimization")` dovrebbe restituire `"SEO"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Search Engine Optimization"), "SEO")`)
}})
```

`build_acronym("Frequently Asked Questions")` dovrebbe restituire `"FAQ"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Frequently Asked Questions"), "FAQ")`)
}})
```

`build_acronym("National Aeronautics and Space Administration")` dovrebbe restituire `"NASA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("National Aeronautics and Space Administration"), "NASA")`)
}})
```

`build_acronym("Federal Bureau of Investigation")` dovrebbe restituire `"FBI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Federal Bureau of Investigation"), "FBI")`)
}})
```

`build_acronym("For your information")` dovrebbe restituire `"FYI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("For your information"), "FYI")`)
}})
```

`build_acronym("By the way")` dovrebbe restituire `"BTW"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("By the way"), "BTW")`)
}})
```

`build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` dovrebbe restituire `"AUHWPOTIMSH"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH")`)
}})
```

# --seed--

## --seed-contents--

```py
def build_acronym(s):

    return s
```

# --solutions--

```py
def build_acronym(s):
    small_words = {"a", "for", "an", "and", "by", "of"}
    words = s.split()
    acronym = ""

    for i, word in enumerate(words):
        if i == 0 or word.lower() not in small_words:
            acronym += word[0].upper()

    return acronym
```
