---
id: 68adce01c0e1144d0a90295a
title: "Sfida 27: ruotare una matrice"
challengeType: 29
dashedName: challenge-27
---

# --description--

Data una matrice (un array di array), ruota la matrice di 90 gradi in senso orario e restituiscila. Per esempio, data `[[1, 2], [3, 4]]`, che appare così:

| 1 | 2 |
|---|---|
| 3 | 4 |

Dovresti restituire `[[3, 1], [4, 2]]`, che appare così:

| 3 | 1 |
|---|---|
| 4 | 2 |

# --hints--

`rotate([[1]])` dovrebbe restituire `[[1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1]]), [[1]])`)
}})
```

`rotate([[1, 2], [3, 4]])` dovrebbe restituire `[[3, 1], [4, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]])`)
}})
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` dovrebbe restituire `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]])`)
}})
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` dovrebbe restituire `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def rotate(matrix):

    return matrix
```

# --solutions--

```py
def rotate(matrix):
    n = len(matrix)
    result = [[0] * n for _ in range(n)]

    for i in range(n):
        for j in range(n):
            result[j][n - 1 - i] = matrix[i][j]

    return result
```
