---
id: 68adce01c0e1144d0a902956
title: "Sfida 25: ripetitore di vocali"
challengeType: 29
dashedName: challenge-25
---

# --description--

Data una stringa, restituisci una nuova versione della stringa in cui ogni vocale è duplicata un numero di volte maggiore di uno rispetto alla vocale precedente incontrata. Per esempio, la prima vocale nella frase rimane invariata. La seconda vocale deve apparire due volte di seguito. La terza vocale deve apparire tre volte di seguito, e così via.

- Le lettere `a`, `e`, `i`, `o` e `u`, sia in maiuscolo che in minuscolo, sono considerate vocali.
- La vocale originale deve mantenere il suo caso.
- Le vocali ripetute devono essere in minuscolo.
- Tutti i caratteri non vocali devono mantenere il loro caso originale.

# --hints--

`repeat_vowels("hello world")` dovrebbe restituire `"helloo wooorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` dovrebbe restituire `"freeeCooodeeeeCaaaaamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` dovrebbe restituire `"AEeIiiOoooUuuuu"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` dovrebbe restituire `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
