---
id: 68adce01c0e1144d0a90295c
title: "Sfida 28: parser di numeri romani"
challengeType: 29
dashedName: challenge-28
---

# --description--

Data una stringa che rappresenta un numero romano, restituisci il suo valore intero.

I numeri romani sono composti dai seguenti simboli e valori:

| Simbolo | Valore |
|---------|--------|
| I       | 1      |
| V       | 5      |
| X       | 10     |
| L       | 50     |
| C       | 100    |
| D       | 500    |
| M       | 1000   |

- I numeri si leggono da sinistra a destra. Se un numero più piccolo precede uno più grande, il valore viene sottratto. Altrimenti, i valori si sommano.

# --hints--

`parse_roman_numeral("III")` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("III"), 3)`)
}})
```

`parse_roman_numeral("IV")` dovrebbe restituire `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("IV"), 4)`)
}})
```

`parse_roman_numeral("XXVI")` dovrebbe restituire `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("XXVI"), 26)`)
}})
```

`parse_roman_numeral("XCIX")` dovrebbe restituire `99`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("XCIX"), 99)`)
}})
```

`parse_roman_numeral("CDLX")` dovrebbe restituire `460`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("CDLX"), 460)`)
}})
```

`parse_roman_numeral("DIV")` dovrebbe restituire `504`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("DIV"), 504)`)
}})
```

`parse_roman_numeral("MMXXV")` dovrebbe restituire `2025`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("MMXXV"), 2025)`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_roman_numeral(numeral):

    return numeral
```

# --solutions--

```py
def parse_roman_numeral(numeral):
    roman_map = {
        "I": 1,
        "V": 5,
        "X": 10,
        "L": 50,
        "C": 100,
        "D": 500,
        "M": 1000
    }

    total = 0
    for i in range(len(numeral)):
        current = roman_map[numeral[i]]
        next_val = roman_map[numeral[i + 1]] if i + 1 < len(numeral) else 0

        if current < next_val:
            total -= current
        else:
            total += current

    return total
```
