---
id: 68af0687ef34c76c28ffa547
title: "Sfida 30: caratteri unici"
challengeType: 29
dashedName: challenge-30
---

# --description--

Data una stringa, determina se tutti i caratteri nella stringa sono unici.

- Le lettere maiuscole e minuscole devono essere considerate caratteri diversi.

# --hints--

`all_unique("abc")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("abc"), True)`)
}})
```

`all_unique("aA")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("aA"), True)`)
}})
```

`all_unique("QwErTy123!@")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("QwErTy123!@"), True)`)
}})
```

`all_unique("~!@#$%^&*()_+")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("~!@#$%^&*()_+"), True)`)
}})
```

`all_unique("hello")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("hello"), False)`)
}})
```

`all_unique("freeCodeCamp")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("freeCodeCamp"), False)`)
}})
```

`all_unique("!@#*$%^&*()aA")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("!@#*$%^&*()aA"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def all_unique(s):

    return s
```

# --solutions--

```py
def all_unique(s):
    seen = set()
    for ch in s:
        if ch in seen:
            return False
        seen.add(ch)
    return True
```
