---
id: 68af0687ef34c76c28ffa54d
title: "Sfida 33: tempo davanti allo schermo"
challengeType: 29
dashedName: challenge-33
---

# --description--

Dato un array di sette numeri interi, che rappresentano il tempo di una settimana, dove ogni numero intero è la quantità di ore trascorse al telefono quel giorno, determina se è troppo tempo davanti allo schermo basandoti su questi vincoli:

- Se un singolo giorno ha 10 ore o più, è troppo.
- Se la media di tre giorni consecutivi è maggiore o uguale a 8 ore, è troppo.
- Se la media dei sette giorni è maggiore o uguale a 6 ore, è troppo.

# --hints--

`too_much_screen_time([1, 2, 3, 4, 5, 6, 7])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 4, 5, 6, 7]), False)`)
}})
```

`too_much_screen_time([7, 8, 8, 4, 2, 2, 3])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([7, 8, 8, 4, 2, 2, 3]), False)`)
}})
```

`too_much_screen_time([5, 6, 6, 6, 6, 6, 6])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([5, 6, 6, 6, 6, 6, 6]), False)`)
}})
```

`too_much_screen_time([1, 2, 3, 11, 1, 3, 4])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 11, 1, 3, 4]), True)`)
}})
```

`too_much_screen_time([1, 2, 3, 10, 2, 1, 0])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 10, 2, 1, 0]), True)`)
}})
```

`too_much_screen_time([3, 3, 5, 8, 8, 9, 4])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([3, 3, 5, 8, 8, 9, 4]), True)`)
}})
```

`too_much_screen_time([3, 9, 4, 8, 5, 7, 6])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([3, 9, 4, 8, 5, 7, 6]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def too_much_screen_time(hours):

    return hours
```

# --solutions--

```py
def too_much_screen_time(hours):
    for h in hours:
        if h >= 10:
            return True

    for i in range(5):
        avg3 = sum(hours[i:i+3]) / 3
        if avg3 >= 8:
            return True

    weekly_avg = sum(hours) / 7
    if weekly_avg >= 6:
        return True

    return False
```
