---
id: 68af0687ef34c76c28ffa54b
title: "Sfida 32: inverti la frase"
challengeType: 29
dashedName: challenge-32
---

# --description--

Data una stringa di parole, restituisci una nuova stringa con le parole in ordine inverso. Per esempio, la prima parola dovrebbe trovarsi alla fine della stringa restituita e l'ultima parola all'inizio della stringa restituita.

- Nella stringa data, le parole possono essere separate da uno o più spazi.
- La stringa restituita deve avere solo uno spazio tra le parole.

# --hints--

`reverse_sentence("world hello")` dovrebbe restituire `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("world hello"), "hello world")`)
}})
```

`reverse_sentence("push commit git")` dovrebbe restituire `"git commit push"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("push commit git"), "git commit push")`)
}})
```

`reverse_sentence("npm  install   apt    sudo")` dovrebbe restituire `"sudo apt install npm"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("npm  install   apt    sudo"), "sudo apt install npm")`)
}})
```

`reverse_sentence("import    default   function  export")` dovrebbe restituire `"export function default import"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("import    default   function  export"), "export function default import")`)
}})
```

# --seed--

## --seed-contents--

```py
def reverse_sentence(sentence):

    return sentence
```

# --solutions--

```py
def reverse_sentence(sentence):

   return " ".join(sentence.split()[::-1])
```
