---
id: 68b06e589bf2273243814771
title: "Sfida 36: regolatore del termostato"
challengeType: 29
dashedName: challenge-36
---

# --description--

Dato la temperatura attuale di una stanza e una temperatura obiettivo, restituisci una stringa che indichi come regolare la temperatura della stanza basandoti su questi vincoli:

- Restituisci `"heat"` se la temperatura attuale è inferiore all'obiettivo.
- Restituisci `"cool"` se la temperatura attuale è superiore all'obiettivo.
- Restituisci `"hold"` se la temperatura attuale è uguale all'obiettivo.

# --hints--

`adjust_thermostat(68, 72)` dovrebbe restituire `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(68, 72), "heat")`)
}})
```

`adjust_thermostat(75, 72)` dovrebbe restituire `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(75, 72), "cool")`)
}})
```

`adjust_thermostat(72, 72)` dovrebbe restituire `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 72), "hold")`)
}})
```

`adjust_thermostat(-20.5, -10.1)` dovrebbe restituire `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(-20.5, -10.1), "heat")`)
}})
```

`adjust_thermostat(100, 99.9)` dovrebbe restituire `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(100, 99.9), "cool")`)
}})
```

`adjust_thermostat(0.0, 0.0)` dovrebbe restituire `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(0.0, 0.0), "hold")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(temp, target):

    return temp
```

# --solutions--

```py
def adjust_thermostat(temp, target):
    if temp < target:
        return "heat"
    elif temp > target:
        return "cool"
    else:
        return "hold"
```
