---
id: 68b06e589bf227324381476f
title: "Sfida 35: frequenza delle parole"
challengeType: 29
dashedName: challenge-35
---

# --description--

Dato un paragrafo, restituisci un array con le tre parole che compaiono più frequentemente.

- Le parole nel paragrafo saranno separate da spazi.
- Ignora il maiuscolo/minuscolo nel paragrafo fornito. Per esempio, considera `Hello` e `hello` come la stessa parola.
- Ignora la punteggiatura nel paragrafo fornito. La punteggiatura consiste in virgole (`,`), punti (`.`) e punti esclamativi (`!`).
- L'array restituito deve contenere tutte parole in minuscolo.
- L'array restituito deve essere in ordine decrescente, con la parola che compare più frequentemente per prima.

# --hints--

`get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` dovrebbe restituire `["coding", "python", "in"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

`get_words("I like coding. I like testing. I love debugging!")` dovrebbe restituire `["i", "like", "coding"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

`get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` dovrebbe restituire `["debug", "test", "deploy"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```
