---
id: 68b06e589bf2273243814773
title: "Sfida 37: maiuscole all'inizio della frase"
challengeType: 29
dashedName: challenge-37
---

# --description--

Dato un paragrafo, restituisci un nuovo paragrafo in cui la prima lettera di ogni frase è maiuscola.

- Tutti gli altri caratteri devono essere preservati.
- Le frasi possono terminare con un punto (`.`), uno o più punti interrogativi (`?`) o uno o più punti esclamativi (`!`).

# --hints--

`capitalize("this is a simple sentence.")` dovrebbe restituire `"This is a simple sentence."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("this is a simple sentence."), "This is a simple sentence.")`)
}})
```

`capitalize("hello world. how are you?")` dovrebbe restituire `"Hello world. How are you?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("hello world. how are you?"), "Hello world. How are you?")`)
}})
```

`capitalize("i did today's coding challenge... it was fun!!")` dovrebbe restituire `"I did today's coding challenge... It was fun!!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!")`)
}})
```

`capitalize("crazy!!!strange???unconventional...sentences.")` dovrebbe restituire `"Crazy!!!Strange???Unconventional...Sentences."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.")`)
}})
```

`capitalize("there's a space before this period . why is there a space before that period ?")` dovrebbe restituire `"There's a space before this period . Why is there a space before that period ?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?")`)
}})
```

# --seed--

## --seed-contents--

```py
def capitalize(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def capitalize(paragraph):
    def repl(match):
        return match.group(1) + match.group(2) + match.group(3).upper()

    result = re.sub(r'([.!?]+)(\s*)([a-z])', repl, paragraph)

    if result and result[0].islower():
        result = result[0].upper() + result[1:]

    return result
```
