---
id: 68b06e589bf2273243814777
title: "Sfida 39: riempi il serbatoio"
challengeType: 29
dashedName: challenge-39
---

# --description--

Dato il volume di un serbatoio di carburante, il livello attuale di carburante e il prezzo per gallone, restituisci il costo per riempire completamente il serbatoio.

- `tankSize` è la capacità totale del serbatoio in galloni.
- `fuelLevel` è la quantità attuale di carburante nel serbatoio in galloni.
- `pricePerGallon` è il costo di un gallone di carburante.
- Il valore restituito dovrebbe essere arrotondato a due cifre decimali nel formato: `"$d.dd"`.

# --hints--

`cost_to_fill(20, 0, 4.00)` dovrebbe restituire `"$80.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` dovrebbe restituire `"$17.50"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` dovrebbe restituire `"$29.25"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` dovrebbe restituire `"$0.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` dovrebbe restituire `"$21.89"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
