---
id: 68b06e589bf2273243814775
title: "Sfida 38: generatore di slug"
challengeType: 29
dashedName: challenge-38
---

# --description--

Data una stringa, restituisci una versione della stringa adatta per un URL usando le seguenti regole:

- Tutte le lettere devono essere minuscole.
- Tutti i caratteri che non sono lettere, numeri o spazi devono essere rimossi.
- Tutti gli spazi devono essere sostituiti con il codice spazio codificato per URL `%20`.
- Spazi consecutivi devono essere sostituiti con un singolo `%20`.
- La stringa restituita non deve avere `%20` all'inizio o alla fine.

# --hints--

`generate_slug("helloWorld")` dovrebbe restituire `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("helloWorld"), "helloworld")`)
}})
```

`generate_slug("hello world!")` dovrebbe restituire `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello world!"), "hello%20world")`)
}})
```

`generate_slug(" hello-world ")` dovrebbe restituire `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug(" hello-world "), "helloworld")`)
}})
```

`generate_slug("hello  world")` dovrebbe restituire `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello  world"), "hello%20world")`)
}})
```

`generate_slug("  ?H^3-1*1]0! W[0%R#1]D  ")` dovrebbe restituire `"h3110%20w0r1d"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_slug(str):

    return str
```

# --solutions--

```py
import re
def generate_slug(s):
    cleaned = re.sub(r'[^a-zA-Z0-9 ]+', '', s)
    cleaned = re.sub(r'\s+', ' ', cleaned).strip()
    return cleaned.lower().replace(' ', '%20')
```
