---
id: 68b1f72371a5ac895ac70a04
title: "Sfida 41: archivio file"
challengeType: 29
dashedName: challenge-41
---

# --description--

Dato un file size, un'unità per il file size e la capacità del disco rigido in gigabyte (GB), restituisci il numero di file che il disco rigido può archiviare usando i seguenti vincoli:

- L'unità per il file size può essere byte (`"B"`), kilobyte (`"KB"`) o megabyte (`"MB"`).
- Restituisci il numero di file interi che il disco può contenere.
- Usa le seguenti conversioni:

| Unità | Equivalente |
|:-----:|:-----------:|
| 1 B   |    1 B      |
| 1 KB  |   1000 B    |
| 1 MB  |   1000 KB   |
| 1 GB  |   1000 MB   |

Per esempio, dati `500`, `"KB"` e `1` come argomenti, determina quanti file da 500 KB possono entrare in un disco rigido da 1 GB.

# --hints--

`number_of_files(500, "KB", 1)` dovrebbe restituire `2000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(500, "KB", 1), 2000)`)
}})
```

`number_of_files(50000, "B", 1)` dovrebbe restituire `20000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(50000, "B", 1), 20000)`)
}})
```

`number_of_files(5, "MB", 1)` dovrebbe restituire `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(5, "MB", 1), 200)`)
}})
```

`number_of_files(4096, "B", 1.5)` dovrebbe restituire `366210`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4096, "B", 1.5), 366210)`)
}})
```

`number_of_files(220.5, "KB", 100)` dovrebbe restituire `453514`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(220.5, "KB", 100), 453514)`)
}})
```

`number_of_files(4.5, "MB", 750)` dovrebbe restituire `166666`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4.5, "MB", 750), 166666)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_files(file_size, file_unit, drive_size_gb):

    return file_size
```

# --solutions--

```py
def number_of_files(file_size, file_unit, drive_size_gb):
    drive_size_bytes = drive_size_gb * 1000 * 1000 * 1000

    if file_unit == "B":
        file_size_bytes = file_size
    elif file_unit == "KB":
        file_size_bytes = file_size * 1000
    else:
        file_size_bytes = file_size * 1000 * 1000

    return int(drive_size_bytes // file_size_bytes)
```
