---
id: 68b1f72371a5ac895ac70a02
title: "Sfida 40: archivio di foto"
challengeType: 29
dashedName: challenge-40
---

# --description--

Dato un formato di foto in megabyte (MB) e la capacità di un disco rigido in gigabyte (GB), restituisci il numero di foto che il disco rigido può archiviare usando le seguenti condizioni:

- 1 gigabyte equivale a 1000 megabyte.
- Restituisci il numero intero di foto che il disco può archiviare.

# --hints--

`number_of_photos(1, 1)` dovrebbe restituire `1000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(1, 1), 1000)`)
}})
```

`number_of_photos(2, 1)` dovrebbe restituire `500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(2, 1), 500)`)
}})
```

`number_of_photos(4, 256)` dovrebbe restituire `64000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(4, 256), 64000)`)
}})
```

`number_of_photos(3.5, 750)` dovrebbe restituire `214285`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 750), 214285)`)
}})
```

`number_of_photos(3.5, 5.5)` dovrebbe restituire `1571`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 5.5), 1571)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_photos(photo_size_mb, drive_size_gb):

    return photo_size_mb
```

# --solutions--

```py
def number_of_photos(photo_size_mb, drive_size_gb):
    drive_size_mb = drive_size_gb * 1000
    return drive_size_mb // photo_size_mb
```
