---
id: 68b1f72371a5ac895ac70a06
title: "Sfida 42: archivio video"
challengeType: 29
dashedName: challenge-42
---

# --description--

Dato un formato video, un’unità per la dimensione del video, la capacità di un hard disk e un’unità per l’hard disk, restituisci il numero di video che l’hard disk può archiviare usando i seguenti vincoli:

- L’unità per la dimensione del video può essere byte (`"B"`), kilobyte (`"KB"`), megabyte (`"MB"`) o gigabyte (`"GB"`).
- Se non viene fornita una delle unità video sopra indicate, restituisci `"Invalid video unit"`.
- L’unità della capacità dell’hard disk può essere gigabyte (`"GB"`) o terabyte (`"TB"`).
- Se non viene fornita una delle unità hard disk sopra indicate, restituisci `"Invalid drive unit"`.
- Restituisci il numero intero di video che l’hard disk può contenere.
- Usa le seguenti conversioni:

| Unità | Equivalente |
|:-----:|:-----------:|
| 1 B   | 1 B         |
| 1 KB  | 1000 B      |
| 1 MB  | 1000 KB     |
| 1 GB  | 1000 MB     |
| 1 TB  | 1000 GB     |

Per esempio, dati `500`, `"MB"`, `100` e `"GB"` come argomenti, determina quanti video da 500 MB possono entrare in un hard disk da 100 GB.

# --hints--

`number_of_videos(500, "MB", 100, "GB")` dovrebbe restituire `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500, "MB", 100, "GB"), 200)`)
}})
```

`number_of_videos(1, "TB", 10, "TB")` dovrebbe restituire `"Invalid video unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1, "TB", 10, "TB"), "Invalid video unit")`)
}})
```

`number_of_videos(2000, "MB", 100000, "MB")` dovrebbe restituire `"Invalid drive unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(2000, "MB", 100000, "MB"), "Invalid drive unit")`)
}})
```

`number_of_videos(500000, "KB", 2, "TB")` dovrebbe restituire `4000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500000, "KB", 2, "TB"), 4000)`)
}})
```

`number_of_videos(1.5, "GB", 2.2, "TB")` dovrebbe restituire `1466`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1.5, "GB", 2.2, "TB"), 1466)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):

    return video_size
```

# --solutions--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):
    video_units = {"KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000}
    drive_units = {"GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000}

    if video_unit not in video_units:
        return "Invalid video unit"
    if drive_unit not in drive_units:
        return "Invalid drive unit"

    video_bytes = video_size * video_units[video_unit]
    drive_bytes = drive_size * drive_units[drive_unit]

    return int(drive_bytes // video_bytes)
```
