---
id: 68b7687dded630607aceccaf
title: "Sfida 47: beccati a superare il limite di velocità"
challengeType: 29
dashedName: challenge-47
---

# --description--

Dato un array di numeri che rappresentano la velocità con cui sono stati osservati veicoli in movimento, e un numero che rappresenta il limite di velocità, restituisci un array con due elementi: il numero di veicoli che superavano il limite e la media di quanto questi veicoli superavano il limite.

- Se non ci sono stati veicoli che superavano il limite, restituisci `[0, 0]`.

# --hints--

`speeding([50, 60, 55], 60)` dovrebbe restituire `[0, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([50, 60, 55], 60), [0, 0])`)
}})
```

`speeding([58, 50, 60, 55], 55)` dovrebbe restituire `[2, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([58, 50, 60, 55], 55), [2, 4])`)
}})
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)` dovrebbe restituire `[4, 8.5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5])`)
}})
```

`speeding([100, 105, 95, 102], 100)` dovrebbe restituire `[2, 3.5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([100, 105, 95, 102], 100), [2, 3.5])`)
}})
```

`speeding([40, 45, 44, 50, 112, 39], 55)` dovrebbe restituire `[1, 57]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57])`)
}})
```

# --seed--

## --seed-contents--

```py
def speeding(speeds, limit):

    return speeds
```

# --solutions--

```py
def speeding(speeds, limit):
    speeding = 0
    total_over = 0

    for speed in speeds:
        if speed > limit:
            speeding += 1
            total_over += (speed - limit)
    
    if speeding == 0:
        return [0, 0]
    
    return [speeding, total_over / speeding]
```
