---
id: 68b7cadffed0e75a517da677
title: "Sfida 54: forza della p@ssw0rd!"
challengeType: 29
dashedName: challenge-54
---

# --description--

Data una stringa password, restituisci `"weak"`, `"medium"` o `"strong"` in base alla forza della password.

Una password viene valutata secondo le seguenti regole:

- È lunga almeno 8 caratteri.
- Contiene sia lettere maiuscole che minuscole.
- Contiene almeno un numero.
- Contiene almeno un carattere speciale da questo insieme: `!`, `@`, `#`, `$`, `%`, `^`, `&` o `*`.

Restituisci `"weak"` se la password soddisfa meno di due regole.
Restituisci `"medium"` se la password soddisfa 2 o 3 regole.
Restituisci `"strong"` se la password soddisfa tutte e 4 le regole.

# --hints--

`check_strength("123456")` dovrebbe restituire `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("123456"), "weak")`)
}})
```

`check_strength("pass!!!")` dovrebbe restituire `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("pass!!!"), "weak")`)
}})
```

`check_strength("Qwerty")` dovrebbe restituire `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("Qwerty"), "weak")`)
}})
```

`check_strength("PASSWORD")` dovrebbe restituire `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD"), "weak")`)
}})
```

`check_strength("PASSWORD!")` dovrebbe restituire `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD!"), "medium")`)
}})
```

`check_strength("PassWord%^!")` dovrebbe restituire `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PassWord%^!"), "medium")`)
}})
```

`check_strength("qwerty12345")` dovrebbe restituire `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("qwerty12345"), "medium")`)
}})
```

`check_strength("S3cur3P@ssw0rd")` dovrebbe restituire `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("S3cur3P@ssw0rd"), "strong")`)
}})
```

`check_strength("C0d3&Fun!")` dovrebbe restituire `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("C0d3&Fun!"), "strong")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_strength(password):

    return password
```

# --solutions--

```py
import re
def check_strength(password):
    rules_met = 0

    if len(password) >= 8:
        rules_met += 1
    if re.search(r'[a-z]', password) and re.search(r'[A-Z]', password):
        rules_met += 1
    if re.search(r'\d', password):
        rules_met += 1
    if re.search(r'[!@#$%^&*]', password):
        rules_met += 1

    if rules_met < 2:
        return "weak"
    elif rules_met < 4:
        return "medium"
    else:
        return "strong"
```
