---
id: 68b7cadffed0e75a517da673
title: "Sfida 52: da binario a decimale"
challengeType: 29
dashedName: challenge-52
---

# --description--

Dato una stringa che rappresenta un numero binario, restituisci il suo equivalente decimale come numero.

Un numero binario usa solo le cifre `0` e `1` per rappresentare qualsiasi numero. Per convertire un numero binario in decimale, moltiplica ogni cifra per una potenza di `2` e somma i risultati. Inizia moltiplicando la cifra più a destra per `2^0`, la cifra successiva a sinistra per `2^1` e così via. Una volta moltiplicate tutte le cifre per una potenza di `2`, somma il risultato.

Per esempio, il numero binario `101` è uguale a `5` in decimale perché:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`to_decimal("101")` dovrebbe restituire `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("101"), 5)`)
}})
```

`to_decimal("1010")` dovrebbe restituire `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010"), 10)`)
}})
```

`to_decimal("10010")` dovrebbe restituire `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("10010"), 18)`)
}})
```

`to_decimal("1010101")` dovrebbe restituire `85`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010101"), 85)`)
}})
```

# --seed--

## --seed-contents--

```py
def to_decimal(binary):

    return binary
```

# --solutions--

```py
def to_decimal(binary):
    decimal = 0
    for i, bit_char in enumerate(reversed(binary)):
        bit = int(bit_char)
        decimal += bit * (2 ** i)
    return decimal
```
