---
id: 68b7687dded630607aceccad
title: "Sfida 46: secondo numero più grande"
challengeType: 29
dashedName: challenge-46
---

# --description--

Dato un array, restituisci il secondo numero distinto più grande.

# --hints--

`second_largest([1, 2, 3, 4])` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([1, 2, 3, 4]), 3)`)
}})
```

`second_largest([20, 139, 94, 67, 31])` dovrebbe restituire `94`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([20, 139, 94, 67, 31]), 94)`)
}})
```

`second_largest([2, 3, 4, 6, 6])` dovrebbe restituire `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([2, 3, 4, 6, 6]), 4)`)
}})
```

`second_largest([10, -17, 55.5, 44, 91, 0])` dovrebbe restituire `55.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([10, -17, 55.5, 44, 91, 0]), 55.5)`)
}})
```

`second_largest([1, 0, -1, 0, 1, 0, -1, 1, 0])` dovrebbe restituire `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(second_largest([1, 0, -1, 0, 1, 0, -1, 1, 0]), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def second_largest(arr):

    return arr
```

# --solutions--

```py
def second_largest(arr):
    unique = list(set(arr))
    unique.sort(reverse=True)
    return unique[1]
```
