---
id: 68b7687dded630607aceccab
title: "Sfida 45: quadrato perfetto"
challengeType: 29
dashedName: challenge-45
---

# --description--

Dato un numero intero, determina se è un quadrato perfetto.

- Un numero è un quadrato perfetto se puoi moltiplicare un intero per se stesso per ottenerlo. Per esempio, 9 è un quadrato perfetto perché puoi moltiplicare 3 per se stesso per ottenerlo.

# --hints--

`is_perfect_square(9)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(9), True)`)
}})
```

`is_perfect_square(49)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(49), True)`)
}})
```

`is_perfect_square(1)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(1), True)`)
}})
```

`is_perfect_square(2)` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(2), False)`)
}})
```

`is_perfect_square(99)` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(99), False)`)
}})
```

`is_perfect_square(-9)` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(-9), False)`)
}})
```

`is_perfect_square(0)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(0), True)`)
}})
```

`is_perfect_square(25281)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(25281), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_perfect_square(n):

    return n
```

# --solutions--

```py
import math
def is_perfect_square(n):
    if n < 0:
        return False
    root = int(math.sqrt(n))
    return root * root == n
```
