---
id: 68b7687dded630607aceccb3
title: "Sfida 49: parser di intestazioni CSV"
challengeType: 29
dashedName: challenge-49
---

# --description--

Data la prima riga di un file CSV (valori separati da virgola), restituisci un array contenente le intestazioni.

- La prima riga di un file CSV contiene intestazioni separate da virgole.
- Rimuovi eventuali spazi bianchi iniziali o finali da ogni intestazione.

# --hints--

`get_headings("name,age,city")` dovrebbe restituire `["name", "age", "city"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("name,age,city"), ["name", "age", "city"])`)
}})
```

`get_headings("first name,last name,phone")` dovrebbe restituire `["first name", "last name", "phone"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("first name,last name,phone"), ["first name", "last name", "phone"])`)
}})
```

`get_headings("username , email , signup date ")` dovrebbe restituire `["username", "email", "signup date"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("username , email , signup date "), ["username", "email", "signup date"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_headings(csv):

    return csv
```

# --solutions--

```py
def get_headings(csv):
    return [h.strip() for h in csv.split(",")]
```
