---
id: 68b7687dded630607aceccb1
title: "Sfida 48: rilevatore di spam"
challengeType: 29
dashedName: challenge-48
---

# --description--

Dato un numero di telefono nel formato `"+A (BBB) CCC-DDDD"`, dove ogni lettera rappresenta una cifra come segue:

- `A` rappresenta il prefisso internazionale e può essere composto da un numero qualsiasi di cifre.
- `BBB` rappresenta il prefisso dell'area ed è sempre composto da tre cifre.
- `CCC` e `DDDD` rappresentano il numero locale e sono rispettivamente di tre e quattro cifre.

Determina se si tratta di un numero spam in base ai seguenti criteri:

- Il prefisso internazionale è più lungo di 2 cifre oppure non inizia con uno zero (`0`).
- Il prefisso dell'area è maggiore di 900 o minore di 200.
- La somma delle prime tre cifre del numero locale appare nelle ultime quattro cifre del numero locale.
- Il numero contiene la stessa cifra ripetuta quattro o più volte di seguito (ignorando i caratteri di formattazione).

# --hints--

`is_spam("+0 (200) 234-0182")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (200) 234-0182"), False)`)
}})
```

`is_spam("+091 (555) 309-1922")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+091 (555) 309-1922"), True)`)
}})
```

`is_spam("+1 (555) 435-4792")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+1 (555) 435-4792"), True)`)
}})
```

`is_spam("+0 (955) 234-4364")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (955) 234-4364"), True)`)
}})
```

`is_spam("+0 (155) 131-6943")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (155) 131-6943"), True)`)
}})
```

`is_spam("+0 (555) 135-0192")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 135-0192"), True)`)
}})
```

`is_spam("+0 (555) 564-1987")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 564-1987"), True)`)
}})
```

`is_spam("+00 (555) 234-0182")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+00 (555) 234-0182"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_spam(number):

    return number
```

# --solutions--

```py
import re
def is_spam(number):
    digits = re.sub(r"\D", "", number)
    match = re.match(r"^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$", number)
    country_code, area_code, ccc, dddd = match.groups()

    if len(country_code) > 2 or not country_code.startswith("0"):
        return True

    area_num = int(area_code)
    if area_num > 900 or area_num < 200:
        return True

    sum_ccc = sum(int(d) for d in ccc)
    if str(sum_ccc) in dddd:
        return True

    if re.search(r"(\d)\1\1\1", digits):
        return True

    return False
```
