---
id: 68c1a929005bf54d342aa8d5
title: "Sfida 58: settimana dello spazio giorno 4: punto di atterraggio"
challengeType: 29
dashedName: challenge-58
---

# --description--

Nel quarto giorno della settimana dello spazio, ti viene fornita una matrice di numeri (un array di array), che rappresentano potenziali punti di atterraggio per il tuo rover. Trova il punto di atterraggio più sicuro basandoti sulle seguenti regole:

- Ogni punto nella matrice conterrà un numero da `0-9`, incluso.
- Qualsiasi `0` rappresenta un potenziale punto di atterraggio.
- Qualsiasi numero diverso da `0` è troppo pericoloso per atterrare. Più il numero è alto, più è pericoloso.
- Il punto più sicuro è definito come la cella `0` le cui celle circostanti (fino a 4 vicini, escludendo le diagonali) hanno il pericolo totale più basso.
- Ignora i vicini fuori dai limiti (gli angoli e i bordi hanno semplicemente meno vicini).
- Restituisci gli indici del punto di atterraggio più sicuro. Ci sarà sempre solo un punto più sicuro.

Per esempio, dato:

```js
[
  [1, 0],
  [2, 0]
]
```

Restituisci `[0, 1]`, gli indici per il `0` nel primo array.

# --hints--

`find_landing_spot([[1, 0], [2, 0]])` dovrebbe restituire `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 0], [2, 0]]), [0, 1])`)
}})
```

`find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` dovrebbe restituire `[1, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1])`)
}})
```

`find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` dovrebbe restituire `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2])`)
}})
```

`find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` dovrebbe restituire `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_landing_spot(matrix):

    return matrix
```

# --solutions--

```py
def find_landing_spot(matrix):
    best_spot = None
    lowest_neighbor_sum = float('inf')

    for i in range(len(matrix)):
        for j in range(len(matrix[i])):
            if matrix[i][j] == 0:
                current_neighbor_sum = 0

                if i > 0:
                    current_neighbor_sum += matrix[i - 1][j]
                if j < len(matrix[i]) - 1:
                    current_neighbor_sum += matrix[i][j + 1]
                if i < len(matrix) - 1:
                    current_neighbor_sum += matrix[i + 1][j]
                if j > 0:
                    current_neighbor_sum += matrix[i][j - 1]

                if current_neighbor_sum < lowest_neighbor_sum:
                    lowest_neighbor_sum = current_neighbor_sum
                    best_spot = [i, j]

    return best_spot
```
