---
id: 68c1a929005bf54d342aa8d6
title: "Sfida 59: settimana dello spazio giorno 5: zona abitabile"
challengeType: 29
dashedName: challenge-59
---

# --description--

Per il quinto giorno della settimana dello spazio, calcolerai la "zona abitabile" di una stella — la regione intorno a una stella dove le condizioni sono "giuste" per far esistere acqua liquida.

Data la massa di una stella, restituisci un array con le distanze di inizio e fine della sua zona abitabile in unità astronomiche.

Per calcolare la zona abitabile:

1. Trova la luminosità della stella elevando la sua massa alla potenza di 3,5.
2. L'inizio della zona è 0,95 volte la radice quadrata della sua luminosità.
3. La fine della zona è 1,37 volte la radice quadrata della sua luminosità.

- Restituisci le distanze arrotondate a due cifre decimali.

Per esempio, data `1` come massa, restituisci `[0.95, 1.37]`.

# --hints--

`goldilocks_zone(1)` dovrebbe restituire `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` dovrebbe restituire `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` dovrebbe restituire `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` dovrebbe restituire `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` dovrebbe restituire `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
