---
id: 68c1a929005bf54d342aa8d3
title: "Sfida 56: settimana dello spazio giorno 2: ricerca di esopianeti"
challengeType: 29
dashedName: challenge-56
---

# --description--

Per il secondo giorno della settimana dello spazio, ti viene data una stringa in cui ogni carattere rappresenta la lettura di luminosità di una stella. Determina se le letture hanno rilevato un esopianeta usando il metodo del transito. Il metodo del transito si verifica quando un pianeta passa davanti a una stella, riducendo la sua luminosità osservata.

- Le letture di luminosità comprendono solo i caratteri `0-9` e `A-Z` dove ogni lettura corrisponde ai seguenti valori numerici:
- I caratteri `0-9` corrispondono ai livelli di luminosità `0-9`.
- I caratteri `A-Z` corrispondono ai livelli di luminosità `10-35`.

Una stella è considerata avere un esopianeta se una qualsiasi lettura singola è minore o uguale all’80% della media di tutte le letture. Per esempio, se la luminosità media di una stella è 10, sarà considerata avere un esopianeta se una qualsiasi lettura singola è 8 o meno.

# --hints--

`has_exoplanet("665544554")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("665544554"), False)`)
}})
```

`has_exoplanet("FGFFCFFGG")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FGFFCFFGG"), True)`)
}})
```

`has_exoplanet("MONOPLONOMONPLNOMPNOMP")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("MONOPLONOMONPLNOMPNOMP"), False)`)
}})
```

`has_exoplanet("FREECODECAMP")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FREECODECAMP"), True)`)
}})
```

`has_exoplanet("9AB98AB9BC98A")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("9AB98AB9BC98A"), False)`)
}})
```

`has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_exoplanet(readings):

    return readings
```

# --solutions--

```py
def has_exoplanet(readings):
    values = [int(c, 36) for c in readings]
    average = sum(values) / len(values)
    threshold = average * 0.8
    for v in values:
        if v <= threshold:
            return True
    return False
```
