---
id: 68c1a929005bf54d342aa8d4
title: "Sfida 57: settimana dello spazio giorno 3: telefono a casa"
challengeType: 29
dashedName: challenge-57
---

# --description--

Per il terzo giorno della settimana dello spazio, ti viene dato un array di numeri che rappresentano le distanze (in chilometri) tra te, i satelliti e il tuo pianeta natale in una rotta di comunicazione. Determina quanto tempo impiegherà un messaggio inviato lungo la rotta per raggiungere il pianeta di destinazione usando le seguenti condizioni:

- Il primo valore nell'array è la distanza dalla tua posizione al primo satellite.
- Ogni valore successivo, tranne l'ultimo, è la distanza fino al satellite successivo.
- L'ultimo valore nell'array è la distanza dall'ultimo satellite al tuo pianeta natale.
- Il messaggio viaggia a 300.000 km/s.
- Ogni satellite che il messaggio **attraversa** aggiunge un ritardo di trasmissione di 0,5 secondi.
- Restituisci un numero arrotondato a 4 cifre decimali, eliminando gli zeri finali.

# --hints--

`send_message([300000, 300000])` dovrebbe restituire `2.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([300000, 300000]), 2.5)`)
}})
```

`send_message([384400, 384400])` dovrebbe restituire `3.0627`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([384400, 384400]), 3.0627)`)
}})
```

`send_message([54600000, 54600000])` dovrebbe restituire `364.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([54600000, 54600000]), 364.5)`)
}})
```

`send_message([1000000, 500000000, 1000000])` dovrebbe restituire `1674.3333`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([1000000, 500000000, 1000000]), 1674.3333)`)
}})
```

`send_message([10000, 21339, 50000, 31243, 10000])` dovrebbe restituire `2.4086`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([10000, 21339, 50000, 31243, 10000]), 2.4086)`)
}})
```

`send_message([802101, 725994, 112808, 3625770, 481239])` dovrebbe restituire `21.1597`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([802101, 725994, 112808, 3625770, 481239]), 21.1597)`)
}})
```

# --seed--

## --seed-contents--

```py
def send_message(route):

    return route
```

# --solutions--

```py
def send_message(route):
    total_distance = sum(route)
    delay = (len(route) - 1) * 0.5
    time = total_distance / 300_000
    total = time + delay
    return round(total, 4)
```
