---
id: 68c497f3aaefc9fd9f1b0e24
title: "Sfida 60: settimana dello spazio giorno 6: fase lunare"
challengeType: 29
dashedName: challenge-60
---

# --description--

Per il sesto giorno della settimana dello spazio, ti verrà fornita una data nel formato `"YYYY-MM-DD"` e dovrai determinare la fase della luna per quel giorno usando le seguenti regole:

Usa un ciclo lunare semplificato di 28 giorni, diviso in quattro fasi uguali:

- `"New"`: giorni 1 - 7
- `"Waxing"`: giorni 8 - 14
- `"Full"`: giorni 15 - 21
- `"Waning"`: giorni 22 - 28

Dopo il giorno 28, il ciclo si ripete con il giorno 1, luna nuova.

- Usa `"2000-01-06"` come luna nuova di riferimento (giorno 1 del ciclo) per determinare la fase del giorno dato.
- Non ti verranno date date precedenti alla data di riferimento.
- Restituisci la fase corretta come stringa.

**Nota:** Il giorno 1 rappresenta il giorno della luna nuova, cioè 0 giorni sono passati dall’ultima luna nuova.

# --hints--

`moon_phase("2000-01-12")` dovrebbe restituire `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` dovrebbe restituire `"Waxing"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` dovrebbe restituire `"Full"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` dovrebbe restituire `"Waning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` dovrebbe restituire `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```
