---
id: 68c497f3aaefc9fd9f1b0e25
title: "Sfida 61: settimana dello spazio giorno 7: carburante per il lancio"
challengeType: 29
dashedName: challenge-61
---

# --description--

Per l'ultimo giorno della settimana dello spazio, ti verrà data la massa in chilogrammi (kg) di un payload che vuoi mandare in orbita. Determina la quantità di carburante necessaria per mandare il tuo payload in orbita seguendo queste regole:

- I razzi richiedono 1 kg di carburante ogni 5 kg di massa che devono sollevare.
- Il carburante stesso ha massa. Quindi quando aggiungi carburante, la massa da sollevare aumenta, il che richiede più carburante, che aumenta la massa, e così via.
- Per calcolare il carburante totale necessario: inizia con la massa del payload, calcola il carburante necessario per quella massa, aggiungi quel carburante alla massa totale e calcola di nuovo. Ripeti questo processo finché il carburante aggiuntivo richiesto è inferiore a 1 kg, poi fermati.
- Ignora la massa del razzo stesso. Calcola solo il carburante necessario per sollevare il payload e il suo carburante.

Per esempio, dato un payload di 50 kg, avresti bisogno di 10 kg di carburante per sollevarlo (payload / 5), il che porta la massa totale a 60 kg, che richiede 12 kg per sollevarla (2 kg aggiuntivi), il che porta la massa totale a 62 kg, che richiede 12,4 kg per sollevarla - 0,4 kg aggiuntivi - che è meno di 1 kg aggiuntivo, quindi ci fermiamo qui. La massa totale da sollevare è 62,4 kg, di cui 50 è il payload iniziale e 12,4 è il carburante.

- Restituisci la quantità di carburante necessaria arrotondata a una cifra decimale.

# --hints--

`launch_fuel(50)` dovrebbe restituire `12.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(50), 12.4)`)
}})
```

`launch_fuel(500)` dovrebbe restituire `124.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(500), 124.8)`)
}})
```

`launch_fuel(243)` dovrebbe restituire `60.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(243), 60.7)`)
}})
```

`launch_fuel(11000)` dovrebbe restituire `2749.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(11000), 2749.8)`)
}})
```

`launch_fuel(6214)` dovrebbe restituire `1553.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(6214), 1553.4)`)
}})
```

# --seed--

## --seed-contents--

```py
def launch_fuel(payload):

    return payload
```

# --solutions--

```py
def launch_fuel(payload):
    total_mass = payload
    additional_fuel = total_mass / 5
    total_fuel = additional_fuel

    while additional_fuel >= 1:
        total_mass += additional_fuel
        additional_fuel = (total_mass / 5) - total_fuel
        total_fuel += additional_fuel

    return round(total_fuel, 1)
```
