---
id: 68c497f3aaefc9fd9f1b0e26
title: "Sfida 62: da esadecimale a decimale"
challengeType: 29
dashedName: challenge-62
---

# --description--

Data una stringa che rappresenta un numero esadecimale (base 16), restituisci il suo valore decimale (base 10) come intero.

L'esadecimale è un sistema numerico che usa 16 cifre:

- `0-9` rappresentano i valori da `0` a `9`.
- `A-F` rappresentano i valori da `10` a `15`.

Ecco una tabella di conversione parziale:

| Esadecimale | Decimale |
|:-----------:|:--------:|
| 0           | 0        |
| 1           | 1        |
| ...         | ...      |
| 9           | 9        |
| A           | 10       |
| ...         | ...      |
| F           | 15       |
| 10          | 16       |
| ...         | ...      |
| 9F          | 159      |
| A0          | 160      |
| ...         | ...      |
| FF          | 255      |
| 100         | 256      |

- La stringa conterrà solo i caratteri `0–9` e `A–F`.

# --hints--

`hex_to_decimal("A")` dovrebbe restituire `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A"), 10)`)
}})
```

`hex_to_decimal("15")` dovrebbe restituire `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("15"), 21)`)
}})
```

`hex_to_decimal("2E")` dovrebbe restituire `46`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("2E"), 46)`)
}})
```

`hex_to_decimal("FF")` dovrebbe restituire `255`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("FF"), 255)`)
}})
```

`hex_to_decimal("A3F")` dovrebbe restituire `2623`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A3F"), 2623)`)
}})
```

# --seed--

## --seed-contents--

```py
def hex_to_decimal(hex):

    return hex
```

# --solutions--

```py
def hex_to_decimal(hex):

    return int(hex, 16)
```
