---
id: 68cae5b538ff798bbd4da004
title: "Sfida 66: rimuovere i tag HTML"
challengeType: 29
dashedName: challenge-66
---

# --description--

Dato un stringa di codice HTML, rimuovi i tag e restituisci il contenuto di testo semplice.

- La stringa di input conterrà solo HTML valido.
- I tag HTML possono essere annidati.
- Rimuovi i tag e tutti gli attributi.

Ad esempio, `'<a href="#">Click here</a>'` dovrebbe restituire `"Click here"`.

# --hints--

`strip_tags('<a href="#">Click here</a>')` dovrebbe restituire `"Click here"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

`strip_tags('<p class="center">Hello <b>World</b>!</p>')` dovrebbe restituire `"Hello World!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

`strip_tags('<img src="cat.jpg" alt="Cat">')` dovrebbe restituire una stringa vuota (`""`).

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

`strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` dovrebbe restituire `sectionsection`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```
