---
id: 68cae5b538ff798bbd4da008
title: "Sfida 70: estrattore di attributi HTML"
challengeType: 29
dashedName: challenge-70
---

# --description--

Dato un stringa di un elemento HTML valido, restituisci gli attributi dell'elemento seguendo questi criteri:

- Ti verrà fornito solo un elemento.
- Gli attributi saranno nel formato: `attribute="value"`.
- Restituisci un array di stringhe con ogni proprietà e valore dell'attributo, separati da una virgola, in questo formato: `["attribute1, value1", "attribute2, value2"]`.
- Restituisci gli attributi nell'ordine in cui sono dati.
- Se non vengono trovati attributi, restituisci un array vuoto.

# --hints--

`extract_attributes('<span class="red"></span>')` dovrebbe restituire `["class, red"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<span class="red"></span>'), ["class, red"])`)
}})
```

`extract_attributes('<meta charset="UTF-8" />')` dovrebbe restituire `["charset, UTF-8"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<meta charset="UTF-8" />'), ["charset, UTF-8"])`)
}})
```

`extract_attributes("<p>Lorem ipsum dolor sit amet</p>")` dovrebbe restituire `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes("<p>Lorem ipsum dolor sit amet</p>"), [])`)
}})
```

`extract_attributes('<input name="email" type="email" required="true" />')` dovrebbe restituire `["name, email", "type, email", "required, true"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"])`)
}})
```

`extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>')` dovrebbe restituire `["id, submit", "class, btn btn-primary"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"])`)
}})
```

# --seed--

## --seed-contents--

```py
def extract_attributes(element):

    return element
```

# --solutions--

```py
import re
def extract_attributes(element):
    pattern = r'([\w-]+)="([^"]*)"'
    matches = re.findall(pattern, element)
    return [f"{attr}, {val}" for attr, val in matches]
```
