---
id: 68cae5b538ff798bbd4da001
title: "Sfida 63: battaglia di parole"
challengeType: 29
dashedName: challenge-63
---

# --description--

Date due frasi che rappresentano la tua squadra e una squadra avversaria, dove ogni parola della tua squadra combatte contro la parola corrispondente della squadra avversaria, determina quale squadra vince seguendo queste regole:

- Le frasi date conterranno sempre lo stesso numero di parole.
- Le parole sono separate da uno spazio singolo e conterranno solo lettere.
- Il valore di ogni parola è la somma delle sue lettere.
- Le lettere `a` a `z` corrispondono ai valori `1` fino a `26`. Per esempio, `a` vale `1`, e `z` vale `26`.
- Una lettera maiuscola raddoppia il valore della lettera. Per esempio, `A` vale `2`, e `Z` vale `52`.
- Le parole combattono in ordine: la prima parola della tua squadra combatte contro la prima parola della squadra avversaria, e così via.
- Una parola vince se il suo valore è maggiore del valore della parola avversaria.
- La squadra con più parole vincenti è la vincitrice.

Restituisci `"We win"` se la tua squadra vince, `"We lose"` se la tua squadra perde, e `"Draw"` se entrambe le squadre hanno lo stesso numero di vittorie.

# --hints--

`battle("hello world", "hello word")` dovrebbe restituire `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` dovrebbe restituire `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` dovrebbe restituire `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` dovrebbe restituire `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` dovrebbe restituire `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` dovrebbe restituire `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` dovrebbe restituire `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
