---
id: 68cae5b538ff798bbd4da009
title: "Sfida 71: calcolatore di mance"
challengeType: 29
dashedName: challenge-71
---

# --description--

Dato il prezzo del tuo pasto e una percentuale di mancia personalizzata, restituisci un array con tre valori di mancia: 15%, 20% e l'importo personalizzato.

- I prezzi saranno forniti nel formato: `"$N.NN"`.
- Le percentuali di mancia personalizzate saranno fornite in questo formato: `"25%"`.
- Restituisci gli importi nello stesso formato `"$N.NN"`, arrotondati a due cifre decimali.

Per esempio, dato un prezzo del pasto `"$10.00"` e un valore di mancia personalizzato `"25%"`, restituisci `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculate_tips("$10.00", "25%")` dovrebbe restituire `["$1.50", "$2.00", "$2.50"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"])`)
}})
```

`calculate_tips("$89.67", "26%")` dovrebbe restituire `["$13.45", "$17.93", "$23.31"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"])`)
}})
```

`calculate_tips("$19.85", "9%")` dovrebbe restituire `["$2.98", "$3.97", "$1.79"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_tips(meal_price, custom_tip):

    return meal_price
```

# --solutions--

```py
def calculate_tips(meal_price, custom_tip):
    meal = float(meal_price[1:])
    custom_percent = float(custom_tip[:-1])

    tips_percent = [15, 20, custom_percent]
    tips = [f"${meal * p / 100:.2f}" for p in tips_percent]

    return tips
```
