---
id: 68cae5b538ff798bbd4da002
title: "Sfida 64: da 24 a 12"
challengeType: 29
dashedName: challenge-64
---

# --description--

Data una stringa che rappresenta un orario nel formato 24 ore di `"HHMM"`, restituisci l'orario nel suo formato equivalente a 12 ore di `"H:MM AM"` o `"H:MM PM"`.

- L'input fornito sarà sempre una stringa di quattro cifre nel formato orario a 24 ore, da `"0000"` a `"2359"`.

# --hints--

`to_12("1124")` dovrebbe restituire `"11:24 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1124"), "11:24 AM")`)
}})
```

`to_12("0900")` dovrebbe restituire `"9:00 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0900"), "9:00 AM")`)
}})
```

`to_12("1455")` dovrebbe restituire `"2:55 PM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1455"), "2:55 PM")`)
}})
```

`to_12("2346")` dovrebbe restituire `"11:46 PM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("2346"), "11:46 PM")`)
}})
```

`to_12("0030")` dovrebbe restituire `"12:30 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0030"), "12:30 AM")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_12(time):

    return time
```

# --solutions--

```py
def convert_hours(hours):
    if hours == 0:
        return 12
    elif hours > 12:
        return hours - 12
    else:
        return hours

def to_12(time):
    hours = int(time[:2])
    minutes = time[2:]
    period = "AM" if hours < 12 else "PM"

    return f"{convert_hours(hours)}:{minutes} {period}"
```
