---
id: 68cae5b538ff798bbd4da006
title: "Sfida 68: mascherare il numero di carta di credito"
challengeType: 29
dashedName: challenge-68
---

# --description--

Dato un stringa di numeri di carta di credito, restituisci una versione mascherata usando le seguenti regole:

- La stringa conterrà quattro insiemi di quattro cifre (`0-9`), con tutti gli insiemi separati da uno spazio singolo o da un singolo trattino (`-`).
- Sostituisci tutti i numeri, tranne gli ultimi quattro, con un asterisco (`*`).
- Lascia invariati gli altri caratteri.

Per esempio, dato `"4012-8888-8888-1881"` restituisci `"****-****-****-1881"`.

# --hints--

`mask("4012-8888-8888-1881")` dovrebbe restituire `"****-****-****-1881"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("4012-8888-8888-1881"), "****-****-****-1881")`)
}})
```

`mask("5105 1051 0510 5100")` dovrebbe restituire `"**** **** **** 5100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("5105 1051 0510 5100"), "**** **** **** 5100")`)
}})
```

`mask("6011 1111 1111 1117")` dovrebbe restituire `"**** **** **** 1117"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("6011 1111 1111 1117"), "**** **** **** 1117")`)
}})
```

`mask("2223-0000-4845-0010")` dovrebbe restituire `"****-****-****-0010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("2223-0000-4845-0010"), "****-****-****-0010")`)
}})
```

# --seed--

## --seed-contents--

```py
def mask(card):

    return card
```

# --solutions--

```py
def mask(card):
    if '-' in card:
        split = card.split('-')
        return f"****-****-****-{split[3]}"
    else:
        split = card.split(' ')
        return f"**** **** **** {split[3]}"
```
