---
id: 68cae5b538ff798bbd4da003
title: "Sfida 65: conteggio delle stringhe"
challengeType: 29
dashedName: challenge-65
---

# --description--

Date due stringhe, determina quante volte la seconda stringa appare nella prima.

- La stringa pattern può sovrapporsi nella prima stringa. Per esempio, `"aaa"` contiene `"aa"` due volte. Le prime due `a` e le seconde due.

# --hints--

`count('abcdefg', 'def')` dovrebbe restituire `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('abcdefg', 'def'), 1)`)
}})
```

`count('hello', 'world')` dovrebbe restituire `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('hello', 'world'), 0)`)
}})
```

`count('mississippi', 'iss')` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('mississippi', 'iss'), 2)`)
}})
```

`count('she sells seashells by the seashore', 'sh')` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('she sells seashells by the seashore', 'sh'), 3)`)
}})
```

`count('101010101010101010101', '101')` dovrebbe restituire `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('101010101010101010101', '101'), 10)`)
}})
```

# --seed--

## --seed-contents--

```py
def count(text, parameter):

    return text
```

# --solutions--

```py
def count(text, pattern):
    if not pattern:
        return 0
    occurrences = 0

    for i in range(len(text) - len(pattern) + 1):
        if text[i:i+len(pattern)] == pattern:
            occurrences += 1

    return occurrences
```
