---
id: 68cae5b538ff798bbd4da00a
title: "Sfida 72: regolatore del termostato 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

Data la temperatura attuale di una stanza in Fahrenheit e una temperatura obiettivo in Celsius, restituisci una stringa che indichi come regolare la temperatura della stanza in base a questi vincoli:

- Restituisci `"Heat: X degrees Fahrenheit"` se la temperatura attuale è inferiore all'obiettivo. Con `X` che indica il numero di gradi in Fahrenheit da riscaldare per raggiungere l'obiettivo, arrotondato a 1 cifra decimale.
- Restituisci `"Cool: X degrees Fahrenheit"` se la temperatura attuale è superiore all'obiettivo. Con `X` che indica il numero di gradi in Fahrenheit da raffreddare per raggiungere l'obiettivo, arrotondato a 1 cifra decimale.
- Restituisci `"Hold"` se la temperatura attuale è uguale all'obiettivo.

Per convertire Celsius in Fahrenheit, moltiplica la temperatura Celsius per 1,8 e aggiungi 32 al risultato (`F = (C * 1.8) + 32`).

# --hints--

`adjust_thermostat(32, 0)` dovrebbe restituire `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` dovrebbe restituire `"Heat: 7.0 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` dovrebbe restituire `"Cool: 7.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` dovrebbe restituire `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` dovrebbe restituire `"Heat: 12.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
