---
id: 68d2ba1468508398389487ce
title: "Sfida 73: parla con saggezza, devi"
challengeType: 29
dashedName: challenge-73
---

# --description--

Data una frase, restituisci una versione che suoni come un consiglio di un insegnante saggio usando le seguenti regole:

- Le parole sono separate da uno spazio singolo.
- Trova la prima occorrenza di una delle seguenti parole nella frase: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Sposta tutte le parole prima e inclusa quella parola alla fine della frase e:
  - conserva l’ordine delle parole quando le sposti.
  - rendile tutte minuscole.
  - aggiungi una virgola e uno spazio prima di esse.
- Metti in maiuscolo la prima lettera della nuova prima parola della frase.
- Tutte le frasi date finiranno con un segno di punteggiatura singolo. Mantieni la punteggiatura originale della frase e spostala alla fine della nuova frase.
- Restituisci la nuova frase, assicurandoti che ci sia uno spazio singolo tra ogni parola e nessuno spazio all’inizio o alla fine della frase.

Per esempio, data `"You must speak wisely."` restituisci `"Speak wisely, you must."`

# --hints--

`wise_speak("You must speak wisely.")` dovrebbe restituire `"Speak wisely, you must."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You must speak wisely."), "Speak wisely, you must.")`)
}})
```

`wise_speak("You can do it!")` dovrebbe restituire `"Do it, you can!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You can do it!"), "Do it, you can!")`)
}})
```

`wise_speak("Do you think you will complete this?")` dovrebbe restituire `"Complete this, do you think you will?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("Do you think you will complete this?"), "Complete this, do you think you will?")`)
}})
```

`wise_speak("All your base are belong to us.")` dovrebbe restituire `"Belong to us, all your base are."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("All your base are belong to us."), "Belong to us, all your base are.")`)
}})
```

`wise_speak("You have much to learn.")` dovrebbe restituire `"Much to learn, you have."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You have much to learn."), "Much to learn, you have.")`)
}})
```

# --seed--

## --seed-contents--

```py
def wise_speak(sentence):

    return sentence
```

# --solutions--

```py
def wise_speak(sentence):
    triggers = ["have", "must", "are", "will", "can"]
    punctuation = sentence[-1]
    words = sentence[:-1].split(" ")

    index = next(i for i, w in enumerate(words) if w in triggers)

    to_move = [w.lower() for w in words[:index + 1]]
    remaining = words[index + 1:]

    if remaining:
        remaining[0] = remaining[0][0].upper() + remaining[0][1:]

    return " ".join(remaining) + ", " + " ".join(to_move) + punctuation
```
