---
id: 68d30845cc08266018fc46bd
title: "Sfida 77: formattatore di durata"
challengeType: 29
dashedName: challenge-77
---

# --description--

Dato un numero intero di secondi, restituisci una stringa che rappresenta la stessa durata nel formato `"H:MM:SS"`, dove `"H"` è il numero di ore, `"MM"` è il numero di minuti e `"SS"` è il numero di secondi. Restituisci il tempo usando le seguenti regole:

- Secondi: devono essere sempre due cifre.
- Minuti: devono omettere gli zeri iniziali quando non sono necessari. Usa `"0"` se la durata è inferiore a un minuto.
- Ore: devono essere incluse solo se sono maggiori di zero.

# --hints--

`format(500)` dovrebbe restituire `"8:20"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(500), "8:20")`)
}})
```

`format(4000)` dovrebbe restituire `"1:06:40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(4000), "1:06:40")`)
}})
```

`format(1)` dovrebbe restituire `"0:01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(1), "0:01")`)
}})
```

`format(5555)` dovrebbe restituire `"1:32:35"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(5555), "1:32:35")`)
}})
```

`format(99999)` dovrebbe restituire `"27:46:39"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(99999), "27:46:39")`)
}})
```

# --seed--

## --seed-contents--

```py
def format(seconds):

    return seconds
```

# --solutions--

```py
def format(seconds):
    h = seconds // 3600
    m = (seconds % 3600) // 60
    s = seconds % 60

    seconds_str = f"{s:02d}"
    minutes_str = str(m)

    if h > 0:
        hours_str = str(h)
        return f"{hours_str}:{minutes_str.zfill(2)}:{seconds_str}"
    else:
        return f"{minutes_str}:{seconds_str}"
```
