---
id: 68d30845cc08266018fc46bc
title: "Sfida 76: DNA complementare"
challengeType: 29
dashedName: challenge-76
---

# --description--

Data una stringa che rappresenta una sequenza di DNA, restituisci il suo filamento complementare usando le seguenti regole:

- Il DNA è composto dalle lettere `"A"`, `"C"`, `"G"` e `"T"`.
- Le lettere `"A"` e `"T"` si completano a vicenda.
- Le lettere `"C"` e `"G"` si completano a vicenda.

Per esempio, data `"ACGT"`, restituisci `"TGCA"`.

# --hints--

`complementary_dna("ACGT")` dovrebbe restituire `"TGCA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ACGT"), "TGCA")`)
}})
```

`complementary_dna("ATGCGTACGTTAGC")` dovrebbe restituire `"TACGCATGCAATCG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ATGCGTACGTTAGC"), "TACGCATGCAATCG")`)
}})
```

`complementary_dna("GGCTTACGATCGAAG")` dovrebbe restituire `"CCGAATGCTAGCTTC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC")`)
}})
```

`complementary_dna("GATCTAGCTAGGCTAGCTAG")` dovrebbe restituire `"CTAGATCGATCCGATCGATC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC")`)
}})
```

# --seed--

## --seed-contents--

```py
def complementary_dna(strand):

    return strand
```

# --solutions--

```py
def complementary_dna(strand):
    complements = { "A": "T", "T": "A", "C": "G", "G": "C" }
    return "".join(complements[n] for n in strand)
```
