---
id: 68d30fc57588d97fd3027b30
title: "Sfida 79: Navigator"
challengeType: 29
dashedName: challenge-79
---

# --description--

Il 28 ottobre 1994 è stato rilasciato Netscape Navigator, che ha aiutato milioni di persone a esplorare il web agli albori.

Dato un array di comandi del browser eseguiti su Netscape Navigator, restituisci la pagina corrente su cui ti trovi dopo aver eseguito tutti i comandi seguendo queste regole:

- Parti sempre dalla pagina `"Home"`, che non sarà inclusa nell'array dei comandi.
- I comandi validi sono:
  - `"Visit Page"`: dove `"Page"` è il nome della pagina che stai visitando. Per esempio, `"Visit About"` ti porta alla pagina `"About"`. Quando visiti una nuova pagina, assicurati di scartare tutta la cronologia avanti che hai.
  - `"Back"`: ti porta alla pagina precedente nella cronologia o rimane sulla pagina corrente se non ce n'è una.
  - `"Forward"`: ti porta avanti nella cronologia alla pagina da cui sei venuto o rimane sulla pagina corrente se non ce n'è una.

Per esempio, dato `["Visit About Us", "Back", "Forward"]`, restituisci `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` dovrebbe restituire `"About Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Back", "Forward"]), "About Us")`)
}})
```

`navigate(["Forward"])` dovrebbe restituire `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Forward"]), "Home")`)
}})
```

`navigate(["Back"])` dovrebbe restituire `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Back"]), "Home")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery"])` dovrebbe restituire `"Gallery"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery"]), "Gallery")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` dovrebbe restituire `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home")`)
}})
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` dovrebbe restituire `"Contact"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact")`)
}})
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` dovrebbe restituire `"Visit Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate(commands):

    return commands
```

# --solutions--

```py
def navigate(commands):
    history = ["Home"]
    current_index = 0

    for command in commands:
        if command.startswith("Visit "):
            history = history[:current_index + 1]
            history.append(command[6:])
            current_index += 1
        elif command == "Back" and current_index > 0:
            current_index -= 1
        elif command == "Forward" and current_index < len(history) - 1:
            current_index += 1

    return history[current_index]
```
