---
id: 68e39ed6106dac2f0a98fd63
title: "Sfida 81: n-esimo numero primo"
challengeType: 29
dashedName: challenge-81
---

# --description--

Un numero primo è un numero intero positivo maggiore di 1 che è divisibile solo per 1 e per se stesso. I primi cinque numeri primi sono `2`, `3`, `5`, `7` e `11`.

Dato un numero intero positivo `n`, restituisci il `n`-esimo numero primo. Per esempio, dato `5` restituisci il 5° numero primo: `11`.

# --hints--

`nth_prime(5)` dovrebbe restituire `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(5), 11)`)
}})
```

`nth_prime(10)` dovrebbe restituire `29`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(10), 29)`)
}})
```

`nth_prime(16)` dovrebbe restituire `53`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(16), 53)`)
}})
```

`nth_prime(99)` dovrebbe restituire `523`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(99), 523)`)
}})
```

`nth_prime(1000)` dovrebbe restituire `7919`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(1000), 7919)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_prime(n):

    return n
```

# --solutions--

```py
def nth_prime(n):
    primes = []
    num = 2

    while len(primes) < n:
        is_prime = True
        for i in range(2, int(num**0.5) + 1):
            if num % i == 0:
                is_prime = False
                break
        if is_prime:
            primes.append(num)
        num += 1

    return primes[n - 1]
```
