---
id: 68e39ed6106dac2f0a98fd66
title: "Sfida 84: infettati"
challengeType: 29
dashedName: challenge-84
---

# --description--

Il 2 novembre 1988 è stato rilasciato il primo grande worm su internet, che ha infettato circa il 10% dei computer connessi a internet in appena un giorno.

In questa sfida, ti viene dato un numero di giorni trascorsi dal rilascio di un worm su internet e devi determinare quanti computer sono infettati seguendo queste regole:

- Al giorno 0, il primo computer è infettato.
- Ogni giorno successivo, il numero di computer infettati raddoppia.
- Ogni 3° giorno, dopo che il virus si è diffuso, viene applicata una patch che riduce del 20% il numero di computer infettati. Arrotonda il numero di computer riparati al numero intero più vicino superiore.

Per esempio, al:

- Giorno 0: 1 computer è infettato in totale.
- Giorno 1: 2 computer sono infettati in totale.
- Giorno 2: 4 computer sono infettati in totale.
- Giorno 3: 8 computer sono infettati in totale. Poi, applica la patch: 8 infettati * 20% = 1,6 riparati. Arrotonda 1,6 a 2. 8 computer infettati - 2 riparati = 6 computer infettati in totale dopo il giorno 3.

Restituisci il numero totale di computer infettati dopo che è passato il numero di giorni indicato.

# --hints--

`infected(1)` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(1), 2)`)
}})
```

`infected(3)` dovrebbe restituire `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(3), 6)`)
}})
```

`infected(8)` dovrebbe restituire `152`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(8), 152)`)
}})
```

`infected(17)` dovrebbe restituire `39808`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(17), 39808)`)
}})
```

`infected(25)` dovrebbe restituire `5217638`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(25), 5217638)`)
}})
```

# --seed--

## --seed-contents--

```py
def infected(days):

    return days
```

# --solutions--

```py
import math
def infected(days):
    infected = 1
    for day in range(1, days + 1):
        infected *= 2
        if day % 3 == 0:
            patched = math.ceil(infected * 0.2)
            infected -= patched

    return infected
```
