---
id: 68ee9e3066cfd4eb2328e8a5
title: "Sfida 86: ricerca di immagini"
challengeType: 29
dashedName: challenge-86
---

# --description--

Il 4 novembre 2001 Google ha lanciato la sua ricerca di immagini, permettendo alle persone di trovare immagini usando termini di ricerca. In questa sfida, imiterai la ricerca di immagini.

Dato un array di nomi di immagini e un termine di ricerca, restituisci un array di nomi di immagini che contengono il termine di ricerca.

- Ignora le maiuscole e minuscole quando confronti i termini di ricerca.
- Restituisci le immagini nello stesso ordine in cui appaiono nell'array di input.

# --hints--

`image_search(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` dovrebbe restituire `["dog.png"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"])`)
}})
```

`image_search(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` dovrebbe restituire `["Sunset.jpg", "sunflower.jpeg"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"])`)
}})
```

`image_search(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` dovrebbe restituire `["Moon.png", "stars.png"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"])`)
}})
```

`image_search(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` dovrebbe restituire `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"])`)
}})
```

# --seed--

## --seed-contents--

```py
def image_search(images, term):

    return images
```

# --solutions--

```py
def image_search(images, term):
    lower_term = term.lower()
    return [img for img in images if lower_term in img.lower()]
```
