---
id: 68ee9e3066cfd4eb2328e8a8
title: "Sfida 89: contare le carte"
challengeType: 29
dashedName: challenge-89
---

# --description--

Un mazzo standard di carte da gioco ha 13 carte uniche per ogni seme. Dato un numero intero che rappresenta il numero di carte da pescare dal mazzo, restituisci il numero di combinazioni uniche di carte che puoi ottenere.

- L'ordine non conta. Pescare la carta A e poi la carta B è la stessa cosa che pescare la carta B e poi la carta A.

Per esempio, dato `52`, restituisci `1`. C'è solo una combinazione di 52 carte da pescare da un mazzo di 52 carte. E dato `2`, restituisci `1326`. Ci sono 1326 combinazioni di carte che puoi ottenere pescando 2 carte dal mazzo.

# --hints--

`combinations(52)` dovrebbe restituire `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` dovrebbe restituire `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` dovrebbe restituire `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` dovrebbe restituire `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` dovrebbe restituire `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` dovrebbe restituire `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```
