---
id: 68ee9e3066cfd4eb2328e8a7
title: "Sfida 88: trovare il giorno della settimana"
challengeType: 29
dashedName: challenge-88
---

# --description--

Data una stringa che rappresenta una data nel formato `YYYY-MM-DD`, restituisci il giorno della settimana.

I giorni validi da restituire sono:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Assicurati di ignorare i fusi orari.

# --hints--

`get_weekday("2025-11-06")` dovrebbe restituire `Thursday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2025-11-06"), "Thursday")`)
}})
```

`get_weekday("1999-12-31")` dovrebbe restituire `Friday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("1999-12-31"), "Friday")`)
}})
```

`get_weekday("1111-11-11")` dovrebbe restituire `Saturday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("1111-11-11"), "Saturday")`)
}})
```

`get_weekday("2112-12-21")` dovrebbe restituire `Wednesday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2112-12-21"), "Wednesday")`)
}})
```

`get_weekday("2345-10-01")` dovrebbe restituire `Monday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2345-10-01"), "Monday")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_weekday(date_string):

    return date_string
```

# --solutions--

```py
import datetime
def get_weekday(date_string):
    days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"]
    year, month, day = map(int, date_string.split("-"))
    date = datetime.date(year, month, day)

    return days[date.weekday() % 7 + 1 if date.weekday() != 6 else 0]
```
