---
id: 68ee9e3066cfd4eb2328e8a6
title: "Sfida 87: generatore di matrici"
challengeType: 29
dashedName: challenge-87
---

# --description--

Dato un numero di righe e un numero di colonne, restituisci una matrice (un array di array) riempita di zeri (`0`) della dimensione specificata.

Per esempio, dati `2` e `3`, restituisci:

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`build_matrix(2, 3)` dovrebbe restituire `[[0, 0, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(2, 3), [[0, 0, 0], [0, 0, 0]])`)
}})
```

`build_matrix(3, 2)` dovrebbe restituire `[[0, 0], [0, 0], [0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(3, 2), [[0, 0], [0, 0], [0, 0]])`)
}})
```

`build_matrix(4, 3)` dovrebbe restituire `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]])`)
}})
```

`build_matrix(9, 1)` dovrebbe restituire `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def build_matrix(rows, cols):

    return rows
```

# --solutions--

```py
def build_matrix(rows, cols):

    return [[0 for _ in range(cols)] for _ in range(rows)]
```
