---
id: 68f6587287ad1f4ad39b0c82
title: "Sfida 96: è il weekend?"
challengeType: 29
dashedName: challenge-96
---

# --description--

Dato un giorno nel formato `"YYYY-MM-DD"`, restituisci il numero di giorni mancanti al weekend.

- Il weekend inizia il sabato.
- Se il giorno dato è sabato o domenica, restituisci `"It's the weekend!"`.
- Altrimenti, restituisci `"X days until the weekend."`, dove `X` è il numero di giorni fino al sabato.
- Se `X` è `1`, usa `"day"` (singolare) invece di `"days"` (plurale).
- Assicurati che il calcolo ignori il fuso orario locale.

# --hints--

`days_until_weekend("2025-11-14")` dovrebbe restituire `"1 day until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-11-14"), "1 day until the weekend.")`)
}})
```

`days_until_weekend("2025-01-01")` dovrebbe restituire `"3 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-01-01"), "3 days until the weekend.")`)
}})
```

`days_until_weekend("2025-12-06")` dovrebbe restituire `"It's the weekend!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-12-06"), "It's the weekend!")`)
}})
```

`days_until_weekend("2026-01-27")` dovrebbe restituire `"4 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-01-27"), "4 days until the weekend.")`)
}})
```

`days_until_weekend("2026-09-07")` dovrebbe restituire `"5 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-09-07"), "5 days until the weekend.")`)
}})
```

`days_until_weekend("2026-11-29")` dovrebbe restituire `"It's the weekend!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-11-29"), "It's the weekend!")`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_weekend(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime

def days_until_weekend(date_string):
    date = datetime.strptime(date_string, "%Y-%m-%d").date()
    day_of_week = date.weekday()

    if day_of_week == 5 or day_of_week == 6:
        return "It's the weekend!"

    days_until_saturday = 5 - day_of_week
    return f"{days_until_saturday} day{'s' if days_until_saturday != 1 else ''} until the weekend."
```
