---
id: 68f6587287ad1f4ad39b0c84
title: "Sfida 98: conteggio dei rettangoli"
challengeType: 29
dashedName: challenge-98
---

# --description--

Dati due numeri interi positivi che rappresentano la larghezza e l'altezza di un rettangolo, determina quanti rettangoli possono essere contenuti in quello dato.

- Conta solo i rettangoli con larghezza e altezza intere.

Per esempio, dati `1` e `3`, restituisci `6`. Tre rettangoli 1x1, due rettangoli 1x2 e un rettangolo 1x3.

# --hints--

`count_rectangles(1, 3)` dovrebbe restituire `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 3), 6)`)
}})
```

`count_rectangles(3, 2)` dovrebbe restituire `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(3, 2), 18)`)
}})
```

`count_rectangles(1, 2)` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 2), 3)`)
}})
```

`count_rectangles(5, 4)` dovrebbe restituire `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(5, 4), 150)`)
}})
```

`count_rectangles(11, 19)` dovrebbe restituire `12540`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(11, 19), 12540)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_rectangles(width, height):

    return width
```

# --solutions--

```py
def count_rectangles(width, height):
    total = 0
    for w in range(1, width + 1):
        for h in range(1, height + 1):
            total += (width - w + 1) * (height - h + 1)
    return total
```
