---
id: 68f6587287ad1f4ad39b0c7d
title: "Sfida 91: ricerca di parole"
challengeType: 29
dashedName: challenge-91
---

# --description--

Data una matrice (un array di array) di singole lettere e una parola da trovare, restituisci gli indici di inizio e fine della parola nella matrice.

- La matrice data sarà riempita con tutte lettere minuscole (`a-z`).
- La parola da trovare sarà sempre presente nella matrice esattamente una volta.
- La parola da trovare sarà sempre in linea retta in una di queste direzioni:
  - da sinistra a destra
  - da destra a sinistra
  - dall’alto verso il basso
  - dal basso verso l’alto

Per esempio, data la matrice:

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

E la parola `"cat"`, restituisci:

```md
[[0, 1], [2, 1]]
```

Dove `[0, 1]` sono gli indici per il `"c"` (inizio della parola), e `[2, 1]` sono gli indici per il `"t"` (fine della parola).

# --hints--

`find_word([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")` dovrebbe restituire `[[0, 1], [2, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_word([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]])`)
}})
```

`find_word([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")` dovrebbe restituire `[[0, 0], [0, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_word([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]])`)
}})
```

`find_word([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")` dovrebbe restituire `[[3, 3], [0, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_word([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]])`)
}})
```

`find_word([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")` dovrebbe restituire `[[1, 3], [1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_word([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_word(matrix, word):

    return matrix
```

# --solutions--

```py
def find_word(matrix, word):
    rows = len(matrix)
    cols = len(matrix[0])
    length = len(word)

    directions = [
        (0, 1),
        (0, -1),
        (1, 0),
        (-1, 0)
    ]

    for r in range(rows):
        for c in range(cols):
            for dr, dc in directions:
                match = True
                for i in range(length):
                    nr = r + dr * i
                    nc = c + dc * i
                    if nr < 0 or nr >= rows or nc < 0 or nc >= cols or matrix[nr][nc] != word[i]:
                        match = False
                        break
                if match:
                    return [[r, c], [r + dr * (length - 1), c + dc * (length - 1)]]
```
