---
id: 68f6587287ad1f4ad39b0c7c
title: "Sfida 90: limite di caratteri"
challengeType: 29
dashedName: challenge-90
---

# --description--

In questa sfida, ti viene data una stringa e devi determinare se può essere inserita in un post sui social media. Restituisci le seguenti stringhe in base alle regole indicate:

- `"short post"` se rientra nel limite di 40 caratteri.
- `"long post"` se supera i 40 caratteri ma rientra nel limite di 80 caratteri.
- `"invalid post"` se è troppo lunga per rientrare in entrambi i limiti.

# --hints--

`can_post("Hello world")` dovrebbe restituire `"short post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("Hello world"), "short post")`)
}})
```

`can_post("This is a longer message but still under eighty characters.")` dovrebbe restituire `"long post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This is a longer message but still under eighty characters."), "long post")`)
}})
```

`can_post("This message is too long to fit into either of the character limits for a social media post.")` dovrebbe restituire `"invalid post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This message is too long to fit into either of the character limits for a social media post."), "invalid post")`)
}})
```

# --seed--

## --seed-contents--

```py
def can_post(message):

    return message
```

# --solutions--

```py
def can_post(message):
    if len(message) <= 40:
        return "short post"
    elif len(message) <= 80:
        return "long post"
    else:
        return "invalid post"
```
