---
id: 68f6587287ad1f4ad39b0c80
title: "Sfida 94: generatore di firme email"
challengeType: 29
dashedName: challenge-94
---

# --description--

Date delle stringhe per il nome di una persona, il titolo e l'azienda, restituisci una firma email come una singola stringa usando le seguenti regole:

- Il nome deve apparire per primo, preceduto da un prefisso che dipende dalla prima lettera del nome. Per i nomi che iniziano con (indipendentemente dal maiuscolo o minuscolo):
  - `A-I`: usa `>>` come prefisso.
  - `J-R`: usa `--` come prefisso.
  - `S-Z`: usa `::` come prefisso.
- Dopo il nome deve seguire una virgola e uno spazio (`, `).
- Il titolo e l'azienda devono seguire la virgola e lo spazio, separati da `" at "` (con spazi intorno).

Per esempio, date `"Quinn Waverly"`, `"Founder and CEO"` e `"TechCo"` restituisci `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generate_signature("Quinn Waverly", "Founder and CEO", "TechCo")` dovrebbe restituire `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo")`)
}})
```

`generate_signature("Alice Reed", "Engineer", "TechCo")` dovrebbe restituire `">>Alice Reed, Engineer at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo")`)
}})
```

`generate_signature("Tina Vaughn", "Developer", "example.com")` dovrebbe restituire `"::Tina Vaughn, Developer at example.com"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com")`)
}})
```

`generate_signature("B. B.", "Product Tester", "AcmeCorp")` dovrebbe restituire `">>B. B., Product Tester at AcmeCorp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp")`)
}})
```

`generate_signature("windstorm", "Cloud Architect", "Atmospheronics")` dovrebbe restituire `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_signature(name, title, company):

    return name
```

# --solutions--

```py
def generate_signature(name, title, company):
    first_letter = name[0].upper()
    if first_letter in "ABCDEFGHI":
        prefix = ">>"
    elif first_letter in "JKLMNOPQR":
        prefix = "--"
    else:
        prefix = "::"

    return f"{prefix}{name}, {title} at {company}"
```
