---
id: 68f6587287ad1f4ad39b0c7e
title: "Sfida 92: estrattore di estensioni"
challengeType: 29
dashedName: challenge-92
---

# --description--

Data una stringa che rappresenta un nome di file, restituisci l'estensione del file.

- L'estensione è la parte del nome del file che segue l'ultimo punto (`.`).
- Se il nome del file non contiene un punto o termina con un punto, restituisci `"none"`.
- L'estensione deve essere restituita così com'è, mantenendo il caso.

# --hints--

`get_extension("document.txt")` dovrebbe restituire `"txt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("document.txt"), "txt")`)
}})
```

`get_extension("README")` dovrebbe restituire `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("README"), "none")`)
}})
```

`get_extension("image.PNG")` dovrebbe restituire `"PNG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("image.PNG"), "PNG")`)
}})
```

`get_extension(".gitignore")` dovrebbe restituire `"gitignore"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension(".gitignore"), "gitignore")`)
}})
```

`get_extension("archive.tar.gz")` dovrebbe restituire `"gz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("archive.tar.gz"), "gz")`)
}})
```

`get_extension("final.draft.")` dovrebbe restituire `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("final.draft."), "none")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_extension(filename):

    return filename
```

# --solutions--

```py
def get_extension(filename):
    last_dot = filename.rfind('.')
    
    if last_dot == -1 or last_dot == len(filename) - 1:
        return "none"
    
    return filename[last_dot + 1:]
```
