---
id: 68f6587287ad1f4ad39b0c7f
title: "Sfida 93: vocali e consonanti"
challengeType: 29
dashedName: challenge-93
---

# --description--

Data una stringa, restituisci un array con il numero di vocali e il numero di consonanti nella stringa.

- Le vocali sono `a`, `e`, `i`, `o`, `u` in qualsiasi caso.
- Le consonanti sono tutte le altre lettere in qualsiasi caso.
- Ignora tutti i caratteri che non sono lettere.

Per esempio, data `"Hello World"`, restituisci `[3, 7]`.

# --hints--

`count("Hello World")` dovrebbe restituire `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("JavaScript")` dovrebbe restituire `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("JavaScript"), [3, 7])`)
}})
```

`count("Python")` dovrebbe restituire `[1, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Python"), [1, 5])`)
}})
```

`count("freeCodeCamp")` dovrebbe restituire `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("freeCodeCamp"), [5, 7])`)
}})
```

`count("Hello, World!")` dovrebbe restituire `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("The quick brown fox jumps over the lazy dog.")` dovrebbe restituire `[11, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

# --seed--

## --seed-contents--

```py
def count(s):

    return s
```

# --solutions--

```py
def count(s):
    vowels = "aeiou"
    consonants = "bcdfghjklmnpqrstvwxyz"
    v = 0
    c = 0

    for char in s:
        ch = char.lower()
        if ch in vowels:
            v += 1
        if ch in consonants:
            c += 1

    return [v, c]
```
