---
id: 68ffb91507a5b645769328cc
title: "Sfida 109: Quanti anni ho di nuovo?"
challengeType: 29
dashedName: challenge-109
---

# --description--

Dato il `YYYY-MM-DD` della data di nascita di una persona, restituisci l'età della persona al 27 novembre 2025.

- Considera che tutte le date di nascita sono valide e precedono il 27 novembre 2025.
- Restituisci l'età come numero intero.
- Assicurati di considerare se la persona ha già compiuto gli anni nel 2025.

# --hints--

`calculate_age("2000-11-20")` dovrebbe restituire `25`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-11-20"), 25)`)
}})
```

`calculate_age("2000-12-01")` dovrebbe restituire `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-12-01"), 24)`)
}})
```

`calculate_age("2014-10-25")` dovrebbe restituire `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2014-10-25"), 11)`)
}})
```

`calculate_age("1994-01-06")` dovrebbe restituire `31`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-01-06"), 31)`)
}})
```

`calculate_age("1994-12-14")` dovrebbe restituire `30`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-12-14"), 30)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_age(birthday):

    return birthday
```

# --solutions--

```py
from datetime import datetime
def calculate_age(birthday):
    today = datetime(2025, 11, 27)
    birth_date = datetime.strptime(birthday, "%Y-%m-%d")

    age = today.year - birth_date.year

    if (today.month, today.day) < (birth_date.month, birth_date.day):
        age -= 1

    return age
```
