---
id: 68ffb91507a5b645769328c6
title: "Sfida 103: mcm"
challengeType: 29
dashedName: challenge-103
---

# --description--

Dati due numeri interi, restituisci il minimo comune multiplo (mcm) dei due numeri.

Il mcm di due numeri è il più piccolo numero intero positivo che è multiplo di entrambi i numeri. Per esempio, dati `4` e `6`, restituisci `12` perché:

- I multipli di `4` sono `4`, `8`, `12` e così via.
- I multipli di `6` sono `6`, `12`, `18` e così via.
- `12` è il numero più piccolo che è multiplo di entrambi.

# --hints--

`lcm(4, 6)` dovrebbe restituire `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(4, 6), 12)`)
}})
```

`lcm(9, 6)` dovrebbe restituire `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(9, 6), 18)`)
}})
```

`lcm(10, 100)` dovrebbe restituire `100`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(10, 100), 100)`)
}})
```

`lcm(13, 17)` dovrebbe restituire `221`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(13, 17), 221)`)
}})
```

`lcm(45, 70)` dovrebbe restituire `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(45, 70), 630)`)
}})
```

# --seed--

## --seed-contents--

```py
def lcm(a, b):

    return a
```

# --solutions--

```py
def lcm(a, b):
    def gcd(x, y):
        return x if y == 0 else gcd(y, x % y)

    return abs(a * b) // gcd(a, b)
```
