---
id: 68ffb91507a5b645769328cb
title: "Sfida 108: BuzzFizz"
challengeType: 29
dashedName: challenge-108
---

# --description--

Dato un array, determina se è una sequenza FizzBuzz corretta da 1 all'ultimo elemento dell'array. Una sequenza è corretta se:

- I numeri che sono multipli di 3 sono sostituiti da "Fizz"
- I numeri che sono multipli di 5 sono sostituiti da "Buzz"
- I numeri che sono multipli sia di 3 che di 5 sono sostituiti da "FizzBuzz"
- Tutti gli altri numeri rimangono come numeri interi in ordine crescente, a partire da 1.
- L'array deve iniziare da 1 e non deve avere elementi mancanti o in eccesso.

# --hints--

`is_fizz_buzz([1, 2, "Fizz", 4])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4]), True)`)
}})
```

`is_fizz_buzz([1, 2, 3, 4])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, 3, 4]), False)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", 7])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", 7]), False)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"]), False)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"]), False)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"]), False)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_fizz_buzz(sequence):

    return sequence
```

# --solutions--

```py
def is_fizz_buzz(sequence):
    for i, val in enumerate(sequence):
        n = i + 1
        if n % 15 == 0 and val != "FizzBuzz":
            return False
        if n % 3 == 0 and n % 5 != 0 and val != "Fizz":
            return False
        if n % 5 == 0 and n % 3 != 0 and val != "Buzz":
            return False
        if n % 3 != 0 and n % 5 != 0 and val != n:
            return False
    return True
```
