---
id: 68ffb91507a5b645769328c5
title: "Sfida 102: parola più lunga"
challengeType: 29
dashedName: challenge-102
---

# --description--

Data una stringa che rappresenta una frase, restituisci la parola più lunga nella frase.

- Le parole sono separate da uno spazio singolo.
- Solo le lettere (`a-z`, senza distinzione tra maiuscole e minuscole) contano ai fini della lunghezza della parola.
- Se ci sono più parole con la stessa lunghezza, restituisci la prima che appare.
- Restituisci la parola così come appare nella stringa data, con la punteggiatura rimossa.

# --hints--

`longest_word("The quick red fox")` dovrebbe restituire `"quick"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("The quick red fox"), "quick")`)
}})
```

`longest_word("Hello coding challenge.")` dovrebbe restituire `"challenge"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Hello coding challenge."), "challenge")`)
}})
```

`longest_word("Do Try This At Home.")` dovrebbe restituire `"This"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Do Try This At Home."), "This")`)
}})
```

`longest_word("This sentence... has commas, ellipses, and an exclamation point!")` dovrebbe restituire `"exclamation"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation")`)
}})
```

`longest_word("A tie? No way!")` dovrebbe restituire `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("A tie? No way!"), "tie")`)
}})
```

`longest_word("Wouldn't you like to know.")` dovrebbe restituire `"Wouldnt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Wouldn't you like to know."), "Wouldnt")`)
}})
```

# --seed--

## --seed-contents--

```py
def longest_word(sentence):

    return sentence
```

# --solutions--

```py
import re
def longest_word(sentence):
    words = sentence.split(" ")
    longest = ""

    for word in words:
        cleaned = re.sub(r'[^a-zA-Z]', '', word)
        if len(cleaned) > len(longest):
            longest = cleaned

    return longest
```
