---
id: 69162d64f96574d9bb629f02
title: "Sfida 117: differenza simmetrica"
challengeType: 29
dashedName: challenge-117
---

# --description--

Dato due array, restituisci un nuovo array contenente la loro differenza simmetrica.

- La differenza simmetrica tra due insiemi è l'insieme dei valori che compaiono in uno dei due insiemi, ma non in entrambi.
- Restituisci i valori nell'ordine in cui compaiono per la prima volta negli array di input.

# --hints--

`difference([1, 2, 3], [3, 4, 5])` dovrebbe restituire `[1, 2, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5])`)
}})
```

`difference(["a", "b"], ["c", "b"])` dovrebbe restituire `["a", "c"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"])`)
}})
```

`difference([1, "a", 2], [2, "b", "a"])` dovrebbe restituire `[1, "b"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"])`)
}})
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` dovrebbe restituire `[2, 4, 6, 8]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8])`)
}})
```

# --seed--

## --seed-contents--

```py
def difference(arr1, arr2):

    return arr1
```

# --solutions--

```py
def difference(arr1, arr2):
    diff = []

    for v in arr1:
        if v not in arr2 and v not in diff:
            diff.append(v)

    for v in arr2:
        if v not in arr1 and v not in diff:
            diff.append(v)

    return diff
```
