---
id: 69162d64f96574d9bb629efe
title: "Sfida 113: miglia in chilometri"
challengeType: 29
dashedName: challenge-113
---

# --description--

Dato un numero che rappresenta una distanza in miglia, restituisci la distanza equivalente in chilometri.

- L'input sarà sempre un numero non negativo.
- 1 miglio equivale a 1,60934 chilometri.
- Arrotonda il risultato a due cifre decimali.
- Rimuovi gli zeri finali non necessari dal risultato arrotondato.

# --hints--

`convert_to_km(1)` dovrebbe restituire `1.61`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(1), 1.61)`)
}})
```

`convert_to_km(21)` dovrebbe restituire `33.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(21), 33.8)`)
}})
```

`convert_to_km(3.5)` dovrebbe restituire `5.63`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(3.5), 5.63)`)
}})
```

`convert_to_km(0)` dovrebbe restituire `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0), 0)`)
}})
```

`convert_to_km(0.621371)` dovrebbe restituire `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0.621371), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_km(miles):

    return miles
```

# --solutions--

```py
def convert_to_km(miles):
    km = miles * 1.60934
    return round(km, 2)
```
