---
id: 69162d64f96574d9bb629f00
title: "Sfida 115: convertitore di elemento di lista ordinata in markdown"
challengeType: 29
dashedName: challenge-115
---

# --description--

Dato una stringa che rappresenta un elemento di lista ordinata in markdown, restituisci la stringa HTML equivalente.

Un elemento valido di lista ordinata in markdown deve:

- Iniziare con zero o più spazi, seguiti da
- Un numero (1 o superiore) e un punto (`.`), seguiti da
- Almeno uno spazio, e poi
- Il testo dell'elemento della lista.

Se la stringa non ha esattamente questo formato, restituisci `"Invalid format"`. Altrimenti, racchiudi il testo dell'elemento della lista in tag `li` e restituisci la stringa.

Per esempio, dato `"1. My item"`, restituisci `"<li>My item</li>"`.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`convert_list_item("1. My item")` dovrebbe restituire `"<li>My item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` dovrebbe restituire `"<li>Another item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` dovrebbe restituire `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` dovrebbe restituire `"<li>list item text</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` dovrebbe restituire `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` dovrebbe restituire `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```
