---
id: 69162d64f96574d9bb629f03
title: "Sfida 118: formattatore di date"
challengeType: 29
dashedName: challenge-118
---

# --description--

Data una data nel formato `"Month day, year"`, restituisci la data nel formato `"YYYY-MM-DD"`.

- Il mese fornito sarà il nome completo del mese in inglese. Per esempio: `"January"`, `"February"`, ecc.
- Nel valore di ritorno, aggiungi zeri iniziali al mese e al giorno se necessario per assicurarti che abbiano due cifre.

Per esempio, data `"December 6, 2025"`, restituisci `"2025-12-06"`.

# --hints--

`format_date("December 6, 2025")` dovrebbe restituire `"2025-12-06"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("December 6, 2025"), "2025-12-06")`)
}})
```

`format_date("January 1, 2000")` dovrebbe restituire `"2000-01-01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("January 1, 2000"), "2000-01-01")`)
}})
```

`format_date("November 11, 1111")` dovrebbe restituire `"1111-11-11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("November 11, 1111"), "1111-11-11")`)
}})
```

`format_date("September 7, 512")` dovrebbe restituire `"512-09-07"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("September 7, 512"), "512-09-07")`)
}})
```

`format_date("May 4, 1950")` dovrebbe restituire `"1950-05-04"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("May 4, 1950"), "1950-05-04")`)
}})
```

`format_date("February 29, 1992")` dovrebbe restituire `"1992-02-29"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("February 29, 1992"), "1992-02-29")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_date(date_string):

    return date_string
```

# --solutions--

```py
def format_date(date_string):
    months = {
        "January": "01",
        "February": "02",
        "March": "03",
        "April": "04",
        "May": "05",
        "June": "06",
        "July": "07",
        "August": "08",
        "September": "09",
        "October": "10",
        "November": "11",
        "December": "12"
    }

    month, day, year = date_string.replace(",", "").split(" ")
    day = day.zfill(2)

    return f"{year}-{months[month]}-{day}"
```
