---
id: 69162d64f96574d9bb629eff
title: "Sfida 114: da camel case a snake case"
challengeType: 29
dashedName: challenge-114
---

# --description--

Data una stringa in camel case, restituisci la versione in snake case della stringa seguendo queste regole:

- La stringa di input conterrà solo lettere (`A-Z` e `a-z`) e inizierà sempre con una lettera minuscola.
- Ogni lettera maiuscola nella stringa camel case indica l'inizio di una nuova parola.
- Converti tutte le lettere in minuscolo.
- Separa le parole con un underscore (`_`).

# --hints--

`to_snake("helloWorld")` dovrebbe restituire `"hello_world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` dovrebbe restituire `"my_variable_name"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` dovrebbe restituire `"freecodecamp_daily_challenges"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
