---
id: 69162d64f96574d9bb629f01
title: "Sfida 116: conteggio delle permutazioni"
challengeType: 29
dashedName: challenge-116
---

# --description--

Data una stringa, restituisci il numero di permutazioni distinte che si possono formare con i suoi caratteri.

- Una permutazione è qualsiasi riordinamento dei caratteri nella stringa.
- Non contare permutazioni duplicate.
- Se la stringa contiene caratteri ripetuti, le disposizioni ripetute devono essere contate una sola volta.
- La stringa conterrà solo lettere (`A-Z`, `a-z`).

Per esempio, data `"abb"`, restituisci `3` perché ci sono tre modi unici per disporre le lettere: `"abb"`, `"bab"` e `"bba"`.

# --hints--

`count_permutations("abb")` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` dovrebbe restituire `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` dovrebbe restituire `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` dovrebbe restituire `39916800`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
