---
id: 69162d64f96574d9bb629efd
title: "Sfida 112: rilevatore di intelligenza artificiale"
challengeType: 29
dashedName: challenge-112
---

# --description--

La sfida di oggi è ispirata al rilascio di ChatGPT il 30 novembre 2022.

Data una stringa composta da una o più frasi, determina se è probabile che sia stata generata da un'intelligenza artificiale usando le seguenti regole:

- Contiene due o più trattini (`-`).
- Contiene due o più insiemi di parentesi (`()`). Il testo può essere all'interno delle parentesi.
- Contiene tre o più parole con 7 o più lettere.

- Le parole sono separate da uno spazio singolo e consistono solo di lettere (`A-Z`). Non includere la punteggiatura o altri caratteri non alfabetici come parte di una parola.

Se la frase data soddisfa una qualsiasi delle regole sopra, restituisci `"AI"`, altrimenti restituisci `"Human"`.

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` dovrebbe restituire `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` dovrebbe restituire `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` dovrebbe restituire `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` dovrebbe restituire `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` dovrebbe restituire `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```
