---
id: 69162d64f96574d9bb629f04
title: "Sfida 119: compressione di stringhe"
challengeType: 29
dashedName: challenge-119
---

# --description--

Data una stringa frase, restituisci una versione compressa della frase in cui le parole duplicate consecutive sono sostituite dalla parola seguita dal numero di volte che si ripete tra parentesi tonde.

- Vengono compresse solo le parole duplicate consecutive.
- Le parole sono separate da spazi singoli.

Per esempio, data `"yes yes yes please"`, restituisci `"yes(3) please"`.

# --hints--

`compress_string("yes yes yes please")` dovrebbe restituire `"yes(3) please"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("yes yes yes please"), "yes(3) please")`)
}})
```

`compress_string("I have have have apples")` dovrebbe restituire `"I have(3) apples"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("I have have have apples"), "I have(3) apples")`)
}})
```

`compress_string("one one three and to the the the the")` dovrebbe restituire `"one(2) three and to the(4)"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("one one three and to the the the the"), "one(2) three and to the(4)")`)
}})
```

`compress_string("route route route route route route tee tee tee tee tee tee")` dovrebbe restituire `"route(6) tee(6)"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)")`)
}})
```

# --seed--

## --seed-contents--

```py
def compress_string(sentence):

    return sentence
```

# --solutions--

```py
def compress_string(sentence):
    words = sentence.split(" ")
    result = []
    count = 1

    for i in range(len(words)):
        if i < len(words) - 1 and words[i] == words[i + 1]:
            count += 1
        else:
            if count > 1:
                result.append(f"{words[i]}({count})")
            else:
                result.append(words[i])
            count = 1

    return " ".join(result)
```
