---
id: 691b559495c5cb5a37b9b480
title: "Sfida 120: da libbre a chilogrammi"
challengeType: 29
dashedName: challenge-120
---

# --description--

Dato un peso in libbre come numero, restituisci la stringa `"(lbs) pounds equals (kgs) kilograms."`.

- Sostituisci `"(lbs)"` con il numero di input.
- Sostituisci `"(kgs)"` con l'input convertito in chilogrammi, arrotondato a due decimali e includendo sempre due cifre decimali nel valore.
- 1 libbra equivale a 0,453592 chilogrammi.
- Se l'input è `1`, usa `"pound"` invece di `"pounds"`.
- Se il valore convertito è `1`, usa `"kilogram"` invece di `"kilograms"`.

# --hints--

`convert_to_kgs(1)` dovrebbe restituire `"1 pound equals 0.45 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(1), "1 pound equals 0.45 kilograms.")`)
}})
```

`convert_to_kgs(0)` dovrebbe restituire `"0 pounds equals 0.00 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(0), "0 pounds equals 0.00 kilograms.")`)
}})
```

`convert_to_kgs(100)` dovrebbe restituire `"100 pounds equals 45.36 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(100), "100 pounds equals 45.36 kilograms.")`)
}})
```

`convert_to_kgs(2.5)` dovrebbe restituire `"2.5 pounds equals 1.13 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.5), "2.5 pounds equals 1.13 kilograms.")`)
}})
```

`convert_to_kgs(2.20462)` dovrebbe restituire `"2.20462 pounds equals 1.00 kilogram."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.20462), "2.20462 pounds equals 1.00 kilogram.")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_kgs(lbs):

    return lbs
```

# --solutions--

```py
def convert_to_kgs(lbs):
    KG_PER_POUND = 0.453592
    kgs = f"{lbs * KG_PER_POUND:.2f}"

    pound_word = "pound" if lbs == 1 else "pounds"
    kilogram_word = "kilogram" if kgs == "1.00" else "kilograms"

    return f"{lbs} {pound_word} equals {kgs} {kilogram_word}."
```
