---
id: 691b559495c5cb5a37b9b489
title: "Sfida 129: parser di citazioni in blocco Markdown"
challengeType: 29
dashedName: challenge-129
---

# --description--

Data una stringa che include una citazione in blocco in Markdown, restituisci la stringa HTML equivalente.

Una citazione in blocco in Markdown è qualsiasi riga che:

- Inizia con zero o più spazi
- Seguita da un segno maggiore di (`>`)
- Poi, uno o più spazi
- E infine, il testo della citazione in blocco.

Restituisci il testo della citazione in blocco racchiuso dai tag HTML `blockquote` di apertura e chiusura.

Per esempio, data `"> This is a quote"`, restituisci `<blockquote>This is a quote</blockquote>`.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`parse_blockquote("> This is a quote")` dovrebbe restituire `"<blockquote>This is a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>")`)
}})
```

`parse_blockquote(" > This is also a quote")` dovrebbe restituire `"<blockquote>This is also a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>")`)
}})
```

`parse_blockquote("  >    So  Is  This")` dovrebbe restituire `"<blockquote>So  Is  This</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```
