---
id: 691b559495c5cb5a37b9b485
title: "Sfida 125: gioco della vita"
challengeType: 29
dashedName: challenge-125
---

# --description--

Data una matrice (array di array) che rappresenta lo stato attuale nel gioco della vita di Conway, restituisci il prossimo stato della matrice usando queste regole:

- Ogni cella è o `1` (viva) o `0` (morta).
- I vicini di una cella sono fino a otto celle circostanti (verticalmente, orizzontalmente e diagonalmente).
- Le celle ai bordi hanno meno di otto vicini.

Regole per aggiornare ogni cella:

- Qualsiasi cella viva con meno di due vicini vivi muore (sottopopolazione).
- Qualsiasi cella viva con due o tre vicini vivi sopravvive.
- Qualsiasi cella viva con più di tre vicini vivi muore (sovrappopolazione).
- Qualsiasi cella morta con esattamente tre vicini vivi diventa viva (riproduzione).

Per esempio, dato:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

restituisci:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Ogni cella si aggiorna in base al numero di vicini vivi. Per esempio, `[0][0]` rimane morta (2 vicini vivi), `[0][1]` rimane viva (2 vicini vivi), `[0][2]` muore (3 vicini vivi), e così via.

# --hints--

`game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` dovrebbe restituire `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]])`)
}})
```

`game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` dovrebbe restituire `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]])`)
}})
```

`game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` dovrebbe restituire `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]])`)
}})
```

`game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` dovrebbe restituire `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def game_of_life(grid):

    return grid
```

# --solutions--

```py
def game_of_life(grid):
    rows, cols = len(grid), len(grid[0])

    def count_live_neighbors(r, c):
        count = 0
        for i in range(r-1, r+2):
            for j in range(c-1, c+2):
                if 0 <= i < rows and 0 <= j < cols and (i != r or j != c):
                    count += grid[i][j]
        return count

    next_state = [row[:] for row in grid]

    for r in range(rows):
        for c in range(cols):
            live_neighbors = count_live_neighbors(r, c)

            if grid[r][c] == 1:
                if live_neighbors < 2 or live_neighbors > 3:
                    next_state[r][c] = 0
            else:
                if live_neighbors == 3:
                    next_state[r][c] = 1

    return next_state
```
