---
id: 691b559495c5cb5a37b9b486
title: "Sfida 126: metti in maiuscolo"
challengeType: 29
dashedName: challenge-126
---

# --description--

Data una stringa title, restituisci una nuova stringa formattata in maiuscolo secondo le seguenti regole:

- Metti in maiuscolo la prima lettera di ogni parola.
- Rendi tutte le altre lettere di ogni parola minuscole.
- Le parole sono sempre separate da uno spazio singolo.

# --hints--

`title_case("hello world")` dovrebbe restituire `"Hello World"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("hello world"), "Hello World")`)
}})
```

`title_case("the quick brown fox")` dovrebbe restituire `"The Quick Brown Fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("the quick brown fox"), "The Quick Brown Fox")`)
}})
```

`title_case("JAVASCRIPT AND PYTHON")` dovrebbe restituire `"Javascript And Python"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("JAVASCRIPT AND PYTHON"), "Javascript And Python")`)
}})
```

`title_case("AvOcAdO tOAst fOr brEAkfAst")` dovrebbe restituire `"Avocado Toast For Breakfast"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast")`)
}})
```

# --seed--

## --seed-contents--

```py
def title_case(title):

    return title
```

# --solutions--

```py
def title_case(title):
    return " ".join(
        w[:1].upper() + w[1:].lower()
        for w in title.split(" ")
    )
```
