---
id: 691b559495c5cb5a37b9b487
title: "Sfida 127: controllo della velocità"
challengeType: 29
dashedName: challenge-127
---

# --description--

Dato la velocità a cui stai viaggiando in miglia all'ora (MPH) e un limite di velocità in chilometri all'ora (KPH), determina se stai superando il limite e se riceverai un avviso o una multa.

- 1 miglio equivale a 1,60934 chilometri.
- Se stai viaggiando a una velocità minore o uguale al limite, restituisci `"Not Speeding"`.
- Se stai viaggiando fino a 5 KPH oltre il limite, restituisci `"Warning"`.
- Se stai viaggiando a più di 5 KPH oltre il limite, restituisci `"Ticket"`.

# --hints--

`speed_check(30, 70)` dovrebbe restituire `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` dovrebbe restituire `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` dovrebbe restituire `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` dovrebbe restituire `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` dovrebbe restituire `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` dovrebbe restituire `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```
