---
id: 691b559495c5cb5a37b9b484
title: "Sfida 124: aggiornamento dell'inventario"
challengeType: 29
dashedName: challenge-124
---

# --description--

Dato un array bidimensionale che rappresenta l'inventario del tuo archivio e un altro array bidimensionale che rappresenta una spedizione ricevuta, restituisci il tuo inventario aggiornato.

- Ogni elemento negli array avrà il formato: `[quantity, "item"]`, dove `quantity` è un intero e `"item"` è una stringa.
- Aggiorna gli elementi nell'inventario aggiungendo la quantità di qualsiasi elemento corrispondente dalla spedizione.
- Se un elemento ricevuto non esiste nell'inventario attuale, aggiungilo come nuova voce alla fine dell'inventario.
- Restituisci l'inventario nell'ordine in cui è stato fornito, con i nuovi elementi alla fine nell'ordine in cui appaiono nella spedizione.

Per esempio, dato un inventario di `[[2, "apples"], [5, "bananas"]]` e una spedizione di `[[1, "apples"], [3, "bananas"]]`, restituisci `[[3, "apples"], [8, "bananas"]]`.

# --hints--

`update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])` dovrebbe restituire `[[3, "apples"], [8, "bananas"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]])`)
}})
```

`update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])` dovrebbe restituire `[[3, "apples"], [8, "bananas"], [4, "oranges"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]])`)
}})
```

`update_inventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])` dovrebbe restituire `[[10, "apples"], [30, "bananas"], [20, "oranges"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]])`)
}})
```

`update_inventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])` dovrebbe restituire `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def update_inventory(inventory, shipment):

    return inventory
```

# --solutions--

```py
def update_inventory(inventory, shipment):
    inventory_map = {item: i for i, (qty, item) in enumerate(inventory)}

    for qty, item in shipment:
        if item in inventory_map:
            index = inventory_map[item]
            inventory[index][0] += qty
        else:
            inventory.append([qty, item])
            inventory_map[item] = len(inventory) - 1

    return inventory
```
