---
id: 6925e2068081f40f549ced1d
title: "Sfida 139: carta, forbice, sasso"
challengeType: 29
dashedName: challenge-139
---

# --description--

Date due stringhe, la prima che rappresenta il giocatore 1 e la seconda il giocatore 2, determina il vincitore di una partita di carta, forbice, sasso.

- Le stringhe di input saranno sempre `"Rock"`, `"Paper"` o `"Scissors"`.
- `"Rock"` batte `"Scissors"`.
- `"Paper"` batte `"Rock"`.
- `"Scissors"` batte `"Paper"`.

Restituisci:

- `"Player 1 wins"` se vince il giocatore 1.
- `"Player 2 wins"` se vince il giocatore 2.
- `"Tie"` se entrambi i giocatori scelgono la stessa opzione.

# --hints--

`rock_paper_scissors("Rock", "Rock")` dovrebbe restituire `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Rock"), "Tie")`)
}})
```

`rock_paper_scissors("Rock", "Paper")` dovrebbe restituire `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Paper"), "Player 2 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Paper")` dovrebbe restituire `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Paper"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Rock", "Scissors")` dovrebbe restituire `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Scissors"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Scissors")` dovrebbe restituire `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Scissors"), "Tie")`)
}})
```

`rock_paper_scissors("Scissors", "Rock")` dovrebbe restituire `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Rock"), "Player 2 wins")`)
}})
```

# --seed--

## --seed-contents--

```py
def rock_paper_scissors(player1, player2):

    return player1
```

# --solutions--

```py
def rock_paper_scissors(player1, player2):
    if player1 == player2:
        return "Tie"
    
    if (player1 == "Rock" and player2 == "Scissors") or \
       (player1 == "Paper" and player2 == "Rock") or \
       (player1 == "Scissors" and player2 == "Paper"):
        return "Player 1 wins"
    else:
        return "Player 2 wins"
```
