---
id: 6925e2068081f40f549ced1b
title: "Sfida 137: generatore di fiocchi di neve"
challengeType: 29
dashedName: challenge-137
---

# --description--

Dato un testo multilinea che usa caratteri di nuova riga (`\n`) per rappresentare un'interruzione di riga, restituisci una nuova stringa in cui ogni riga è specchiata orizzontalmente e allegata alla fine della riga originale.

- Specchia una riga invertendo tutti i suoi caratteri, inclusi gli spazi.

Per esempio, dato `"* \n *\n* "`, che stampa sulla console:

```sh
* 
 *
* 
```

Restituisci `"*  *\n ** \n*  *"`, che stampa sulla console:

```sh
*  *
 ** 
*  *
```

Presta molta attenzione agli spazi bianchi nelle stringhe date e restituite. Assicurati di non eliminarne nessuno.

# --hints--

`generate_snowflake("* \n *\n* ")` dovrebbe restituire `"*  *\n ** \n*  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` dovrebbe restituire `"X=~~=X"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` dovrebbe restituire `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` dovrebbe restituire `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` dovrebbe restituire `"*  --  *\n * -- * \n*  --  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```
