---
id: 6925e2068081f40f549ced1a
title: "Sfida 136: parser di immagini in markdown"
challengeType: 29
dashedName: challenge-136
---

# --description--

Dato un testo che rappresenta un'immagine in Markdown, restituisci la stringa HTML equivalente.

Un'immagine in Markdown ha il formato seguente: `"![alt text](image_url)"`. Dove:

- `alt text` è la descrizione dell'immagine (il valore dell'attributo `alt`).
- `image_url` è l'URL sorgente dell'immagine (il valore dell'attributo `src`).

Restituisci una stringa con il tag HTML `img` con l'attributo `src` impostato sull'URL dell'immagine e l'attributo `alt` impostato sul testo alternativo.

Per esempio, dato `"![Cute cat](cat.png)"` restituisci `'<img src="cat.png" alt="Cute cat">'`;

- Assicurati che il tag, l'ordine degli attributi, gli spazi e l'uso delle virgolette siano esattamente come sopra.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi — controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`parse_image("![Cute cat](cat.png)")` dovrebbe restituire `'<img src="cat.png" alt="Cute cat">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` dovrebbe restituire `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` dovrebbe restituire `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```
