---
id: 6925e2068081f40f549ced1c
title: "Sfida 138: somma dei divisori"
challengeType: 29
dashedName: challenge-138
---

# --description--

Dato un numero intero positivo, restituisci la somma di tutti i suoi divisori.

- Un divisore è un qualsiasi intero che divide il numero senza resto (il resto è `0`).
- Conta ogni divisore una sola volta.

Per esempio, dato `6`, restituisci `12` perché i divisori di `6` sono `1`, `2`, `3` e `6`, e la somma di questi è `12`.

# --hints--

`sum_divisors(6)` dovrebbe restituire `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` dovrebbe restituire `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` dovrebbe restituire `56`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` dovrebbe restituire `224`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` dovrebbe restituire `806`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` dovrebbe restituire `23520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```
