---
id: 69272dcf1c24b44fd79137c5
title: "Sfida 142: somma la stringa"
challengeType: 29
dashedName: challenge-142
---

# --description--

Data una stringa contenente cifre e altri caratteri, restituisci la somma di tutti i numeri presenti nella stringa.

- Considera le cifre consecutive come un unico numero. Per esempio, `"13"` conta come 13, non come 1 + 3.
- Ignora tutti i caratteri non numerici.

# --hints--

`string_sum("3apples2bananas")` dovrebbe restituire `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("3apples2bananas"), 5)`)
}})
```

`string_sum("10cats5dogs2birds")` dovrebbe restituire `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("10cats5dogs2birds"), 17)`)
}})
```

`string_sum("125344")` dovrebbe restituire `125344`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("125344"), 125344)`)
}})
```

`string_sum("a1b20c300")` dovrebbe restituire `321`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a1b20c300"), 321)`)
}})
```

`string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` dovrebbe restituire `1653`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653)`)
}})
```

# --seed--

## --seed-contents--

```py
def string_sum(s):

    return s
```

# --solutions--

```py
import re
def string_sum(s):
    numbers = re.findall(r'\d+', s)
    return sum(int(num) for num in numbers)
```
