---
id: 69272dcf1c24b44fd79137c3
title: "Sfida 140: SCREAMING_SNAKE_CASE"
challengeType: 29
dashedName: challenge-140
---

# --description--

Data una stringa che rappresenta il nome di una variabile, restituisci il nome della variabile convertito in SCREAMING_SNAKE_CASE.

I nomi delle variabili forniti saranno scritti in uno dei seguenti formati:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

Nei formati sopra, le parole sono separate da un underscore (`_`), un trattino (`-`) oppure una nuova parola inizia con una lettera maiuscola.

Per convertire in SCREAMING_SNAKE_CASE:

- Trasforma tutte le lettere in maiuscolo
- Separa le parole con un underscore (`_`)

# --hints--

`to_screaming_snake_case("userEmail")` dovrebbe restituire `"USER_EMAIL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("userEmail"), "USER_EMAIL")`)
}})
```

`to_screaming_snake_case("UserPassword")` dovrebbe restituire `"USER_PASSWORD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("UserPassword"), "USER_PASSWORD")`)
}})
```

`to_screaming_snake_case("user_id")` dovrebbe restituire `"USER_ID"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user_id"), "USER_ID")`)
}})
```

`to_screaming_snake_case("user-address")` dovrebbe restituire `"USER_ADDRESS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user-address"), "USER_ADDRESS")`)
}})
```

`to_screaming_snake_case("username")` dovrebbe restituire `"USERNAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("username"), "USERNAME")`)
}})
```

`to_screaming_snake_case("my_variable_name")` dovrebbe restituire `"MY_VARIABLE_NAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("my_variable_name"), "MY_VARIABLE_NAME")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_screaming_snake_case(variable_name):

    return variable_name
```

# --solutions--

```py
import re
def to_screaming_snake_case(variable_name):
    temp = re.sub(r'[-_]+', ' ', variable_name)
    temp = re.sub(r'([a-z0-9])([A-Z])', r'\1 \2', temp)
    words = temp.strip().split()
    return '_'.join(words).upper()
```
