---
id: 69272dcf1c24b44fd79137c4
title: "Sfida 141: carburante per il decollo"
challengeType: 29
dashedName: challenge-141
---

# --description--

Dato il numero di galloni di carburante attualmente nel tuo aereo e il numero richiesto di litri di carburante per raggiungere la destinazione, determina quanti galloni di carburante aggiuntivi dovresti aggiungere.

- 1 gallone equivale a 3,78541 litri.
- Se l'aereo ha già abbastanza carburante, restituisci `0`.
- Puoi aggiungere solo galloni interi.
- Non includere decimali nel numero restituito.

# --hints--

`fuel_to_add(0, 1)` dovrebbe restituire `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(0, 1), 1)`)
}})
```

`fuel_to_add(5, 40)` dovrebbe restituire `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(5, 40), 6)`)
}})
```

`fuel_to_add(10, 30)` dovrebbe restituire `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(10, 30), 0)`)
}})
```

`fuel_to_add(896, 20500)` dovrebbe restituire `4520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(896, 20500), 4520)`)
}})
```

`fuel_to_add(1000, 50000)` dovrebbe restituire `12209`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(1000, 50000), 12209)`)
}})
```

# --seed--

## --seed-contents--

```py
def fuel_to_add(current_gallons, required_liters):

    return current_gallons
```

# --solutions--

```py
import math
def fuel_to_add(current_gallons, required_liters):
    liters_per_gallon = 3.78541
    current_liters = current_gallons * liters_per_gallon
    if current_liters >= required_liters:
        return 0
    liters_needed = required_liters - current_liters
    return math.ceil(liters_needed / liters_per_gallon)
```
