---
id: 69272dcf1c24b44fd79137c6
title: "Sfida 143: parser per corsivo in markdown"
challengeType: 29
dashedName: challenge-143
---

# --description--

Data una stringa che può includere del testo in corsivo in Markdown, restituisci la stringa HTML equivalente.

- Il testo in corsivo in Markdown è qualsiasi testo che inizia e finisce con un singolo asterisco (`*`) o un singolo underscore (`_`).
- Non possono esserci spazi tra il testo e l’asterisco o l’underscore, ma possono esserci spazi all’interno del testo stesso.
- Converti tutte le occorrenze in corsivo in tag HTML `i` e restituisci la stringa.

Per esempio, data `"*This is italic*"`, restituisci `"<i>This is italic</i>"`.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`parse_italics("*This is italic*")` dovrebbe restituire `"<i>This is italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is italic*"), "<i>This is italic</i>")`)
}})
```

`parse_italics("_This is also italic_")` dovrebbe restituire `"<i>This is also italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_This is also italic_"), "<i>This is also italic</i>")`)
}})
```

`parse_italics("*This is not italic *")` dovrebbe restituire `"*This is not italic *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is not italic *"), "*This is not italic *")`)
}})
```

`parse_italics("_ This is also not italic_")` dovrebbe restituire `"_ This is also not italic_"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_ This is also not italic_"), "_ This is also not italic_")`)
}})
```

`parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog.")` dovrebbe restituire `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_italics(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_italics(markdown):
    pattern = r'(\*|_)([^\s][^]*?[^\s])\1'
    pattern = r'(\*|_)([^\s].*?[^\s])\1'
    return re.sub(pattern, r'<i>\2</i>', markdown)
```
