---
id: 69306364df283fcaff2e1ad9
title: "Sfida 148: pressione degli pneumatici"
challengeType: 29
dashedName: challenge-148
---

# --description--

Dato un array con quattro numeri che rappresentano le pressioni in psi dei quattro pneumatici del tuo veicolo, e un altro array di due numeri che rappresentano la pressione minima e massima per i tuoi pneumatici in bar, restituisci un array di quattro stringhe che descrivano lo stato di ciascun pneumatico.

- 1 bar equivale a 14,5038 psi.

Restituisci un array con i seguenti valori per ogni pneumatico:

- `"Low"` se la pressione dello pneumatico è inferiore al minimo consentito.
- `"Good"` se è compresa tra il minimo e il massimo consentito.
- `"High"` se è superiore al massimo consentito.

# --hints--

`tire_status([32, 28, 35, 29], [2, 3])` dovrebbe restituire `["Good", "Low", "Good", "Low"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"])`)
}})
```

`tire_status([32, 28, 35, 30], [2, 2.3])` dovrebbe restituire `["Good", "Low", "High", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"])`)
}})
```

`tire_status([29, 26, 31, 28], [2.1, 2.5])` dovrebbe restituire `["Low", "Low", "Good", "Low"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"])`)
}})
```

`tire_status([31, 31, 30, 29], [1.5, 2])` dovrebbe restituire `["High", "High", "High", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"])`)
}})
```

`tire_status([30, 28, 30, 29], [1.9, 2.1])` dovrebbe restituire `["Good", "Good", "Good", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"])`)
}})
```

# --seed--

## --seed-contents--

```py
def tire_status(pressures_psi, range_bar):

    return pressures_psi
```

# --solutions--

```py
def tire_status(pressures_psi, range_bar):
    psi_to_bar = 1 / 14.5038
    min_bar, max_bar = range_bar
    result = []

    for psi in pressures_psi:
        bar = psi * psi_to_bar
        if bar < min_bar:
            result.append("Low")
        elif bar > max_bar:
            result.append("High")
        else:
            result.append("Good")
    return result
```
