---
id: 69306364df283fcaff2e1ad7
title: "Sfida 146: sedile per mancini a tavola"
challengeType: 29
dashedName: challenge-146
---

# --description--

Data una matrice 4x2 (array di array) che rappresenta la disposizione dei posti a tavola per un pasto, determina quanti sedili può occupare una persona mancina.

- Una persona mancina non può sedersi dove si troverebbe una persona destrimane nel sedile immediatamente a sinistra.

Nella matrice data:

- Un `"R"` è un sedile occupato da una persona destrimane.
- Un `"L"` è un sedile occupato da una persona mancina.
- Un `"U"` è un sedile libero.
- Solo i sedili liberi sono disponibili per sedersi.
- I sedili nella fila superiore sono rivolti verso il "basso", e i sedili nella fila inferiore sono rivolti verso l'"alto" (come a una tavola), quindi sinistra e destra sono relativi all'orientamento del sedile.
- I sedili agli angoli hanno solo un sedile accanto.

Per esempio, nella matrice data:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

Il sedile in alto a sinistra non può essere occupato perché c'è una persona destrimane a sinistra. Gli altri due sedili liberi possono essere occupati perché non c'è una persona destrimane a sinistra.

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` dovrebbe restituire `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` dovrebbe restituire `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` dovrebbe restituire `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` dovrebbe restituire `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```
