---
id: 69306364df283fcaff2e1ada
title: "Sfida 149: vOwElcAsE"
challengeType: 29
dashedName: challenge-149
---

# --description--

Data una stringa, restituisci una nuova stringa in cui tutte le vocali sono convertite in maiuscolo e tutti gli altri caratteri alfabetici sono convertiti in minuscolo.

- Le vocali sono `"a"`, `"e"`, `"i"`, `"o"` e `"u"` in qualsiasi caso.
- I caratteri non alfabetici devono rimanere invariati.

# --hints--

`vowel_case("vowelcase")` dovrebbe restituire `"vOwElcAsE"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("vowelcase"), "vOwElcAsE")`)
}})
```

`vowel_case("coding is fun")` dovrebbe restituire `"cOdIng Is fUn"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("coding is fun"), "cOdIng Is fUn")`)
}})
```

`vowel_case("HELLO, world!")` dovrebbe restituire `"hEllO, wOrld!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HELLO, world!"), "hEllO, wOrld!")`)
}})
```

`vowel_case("git cherry-pick")` dovrebbe restituire `"gIt chErry-pIck"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("git cherry-pick"), "gIt chErry-pIck")`)
}})
```

`vowel_case("HEAD~1")` dovrebbe restituire `"hEAd~1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HEAD~1"), "hEAd~1")`)
}})
```

# --seed--

## --seed-contents--

```py
def vowel_case(s):

    return s
```

# --solutions--

```py
def vowel_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char in vowels:
            result += char.upper()
        elif char.isalpha():
            result += char.lower()
        else:
            result += char

    return result
```
