---
id: 69373793f5a867f769cde13a
title: "Sfida 155: pianta il raccolto"
challengeType: 29
dashedName: challenge-155
---

# --description--

Dato un numero intero che rappresenta la dimensione del tuo campo agricolo, e `"acres"` o `"hectares"` che rappresentano l’unità di misura della dimensione del tuo campo, e un tipo di raccolto, determina quante piante di quel tipo puoi piantare nel tuo campo.

- 1 acro corrisponde a 4046,86 metri quadrati.
- 1 ettaro corrisponde a 10.000 metri quadrati.

Ecco una lista di raccolti che saranno forniti come input e lo spazio che occupa una singola pianta:

|Raccolto|Spazio per pianta|
|-|-|
|`"corn"`|1 metro quadrato|
|`"wheat"`|0,1 metri quadrati|
|`"soybeans"`|0,5 metri quadrati|
|`"tomatoes"`|0,25 metri quadrati|
|`"lettuce"`|0,2 metri quadrati|

Restituisci il numero di piante che possono essere piantate nel campo, arrotondato per difetto all’intera pianta più vicina.

# --hints--

`get_number_of_plants(1, "acres", "corn")` dovrebbe restituire `4046`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(1, "acres", "corn"), 4046)`)
}})
```

`get_number_of_plants(2, "hectares", "lettuce")` dovrebbe restituire `100000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(2, "hectares", "lettuce"), 100000)`)
}})
```

`get_number_of_plants(20, "acres", "soybeans")` dovrebbe restituire `161874`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(20, "acres", "soybeans"), 161874)`)
}})
```

`get_number_of_plants(3.75, "hectares", "tomatoes")` dovrebbe restituire `150000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(3.75, "hectares", "tomatoes"), 150000)`)
}})
```

`get_number_of_plants(16.75, "acres", "tomatoes")` dovrebbe restituire `271139`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(16.75, "acres", "tomatoes"), 271139)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_of_plants(field_size, unit, crop):

    return field_size
```

# --solutions--

```py
import math
def get_number_of_plants(field_size, unit, crop):
    crop_space = {
        "corn": 1,
        "wheat": 0.1,
        "soybeans": 0.5,
        "tomatoes": 0.25,
        "lettuce": 0.2
    }

    field_in_sqm = field_size * 4046.86 if unit == "acres" else field_size * 10000

    return math.floor(field_in_sqm / crop_space[crop.lower()])
```
