---
id: 69373793f5a867f769cde139
title: "Sfida 154: par per la buca"
challengeType: 29
dashedName: challenge-154
---

# --description--

Dati due numeri interi, il par per una buca da golf e il numero di tratti che un golfista ha impiegato su quella buca, restituisci il punteggio del golfista usando i termini del golf.

Restituisci:

- `"Hole in one!"` se ha impiegato un tratto.
- `"Eagle"` se ha impiegato due tratti in meno rispetto al par.
- `"Birdie"` se ha impiegato un tratto in meno rispetto al par.
- `"Par"` se ha impiegato lo stesso numero di tratti del par.
- `"Bogey"` se ha impiegato un tratto in più rispetto al par.
- `"Double bogey"` se ha impiegato due tratti in più rispetto al par.

# --hints--

`golf_score(3, 3)` dovrebbe restituire `"Par"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 3), "Par")`)
}})
```

`golf_score(4, 3)` dovrebbe restituire `"Birdie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 3), "Birdie")`)
}})
```

`golf_score(3, 1)` dovrebbe restituire `"Hole in one!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 1), "Hole in one!")`)
}})
```

`golf_score(5, 7)` dovrebbe restituire `"Double bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 7), "Double bogey")`)
}})
```

`golf_score(4, 5)` dovrebbe restituire `"Bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 5), "Bogey")`)
}})
```

`golf_score(5, 3)` dovrebbe restituire `"Eagle"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 3), "Eagle")`)
}})
```

# --seed--

## --seed-contents--

```py
def golf_score(par, strokes):

    return par
```

# --solutions--

```py
def golf_score(par, strokes):
    if strokes == 1:
        return "Hole in one!"

    diff = strokes - par

    if diff == -2:
        return "Eagle"
    if diff == -1:
        return "Birdie"
    if diff == 0:
        return "Par"
    if diff == 1:
        return "Bogey"
    if diff == 2:
        return "Double bogey"
```
