---
id: 69373793f5a867f769cde137
title: "Sfida 152: numero primo circolare"
challengeType: 29
dashedName: challenge-152
---

# --description--

Dato un numero intero, determina se è un numero primo circolare.

Un numero primo circolare è un intero per cui tutte le rotazioni delle sue cifre sono a loro volta numeri primi.

Per esempio, `197` è un numero primo circolare perché tutte le rotazioni delle sue cifre: `197`, `971` e `719`, sono numeri primi.

# --hints--

`is_circular_prime(197)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(197), True)`)
}})
```

`is_circular_prime(23)` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(23), False)`)
}})
```

`is_circular_prime(13)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(13), True)`)
}})
```

`is_circular_prime(89)` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(89), False)`)
}})
```

`is_circular_prime(1193)` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(1193), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_circular_prime(n):

    return n
```

# --solutions--

```py
import math
def is_prime(n):
    if n < 2:
        return False
    for i in range(2, int(math.isqrt(n)) + 1):
        if n % i == 0:
            return False
    return True

def rotations(n):
    s = str(n)
    return [int(s[i:] + s[:i]) for i in range(len(s))]

def is_circular_prime(n):
    return all(is_prime(rot) for rot in rotations(n))
```
