---
id: 6939b873185d8e00d453563d
title: "Sfida 158: scambio di array"
challengeType: 29
dashedName: challenge-158
---

# --description--

Dato un array con due valori, restituisci un array con i valori scambiati.

Per esempio, dato `["A", "B"]` restituisci `["B", "A"]`.

# --hints--

`array_swap(["A", "B"])` dovrebbe restituire `["B", "A"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["A", "B"]), ["B", "A"])`)
}})
```

`array_swap([25, 20])` dovrebbe restituire `[20, 25]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([25, 20]), [20, 25])`)
}})
```

`array_swap([True, False])` dovrebbe restituire `[False, True]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([True, False]), [False, True])`)
}})
```

`array_swap(["1", 1])` dovrebbe restituire `[1, "1"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["1", 1]), [1, "1"])`)
}})
```

# --seed--

## --seed-contents--

```py
def array_swap(arr):

    return arr
```

# --solutions--

```py
def array_swap(arr):

    return [arr[1], arr[0]]
```
