---
id: 6939b873185d8e00d453563c
title: "Sfida 157: parser di collegamenti markdown"
challengeType: 29
dashedName: challenge-157
---

# --description--

Dato un stringa di un collegamento in Markdown, restituisci la stringa HTML equivalente.

Un'immagine Markdown ha il formato seguente: `"[link_text](link_url)"`. Restituisci la stringa del tag HTML `a` con l'insieme `href` impostato su `link_url` e `link_text` come contenuto del tag.

Per esempio, dato `"[freeCodeCamp](https://freecodecamp.org/)"` restituisci `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi — controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

`parse_link("[freeCodeCamp](https://freecodecamp.org/)")` dovrebbe restituire `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>')`)
}})
```

`parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` dovrebbe restituire `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>')`)
}})
```

`parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` dovrebbe restituire `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_link(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_link(markdown):
    match = re.match(r'\[(.*?)\]\((.*?)\)', markdown)
    text, url = match.groups()
    return f'<a href="{url}">{text}</a>'
```
