---
id: 694596b0585c11170ac7c7fb
title: "Sfida 163: caso consonante"
challengeType: 29
dashedName: challenge-163
---

# --description--

Data una stringa che rappresenta il nome di una variabile, convertila in caso consonante seguendo queste regole:

- Tutte le consonanti devono essere convertite in maiuscolo.
- Tutte le vocali (`a`, `e`, `i`, `o`, `u` in qualsiasi caso) devono essere convertite in minuscolo.
- Tutti i trattini (`-`) devono essere convertiti in underscore (`_`).

# --hints--

`to_consonant_case("helloworld")` dovrebbe restituire `"HeLLoWoRLD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("helloworld"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("HELLOWORLD")` dovrebbe restituire `"HeLLoWoRLD"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("HELLOWORLD"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("_hElLO-WOrlD-")` dovrebbe restituire `"_HeLLo_WoRLD_"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_hElLO-WOrlD-"), "_HeLLo_WoRLD_")`)
}})
```

`to_consonant_case("_~-generic_~-variable_~-name_~-here-~_")` dovrebbe restituire `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_consonant_case(s):

    return s
```

# --solutions--

```py
def to_consonant_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char == "-":
            result += "_"
        elif char.isalpha():
            if char in vowels:
                result += char.lower()
            else:
                result += char.upper()
        else:
            result += char

    return result
```
