---
id: 694596b0585c11170ac7c7fc
title: "Sfida 164: parser di codice inline in markdown"
challengeType: 29
dashedName: challenge-164
---

# --description--

Dato un stringa di Markdown che include uno o più blocchi di codice inline, restituisci la stringa HTML equivalente.

I blocchi di codice inline in Markdown usano un singolo backtick (``` ` ```) all'inizio e alla fine del testo del blocco di codice.

Restituisci la stringa data con tutti i blocchi di codice convertiti in tag HTML `code`.

Per esempio, data la stringa ```"Use `let` to declare the variable."```, return `"Usa <code>let</code> per dichiarare la variabile."``.

Nota: la console potrebbe non mostrare i tag HTML nelle stringhe quando registra messaggi. Controlla la console del browser per vedere i log con i tag inclusi.

# --hints--

```parse_inline_code("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Usa \`let\` per dichiarare la variabile."), "Usa <code>let</code> per dichiarare la variabile.")`)
}})
```

```parse_inline_code("Usa `let` o `const` per dichiarare una variabile.")``` should return `"Usa <code>let</code> o <code>const</code> per dichiarare una variabile."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` or \`const\` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.")`)
}})
```

```parse_inline_code("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Esegui \`npm install\` poi \`npm start\`."), "Esegui <code>npm install</code> poi <code>npm start</code>.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_inline_code(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_inline_code(markdown):
    return re.sub(r"`([^`]+)`", r"<code>\1</code>", markdown)
```
