---
id: 694596b0585c11170ac7c7fa
title: "Sfida 162: consumo di energia"
challengeType: 29
dashedName: challenge-162
---

# --description--

Dato il numero di calorie bruciate durante un allenamento e il numero di wattora consumati dai tuoi dispositivi elettronici durante quell’allenamento, determina quale dei due ha usato più energia.

Per confrontarli, converti entrambi i valori in joule usando le seguenti conversioni:

- 1 caloria equivale a 4184 joule.
- 1 wattora equivale a 3600 joule.

Restituisci:

- `"Workout"` se l’allenamento ha usato più energia.
- `"Devices"` se il dispositivo ha usato più energia.
- `"Equal"` se entrambi hanno usato la stessa quantità di energia.

# --hints--

`compare_energy(250, 50)` dovrebbe restituire `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` dovrebbe restituire `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` dovrebbe restituire `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` dovrebbe restituire `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` dovrebbe restituire `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` dovrebbe restituire `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```
