---
id: 696655d24b614176d4c9b78c
title: "Sfida 169: mini fizzbuzz"
challengeType: 29
dashedName: challenge-169
---

# --description--

Dato un numero intero, restituisci una stringa basata sulle seguenti regole:

- Se il numero è divisibile per 3, restituisci `"Fizz"`.
- Se il numero è divisibile per 5, restituisci `"Buzz"`.
- Se il numero è divisibile sia per 3 che per 5, restituisci `"FizzBuzz"`.
- Altrimenti, restituisci il numero dato come stringa.

# --hints--

`fizz_buzz_mini(3)` dovrebbe restituire `"Fizz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(3), "Fizz")`)
}})
```

`fizz_buzz_mini(4)` dovrebbe restituire `"4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(4), "4")`)
}})
```

`fizz_buzz_mini(35)` dovrebbe restituire `"Buzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(35), "Buzz")`)
}})
```

`fizz_buzz_mini(75)` dovrebbe restituire `"FizzBuzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(75), "FizzBuzz")`)
}})
```

`fizz_buzz_mini(98)` dovrebbe restituire `"98"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(98), "98")`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_mini(n):

    return n
```

# --solutions--

```py
def fizz_buzz_mini(n):
    if n % 3 == 0 and n % 5 == 0:
        return "FizzBuzz"
    if n % 3 == 0:
        return "Fizz"
    if n % 5 == 0:
        return "Buzz"
    return str(n)
```
