---
id: 696655d24b614176d4c9b78d
title: "Sfida 170: giorno dispari o pari"
challengeType: 29
dashedName: challenge-170
---

# --description--

Dato un timestamp (numero di millisecondi dall'epoca Unix), restituisci:

- `"odd"` se il giorno del mese per quel timestamp è dispari.
- `"even"` se il giorno del mese per quel timestamp è pari.

Per esempio, dato `1769472000000`, un timestamp per il 27 gennaio 2026, restituisci `"odd"` perché il numero del giorno (`27`) è un numero dispari.

Nota: il timestamp è in millisecondi e devi usare la data nel fuso orario UTC, non nel tuo orario locale.

# --hints--

`odd_or_even_day(1769472000000)` dovrebbe restituire `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` dovrebbe restituire `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` dovrebbe restituire `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` dovrebbe restituire `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` dovrebbe restituire `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```
