---
id: 696655d24b614176d4c9b78b
title: "Sfida 168: immagine scalata"
challengeType: 29
dashedName: challenge-168
---

# --description--

Data una stringa che rappresenta la larghezza e l'altezza di un'immagine, e un numero per scalare l'immagine, restituisci la larghezza e l'altezza scalate.

- La stringa di input è nel formato `"WxH"`. Per esempio, `"800x600"`.
- La scala è un numero per moltiplicare la larghezza e l'altezza.

Restituisci le dimensioni scalate nello stesso formato `"WxH"`.

# --hints--

`scale_image("800x600", 2)` dovrebbe restituire `"1600x1200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("800x600", 2), "1600x1200")`)
}})
```

`scale_image("100x100", 10)` dovrebbe restituire `"1000x1000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("100x100", 10), "1000x1000")`)
}})
```

`scale_image("1024x768", 0.5)` dovrebbe restituire `"512x384"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("1024x768", 0.5), "512x384")`)
}})
```

`scale_image("300x200", 1.5)` dovrebbe restituire `"450x300"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("300x200", 1.5), "450x300")`)
}})
```

# --seed--

## --seed-contents--

```py
def scale_image(size, scale):

    return size
```

# --solutions--

```py
def scale_image(size, scale):
    width, height = map(int, size.split("x"))

    new_width = round(width * scale)
    new_height = round(height * scale)

    return f"{new_width}x{new_height}"
```
