---
id: 696655d24b614176d4c9b789
title: "Sfida 166: validatore esadecimale"
challengeType: 29
dashedName: challenge-166
---

# --description--

Data una stringa, determina se è un colore esadecimale CSS valido. Un colore esadecimale CSS valido deve:

- Iniziare con un `#` e
- essere seguito da 3 o 6 caratteri esadecimali.

I caratteri esadecimali sono i numeri `0` fino a `9` e le lettere `a` fino a `f` (non fa differenza tra maiuscole e minuscole).

# --hints--

`is_valid_hex("#123")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#123"), True)`)
}})
```

`is_valid_hex("#123abc")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#123abc"), True)`)
}})
```

`is_valid_hex("#ABCDEF")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#ABCDEF"), True)`)
}})
```

`is_valid_hex("#0a1B2c")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#0a1B2c"), True)`)
}})
```

`is_valid_hex("#12G")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#12G"), False)`)
}})
```

`is_valid_hex("#1234567")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#1234567"), False)`)
}})
```

`is_valid_hex("#12 3")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#12 3"), False)`)
}})
```

`is_valid_hex("fff")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("fff"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_hex(s):

    return s
```

# --solutions--

```py
def is_valid_hex(s):
    if not s.startswith("#"):
        return False

    hex_part = s[1:]
    if len(hex_part) not in (3, 6):
        return False

    for char in hex_part:
        if not char.lower() in "0123456789abcdef":
            return False

    return True
```
