---
id: 69738771fb5a7b8b24cca2a4
title: "Sfida 178: troncare il testo"
challengeType: 29
dashedName: challenge-178
---

# --description--

Data una stringa, restituiscila così com'è se è lunga 20 caratteri o meno. Se è più lunga di 20 caratteri, tronca la stringa ai primi 17 caratteri e aggiungi `"..."` alla fine (così la lunghezza totale sarà di 20 caratteri) e restituisci il risultato.

# --hints--

`truncate_text("Hello, world!")` dovrebbe restituire `"Hello, world!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Hello, world!"), "Hello, world!")`)
}})
```

`truncate_text("This string should get truncated.")` dovrebbe restituire `"This string shoul..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("This string should get truncated."), "This string shoul...")`)
}})
```

`truncate_text("Exactly twenty chars")` dovrebbe restituire `"Exactly twenty chars"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Exactly twenty chars"), "Exactly twenty chars")`)
}})
```

`truncate_text(".....................")` dovrebbe restituire `"...................."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("....................."), "....................")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(text):

    return text
```

# --solutions--

```py
def truncate_text(text):
    if len(text) <= 20:
        return text
    return text[:17] + "..."
```
