---
id: 69738771fb5a7b8b24cca29f
title: "Sfida 173: mosse valide del pedone"
challengeType: 29
dashedName: challenge-173
---

# --description--

Data la posizione di uno dei tuoi pedoni su una scacchiera, restituisci un array con tutte le caselle valide in cui può muoversi in ordine crescente.

Una scacchiera standard è 8x8, con le colonne etichettate `A` fino a `H` (da sinistra a destra) e le righe etichettate `1` fino a `8` (dal basso verso l’alto). Ha questo aspetto:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Per questa sfida:

- Sei il giocatore in basso sulla scacchiera.
- I pedoni possono muoversi solo di una casella "in su".
- A meno che il pedone non sia nella riga di partenza (riga 2), allora può muoversi di una o due caselle in su.

Per esempio, dato `"D4"`, restituisci `["D5"]`, l’unica casella in cui il pedone può muoversi.
Dato `"B2"`, restituisci `["B3", "B4"]`, perché si trova nella riga di partenza e deve essere in ordine crescente.

# --hints--

`find_pawn_moves("D4")` dovrebbe restituire `["D5"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("D4"), ["D5"])`)
}})
```

`find_pawn_moves("B2")` dovrebbe restituire `["B3", "B4"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("B2"), ["B3", "B4"])`)
}})
```

`find_pawn_moves("A7")` dovrebbe restituire `["A8"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("A7"), ["A8"])`)
}})
```

`find_pawn_moves("G2")` dovrebbe restituire `["G3", "G4"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("G2"), ["G3", "G4"])`)
}})
```

`find_pawn_moves("E3")` dovrebbe restituire `["E4"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("E3"), ["E4"])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_pawn_moves(position):

    return position
```

# --solutions--

```py
def find_pawn_moves(position):
    column = position[0].upper()
    row = int(position[1])
    moves = []

    if row >= 8:
        return moves

    moves.append(f"{column}{row + 1}")

    if row == 2:
        moves.append(f"{column}{row + 2}")

    return moves
```
