---
id: 69738771fb5a7b8b24cca2a0
title: "Sfida 174: trovare il segno zodiacale"
challengeType: 29
dashedName: challenge-174
---

# --description--

Dato una stringa di data nel formato `"YYYY-MM-DD"`, restituisci il segno zodiacale per quella data usando il seguente grafico:

| Intervallo di date | Segno zodiacale |
| - | - |
| 21 marzo - 19 aprile | `"Aries"` |
| 20 aprile - 20 maggio | `"Taurus"` |
| 21 maggio - 20 giugno | `"Gemini"` |
| 21 giugno - 22 luglio | `"Cancer"` |
| 23 luglio - 22 agosto | `"Leo"` |
| 23 agosto - 22 settembre | `"Virgo"` |
| 23 settembre - 22 ottobre | `"Libra"` |
| 23 ottobre - 21 novembre | `"Scorpio"` |
| 22 novembre - 21 dicembre | `"Sagittarius"` |
| 22 dicembre - 19 gennaio | `"Capricorn"` |
| 20 gennaio - 18 febbraio | `"Aquarius"` |
| 19 febbraio - 20 marzo | `"Pisces"` |

- I segni zodiacali si basano solo sul mese e sul giorno, puoi ignorare l'anno.

# --hints--

`get_sign("2026-01-31")` dovrebbe restituire `"Aquarius"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2026-01-31"), "Aquarius")`)
}})
```

`get_sign("2001-06-10")` dovrebbe restituire `"Gemini"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2001-06-10"), "Gemini")`)
}})
```

`get_sign("1985-09-07")` dovrebbe restituire `"Virgo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1985-09-07"), "Virgo")`)
}})
```

`get_sign("2023-03-19")` dovrebbe restituire `"Pisces"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2023-03-19"), "Pisces")`)
}})
```

`get_sign("2045-11-05")` dovrebbe restituire `"Scorpio"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2045-11-05"), "Scorpio")`)
}})
```


`get_sign("1985-12-06")` dovrebbe restituire `"Sagittarius"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1985-12-06"), "Sagittarius")`)
}})
```

`get_sign("2025-12-30")` dovrebbe restituire `"Capricorn"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2025-12-30"), "Capricorn")`)
}})
```

`get_sign("2018-10-08")` dovrebbe restituire `"Libra"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2018-10-08"), "Libra")`)
}})
```

`get_sign("1958-05-04")` dovrebbe restituire `"Taurus"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1958-05-04"), "Taurus")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_sign(date_str):

    return date_str
```

# --solutions--

```py
def get_sign(date_str):
    _, month_str, day_str = date_str.split("-")
    month = int(month_str)
    day = int(day_str)

    if (month == 3 and day >= 21) or (month == 4 and day <= 19):
        return "Aries"
    if (month == 4 and day >= 20) or (month == 5 and day <= 20):
        return "Taurus"
    if (month == 5 and day >= 21) or (month == 6 and day <= 20):
        return "Gemini"
    if (month == 6 and day >= 21) or (month == 7 and day <= 22):
        return "Cancer"
    if (month == 7 and day >= 23) or (month == 8 and day <= 22):
        return "Leo"
    if (month == 8 and day >= 23) or (month == 9 and day <= 22):
        return "Virgo"
    if (month == 9 and day >= 23) or (month == 10 and day <= 22):
        return "Libra"
    if (month == 10 and day >= 23) or (month == 11 and day <= 21):
        return "Scorpio"
    if (month == 11 and day >= 22) or (month == 12 and day <= 21):
        return "Sagittarius"
    if (month == 12 and day >= 22) or (month == 1 and day <= 19):
        return "Capricorn"
    if (month == 1 and day >= 20) or (month == 2 and day <= 18):
        return "Aquarius"
    if (month == 2 and day >= 19) or (month == 3 and day <= 20):
        return "Pisces"
```
