---
id: 69738771fb5a7b8b24cca2a2
title: "Sfida 176: il giorno della marmotta"
challengeType: 29
dashedName: challenge-176
---

# --description--

Oggi è il giorno della marmotta, in cui una marmotta predice il tempo in base al fatto che veda o meno la sua ombra.

Dato un valore che rappresenta l'apparizione della marmotta, restituisci la previsione corretta:

- Se il valore dato è il booleano `true` (la marmotta ha visto la sua ombra), restituisci `"Looks like we'll have six more weeks of winter."`.
- Se il valore è il booleano `false` (la marmotta non ha visto la sua ombra), restituisci `"It's going to be an early spring."`.
- Se il valore è qualsiasi altro (la marmotta non si è presentata), restituisci `"No prediction this year."`.

# --hints--

`groundhog_day_prediction(True)` dovrebbe restituire `"Looks like we'll have six more weeks of winter."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(True), "Looks like we'll have six more weeks of winter.")`)
}})
```

`groundhog_day_prediction(False)` dovrebbe restituire `"It's going to be an early spring."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(False), "It's going to be an early spring.")`)
}})
```

`groundhog_day_prediction(None)` dovrebbe restituire `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(None), "No prediction this year.")`)
}})
```

`groundhog_day_prediction(" ")` dovrebbe restituire `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(" "), "No prediction this year.")`)
}})
```

`groundhog_day_prediction("True")` dovrebbe restituire `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction("True"), "No prediction this year.")`)
}})
```

# --seed--

## --seed-contents--

```py
def groundhog_day_prediction(appearance):

    return appearance
```

# --solutions--

```py
def groundhog_day_prediction(appearance):
    if appearance is True:
        return "Looks like we'll have six more weeks of winter."
    if appearance is False:
        return "It's going to be an early spring."
    return "No prediction this year."
```
