---
id: 697a49e9860d24853adef680
title: "Sfida 195: giochi invernali 2026 giorno 16: curling"
challengeType: 29
dashedName: challenge-195
---

# --description--

Data una matrice 5x5 che rappresenta la "casa" alla fine di un round di curling, determina quale squadra segna e quanti punti ottiene.

La disposizione:

- La cella centrale (indice `[2, 2]`) è il "button".
- Le 8 celle che circondano direttamente il button rappresentano l'anello 1.
- Le 16 celle sul bordo esterno della casa rappresentano l'anello 2.

Nella matrice data:

- `"."` rappresenta uno spazio vuoto.
- `"R"` rappresenta uno spazio con una pietra rossa.
- `"Y"` rappresenta uno spazio con una pietra gialla.

Regole di punteggio:

- Solo una squadra può segnare per round.
- La squadra con la pietra più vicina al button segna.
- La squadra che segna ottiene 1 punto per ogni sua pietra che è più vicina al button della pietra più vicina dell'avversario.
- Più basso è il numero dell'anello, più la pietra è vicina al centro.
- Se la pietra più vicina di entrambe le squadre è alla stessa distanza dal centro, nessuna squadra segna.

Restituisci:

- Una stringa nel formato `"team: number_of_points"`. es.: `"R: 2"`.
- oppure `"No points awarded"` se nessuna squadra ha segnato punti.

Per esempio, dato:

```js
[
  [".", ".", "R", ".", "."],
  [".", "R", ".", ".", "."],
  ["Y", ".", ".", ".", "."],
  [".", "R", ".", ".", "."],
  [".", ".", ".", ".", "."]
]
```

Restituisci `"R: 2"`. Le due pietre rosse nell'anello 1 sono pari come distanza più vicina e sono le uniche due pietre più vicine rispetto alla pietra gialla più vicina.

# --hints--

`score_curling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]])` dovrebbe restituire `"R:2"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]]), "R: 2")`)
}})
```

`score_curling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]])` dovrebbe restituire `"Y: 3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]]), "Y: 3")`)
}})
```

`score_curling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]])` dovrebbe restituire `"No points awarded"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]]), "No points awarded")`)
}})
```

`score_curling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]])` dovrebbe restituire `"R: 4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]]), "R: 4")`)
}})
```

`score_curling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]])` dovrebbe restituire `"Y: 1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]]), "Y: 1")`)
}})
```

`score_curling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]])` dovrebbe restituire `"No points awarded"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(score_curling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]]), "No points awarded")`)
}})
```

# --seed--

## --seed-contents--

```py
def score_curling(house):

    return house
```

# --solutions--

```py
def score_curling(house):
  stones = {'R': [], 'Y': []}
  
  for i in range(5):
    for j in range(5):
      cell = house[i][j]
      if cell == 'R' or cell == 'Y':
        distance = max(abs(i - 2), abs(j - 2))
        stones[cell].append(distance)
  
  red_closest = min(stones['R']) if stones['R'] else float('inf')
  yellow_closest = min(stones['Y']) if stones['Y'] else float('inf')
  
  if red_closest == float('inf') and yellow_closest == float('inf'):
    return "No points awarded"
  
  if red_closest == yellow_closest:
    return "No points awarded"
  
  if red_closest < yellow_closest:
    winner = 'R'
    opponent_closest = yellow_closest
    winning_stones = stones['R']
  else:
    winner = 'Y'
    opponent_closest = red_closest
    winning_stones = stones['Y']
  
  points = sum(1 for d in winning_stones if d < opponent_closest)
  
  return "No points awarded" if points == 0 else f"{winner}: {points}"
```
