---
id: 697a49e9860d24853adef67c
title: "Sfida 191: giochi invernali 2026 giorno 12: bob"
challengeType: 29
dashedName: challenge-191
---

# --description--

Dato un array che rappresenta i pesi degli atleti di una squadra di bob e un numero che rappresenta il peso del bob, determina se la squadra è idonea a gareggiare.

- La lunghezza dell'array determina la dimensione della squadra: squadre da 1, 2 o 4 persone.
- Tutti i valori di peso forniti sono in chilogrammi (kg).

Il bob (slitta da solo) deve avere un peso minimo di:

- 162 kg per una squadra da 1 persona
- 170 kg per una squadra da 2 persone
- 210 kg per una squadra da 4 persone

Il peso totale del bob (atleti più slitta) non deve superare:

- 247 kg per una squadra da 1 persona
- 390 kg per una squadra da 2 persone
- 630 kg per una squadra da 4 persone

Restituisci `"Eligible"` se la squadra soddisfa tutti i requisiti, oppure `"Not Eligible"` se la squadra non soddisfa uno o più requisiti.

# --hints--

`check_eligibility([78], 165)` dovrebbe restituire `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([78], 165), "Eligible")`)
}})
```

`check_eligibility([80], 160)` dovrebbe restituire `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 160), "Not Eligible")`)
}})
```

`check_eligibility([80], 170)` dovrebbe restituire `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 170), "Not Eligible")`)
}})
```

`check_eligibility([85, 90], 170)` dovrebbe restituire `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 90], 170), "Eligible")`)
}})
```

`check_eligibility([85, 95], 168)` dovrebbe restituire `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 95], 168), "Not Eligible")`)
}})
```

`check_eligibility([112, 97], 185)` dovrebbe restituire `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([112, 97], 185), "Not Eligible")`)
}})
```

`check_eligibility([110, 102, 90, 106], 222)` dovrebbe restituire `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([110, 102, 90, 106], 222), "Eligible")`)
}})
```

`check_eligibility([106, 99, 90, 88], 205)` dovrebbe restituire `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 90, 88], 205), "Not Eligible")`)
}})
```

`check_eligibility([106, 99, 103, 96], 227)` dovrebbe restituire `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 103, 96], 227), "Not Eligible")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_eligibility(athlete_weights, sled_weight):

    return athlete_weights
```

# --solutions--

```py
def check_eligibility(athlete_weights, sled_weight):
    team_size = len(athlete_weights)

    rules = {
        1: {"min_sled": 162, "max_total": 247},
        2: {"min_sled": 170, "max_total": 390},
        4: {"min_sled": 210, "max_total": 630}
    }

    min_sled = rules[team_size]["min_sled"]
    max_total = rules[team_size]["max_total"]

    total_weight = sum(athlete_weights) + sled_weight

    if sled_weight < min_sled:
        return "Not Eligible"

    if total_weight > max_total:
        return "Not Eligible"

    return "Eligible"
```
