---
id: 697a49e6ff50d756c9b69365
title: "Sfida 188: Giochi invernali 2026 giorno 9: skeleton"
challengeType: 29
dashedName: challenge-188
---

# --description--

Data una stringa che rappresenta le curve di una pista di skeleton, determina la difficoltà della pista.

- La stringa data conterrà solo le lettere:
  - `"L"` per una curva a sinistra
  - `"R"` per una curva a destra
  - `"S"` per un tratto dritto

- Ogni cambio di direzione aggiunge 15 punti (un `"L"` seguito da un `"R"` o viceversa).
- Tutte le altre curve aggiungono 5 punti ciascuna (tutti gli altri caratteri `"L"` o `"R"`).
- I tratti dritti aggiungono 0 punti.

La difficoltà della pista si basa sul punteggio totale. Restituisci:

- `"Easy"` se il totale è da 0 a 100
- `"Medium"` se il totale è da 101 a 200
- `"Hard"` se il totale è superiore a 200

# --hints--

`get_difficulty("SLSLLSRRLSRLRL")` dovrebbe restituire `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLSLLSRRLSRLRL"), "Easy")`)
}})
```

`get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` dovrebbe restituire `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard")`)
}})
```

`get_difficulty("SRRRRLSLLRLRSSRLSRL")` dovrebbe restituire `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRRRRLSLLRLRSSRLSRL"), "Medium")`)
}})
```

`get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` dovrebbe restituire `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard")`)
}})
```

`get_difficulty("SLLSSLRLSLSLRSLSSLRL")` dovrebbe restituire `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium")`)
}})
```

`get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` dovrebbe restituire `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_difficulty(track):

    return track
```

# --solutions--

```py
def get_difficulty(track):
  score = 0

  for i in range(len(track)):
    current = track[i]
    previous = track[i - 1] if i > 0 else None

    if current == 'S':
      score += 0
    elif current == 'L' or current == 'R':
      if previous and previous != 'S' and previous != current:
        score += 15
      else:
        score += 5

  if score <= 100:
    return 'Easy'
  elif score <= 200:
    return 'Medium'
  else:
    return 'Hard'
```
