---
id: 697a49e9860d24853adef67d
title: "Sfida 192: giochi invernali 2026 giorno 13: combinata nordica"
challengeType: 29
dashedName: challenge-192
---

# --description--

Dato un array di punteggi di salto per gli atleti, calcola i loro tempi di ritardo di partenza per la parte di sci di fondo della combinata nordica.

L’atleta con il punteggio di salto più alto parte per primo (ritardo di 0 secondi). Tutti gli altri atleti partono più tardi in base a quanto il loro punteggio di salto è inferiore rispetto al miglior salto.

Per calcolare il ritardo di ogni atleta, sottrai il punteggio di salto dell’atleta dal miglior punteggio di salto complessivo e moltiplica il risultato per 1,5. Arrotonda il ritardo all’intero più vicino superiore.

# --hints--

`calculate_start_delays([120, 110, 125])` dovrebbe restituire `[8, 23, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([120, 110, 125]), [8, 23, 0])`)
}})
```

`calculate_start_delays([118, 125, 122, 120])` dovrebbe restituire `[11, 0, 5, 8]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([118, 125, 122, 120]), [11, 0, 5, 8])`)
}})
```

`calculate_start_delays([100, 105, 95, 110, 120, 115, 108])` dovrebbe restituire `[30, 23, 38, 15, 0, 8, 18]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18])`)
}})
```

`calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` dovrebbe restituire `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_start_delays(jump_scores):

    return jump_scores
```

# --solutions--

```py
import math
def calculate_start_delays(jump_scores):
    best_jump = max(jump_scores)
    return [math.ceil((best_jump - score) * 1.5) for score in jump_scores]
```
