---
id: 697a49e9860d24853adef67b
title: "Sfida 190: Giochi invernali 2026 giorno 11: hockey su ghiaccio"
challengeType: 29
dashedName: challenge-190
---

# --description--

Dato un array di 6 squadre di hockey su ghiaccio e i loro record dopo le partite del girone all’italiana, determina gli abbinamenti per il turno delle semifinali.

- Ogni elemento dell’array conterrà una squadra e il suo record nel formato `"TEAM: W-OTW-OTL-L"`. Dove:
  - `"W"` è il numero di vittorie nei tempi regolamentari, che valgono 3 punti ciascuna
  - `"OTW"` è il numero di vittorie ai tempi supplementari, che valgono 2 punti ciascuna
  - `"OTL"` è il numero di sconfitte ai tempi supplementari, che valgono 1 punto ciascuna
  - `"L"` è il numero di sconfitte, che valgono 0 punti ciascuna
  
Per esempio, `"FIN: 2-2-1-0"` avrebbe 11 punti sommando il suo record.
  
Trova il numero totale di punti per ogni squadra e restituisci `"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."`. Per esempio, `"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`get_semifinal_matchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])` dovrebbe restituire `"The semi-final games will be FIN vs SWE and CAN vs USA."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.")`)
}})
```

`get_semifinal_matchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])` dovrebbe restituire `"The semi-final games will be USA vs CZE and CAN vs FIN."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.")`)
}})
```

`get_semifinal_matchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])` dovrebbe restituire `"The semi-final games will be CAN vs ITA and USA vs CZE."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.")`)
}})
```

`get_semifinal_matchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])` dovrebbe restituire `"The semi-final games will be JPN vs ITA and AUT vs KOR."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_semifinal_matchups(teams):

    return teams
```

# --solutions--

```py
def get_semifinal_matchups(teams):
    team_points = []
    for team_str in teams:
        name, record = team_str.split(": ")
        W, OTW, OTL, L = map(int, record.split("-"))
        points = W * 3 + OTW * 2 + OTL * 1 + L * 0
        team_points.append((name, points))

    team_points.sort(key=lambda x: x[1], reverse=True)
    first, second, third, fourth = team_points[:4]

    return f"The semi-final games will be {first[0]} vs {fourth[0]} and {second[0]} vs {third[0]}."
```
