---
id: 697a49e6ff50d756c9b6935f
title: "Sfida 182: giochi invernali 2026 giorno 3: biathlon"
challengeType: 29
dashedName: challenge-182
---

# --description--

Dato un array di numeri interi, dove ogni valore rappresenta il numero di obiettivi colpiti in un singolo round di biathlon, restituisci la distanza totale di penalità che l'atleta deve sciare.

- Ogni round consiste in 5 obiettivi.
- Ogni obiettivo mancato comporta un loop di penalità di 150 metri.

# --hints--

`calculate_penalty_distance([4, 4])` dovrebbe restituire `300`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 4]), 300)`)
}})
```

`calculate_penalty_distance([5, 5])` dovrebbe restituire `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 5]), 0)`)
}})
```

`calculate_penalty_distance([4, 5, 3, 5])` dovrebbe restituire `450`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 5, 3, 5]), 450)`)
}})
```

`calculate_penalty_distance([5, 4, 5, 5])` dovrebbe restituire `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 4, 5, 5]), 150)`)
}})
```

`calculate_penalty_distance([4, 3, 0, 3])` dovrebbe restituire `1500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 3, 0, 3]), 1500)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_penalty_distance(rounds):

    return rounds
```

# --solutions--

```py
def calculate_penalty_distance(rounds):
    total_penalty = 0

    for hits in rounds:
        misses = 5 - hits
        total_penalty += misses * 150

    return total_penalty
```
