---
id: 697a49e9860d24853adef67e
title: "Sfida 193: giochi invernali 2026 giorno 14: sci alpinismo"
challengeType: 29
dashedName: challenge-193
---

# --description--

Dato lo spessore della neve e la pendenza di una montagna, determina se c'è rischio di valanga.

- I valori dello spessore della neve sono `"Shallow"`, `"Moderate"` o `"Deep"`.
- I valori della pendenza sono `"Gentle"`, `"Steep"` o `"Very Steep"`.

Restituisci `"Safe"` o `"Risky"` in base a questa tabella:

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalanche_risk("Shallow", "Gentle")` dovrebbe restituire `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Shallow", "Steep")` dovrebbe restituire `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Steep"), "Safe")`)
}})
```

`avalanche_risk("Shallow", "Very Steep")` dovrebbe restituire `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Very Steep"), "Safe")`)
}})
```

`avalanche_risk("Moderate", "Gentle")` dovrebbe restituire `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Moderate", "Steep")` dovrebbe restituire `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Steep"), "Risky")`)
}})
```

`avalanche_risk("Moderate", "Very Steep")` dovrebbe restituire `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Very Steep"), "Risky")`)
}})
```

`avalanche_risk("Deep", "Gentle")` dovrebbe restituire `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Deep", "Steep")` dovrebbe restituire `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Steep"), "Risky")`)
}})
```

`avalanche_risk("Deep", "Very Steep")` dovrebbe restituire `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Very Steep"), "Risky")`)
}})
```

# --seed--

## --seed-contents--

```py
def avalanche_risk(snow_depth, slope):

    return snow_depth
```

# --solutions--

```py
def avalanche_risk(snow_depth, slope):
    risk_table = {
        "Gentle": {
            "Shallow": "Safe",
            "Moderate": "Safe",
            "Deep": "Safe",
        },
        "Steep": {
            "Shallow": "Safe",
            "Moderate": "Risky",
            "Deep": "Risky",
        },
        "Very Steep": {
            "Shallow": "Safe",
            "Moderate": "Risky",
            "Deep": "Risky",
        },
    }

    return risk_table[slope][snow_depth]
```
