---
id: 697a49e6ff50d756c9b6935e
title: "Sfida 181: giochi invernali 2026 giorno 2: snowboard"
challengeType: 29
dashedName: challenge-181
---

# --description--

Dato lo stance iniziale di uno snowboarder e una rotazione in gradi, determina lo stance di atterraggio.

- Lo stance di uno snowboarder è o `"Regular"` o `"Goofy"`.
- Le rotazioni dei trick sono multipli di 90 gradi. Un valore positivo indica rotazione in senso orario, un valore negativo indica rotazione in senso antiorario.
- Lo stance di atterraggio cambia ogni 180 gradi di rotazione.

Per esempio, dato `"Regular"` e `90`, restituisci `"Regular"`. Dato `"Regular"` e `180` gradi, restituisci `"Goofy"`.

# --hints--

`get_landing_stance("Regular", 90)` dovrebbe restituire `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` dovrebbe restituire `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` dovrebbe restituire `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` dovrebbe restituire `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` dovrebbe restituire `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` dovrebbe restituire `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
