---
id: 697a49e6ff50d756c9b69360
title: "Sfida 183: giochi invernali 2026 giorno 4: salto con gli sci"
challengeType: 29
dashedName: challenge-183
---

# --description--

Dato il punteggio per la distanza, il punteggio per lo stile, un valore di compensazione del vento e un bonus per il punto K, calcola il tuo punteggio per il salto con gli sci e determina se hai vinto una medaglia o no.

- Il tuo punteggio si calcola sommando i quattro valori sopra indicati.

I punteggi totali attuali degli altri saltatori sono:

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- Se il tuo punteggio è il migliore, restituisci `"Gold"`
- Se è il secondo migliore, restituisci `"Silver"`
- Se è il terzo migliore, restituisci `"Bronze"`
- Altrimenti, restituisci `"No Medal"`

# --hints--

`ski_jump_medal(125.0, 58.0, 0.0, 6.0)` dovrebbe restituire `"Gold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(125.0, 58.0, 0.0, 6.0), "Gold")`)
}})
```

`ski_jump_medal(119.0, 50.0, 1.0, 4.0)` dovrebbe restituire `"Bronze"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(119.0, 50.0, 1.0, 4.0), "Bronze")`)
}})
```

`ski_jump_medal(122.0, 52.0, -1.0, 4.0)` dovrebbe restituire `"Silver"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(122.0, 52.0, -1.0, 4.0), "Silver")`)
}})
```

`ski_jump_medal(118.0, 50.5, -1.5, 4.0)` dovrebbe restituire `"No Medal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(118.0, 50.5, -1.5, 4.0), "No Medal")`)
}})
```

`ski_jump_medal(124.0, 50.5, 2.0, 5.0)` dovrebbe restituire `"Gold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(124.0, 50.5, 2.0, 5.0), "Gold")`)
}})
```

`ski_jump_medal(119.0, 49.5, 0.0, 3.0)` dovrebbe restituire `"No Medal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(119.0, 49.5, 0.0, 3.0), "No Medal")`)
}})
```

# --seed--

## --seed-contents--

```py
def ski_jump_medal(distance_points, style_points, wind_comp, k_point_bonus):

    return distance_points
```

# --solutions--

```py
def ski_jump_medal(distance_points, style_points, wind_comp, k_point_bonus):
    my_score = distance_points + style_points + wind_comp + k_point_bonus
    other_scores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0]
    all_scores = other_scores + [my_score]
    all_scores.sort(reverse=True)
    rank = all_scores.index(my_score) + 1

    if rank == 1:
        return "Gold"
    elif rank == 2:
        return "Silver"
    elif rank == 3:
        return "Bronze"
    else:
        return "No Medal"
```
