---
id: 697a49e6ff50d756c9b69362
title: "Sfida 185: giochi invernali 2026 giorno 6: pattinaggio artistico"
challengeType: 29
dashedName: challenge-185
---

# --description--

Dato un array di punteggi dei giudici e penalità opzionali, calcola il punteggio finale per una routine di pattinaggio artistico.

Il primo argomento è un array di 10 punteggi dei giudici, ciascuno un numero da 0 a 10. Rimuovi il punteggio più alto e quello più basso e somma gli 8 punteggi rimanenti per ottenere il punteggio base.

Qualsiasi argomento aggiuntivo passato alla funzione sono penalità. Sottrai tutte le penalità dal punteggio base per ottenere il punteggio finale.

# --hints--

`compute_score([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1)` dovrebbe restituire `64`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1), 64)`)
}})
```

`compute_score([10, 10, 10, 10, 10, 10, 10, 10, 10, 10])` dovrebbe restituire `80`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 10, 10, 10, 10, 10, 10, 10, 10, 10]), 80)`)
}})
```

`compute_score([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1)` dovrebbe restituire `67`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1), 67)`)
}})
```

`compute_score([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0)` dovrebbe restituire `67.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0), 67.5)`)
}})
```

`compute_score([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5)` dovrebbe restituire `59`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5), 59)`)
}})
```

# --seed--

## --seed-contents--

```py
def compute_score(judge_scores, *penalties):

    return judge_scores
```

# --solutions--

```py
def compute_score(judge_scores, *penalties):
    sorted_scores = sorted(judge_scores)
    base_score = sum(sorted_scores[1:9])
    total_penalty = sum(penalties)

    return base_score - total_penalty
```
