---
id: 697a49e6ff50d756c9b69366
title: "Sfida 189: giochi invernali 2026 giorno 10: sci alpino"
challengeType: 29
dashedName: challenge-189
---

# --description--

Dato il dislivello verticale, la distanza orizzontale e il tipo di una pista da sci, determina la valutazione della difficoltà della pista.

Per determinare la valutazione:

- Calcola la pendenza della pista dividendo il dislivello per la distanza.
- Poi, calcola la pendenza aggiustata in base al tipo di pista:
  - `"Downhill"`: moltiplica la pendenza per 1.2
  - `"Slalom"`: moltiplica la pendenza per 0.9
  - `"Giant Slalom"`: moltiplica la pendenza per 1.0

Restituisci:

- `"Green"` se la pendenza aggiustata è minore o uguale a 0.1
- `"Blue"` se la pendenza aggiustata è maggiore di 0.1 e minore o uguale a 0.25
- `"Black"` se la pendenza aggiustata è maggiore di 0.25


# --hints--

`get_hill_rating(95, 900, "Slalom")` dovrebbe restituire `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(95, 900, "Slalom"), "Green")`)
}})
```

`get_hill_rating(620, 2800, "Downhill")` dovrebbe restituire `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(620, 2800, "Downhill"), "Black")`)
}})
```

`get_hill_rating(420, 1680, "Giant Slalom")` dovrebbe restituire `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(420, 1680, "Giant Slalom"), "Blue")`)
}})
```

`get_hill_rating(250, 3000, "Downhill")` dovrebbe restituire `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(250, 3000, "Downhill"), "Green")`)
}})
```

`get_hill_rating(110, 900, "Slalom")` dovrebbe restituire `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(110, 900, "Slalom"), "Blue")`)
}})
```

`get_hill_rating(380, 1500, "Giant Slalom")` dovrebbe restituire `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(380, 1500, "Giant Slalom"), "Black")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_hill_rating(drop, distance, hill_type):

    return drop
```

# --solutions--

```py
def get_hill_rating(drop, distance, hill_type):
    steepness = drop / distance

    if hill_type == "Downhill":
        steepness *= 1.2
    elif hill_type == "Slalom":
        steepness *= 0.9
    elif hill_type == "Giant Slalom":
        steepness *= 1.0

    if steepness <= 0.1:
        return "Green"
    elif steepness <= 0.25:
        return "Blue"
    else:
        return "Black"
```
