---
id: 698a1a863194f1f4e63f645e
title: "Sfida 207: spazio più piccolo"
challengeType: 29
dashedName: challenge-207
---

# --description--

Data una stringa, restituisci la sottostringa compresa tra due caratteri identici che hanno il numero più piccolo di caratteri tra di loro (spazio più piccolo).

- Ci sarà sempre almeno una coppia di caratteri corrispondenti.
- La sottostringa restituita deve escludere i caratteri corrispondenti.
- Se due o più spazi hanno la stessa lunghezza, restituisci i caratteri del primo.

Per esempio, data `"ABCDAC"`, restituisci `"DA"` perché:

- Solo `"A"` e `"C"` si ripetono nella stringa.
- Il numero di caratteri tra i due caratteri `"A"` è 3, e tra i caratteri `"C"` è 2.
- Quindi restituisci la stringa tra i due caratteri `"C"`.

# --hints--

`smallest_gap("ABCDAC")` dovrebbe restituire `"DA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("ABCDAC"), "DA")`)
}})
```

`smallest_gap("racecar")` dovrebbe restituire `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("racecar"), "e")`)
}})
```

`smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` dovrebbe restituire `"#q6e&rkf(p"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p")`)
}})
```

`smallest_gap("Hello World")` dovrebbe restituire `""`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("Hello World"), "")`)
}})
```

`smallest_gap("The quick brown fox jumps over the lazy dog.")` dovrebbe restituire `"fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("The quick brown fox jumps over the lazy dog."), "fox")`)
}})
```

# --seed--

## --seed-contents--

```py
def smallest_gap(s):

    return s
```

# --solutions--

```py
def smallest_gap(s):
    min_gap = float("inf")
    result = ""

    for i in range(len(s)):
        for j in range(i + 1, len(s)):
            if s[i] == s[j]:
                gap = j - i - 1

                if gap < min_gap:
                    min_gap = gap
                    result = s[i + 1:j]

                break

    return result
```
