---
id: 698a1a73ade5ac0e19180fa4
title: "Sfida 201: spostamento di una matrice"
challengeType: 29
dashedName: challenge-201
---

# --description--

Data una matrice (array di array) di numeri e un intero, sposta tutti i valori nella matrice della quantità indicata.

- Uno spostamento positivo muove i valori verso destra.
- Uno spostamento negativo muove i valori verso sinistra.

I valori devono avvolgersi:

- Considera la matrice come una sequenza continua di valori.
- Quando un valore supera la fine di una riga, continua all'inizio della riga successiva.
- Quando un valore supera l'ultima posizione nella matrice, si avvolge alla prima posizione.
- Lo stesso vale al contrario quando si sposta a sinistra.

Per esempio, dato:

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

con uno spostamento di `1`, sposta tutti i numeri di uno verso destra:

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shift_matrix([[1, 2, 3], [4, 5, 6]], 1)` dovrebbe restituire `[[6, 1, 2], [3, 4, 5]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6]], -1)` dovrebbe restituire `[[2, 3, 4], [5, 6, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)` dovrebbe restituire `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)` dovrebbe restituire `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]])`)
}})
```

`shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)` dovrebbe restituire `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]])`)
}})
```

`shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)` dovrebbe restituire `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]])`)
}})
```

# --seed--

## --seed-contents--

```py
def shift_matrix(matrix, shift):

    return matrix
```

# --solutions--

```py
def shift_matrix(matrix, shift):
    rows = len(matrix)
    cols = len(matrix[0])
    flat = [num for row in matrix for num in row]
    n = len(flat)
    s = shift % n
    shifted = flat[-s:] + flat[:-s] if s != 0 else flat[:]
    result = []
    for i in range(rows):
        result.append(shifted[i*cols:(i+1)*cols])

    return result
```
