---
id: 698a1a73ade5ac0e19180fa7
title: "Sfida 204: somma le lettere"
challengeType: 29
dashedName: challenge-204
---

# --description--

Data una stringa, restituisci la somma delle sue lettere.

- Le lettere sono da A a Z in maiuscolo o minuscolo
- I valori delle lettere sono: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26
- Le lettere maiuscole e minuscole hanno lo stesso valore.
- Ignora tutti i caratteri che non sono lettere.

# --hints--

`sum_letters("Hello")` dovrebbe restituire `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("Hello"), 52)`)
}})
```

`sum_letters("freeCodeCamp")` dovrebbe restituire `94`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("freeCodeCamp"), 94)`)
}})
```

`sum_letters("The quick brown fox jumps over the lazy dog.")` dovrebbe restituire `473`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("The quick brown fox jumps over the lazy dog."), 473)`)
}})
```

`sum_letters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` dovrebbe restituire `1681`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681)`)
}})
```

`sum_letters("</404>")` dovrebbe restituire `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("</404>"), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_letters(s):

    return s
```

# --solutions--

```py
def sum_letters(s):
    total = 0
    for char in s:
        upper = char.upper()
        if 'A' <= upper <= 'Z':
            total += ord(upper) - ord('A') + 1
    return total
```
