---
id: 698a1a73ade5ac0e19180fa3
title: "Sfida 200: conteggio di lettere e numeri"
challengeType: 29
dashedName: challenge-200
---

# --description--

Data una stringa, restituisci un messaggio con il conteggio di quante lettere e numeri contiene.

- Le lettere sono `A-Z` e `a-z`.
- I numeri sono `0-9`.
- Ignora tutti gli altri caratteri.

Restituisci `"The string has X letters and Y numbers."`, dove `"X"` è il conteggio delle lettere e `"Y"` è il conteggio dei numeri. Se uno dei due conteggi è 1, usa la forma singolare per quell’elemento. Per esempio: `"1 letter"` invece di `"1 letters"` e `"1 number"` invece di `"1 numbers"`.

# --hints--

`count_letters_and_numbers("helloworld123")` dovrebbe restituire `"The string has 10 letters and 3 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("helloworld123"), "The string has 10 letters and 3 numbers.")`)
}})
```

`count_letters_and_numbers("Catch 22")` dovrebbe restituire `"The string has 5 letters and 2 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("Catch 22"), "The string has 5 letters and 2 numbers.")`)
}})
```

`count_letters_and_numbers("A1!")` dovrebbe restituire `"The string has 1 letter and 1 number."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("A1!"), "The string has 1 letter and 1 number.")`)
}})
```

`count_letters_and_numbers("12345")` dovrebbe restituire `"The string has 0 letters and 5 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("12345"), "The string has 0 letters and 5 numbers.")`)
}})
```

`count_letters_and_numbers("password")` dovrebbe restituire `"The string has 8 letters and 0 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("password"), "The string has 8 letters and 0 numbers.")`)
}})
```

# --seed--

## --seed-contents--

```py
def count_letters_and_numbers(s):

    return s
```

# --solutions--

```py
def count_letters_and_numbers(s):
    letter_count = 0
    number_count = 0

    for char in s:
        if ('A' <= char <= 'Z') or ('a' <= char <= 'z'):
            letter_count += 1
        elif '0' <= char <= '9':
            number_count += 1

    letter_word = "letter" if letter_count == 1 else "letters"
    number_word = "number" if number_count == 1 else "numbers"

    return f"The string has {letter_count} {letter_word} and {number_count} {number_word}."
```
