---
id: 698a1a73ade5ac0e19180fa8
title: "Sfida 205: conteggio dei cubi perfetti"
challengeType: 29
dashedName: challenge-205
---

# --description--

Dati due numeri interi, determina quanti cubi perfetti esistono nell'intervallo compreso tra i due numeri inclusi.

- Il numero più basso non è garantito essere il primo argomento.
- Un numero è un cubo perfetto se esiste un numero intero (`n`) tale che `n * n * n = number`. Per esempio, 27 è un cubo perfetto perché `3 * 3 * 3 = 27`.

# --hints--

`count_perfect_cubes(3, 30)` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(3, 30), 2)`)
}})
```

`count_perfect_cubes(1, 30)` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(1, 30), 3)`)
}})
```

`count_perfect_cubes(30, 0)` dovrebbe restituire `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(30, 0), 4)`)
}})
```

`count_perfect_cubes(-64, 64)` dovrebbe restituire `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(-64, 64), 9)`)
}})
```

`count_perfect_cubes(9214, -8127)` dovrebbe restituire `41`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(9214, -8127), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_perfect_cubes(a, b):

    return a
```

# --solutions--

```py
def count_perfect_cubes(a, b):
    start, end = min(a, b), max(a, b)
    count = 0

    n = 0
    while n ** 3 <= end:
        if n ** 3 >= start:
            count += 1
        n += 1

    n = -1
    while n ** 3 >= start:
        if n ** 3 <= end:
            count += 1
        n -= 1

    return count
```
