---
id: 698a1a73ade5ac0e19180fa9
title: "Sfida 206: valori delle carte da gioco"
challengeType: 29
dashedName: challenge-206
---

# --description--

Dato un array di carte da gioco, restituisci un nuovo array con il valore numerico di ogni carta.

Valori delle carte:

- Un asso (`"A"`) ha valore 1.
- Le carte numerate (`"2"` - `"10"`) hanno il loro valore nominale: rispettivamente da 2 a 10.
- Le figure: jack (`"J"`), queen (`"Q"`) e king (`"K"`) valgono ciascuna 10.

Semi:

- Ogni carta ha un seme: spade (`"S"`), fiori (`"C"`), quadri (`"D"`) o cuori (`"H"`).

Formato della carta:

- Ogni carta è rappresentata come una stringa: `"valueSuit"`. Per esempio: `"AS"` è l'asso di spade, `"10H"` è il dieci di cuori e `"QC"` è la regina di fiori.

# --hints--

`card_values(["3H", "4D", "5S"])` dovrebbe restituire `[3, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["3H", "4D", "5S"]), [3, 4, 5])`)
}})
```

`card_values(["AS", "10S", "10H", "6D", "7D"])` dovrebbe restituire `[1, 10, 10, 6, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7])`)
}})
```

`card_values(["8D", "QS", "2H", "JC", "9C"])` dovrebbe restituire `[8, 10, 2, 10, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9])`)
}})
```

`card_values(["AS", "KS"])` dovrebbe restituire `[1, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "KS"]), [1, 10])`)
}})
```

`card_values(["10H", "JH", "QH", "KH", "AH"])` dovrebbe restituire `[10, 10, 10, 10, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def card_values(cards):

    return cards
```

# --solutions--

```py
def card_values(cards):
    values = []
    for card in cards:
        value_str = card[:-1]

        if value_str == "A":
            value = 1
        elif value_str in ["J", "Q", "K"]:
            value = 10
        else:
            value = int(value_str)

        values.append(value)
    return values
```
