---
id: 698a1a73ade5ac0e19180fa6
title: "Sfida 203: array appiattito"
challengeType: 29
dashedName: challenge-203
---

# --description--

Dato un array, determina se è appiattito.

- Un array è appiattito se nessuno dei suoi elementi è un array.

# --hints--

`is_flat([1, 2, 3, 4])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, 2, 3, 4]), True)`)
}})
```

`is_flat([1, [2, 3], 4])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, [2, 3], 4]), False)`)
}})
```

`is_flat([1, 0, False, True, "a", "b"])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, 0, False, True, "a", "b"]), True)`)
}})
```

`is_flat(["a", [0], "b", True])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat(["a", [0], "b", True]), False)`)
}})
```

`is_flat([1, [2, [3, [4, [5]]]], 6])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, [2, [3, [4, [5]]]], 6]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_flat(arr):

    return arr
```

# --solutions--

```py
def is_flat(arr):
    for el in arr:
        if isinstance(el, list):
            return False
    return True
```
