---
id: 698a1a863194f1f4e63f6460
title: "Sfida 209: dimensione dell'elemento"
challengeType: 29
dashedName: challenge-209
---

# --description--

Dato un valore per la dimensione della finestra, la larghezza di un elemento in unità `"vw"` di larghezza viewport e l'altezza di un elemento in unità `"vh"` di altezza viewport, determina la dimensione dell'elemento in pixel.

- La dimensione della finestra fornita e la dimensione dell'elemento restituita sono stringhe nel formato `"width x height"`, `"1200 x 800"` per esempio.

- Le unità `"vw"` rappresentano la percentuale della larghezza della finestra. `"50vw"`, per esempio, è il 50% della larghezza della finestra.
- Le unità `"vh"` rappresentano la percentuale dell'altezza della finestra. `"50vh"`, per esempio, è il 50% dell'altezza della finestra.

# --hints--

`get_element_size("1200 x 800", "50vw", "50vh")` dovrebbe restituire `"600 x 400"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "50vw", "50vh"), "600 x 400")`)
}})
```

`get_element_size("320 x 480", "25vw", "50vh")` dovrebbe restituire `"80 x 240"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("320 x 480", "25vw", "50vh"), "80 x 240")`)
}})
```

`get_element_size("1000 x 500", "7vw", "3vh")` dovrebbe restituire `"70 x 15"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1000 x 500", "7vw", "3vh"), "70 x 15")`)
}})
```

`get_element_size("1920 x 1080", "95vw", "100vh")` dovrebbe restituire `"1824 x 1080"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1920 x 1080", "95vw", "100vh"), "1824 x 1080")`)
}})
```

`get_element_size("1200 x 800", "0vw", "0vh")` dovrebbe restituire `"0 x 0"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "0vw", "0vh"), "0 x 0")`)
}})
```

`get_element_size("1440 x 900", "100vw", "114vh")` dovrebbe restituire `"1440 x 1026"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1440 x 900", "100vw", "114vh"), "1440 x 1026")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_element_size(window_size, element_vw, element_vh):

    return window_size
```

# --solutions--

```py
def get_element_size(window_size, element_vw, element_vh):
    window_width, window_height = map(int, window_size.split(" x "))

    vw = int(element_vw.replace("vw", ""))
    vh = int(element_vh.replace("vh", ""))

    width_px = round((vw / 100) * window_width)
    height_px = round((vh / 100) * window_height)

    return f"{width_px} x {height_px}"
```
