---
id: 698a1a73ade5ac0e19180fa0
title: "Sfida 197: compatibilitﾃ dei gruppi sanguigni"
challengeType: 29
dashedName: challenge-197
---

# --description--

Dato un gruppo sanguigno del donatore e un gruppo sanguigno del ricevente, determina se il donatore puﾃｲ donare sangue al ricevente.

Ogni gruppo sanguigno ﾃｨ composto da:

- Una lettera: `"A"`, `"B"`, `"AB"` o `"O"`
- E un fattore Rh: `"+"` o `"-"`

I gruppi sanguigni saranno una delle lettere valide seguite da un fattore Rh. Per esempio, `"AB+"` e `"O-"` sono gruppi sanguigni validi.

Regole per la lettera:

- `"O"` puﾃｲ donare ad altri gruppi con la stessa lettera.
- `"A"` puﾃｲ donare a `"A"` e `"AB"`.
- `"B"` puﾃｲ donare a `"B"` e `"AB"`.
- `"AB"` puﾃｲ donare solo a `"AB"`.

Regole per il fattore Rh:

- Negativo (`"-"`) puﾃｲ donare sia a `"-"` che a `"+"`.
- Positivo (`"+"`) puﾃｲ donare solo a `"+"`.

Sia la regola della lettera che quella del fattore Rh devono essere rispettate affinchﾃｩ un donatore possa donare al ricevente.

# --hints--

`can_donate("B+", "B+")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B+", "B+"), True)`)
}})
```

`can_donate("O-", "AB-")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "AB-"), True)`)
}})
```

`can_donate("O+", "A-")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "A-"), False)`)
}})
```

`can_donate("A+", "AB+")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A+", "AB+"), True)`)
}})
```

`can_donate("A-", "B-")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A-", "B-"), False)`)
}})
```

`can_donate("B-", "AB+")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "AB+"), True)`)
}})
```

`can_donate("B-", "A+")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "A+"), False)`)
}})
```

`can_donate("O-", "O+")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "O+"), True)`)
}})
```

`can_donate("O+", "O-")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "O-"), False)`)
}})
```

`can_donate("AB+", "AB-")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("AB+", "AB-"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_donate(donor, recipient):

    return donor
```

# --solutions--

```py
def can_donate(donor, recipient):
    donor_type = donor[:-1]
    donor_rh = donor[-1]

    recipient_type = recipient[:-1]
    recipient_rh = recipient[-1]

    abo_compatibility = {
        "O": ["A", "B", "AB", "O"],
        "A": ["A", "AB"],
        "B": ["B", "AB"],
        "AB": ["AB"]
    }

    abo_match = recipient_type in abo_compatibility[donor_type]
    rh_match = (
        donor_rh == "-" or
        (donor_rh == "+" and recipient_rh == "+")
    )

    return abo_match and rh_match
```
