---
id: 6994cff2290543b3aec9f510
title: "Sfida 211: somma di un array"
challengeType: 29
dashedName: challenge-211
---

# --description--

Dato un array di numeri, restituisci la somma di tutti i numeri.

# --hints--

`sum_array([1, 2, 3, 4, 5])` dovrebbe restituire `15`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([1, 2, 3, 4, 5]), 15)`)
}})
```

`sum_array([42])` dovrebbe restituire `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([42]), 42)`)
}})
```

`sum_array([5, -2, 7, -3])` dovrebbe restituire `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([5, -2, 7, -3]), 7)`)
}})
```

`sum_array([203, 145, -129, 6293, 523, -919, 845, 2434])` dovrebbe restituire `9395`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395)`)
}})
```

`sum_array([0, 0])` dovrebbe restituire `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([0, 0]), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_array(numbers):

    return numbers
```

# --solutions--

```py
def sum_array(numbers):

    return sum(numbers)
```
