---
id: 699c8e045ee7cb94ed2322dd
title: "Sfida 222: ombre all'equinozio"
challengeType: 29
dashedName: challenge-222
---

# --description--

Oggi è l'equinozio, quando il sole è direttamente sopra l'equatore e perfettamente allo zenit a mezzogiorno. Data un'ora, determina l'ombra proiettata da un palo verticale alto 4 piedi.

- L'ora sarà una stringa in formato 24 ore `"HH:MM"` (per esempio, `"15:00"` indica le 15:00).
- Ti verrà data solo un'ora con incrementi di 30 minuti.

Regole:

- Il sole sorge alle 6:00 direttamente `"east"` e tramonta alle 18:00 direttamente `"west"`.
- Un'ombra punta sempre nella direzione opposta al sole.
- La lunghezza dell'ombra (in piedi) è il numero di ore di distanza da mezzogiorno elevato al cubo.
- Non c'è ombra prima dell'alba (prima delle 6:00), dopo il tramonto (18:00 o oltre) o a mezzogiorno.

Restituisci:

- Se esiste un'ombra, restituisci `"(length)ft (direction)"`. Per esempio, `"8ft west"`.
- Altrimenti, restituisci `"No shadow"`.

Per esempio, data `"10:00"`, restituisci `"8ft west"` perché le 10:00 sono 2 ore da mezzogiorno, quindi 2<sup>3</sup> = 8 piedi, e l'ombra punta a ovest perché il sole è a est alle 10:00.

# --hints--

`get_shadow("10:00")` dovrebbe restituire `"8ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` dovrebbe restituire `"27ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` dovrebbe restituire `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` dovrebbe restituire `"166.375ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` dovrebbe restituire `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` dovrebbe restituire `"216ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` dovrebbe restituire `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` dovrebbe restituire `"91.125ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` dovrebbe restituire `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
