---
id: 699c8e045ee7cb94ed2322db
title: "Sfida 220: numero più grande"
challengeType: 29
dashedName: challenge-220
---

# --description--

Data una stringa di numeri separati da vari segni di punteggiatura, restituisci il numero più grande.

- La stringa data conterrà solo numeri e separatori.
- I separatori possono essere virgole (`","`), punti esclamativi (`"!"`), punti interrogativi (`"?"`), due punti (`":"`) o punti e virgola (`";"`).

# --hints--

`largest_number("1,2")` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("1,2"), 2)`)
}})
```

`largest_number("4;15:60,26?52!0")` dovrebbe restituire `60`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("4;15:60,26?52!0"), 60)`)
}})
```

`largest_number("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` dovrebbe restituire `-402`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402)`)
}})
```

`largest_number("12;-50,99.9,49.1!-10.1?88?16")` dovrebbe restituire `99.9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("12;-50,99.9,49.1!-10.1?88?16"), 99.9)`)
}})
```

# --seed--

## --seed-contents--

```py
def largest_number(s):

    return s
```

# --solutions--

```py
import re
def largest_number(s):
    numbers = [float(n) for n in re.split(r'[,!?:;]', s)]
    return max(numbers)
```
