---
id: 699c8e045ee7cb94ed2322d5
title: "Sfida 214: validatore di catena di domino"
challengeType: 29
dashedName: challenge-214
---

# --description--

Dato un array 2D che rappresenta una sequenza di domino, determina se forma una catena valida.

- Ogni elemento nell'array rappresenta un domino ed è un array di due numeri da 1 a 6 (inclusi).
- Perché la catena sia valida, il secondo numero di ogni domino deve corrispondere al primo numero del domino successivo.
- Il primo numero del primo domino e l'ultimo numero dell'ultimo domino non devono corrispondere a nulla.

# --hints--

`is_valid_domino_chain([[1, 3], [3, 6], [6, 5]])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[1, 3], [3, 6], [6, 5]]), True)`)
}})
```

`is_valid_domino_chain([[6, 2], [3, 4], [4, 1]])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[6, 2], [3, 4], [4, 1]]), False)`)
}})
```

`is_valid_domino_chain([[2, 5], [5, 6], [5, 1]])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 5], [5, 6], [5, 1]]), False)`)
}})
```

`is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]), True)`)
}})
```

`is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_domino_chain(dominoes):

    return dominoes
```

# --solutions--

```py
def is_valid_domino_chain(dominoes):
    for i in range(len(dominoes) - 1):
        if dominoes[i][1] != dominoes[i + 1][0]:
            return False

    return True
```
