---
id: 699c8e045ee7cb94ed2322d8
title: "Sfida 217: pezzi degli scacchi catturati"
challengeType: 29
dashedName: challenge-217
---

# --description--

Dato un array di stringhe che rappresentano i pezzi degli scacchi che hai ancora sulla scacchiera, calcola il valore dei pezzi che il tuo avversario ha catturato.

Negli scacchi, inizi con 16 pezzi:

| Pezzo  | Abbreviazione | Quantità | Valore |
| ------ | ------------- | -------- | ------ |
| Pedone | `"P"`     | 8        | 1      |
| Torre  | `"R"`     | 2        | 5      |
| Cavallo| `"N"`     | 2        | 3      |
| Alfiere| `"B"`     | 2        | 3      |
| Regina | `"Q"`     | 1        | 9      |
| Re     | `"K"`     | 1        | 0      |

- L'array dato conterrà solo le abbreviazioni sopra indicate.
- Qualsiasi dei 16 pezzi non incluso nell'array dato è stato catturato.
- Restituisci il valore totale di tutti i pezzi catturati, a meno che...
- Se il Re è stato catturato, restituisci `"Checkmate"`.

# --hints--

`get_captured_value(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` dovrebbe restituire `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8)`)
}})
```

`get_captured_value(["P", "P", "P", "P", "P", "R", "B", "K"])` dovrebbe restituire `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "P", "P", "P", "P", "R", "B", "K"]), 26)`)
}})
```

`get_captured_value(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` dovrebbe restituire `16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16)`)
}})
```

`get_captured_value(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` dovrebbe restituire `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4)`)
}})
```

`get_captured_value(["P", "K"])` dovrebbe restituire `38`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "K"]), 38)`)
}})
```

`get_captured_value(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` dovrebbe restituire `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0)`)
}})
```

`get_captured_value(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` dovrebbe restituire `"Checkmate"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_captured_value(pieces):

    return pieces
```

# --solutions--

```py
def get_captured_value(pieces):
    if "K" not in pieces:
        return "Checkmate"

    values = {"P":1, "R":5, "N":3, "B":3, "Q":9, "K":0}
    remaining = "PPPPPPPPRRNNBBQK"

    for p in pieces:
        remaining = remaining.replace(p, "", 1)

    return sum(values[p] for p in remaining)
```
