---
id: 699c8e045ee7cb94ed2322da
title: "Sfida 219: traguardi anniversari"
challengeType: 29
dashedName: challenge-219
---

# --description--

Dato un numero intero che rappresenta gli anni di matrimonio di una coppia, restituisci il loro traguardo anniversario più recente secondo questa tabella:

| Anni di matrimonio | Traguardo    |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- Se non hanno raggiunto il primo traguardo, restituisci `"Newlyweds"`.

# --hints--

`get_milestone(0)` dovrebbe restituire `"Newlyweds"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(0), "Newlyweds")`)
}})
```

`get_milestone(1)` dovrebbe restituire `"Paper"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(1), "Paper")`)
}})
```

`get_milestone(8)` dovrebbe restituire `"Wood"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(8), "Wood")`)
}})
```

`get_milestone(10)` dovrebbe restituire `"Tin"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(10), "Tin")`)
}})
```

`get_milestone(26)` dovrebbe restituire `"Silver"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(26), "Silver")`)
}})
```

`get_milestone(45)` dovrebbe restituire `"Ruby"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(45), "Ruby")`)
}})
```

`get_milestone(50)` dovrebbe restituire `"Gold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(50), "Gold")`)
}})
```

`get_milestone(64)` dovrebbe restituire `"Diamond"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(64), "Diamond")`)
}})
```

`get_milestone(71)` dovrebbe restituire `"Platinum"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(71), "Platinum")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_milestone(years):

    return years
```

# --solutions--

```py
def get_milestone(years):
    if years < 1:
        return "Newlyweds"

    milestones = [
        (1, "Paper"),
        (5, "Wood"),
        (10, "Tin"),
        (25, "Silver"),
        (40, "Ruby"),
        (50, "Gold"),
        (60, "Diamond"),
        (70, "Platinum")
    ]

    for year, milestone in reversed(milestones):
        if years >= year:
            return milestone
```
