---
id: 699c8e045ee7cb94ed2322dc
title: "Sfida 221: matrice invertita"
challengeType: 29
dashedName: challenge-221
---

# --description--

Data una matrice (un array di array) riempita con due valori distinti, restituisci una nuova matrice in cui tutte le occorrenze di un valore sono scambiate con l'altro.

Ad esempio, dato:

```js
[
  ["a", "b"],
  ["a", "a"]
]
```

Restituisci:

```js
[
  ["b", "a"],
  ["b", "b"]
]
```

# --hints--

`invert_matrix([["a", "b"], ["a", "a"]])` dovrebbe restituire `[["b", "a"], ["b", "b"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(invert_matrix([["a", "b"], ["a", "a"]]), [["b", "a"], ["b", "b"]])`)
}})
```

`invert_matrix([[1, 0, 1], [1, 1, 1], [0, 1, 0]])` dovrebbe restituire `[[0, 1, 0], [0, 0, 0], [1, 0, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(invert_matrix([[1, 0, 1], [1, 1, 1], [0, 1, 0]]), [[0, 1, 0], [0, 0, 0], [1, 0, 1]])`)
}})
```

`invert_matrix([["apple", "banana", "banana", "apple"], ["banana", "apple", "apple", "banana"], ["banana", "banana", "banana", "apple"]])` dovrebbe restituire `[["banana", "apple", "apple", "banana"], ["apple", "banana", "banana", "apple"], ["apple", "apple", "apple", "banana"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(invert_matrix([["apple", "banana", "banana", "apple"], ["banana", "apple", "apple", "banana"], ["banana", "banana", "banana", "apple"]]), [["banana", "apple", "apple", "banana"], ["apple", "banana", "banana", "apple"], ["apple", "apple", "apple", "banana"]])`)
}})
```

`invert_matrix([[6, 7, 7, 7, 6], [7, 6, 7, 6, 7], [7, 7, 6, 7, 7], [7, 6, 7, 6, 7], [6, 7, 7, 7, 6]])` dovrebbe restituire `[[7, 6, 6, 6, 7], [6, 7, 6, 7, 6], [6, 6, 7, 6, 6], [6, 7, 6, 7, 6], [7, 6, 6, 6, 7]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(invert_matrix([[6, 7, 7, 7, 6], [7, 6, 7, 6, 7], [7, 7, 6, 7, 7], [7, 6, 7, 6, 7], [6, 7, 7, 7, 6]]), [[7, 6, 6, 6, 7], [6, 7, 6, 7, 6], [6, 6, 7, 6, 6], [6, 7, 6, 7, 6], [7, 6, 6, 6, 7]])`)
}})
```

`invert_matrix([[1.2, 2.1, 2.1, 2.1], [2.1, 1.2, 2.1, 1.2], [1.2, 1.2, 2.1, 2.1]])` dovrebbe restituire `[[2.1, 1.2, 1.2, 1.2], [1.2, 2.1, 1.2, 2.1], [2.1, 2.1, 1.2, 1.2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(invert_matrix([[1.2, 2.1, 2.1, 2.1], [2.1, 1.2, 2.1, 1.2], [1.2, 1.2, 2.1, 2.1]]), [[2.1, 1.2, 1.2, 1.2], [1.2, 2.1, 1.2, 2.1], [2.1, 2.1, 1.2, 1.2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def invert_matrix(matrix):

    return matrix
```

# --solutions--

```py
def invert_matrix(matrix):
    flat = [cell for row in matrix for cell in row]
    values = list(set(flat))
    val1, val2 = values

    return [[val2 if cell == val1 else val1 for cell in row] for row in matrix]
```
