---
id: 699c8e045ee7cb94ed2322d6
title: "Sfida 215: calcolatore della tariffa del parcheggio"
challengeType: 29
dashedName: challenge-215
---

# --description--

Date due stringhe che rappresentano l'orario in cui hai parcheggiato la macchina e l'orario in cui l'hai ritirata, calcola la tariffa del parcheggio.

- Le stringhe date saranno nel formato `"HH:MM"` usando un orologio a 24 ore. Quindi `"14:00"` è per esempio le 14:00.
- La prima stringa sarà l'orario in cui hai parcheggiato la macchina, la seconda sarà l'orario in cui l'hai ritirata.
- Se l'orario di ritiro è precedente a quello di entrata, significa che il parcheggio è durato oltre la mezzanotte nel giorno successivo.

Regole della tariffa:

- Ogni ora di parcheggio costa 3$.
- Le ore parziali si arrotondano per eccesso all'ora intera successiva.
- Se il parcheggio dura tutta la notte (oltre la mezzanotte), si applica una tariffa aggiuntiva di 10$ per la notte.
- C'è una tariffa minima di 5$ (usata solo se il totale sarebbe inferiore a 5$).

Restituisci il costo totale nel formato `"$cost"`, `"$5"` per esempio.

# --hints--

`calculate_parking_fee("09:00", "11:00")` dovrebbe restituire `"$6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("09:00", "11:00"), "$6")`)
}})
```

`calculate_parking_fee("10:00", "10:30")` dovrebbe restituire `"$5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("10:00", "10:30"), "$5")`)
}})
```

`calculate_parking_fee("08:10", "10:45")` dovrebbe restituire `"$9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("08:10", "10:45"), "$9")`)
}})
```

`calculate_parking_fee("14:40", "23:10")` dovrebbe restituire `"$27"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("14:40", "23:10"), "$27")`)
}})
```

`calculate_parking_fee("18:15", "01:30")` dovrebbe restituire `"$34"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("18:15", "01:30"), "$34")`)
}})
```

`calculate_parking_fee("11:11", "11:10")` dovrebbe restituire `"$82"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("11:11", "11:10"), "$82")`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_parking_fee(park_time, pickup_time):

    return park_time
```

# --solutions--

```py
import math

def calculate_parking_fee(park_time, pickup_time):
    def to_minutes(time_str):
        hours, minutes = map(int, time_str.split(":"))
        return hours * 60 + minutes

    entry_minutes = to_minutes(park_time)
    exit_minutes = to_minutes(pickup_time)

    overnight = False

    if exit_minutes < entry_minutes:
        overnight = True
        total_minutes = (24 * 60 - entry_minutes) + exit_minutes
    else:
        total_minutes = exit_minutes - entry_minutes

    hours = math.ceil(total_minutes / 60)

    cost = hours * 3

    if overnight:
        cost += 10

    if cost < 5:
        cost = 5

    return f"${cost}"
```
