---
id: 69a890af247de743333bd4ce
title: "Sfida 225: nessuna ripetizione consecutiva"
challengeType: 29
dashedName: challenge-225
---

# --description--

Data una stringa, determina se non contiene caratteri ripetuti.

- Una stringa non ha ripetizioni se non presenta lo stesso carattere due o più volte di seguito.

# --hints--

`has_no_repeats("hi world")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("hi world"), True)`)
}})
```

`has_no_repeats("hello world")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("hello world"), False)`)
}})
```

`has_no_repeats("abcdefghijklmnopqrstuvwxyz")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("abcdefghijklmnopqrstuvwxyz"), True)`)
}})
```

`has_no_repeats("freeCodeCamp")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("freeCodeCamp"), False)`)
}})
```

`has_no_repeats("The quick brown fox jumped over the lazy dog.")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("The quick brown fox jumped over the lazy dog."), True)`)
}})
```

`has_no_repeats("Mississippi")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("Mississippi"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_no_repeats(s):

    return s
```

# --solutions--

```py
def has_no_repeats(s):
    for i in range(1, len(s)):
        if s[i] == s[i-1]:
            return False
    return True
```
