---
id: 69a890af247de743333bd4cd
title: "Sfida 224: rilevatore di tostatura del caffè"
challengeType: 29
dashedName: challenge-224
---

# --description--

Data una stringa che rappresenta i chicchi usati per preparare una tazza di caffè, determina la tostatura della tazza.

- La stringa fornita conterrà i seguenti caratteri, ognuno rappresentante un tipo di chicco:

  - Un apostrofo (`'`) è un chicco a tostatura chiara che vale 1 punto ciascuno.
  - Un trattino (`-`) è un chicco a tostatura media che vale 2 punti ciascuno.
  - Un punto (`.`) è un chicco a tostatura scura che vale 3 punti ciascuno.

- Il livello di tostatura si determina dalla media di tutti i chicchi.

Restituisci:

- `"Light"` se la media è inferiore a 1,75.
- `"Medium"` se la media è compresa tra 1,75 e 2,5.
- `"Dark"` se la media è superiore a 2,5.

# --hints--

`detect_roast("''-''''''-'-''--''''")` dovrebbe restituire `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("''-''''''-'-''--''''"), "Light")`)
}})
```

`detect_roast(".'-''-''..'''.-.-''-")` dovrebbe restituire `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".'-''-''..'''.-.-''-"), "Medium")`)
}})
```

`detect_roast("--.''--'-''.--..-.--")` dovrebbe restituire `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("--.''--'-''.--..-.--"), "Medium")`)
}})
```

`detect_roast("-...'-......-..-...-")` dovrebbe restituire `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-...'-......-..-...-"), "Dark")`)
}})
```

`detect_roast(".--.-..-......----.'")` dovrebbe restituire `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".--.-..-......----.'"), "Medium")`)
}})
```

`detect_roast("..-..-..-..-....-.-.")` dovrebbe restituire `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("..-..-..-..-....-.-."), "Dark")`)
}})
```

`detect_roast("-'-''''''..-'.''-'.'")` dovrebbe restituire `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-'-''''''..-'.''-'.'"), "Light")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_roast(beans):

    return beans
```

# --solutions--

```py
def detect_roast(beans):
    total = 0

    for bean in beans:
        if bean == "'":
            total += 1
        elif bean == "-":
            total += 2
        elif bean == ".":
            total += 3

    avg = total / len(beans)

    if avg < 1.75:
        return "Light"
    elif avg <= 2.5:
        return "Medium"
    else:
        return "Dark"
```
