---
id: 69a890af247de743333bd4d2
title: "Sfida 229: troncare il testo 2"
challengeType: 29
dashedName: challenge-229
---

# --description--

Data una stringa, restituisci una nuova stringa troncata in modo che la larghezza totale dei caratteri non superi 50 unità.

Ogni carattere ha una larghezza specifica:

| Lettere | Larghezza |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

La tabella sopra include tutte le lettere maiuscole e minuscole. Inoltre:

- Gli spazi (`" "`) hanno una larghezza di 2
- I punti (`"."`) hanno una larghezza di 1

- Se la stringa data è di 50 unità o meno, restituisci la stringa così com'è, altrimenti
- Tronca la stringa e aggiungi tre punti alla fine (`"..."`) in modo che la larghezza totale, inclusi i tre punti, sia il più possibile vicina a 60 unità senza superarla.

# --hints--

`truncate_text("The quick brown fox")` dovrebbe restituire `"The quick brown f..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The quick brown fox"), "The quick brown f...")`)
}})
```

`truncate_text("The silky smooth sloth")` dovrebbe restituire `truncate_text("The silky smooth sloth")`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The silky smooth sloth"), truncate_text("The silky smooth sloth"))`)
}})
```

`truncate_text("THE LOUD BRIGHT BIRD")` dovrebbe restituire `"THE LOUD BRIG..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...")`)
}})
```

`truncate_text("The fast striped zebra")` dovrebbe restituire `"The fast striped z..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The fast striped zebra"), "The fast striped z...")`)
}})
```

`truncate_text("The big black bear")` dovrebbe restituire `"The big black bear"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The big black bear"), "The big black bear")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(s):

    return s
```

# --solutions--

```py
def truncate_text(s):
    MAX_WIDTH = 50
    ELLIPSIS = "..."
    ELLIPSIS_WIDTH = 3
    TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH

    char_groups = {
        "ilI.": 1,
        "fjrt ": 2,
        "abcdeghkmnopqrstuvwxyzJL": 3,
        "ABCDEFGHKMNOPQRSTUVWXYZ": 4
    }

    def get_char_width(char):
        for key, width in char_groups.items():
            if char in key:
                return width
        return 3

    def string_width(string):
        total = 0
        for char in string:
            total += get_char_width(char)
        return total

    if string_width(s) <= MAX_WIDTH:
        return s

    result = ""
    total_width = 0

    for char in s:
        char_width = get_char_width(char)

        if total_width + char_width > TRUNCATE_LIMIT:
            break

        result += char
        total_width += char_width

    return result + ELLIPSIS
```
