---
id: 69a890af247de743333bd4d1
title: "Sfida 228: serata al cinema"
challengeType: 29
dashedName: challenge-228
---

# --description--

Dato un giorno della settimana come stringa, un'altra stringa per l'orario dello spettacolo e un numero intero di biglietti, restituisci il costo totale dei biglietti per quella proiezione.

Il giorno fornito sarà uno tra:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

L'orario dello spettacolo sarà fornito nel formato `"H:MMam"` o `"H:MMpm"`. Per esempio `"10:00am"` o `"10:00pm"`.

Restituisci il costo totale nel formato `"$D.CC"` seguendo queste regole:

- Weekend (venerdì - domenica): 12,00 $ per biglietto.
- Giorni feriali (lunedì - giovedì): 10,00 $ per biglietto.
- Matinée (prima delle 17:00): sconto di 2,00 $ per biglietto (tranne il martedì).
- Martedì: tutti i biglietti costano 5,00 $ ciascuno.

# --hints--

`get_movie_night_cost("Saturday", "10:00pm", 1)` dovrebbe restituire `"$12.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Saturday", "10:00pm", 1), "$12.00")`)
}})
```

`get_movie_night_cost("Sunday", "10:00am", 1)` dovrebbe restituire `"$10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Sunday", "10:00am", 1), "$10.00")`)
}})
```

`get_movie_night_cost("Tuesday", "7:20pm", 2)` dovrebbe restituire `"$10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "7:20pm", 2), "$10.00")`)
}})
```

`get_movie_night_cost("Wednesday", "5:40pm", 3)` dovrebbe restituire `"$30.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Wednesday", "5:40pm", 3), "$30.00")`)
}})
```

`get_movie_night_cost("Monday", "11:50am", 4)` dovrebbe restituire `"$32.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Monday", "11:50am", 4), "$32.00")`)
}})
```

`get_movie_night_cost("Friday", "4:30pm", 5)` dovrebbe restituire `"$50.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Friday", "4:30pm", 5), "$50.00")`)
}})
```

`get_movie_night_cost("Tuesday", "11:30am", 1)` dovrebbe restituire `"$5.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "11:30am", 1), "$5.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):

    return day
```

# --solutions--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):
    if day == "Tuesday":
        price = 5
    else:
        if day in ["Friday", "Saturday", "Sunday"]:
            price = 12
        else:
            price = 10

        hour = int(showtime.split(":")[0])
        period = showtime[-2:]

        hour24 = hour

        if period == "pm" and hour != 12:
            hour24 += 12
        if period == "am" and hour == 12:
            hour24 = 0

        if hour24 < 17:
            price -= 2

    total = price * number_of_tickets

    return f"${total:.2f}"
```
