---
id: 69a890af247de743333bd4d0
title: "Sfida 227: tempo di raffreddamento"
challengeType: 29
dashedName: challenge-227
---

# --description--

Date due date temporali, la prima che indica quando un utente ha terminato un esame e la seconda che indica l'ora attuale, determina se l'utente può sostenere di nuovo un esame.

- Entrambe le date saranno fornite nel formato: `"YYYY-MM-DDTHH:MM:SS"`, per esempio `"2026-03-25T14:00:00"`. Nota che l'ora è nel formato 24 ore.
- Un utente deve aspettare almeno 48 ore prima di ripetere un esame.

# --hints--

`can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00"), True)`)
}})
```

`can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00"), False)`)
}})
```

`can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00"), True)`)
}})
```

`can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_retake(finish_time, current_time):

    return finish_time
```

# --solutions--

```py
from datetime import datetime, timedelta

def can_retake(finish_time, current_time):
    fmt = "%Y-%m-%dT%H:%M:%S"
    finished_dt = datetime.strptime(finish_time, fmt)
    current_dt = datetime.strptime(current_time, fmt)
    diff = current_dt - finished_dt
    cooldown = timedelta(hours=48)

    return diff >= cooldown
```
