---
id: 69b1028d6e265413d0198a2c
title: "Sfida 233: sveglia"
challengeType: 29
dashedName: challenge-233
---

# --description--

Dato una stringa che rappresenta l'ora in cui hai impostato la sveglia e una stringa che rappresenta l'ora in cui ti sei effettivamente svegliato, determina se ti sei svegliato presto, in orario o in ritardo.

- Entrambi gli orari saranno forniti nel formato 24 ore `"HH:MM"`.

Restituisci:

- `"early"` se ti sei svegliato prima dell'ora della sveglia.
- `"on time"` se ti sei svegliato all'ora della sveglia o entro la finestra di 10 minuti di snooze dopo l'ora della sveglia.
- `"late"` se ti sei svegliato più di 10 minuti dopo l'ora della sveglia.

Entrambi gli orari sono nello stesso giorno.

# --hints--

`alarm_check("07:00", "06:45")` dovrebbe restituire `"early"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("07:00", "06:45"), "early")`)
}})
```

`alarm_check("06:30", "06:30")` dovrebbe restituire `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("06:30", "06:30"), "on time")`)
}})
```

`alarm_check("08:10", "08:15")` dovrebbe restituire `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("08:10", "08:15"), "on time")`)
}})
```

`alarm_check("09:30", "09:45")` dovrebbe restituire `"late"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("09:30", "09:45"), "late")`)
}})
```

`alarm_check("08:15", "08:25")` dovrebbe restituire `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("08:15", "08:25"), "on time")`)
}})
```

`alarm_check("05:45", "05:56")` dovrebbe restituire `"late"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("05:45", "05:56"), "late")`)
}})
```

`alarm_check("04:30", "04:00")` dovrebbe restituire `"early"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("04:30", "04:00"), "early")`)
}})
```

# --seed--

## --seed-contents--

```py
def alarm_check(alarm_time, wake_time):

    return alarm_time
```

# --solutions--

```py
def alarm_check(alarm_time, wake_time):
    alarm_h, alarm_m = map(int, alarm_time.split(":"))
    wake_h, wake_m = map(int, wake_time.split(":"))

    alarm_minutes = alarm_h * 60 + alarm_m
    wake_minutes = wake_h * 60 + wake_m
    diff = wake_minutes - alarm_minutes

    if diff < 0:
        return "early"
    if diff <= 10:
        return "on time"
    return "late"
```
