---
id: 69b1028d6e265413d0198a2f
title: "Sfida 236: cronologia del browser"
challengeType: 29
dashedName: challenge-236
---

# --description--

Dato un array di comandi del browser, restituisci un array con due valori: la cronologia come array di URL e l'indice della pagina corrente.

I comandi validi sono:

- `"URL"` - Dove URL è un indirizzo web (ad esempio `"freecodecamp.org"`). Naviga all'URL indicato, lo aggiunge alla cronologia nella posizione successiva e scarta qualsiasi cronologia futura.
- `"Back"` - si sposta alla pagina precedente nella cronologia, oppure rimane sulla pagina corrente se non ce n'è una.
- `"Forward"` - si sposta alla pagina successiva nella cronologia, oppure rimane sulla pagina corrente se non ce n'è una.

Per esempio, dato `["freecodecamp.org", "freecodecamp.org/learn", "Back"]`, restituisci `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

# --hints--

`get_browser_history(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` dovrebbe restituire `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` dovrebbe restituire `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"])` dovrebbe restituire `[["example.com", "example.com/contact", "example.com/blog"], 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` dovrebbe restituire `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Back", "Back"])` dovrebbe restituire `[["example.com", "example.com/about"], 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Forward"])` dovrebbe restituire `[["example.com", "example.com/about"], 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_browser_history(commands):

    return commands
```

# --solutions--

```py
def get_browser_history(commands):
    stack = []
    index = -1

    for command in commands:
        if command == "Back":
            if index > 0:
                index -= 1
        elif command == "Forward":
            if index < len(stack) - 1:
                index += 1
        else:
            del stack[index + 1:]
            stack.append(command)
            index += 1

    return [stack, index]
```
