---
id: 69b5b2be76ec8135a7fbe973
title: "Sfida 244: attacco di torre e alfiere"
challengeType: 29
dashedName: challenge-244
---

# --description--

Dato un valore stringa per la posizione di una torre su una scacchiera e un altro per la posizione di un alfiere, determina se un pezzo può attaccare l'altro.

Una scacchiera standard è 8x8, con le colonne etichettate `A` fino a `H` (da sinistra a destra) e le righe etichettate `1` fino a `8` (dal basso verso l'alto). Si presenta così:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- Le torri possono muoversi di quante caselle vogliono in orizzontale o verticale.
- Gli alfieri possono muoversi di quante caselle vogliono in qualsiasi direzione diagonale.
- Un pezzo può attaccare un altro se può muoversi nella posizione occupata da quell'altro pezzo.

Restituisci:

- `"rook"` se la torre può attaccare l'alfiere.
- `"bishop"` se l'alfiere può attaccare la torre.
- `"neither"` se nessuno dei due pezzi può attaccare l'altro.

# --hints--

`rook_bishop_attack("A1", "A5")` dovrebbe restituire `"rook"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("A1", "A5"), "rook")`)
}})
```

`rook_bishop_attack("C3", "F6")` dovrebbe restituire `"bishop"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("C3", "F6"), "bishop")`)
}})
```

`rook_bishop_attack("D4", "D7")` dovrebbe restituire `"rook"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("D4", "D7"), "rook")`)
}})
```

`rook_bishop_attack("B7", "H1")` dovrebbe restituire `"bishop"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B7", "H1"), "bishop")`)
}})
```

`rook_bishop_attack("B3", "C5")` dovrebbe restituire `"neither"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B3", "C5"), "neither")`)
}})
```

`rook_bishop_attack("G3", "E8")` dovrebbe restituire `"neither"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("G3", "E8"), "neither")`)
}})
```

# --seed--

## --seed-contents--

```py
def rook_bishop_attack(rook, bishop):

    return rook
```

# --solutions--

```py
def rook_bishop_attack(rook, bishop):
    col_diff = abs(ord(rook[0]) - ord(bishop[0]))
    row_diff = abs(int(rook[1]) - int(bishop[1]))

    if rook[0] == bishop[0] or rook[1] == bishop[1]:
        return "rook"
    if col_diff == row_diff:
        return "bishop"
    return "neither"
```
