---
id: 69b58ce40693f140c84c8559
title: "Sfida 240: caratteri palindromi"
challengeType: 29
dashedName: challenge-240
---

# --description--

Data una stringa, determina se è un palindromo e restituisci il carattere centrale (se la lunghezza è dispari) o i due caratteri centrali (se la lunghezza è pari).

- Un palindromo è una stringa che si legge uguale da sinistra a destra e da destra a sinistra.
- Se non è un palindromo, restituisci `"none"`.

# --hints--

`palindrome_locator("racecar")` dovrebbe restituire `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("racecar"), "e")`)
}})
```

`palindrome_locator("level")` dovrebbe restituire `"v"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("level"), "v")`)
}})
```

`palindrome_locator("freecodecamp")` dovrebbe restituire `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("freecodecamp"), "none")`)
}})
```

`palindrome_locator("noon")` dovrebbe restituire `"oo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("noon"), "oo")`)
}})
```

`palindrome_locator("11100111")` dovrebbe restituire `"00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("11100111"), "00")`)
}})
```

# --seed--

## --seed-contents--

```py
def palindrome_locator(s):

    return s
```

# --solutions--

```py
def palindrome_locator(s):
    if s != s[::-1]:
        return "none"
    mid = len(s) // 2
    return s[mid] if len(s) % 2 == 1 else s[mid - 1] + s[mid]
```
