---
id: 69b58ce40693f140c84c855a
title: "Sfida 241: validatore fizzbuzz"
challengeType: 29
dashedName: challenge-241
---

# --description--

Dato un array di numeri interi sequenziali, con i multipli di 3 e 5 sostituiti, determina se è una sequenza FizzBuzz valida.

In una sequenza FizzBuzz valida:

- I multipli di 3 sono sostituiti con "Fizz".
- I multipli di 5 sono sostituiti con "Buzz".
- I multipli sia di 3 che di 5 sono sostituiti con "FizzBuzz".
- Tutti gli altri numeri rimangono numeri interi.

# --hints--

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz"])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz"]), True)`)
}})
```

`is_fizz_buzz([13, 14, "FizzBuzz", 16, 17])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([13, 14, "FizzBuzz", 16, 17]), True)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, 5])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, 5]), False)`)
}})
```

`is_fizz_buzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]), True)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", "Buzz", 5])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", "Buzz", 5]), False)`)
}})
```

`is_fizz_buzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]), False)`)
}})
```

`is_fizz_buzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_fizz_buzz(arr):

    return arr
```

# --solutions--

```py
def is_fizz_buzz(arr):
    start = None
    for i, val in enumerate(arr):
        if isinstance(val, int):
            start = val - i
            break

    if start is None:
        return False

    for i in range(len(arr)):
        n = start + i
        if n % 15 == 0:
            expected = "FizzBuzz"
        elif n % 3 == 0:
            expected = "Fizz"
        elif n % 5 == 0:
            expected = "Buzz"
        else:
            expected = n
        if arr[i] != expected:
            return False

    return True
```
