---
id: 69bc6cb30c1d112a2e110a07
title: "Sfida 250: chiave nascosta"
challengeType: 29
dashedName: challenge-250
---

# --description--

Benvenuto alla 250ª sfida quotidiana!

Data una stringa codificata, decodificala usando una chiave di cifratura e restituiscila.

Per trovare la chiave:

- Guarda tutte le sfide quotidiane fino a oggi il cui numero di sfida è un multiplo di 25 (inclusa questa).
- Prendi la prima lettera di ciascuno di quei titoli di sfida e combinane una stringa. Se il titolo inizia con un carattere non alfabetico, trova la sua prima lettera.

Per decodificare il messaggio, esamina ogni lettera nel messaggio codificato e:

  1. Guarda la lettera corrispondente nella chiave (ripeti la chiave se il messaggio è più lungo della chiave).
  2. Converti la lettera della chiave nel suo numero corrispondente: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
  3. Sposta la lettera codificata indietro nell'alfabeto di quel numero.
  4. Se lo spostamento va oltre `"A"`, torna indietro fino a `"Z"`.

Per esempio, se il messaggio codificato inizia con `"Y"` e la prima lettera della chiave è `"V"` (22), sposta `"Y"` indietro di 22 posizioni per ottenere `"C"`. Ripeti questo processo per ogni lettera per decodificare il messaggio completo.

- Solo le lettere vengono spostate, gli spazi vengono restituiti così come sono.
- Tutte le lettere date e restituite sono maiuscole.

# --hints--

`decode("YAVJYNXE")` dovrebbe restituire `"CONGRATS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")` dovrebbe restituire `"CODING LEGEND"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")` dovrebbe restituire `"YOU ARE AWESOME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")` dovrebbe restituire `"KEEP GOING"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` dovrebbe restituire `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```
