---
id: 69c5f3d787b1725d5f00c8bd
title: "Sfida 259: esplosione fizzbuzz"
challengeType: 29
dashedName: challenge-259
---

# --description--

Dato un numero intero, restituisci il numero di passaggi necessari per trasformare la parola `"fizzbuzz"` in una stringa con almeno il numero dato di `"z"` usando le seguenti regole:

- Inizia con la stringa `"fizzbuzz"`.
- Ad ogni passaggio, applica le regole standard di FizzBuzz usando la posizione della lettera nella stringa (la prima `"f"` è in posizione 1).
  - Se la posizione della lettera è divisibile per 3, sostituisci la lettera con `"fizz"`
  - Se è divisibile per 5, sostituisci la lettera con `"buzz"`
  - Se è divisibile per 3 e 5, sostituisci la lettera con `"fizzbuzz"`

Quindi dopo 1 passaggio, `"fizzbuzz"` diventa `"fifizzzbuzzfizzzz"`, che contiene 9 `"z"`.

# --hints--

`explode_fizzbuzz(9)` dovrebbe restituire `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

`explode_fizzbuzz(15)` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

`explode_fizzbuzz(51)` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

`explode_fizzbuzz(52)` dovrebbe restituire `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

`explode_fizzbuzz(359)` dovrebbe restituire `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

`explode_fizzbuzz(789)` dovrebbe restituire `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

`explode_fizzbuzz(54482)` dovrebbe restituire `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

`explode_fizzbuzz(1000000)` dovrebbe restituire `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```
