---
id: 69c5f3d787b1725d5f00c8ba
title: "Sfida 256: direzione del tempo più vicina"
challengeType: 29
dashedName: challenge-256
---

# --description--

Date due ore, determina se puoi passare dalla prima alla seconda più velocemente andando avanti o indietro.

- Le ore sono date nel formato 24 ore (`"HH:MM"`)
- L'orologio fa il giro completo (23:59 passa a 00:00 andando avanti, e 00:00 passa a 23:59 andando indietro)

Restituisci:

- `"forward"` se andare avanti è più breve
- `"backward"` se andare indietro è più breve
- `"equal"` se entrambe le direzioni richiedono lo stesso tempo

# --hints--

`get_direction("10:00", "12:00")` dovrebbe restituire `"forward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("10:00", "12:00"), "forward")`)
}})
```

`get_direction("11:00", "05:00")` dovrebbe restituire `"backward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("11:00", "05:00"), "backward")`)
}})
```

`get_direction("00:00", "12:00")` dovrebbe restituire `"equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("00:00", "12:00"), "equal")`)
}})
```

`get_direction("15:45", "01:10")` dovrebbe restituire `"forward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("15:45", "01:10"), "forward")`)
}})
```

`get_direction("03:30", "19:50")` dovrebbe restituire `"backward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("03:30", "19:50"), "backward")`)
}})
```

`get_direction("06:30", "18:30")` dovrebbe restituire `"equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("06:30", "18:30"), "equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_direction(time1, time2):

    return time1
```

# --solutions--

```py
def get_direction(time1, time2):
    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    start = to_minutes(time1)
    end = to_minutes(time2)
    forward = (end - start + 1440) % 1440
    backward = (start - end + 1440) % 1440

    if forward < backward:
        return 'forward'
    if backward < forward:
        return 'backward'
    return 'equal'
```
