---
id: 69c6ff713a52713463aa7927
title: "Sfida 261: parole numeriche"
challengeType: 29
dashedName: challenge-261
---

# --description--

Dato un numero intero da 0 a 99, restituisci la sua rappresentazione in parole inglesi.

- `0` restituisce `"zero"`.
- I numeri da 1 a 19 hanno nomi unici (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- I multipli di 10 da 20 a 90 hanno nomi propri (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- I numeri da 21 a 99 che non sono multipli di 10 si scrivono come due parole unite da un trattino. Per esempio `"forty-two"` e `"fifty-three"`.

# --hints--

`get_number_words(0)` dovrebbe restituire `"zero"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(0), "zero")`)
}})
```

`get_number_words(10)` dovrebbe restituire `"ten"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(10), "ten")`)
}})
```

`get_number_words(19)` dovrebbe restituire `"nineteen"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(19), "nineteen")`)
}})
```

`get_number_words(30)` dovrebbe restituire `"thirty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(30), "thirty")`)
}})
```

`get_number_words(53)` dovrebbe restituire `"fifty-three"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(53), "fifty-three")`)
}})
```

`get_number_words(7)` dovrebbe restituire `"seven"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(7), "seven")`)
}})
```

`get_number_words(12)` dovrebbe restituire `"twelve"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(12), "twelve")`)
}})
```

`get_number_words(60)` dovrebbe restituire `"sixty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(60), "sixty")`)
}})
```

`get_number_words(67)` dovrebbe restituire `"sixty-seven"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(67), "sixty-seven")`)
}})
```

`get_number_words(98)` dovrebbe restituire `"ninety-eight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(98), "ninety-eight")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_words(n):

    return n
```

# --solutions--

```py
def get_number_words(n):
    ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
            "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
            "seventeen","eighteen","nineteen"]
    tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"]

    if n < 20:
        return ones[n]
    if n % 10 == 0:
        return tens[n // 10]
    return tens[n // 10] + "-" + ones[n % 10]
```
