---
id: 69c6ff713a52713463aa7928
title: "Sfida 262: parser di query URL"
challengeType: 29
dashedName: challenge-262
---

# --description--

Dato un URL che contiene una query string, analizza la query string in un oggetto (o dizionario) di coppie chiave-valore.

- La query string inizia dopo `"?"`,
- ogni parametro è separato da `"&"`,
- ogni coppia chiave/valore è separata da `"="`

Per esempio, dato `"https://example.com/search?name=Alice&age=30"`, restituisci:

```json
{
  "name": "Alice",
  "age": "30"
}
```

Tutti i valori devono essere restituiti come stringhe.

# --hints--

`parse_url_query("https://example.com/search?name=Alice&age=30")` dovrebbe restituire `{"name": "Alice", "age": "30"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"})`)
}})
```

`parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python")` dovrebbe restituire `{"skill": "programming", "language": "python"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"})`)
}})
```

`parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2")` dovrebbe restituire `{"category": "books", "sort": "asc", "page": "2"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"})`)
}})
```

`parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now")` dovrebbe restituire `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_url_query(url):

    return url
```

# --solutions--

```py
def parse_url_query(url):
    query_string = url.split('?')[1]
    return dict(param.split('=') for param in query_string.split('&'))
```
