---
id: 69cfca90e8a0a6d4d6871c51
title: "Sfida 267: convertitore di parsec"
challengeType: 29
dashedName: challenge-267
---

# --description--

In una galassia lontana, i parsec vengono usati per misurare sia il tempo sia la distanza. Dato un numero intero di parsec, restituisci il suo equivalente in tempo o distanza.

- Se l'intero dato è dispari, rappresenta il tempo. Se è pari, rappresenta la distanza.

Usa questi tassi di conversione:

| Parsec | Tempo/Distanza |
| - | - |
| 1 | 2 ore |
| 2 | 6 anni luce |

Restituisci il valore convertito come un intero.

# --hints--

`convert_parsecs(1)` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(1), 2)`)
}})
```

`convert_parsecs(2)` dovrebbe restituire `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(2), 6)`)
}})
```

`convert_parsecs(31)` dovrebbe restituire `62`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(31), 62)`)
}})
```

`convert_parsecs(88)` dovrebbe restituire `264`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(88), 264)`)
}})
```

`convert_parsecs(17)` dovrebbe restituire `34`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(17), 34)`)
}})
```

`convert_parsecs(14)` dovrebbe restituire `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(14), 42)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_parsecs(parsecs):

    return parsecs
```

# --solutions--

```py
def convert_parsecs(parsecs):
    if parsecs % 2 != 0:
        return parsecs * 2
    else:
        return parsecs * 3
```
