---
id: 69f35a5bb823ed620fcb7cb8
title: "Sfida 279: la catena di domino più lunga"
challengeType: 29
dashedName: challenge-279
---

# --description--

Dato un array 2D che rappresenta un insieme di domino, restituisci la catena valida più lunga.

- Ogni domino è una coppia di numeri da 0 a 6, ad esempio `[3, 2]`.
- Una catena è valida quando il secondo numero di ogni domino corrisponde al primo numero del successivo.
- Il primo numero del primo domino e il secondo numero dell'ultimo non devono corrispondere a nulla.
- Qualsiasi domino può essere capovolto, quindi `[3, 2]` può essere giocato come `[2, 3]`.
- Esiste sempre esattamente una catena valida più lunga.

Ad esempio, dato `[[1, 2], [4, 5], [2, 3]]`, restituisci `[[1, 2], [2, 3]]`.

# --hints--

`get_longest_chain([[1, 2], [4, 5], [2, 3]])` dovrebbe restituire `[[1, 2], [2, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
result = get_longest_chain([[1, 2], [4, 5], [2, 3]])
TestCase().assertIn(result, [[[1, 2], [2, 3]], [[3, 2], [2, 1]]])`)
}})
```

`get_longest_chain([[2, 1], [4, 3], [5, 3]])` dovrebbe restituire `[[4, 3], [3, 5]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
result = get_longest_chain([[2, 1], [4, 3], [5, 3]])
TestCase().assertIn(result, [[[4, 3], [3, 5]], [[5, 3], [3, 4]]])`)
}})
```

`get_longest_chain([[1, 2], [3, 4], [2, 3], [4, 0]])` dovrebbe restituire `[[1, 2], [2, 3], [3, 4], [4, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
result = get_longest_chain([[1, 2], [3, 4], [2, 3], [4, 0]])
TestCase().assertIn(result, [[[1, 2], [2, 3], [3, 4], [4, 0]], [[0, 4], [4, 3], [3, 2], [2, 1]]])`)
}})
```

`get_longest_chain([[6, 6], [6, 1], [1, 1], [0, 3], [2, 3], [4, 1], [5, 6]])` dovrebbe restituire `[[4, 1], [1, 1], [1, 6], [6, 6], [6, 5]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
result = get_longest_chain([[6, 6], [6, 1], [1, 1], [0, 3], [2, 3], [4, 1], [5, 6]])
TestCase().assertIn(result, [[[4, 1], [1, 1], [1, 6], [6, 6], [6, 5]], [[5, 6], [6, 6], [6, 1], [1, 1], [1, 4]]])`)
}})
```

`get_longest_chain([[0, 4], [3, 3], [0, 3], [5, 6], [4, 5], [4, 2], [5, 5], [1, 2], [4, 4]])` dovrebbe restituire `[[3, 3], [3, 0], [0, 4], [4, 4], [4, 5], [5, 5], [5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
result = get_longest_chain([[0, 4], [3, 3], [0, 3], [5, 6], [4, 5], [4, 2], [5, 5], [1, 2], [4, 4]])
TestCase().assertIn(result, [[[3, 3], [3, 0], [0, 4], [4, 4], [4, 5], [5, 5], [5, 6]], [[6, 5], [5, 5], [5, 4], [4, 4], [4, 0], [0, 3], [3, 3]]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_longest_chain(dominoes):

    return dominoes
```

# --solutions--

```py
def get_longest_chain(dominoes):
    def search(chain, remaining):
        best = chain
        last = chain[-1][1]
        for i, (a, b) in enumerate(remaining):
            rest = remaining[:i] + remaining[i+1:]
            if a == last:
                result = search(chain + [[a, b]], rest)
                if len(result) > len(best):
                    best = result
            if b == last and a != b:
                result = search(chain + [[b, a]], rest)
                if len(result) > len(best):
                    best = result
        return best

    best = []
    for i, (a, b) in enumerate(dominoes):
        rest = dominoes[:i] + dominoes[i+1:]
        r1 = search([[a, b]], rest)
        if len(r1) > len(best):
            best = r1
        if a != b:
            r2 = search([[b, a]], rest)
            if len(r2) > len(best):
                best = r2
    return best
```
