---
id: 69f8c998d78ad3171a0713bc
title: "Sfida 289: somma delle differenze"
challengeType: 29
dashedName: challenge-289
---

# --description--

Dato un array di numeri, restituisci la somma delle differenze tra ogni numero e quello che lo segue.

Per esempio, dato `[1, 3, 4]`, restituisci `3` (2 + 1).

# --hints--

`sum_of_differences([1, 3, 4])` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([1, 3, 4]), 3)`)
}})
```

`sum_of_differences([5, -3, 3, 9, 10])` dovrebbe restituire `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([5, -3, 3, 9, 10]), 5)`)
}})
```

`sum_of_differences([9, 6, 15, -20, 33, 14, 25, 16, -7])` dovrebbe restituire `-16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16)`)
}})
```

`sum_of_differences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])` dovrebbe restituire `25`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_of_differences(arr):

    return arr
```

# --solutions--

```py
def sum_of_differences(arr):
    return sum(arr[i + 1] - arr[i] for i in range(len(arr) - 1))
```
