---
id: 69f8c998d78ad3171a0713bd
title: "Sfida 290: festa della pizza"
challengeType: 29
dashedName: challenge-290
---

# --description--

Dato un array con le ore lavorate oggi da ogni persona, restituisci il numero di pizze da ordinare per una festa della pizza.

- Dividi le ore lavorate di ogni persona per 3 per ottenere il numero di fette.
- Non puoi mangiare una fetta parziale, quindi arrotonda per eccesso il numero di fette di ogni persona al numero intero più vicino.
- Ogni persona riceve almeno due fette.
- Ogni pizza ha 8 fette. Arrotonda per eccesso il numero totale di pizze al numero intero più vicino.

# --hints--

`get_pizzas_to_order([8, 8, 8])` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8]), 2)`)
}})
```

`get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10])` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10]), 3)`)
}})
```

`get_pizzas_to_order([1, 2, 3, 4, 5])` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([1, 2, 3, 4, 5]), 2)`)
}})
```

`get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8])` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8]), 3)`)
}})
```

`get_pizzas_to_order([9, 9, 6])` dovrebbe restituire `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([9, 9, 6]), 1)`)
}})
```

`get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0])` dovrebbe restituire `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_pizzas_to_order(hours_worked):

    return hours_worked
```

# --solutions--

```py
import math

def get_pizzas_to_order(hours_worked):
    total_slices = sum(max(math.ceil(h / 3), 2) for h in hours_worked)
    return math.ceil(total_slices / 8)
```
