---
id: 69f8c998d78ad3171a0713be
title: "Challenge 291: FizzBuzz Count"
challengeType: 29
dashedName: challenge-291
---

# --description--

Dato un numero di inizio e uno di fine, conta il numero di apparizioni di fizz e buzz nell'intervallo (inclusivo).

- I numeri divisibili per 3 contano come fizz.
- I numeri divisibili per 5 contano come buzz.
- I numeri divisibili sia per 3 che per 5 contano sia come fizz che come buzz.

Restituisci un oggetto o un dizionario con i conteggi nel formato: `{ fizz, buzz }`.

# --hints--

`fizz_buzz_count(1, 11)` should return `{"fizz": 3, "buzz": 2}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(1, 11), {"fizz": 3, "buzz": 2})`)
}})
```

`fizz_buzz_count(14, 41)` should return `{"fizz": 9, "buzz": 6}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(14, 41), {"fizz": 9, "buzz": 6})`)
}})
```

`fizz_buzz_count(24, 100)` should return `{"fizz": 26, "buzz": 16}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(24, 100), {"fizz": 26, "buzz": 16})`)
}})
```

`fizz_buzz_count(-635, -14)` should return `{"fizz": 207, "buzz": 125}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-635, -14), {"fizz": 207, "buzz": 125})`)
}})
```

`fizz_buzz_count(-5432, 6789)` should return `{"fizz": 4074, "buzz": 2444}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-5432, 6789), {"fizz": 4074, "buzz": 2444})`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_count(start, end):

    return start
```

# --solutions--

```py
def fizz_buzz_count(start, end):
    fizz = sum(1 for i in range(start, end + 1) if i % 3 == 0)
    buzz = sum(1 for i in range(start, end + 1) if i % 5 == 0)
    return {"fizz": fizz, "buzz": buzz}
```
