---
id: 69f8c998d78ad3171a0713bf
title: "Sfida 292: rapporto d'aspetto più ampio"
challengeType: 29
dashedName: challenge-292
---

# --description--

Date due stringhe che indicano dimensioni diverse di un'immagine, restituisci il rapporto d'aspetto dell'immagine con il rapporto larghezza-altezza maggiore.

- Le stringhe fornite saranno nel formato `"WxH"`, per esempio `"1920x1080"`.
- Il rapporto d'aspetto è il rapporto tra larghezza e altezza, ridotto ai numeri interi più bassi. Per esempio, `"1920x1080"` si riduce a `"16:9"`.
- Restituisci una stringa nel formato `"W:H"`, per esempio `"16:9"`.

# --hints--

`get_wider_aspect_ratio("1920x1080", "800x600")` dovrebbe restituire `"16:9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1920x1080", "800x600"), "16:9")`)
}})
```

`get_wider_aspect_ratio("1080x1350", "2048x1536")` dovrebbe restituire `"4:3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1080x1350", "2048x1536"), "4:3")`)
}})
```

`get_wider_aspect_ratio("640x480", "2440x1220")` dovrebbe restituire `"2:1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("640x480", "2440x1220"), "2:1")`)
}})
```

`get_wider_aspect_ratio("360x640", "1080x1920")` dovrebbe restituire `"9:16"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("360x640", "1080x1920"), "9:16")`)
}})
```

`get_wider_aspect_ratio("3440x1440", "2048x858")` dovrebbe restituire `"43:18"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("3440x1440", "2048x858"), "43:18")`)
}})
```

`get_wider_aspect_ratio("12345x61234", "12534x51234")` dovrebbe restituire `"2089:8539"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("12345x61234", "12534x51234"), "2089:8539")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_wider_aspect_ratio(a, b):

    return a
```

# --solutions--

```py
from math import gcd

def get_wider_aspect_ratio(a, b):
    def parse(s):
        w, h = map(int, s.split("x"))
        g = gcd(w, h)
        return w / h, f"{w // g}:{h // g}"

    ratio_a, simplified_a = parse(a)
    ratio_b, simplified_b = parse(b)

    return simplified_a if ratio_a >= ratio_b else simplified_b
```
