---
id: 69f8c998d78ad3171a0713b9
title: "Sfida 286: problemi aperti"
challengeType: 29
dashedName: challenge-286
---

# --description--

Dato un array di numeri di problemi e un altro array di numeri di pull request (PR), restituisci un array dei problemi che rimangono aperti dopo che tutte le PR sono state unite.

- Una PR chiude un problema se le loro cifre sono una rotazione l’una dell’altra. Per esempio, il problema 123 sarebbe chiuso dalla PR 231 o 312.
- Una PR non chiude un problema con lo stesso numero esatto. Per esempio, la PR 123 non chiude il problema 123. Quindi un problema con tutte le cifre uguali non può essere chiuso.
- Entrambi i numeri possono avere gli zeri iniziali rimossi. Per esempio, la PR 201 chiuderebbe il problema 12 (012, una rotazione di 201). Allo stesso modo, il problema 201 sarebbe chiuso dalla PR 12.

Restituisci i problemi aperti rimanenti nell’ordine in cui sono stati dati.

# --hints--

`get_open_issues([123, 234], [231])` dovrebbe restituire `[234]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 234], [231]), [234])`)
}})
```

`get_open_issues([123, 345, 16], [345, 231])` dovrebbe restituire `[345, 16]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 345, 16], [345, 231]), [345, 16])`)
}})
```

`get_open_issues([456, 332, 12, 15], [201, 945, 180])` dovrebbe restituire `[456, 332, 15]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15])`)
}})
```

`get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` dovrebbe restituire `[115, 296, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24])`)
}})
```

`get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` dovrebbe restituire `[95, 395, 754, 296, 709, 237, 1802]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_open_issues(issues, prs):

    return issues
```

# --solutions--

```py
def get_open_issues(issues, prs):
    def get_rotations(n):
        s = str(n)
        rotations = set()
        for i in range(len(s)):
            rotated = s[i:] + s[:i]
            rotations.add(int(rotated))
        return rotations

    def are_rotations(a, b):
        if a == b:
            return False
        return b in get_rotations(a)

    return [issue for issue in issues if not any(are_rotations(pr, issue) or are_rotations(issue, pr) for pr in prs)]
```
