---
id: 69f90c1329a94b37e2a2086c
title: "Sfida 293: la mano migliore"
challengeType: 29
dashedName: challenge-293
---

# --description--

Dato un array di cinque stringhe che rappresentano carte da gioco, restituisci il nome della mano migliore.

- Ogni carta è rappresentata come una stringa di due caratteri: il valore seguito dal seme, `"2h"` per esempio.
  - I valori, dal più basso al più alto, sono: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"` e `"A"`.
  - I semi sono: `"h"`, `"d"`, `"c"` e `"s"`.
- Gli assi (`"A"`) possono essere usati come alti o bassi in una scala.

Le mani, in ordine dal peggiore al migliore, sono:

| Nome | Descrizione |
| - | - |
| `"High Card"` | Nessuna coppia o meglio |
| `"Pair"` | Due carte dello stesso valore |
| `"Two Pair"` | Due carte di un valore e due di un altro |
| `"Three of a Kind"` | Tre carte dello stesso valore |
| `"Straight"` | Cinque valori consecutivi |
| `"Flush"` | Cinque carte dello stesso seme |
| `"Full House"` | Tre carte di un valore e due di un altro |
| `"Four of a Kind"` | Quattro carte dello stesso valore |
| `"Straight Flush"` | Cinque valori consecutivi dello stesso seme |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` dello stesso seme |

Restituisci il nome della mano migliore.

# --hints--

`get_best_hand(["7s", "7h", "7d", "2c", "5h"])` dovrebbe restituire `"Three of a Kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind")`)
}})
```

`get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"])` dovrebbe restituire `"Full House"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House")`)
}})
```

`get_best_hand(["2h", "5h", "7h", "9h", "Jh"])` dovrebbe restituire `"Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["2h", "5h", "7h", "9h", "Jh"]), "Flush")`)
}})
```

`get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"])` dovrebbe restituire `"Four of a Kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind")`)
}})
```

`get_best_hand(["Ts", "Th", "9d", "9c", "8h"])` dovrebbe restituire `"Two Pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair")`)
}})
```

`get_best_hand(["9c", "8c", "7c", "6c", "5c"])` dovrebbe restituire `"Straight Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush")`)
}})
```

`get_best_hand(["As", "Kh", "Jd", "8c", "5h"])` dovrebbe restituire `"High Card"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Kh", "Jd", "8c", "5h"]), "High Card")`)
}})
```

`get_best_hand(["As", "2h", "3d", "4c", "5h"])` dovrebbe restituire `"Straight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "2h", "3d", "4c", "5h"]), "Straight")`)
}})
```

`get_best_hand(["Ts", "Th", "7c", "6d", "5h"])` dovrebbe restituire `"Pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "7c", "6d", "5h"]), "Pair")`)
}})
```

`get_best_hand(["As", "Ks", "Qs", "Js", "Ts"])` dovrebbe restituire `"Royal Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_best_hand(cards):

    return cards
```

# --solutions--

```py
def get_best_hand(cards):
    rank_order = "23456789TJQKA"

    ranks = [c[0] for c in cards]
    suits = [c[1] for c in cards]

    is_flush = len(set(suits)) == 1
    rank_indices = sorted(rank_order.index(r) for r in ranks)

    def is_straight(indices):
        is_regular = all(indices[i] == indices[i - 1] + 1 for i in range(1, len(indices)))
        is_low_ace = indices == [0, 1, 2, 3, 12]
        return is_regular or is_low_ace

    straight = is_straight(rank_indices)

    counts = {}
    for r in ranks:
        counts[r] = counts.get(r, 0) + 1
    freq = sorted(counts.values(), reverse=True)

    if is_flush and rank_indices == [8, 9, 10, 11, 12]:
        return "Royal Flush"
    if is_flush and straight:
        return "Straight Flush"
    if freq[0] == 4:
        return "Four of a Kind"
    if freq[0] == 3 and freq[1] == 2:
        return "Full House"
    if is_flush:
        return "Flush"
    if straight:
        return "Straight"
    if freq[0] == 3:
        return "Three of a Kind"
    if freq[0] == 2 and freq[1] == 2:
        return "Two Pair"
    if freq[0] == 2:
        return "Pair"
    return "High Card"
```
