---
id: 6a0dcd03ee4e68698080ef67
title: "Sfida 303: coinquilini"
challengeType: 29
dashedName: challenge-303
---

# --description--

Dato un array di persone e del loro gruppo di coinquilini, restituisci le assegnazioni delle stanze per un soggiorno in hotel seguendo queste regole:

- Ogni persona ha una proprietà `name` e una `group`:

```json
[
  { "name": "Alice", "group": "A" },
  { "name": "Bob", "group": "B" },
  { "name": "Carol", "group": "A" }
]
```

- Le persone possono condividere la stanza solo con qualcuno dello stesso gruppo e sono abbinate nell'ordine in cui sono fornite.
- Restituisci un array di stringhe con i nomi separati da `" and "` per una stanza condivisa, e solo il nome per una stanza singola. I nomi devono apparire nell'ordine in cui sono stati abbinati. Per l'esempio sopra, restituisci `["Alice and Carol", "Bob"]`.

# --hints--

`get_roommates([{ "name": "Alice", "group": "A" }, { "name": "Bob", "group": "B" }, { "name": "Carol", "group": "A" }])` dovrebbe restituire `["Alice and Carol", "Bob"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_roommates([{"name": "Alice", "group": "A"}, {"name": "Bob", "group": "B"}, {"name": "Carol", "group": "A"}])), sorted(["Alice and Carol", "Bob"]))`)
}})
```

`get_roommates([{ "name": "John", "group": "C" }, { "name": "Julia", "group": "C" }, { "name": "Jim", "group": "C" }])` dovrebbe restituire `["John and Julia", "Jim"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_roommates([{"name": "John", "group": "C"}, {"name": "Julia", "group": "C"}, {"name": "Jim", "group": "C"}])), sorted(["John and Julia", "Jim"]))`)
}})
```

`get_roommates([{ "name": "Adam", "group": "D" }, { "name": "Abraham", "group": "E" }, { "name": "Austin", "group": "E" }, { "name": "Augustus", "group": "D" }, { "name": "Angelica", "group": "D" }, { "name": "Aaron", "group": "E" }])` dovrebbe restituire `["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_roommates([{"name": "Adam", "group": "D"}, {"name": "Abraham", "group": "E"}, {"name": "Austin", "group": "E"}, {"name": "Augustus", "group": "D"}, {"name": "Angelica", "group": "D"}, {"name": "Aaron", "group": "E"}])), sorted(["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"]))`)
}})
```

`get_roommates([{ "name": "Frank", "group": "A" }, { "name": "Emitt", "group": "B" }, { "name": "Daria", "group": "F" }, { "name": "Charles", "group": "D" }, { "name": "Bailey", "group": "A" }, { "name": "Albert", "group": "F" }])` dovrebbe restituire `["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_roommates([{"name": "Frank", "group": "A"}, {"name": "Emitt", "group": "B"}, {"name": "Daria", "group": "F"}, {"name": "Charles", "group": "D"}, {"name": "Bailey", "group": "A"}, {"name": "Albert", "group": "F"}])), sorted(["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"]))`)
}})
```

`get_roommates([{ "name": "Kevin", "group": "A" }, { "name": "Yuri", "group": "A" }, { "name": "Hugo", "group": "B" }, { "name": "Violet", "group": "A" }, { "name": "Brett", "group": "A" }, { "name": "Wayne", "group": "B" }])` dovrebbe restituire `["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_roommates([{"name": "Kevin", "group": "A"}, {"name": "Yuri", "group": "A"}, {"name": "Hugo", "group": "B"}, {"name": "Violet", "group": "A"}, {"name": "Brett", "group": "A"}, {"name": "Wayne", "group": "B"}])), sorted(["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def get_roommates(people):

    return people
```

# --solutions--

```py
def get_roommates(people):
    groups = {}
    for person in people:
        groups.setdefault(person["group"], []).append(person["name"])
    result = []
    for members in groups.values():
        for i in range(0, len(members), 2):
            result.append(" and ".join(members[i:i + 2]))
    return result
```
