---
id: 6a0dcd03ee4e68698080ef66
title: "Sfida 302: Jet lag"
challengeType: 29
dashedName: challenge-302
---

# --description--

Dato un città di partenza, una città di arrivo, una durata del volo in ore e una direzione di viaggio, restituisci il numero di ore di jet lag che il viaggiatore sta sperimentando.

Le città fornite saranno prese dalla seguente lista che include il loro offset UTC:

| Città | Offset |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Per calcolare le ore di jet lag:

1. Trova la differenza di fuso orario in ore tra le due città.
2. Determina il moltiplicatore di direzione. Se si viaggia `"east"`, è 1,5, altrimenti è 1,0.
3. Ottieni le ore di jet lag con la formula: differenza di fuso orario + (durata del volo * 0,1) * moltiplicatore di direzione

Restituisci le ore di jet lag arrotondate a una cifra decimale.

# --hints--

`get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east")` dovrebbe restituire `6.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east"), 6.5)`)
}})
```

`get_jet_lag_hours("London", "New York", 8, "west")` dovrebbe restituire `5.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("London", "New York", 8, "west"), 5.8)`)
}})
```

`get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east")` dovrebbe restituire `1.6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east"), 1.6)`)
}})
```

`get_jet_lag_hours("Dubai", "London", 7, "west")` dovrebbe restituire `4.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Dubai", "London", 7, "west"), 4.7)`)
}})
```

`get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west")` dovrebbe restituire `17.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west"), 17.5)`)
}})
```

`get_jet_lag_hours("Tokyo", "Dubai", 9, "west")` dovrebbe restituire `5.9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Tokyo", "Dubai", 9, "west"), 5.9)`)
}})
```

`get_jet_lag_hours("New York", "Istanbul", 10, "east")` dovrebbe restituire `9.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("New York", "Istanbul", 10, "east"), 9.5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):

    return departure_city
```

# --solutions--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):
    offsets = {
        "Los Angeles": -8,
        "New York": -5,
        "London": 0,
        "Istanbul": 3,
        "Dubai": 4,
        "Hong Kong": 8,
        "Tokyo": 9
    }
    timezone_diff = abs(offsets[arrival_city] - offsets[departure_city])
    multiplier = 1.5 if direction == "east" else 1.0
    score = timezone_diff + (flight_duration * 0.1) * multiplier
    return round(score, 1)
```
