---
id: 6a0dcd03ee4e68698080ef6c
title: "Sfida 308: validatore di carte di credito"
challengeType: 29
dashedName: challenge-308
---

# --description--

Data una stringa di cifre che rappresenta un numero di carta di credito, determina se è un numero valido usando il seguente metodo:

- Partendo dalla penultima cifra, raddoppia ogni altra cifra procedendo verso sinistra.
- Se raddoppiare una cifra produce un numero maggiore di 9, sottrai 9.
- Somma tutte le cifre (raddoppiate e non).
- Se il totale è divisibile per 10, il numero è valido.

# --hints--

`is_valid_card("4532015112830366")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830366"), True)`)
}})
```

`is_valid_card("5425233430109903")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("5425233430109903"), True)`)
}})
```

`is_valid_card("371449635398431")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("371449635398431"), True)`)
}})
```

`is_valid_card("6011111111111117")` dovrebbe restituire `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("6011111111111117"), True)`)
}})
```

`is_valid_card("4532015112830367")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830367"), False)`)
}})
```

`is_valid_card("1234567890123456")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("1234567890123456"), False)`)
}})
```

`is_valid_card("4532015112830368")` dovrebbe restituire `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830368"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_card(number):

    return number
```

# --solutions--

```py
def is_valid_card(number):
    total = 0
    for i, digit in enumerate(reversed(number)):
        d = int(digit)
        if i % 2 == 1:
            d *= 2
            if d > 9:
                d -= 9
        total += d
    return total % 10 == 0
```
