---
id: 6a0dcd03ee4e68698080ef6b
title: "Sfida 307: regolamenti di zonizzazione"
challengeType: 29
dashedName: challenge-307
---

# --description--

Dato un array bidimensionale (array di array) che rappresenta la disposizione degli edifici di una città, restituisci le coordinate di tutti gli edifici che violano le regole di zonizzazione.

Ogni cella nella griglia contiene una delle etichette della tabella sottostante. Un edificio viola le regole se uno qualsiasi dei suoi (fino a) 4 vicini, orizzontali o verticali, è di un tipo a cui non può essere adiacente.

| Etichetta | Tipo | Non può essere adiacente a |
| - | - | - |
| `"i"` | industriale | `"R"`, `"I"` |
| `"A"` | agricolo | `"C"` |
| `"R"` | residenziale | `"i"`, `"C"` |
| `"I"` | istituzionale | `"i"` |
| `"C"` | commerciale | `"R"`, `"A"` |
| `""` (stringa vuota) | non sviluppato | nessuna restrizione |

Restituisci le coordinate di tutte le celle in violazione come un array di coppie `[row, col]`, in qualsiasi ordine. Se non ci sono violazioni, restituisci un array vuoto.

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` dovrebbe restituire `[[0, 0], [0, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` dovrebbe restituire `[[0, 1], [1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` dovrebbe restituire `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` dovrebbe restituire `[[0, 1], [0, 2], [0, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` dovrebbe restituire `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
