---
id: 6a0dcd03ee4e68698080ef68
title: "Sfida 304: organizzazione dell'itinerario"
challengeType: 29
dashedName: challenge-304
---

# --description--

Dato un array di almeno due fermate opzionali per una gita di un giorno, restituisci il numero di disposizioni valide dell'itinerario.

L'itinerario include sempre `"breakfast"`, `"lunch"` e `"dinner"`, che non saranno passati come argomenti. Le fermate opzionali possono essere collocate ovunque nell'itinerario, rispettando le seguenti regole:

- `"breakfast"` è sempre la prima, con almeno una fermata prima di `"lunch"`.
- `"lunch"` deve comparire prima di `"dinner"`, con almeno una fermata in mezzo.
- Al massimo, una fermata opzionale può comparire dopo `"dinner"`.

Restituisci il numero di disposizioni valide.

# --hints--

`get_itinerary_count(["library", "park"])` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` dovrebbe restituire `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` dovrebbe restituire `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` dovrebbe restituire `840`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` dovrebbe restituire `55440`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
