---
id: 6a15cadf5f240d05a2649556
title: "Sfida 312: costo dello streaming"
challengeType: 29
dashedName: challenge-312
---

# --description--

Dato un array che rappresenta i film nel carrello del tuo servizio di streaming e una stringa per il tuo livello di abbonamento, restituisci il costo totale dei film.

Ogni elemento nel carrello è un oggetto con un `"format"` (`"HD"` o `"4K"`) e un `"type"` (`"rent"` o `"buy"`). I loro costi sono:

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

Applica i seguenti sconti in base al livello di abbonamento:

- `"none"`: prezzo pieno
- `"basic"`: 10% di sconto
- `"premium"`: 25% di sconto

Restituisci il costo totale arrotondato a due decimali nel formato `"$D.CC"`.

# --hints--

`get_streaming_bill([{ "format": "HD", "type": "rent" }], "none")` dovrebbe restituire `"$3.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}], "none"), "$3.99")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "buy" }], "premium")` dovrebbe restituire `"$9.74"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "buy"}], "premium"), "$9.74")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "rent" }, { "format": "HD", "type": "rent" }, { "format": "HD", "type": "buy" }], "basic")` dovrebbe restituire `"$18.87"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "HD", "type": "rent"}, {"format": "HD", "type": "buy"}], "basic"), "$18.87")`)
}})
```

`get_streaming_bill([{ "format": "4K", "type": "buy" }, { "format": "4K", "type": "buy" }, { "format": "4K", "type": "buy" }], "premium")` dovrebbe restituire `"$44.98"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "4K", "type": "buy"}, {"format": "4K", "type": "buy"}, {"format": "4K", "type": "buy"}], "premium"), "$44.98")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "rent" }, { "format": "4K", "type": "rent" }, { "format": "HD", "type": "buy" }, { "format": "4K", "type": "buy" }], "none")` dovrebbe restituire `"$42.96"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "4K", "type": "rent"}, {"format": "HD", "type": "buy"}, {"format": "4K", "type": "buy"}], "none"), "$42.96")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "rent" }, { "format": "4K", "type": "rent" }, { "format": "HD", "type": "buy" }, { "format": "4K", "type": "buy" }], "basic")` dovrebbe restituire `"$38.66"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "4K", "type": "rent"}, {"format": "HD", "type": "buy"}, {"format": "4K", "type": "buy"}], "basic"), "$38.66")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_streaming_bill(cart, subscription):

    return cart
```

# --solutions--

```py
def get_streaming_bill(cart, subscription):
    pricing = {
        "HD": {"rent": 3.99, "buy": 12.99},
        "4K": {"rent": 5.99, "buy": 19.99},
    }
    discounts = {"none": 1, "basic": 0.9, "premium": 0.75}
    multiplier = discounts[subscription]
    total = sum(pricing[item["format"]][item["type"]] * multiplier for item in cart)
    return f"${total:.2f}"
```
