---
id: 6a151d32d772271dd2448c2d
title: "Sfida 311: incantatore"
challengeType: 29
dashedName: challenge-311
---

# --description--

Data una stringa di codici di incantesimi che stai lanciando, calcola il punteggio totale.

Ogni carattere nella stringa rappresenta un incantesimo:

| Codice | Incantesimo | Categoria | Punteggio base |
| - | - | - | - |
| `"f"` | Fuoco | Distruzione | 3 |
| `"l"` | Fulmine | Distruzione | 3 |
| `"i"` | Ghiaccio | Controllo | 2 |
| `"w"` | Vento | Controllo | 2 |
| `"h"` | Cura | Restauro | 1 |
| `"s"` | Scudo | Restauro | 1 |

Un moltiplicatore combo viene applicato in base a quanti incantesimi consecutivi sono stati lanciati da categorie diverse:

- Il primo incantesimo vale sempre il punteggio base.
- Ogni incantesimo consecutivo di categoria diversa dalla precedente aumenta il moltiplicatore di 1.
- Lanciare un incantesimo della stessa categoria del precedente azzera il moltiplicatore a 1.
- Il punteggio di ogni incantesimo è il suo punteggio base moltiplicato per il moltiplicatore corrente.

Restituisci il punteggio totale dalla sequenza di incantesimi.

# --hints--

`cast("fihwl")` dovrebbe restituire `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` dovrebbe restituire `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` dovrebbe restituire `37`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` dovrebbe restituire `50`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` dovrebbe restituire `101`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
