---
id: 6a15cadf5f240d05a264955b
title: "Sfida 317: calcolatore BMI"
challengeType: 29
dashedName: challenge-317
---

# --description--

Dato un peso in libbre e un'altezza in pollici, restituisci il BMI (indice di massa corporea) arrotondato a una cifra decimale.

Per ottenere il BMI: dividi il peso per l'altezza al quadrato, poi moltiplica il risultato per 703.

# --hints--

`calculate_bmi(180, 70)` dovrebbe restituire `25.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(180, 70), 25.8)`)
}})
```

`calculate_bmi(140, 64)` dovrebbe restituire `24.0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(140, 64), 24.0)`)
}})
```

`calculate_bmi(160, 76)` dovrebbe restituire `19.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(160, 76), 19.5)`)
}})
```

`calculate_bmi(200, 60)` dovrebbe restituire `39.1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(200, 60), 39.1)`)
}})
```

`calculate_bmi(150, 68)` dovrebbe restituire `22.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(150, 68), 22.8)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_bmi(weight, height):

    return weight
```

# --solutions--

```py
def calculate_bmi(weight, height):
    return round((weight / (height * height)) * 703, 1)
```
