---
id: 6a15cadf5f240d05a2649558
title: "Sfida 314: fattorizzazione in fattori primi"
challengeType: 29
dashedName: challenge-314
---

# --description--

Dato un intero maggiore di 1, restituisci la sua fattorizzazione in fattori primi come un array di numeri in ordine crescente.

La fattorizzazione in fattori primi è l'insieme di numeri primi che moltiplicati insieme producono l'intero dato. Ogni numero ha esattamente un insieme. Per esempio, la fattorizzazione in fattori primi di 20 è `[2, 2, 5]` perché 2 * 2 * 5 = 20.

Se l'intero dato è esso stesso primo, restituiscilo in un array con un solo elemento.

# --hints--

`prime_factorization(20)` dovrebbe restituire `[2, 2, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(20), [2, 2, 5])`)
}})
```

`prime_factorization(17)` dovrebbe restituire `[17]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(17), [17])`)
}})
```

`prime_factorization(15)` dovrebbe restituire `[3, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(15), [3, 5])`)
}})
```

`prime_factorization(35)` dovrebbe restituire `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(35), [5, 7])`)
}})
```

`prime_factorization(999)` dovrebbe restituire `[3, 3, 3, 37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(999), [3, 3, 3, 37])`)
}})
```

`prime_factorization(360)` dovrebbe restituire `[2, 2, 2, 3, 3, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(360), [2, 2, 2, 3, 3, 5])`)
}})
```

`prime_factorization(510510)` dovrebbe restituire `[2, 3, 5, 7, 11, 13, 17]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(510510), [2, 3, 5, 7, 11, 13, 17])`)
}})
```

# --seed--

## --seed-contents--

```py
def prime_factorization(n):

    return n
```

# --solutions--

```py
def prime_factorization(n):
    factors = []
    divisor = 2
    while divisor * divisor <= n:
        while n % divisor == 0:
            factors.append(divisor)
            n //= divisor
        divisor += 1
    if n > 1:
        factors.append(n)
    return factors
```
