---
id: 6a15cadf5f240d05a264955e
title: "Sfida 320: banca del sangue"
challengeType: 29
dashedName: challenge-320
---

# --description--

Dato un array dell'inventario di una banca del sangue e un array di richieste di tipo di sangue dei pazienti, restituisci una stringa nel formato `"X of Y patients served"`. Dove `X` è il numero massimo di pazienti che possono ricevere sangue dall'inventario della banca e `Y` è il numero totale di pazienti.

Ogni elemento in entrambi gli array è uno dei seguenti tipi di sangue: `"AB"`, `"A"`, `"B"` o `"O"`.

Regole di compatibilità:

- `"AB"` può ricevere da qualsiasi tipo di sangue.
- `"A"` può ricevere da `"A"` e `"O"`.
- `"B"` può ricevere da `"B"` e `"O"`.
- `"O"` può ricevere solo da `"O"`.

Le voci duplicate negli array dati rappresentano la quantità.

# --hints--

`triage_blood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` dovrebbe restituire `"4 of 4 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served")`)
}})
```

`triage_blood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` dovrebbe restituire `"3 of 5 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served")`)
}})
```

`triage_blood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` dovrebbe restituire `"4 of 5 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served")`)
}})
```

`triage_blood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` dovrebbe restituire `"4 of 4 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served")`)
}})
```

`triage_blood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` dovrebbe restituire `"8 of 11 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served")`)
}})
```

`triage_blood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` dovrebbe restituire `"13 of 14 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_blood(bank, patients):

    return bank
```

# --solutions--

```py
def triage_blood(bank, patients):
    b = {}
    p = {}
    for t in bank:
        b[t] = b.get(t, 0) + 1
    for t in patients:
        p[t] = p.get(t, 0) + 1

    saved = 0

    def serve(patient, donors):
        nonlocal saved
        for donor in donors:
            n = min(p.get(patient, 0), b.get(donor, 0))
            p[patient] = p.get(patient, 0) - n
            b[donor] = b.get(donor, 0) - n
            saved += n

    serve("O", ["O"])
    serve("A", ["A", "O"])
    serve("B", ["B", "O"])
    serve("AB", ["AB", "A", "B", "O"])
    return f"{saved} of {len(patients)} patients served"
```
