---
id: 6a1d9f98e819ed70a0e994de
title: "Sfida 332: smistamento dei problemi"
challengeType: 29
dashedName: challenge-332
---

# --description--

Dato un numero di millisecondi dall'ultimo post su un problema e l'ultimo messaggio pubblicato sul problema, determina cosa dovresti fare con il problema secondo queste regole:

- Se l'ultimo messaggio risale a meno di 7 giorni fa, restituisci `"leave it"`
- Se l'ultimo messaggio risale a 7 o più giorni fa e il suo contenuto contiene `"bump"` (non sensibile alle maiuscole), restituisci `"close it"`
- Altrimenti, restituisci `"bump it"`

# --hints--

`triage_issue(86400000, "Lets fix it")` dovrebbe restituire `"leave it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(86400000, "Lets fix it"), "leave it")`)
}})
```

`triage_issue(1209600000, "still waiting")` dovrebbe restituire `"bump it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(1209600000, "still waiting"), "bump it")`)
}})
```

`triage_issue(864000000, "bump")` dovrebbe restituire `"close it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(864000000, "bump"), "close it")`)
}})
```

`triage_issue(604800000, "Do we still want this?")` dovrebbe restituire `"bump it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Do we still want this?"), "bump it")`)
}})
```

`triage_issue(604800000, "Bumping this")` dovrebbe restituire `"close it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Bumping this"), "close it")`)
}})
```

`triage_issue(345600000, "I'll make a PR")` dovrebbe restituire `"leave it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(345600000, "I'll make a PR"), "leave it")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(ms, message):

    return ms
```

# --solutions--

```py
def triage_issue(ms, message):
    seven_days = 7 * 24 * 60 * 60 * 1000
    if ms < seven_days:
        return "leave it"
    if "bump" in message.lower():
        return "close it"
    return "bump it"
```
