---
id: 6a1d9f98e819ed70a0e994df
title: "Sfida 333: Smistamento problemi 2"
challengeType: 29
dashedName: challenge-333
---

# --description--

Dato un titolo di un problema e un array di etichette attuali, restituisci un array aggiornato di etichette basato sulle seguenti regole:

Se il problema non ha etichette, aggiungi:

- `"bug"` e `"needs triage"` se il titolo contiene `"error"` o `"bug"`
- `"enhancement"` e `"discussing"` se il titolo contiene `"feature"` o `"add"`

Altrimenti, se le etichette date contengono:

- `"needs triage"` e il titolo contiene `"simple"` o `"easy"`, rimuovi `"needs triage"` e aggiungi `"good first issue"`
- `"discussing"` e il titolo contiene `"planned"` o `"next"`, rimuovi `"discussing"` e aggiungi `"on the roadmap"`
- Altrimenti, se è presente `"needs triage"` o `"discussing"`, rimuovila e aggiungi `"help wanted"`

Se il titolo contiene:

- `"security"`, aggiungi un'etichetta `"critical"`

# --hints--

`triage_issue("app crashes with error", [])` dovrebbe restituire `["bug", "needs triage"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", [])), sorted(["bug", "needs triage"]))`)
}})
```

`triage_issue("app crashes with error", ["bug", "needs triage"])` dovrebbe restituire `["bug", "help wanted"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", ["bug", "needs triage"])), sorted(["bug", "help wanted"]))`)
}})
```

`triage_issue("add dark mode", [])` dovrebbe restituire `["enhancement", "discussing"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", [])), sorted(["enhancement", "discussing"]))`)
}})
```

`triage_issue("add dark mode", ["enhancement", "discussing"])` dovrebbe restituire `["enhancement", "help wanted"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted"]))`)
}})
```

`triage_issue("xss security bug", [])` dovrebbe restituire `["bug", "needs triage", "critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("xss security bug", [])), sorted(["bug", "needs triage", "critical"]))`)
}})
```

`triage_issue("security vulnerability in auth", [])` dovrebbe restituire `["critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("security vulnerability in auth", [])), sorted(["critical"]))`)
}})
```

`triage_issue("easy a11y fix", ["bug", "needs triage"])` dovrebbe restituire `["bug", "good first issue"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("easy a11y fix", ["bug", "needs triage"])), sorted(["bug", "good first issue"]))`)
}})
```

`triage_issue("planned api migration", ["enhancement", "discussing"])` dovrebbe restituire `["enhancement", "on the roadmap"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("planned api migration", ["enhancement", "discussing"])), sorted(["enhancement", "on the roadmap"]))`)
}})
```

`triage_issue("improve security", ["enhancement", "discussing"])` dovrebbe restituire `["enhancement", "help wanted", "critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("improve security", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted", "critical"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(title, labels):

    return title
```

# --solutions--

```py
def triage_issue(title, labels):
    result = list(labels)
    t = title.lower()

    def has(label):
        return label in result

    def add(label):
        if not has(label):
            result.append(label)

    def remove(label):
        if label in result:
            result.remove(label)

    if len(result) == 0:
        if "error" in t or "bug" in t:
            add("bug")
            add("needs triage")
        elif "feature" in t or "add" in t:
            add("enhancement")
            add("discussing")
    else:
        if has("needs triage") and ("simple" in t or "easy" in t):
            remove("needs triage")
            add("good first issue")
        elif has("discussing") and ("planned" in t or "next" in t):
            remove("discussing")
            add("on the roadmap")
        elif has("needs triage") or has("discussing"):
            remove("needs triage")
            remove("discussing")
            add("help wanted")

    if "security" in t:
        add("critical")

    return result
```
