---
id: 6a1d9f98e819ed70a0e994da
title: "Sfida 328: la routine di Kaprekar"
challengeType: 29
dashedName: challenge-328
---

# --description--

Dato un numero di 4 cifre, restituisci il numero di volte che devi applicare la routine di Kaprekar fino a raggiungere 6174.

La routine di Kaprekar funziona così:

- Disponi le cifre in ordine decrescente per formare il numero più grande
- Disponi le cifre in ordine crescente per formare il numero più piccolo (aggiungi zeri iniziali se necessario)
- Sottrai il numero più piccolo da quello più grande
- Ripeti con il nuovo numero

# --hints--

`kaprekar(1234)` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(1234), 3)`)
}})
```

`kaprekar(2025)` dovrebbe restituire `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(2025), 6)`)
}})
```

`kaprekar(7173)` dovrebbe restituire `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(7173), 4)`)
}})
```

`kaprekar(3164)` dovrebbe restituire `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(3164), 7)`)
}})
```

`kaprekar(8082)` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(8082), 2)`)
}})
```

# --seed--

## --seed-contents--

```py
def kaprekar(n):

    return n
```

# --solutions--

```py
def kaprekar(n):
    steps = 0
    current = n
    while current != 6174:
        digits = list(str(current).zfill(4))
        desc = int("".join(sorted(digits, reverse=True)))
        asc = int("".join(sorted(digits)))
        current = desc - asc
        steps += 1
    return steps
```
