---
id: 6a1d9f98e819ed70a0e994e2
title: "Sfida 336: compatibilità dell'oroscopo"
challengeType: 29
dashedName: challenge-336
---

# --description--

Date due stringhe di segni zodiacali, restituisci la loro percentuale di compatibilità.

I segni sono disposti in una ruota di 12 posizioni in questo ordine: `"Aries"`, `"Taurus"`, `"Gemini"`, `"Cancer"`, `"Leo"`, `"Virgo"`, `"Libra"`, `"Scorpio"`, `"Sagittarius"`, `"Capricorn"`, `"Aquarius"`, `"Pisces"`, tornando a `"Aries"` dopo `"Pisces"`. Trova la distanza più breve tra i due segni e restituisci la compatibilità:

| Distanza | Compatibilità |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscope_match("Libra", "Sagittarius")` dovrebbe restituire `"80%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Libra", "Sagittarius"), "80%")`)
}})
```

`horoscope_match("Gemini", "Scorpio")` dovrebbe restituire `"20%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Gemini", "Scorpio"), "20%")`)
}})
```

`horoscope_match("Pisces", "Aries")` dovrebbe restituire `"40%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Pisces", "Aries"), "40%")`)
}})
```

`horoscope_match("Capricorn", "Cancer")` dovrebbe restituire `"50%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Capricorn", "Cancer"), "50%")`)
}})
```

`horoscope_match("Aquarius", "Aquarius")` dovrebbe restituire `"100%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Aquarius", "Aquarius"), "100%")`)
}})
```

`horoscope_match("Virgo", "Taurus")` dovrebbe restituire `"90%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Virgo", "Taurus"), "90%")`)
}})
```

`horoscope_match("Leo", "Scorpio")` dovrebbe restituire `"30%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Leo", "Scorpio"), "30%")`)
}})
```

# --seed--

## --seed-contents--

```py
def horoscope_match(sign1, sign2):

    return sign1
```

# --solutions--

```py
def horoscope_match(sign1, sign2):
    signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"]
    a = signs.index(sign1)
    b = signs.index(sign2)
    diff = abs(a - b)
    distance = min(diff, 12 - diff)
    compatibility = {0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%"}
    return compatibility[distance]
```
