---
id: 6a1d9f98e819ed70a0e994e0
title: "Sfida 334: resto esatto"
challengeType: 29
dashedName: challenge-334
---

# --description--

Dato un importo intero in centesimi, restituisci il numero di modi distinti per dare il resto esatto usando pennies (1 cent), nickels (5 cent), dimes (10 cent) e quarters (25 cent).

# --hints--

`exact_change(3)` dovrebbe restituire `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(3), 1)`)
}})
```

`exact_change(9)` dovrebbe restituire `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(9), 2)`)
}})
```

`exact_change(17)` dovrebbe restituire `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(17), 6)`)
}})
```

`exact_change(39)` dovrebbe restituire `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(39), 24)`)
}})
```

`exact_change(61)` dovrebbe restituire `73`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(61), 73)`)
}})
```

`exact_change(99)` dovrebbe restituire `213`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(99), 213)`)
}})
```

# --seed--

## --seed-contents--

```py
def exact_change(amount):

    return amount
```

# --solutions--

```py
def exact_change(amount):
    coins = [1, 5, 10, 25]
    dp = [0] * (amount + 1)
    dp[0] = 1
    for coin in coins:
        for i in range(coin, amount + 1):
            dp[i] += dp[i - coin]
    return dp[amount]
```
