---
id: 6a2037a68a0bc2aef0006001
title: "Sfida 339: sotto-array di un array"
challengeType: 29
dashedName: challenge-339
---

# --description--

Dato un array e una dimensione per i sotto-array, restituisci l'array diviso in sotto-array di quella dimensione.

- L'ultimo sotto-array può essere più piccolo se l'array non si divide in modo uniforme.

# --hints--

`chunk_array([1, 2, 3, 4, 5, 6], 3)` dovrebbe restituire `[[1, 2, 3], [4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]])`)
}})
```

`chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` dovrebbe restituire `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]])`)
}})
```

`chunk_array([1, 2, 3, 4, 5], 3)` dovrebbe restituire `[[1, 2, 3], [4, 5]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]])`)
}})
```

`chunk_array(["a", "b", "c", "d", "e"], 1)` dovrebbe restituire `[["a"], ["b"], ["c"], ["d"], ["e"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]])`)
}})
```

`chunk_array([1, 2, 3], 5)` dovrebbe restituire `[[1, 2, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3], 5), [[1, 2, 3]])`)
}})
```

# --seed--

## --seed-contents--

```py
def chunk_array(arr, size):

    return arr
```

# --solutions--

```py
def chunk_array(arr, size):
    return [arr[i:i + size] for i in range(0, len(arr), size)]
```
