---
id: 6a2037a68a0bc2aef0006002
title: "Sfida 340: convertitore in Pig Latin"
challengeType: 29
dashedName: challenge-340
---

# --description--

Data una stringa, convertila in Pig Latin usando le seguenti regole:

- Se una parola inizia con una vocale (`"a"`, `"e"`, `"i"`, `"o"` o `"u"`), aggiungi `"way"` alla fine. Per esempio, `"universe"` si converte in `"universeway"`.
- Se una parola inizia con una o più consonanti, spostale alla fine e aggiungi `"ay"`. Per esempio, `"hello"` si converte in `"ellohay"`.
- Conserva il caso della prima lettera. Per esempio, `"Hello"` si converte in `"Ellohay"`.

# --hints--

`pig_latin("universe")` dovrebbe restituire `"universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("universe"), "universeway")`)
}})
```

`pig_latin("hello")` dovrebbe restituire `"ellohay"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello"), "ellohay")`)
}})
```

`pig_latin("hello universe")` dovrebbe restituire `"ellohay universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello universe"), "ellohay universeway")`)
}})
```

`pig_latin("Hello universe")` dovrebbe restituire `"Ellohay universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Hello universe"), "Ellohay universeway")`)
}})
```

`pig_latin("Pig Latin is fun")` dovrebbe restituire `"Igpay Atinlay isway unfay"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Pig Latin is fun"), "Igpay Atinlay isway unfay")`)
}})
```

`pig_latin("The quick brown fox jumped over the lazy dog")` dovrebbe restituire `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday")`)
}})
```

# --seed--

## --seed-contents--

```py
def pig_latin(s):

    return s
```

# --solutions--

```py
def pig_latin(s):
    vowels = "aeiou"

    def convert_word(word):
        lower_word = word.lower()
        if lower_word[0] in vowels:
            return word + "way"
        i = 0
        while i < len(word) and lower_word[i] not in vowels:
            i += 1
        consonants = word[:i]
        rest = word[i:]
        result = rest + consonants + "ay"
        if word[0].isupper():
            return result[0].upper() + result[1:].lower()
        return result

    return " ".join(convert_word(word) for word in s.split(" "))
```
