---
id: 6a2037a68a0bc2aef0006007
title: "Sfida 345: fusione di parole"
challengeType: 29
dashedName: challenge-345
---

# --description--

Date due parole, restituisci una nuova parola combinando la prima metà della prima parola con la seconda metà della seconda parola.

- Per le parole di lunghezza dispari, la prima metà è la metà più corta.

# --hints--

`blend_words("turtle", "toucan")` dovrebbe restituire `"turcan"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("turtle", "toucan"), "turcan")`)
}})
```

`blend_words("chipmunk", "flamingo")` dovrebbe restituire `"chipingo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("chipmunk", "flamingo"), "chipingo")`)
}})
```

`blend_words("falcon", "pelican")` dovrebbe restituire `"falican"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("falcon", "pelican"), "falican")`)
}})
```

`blend_words("hyena", "iguana")` dovrebbe restituire `"hyana"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("hyena", "iguana"), "hyana")`)
}})
```

`blend_words("scorpion", "gorilla")` dovrebbe restituire `"scorilla"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("scorpion", "gorilla"), "scorilla")`)
}})
```

`blend_words("platypus", "wolverine")` dovrebbe restituire `"platerine"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("platypus", "wolverine"), "platerine")`)
}})
```

# --seed--

## --seed-contents--

```py
def blend_words(word1, word2):

    return word1
```

# --solutions--

```py
def blend_words(word1, word2):
    first_half = word1[:len(word1) // 2]
    second_half = word2[len(word2) // 2:]
    return first_half + second_half
```
