---
id: 6a2037a68a0bc2aef0006000
title: "Sfida 338: calcolatore dell'età degli animali domestici"
challengeType: 29
dashedName: challenge-338
---

# --description--

Dato un tipo di animale domestico e la sua età in anni umani, restituisci l'età equivalente in anni dell'animale usando la seguente tabella di conversione:

| Animale | Moltiplicatore |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`pet_years("dog", 5)` dovrebbe restituire `35`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("dog", 5), 35)`)
}})
```

`pet_years("cat", 9)` dovrebbe restituire `54`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("cat", 9), 54)`)
}})
```

`pet_years("rabbit", 3)` dovrebbe restituire `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("rabbit", 3), 24)`)
}})
```

`pet_years("hamster", 4)` dovrebbe restituire `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("hamster", 4), 120)`)
}})
```

`pet_years("guinea pig", 5)` dovrebbe restituire `60`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("guinea pig", 5), 60)`)
}})
```

`pet_years("goldfish", 2)` dovrebbe restituire `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("goldfish", 2), 12)`)
}})
```

`pet_years("bird", 1)` dovrebbe restituire `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("bird", 1), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def pet_years(pet, age):

    return pet
```

# --solutions--

```py
def pet_years(pet, age):
    table = {
        "dog": 7,
        "cat": 6,
        "rabbit": 8,
        "hamster": 30,
        "guinea pig": 12,
        "goldfish": 6,
        "bird": 5,
    }
    return age * table[pet]
```
