---
id: 6a2037a68a0bc2aef0005fff
title: "Sfida 337: contatore di segni"
challengeType: 29
dashedName: challenge-337
---

# --description--

Data una stringa di segni di conteggio, restituisci il totale rappresentato.

- Ogni barra verticale `"|"` rappresenta un conteggio.
- Ogni quinto segno è rappresentato da una barra obliqua in avanti `"/"`, completando un gruppo di cinque (`"||||/"`).
- I gruppi sono separati da uno spazio.

# --hints--

`get_tally_count("||||")` dovrebbe restituire `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||"), 4)`)
}})
```

`get_tally_count("||||/")` dovrebbe restituire `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/"), 5)`)
}})
```

`get_tally_count("||||/ |||")` dovrebbe restituire `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ |||"), 8)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||")` dovrebbe restituire `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||"), 17)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` dovrebbe restituire `41`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_tally_count(s):

    return s
```

# --solutions--

```py
def get_tally_count(s):
    return len(s.replace(" ", ""))
```
