---
id: 6a2037a68a0bc2aef0006008
title: "Sfida 346: Salvadanaio"
challengeType: 29
dashedName: challenge-346
---

# --description--

Dato un oggetto che rappresenta un salvadanaio, restituisci il valore totale come una stringa formattata come `"$D.CC"`.

L'oggetto può contenere uno qualsiasi dei seguenti:

| Moneta | Valore |
|------|-------|
| `pennies` | $0.01 |
| `nickels` | $0.05 |
| `dimes` | $0.10 |
| `quarters` | $0.25 |

# --hints--

`piggy_bank({"pennies": 3, "nickels": 5, "dimes": 2, "quarters": 6})` dovrebbe restituire `"$1.98"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"pennies": 3, "nickels": 5, "dimes": 2, "quarters": 6}), "$1.98")`)
}})
```

`piggy_bank({"pennies": 1, "nickels": 1, "dimes": 1, "quarters": 1})` dovrebbe restituire `"$0.41"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"pennies": 1, "nickels": 1, "dimes": 1, "quarters": 1}), "$0.41")`)
}})
```

`piggy_bank({"nickels": 8, "dimes": 6, "quarters": 5})` dovrebbe restituire `"$2.25"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"nickels": 8, "dimes": 6, "quarters": 5}), "$2.25")`)
}})
```

`piggy_bank({})` dovrebbe restituire `"$0.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({}), "$0.00")`)
}})
```

`piggy_bank({"pennies": 146, "nickels": 11, "dimes": 0, "quarters": 19})` dovrebbe restituire `"$6.76"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"pennies": 146, "nickels": 11, "dimes": 0, "quarters": 19}), "$6.76")`)
}})
```

# --seed--

## --seed-contents--

```py
def piggy_bank(coins):

    return coins
```

# --solutions--

```py
def piggy_bank(coins):
    values = {"pennies": 0.01, "nickels": 0.05, "dimes": 0.10, "quarters": 0.25}
    total = sum(values[coin] * count for coin, count in coins.items())
    return f"${total:.2f}"
```
