---
id: 6a22d77ddf034bc4e35b1d5a
title: "Sfida 353: valutazione del rapporto di contrasto 2"
challengeType: 29
dashedName: challenge-353
---

# --description--

Dato due valori di luminanza relativa e un booleano che indica se il testo è grande, restituisci la valutazione del contrasto WCAG usando il seguente metodo:

Calcola il rapporto di contrasto aggiungendo 0.05 a ciascun valore di luminanza, poi dividendo il valore più chiaro per quello più scuro. Il valore più chiaro sarà sempre il primo argomento.

Restituisci la valutazione basata sul rapporto di contrasto usando la seguente tabella:

| Valutazione | Testo normale | Testo grande |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | sotto 4.5 | sotto 3.0 |

# --hints--

`get_contrast_rating(1.0, 0.0, False)` dovrebbe restituire `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(1.0, 0.0, False), "AAA")`)
}})
```

`get_contrast_rating(0.9015, 0.1364, False)` dovrebbe restituire `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.9015, 0.1364, False), "AA")`)
}})
```

`get_contrast_rating(0.8965, 0.1628, False)` dovrebbe restituire `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.8965, 0.1628, False), "Fail")`)
}})
```

`get_contrast_rating(0.7469, 0.0957, True)` dovrebbe restituire `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7469, 0.0957, True), "AAA")`)
}})
```

`get_contrast_rating(0.7489, 0.2018, True)` dovrebbe restituire `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7489, 0.2018, True), "AA")`)
}})
```

`get_contrast_rating(0.6571, 0.1974, True)` dovrebbe restituire `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.6571, 0.1974, True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(l1, l2, is_large_text):

    return l1
```

# --solutions--

```py
def get_contrast_rating(l1, l2, is_large_text):
    ratio = (l1 + 0.05) / (l2 + 0.05)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
