---
id: 6a22d77ddf034bc4e35b1d5b
title: "Sfida 354: valutazione del contrasto 3"
challengeType: 29
dashedName: challenge-354
---

# --description--

Dato due array che rappresentano valori RGB e un booleano che indica se il testo è grande, restituisci la valutazione del contrasto WCAG usando il seguente metodo:

Per prima cosa, converti ogni valore RGB in luminanza relativa:

- Dividi ogni canale `[R, G, B]` per 255 per ottenere un valore tra `0` e `1`
- Applica la formula di correzione gamma a ogni canale:
  - Se il valore del canale è minore o uguale a `0.04045`: `channel / 12.92`
  - Altrimenti: `((channel + 0.055) / 1.055) ^ 2.4`
- Calcola la luminanza: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Poi, calcola il rapporto di contrasto aggiungendo 0.05 a ogni valore di luminanza, quindi dividendo quello più chiaro per quello più scuro. Quello più chiaro sarà sempre il primo argomento.

Restituisci la valutazione basata sul rapporto di contrasto usando la seguente tabella:

| Valutazione | Testo normale | Testo grande |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | sotto 4.5 | sotto 3.0 |

# --hints--

`get_contrast_rating([255, 255, 255], [0, 0, 0], False)` dovrebbe restituire `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([255, 255, 255], [0, 0, 0], False), "AAA")`)
}})
```

`get_contrast_rating([215, 188, 188], [55, 55, 55], False)` dovrebbe restituire `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([215, 188, 188], [55, 55, 55], False), "AA")`)
}})
```

`get_contrast_rating([143, 144, 210], [46, 47, 61], False)` dovrebbe restituire `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([143, 144, 210], [46, 47, 61], False), "Fail")`)
}})
```

`get_contrast_rating([167, 167, 210], [53, 10, 53], True)` dovrebbe restituire `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([167, 167, 210], [53, 10, 53], True), "AAA")`)
}})
```

`get_contrast_rating([135, 147, 155], [60, 70, 90], True)` dovrebbe restituire `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([135, 147, 155], [60, 70, 90], True), "AA")`)
}})
```

`get_contrast_rating([125, 210, 195], [105, 130, 90], True)` dovrebbe restituire `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([125, 210, 195], [105, 130, 90], True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):

    return rgb1
```

# --solutions--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):
    def to_luminance(rgb):
        channels = []
        for c in rgb:
            c = c / 255
            if c <= 0.04045:
                channels.append(c / 12.92)
            else:
                channels.append(((c + 0.055) / 1.055) ** 2.4)
        return 0.2126 * channels[0] + 0.7152 * channels[1] + 0.0722 * channels[2]

    l1 = to_luminance(rgb1)
    l2 = to_luminance(rgb2)
    ratio = (l1 + 0.05) / (l2 + 0.05)

    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
