---
id: 6a22d77ddf034bc4e35b1d54
title: "Sfida 347: teoria dei giochi"
challengeType: 29
dashedName: challenge-347
---

# --description--

Date due stringhe di uguale lunghezza che rappresentano le strategie di due giocatori in un gioco, restituisci i punteggi come un array `[player1, player2]`.

- Le stringhe date conterranno solo una di due lettere: `"C"` (cooperare) o `"D"` (tradire).
- Ogni carattere rappresenta un turno, valutato come segue:
  - Se entrambi i giocatori cooperano, ciascuno ottiene 3 punti.
  - Se entrambi i giocatori tradiscono, ciascuno ottiene 1 punto.
  - Se un giocatore tradisce e l'altro coopera, il traditore ottiene 5 punti e il cooperante 0.

# --hints--

`play_game("CCCC", "CCCC")` dovrebbe restituire `[12, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCC", "CCCC"), [12, 12])`)
}})
```

`play_game("DDDD", "DDDD")` dovrebbe restituire `[4, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDDD", "DDDD"), [4, 4])`)
}})
```

`play_game("CCDD", "CDDD")` dovrebbe restituire `[5, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCDD", "CDDD"), [5, 10])`)
}})
```

`play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD")` dovrebbe restituire `[24, 34]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34])`)
}})
```

`play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` dovrebbe restituire `[66, 21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21])`)
}})
```

# --seed--

## --seed-contents--

```py
def play_game(p1, p2):

    return p1
```

# --solutions--

```py
def play_game(p1, p2):
    scores = [0, 0]
    for i in range(len(p1)):
        if p1[i] == 'C' and p2[i] == 'C':
            scores[0] += 3
            scores[1] += 3
        elif p1[i] == 'D' and p2[i] == 'D':
            scores[0] += 1
            scores[1] += 1
        elif p1[i] == 'D' and p2[i] == 'C':
            scores[0] += 5
        else:
            scores[1] += 5
    return scores
```
