---
id: 6a22d77ddf034bc4e35b1d5c
title: "Sfida 355: codice Morse"
challengeType: 29
dashedName: challenge-355
---

# --description--

Data una stringa in codice Morse, restituisci il messaggio decodificato usando la seguente tabella:

| Codice | Lettera | Codice | Lettera |
|------|--------|------|--------|
| `.-` | `A` | `-.` | `N` |
| `-...` | `B` | `---` | `O` |
| `-.-.` | `C` | `.--.` | `P` |
| `-..` | `D` | `--.-` | `Q` |
| `.` | `E` | `.-.` | `R` |
| `..-.` | `F` | `...` | `S` |
| `--.` | `G` | `-` | `T` |
| `....` | `H` | `..-` | `U` |
| `..` | `I` | `...-` | `V` |
| `.---` | `J` | `.--` | `W` |
| `-.-` | `K` | `-..-` | `X` |
| `.-..` | `L` | `-.--` | `Y` |
| `--` | `M` | `--..` | `Z` |

- Le lettere sono separate da uno spazio singolo
- Le parole sono separate da tre spazi

# --hints--

`decode_morse("--..")` dovrebbe restituire `"Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_morse("--.."), "Z")`)
}})
```

`decode_morse("... --- ...")` dovrebbe restituire `"SOS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_morse("... --- ..."), "SOS")`)
}})
```

`decode_morse("..-. .-. . . -.-. --- -.. . -.-. .- -- .--.")` dovrebbe restituire `"FREECODECAMP"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_morse("..-. .-. . . -.-. --- -.. . -.-. .- -- .--."), "FREECODECAMP")`)
}})
```

`decode_morse(".... . .-.. .-.. ---   .-- --- .-. .-.. -..")` dovrebbe restituire `"HELLO WORLD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_morse(".... . .-.. .-.. ---   .-- --- .-. .-.. -.."), "HELLO WORLD")`)
}})
```

`decode_morse("- .... .   --.- ..- .. -.-. -.-   -... .-. --- .-- -.   ..-. --- -..-   .--- ..- -- .--. . -..   --- ...- . .-.   - .... .   .-.. .- --.. -.--   -.. --- --.")` dovrebbe restituire `"THE QUICK BROWN FOX JUMPED OVER THE LAZY DOG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_morse("- .... .   --.- ..- .. -.-. -.-   -... .-. --- .-- -.   ..-. --- -..-   .--- ..- -- .--. . -..   --- ...- . .-.   - .... .   .-.. .- --.. -.--   -.. --- --."), "THE QUICK BROWN FOX JUMPED OVER THE LAZY DOG")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode_morse(code):

    return code
```

# --solutions--

```py
def decode_morse(code):
    table = {
        '.-': 'A', '-...': 'B', '-.-.': 'C', '-..': 'D', '.': 'E',
        '..-.': 'F', '--.': 'G', '....': 'H', '..': 'I', '.---': 'J',
        '-.-': 'K', '.-..': 'L', '--': 'M', '-.': 'N', '---': 'O',
        '.--.': 'P', '--.-': 'Q', '.-.': 'R', '...': 'S', '-': 'T',
        '..-': 'U', '...-': 'V', '.--': 'W', '-..-': 'X', '-.--': 'Y',
        '--..': 'Z'
    }
    return ' '.join(
        ''.join(table[c] for c in word.split(' '))
        for word in code.split('   ')
    )
```
