---
id: 6a22d77ddf034bc4e35b1d56
title: "Sfida 349: segnale cellulare"
challengeType: 29
dashedName: challenge-349
---

# --description--

Data una griglia contenente tre letture di torri cellulari, determina la posizione del telefono.

- Ogni cella nella griglia è `0` (nessuna torre) oppure un numero intero positivo che rappresenta il numero di celle dal telefono, misurato in linea retta: orizzontale, verticale o diagonale.
- Restituisci la posizione `[row, col]` della cella che si trova alla distanza corretta da tutte e tre le torri.
- C'è sempre esattamente una soluzione.

# --hints--

`find_signal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` dovrebbe restituire `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2])`)
}})
```

`find_signal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` dovrebbe restituire `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1])`)
}})
```

`find_signal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` dovrebbe restituire `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2])`)
}})
```

`find_signal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` dovrebbe restituire `[3, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4])`)
}})
```

`find_signal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` dovrebbe restituire `[3, 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_signal(grid):

    return grid
```

# --solutions--

```py
def find_signal(grid):
    rows = len(grid)
    cols = len(grid[0])

    towers = []
    for r in range(rows):
        for c in range(cols):
            if grid[r][c] != 0:
                towers.append((r, c, grid[r][c]))

    def is_valid(r, c, tr, tc, dist):
        dr = abs(r - tr)
        dc = abs(c - tc)
        return (dr == dist or dc == dist) and (dr == 0 or dc == 0 or dr == dc)

    for r in range(rows):
        for c in range(cols):
            if grid[r][c] != 0:
                continue
            if all(is_valid(r, c, tr, tc, dist) for tr, tc, dist in towers):
                return [r, c]
```
