---
id: 6a26df3e2988bcdded204897
title: "Sfida 361: tempo parlato"
challengeType: 29
dashedName: challenge-361
---

# --description--

Dato l'angolo delle lancette delle ore e dei minuti di un orologio analogico in gradi (in senso orario da 12), restituisci l'ora in inglese parlato.

Converti l'angolo della lancetta dei minuti in minuti (360° = 60 minuti), poi usa le seguenti regole:

| Minuti | Parlato |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (escluso 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (escluso 45) | `"X minutes to Z"` (dove `X` è 60 - minuti) |

Dove `Y` è l'ora corrente e `Z` è l'ora successiva, entrambe ricavate dall'angolo della lancetta delle ore (360° = 12 ore).

Nota: gli angoli delle lancette potrebbero non corrispondere esattamente a un numero, considera di arrotondarli in qualche modo.

# --hints--

`get_spoken_time(90, 0)` dovrebbe restituire `"3 o'clock"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(90, 0), "3 o'clock")`)
}})
```

`get_spoken_time(160, 120)` dovrebbe restituire `"20 minutes past 5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(160, 120), "20 minutes past 5")`)
}})
```

`get_spoken_time(255, 180)` dovrebbe restituire `"half past 8"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(255, 180), "half past 8")`)
}})
```

`get_spoken_time(67.5, 92)` dovrebbe restituire `"quarter past 2"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(67.5, 92), "quarter past 2")`)
}})
```

`get_spoken_time(200, 240)` dovrebbe restituire `"20 minutes to 7"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(200, 240), "20 minutes to 7")`)
}})
```

`get_spoken_time(322.5, 273)` dovrebbe restituire `"quarter to 11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(322.5, 273), "quarter to 11")`)
}})
```

`get_spoken_time(117.5, 335)` dovrebbe restituire `"5 minutes to 4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(117.5, 335), "5 minutes to 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_time(hour_angle, minute_angle):

    return hour_angle
```

# --solutions--

```py
def get_spoken_time(hour_angle, minute_angle):
    minutes = int(minute_angle / 6)
    hour = int(hour_angle / 30) or 12
    next_hour = (hour % 12) + 1

    if minutes == 0:
        return f"{hour} o'clock"
    if minutes == 15:
        return f"quarter past {hour}"
    if minutes == 30:
        return f"half past {hour}"
    if minutes == 45:
        return f"quarter to {next_hour}"
    if minutes < 30:
        return f"{minutes} minutes past {hour}"
    return f"{60 - minutes} minutes to {next_hour}"
```
