---
id: 6a26df95efa55a2524399744
title: "Sfida 363: Riempimento secchio 2"
challengeType: 29
dashedName: challenge-363
---

# --description--

Dato un array bidimensionale di stringhe di colori a singola lettera e un colore obiettivo, restituisci il numero minimo di "clic" di riempimento flood necessari per rendere l'intera griglia del colore obiettivo.

- Ogni clic cambia il colore della cella cliccata e di tutta la regione di celle connesse dello stesso colore con il colore obiettivo.
- Le celle sono connesse orizzontalmente e verticalmente (non diagonalmente).

# --hints--

`bucket_fill([["R", "R"], ["R", "R"]], "G")` dovrebbe restituire `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["R", "R"], ["R", "R"]], "G"), 1)`)
}})
```

`bucket_fill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B")` dovrebbe restituire `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B"), 0)`)
}})
```

`bucket_fill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R")` dovrebbe restituire `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R"), 3)`)
}})
```

`bucket_fill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P")` dovrebbe restituire `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P"), 5)`)
}})
```

`bucket_fill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y")` dovrebbe restituire `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y"), 12)`)
}})
```

# --seed--

## --seed-contents--

```py
def bucket_fill(grid, target_color):

    return grid
```

# --solutions--

```py
def bucket_fill(grid, target_color):
    rows = len(grid)
    cols = len(grid[0])
    visited = set()

    def dfs(r, c, color):
        if r < 0 or r >= rows or c < 0 or c >= cols:
            return
        if (r, c) in visited:
            return
        if grid[r][c] != color:
            return
        visited.add((r, c))
        dfs(r + 1, c, color)
        dfs(r - 1, c, color)
        dfs(r, c + 1, color)
        dfs(r, c - 1, color)

    clicks = 0
    for r in range(rows):
        for c in range(cols):
            if (r, c) not in visited and grid[r][c] != target_color:
                dfs(r, c, grid[r][c])
                clicks += 1

    return clicks
```
