---
id: 6a26df3e2988bcdded204893
title: "Sfida 357: catena alimentare"
challengeType: 29
dashedName: challenge-357
---

# --description--

Dato un array di coppie `[predator, prey]`, restituisci la catena alimentare dall'apice del predatore fino al fondo.

- L'apice del predatore è l'animale che non è mai preda di un altro animale.
- Restituisci la catena come un array di stringhe.

# --hints--

`get_food_chain([["cat", "mouse"]])` dovrebbe restituire `["cat", "mouse"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["cat", "mouse"]]), ["cat", "mouse"])`)
}})
```

`get_food_chain([["wolf", "deer"], ["deer", "grass"]])` dovrebbe restituire `["wolf", "deer", "grass"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"])`)
}})
```

`get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` dovrebbe restituire `["hawk", "snake", "frog", "fly"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"])`)
}})
```

`get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` dovrebbe restituire `["eagle", "fox", "rabbit", "grass"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"])`)
}})
```

`get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` dovrebbe restituire `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_food_chain(pairs):

    return pairs
```

# --solutions--

```py
def get_food_chain(pairs):
    prey_set = {p for _, p in pairs}
    chain_map = dict(pairs)
    apex = next(predator for predator, _ in pairs if predator not in prey_set)
    chain = [apex]
    while chain[-1] in chain_map:
        chain.append(chain_map[chain[-1]])
    return chain
```
