---
id: 6a26df3e2988bcdded204894
title: "Sfida 358: traduttore di emoji"
challengeType: 29
dashedName: challenge-358
---

# --description--

Data una stringa di emoji, restituisci la frase usando la seguente tabella:

| Emoji | Parola |
|-------|--------|
| 👶 | `"baby"` |
| 🐱 | `"cat"` |
| 🐕 | `"dog"` |
| 🐟 | `"fish"` |
| 🥵 | `"hot"` |
| 🧊 | `"ice"` |
| 🪨 | `"rock"` |
| 🦈 | `"shark"` |
| 🍲 | `"soup"` |
| ⭐ | `"star"` |

Restituisci le parole separate da spazi.

# --hints--

`get_emoji_phrase("🪨⭐")` dovrebbe restituire `"rock star"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🪨⭐"), "rock star")`)
}})
```

`get_emoji_phrase("🥵🐕")` dovrebbe restituire `"hot dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🥵🐕"), "hot dog")`)
}})
```

`get_emoji_phrase("👶🦈")` dovrebbe restituire `"baby shark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("👶🦈"), "baby shark")`)
}})
```

`get_emoji_phrase("⭐🐟")` dovrebbe restituire `"star fish"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("⭐🐟"), "star fish")`)
}})
```

`get_emoji_phrase("🧊🧊👶")` dovrebbe restituire `"ice ice baby"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🧊🧊👶"), "ice ice baby")`)
}})
```

`get_emoji_phrase("🐱🐟🍲")` dovrebbe restituire `"cat fish soup"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🐱🐟🍲"), "cat fish soup")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_emoji_phrase(s):

    return s
```

# --solutions--

```py
def get_emoji_phrase(s):
    table = {
        '👶': 'baby',
        '🐱': 'cat',
        '🐕': 'dog',
        '🐟': 'fish',
        '🥵': 'hot',
        '🧊': 'ice',
        '🪨': 'rock',
        '🦈': 'shark',
        '🍲': 'soup',
        '⭐': 'star',
    }
    return ' '.join(table[emoji] for emoji in s)
```
