---
id: 619665c9abd72906f3ad30f9
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Utaunda Flappy Penguin mwenye furaha, na pia utaendelea kufanya mazoezi ya mabadiliko ya CSS na michoro ya kusisimua.

Anza kwa kuunganisha karatasi yako ya mitindo na ukurasa.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.exists(document.querySelector('link'));
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `rel` yenye thamani `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `href` yenye thamani `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
--fcc-editable-region--
    
--fcc-editable-region--
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css

```
