---
id: 65648f4b2281ba50051ae39c
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

次に、`for` ループを `i` をループ変数として書いてください。`range()` 関数を使って `length` の値まで反復してください。

ループの中で、加算代入演算子を使って `all_characters` からのランダムな文字を現在の `password` の値に追加してください。そのために `choice()` モジュールの `secrets` 関数を使います。

# --hints--

`for` を反復する `range(length)` ループを書く必要があります。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

加算代入演算子 `+=` を使って `all_characters` からのランダムな文字を現在の `password` の値に追加してください。

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password
    
--fcc-editable-region--
```
