---
id: 6564a767a7241362f7d8d664
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

単独のアンダースコアは、気にしない値やコードで使わない値を表すために使われます。
繰り返し変数は実際には使われていません。

`i` 変数を単一のアンダースコアに変更してください。

# --hints--

`i` 変数を `_` に変更してください。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/for\s+_\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--
    password = ''
    # Generate password
    for i in range(length):
        password += secrets.choice(all_characters)
    
--fcc-editable-region--
```
