---
id: 65646ffeaed2d238c562a014
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

`import` 文をコードの先頭に置くのが一般的な慣例です。さらに、複数の `import` 文がある場合は、可読性を高めるためにアルファベット順に並べてください。

コードの先頭で、`random` モジュールをインポートしてください。

# --hints--

`random` モジュールをインポートする必要があります。

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
