---
id: 6564b8c9349bd76dc037967b
title: ステップ 29
challengeType: 20
dashedName: step-29
---

# --description--

`search()` モジュールの `re` 関数は、引数として渡された文字列を解析し、正規表現パターンが文字列に最初にマッチする場所を探します。

`quote` という変数を宣言し、文字列 `'Not all those who wander are lost.'` をこの変数に代入してください。その後、`pattern.search(quote)` の結果を出力してください。

# --hints--

`quote` という変数が必要です。

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

提供された文字列を新しい `quote` 変数に代入してください。

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

`pattern.search(quote)` を出力してください。

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
