---
id: 656475bbf1c2573de1d2c69c
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

コードを実行するたびに、`all_characters` 文字列からランダムな文字が表示されるはずです。これはランダムなパスワードを作成するためにまさに達成したいことです。

しかし、`random` が依存しているアルゴリズムは、生成される疑似乱数を予測可能にします。したがって、`random` モジュールは一般的なアプリケーションには適していますが、その決定論的な性質のため暗号目的には使用できません。

`random` をインポートする代わりに、`secrets` モジュールをインポートしてください。そして `print()` の呼び出しを `secrets.choice(all_characters)` を使うように変更してください。

# --hints--

`secrets` モジュールの代わりに `random` モジュールをインポートしてください。

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

既存の `print()` 呼び出しを修正して `secrets.choice(all_characters)` を出力するようにしてください。

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
