---
id: 6564cbb6311a40783b3f5de6
title: ステップ 36
challengeType: 20
dashedName: step-36
---

# --description--

文字クラスでは、マッチさせたい文字の範囲を指定することもできます。開始文字と終了文字をハイフンで区切って指定してください。`-`。文字はUnicode順に並びます。

`pattern`変数の小文字に続く`t`の任意の文字にマッチするように、`quote`変数を修正してください。そのために`a`から`z`までの文字範囲を使います。

# --hints--

`pattern`クラスを使って、小文字に続く`t`内の任意の`quote`にマッチするように`[a-z]`変数を修正してください。

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
