---
id: 656469e62377ab34a5bcb8d1
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

`digits` と `symbols` という名前の新しい変数を2つ宣言し、それぞれに `string.digits` と `string.punctuation` を代入してください。

# --hints--

`digits` という名前の変数を宣言してください。

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

`string.digits` を `digits` 変数に代入してください。

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

`symbols` という名前の変数を宣言してください。

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

`string.punctuation` を `symbols` 変数に代入してください。

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
