---
id: 656485a8a3496d4a36b1496a
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

これら3つの変数は、パスワードを生成するときに選択可能な文字を構成しています。

それらの直前に、`Define the possible characters for the password` というコメントを追加してください。

# --hints--

コメントは `letters` 変数のすぐ上に追加してください。

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
