---
id: 6564642ce403fe29cc23fee7
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

Pythonのモジュールとは、特定のタスクを実行するために設計されたコードを含むファイルです。Python標準ライブラリには、多くのモジュールが含まれており、これらをインポートしてコード内で使用できます。これは、`import` 文の後にモジュール名を続けて行います。

このプロジェクトは、`string` モジュールをインポートすることから始めてください。

# --hints--

`string` モジュールをインポートしてください。

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
