---
id: 6564683821b2ee3174e7250e
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

インポートしたモジュール内で定義されたユーティリティには、ドット記法でアクセスできます。ドット記法は、モジュール名の後にドットとユーティリティ名を付けることで機能します。例えば、`ascii_lowercase` 定数にアクセスする方法は次の通りです。

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

このプロジェクトでは、`string` モジュールから異なる定数を使います。

`letters` という新しい変数を宣言し、この変数に `string.ascii_letters` を代入してください。

# --hints--

`letters` という名前の変数を宣言してください。

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

`string.ascii_letters` を `letters` 変数に代入してください。

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
