---
id: 65649a631440b757f9f7b428
title: ステップ 28
challengeType: 20
dashedName: step-28
---

# --description--

正規表現（regex）は、文字列の中で特定の文字の組み合わせにマッチさせるためのパターンです。これは、文字列の中でパターンをマッチさせたり見つけたりしてバリデーションや文字の置換などを行う場合に、`if`/`else`の条件文の有効な代替手段となります。

`compile()`モジュールの`re`関数は、引数として渡された文字列を正規表現オブジェクトにコンパイルし、他の`re`メソッドで使用できるようにします。

新しい`pattern`変数を宣言し、この変数に`re.compile('i')`の値を代入してください。

# --hints--

`pattern`変数を宣言してください。

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

`re.compile('i')`変数に`pattern`を代入してください。

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
