---
id: 656474314bf37d3c83a85143
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`choice()` モジュールの `random` 関数はシーケンスを受け取り、そのシーケンスのランダムな要素を返します。

`print()` 関数を使い、`choice()` を引数として渡すように `all_characters` の呼び出しを修正してください。

# --hints--

既存の `print()` の呼び出しを修正して、`random.choice(all_characters)` を出力するようにしてください。

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
