---
id: 65646ab6bced52360f44670f
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`all_characters` という名前の変数を宣言し、既存の変数を連結した結果を代入してください。

# --hints--

`all_characters` という名前の変数を宣言する必要があります。

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

`letters`、`digits`、`symbols` を連結し、その結果を `all_characters` 変数に代入してください。

```js
({ test: () => assert(runPython(`
    import string
    from itertools import permutations
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == all_characters for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
