---
id: 6564d97a5ef99783216a6229
title: ステップ 44
challengeType: 20
dashedName: step-44
---

# --description--

Python は *raw* 文字列と呼ばれる特別なタイプの文字列を提供しています。raw 文字列は `r` を接頭辞として付けます。通常の文字列との主な違いはバックスラッシュ文字の扱いにあります。raw 文字列ではバックスラッシュはエスケープ文字としてではなく、そのままの文字として扱われます。正規表現を書くときは、通常多くの `\` 文字を含むため、raw 文字列を使うことが推奨されます。

`pattern` 文字列を `r` を接頭辞として付けることで raw 文字列に変えてください。

# --hints--

`pattern` 変数を `r'\.'` に変更してください。

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
