---
id: 656477845006313fbfea0ad1
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

効果は`random.choice()`と同じに見えるかもしれませんが、`secrets`はオペレーティングシステムが提供できる最も安全な乱数源を保証します。

では、`print()`の呼び出しを2つ削除してください。

# --hints--

`print()`の呼び出しを2つ削除してください。

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
