---
id: 6564d27c8e4c197d367c3b1c
title: ステップ 39
challengeType: 20
dashedName: step-39
---

# --description--

`constraints` リストに新しいタプルを追加してください。最初の要素には `lowercase` を使い、2番目の要素には小文字1文字にマッチする正規表現パターンを使ってください。

# --hints--

`constraints` リストに2つ目のタプルを追加する必要があります。最初の要素には `lowercase` パラメータを使い、2番目の要素には `[a-z]` を使ってください。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\[0-9\]\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*("|')\[a-z\]\2\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, '[0-9]')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = '[^a-z]t'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
