---
id: 656489e0a5d0954c1a64fc93
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`all_characters` 変数を出力して、どのような内容か確認してください。

# --hints--

`all_characters` 変数を出力してください。

```js
({ test: () => assert.match(code, /^print\s*\(\s*all_characters\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
# Combine all characters
all_characters = letters + digits + symbols

--fcc-editable-region--
```
