---
id: 6564d096a55d707bd77ab67b
title: ステップ 38
challengeType: 20
dashedName: step-38
---

# --description--

キャレット `^` は文字クラスの先頭に置くと、そのクラスで指定された文字以外のすべての文字にマッチします。

文字クラスの最初の文字として `^` を追加して、どうなるか確認してください。

# --hints--

文字クラスの最初の文字として `^` を追加してください。

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[^a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '[a-z]t'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
