---
id: 6564883669b5af4b69f794cc
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

あなたの `all_characters` 変数は、小文字と大文字のすべての文字、10個の数字すべて、そしていくつかの特殊文字で構成された文字列です。

その直前に、`Combine all characters` と書かれたコメントを追加してください。

# --hints--

コメントは `all_characters` 変数のすぐ上に追加してください。

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
