---
id: 6564dd65c3c2fa873a83d213
title: ステップ 47
challengeType: 20
dashedName: step-47
---

# --description--

文字クラスでは、大括弧の中に複数の範囲を続けて書くことで、それらの範囲を組み合わせられます（追加の文字なしで）。例えば、`[a-d3-6]` は `[a-d]` と `[3-6]` の組み合わせです。

では、最後の正規表現パターンを修正して、英数字以外の任意の文字にマッチするようにしてください。`a-z`、`A-Z`、`0-9` の範囲を1つの文字クラスにまとめ、先頭に `^` を追加してパターンを否定してください。

# --hints--

4つ目のパターンを完成させて、英数字以外の任意の文字にマッチさせてください。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\[\^a-zA-Z0-9\]\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
