---
id: 656470d517833a39bb8b5608
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

`random` モジュールには疑似乱数生成器が含まれています。その多くの機能は `random()` 関数に依存しており、この関数は `0.0`（含む）から `1.0`（含まない）までの範囲の浮動小数点数を返します。

`random()` モジュールから `random` 関数を呼び出し、その結果を出力してください。

# --hints--

`random.random()` を呼び出した結果を出力してください。

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
