---
id: 6564c67db26c417561ab510d
title: ステップ 31
challengeType: 20
dashedName: step-31
---

# --description--

出力からわかるように、正規表現は文字列内の最初の `l` にマッチしています。

パターンでは、文字の後に量指定子を追加して、その文字が何回繰り返されるかを指定できます。たとえば、`+` という量指定子は、1回以上繰り返されることを意味します。

パターンに `+` という量指定子を追加してください。

# --hints--

`pattern` 変数を `re.compile('l+')` に変更してください。

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
