---
id: 656ef89b4b486b7a16077864
title: ステップ 26
challengeType: 20
dashedName: step-26
---

# --description--

タプルはPythonのもう一つの組み込みデータ構造です。タプルはリストに非常によく似ていますが、大括弧の代わりに丸括弧 `()` で定義されます。また、タプルはリストとは異なり変更不可能です。

```py
my_tuple = ('larch', 1, True)
```

あなたの`constraints`リストはタプルをストアします。各タプルの最初の項目は制約のパラメータになります。

`constraints`リストの代入を修正して、リストにタプルを追加してください。最初の項目には`nums`を使い、2番目の項目には空の文字列を使ってください。

# --hints--

`constraints`リストにタプルを追加する必要があります。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

`(nums, '')`リストに`constraints`を追加してください。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
