---
id: 656ef54070c72161e6feb90f
title: ステップ 22
challengeType: 20
dashedName: step-22
---

# --description--

すべて問題なさそうですが、生成されたパスワードが特定の機能を満たしているかどうかを確認する方法があると便利です。例えば、特殊文字、数字、または大文字・小文字の最小数などです。これについてはすぐに対応します。

今は、コードの最後の2行をコメントアウトしてください。

# --hints--

コードの最後の2行をコメントにしてください。

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
