---
id: 656ef5bc5c0cd464be1df675
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

次に、生成されるパスワードの制約として機能するパラメータを関数に追加します。

既存の`nums`パラメータの後に、順番に`special_chars`、`uppercase`、`lowercase`、`length`を追加して関数宣言を修正してください。

# --hints--

関数は`length`、`nums`、`special_chars`、`uppercase`、`lowercase`をパラメータとして受け取る必要があります。順番が重要です。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password

# new_password = generate_password(8)
# print(new_password)
```
