---
id: 656f0d6bc60b58b9777a36c9
title: ステップ 34
challengeType: 20
dashedName: step-34
---

# --description--

文字クラスは大括弧で示され、その括弧内に指定された文字のうち1文字にマッチします。たとえば、`ma[dnt]` は `mad`、`man`、または `mat` のいずれかにマッチします。

パターンを修正して、`w` の後に `h` または `a` のいずれかが続くものにマッチさせてください。

# --hints--

パターンは文字クラスを使って、`w` の後に `h` または `a` のいずれかが続くものにマッチさせる必要があります。

```js
({ test: () => assert(runPython(`
    pattern == "w[ha]" or pattern == "w[ah]"    
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
