---
id: 656f0bc259c1f6b57486ed68
title: ステップ 33
challengeType: 20
dashedName: step-33
---

# --description--

生成されたパスワードが必要な機能を満たしているか確認するために、`findall()` モジュールの `re` 関数を使います。これは `search` と似ていますが、マッチしたパターンのすべての出現箇所をリストで返します。

`search()` の呼び出しを `findall()` に置き換えて、出力を確認してください。

# --hints--

既存の `print()` の呼び出しを修正し、`search()` を `findall()` に置き換えてください。

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.search(pattern, quote))
--fcc-editable-region--
```
