---
id: 656f10458134d4c4e283a2f1
title: ステップ 35
challengeType: 20
dashedName: step-35
---

# --description--

制約のタプル内の空の文字列を、1桁の数字にマッチする正規表現パターンに変えましょう。`0`から`9`までの各数字を指定する文字クラスを使ってください。

# --hints--

`(nums, '[0123456789]')`リストの中に`constraints`タプルがあるはずです。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\s*\(\s*nums\s*,\s*("|')\[\d{10}\]\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, '')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
