---
id: 6577319039f4f7de9251b822
title: ステップ 61
challengeType: 20
dashedName: step-61
---

# --description--

ループとカウンター変数を使う代わりに、次のいくつかのステップで実装する別の方法で同じ結果を得ることができます。

`all()` は組み込みの Python 関数で、指定されたイテラブル内のすべての要素が `True` と評価される場合に `True` を返します。そうでなければ `False` を返します。

既存の `for` ループと 2 つの `if` 文を、1 つの `if` 文に置き換えてください。`if` の条件には `all()` 関数の呼び出しを使い、引数として空のリストを渡してください。

# --hints--

既存の `for` ループと 2 つの `if` 文を、1 つの `if` 文に置き換える必要があります。

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

新しい `if` 条件は `all([])` にしてください。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

新しい `break` の本体には `if` が含まれているはずです。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
            
        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
