---
id: 6577333feab1e8e927014f03
title: ステップ 63
challengeType: 20
dashedName: step-63
---

# --description--

`all([expression for i in iterable])`を持つということは、`expression`の各`i`に対して`iterable`を評価することで新しいリストが作成されることを意味します。`all()`関数が新しく作成されたリストを反復した後、そのリストはもう必要ないため自動的に削除されます。

メモリはジェネレータ式を使うことで節約できます。ジェネレータ式はリスト内包表記の構文に従いますが、大括弧の代わりに丸括弧を使います。

リスト内包表記の大括弧を削除して、ジェネレータ式に変更してください。

# --hints--

リスト内包表記の大括弧を削除して、ジェネレータ式に変更してください。

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        if all(
            [
                constraint <= len(re.findall(pattern, password))
                for constraint, pattern in constraints
            ]
        ):
--fcc-editable-region--
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
