---
id: 657732654845d8e2fb1217e6
title: ステップ 62
challengeType: 20
dashedName: step-62
---

# --description--

現在、`all()` は空のリストを引数として受け取っています。内包表記の構文を使ってその空のリストを埋め、`constraint <= len(re.findall(pattern, password))` リスト内の各 `constraint`-`pattern` タプルに対して `constraints` 式を評価した結果をリストに格納してください。

こうすることで、すべての条件が満たされた後にのみ `while` ループを抜けることができます。

# --hints--

`[constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints]` を `all()` 関数に渡す必要があります。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("all([constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints])")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        if all([]):
--fcc-editable-region--        
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
