---
id: 657db4cb77190e33a20e852a
title: ステップ 69
challengeType: 20
dashedName: step-69
---

# --description--

デフォルト引数とキーワード引数を組み合わせると、関数が要求するよりも少ない引数を明示的に渡すことができます。関数呼び出しで明示的に渡されなかった引数は、デフォルト値を取ります。

`generate_password()` の呼び出しを修正して、`length=8` のみを渡してください。

# --hints--

`length=8` の呼び出しには `generate_password()` のみを渡してください。

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*length\s*=\s*8\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length=16, nums=1, special_chars=1, uppercase=1, lowercase=1):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    
--fcc-editable-region--
new_password = generate_password(nums=1, length=8, special_chars=1, uppercase=1, lowercase=1)
print(new_password)
--fcc-editable-region--
```
