---
id: 689f7996426c1534fa5bea66
title: Construye un Traductor de Puntuaciones de Golf
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

En el juego de Golf, cada hoyo tiene un `par`, que significa el número promedio de `strokes` que se espera que un golfista haga para hundir la bola en el hoyo y completar la jugada. Dependiendo de qué tan por encima o por debajo del `par` estén tus `strokes`, hay un apodo diferente.

En este laboratorio, escribirás una función que convierte el `par` y los `strokes` a su apodo. 

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función llamada `golfScore`.
1. `golfScore` debe recibir dos argumentos numéricos, que son el par del campo y la cantidad de golpes realizados.
1. `golfScore` debería devolver una cadena.
1. `golfScore` debería devolver `"Hole-in-one!"` si `strokes` es `1`.
1. `golfScore` debería devolver `"Eagle"` si `strokes` es menor o igual a `par` menos `2`.
1. `golfScore` debería devolver `"Birdie"` si `strokes` es igual a `par` menos `1`.
1. `golfScore` debería devolver `"Par"` si `strokes` es igual a `par`.
1. `golfScore` debería devolver `"Bogey"` si `strokes` es igual a `par` más `1`.
1. `golfScore` debería devolver `"Double Bogey"` si `strokes` es igual a `par` más `2`.
1. `golfScore` debería devolver `"Go Home!"` si `strokes` es mayor o igual a `par` más `3`.

# --hints--

Debes crear una función llamada `golfScore`.

```js
assert.isFunction(golfScore);
```

`golfScore` debe tomar dos parámetros.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` debería devolver una cadena

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)` debe devolver la cadena `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` debería devolver la cadena `Hole-in-one!`

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` debe devolver la cadena `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` debería devolver la cadena `Hole-in-one!`

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` debe devolver la cadena `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` debe devolver la cadena `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` debería devolver la cadena `Birdie`

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` debe devolver la cadena `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` debería devolver la cadena `Birdie`

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` debería devolver la cadena `Par`

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` debe devolver la cadena `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` debe devolver la cadena `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` debería devolver la cadena `Bogey`

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` debe devolver la cadena `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` debería devolver la cadena `Bogey`

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` debería devolver la cadena `Double Bogey`

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` debe devolver la cadena `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` debería devolver la cadena `Double Bogey`

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` debería devolver la cadena `Go Home!`

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` debería devolver la cadena `Go Home!`

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` debe devolver la cadena `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
