---
id: 65f0694fb2296f3caadf8347
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

向量 $\mathbf{a}$ 的長度，或範數，通常顯示為 $\\| \mathbf{a} \\|$。它可以計算為其平方分量總和的平方根：
\\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2}
 \\]

計算向量範數並從你的 `norm` 方法傳回結果。

# --hints--

你的 `norm` 方法應該使用提供的公式來計算並傳回一個數字，表述目前向量實體的範數。不要近似該值。

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
