---
id: 65f03d9f92eac9183a4d3281
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

向量可以由歐幾里得平面上的兩個座標 `x` 和 `y` 定義。軸的原點與點 `(x, y)` 之間的距離將是其長度或範數。而向量方向將指向 `(x, y)`。

<img alt="a 2-dimensional vector of coordinates (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

在 `Vector` 類別中，創建一個 `__init__` 方法並給它三個參數，`self`、`x` 和 `y`。

# --hints--

你應該在 `Vector` 類別中定義一個 `__init__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

你的 `__init__` 方法應該接受三個參數：`self`、`x` 和 `y`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
