---
id: 65f056a405239e1dc4cc2854
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

現在，將 `y` 指定到 `Vector` 物件的 `y` 屬性。

# --hints--

你應該將 `y` 指定到 `self.y`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").find_variable("self.y").is_equivalent("self.y = y")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
--fcc-editable-region--
```
