---
id: 65f055b9190fc41ca35549b8
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

Python 提供各種方法，其名稱中包含前置和後置的雙底線。你可能已經熟悉其中一些，例如 `__init__` 和 `__str__`。這些遵循 `__<name>__` 命名樣式的方法稱為特殊方法、魔術方法，或 *dunder*（代表雙底線）方法。

在類別中定義特殊方法會影響該類別的行為。它們會在特定情況下被自動呼叫（例如，`__init__` 在具現化實體時，`__str__` 在物件被列印或傳遞給 `str()` 時）。在這個專案中，你將學習一些最常用的特殊方法。

目前，將 `x` 指定到 `Vector` 物件的 `x` 屬性。

# --hints--

你應該將 `x` 指定到 `self.x`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
