---
id: 65f01969115f933073b6be03
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

向量是一種具有長度（或大小）和方向的物件，且無法用單一數字表述。在物理學中，向量通常用來表述力、速度、加速度及其他量。

在這個專案中你將建置一個向量空間，一個在該集合中元素之間定義一系列操作的集合。你很快就會學到所有細節。

現在，透過宣告一個名為 `Vector` 的空類別來開始專案。

# --hints--

你應該宣告一個名為 `Vector` 的空類別。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
