---
id: 65f055b9190fc41ca35549b8
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

Python 提供了各種方法，它們的名稱中包含前導和尾隨雙下劃線。 你可能已經熟悉一些，例如 `__init__` 和 `__str__`。 這些遵循 `__<name>__` 命名模式的方法被稱爲特殊方法、魔術方法或 *dunder* （代表雙下劃線）方法。

在類中定義特殊方法會影響該類的行爲。 它們在特定情況下在後臺被調用（例如，實例化期間爲 `__init__`，打印對象或將對象傳遞給 `str()` 時爲 `__str__`）。 在這個項目中，你將學習一些最常用的特殊方法。

現在，將 `x` 分配給 `Vector` 對象的 `x` 屬性。

# --hints--

你應該將 `x` 分配給 `self.x`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
