---
id: 65f06a8e5a57673d700c79c3
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

在 `Vector` 類別外，創建一個 `Vector` 實體，傳入整數 `2` 和 `3` 作為 `x` 和 `y` 值，並將其指定給名為 `v1` 的變數。

# --hints--

你應該宣告一個變數 `v1` 並將它指定為 `Vector(2, 3)`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
