---
id: 65f07c9b1ffb814d856dcffc
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

測試 `norm` 是否如預期運作，方法是列印該方法傳回的值。

# --hints--

你應該對 `v1` 呼叫 `norm()` 並列印結果。

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1.norm())")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
--fcc-editable-region--
```
