---
id: 65f058e06f34fd1f0ee6e55d
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

接著，在 `Vector` 類別中，創建一個空的 `norm` 方法並給它一個 `self` 參數。

# --hints--

你的 `Vector` 類別應該有一個 `norm` 方法。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

你的 `norm` 方法應該有一個 `self` 參數。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
