---
id: 65f07c9b1ffb814d856dcffc
title: 步驟 8
challengeType: 20
dashedName: step-8
---

# --description--

通過打印該方法返回的值來測試 `norm` 是否按預期工作。

# --hints--

你應該在 `v1` 上調用 `norm()` 並打印結果。

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1.norm())")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
--fcc-editable-region--
```
