---
id: 65f3fe07cc763212efe91285
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

一個向量可以有數量為 `n` 的維度（組件）。以下是 3 維向量的表述：

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

到目前為止，你已經創建了一個 2 維向量。你想要能夠表現具有不同維度數的向量，而不必為每個特定情況重寫必要的程式碼。為此，你將使用繼承。

首先將 `Vector` 類別重新命名為 `R2Vector`，以指定此類別將處理二維向量。記得也要修改 `v1` 的具現體。

# --hints--

你應該將 `Vector` 類別重新命名為 `R2Vector`。記得也要修改 `v1` 的具現化實體。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
