---
id: 65f3fa097f26b510db6c710b
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

太好了，`norm` 正常運作。現在，如果你嘗試列印 `v1`，你會得到物件的預設字串表現（類似 `<__main__.Vector object at 0x11eb778>`）。

在 `Vector` 類別內，宣告一個空的 `__str__` 方法來實作可讀的字串表現。記得給它一個 `self` 參數。

注意不要列印 `v1`，直到 `__str__` 傳回字串，否則你會得到 `TypeError`，因為 `__str__` 必須始終傳回字串。

# --hints--

你應該在 `Vector` 類別中定義一個 `__str__` 方法。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

你的 `__str__` 方法應該有一個 `self` 參數。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
