---
id: 65f3fa60a93b84110b3f1708
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

你想要傳回一個有意義的字串表現。 例如，在 `v1` 的情況下，你想要傳回一個包含向量組件值的元組字串：`(2, 3)`。

從 `__str__` 方法傳回一個字串，表述向量為包含向量組件依序排列的元組。然後，離開 `Vector` 類別並列印 `v1` 以檢查結果。

# --hints--

你的 `__str__` 方法應該傳回一個字串，該字串表述向量為一個包含向量組件依序排列的元組。

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

你應該列印 `v1`。

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
