---
id: 65f4535bbdb28d436ff3ddc9
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

如你從輸出中所見，`__dict__` 包含你的實體屬性的值。你將不會明確地添加 `self.x` 和 `self.y` 的平方，而是將迭代儲存在 `__dict__` 中的值來計算範數的值。

在 `norm` 方法主體內，將小括弧的內容取代為一個產生器表達式，該表達式將 `self.__dict__.values()` 中的每個值 `val` 提升到 `2` 次方。並且，將該產生器表達式作為引數傳遞給 `sum` 函式，以便在計算平方根之前將所有平方值相加。

# --hints--

你應該從 `norm()` 方法傳回 `sum(val**2 for val in self.__dict__.values())**0.5`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in self.__dict__.values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)

```
