---
id: 65f4520e363e2642f8112e33
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

如你所見，有些地方不太對勁。從 `R2Vector` 繼承的 `norm` 和 `__str__` 方法無法適應三維向量。它們的實作必須更具彈性。

Python 中的每個物件都有一個名為 `__dict__` 的特殊屬性，它是一個字典，用來儲存物件屬性。

移除現有的 `print` 呼叫。然後，列印你的 `v1` 和 `v2` 向量的 `__dict__` 屬性，以查看它們的樣子。

# --hints--

你不應該在你的程式碼中同時有 `print(v2)` 和 `print(v2.norm())`。

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

你應該列印 `v1` 和 `v2` 的 `__dict__` 屬性。

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
