---
id: 65f412c208c3791fee305acf
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

在 Python 中，你可以透過在函式或方法的標記中添加一個 `*` 作為額外的引數來強制使用關鍵字引數。修改兩個 `__init__` 方法，將 `*` 添加為第二個參數（在 `self` 之後）。放在那之後的每個參數，在函式／方法的呼叫中都必須使用關鍵字引數。

這表示你也需要修改 `super().__init__(x, y)` 呼叫。方法是將 `x` 設為值 `x`，並將 `y` 設為值 `y`。

最後，透過使用關鍵字引數來修改 `v1` 和 `v2` 的具現體。

# --hints--

你應該在你的兩個 `__init__` 方法中的第二個參數添加 `*` 來強制使用關鍵字引數。

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

你應該修改你的 `super.__init__(x, y)` 呼叫以使用關鍵字引數。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

你應該將 `v1` 的指派修改為 `R2Vector(x=2, y=3)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

你應該將 `v2` 的指派修改為 `R3Vector(x=2, y=2, z=3)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
