---
id: 65f4148dea0f802040225e0c
title: 步驟 18
challengeType: 20
dashedName: step-18
---

# --description--

移除現有的 `print` 呼叫。然後，如同之前對 `v1` 所做的，列印 `v2` 和 `v2.norm()` 傳回的值。

# --hints--

你不應該在你的程式碼中同時有 `print(v1)` 和 `print(v1.norm())`。

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v1)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v1.norm())")`));
    }
})
```

你應該列印 `v2` 和 `v2.norm()`。

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_call("print(v2)")`));
      assert(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
print(v1.norm())
print(v1)

v2 = R3Vector(x=2, y=2, z=3)
--fcc-editable-region--
```
