---
id: 65f40051d6b09a139f253e8e
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

繼承讓你能從現有的類別定義一個類別。新的類別稱為子類別，繼承現有類別（稱為父類別）的所有方法和屬性。

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass
    
Oak().sprout() # Output: Making new leaves!
```

在上述範例中，子類別 `Oak` 繼承自 `Tree`，並從父類別 `Tree` 繼承了 `sprout` 方法。

創建一個名為 `R3Vector` 的新類別，並依照上述範例使其繼承自 `R2Vector` 類別。

# --hints--

你應該定義一個新的類別 `R3Vector`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

你的新類別 `R3Vector` 應該繼承自 `R2Vector`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
