---
id: 65f40401e6ef53173c04e27d
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

在你的新類別中，宣告一個 `__init__` 方法。為它指定 `self`、`x`、`y` 和 `z` 作為參數。

# --hints--

你應該在你的 `R3Vector` 類別中定義一個 `__init__` 方法。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").has_function("__init__")
    `))
})
```

你的 `__init__` 方法應該有四個參數 `self`、`x`、`y` 和 `z`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, x, y, z")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
