---
id: 65f407ea37ad6e181b90462e
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

你可以像之前一樣將每個 `i` 參數指定給 `self.i`。雖然在這種情況下你只有少量程式碼需要重複，但避免這種重複的一種方法是使用 `super()` 函式。`super()` 讓你能夠隱式地取用父類別：`super().__init__(x, y)` 會呼叫父類別的 `__init__` 方法。

為你的 `__init__` 方法添加一個 `super().__init__(x, y)` 呼叫。

# --hints--

你應該在你的 `__init__` 方法中有一個 `super().__init__(x, y)` 呼叫。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
