---
id: 65f412c208c3791fee305acf
title: Step 17
challengeType: 20
dashedName: step-25
---

# --description--

在 Python 中，你可以通過將 `*` 作爲附加參數添加到函數或方法簽名來強制使用僅限關鍵字參數。 修改兩個 `__init__` 方法，添加 `*` 作爲第二個參數（在 `self` 之後）。 此後放置的每個參數都需要在函數/方法調用中使用關鍵字參數。

這意味着你也需要修改 `super().__init__(x, y)` 調用。 通過賦予 `x` 值 `x`，以及賦予 `y` 值 `y` 來實現。

最後，使用關鍵字參數修改 `v1` 和 `v2` 的實例。

# --hints--

你應該通過在兩個 `__init__` 方法中添加 `*` 作爲第二個參數來強制執行關鍵字參數。

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

你應該修改 `super.__init__(x, y)` 調用以使用關鍵字參數。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

你應該將 `v1` 的賦值修改爲 `R2Vector(x=2, y=3)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

你應該將 `v2` 的賦值修改爲 `R3Vector(x=2, y=2, z=3)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
