---
id: 65f876d17832001e8e1abb05
title: 步驟 28
challengeType: 20
dashedName: step-28
---

# --description--

宣告一個變數 `arg_list`。將一個列表推導式的值賦給它，該推導式對 `vars(self).items()` 中的每個鍵值配對的 `key` 和 `val` 進行迭代，並在每次迭代時計算 f 字串 `f'{key}={val}'`。

# --hints--

你應該宣告一個變數 `arg_list`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_variable("arg_list")`)) })
```

你應該將一個迭代 `vars(self).items()` 的列表推導式指定給 `arg_list`。

```js
({ test: () => runPython(`
import ast
var = _Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list")
assert isinstance(var.tree.value, ast.ListComp), "Expected arg_list to be a list comprehension"
assert (actual := var.find_comp_iters()[0].is_equivalent("vars(self).items()")), f"Expected vars(self).items(), got {actual}"
`) })
```

指定給 `arg_list` 的列表推導式應該使用 `key` 和 `val` 來迭代 `vars(self).items()`。

```js
({ test: () => assert(runPython(`
_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_targets()[0].is_equivalent("key, val")
`)) })
```

指定給 `arg_list` 的列表推導式應該為 `vars(self).items()` 中的每個鍵值對計算字串 `f'{key}={val}'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_expr().is_equivalent("f'{key}={val}'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
