---
id: 65f8709620b2ce1a62608f5a
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

當你需要從字串輸入動態存取某些屬性時，內建的 `getattr()` 函式（程式）就是你需要的。它將一個物件作為第一個引數，並將包含屬性名稱的字串作為第二個引數。

開始修正 `__str__` 方法，將 `__str__()` 傳回的字串替換為一個產生器表達式，該表達式迭代物件的屬性並對每個屬性 `i` 呼叫 `getattr()`。

# --hints--

你應該傳回一個對 `vars(self)` 進行迭代的產生器表達式。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

你應該傳回一個使用 `i` 作為迭代變數的產生器表達式。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

你應該傳回一個產生器表達式，該表達式對 `vars(self)` 中的每個項目 `i` 呼叫 `getattr(self, i)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
