---
id: 65f8709620b2ce1a62608f5a
title: 步驟-23
challengeType: 20
dashedName: 步驟-23
---

# --description--

當你需要從字符串輸入開始動態訪問某些屬性時，內置的 `getattr()` 函數就是你所需要的。 它將對象作爲其第一個參數，並將包含屬性名稱的字符串作爲其第二個屬性。

開始修復 `__str__` 方法，方法是將 `__str__()` 返回的字符串替換爲一個生成器表達式，該生成器表達式遍歷對象屬性併爲每個屬性 `i` 調用 `getattr()`。

# --hints--

你應該返回一個迭代 `vars(self)` 的生成器表達式。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

你應該返回一個使用 `i` 作爲迭代變量的生成器表達式。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

你應該返回一個生成器表達式，該表達式爲 `vars(self)` 中的每個項目 `i` 調用 `getattr(self, i)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
