---
id: 65f8720685ec351abef26740
title: 第 24 步
challengeType: 20
dashedName: step-24
---

# --description--

為了取得正確的字串表述，呼叫 `tuple()` 建構式並將你在前一步撰寫的產生器表達式作為引數傳入。

# --hints--

你應該將 `__str__` 方法傳回的表達式傳遞給 `tuple()` 建構式。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").has_return("tuple(getattr(self, i) for i in vars(self))")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return (getattr(self, i) for i in vars(self))
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
