---
id: 65f8749b13774b1d2e4a7fba
title: 步驟 27
challengeType: 20
dashedName: step-27
---

# --description--

當 `__str__` 方法傳回物件的可讀字串表述時，`__repr__` 方法應該傳回用來實例化該物件的字串。

在 `R2Vector` 類別中為 `self` 參數添加一個 `__repr__` 方法。

# --hints--

你的 `R2Vector` 類別應該有一個 `__repr__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__repr__")`)) })
```

你的 `__repr__` 方法應該有一個 `self` 參數。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
