---
id: 65f9cb678070ca1668898c70
title: 步驟 44
challengeType: 20
dashedName: step-44
---

# --description--

就像在兩個物件相加時會在底層呼叫 `__add__` 一樣，當進行減法時會隱式呼叫 `__sub__` 方法。

現在，定義一個空的 `__sub__` 方法並給予兩個參數：`self` 和 `other`。在你的新方法中，建立一個 if 述語來檢查 `self` 和 `other` 是否不屬於同一個類別，並傳回 `NotImplemented`，如同你之前所做的。

# --hints--

你應該在 `R2Vector` 類別中定義一個 `__sub__` 方法。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__sub__")
    `))
})
```

你的 `__sub__` 方法應該有兩個參數，`self` 和 `other`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").has_args("self, other")
    `))
})
```

你應該創建一個 `if` 述語來檢查 `self` 和 `other` 是否不屬於同一個類別。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
    `))
})
```

你應該從 `if` 區塊內傳回 `NotImplemented`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
