---
id: 65f9b17a638f0d0dcce8c354
title: 步驟 38
challengeType: 20
dashedName: step-38
---

# --description--

要創建向量空間，你需要定義向量在幾種情況下的行爲方式。 可以添加向量，形成新的向量。

可以實現特殊方法 `__add__` 來覆蓋使用 `+` 運算符將兩個對象相加時默認發生的情況。

現在，嘗試添加兩個 `R2Vector` 或 `R3Vector` 實例將引發異常。 在 `R2Vector` 類中創建一個空的 `__add__` 方法，併爲其提供兩個參數：`self` 和 `other`。

# --hints--

你應該在 `R2Vector` 類中定義一個 `__add__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

你的 `__add__` 方法應該有兩個參數，`self` 和 `other`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--

--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
