---
id: 65f9b17a638f0d0dcce8c354
title: 第 38 步
challengeType: 20
dashedName: step-38
---

# --description--

要創建一個向量空間，你需要定義向量在多種情況下應該如何表現。向量可以相加，形成一個新的向量。

可以實作特殊行為 `__add__` 來改寫使用 `+` 運算子將兩個物件相加時的預設行為。

目前，嘗試將兩個 `R2Vector` 或 `R3Vector` 的實體相加會引發例外。在 `R2Vector` 類別中創建一個空的 `__add__` 方法，並給它兩個參數：`self` 和 `other`。

# --hints--

你應該在 `R2Vector` 類別中定義一個 `__add__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

你的 `__add__` 方法應該有兩個參數，`self` 和 `other`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
