---
id: 65f93d4647ce2670dc6f095c
title: 步驟 31
challengeType: 20
dashedName: step-31
---

# --description--

`__repr__` 方法會在底層被 `repr` 函式呼叫。透過將 `f'\nrepr = {repr(v1)}'` 和 `f'\nrepr = {repr(v2)}'` 分別作為你第一個和第二個 `print` 呼叫的第二個引數，來驗證 `__repr__` 是否正常運作。

# --hints--

你應該將 `f'\nrepr = {repr(v1)}'` 作為第二個引數添加到你的 `print(f'v1 = {v1}')` 呼叫中。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v1 = {v1}', f'\\\\nrepr = {repr(v1)}')")`)) })
```

你應該將 `f'\nrepr = {repr(v2)}'` 作為第二個引數添加到你的 `print(f'v2 = {v2}')` 呼叫中。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v2 = {v2}', f'\\\\nrepr = {repr(v2)}')")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
--fcc-editable-region--
```
