---
id: 65f9bc10a9fb1612e066e9e8
title: 步驟 42
challengeType: 20
dashedName: step-42
---

# --description--

當將兩個向量相加時，一個向量的每個組件會加到另一個向量的相同組件上。 例如，將 `(1, 2)` 和 `(2, 4)` 相加會產生第三個向量 `(3, 6)`，其中 `3` 是 `1` 和 `2` 的和，`6` 是 `2` 和 `4` 的和。

`kwargs` 字典將包含用於實例化與兩個相加向量相同類別的新向量所需的鍵值對組。

將空字典轉換為一個字典推導式，該推導式迭代 `vars(self)`，並為每個鍵（`i`）創建一個鍵值對，其中鍵是當前迭代的相同鍵，其值是兩個運算元中 `i` 屬性值的總和。

# --hints--

你應該將指定給 `kwargs` 的空字典轉換成一個對 `vars(self)` 進行迭代的字典推導式。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

指定給 `kwargs` 的字典推導式應該使用變數 `i` 來迭代 `vars(self)`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

指定給 `kwargs` 的字典推導式應該使用 `i` 作為鍵。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

指定給 `kwargs` 的字典推導式應該為 `vars(self)` 中的每個 `i`，將 `self` 的 `i` 屬性與 `other` 的 `i` 屬性相加後指定給 `i` 鍵。為此請使用 `getattr()` 函式（程式）。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
