---
id: 65f9dd6e5a08af19c196c2df
title: 步驟 45
challengeType: 20
dashedName: step-45
---

# --description--

一個向量與另一個向量相減的結果是通過計算它們各個分量的差值而得到的。 例如，從 `(7, 3)` 中減去 `(2, 4)` 會生成第三個向量 `(5, -1)`，其中 `5` 是 `7` 和 `2` 之間的差值，而 `-1` 是 `3` 和 `4` 之間的差值。

就像之前在 `__add__` 方法中所做的那樣，在 `if` 語句後聲明一個 `kwargs` 變量。 爲其分配一個字典推導，該字典推導遍歷對象屬性併爲每個鍵 `i` 創建一個鍵值對，其中鍵與當前迭代的鍵相同，其值是兩個操作數的 `i` 屬性值之間的差值。

# --hints--

你應該聲明一個變量 `kwargs` 併爲其分配一個遍歷對象屬性的字典理解。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

分配給 `kwargs` 的字典理解應該使用變量 `i` 來迭代對象屬性。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

分配給 `kwargs` 的字典理解應該使用 `i` 作爲鍵。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

分配給 `kwargs` 的字典理解應該將 `self` 的 `i` 屬性與 `other` 的 `i` 屬性之間的差異分配給每個 `i` 的 `i` 鍵。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
