---
id: 65f9b486989cb90ff3e77ac8
title: 步驟 39
challengeType: 20
dashedName: step-39
---

# --description--

你的 `__add__` 方法中的 `other` 參數表述加法操作中置於 `+` 運算子右側的運算元。

你想要驗證右側的運算元不是與左側運算元（即 `self`）相同類別的物件。創建一個 `if` 述語，透過檢查運算元的型別來驗證此點。

# --hints--

你應該創建一個 `if` 述語來檢查 `type(self)` 和 `type(other)` 是否不同。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
