---
id: 65f9b9710cca621244d3bde1
title: 步驟 40
challengeType: 20
dashedName: step-40
---

# --description--

在 Python 中，`NotImplemented` 是一個特殊值，用來顯示某個操作行為在特定情況下尚未被實作。

`NotImplemented` 不會立即引發例外。相反地，它會發出信號，詢問另一個運算元如何執行該操作。如果無法滿足該請求，預設會傳回 `TypeError`。

因為你只想在兩個向量屬於相同類別時才能相加，所以從你在前一步驟中創建的 `if` 述語中傳回 `NotImplemented`。

# --hints--

你應該從 `if` 區塊內傳回 `NotImplemented`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
