---
id: 65f93b67169a9c703264458a
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

由於該方法應該傳回用於實例化 `R2Vector` 以及繼承時的 `R3Vector` 物件的字串，因此你不能建置指定類別名稱的字串。

你可以使用 `__class__.__name__` 存取類別的名稱。為你的 `__repr__` 方法添加一個 `return` 述語，並傳回建立該物件所需的字串。

# --hints--

你應該使用 `__class__.__name__` 來建造用於實例化向量物件並從 `__repr__` 方法傳回的 f-string。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_return("f'{self.__class__.__name__}({args})'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
