---
id: 65f9dd6e5a08af19c196c2df
title: 步驟 45
challengeType: 20
dashedName: step-45
---

# --description--

從一個向量中減去另一個向量所得到的向量，是透過計算它們每個組件的差值來取得的。<dfn>例如</dfn>，從 `(7, 3)` 中減去 `(2, 4)` 會產生第三個向量 `(5, -1)`，其中 `5` 是 `7` 與 `2` 的差值，`-1` 是 `3` 與 `4` 的差值。

如同你之前在 `__add__` 方法中所做的，在 `if` 述語之後宣告一個 `kwargs` 變數。為它指定一個字典推導式，該推導式迭代物件屬性並為每個鍵 `i` 產生一個鍵值配對，其中鍵與當前迭代的鍵相同，而其值是兩個運算元中 `i` 屬性值的差異。

# --hints--

你應該宣告一個變數 `kwargs` 並指定一個字典推導式，該推導式會迭代物件屬性。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

指定給 `kwargs` 的字典推導式應該使用變數 `i` 來迭代物件屬性。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

指定給 `kwargs` 的字典推導式應該使用 `i` 作為鍵。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

指定給 `kwargs` 的字典推導式應該為每個 `i`，將 `self` 的 `i` 屬性與 `other` 的 `i` 屬性之差指定給 `i` 鍵。

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
