---
id: 65fa9e1b6c6db919385359ec
title: 步驟 51
challengeType: 20
dashedName: step-51
---

# --description--

可以實作特殊方法 `__mul__` 來指定當目前的實體與另一個物件相乘時應該發生什麼。

在 `R2Vector` 類別中創建一個空的 `__mul__` 方法，並給它兩個參數：`self` 和 `other`。

# --hints--

你應該在 `R2Vector` 類別中定義一個 `__mul__` 方法。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__mul__")
    `))
})
```

你的 `__mul__` 方法應該有兩個參數，`self` 和 `other`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__mul__").has_args("self, other")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) - getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')

```
