---
id: 66682150151af29efec9727d
title: 步驟 33
challengeType: 20
dashedName: step-33
---

# --description--

每當你嘗試存取實體屬性時，`__getattribute__` 方法會在背後被呼叫。如果在實體層級找不到該屬性，該方法會在類別層級搜尋，最終會在父類別中搜尋。

在你的類別內，定義一個帶有兩個參數 `self` 和 `attr` 的 `__getattribute__` 方法，並讓它傳回字串 `'calling __getattribute__'`。你將稍後改寫預設的實作，以便觀察屬性存取的運作方式。

# --hints--

你應該在 `R2Vector` 類別中定義一個名為 `__getattribute__` 的方法，該方法有兩個參數，分別是 `self` 和 `attr`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_args("self, attr")`)) })
```

你的 `__getattribute__` 方法應該傳回字串 `'calling __getattribute__'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_return("'calling __getattribute__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

```
