---
id: 66680ddfd0f8c76782923cb0
title: 第 22 步
challengeType: 20
dashedName: step-22
---

# --description--

`norm()` 方法正在傳回正確的值，但你仍然可以改進一點：可讀性。 

內建函式 `vars()` 以物件作為其引數，並傳回該物件的 `__dict__` 屬性。

與其直接存取 `self` 的 `__dict__` 屬性，不如修改 `norm` 方法以使用 `vars()` 函式。

# --hints--

你應該修改你的 `norm` 方法，使用 `vars(self)` 取代 `self.__dict__`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in vars(self).values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
