---
id: 666832e427d70bc5219dc62a
title: 第 36 步
challengeType: 20
dashedName: step-36
---

# --description--

如你從輸出中所見，雖然你在類別中定義了 `__getattr__`，但此方法尚未被呼叫。這是因為預設的屬性存取是透過 `__getattribute__` 進行的。因此，你可以看到屬性值被列印在終端機上。

現在修改最後兩行程式碼，在你的兩個 `print()` 呼叫中存取 `v1` 的 `z` 屬性。這次，你將存取 `v1` 沒有的屬性。因此，`__getattribute__` 會引發錯誤，並且會呼叫 `__getattr__`。

# --hints--

你應該修改你的程式碼，使用句點運算子和 `getattr()` 函式列印 `v1` 的 `z` 屬性。

```js
({ test: () => assert(runPython(`
n = _Node(_code)
(n.has_call("print(v1.z)") and n.has_call("print(getattr(v1, 'z'))")) or n.has_call("print(v1.z, getattr(v1, 'z'))") or n.has_call("print(getattr(v1, 'z'), v1.z)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __getattr__(self, attr):                
        return 'calling __getattr__'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')
--fcc-editable-region--
print(v1.x)
print(getattr(v1, 'x'))
--fcc-editable-region--
```
