---
id: 666827a6fd0dbaafe8330ea6
title: 步驟 35
challengeType: 20
dashedName: step-35
---

# --description--

`__getattr__` 是另一個在存取屬性時扮演角色的特殊方法。

`__getattribute__` 的預設實作是在請求的屬性不是實體屬性或不在類別樹中時引發 `AttributeError`。

在那種情況下，如果類別定義了 `__getattr__`，就會呼叫它。你可以將它視為當一般屬性存取程序失敗時的一種備援機制。

將 `__getattribute__` 方法轉成 `__getattr__`，並將該方法傳回的字串修改為 `'calling __getattr__'`。

# --hints--

你應該在 `R2Vector` 類別中定義一個名為 `__getattr__` 的方法，該方法有兩個參數，分別是 `self` 和 `attr`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_args("self, attr")`)) })
```

你的 `__getattr__` 方法應該傳回字串 `'calling __getattr__'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_return("'calling __getattr__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __getattribute__(self, attr):                
        return 'calling __getattribute__'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

print(v1.x)
print(getattr(v1, 'x'))

```
