---
id: 579e2a2c335b9d72dd32e05c
title: "implementare l'algoritmo slice e splice"
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

In questo laboratorio, dovrai creare un algoritmo per unire due array.

**Obiettivo:** soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Crea una funzione `frankenSplice` che accetta due array e un indice.
1. Copia ogni elemento del primo array nel secondo array, in ordine, a partire dall'indice dato, e restituisci l'array risultante.
1. Gli array di input devono rimanere invariati dopo l'esecuzione della funzione.

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)` dovrebbe restituire `[4, 1, 2, 3, 5]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)` dovrebbe restituire `["a", 1, 2, "b"]`.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` dovrebbe restituire `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

Tutti gli elementi del primo array devono essere aggiunti al secondo array nel loro ordine originale. `frankenSplice([1, 2, 3, 4], [], 0)` dovrebbe restituire `[1, 2, 3, 4]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

Il primo array deve rimanere invariato dopo l'esecuzione della funzione.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

Il secondo array deve rimanere invariato dopo l'esecuzione della funzione.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
