---
id: 67d83df6f82eda3868dd2a84
title: crea un personaggio per un rpg
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

In questo laboratorio praticherai le basi di Python creando una piccola app che genera un personaggio per un'avventura RPG.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `create_character`.
1. La funzione dovrebbe accettare, in ordine, un nome del personaggio seguito da tre statistiche: forza, intelligenza e carisma.
1. Il nome del personaggio deve essere validato:
   - Se il nome del personaggio non è una stringa, la funzione dovrebbe restituire `The character name should be a string`.
   - Se il nome del personaggio è una stringa vuota, la funzione dovrebbe restituire `The character should have a name`.
   - Se il nome del personaggio è più lungo di 10 caratteri, la funzione dovrebbe restituire `The character name is too long`.
   - Se il nome del personaggio contiene spazi, la funzione dovrebbe restituire `The character name should not contain spaces`.
1. Anche le statistiche devono essere validate:
   - Se una o più statistiche non sono numeri interi, la funzione dovrebbe restituire `All stats should be integers`.
   - Se una o più statistiche sono inferiori a 1, la funzione dovrebbe restituire `All stats should be no less than 1`.
   - Se una o più statistiche sono superiori a 4, la funzione dovrebbe restituire `All stats should be no more than 4`.
   - Se la somma di tutte le statistiche è diversa da 7, la funzione dovrebbe restituire `The character should start with 7 points`.
1. Se tutti i valori superano la verifica, la funzione dovrebbe restituire una stringa con quattro righe:
   - la prima riga deve contenere il nome del personaggio
   - le righe 2-4 devono iniziare con l'abbreviazione della statistica, `STR`, `INT` o `CHA` (in questo ordine), poi uno spazio, e poi un numero di punti pieni (`●`) pari al valore della statistica, e un numero di punti vuoti (`○`) per arrivare a 10. Esempio: se il valore della forza è 3 ci devono essere 3 punti pieni seguiti da 7 punti vuoti. I punti sono forniti nell'editor.

Ecco la stringa che dovrebbe restituire `create_character('ren', 4, 2, 1)`:

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

NOTA: mentre `str` e `int` sono abbreviazioni comuni per le statistiche, ricorda che sono parole chiave riservate in Python e non dovrebbero essere usate come nomi di variabili.

# --hints--

Dovresti avere una funzione chiamata `create_character`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

Quando `create_character` viene chiamata con un primo argomento che non è una stringa dovrebbe restituire `The character name should be a string`.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

Quando `create_character` viene chiamata con un primo argomento che è una stringa non dovrebbe restituire `The character name should be a string`.

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

Quando `create_character` viene chiamata con un primo argomento che è una stringa vuota, dovrebbe restituire `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

Quando `create_character` viene chiamata con un primo argomento che non è una stringa vuota, non dovrebbe restituire `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

Quando `create_character` viene chiamata con un primo argomento più lungo di 10 caratteri dovrebbe restituire `The character name is too long`.

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

La funzione `create_character` non dovrebbe dire che il personaggio è troppo lungo quando non è più lungo di 10 caratteri.

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

Quando `create_character` viene chiamata con un primo argomento che contiene uno spazio dovrebbe restituire `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

Quando `create_character` viene chiamata con un primo argomento che non contiene spazi non dovrebbe restituire `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

Quando `create_character` viene chiamata con un secondo, terzo o quarto argomento che non è un numero intero dovrebbe restituire `All stats should be integers`.

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

Quando `create_character` viene chiamata con un secondo, terzo e quarto argomento che sono tutti numeri interi non dovrebbe restituire `All stats should be integers`.

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

Quando `create_character` viene chiamata con un secondo, terzo o quarto argomento inferiore a `1` dovrebbe restituire `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

Quando `create_character` viene chiamata con un secondo, terzo e quarto argomento che sono tutti almeno `1` non dovrebbe restituire `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

Quando `create_character` viene chiamata con un secondo, terzo o quarto argomento superiore a `4` dovrebbe restituire `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

Quando `create_character` viene chiamata con un secondo, terzo e quarto argomento che sono tutti al massimo `4` non dovrebbe restituire `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

Quando `create_character` viene chiamata con un secondo, terzo o quarto argomento che non sommano a `7` dovrebbe restituire `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

Quando `create_character` viene chiamata con un secondo, terzo e quarto argomento che sommano a `7` non dovrebbe restituire `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` dovrebbe restituire `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

Quando `create_character` viene chiamata con valori validi dovrebbe mostrare le statistiche del personaggio come richiesto.

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
