---
id: 68e76b1ce8216249d04df670
title: Створіть клас планети
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити клас з назвою `Planet`.
1. Клас `Planet` має мати метод `__init__`, який:
   - Має чотири параметри: `self`, `name`, `planet_type` та `star`.
   - Викликає `TypeError` з повідомленням `name, planet type, and star must be strings`, якщо будь-який із переданих аргументів не є рядком.
   - Викликає `ValueError` з повідомленням `name, planet_type, and star must be non-empty strings`, якщо будь-який із переданих аргументів є порожнім рядком.
   - Призначає передані значення атрибутам екземпляра `name`, `planet_type` та `star`.
1. Клас `Planet` має мати метод `orbit`, який повертає рядок у форматі `{name} is orbiting around {star}...`.
1. Клас `Planet` має мати метод `__str__`, який повертає рядок у форматі `Planet: {name} | Type: {planet_type} | Star: {star}`.
1. Ви повинні створити три екземпляри класу `Planet` з іменами `planet_1`, `planet_2` та `planet_3`.
1. Ви повинні вивести кожен об’єкт планети, щоб побачити вихідні дані методу `__str__`.
1. Ви повинні викликати метод `orbit` для кожного об’єкта планети та вивести результат.

# --hints--

Ви повинні створити клас з назвою `Planet`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

Клас `Planet` має мати метод `__init__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

Метод `__init__` має мати чотири параметри: `self`, `name`, `planet_type` та `star` у цьому порядку.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

Метод `__init__` має викликати `TypeError` з повідомленням `name, planet type, and star must be strings`, якщо перший переданий аргумент не є рядком.

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Метод `__init__` має викликати `TypeError` з повідомленням `name, planet type, and star must be strings`, якщо другий переданий аргумент не є рядком.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Метод `__init__` має викликати `TypeError` з повідомленням `name, planet type, and star must be strings`, якщо третій переданий аргумент не є рядком.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Метод `__init__` має викликати `ValueError` з повідомленням `name, planet_type, and star must be non-empty strings`, якщо перший переданий аргумент є порожнім рядком.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Метод `__init__` має викликати `ValueError` з повідомленням `name, planet_type, and star must be non-empty strings`, якщо другий переданий аргумент є порожнім рядком.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Метод `__init__` має викликати `ValueError` з повідомленням `name, planet_type, and star must be non-empty strings`, якщо третій переданий аргумент є порожнім рядком.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Ви повинні призначити `name` атрибуту `self.name` у методі `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

Ви повинні призначити `planet_type` атрибуту `self.planet_type` у методі `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

Ви повинні призначити `star` атрибуту `self.star` у методі `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

Клас `Planet` має мати метод `orbit`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

Метод `orbit` має повертати рядок у форматі `{name} is orbiting around {star}...`.

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

Клас `Planet` має мати метод `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

Метод `__str__` має повертати рядок у форматі `Planet: {name} | Type: {planet_type} | Star: {star}`.

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

Ви повинні створити три екземпляри класу `Planet` з іменами `planet_1`, `planet_2` та `planet_3`.

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

Ви повинні вивести кожен об’єкт планети, щоб побачити вихідні дані методу `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

Ви повинні викликати метод `orbit` для кожного об’єкта планети та вивести результат.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
