---
id: 698c55e7aaacfae48a227459
title: خطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

أضف عنصر `h2` جديدًا أسفل عنصر `audio` مع النص `Cruising for a Musing`.

أسفل عنصر `h2`، أضف عنصر `p` مع النص `Artist: Quincy Larson`.

# --hints--

يجب أن يكون لديك عنصر `h2` ثاني.

```js
assert.exists(document.querySelector('h2:nth-of-type(2)'));
```

يجب أن يحتوي عنصر `h2` الثاني على النص `Cruising for a Musing`.

```js
const h2 = document.querySelector('h2:nth-of-type(2)');
assert.strictEqual(h2?.textContent.trim(), 'Cruising for a Musing');
```

يجب أن يكون لديك عنصر `p` ثاني.

```js
assert.exists(document.querySelector('p:nth-of-type(2)'));
```

يجب أن يحتوي عنصر `p` الثاني على النص `Artist: Quincy Larson`.

```js
const p = document.querySelector('p:nth-of-type(2)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
