---
id: 698c5710136539797f296098
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

أسفل عنصر `p`، أضف عنصر `audio`، وأعطه خاصية `src` بقيمة `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`، وأضف خاصيتي `loop` و`controls`.

# --hints--

يجب أن يكون لديك عنصر `audio` ثاني.

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

يجب أن يكون عنصر `audio` الجديد أسفل عنصر `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

يجب أن يحتوي عنصر `audio` على خاصية `src`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

يجب أن تحتوي خاصية `src` في عنصر `audio` على القيمة `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

يجب أن يحتوي عنصر `audio` على خاصية `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

يجب أن يحتوي عنصر `audio` على خاصية `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
