---
id: 698c4e69abe14a9746c72312
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستبني مشغل موسيقى باستخدام HTML. تم توفير قالب جاهز لـ HTML لك.

أنشئ عنصر `h1` وأعطه النص `freeCodeCamp Tunes`.

# --hints--

يجب أن يكون لديك عنصر `h1`.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن يحتوي عنصر `h1` على النص `freeCodeCamp Tunes`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'freeCodeCamp Tunes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
