---
id: 698c4f4e965bd853ad5084e1
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أسفل العنصر `h1`، أضف عنصر `h2` يحتوي على النص `Can't Stay Down`، هذا هو عنوان الأغنية الأولى.

أسفل عنصر `h2` أضف عنصر `p` يحتوي على النص `Artist: Quincy Larson`.

# --hints--

يجب أن يكون لديك عنصر `h2`.

```js
assert.exists(document.querySelector('h2'));
```

يجب أن يكون عنصر `h2` أسفل عنصر `h1`.

```js
assert.exists(document.querySelector('h1 + h2'));
```

يجب أن يحتوي عنصر `h2` على النص `Can't Stay Down`.

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

يجب أن يكون لديك عنصر `p`.

```js
assert.exists(document.querySelector('p'));
```

يجب أن يحتوي عنصر `p` على النص `Artist: Quincy Larson`.

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
