---
id: 698c5229e47d57c6dfaa91bd
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

تُوفر الخاصية `controls` عناصر تحكم في التشغيل تشمل إيقاف/استئناف التشغيل، والتنقل، والتحكم في مستوى الصوت لعنصر `audio`.

الخاصية `controls` هي خاصية قيمة منطقية ولا تحتاج إلى قيمة.

أضف الخاصية `controls` إلى عنصر `audio`.

الآن يجب أن ترى عنصر `audio` معروضًا على الصفحة.

# --hints--

يجب أن يحتوي عنصر `audio` على الخاصية `controls`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
