---
id: 698c50535b54b2d75fa985fb
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

خاصية `loop` تعيد تشغيل الصوت بمجرد انتهاء التشغيل. فكر في ميم إنترنت يعيد التشغيل بشكل متكرر. حذف خاصية `loop` يجعل الصوت يُشغل مرة واحدة فقط.

خاصية `loop` هي خاصية قيمة منطقية ولا تحتاج إلى قيمة.

أضف خاصية `loop` إلى عنصر `audio`.

# --hints--

يجب أن يحتوي عنصر `audio` على خاصية `loop`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
