---
id: 698c4fe24f1f1a17be74a23e
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

بعد ذلك، أنشئ عنصر `audio` أسفل عنصر `p`. خلال الخطوات التالية، ستضيف الخصائص اللازمة لجعل عنصر `audio` يشغل الموسيقى.

# --hints--

يجب أن يكون لديك عنصر `audio`.

```js
assert.exists(document.querySelector('audio'));
```

يجب أن يكون عنصر `audio` الجديد أسفل عنصر `p`.

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
