---
id: 698c58745b73289bb6c89ece
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

الأغنية الأخيرة بعنوان `Scratching the Surface`، والفنان هو Quincy Larson، والملف المستخدم هو `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

أضف الأغنية الأخيرة لإكمال مشغل الموسيقى باتباع نفس النمط الذي استخدمته للأغنيتين الأوليين.

# --hints--

يجب أن يكون لديك عنصر `h2` ثالث.

```js
assert.exists(document.querySelector('h2:nth-of-type(3)'));
```

يجب أن يحتوي عنصر `h2` الثالث على النص `Scratching the Surface`.

```js
const h2 = document.querySelector('h2:nth-of-type(3)');
assert.strictEqual(h2?.textContent.trim(), 'Scratching the Surface');
```

يجب أن يكون لديك عنصر `p` ثالث.

```js
assert.exists(document.querySelector('p:nth-of-type(3)'));
```

يجب أن يحتوي عنصر `p` الثالث على النص `Artist: Quincy Larson`.

```js
const p = document.querySelector('p:nth-of-type(3)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

يجب أن يكون لديك عنصر `audio` ثالث.

```js
assert.exists(document.querySelector('audio:nth-of-type(3)'));
```

يجب أن يكون عنصر `audio` الجديد أسفل عنصر `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(3) + audio:nth-of-type(3)'));
```

يجب أن يحتوي عنصر `audio` على خاصية `src`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('src'));
```

يجب أن تحتوي خاصية `src` في عنصر `audio` على القيمة `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

يجب أن يحتوي عنصر `audio` على الخاصية `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('loop'));
```

يجب أن يحتوي عنصر `audio` على الخاصية `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

  <h2>Scratching the Surface</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3" loop controls></audio>
  
</body>
</html>
```
