---
id: 69de4308e652aaeb5ccdb03e
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

На цьому практичному занятті ви попрактикуєтеся у роботі з методами рядків у Python.

Почніть зі створення змінної з ім’ям `email` і призначте їй рядок `alice.johnson@company.com`.

# --hints--

Ви повинні створити змінну з ім’ям `email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

Ви повинні призначити рядок `alice.johnson@company.com` змінній `email`. Не забудьте взяти текст у одинарні або подвійні лапки.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
