---
id: 69e20a0451c7a3ef189514d3
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер, коли ви знаєте позицію `@` в е-пошті, ви можете витягти ім'я користувача, яке розташоване перед знаком `@`.

На попередньому практичному занятті ви витягували частину рядка за допомогою зрізів. Нагадаємо, синтаксис такий:

```python
text = 'hello world'
print(text[:5])  # hello
```

Використайте зріз на `email`, щоб витягти все, що знаходиться перед `at_position`, і збережіть результат у змінній з ім’ям `username`.

Потім виведіть `username`, щоб вихідні дані були у форматі `Username: alice.johnson`.

# --hints--

Ви повинні створити змінну з ім’ям `username`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

Ви повинні використати зріз з `at_position` на `email`, щоб призначити `username` значення.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

Ви повинні вивести `username`, щоб вихідні дані були у форматі `Username: alice.johnson`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
