---
id: 69e20876baed33de636b989a
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Іноді ви можете захотіти знайти позицію певного символу всередині рядка. Python надає вбудований метод `.find()` для цієї мети. Він повертає ціле число, яке представляє індекс (позицію) першої появи заданої підрядка.

Ось приклад:

```py
text = 'hello'
print(text.find('e'))  # 1
```

Тепер викличте метод `.find()` для `email` і передайте йому `'@'`, щоб знайти позицію знака @ у рядку `email`. Призначте результат змінній з ім’ям `at_position`.

Потім виведіть `at_position`, щоб вихідні дані були: `Position of @: 13`.

# --hints--

Ви повинні створити змінну з ім’ям `at_position`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

Ви повинні призначити `at_position` значення `email.find('@')`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

Ви повинні вивести `at_position`, щоб вихідні дані були `Position of @: 13`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
