---
id: 69e6284f3b7659c5a05c2a79
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Ви можете перевірити, чи існує підрядок у рядку, використовуючи ключове слово `in`. Воно повертає `True`, якщо підрядок існує, і `False` в іншому випадку. Ось приклад:

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

Створіть змінну з назвою `has_at`, яка перевіряє, чи символ `@` є в `email`.

Потім виведіть `has_at`, щоб вихідні дані були: `Has @ symbol: True`.

# --hints--

Ви повинні створити змінну з назвою `has_at`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

Значення вашої змінної `has_at` має бути `'@' in email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

Ви повинні вивести `has_at`, щоб вихідні дані були `Has @ symbol: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
