---
id: 69e6299710319e7cb8ebdc37
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Реальні дані часто бувають неохайними з зайвими пробілами та непослідовним регістром.

Створіть змінну `messy_name` і призначте їй рядок `sARaH dAVis` з двома пробілами на початку та двома пробілами в кінці. Потім виведіть її так, щоб вихідні дані виглядали так:

```md
Original name: '  sARaH dAVis  '
```

Ви очистите це в наступних кроках.

# --hints--

Ви повинні створити змінну з ім’ям `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

Ви повинні призначити <code>'&nbsp;&nbsp;sARaH dAVis&nbsp;&nbsp;'</code> змінній `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

Ви повинні вивести `messy_name` так, щоб вихідні дані були `Original name: '  sARaH dAVis  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
