---
id: 69e62ce96df237884da25282
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Метод `.lower()` перетворює всі символи в рядку на нижній регістр. Це корисно для нормалізації е-пошти:

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

Ви можете викликати одночасно `.strip()` і `.lower()` для рядка таким чином:

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

Створіть змінну `messy_email` і призначте їй рядок `SARAH.DAVIS@COMPANY.COM` з двома пробілами на початку і двома пробілами в кінці.

Також викличте `.strip().lower()` для `messy_email` і призначте результат змінній `clean_email`.

Потім виведіть обидва, щоб вихідні дані були такими:

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

Ви повинні створити змінну з ім’ям `messy_email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_email')
`) })
```

Ви повинні призначити <code>'&nbsp;&nbsp;SARAH.DAVIS@COMPANY.COM&nbsp;&nbsp;'</code> змінній `messy_email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`) })
```

Ви повинні створити змінну з ім’ям `clean_email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_email')
`) })
```

Ви повинні викликати `.strip().lower()` для `messy_email` і призначити результат змінній `clean_email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_email').is_equivalent("clean_email = messy_email.strip().lower()")
`) })
```

Ви повинні вивести `messy_email`, щоб вихідні дані були `Original email: '  SARAH.DAVIS@COMPANY.COM  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original email: '{messy_email}'\\")",
    "print(\\"Original email: '\\" + messy_email + \\"'\\")",
    "print(\\"Original email: '\\", messy_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Ви повинні вивести `clean_email`, щоб вихідні дані були `Cleaned email: 'sarah.davis@company.com'`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned email: '{clean_email}'\\")",
    "print(\\"Cleaned email: '\\" + clean_email + \\"'\\")",
    "print(\\"Cleaned email: '\\", clean_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
