---
id: 69e73a141addc6f442924b0d
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Ви вже використали `.find()` і зрізи для змінної `email`. Тепер ви застосуєте ту саму техніку до `clean_email`.

Знайдіть позицію символу `@` у `clean_email` і призначте її змінній `at_pos`. Потім використайте зріз, щоб витягти все, що перед `@`, і призначте це змінній `email_user`.

Потім виведіть `clean_email` і `email_user`, щоб вихідні дані виглядали так:

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

Ви повинні створити змінну з ім’ям `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

Ви повинні викликати `.find('@')` для `clean_email` і призначити результат змінній `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

Ви повинні створити змінну з ім’ям `email_user`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

Ви повинні використати зріз із `at_pos` для `clean_email` і призначити результат змінній `email_user`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

Ви повинні вивести `clean_email`, щоб вихідні дані були `Email: sarah.davis@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Ви повинні вивести `email_user`, щоб вихідні дані були `Username: sarah.davis`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
