---
id: 69e736c39dfa403496e8071e
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Метод `.upper()` перетворює всі символи на верхній регістр. Це корисно для бейджиків із іменами або заголовків:

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

Створіть змінну з ім’ям `name_badge` і призначте їй результат виклику методу `.upper()` на `clean_name`.

Потім виведіть `name_badge`, щоб вихідні дані були: `Name badge: SARAH DAVIS`.

# --hints--

Ви повинні створити змінну з ім’ям `name_badge`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('name_badge')
`) })
```

Ви повинні викликати `.upper()` на `clean_name` і призначити результат змінній `name_badge`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`) })
```

Ви повинні вивести `name_badge`, щоб вихідні дані були `Name badge: SARAH DAVIS`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Name badge: {name_badge}')",
    "print('Name badge: ' + name_badge)",
    "print('Name badge:', name_badge)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

--fcc-editable-region--

--fcc-editable-region--
```
