---
id: 6a71a8df10e823e77358d896
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Створіть ще одну змінну з назвою `has_com`, яка перевіряє, чи є `.com` у `email`.

Потім виведіть `has_com`, щоб вихідні дані були: `Has .com: True`.

# --hints--

Ви повинні створити змінну з назвою `has_com`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_com')
`) })
```

Значення вашої змінної `has_com` має бути `'.com' in email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_com').is_equivalent("has_com = '.com' in email")
`) })
```

Ви повинні вивести `has_com`, щоб вихідні дані були `Has .com: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has .com: {has_com}')",
    "print('Has .com: ' + str(has_com))",
    "print('Has .com:', has_com)",
    "print('Has .com:', str(has_com))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')
--fcc-editable-region--

--fcc-editable-region--
```
