---
id: 6a86b9e64da746838e67fb13
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Метод `.strip()` видаляє пробіли з обох кінців рядка.

```python
messy = '  hello world  '
clean = messy.strip()
print(clean)  # Output: hello world
```

Ви можете поєднати `.strip()` і `.title()`, щоб видалити пробіли з обох кінців рядка та зробити першу літеру кожного слова великою:

```python
messy = '  hello world  '
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

Оновіть значення `clean_name`, щоб викликати `.strip()` на `messy_name` перед `.title()`.

Після цього виведіть `clean_name`, щоб вихідні дані були такими: `Cleaned name: 'Sarah Davis'`.

# --hints--

Ви повинні оновити `clean_name`, щоб викликати `.strip()` на `messy_name` перед `.title()`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.strip().title()")
`) })
```

Ви повинні вивести `clean_name`, щоб вихідні дані були такими: `Cleaned name: 'Sarah Davis'`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned name: '{clean_name}'\\")",
    "print(\\"Cleaned name: '\\" + clean_name + \\"'\\")",
    "print(\\"Cleaned name: '\\", clean_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--
clean_name = messy_name.title()
--fcc-editable-region--
```
