---
id: 6a870cc60c23f0332a2fe788
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Пам’ятайте з попереднього практичного заняття, що ви можете використовувати f-рядок, щоб вставити змінну всередину рядка, написавши `f` перед відкриваючою лапкою і обгорнувши змінну в `{}`.

Ось нагадування, як це працює:

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

Тепер виведіть `email`, щоб вихідні дані виглядали так: `Email: alice.johnson@company.com`.

# --hints--

Ви повинні вивести `email`, щоб вихідні дані виглядали як `Email: alice.johnson@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
