---
id: 61695ab9f6ffe951c16d03dd
title: 步驟2
challengeType: 0
dashedName: step-2
---

# --description--

Nest a `head` element within the `html` element. Just after the `head` element, add a `body` element.

# --hints--

應該有一個 `head` 開始標籤。

```js
assert(code.match(/<head\s*>/i));
```

應該有一個 `head` 結束標籤。

```js
assert(code.match(/<\/head\s*>/i));
```

應該有一個 `body` 開始標籤。

```js
assert(code.match(/<body\s*>/i));
```

應該有一個 `body` 結束標籤。

```js
assert(code.match(/<\/body\s*>/i));
```

應該只有一個 `head` 開始標籤。

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

應該只有一個 `head` 結束標籤。

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

應該只有一個 `body` 開始標籤。

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

應該只有一個 `body` 結束標籤。

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

`head` 元素內容應該爲空。

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

`body` 元素內容應該爲空。

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

`body` 元素應該在 `head` 元素之後。

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

`head` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
