---
id: 616965351e74d4689eb6de30
title: 步驟5
challengeType: 0
dashedName: step-5
---

# --description--

You can have multiple `meta` elements on a web page. Each `meta` element adds information about the page that cannot be expressed by other HTML elements.

在 `head` 中添加另一個 `meta` 元素。 給它一個 `name` 屬性，值爲 `"viewport"`，和一個 `content` 屬性，值爲 `"width=device-width, initial-scale=1.0"`，這樣你的頁面在各種設備上看起來是一樣的。

# --hints--

應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

你的新 `meta` 元素應該是一個空元素，它不應該有結束標籤 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```


你的新 `meta` 元素應將 `name` 屬性設置爲 `"viewport"`，並將 `content` 屬性設置爲 `"width=device-width, initial-scale=1.0"`。

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

Your two `meta` elements should be inside the `head` element.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
