---
id: 61696ef7ac756c829f9e4048
title: 步驟 7
challengeType: 0
dashedName: step-7
---

# --description--

在這個項目中，你將使用外部 CSS 文件來設置頁面樣式。 我們已經爲你創建了一個 `styles.css` 文件。 但在你使用它之前，你需要將它鏈接到頁面上。

在 `head` 元素中嵌套一個 `link` 元素。 給它一個 `rel` 屬性，設置爲 `"stylesheet"`，和一個 `href` 屬性，設置爲 `"styles.css"`。

# --hints--

你應該有一個 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

你的 `link` 元素應該是一個空元素，它不應該有結束標籤 `</link>`。

```js
assert.notMatch(code, /<\/link>/i);
```

你的 `link` 元素應該具有值爲 `"stylesheet"` 的 `rel` 屬性。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素應該具有值爲 `"styles.css"` 的 `href` 屬性。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

你的 `link` 元素應該在你的 `head` 元素中。

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
