---
id: 61695d1fbc003856628bf561
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

`charset` 屬性指定了文檔使用的字符編碼。 `utf-8`（Unicode Transforming Format - 8 bit）是一種用於電子通信的字符編碼標準。

在 `head` 元素內，嵌套一個 `meta` 元素，其中 `charset` 屬性值爲 `"utf-8"`。

# --hints--

你應該有一個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

你的 `meta` 元素應該是一個空元素，它不應該有結束標籤 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 元素應該有一個 `charset` 屬性，其值爲 `"utf-8"`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

你的 `meta` 元素應該在 `head` 元素中。

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
