---
id: 616965351e74d4689eb6de30
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

你可以在一個網頁上有多個 `meta` 元素。 每個 `meta` 元素都會添加無法通過其他 HTML 元素來表達的頁面信息。

在 `head` 中添加另一個 `meta` 元素。 給它一個 `name` 屬性，值爲 `"viewport"`，和一個 `content` 屬性，值爲 `"width=device-width, initial-scale=1.0"`，這樣你的頁面在各種設備上看起來是一樣的。

# --hints--

你應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

你的新 `meta` 元素應該是一個空元素，它不應該有結束標籤 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```


你的新 `meta` 元素應將 `name` 屬性設置爲 `"viewport"`，並將 `content` 屬性設置爲 `"width=device-width, initial-scale=1.0"`。

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

你的新 `meta` 元素應該在 `head` 元素中。

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta\s+name\s*=\s*('|"|`)\s*viewport\s*\1\s+content\s*=\s*\1\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
