---
id: 61329d52e5010e08d9b9d66b
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

アクセシビリティとSEOにおいてもう一つ重要な`meta`要素は`description`定義です。`content`属性の値は検索エンジンによってページの説明として使われます。

`meta`属性を`name`に設定した`description`要素を追加し、有用な`content`属性を付けてください。

# --hints--

新しい`meta`要素を`head`に追加してください。

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

`meta`に`name`の`description`属性を付けてください。

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

`meta`に`content`属性を付けてください。

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

`content`属性の値は165文字を超えないようにしてください。_これはGoogleの最大説明長です。_

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
