---
id: 613297a923965e0703b64796
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`meta` 要素にはすでに馴染みがあるかもしれません。これはページのタイトル、説明、キーワード、作者などの情報を指定するために使われます。

ページに適切な `meta` 要素を `charset` の値で追加してください。

`charset` 属性はページの文字エンコーディングを指定し、現在ではほとんどのブラウザで `UTF-8` のみがサポートされています。

# --hints--

`meta` 要素は `head` 要素の中に作成してください。

```js
assert.exists(document.querySelector('head > meta'));
```

`meta` タグに `charset` として `UTF-8` を指定してください。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
