---
id: 6133d11ef548f51f876149e3
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

ナビゲーションはアクセシビリティの重要な部分であり、スクリーンリーダーはページの構造を提供することに依存しています。これはセマンティックなHTML要素で実現されます。

ページに`header`と`main`要素を追加してください。

`header`要素はページの紹介とナビゲーションメニューの提供に使われます。

`main`要素はページの主要なコンテンツを含みます。

# --hints--

`header`に`body`要素を追加してください。

```js
assert.exists(document.querySelector('body > header'));
```

`main`に`body`要素を追加してください。

```js
assert.exists(document.querySelector('body > main'));
```

`header`要素は`main`要素の前に配置してください。

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    
  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
