---
id: 614090d5a22b6f0a5a6b464c
title: ステップ 13
challengeType: 0
dashedName: step-13
---

<!-- TODO: I purposefully added the `nav` styles without Camper input -->

# --description--

子セレクターコンビネーター `>` は、2つのセレクターの間に使い、2番目のセレクターに一致し、かつ1番目のセレクターの直接の子要素である要素だけを対象にします。

これは、深くネストされた要素がある場合に、スタイルのスコープを制御したいときに役立ちます。

`>` セレクターを使って `nav` 要素内の番号なしリスト要素を対象にし、_Flexbox_ を使って子要素を均等に配置してください。

# --hints--

`nav > ul` セレクターを使うべきです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('nav > ul'));
```

`nav > ul` 要素に `display` を `flex` に設定するべきです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('nav > ul')?.display, 'flex');
```

`nav > ul` 要素に `justify-content` を `space-evenly` に設定するべきです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('nav > ul')?.justifyContent, 'space-evenly');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

--fcc-editable-region--

--fcc-editable-region--
```
