---
id: 5d5a813321b9e3db6c106a46
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript es un poderoso lenguaje que permite construir sitios web interactivos.

*Nota*: Para los siguientes proyectos en esta guía de aprendizaje, necesitarás conocimientos básicos de HTML y CSS. Si eres nueva en HTML y CSS, por favor ve a la [Certificación de Diseño Web Responsivo](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

Para comenzar, crea una plantilla estándar de HTML con `DOCTYPE`, `html`, `head` y `body`, añadiendo a continuación una etiqueta `meta` para `charset`. Añade un elemento `title` con el texto `RPG - Dragon Repeller`. Incluye una etiqueta `link` para la hoja de estilos, enlazando el fichero `styles.css`.

Por último, anidado en el `body`, crea un elemento `div` con el atributo `id` establecido a `game`.

# --hints--

Tu código debe tener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debes incluir un espacio después de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debes definir el tipo de documento a `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

El elemento `html` debe tener un tag de apertura. No olvides el atributo `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

El elemento `html` debe tener un tag de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

La declaración `DOCTYPE` debe estar al comienzo del documento HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Debe haber un tag de apertura para `head`.

```js
assert(code.match(/<head\s*>/i));
```

Debe haber un tag de cierre para `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Debe haber un tag de apertura para `body`.

```js
assert(code.match(/<body\s*>/i));
```

Debe haber un tag de cierre para `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Los elementos `head` y `body` deben ser hermanos.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

El elemento `head` debe estar anidado en el elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

El elemento `body` debe estar anidado en el elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Tu código debe tener un elemento `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

El elemento `meta` debe tener un atributo `charset` con el valor `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Debes tener una etiqueta `title` de cierre.

```js
assert(code.match(/<\/title\s*>/i));
```

Tu código debe tener un elemento `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Tu código debe tener un elemento `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

El elemento `div` debe tener un atributo `id` con el valor `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

El elemento `div` debe estar anidado en el elemento `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
