---
id: 62a115879a6d51422652cbfc
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

En tu juego de rol, los usuarios podrán seguir sus estadísticas, comprar armas y luchar contra monstruos. Antes de poder continuar con la parte interactiva de JavaScript del juego, primero necesitas crear los elementos HTML que mostrarán la información del juego.

Create four `div` elements within your `#game` element. Give them the following respective `id` values, in order: `stats`, `controls`, `monsterStats`, and `text`.

# --hints--

You should create four new `div` elements.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

You should give one of the new `div` elements an `id`  of `stats`.

```js
assert.exists(document.querySelector('div#stats'));
```

You should give one of the new `div` elements an `id` of `controls`.

```js
assert.exists(document.querySelector('div#controls'));
```

You should give one of the new `div` elements an `id` of `monsterStats`.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

You should give one of the new `div` elements an `id` of `text`.

```js
assert.exists(document.querySelector('div#text'));
```

You should place the new `div` elements in the correct order.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

You should place the new `div` elements within the `#game` element.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
