---
id: 62a23d1c5f1c93161f3582ae
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Al igual que el elemento `#stats`, el elemento `#monsterStats` necesita dos elementos `span`. Dales la clase `stat` y dale al primer elemento el texto `Monster Name:` y al segundo el texto `Health:`. Después del texto de cada uno, añade un elemento `strong` con un elemento `span` anidado vacío. Dale al primer elemento `span` interno un `id` de `monsterName` y al segundo elemento `span` interno un `id` de `monsterHealth`.

# --hints--

Tu elemento `#monsterStats` debe tener dos elementos `span`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

Los nuevos elementos `span` deben tener un valor de `class` `stat`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

Tu primer elemento `span` debe tener un elemento `strong` con un elemento `span` anidado vacío.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Tu segundo elemento `span` debe tener un elemento `strong` con un elemento `span` anidado vacío.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Tu primer elemento `span` debe tener el texto `Monster Name:`. Asegúrete de que el espaciado es correcto.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+/);
```

Tu segundo elemento `span` debe tener el texto `Health:`. Asegúrete de que el espaciado es correcto.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[1].innerHTML, /Health:\s+/);
```

Los elementos `strong` y `span` deben ir después del texto.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+<strong>/);
assert.match(spans[1].innerHTML, /Health:\s+<strong>/);
```

Tu primer elemento `span` interno debe tener un `id` de `monsterName`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[0]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterName');
```

Tu segundo elemento `span` interno debe tener un `id` de `monsterHealth`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[1]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterHealth');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
--fcc-editable-region--
      <div id="monsterStats">

      </div>
--fcc-editable-region--
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
