---
id: 62a257659d0d1e2456f24ba2
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Antes de empezar a escribir el código del proyecto, debe moverlo a su propio archivo para mantener las cosas organizadas.

Elimina la línea `console.log("Hello World");`. A continuación, asigne a su elemento `script` ahora vacío un atributo `src` establecido en `./script.js`.

# --hints--

No debes tener la línea `console.log("Hello World");` en el código.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

El elemento `script` debe tener un atributo `src` establecido en `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
