---
id: 62a2509ba163e020bb9d84ea
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Ahora puedes empezar a escribir tu código JavaScript. Comienza creando un elemento `script`. Este elemento se utiliza para cargar código JavaScript en tu archivo HTML.

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

Debe tener un elemento `script`.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 1);
```

Tu elemento `script` debe tener una etiqueta de apertura.

```js
assert.match(code, /<script\s*>/i);
```

Tu elemento `script` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
