---
id: 62a3bec30ea7f941412512dc
title: Paso 29
challengeType: 0
dashedName: step-29
---

# --description--

Al igual que hiciste con los botones, crea variables para los siguientes `id` y usa `querySelector()` para darles el elemento como valor:

`text`, `xpText`, `healthText`, `goldText`, `monsterStats`, y `monsterName`.

Recuerda declararlas con la palabra reservada `const`, y nombra las variables para que coincidan con cada `id`.

# --hints--

Debe declarar una variable `text` usando `const`.

```js
assert.match(code, /const\s+text/);
```

Tu variable `text` debe tener el valor de tu elemento `#text`.

```js
assert.deepEqual(text, document.querySelector('#text'));
```

Debes declarar una variable `xpText` usando `const`.

```js
assert.match(code, /const\s+xpText/);
```

Tu variable `xpText` debe tener el valor de tu elemento `#xpText`.

```js
assert.deepEqual(xpText, document.querySelector('#xpText'));
```

Debe declarar una variable `healthText` usando `const`.

```js
assert.match(code, /const\s+healthText/);
```

Tu variable `healthText` debe tener el valor de tu elemento `#healthText`.

```js
assert.deepEqual(healthText, document.querySelector('#healthText'));
```

Debe declarar una variable `goldText` usando `const`.

```js
assert.match(code, /const\s+goldText/);
```

Tu variable `goldText` debe tener el valor de tu elemento `#goldText`.

```js
assert.deepEqual(goldText, document.querySelector('#goldText'));
```

Debes declarar una variable `monsterStats` usando `const`.

```js
assert.match(code, /const\s+monsterStats/);
```

Tu variable `monsterStats` debe tener el valor de tu elemento `#monsterStats`.

```js
assert.deepEqual(monsterStats, document.querySelector('#monsterStats'));
```

Debes declarar una variable `monsterName` usando `const`.

```js
assert.match(code, /const\s+monsterName/);
```

Tu variable `monsterName` debe tener el valor de tu elemento `#monsterName`.

```js
assert.deepEqual(monsterName, document.querySelector('#monsterName'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");

--fcc-editable-region--
```
