---
id: 62a3b41c9494f937560640ab
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Declara dos variables más, llamadas `monsterHealth` e `inventory`, pero no las inicialices.

Para tu variable `inventory`, como valor, asígnale un arreglo con la cadena de texto `"stick"`.

Recuerda que trabajaste con arreglos en el proyecto anterior de esta manera:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

Debes usar `let` para declarar tu variable `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Debes usar camelCase para declarar tu variable `monsterHealth`.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` no debe tener un valor.

```js
assert.isUndefined(monsterHealth);
```

No debes asignarle un valor a tu variable `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

Debes usar `let` para declarar tu variable `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

Tu variable `inventory` aún debe ser un arreglo.

```js
assert.isArray(inventory);
```

Tu variable `inventory` debe tener un solo valor.

```js
assert.lengthOf(inventory, 1);
```

Tu variable `inventory` debe incluir la cadena `"stick"`.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
