---
id: 62a3b79d520a7f3d0e25afd6
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript interactúa con HTML usando el Modelo de Objetos del Documento - DOM, <dfn>Document Object Model</dfn>. EL DOM es un árbol de objetos que representa al HTML. Puedes acceder al HTML usando el objeto `document`, el cual representa todo tu documento HTML.

Una forma de encontrar elementos específicos en tu HTML es usar el método `querySelector()`. El método `querySelector()` toma un selector CSS como argumento y devuelve el primer elemento que coincida con ese selector. Por ejemplo, para encontrar el elemento `<h1>` en tu HTML, escribirías:

```js
let h1 = document.querySelector("h1");
```

Puedes notar que `h1` es una cadena y coincide con el selector CSS que usarías.

Crea una variable `button1` y usa `querySelector()` para asignarle tu elemento con el `id` con el valor `button1`. Recuerda que los selectores CSS de `id` tienen el prefijo `#`.

# --hints--

Debes usar `let` para declarar una variable `button1`.

```js
assert.match(code, /let\s+button1/);
```

Debes usar `document.querySelector()`.

```js
assert.match(code, /document\.querySelector/);
```

Debes usar el selector `"#button1"`.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
