---
id: 62a3a75d8466a12e009eff76
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Inicializa otra variable llamada `health` (salud) con un valor de `100` y una variable llamada `gold` (oro) con un valor de `50`.

# --hints--

Debes usar `let` para declarar una variable llamada `health`.

```js
assert.match(code, /let\s+health/);
```

Debes inicializar una variable llamada `health` con un valor de `100`.

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

Debes usar `let` para declarar una variable llamada `gold`.

```js
assert.match(code, /let\s+gold/);
```

Debes inicializar una variable llamada `gold` con un valor de `50`.

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health` debe tener un valor de `100`.

```js
assert.equal(health, 100);
```

`gold` debe tener un valor de `50`.

```js
assert.equal(gold, 50);
```

`xp` debe tener un valor de `0`.

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
