---
id: 62a8d81f539f004776dd9b1e
title: Paso 110
challengeType: 0
dashedName: step-110
---

# --description--

Ahora puedes iniciar el código para luchar contra monstruos. Para mantener tu código organizado, tu función `fightDragon` se ha movido para que esté cerca de las otras funciones `fight`.

Debajo de tu arreglo `weapons`, define una variable `monsters` y asígnale un arreglo. Establece que ese arreglo tenga tres objetos, cada uno con propiedades `name`, `level` y `health`. Los valores del primer objeto deben ser `"slime"`, `2` y `15`, en orden. El segundo debe ser `"fanged beast"`, `8` y `60`. El tercero debe ser `"dragon"`, `20` y `300`.

# --hints--

Debes usar `const` para declarar una variable `monsters`.

```js
assert.match(code, /const\s+monsters/);
```

Tu variable `monsters` debe ser un arreglo.

```js
assert.isArray(monsters);
```

Tu variable `monsters` debe tener 3 valores.

```js
assert.lengthOf(monsters, 3);
```

Tu arreglo `monsters` debe tener 3 objetos.

```js
assert(monsters.every(val => typeof val === "object"));
```

El primer valor en tu arreglo `monsters` debe ser un objeto con una propiedad `name` establecida en `"slime"`.

```js
assert.equal(monsters[0].name, "slime");
```

El primer valor en tu arreglo `monsters` debe ser un objeto con una propiedad `level` establecida en `2`.

```js
assert.equal(monsters[0].level, 2);
```

El primer valor en tu arreglo `monsters` debe ser un objeto con una propiedad `health` establecida en `15`.

```js
assert.equal(monsters[0].health, 15);
```

El segundo valor en tu arreglo `monsters` debe ser un objeto con una propiedad `name` establecida en `"fanged beast"`.

```js
assert.equal(monsters[1].name, "fanged beast");
```

El segundo valor en tu arreglo `monsters` debe ser un objeto con una propiedad `level` establecida en `8`.

```js
assert.equal(monsters[1].level, 8);
```

El segundo valor en tu arreglo `monsters` debe ser un objeto con una propiedad `health` establecida en `60`.

```js
assert.equal(monsters[1].health, 60);
```

El tercer valor en tu arreglo `monsters` debe ser un objeto con una propiedad `name` establecida en `"dragon"`.

```js
assert.equal(monsters[2].name, "dragon");
```

El tercer valor en tu arreglo `monsters` debe ser un objeto con una propiedad `level` establecida en `20`.

```js
assert.equal(monsters[2].level, 20);
```

El tercer valor en tu arreglo `monsters` debe ser un objeto con una propiedad `health` establecida en `300`.

```js
assert.equal(monsters[2].health, 300);
```

No debes agregar propiedades extra a tus objetos.

```js
assert.deepEqual(monsters, [
  {name: "slime", level: 2, health: 15},
  {name: "fanged beast", level: 8, health: 60},
  {name: "dragon", level: 20, health: 300}
])
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
const weapons = [
  { name: 'stick', power: 5 },
  { name: 'dagger', power: 30 },
  { name: 'claw hammer', power: 50 },
  { name: 'sword', power: 100 }
];
--fcc-editable-region--

--fcc-editable-region--
const locations = [
  {
    name: "town square",
    "button text": ["Go to store", "Go to cave", "Fight dragon"],
    "button functions": [goStore, goCave, fightDragon],
    text: "You are in the town square. You see a sign that says \"Store\"."
  },
  {
    name: "store",
    "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
    "button functions": [buyHealth, buyWeapon, goTown],
    text: "You enter the store."
  },
  {
    name: "cave",
    "button text": ["Fight slime", "Fight fanged beast", "Go to town square"],
    "button functions": [fightSlime, fightBeast, goTown],
    text: "You enter the cave. You see some monsters."
  }
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  update(locations[2]);
}

function buyHealth() {
  if (gold >= 10) {
    gold -= 10;
    health += 10;
    goldText.innerText = gold;
    healthText.innerText = health;
  } else {
    text.innerText = "You do not have enough gold to buy health.";
  }
}

function buyWeapon() {
  if (currentWeaponIndex < weapons.length - 1) {
    if (gold >= 30) {
      gold -= 30;
      currentWeaponIndex++;
      goldText.innerText = gold;
      let newWeapon = weapons[currentWeaponIndex].name;
      text.innerText = "You now have a " + newWeapon + ".";
      inventory.push(newWeapon);
      text.innerText += " In your inventory you have: " + inventory;
    } else {
      text.innerText = "You do not have enough gold to buy a weapon.";
    }
  } else {
    text.innerText = "You already have the most powerful weapon!";
    button2.innerText = "Sell weapon for 15 gold";
    button2.onclick = sellWeapon;
  }
}

function sellWeapon() {
  if (inventory.length > 1) {
    gold += 15;
    goldText.innerText = gold;
    let currentWeapon = inventory.shift();
    text.innerText = "You sold a " + currentWeapon + ".";
    text.innerText += " In your inventory you have: " + inventory;
  } else {
    text.innerText = "Don't sell your only weapon!";
  }
}

function fightSlime() {

}

function fightBeast() {

}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
