---
id: 62b46e3a8d4be31be5af793d
title: Paso 16
challengeType: 0
dashedName: step-16
---

# --description--

Nos hemos encontrado con un pequeño problema. Estás intentando consultar tu página para un elemento de botón, pero tu etiqueta `script` está en la etiqueta `head` de tu HTML. Esto significa que tu código se ejecuta antes de que el navegador haya terminado de leer el HTML, y tu `document.querySelector()` no verá el botón, porque el navegador no lo ha procesado aún.

To fix this, move your `script` element out of the `head` element, and place it at the end of your `body` element (just before the closing `</body>` tag).

# --hints--

Tu elemento `script` no debes estar en tu elemento `head`.

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

Tu elemento `script` debe estar al final de tu elemento `body`.

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
