---
id: 6762ec275cef87635acc4fe3
title: Costruisci una drum machine
challengeType: 25
dashedName: build-drum-machine
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `div` con un `id` di `drum-machine` che contiene tutti gli altri elementi.

2. All’interno dell’elemento `#drum-machine` dovresti avere un altro `div` con un `id` di `pad-bank`.

3. All’interno dell’elemento `#drum-machine` dovresti avere un elemento `p` con un `id` di `display`.

4. All’interno del tuo elemento `#pad-bank` dovresti avere nove elementi drum pad cliccabili, ognuno con una classe `drum-pad`, un `id` unico che descrive il clip audio che il drum pad attiverà, e un testo interno che corrisponde a uno dei seguenti tasti sulla tastiera: `Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`. I drum pad DEVONO essere in questo ordine.

5. Ogni `.drum-pad` dovrebbe avere un elemento `audio` che ha una classe `clip`, un attributo `src` che punta a un clip audio, e un `id` corrispondente al testo interno del suo elemento genitore `.drum-pad` (es. `id="Q"`, `id="W"`, `id="E"` ecc.).

6. Quando clicchi su un elemento `.drum-pad`, il clip audio contenuto nel suo elemento audio figlio dovrebbe essere attivato.

7. Quando premi il tasto trigger associato a ogni `.drum-pad`, il clip audio contenuto nel suo elemento audio figlio dovrebbe essere attivato (es. premendo il tasto `Q` dovrebbe attivare il drum pad che contiene la stringa `Q`, premendo il tasto `W` dovrebbe attivare il drum pad che contiene la stringa `W`, ecc.).

8. Quando un `.drum-pad` viene attivato, dovresti mostrare una stringa che descrive il clip audio associato come testo interno dell’elemento `#display` (ogni stringa deve essere unica).

Alcuni campioni audio che puoi usare per la tua drum machine si trovano su `https://cdn.freecodecamp.org/curriculum/drum/<fileName>`, dove `<fileName>` è il seguente:

| nome del drum | `fileName` |
| --- | --- |
| Heater 1 | `Heater-1.mp3` |
| Heater 2 | `Heater-2.mp3` |
| Heater 3 | `Heater-3.mp3` |
| Heater 4 | `Heater-4_1.mp3` |
| Clap | `Heater-6.mp3` |
| Open-HH | `Dsc_Oh.mp3` |
| Kick-n'-Hat | `Kick_n_Hat.mp3` |
| Kick | `RP4_KICK_1.mp3` |
| Closed-HH | `Cev_H2.mp3` |

# --hints--

Dovresti avere un elemento `div` con un `id` di `drum-machine` che contiene tutti gli altri elementi.

```js
assert.isNotNull(document.querySelector('div#drum-machine'));
```

All’interno dell’elemento `#drum-machine` dovresti avere un altro `div` con un `id` di `pad-bank`.

```js
assert.isNotNull(document.querySelector('#drum-machine div#pad-bank'));
```

All’interno dell’elemento `#drum-machine` dovresti avere un elemento `p` con un `id` di `display`.

```js
assert.isNotNull(document.querySelector('#drum-machine p#display'));
```

All’interno del tuo elemento `#pad-bank` dovresti avere nove elementi drum pad cliccabili, ognuno con una classe `drum-pad`.

```js
assert.lengthOf(document.querySelectorAll('#pad-bank .drum-pad'), 9);
```

Ogni `.drum-pad` dovrebbe essere un elemento `button`.

```js
const pads = document.querySelectorAll('.drum-pad');
pads.forEach(pad => {
  assert.equal(pad.tagName, 'BUTTON');
});
```

Ogni `.drum-pad` dovrebbe avere una delle seguenti lettere come `innerText`, in ordine: `Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`.

```js
const pads = document.querySelectorAll('.drum-pad');
const letters = [`Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`];
for (let i = 0; i < 9; i++) {
  assert.equal(pads[i].innerText.trim(), letters[i]);
}
```

Ogni `.drum-pad` dovrebbe avere un elemento `audio` che ha una classe `clip`, un attributo `src` che punta a un clip audio, e un `id` corrispondente al testo interno del suo elemento genitore `.drum-pad` (es. `id="Q"`, `id="W"`, `id="E"` ecc.).

```js
const letters = [`Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`];
const audios = document.querySelectorAll('#drum-machine .drum-pad audio');

for (let i = 0; i < 9; i++) {
  assert.equal(audios[i].id, letters[i]);
  assert.isTrue(audios[i].classList.contains('clip'));
  assert.isTrue(audios[i].hasAttribute('src'));
}
```

Quando clicchi su un elemento `.drum-pad`, il clip audio contenuto nel suo elemento audio figlio dovrebbe essere attivato.

```js
const pads = document.querySelectorAll('.drum-pad');
assert.isNotEmpty(pads);
pads.forEach(el => {
  const audio = el.querySelector('audio');
  el.click();
  assert.isFalse(audio.paused);
  audio.pause();
});
```

Quando premi uno dei tasti Q, W, E, A, S, D, Z, X, C sulla tastiera, l’elemento `audio` corrispondente dovrebbe riprodurre il suono corrispondente.

```js
const letters = [`Q`, `W`, `E`, `A`, `S`, `D`, `Z`, `X`, `C`];
const audios = document.querySelectorAll('#drum-machine .drum-pad audio');

for (let i = 0; i < 9; i++) {
  const audio = audios[i];
  document.dispatchEvent(
    new KeyboardEvent('keydown', { key: letters[i].toLowerCase() })
  );
  assert.isFalse(audio.paused);
  audio.pause();
}
```

Quando un `.drum-pad` viene attivato, dovresti mostrare una stringa che descrive il clip audio associato come testo interno dell’elemento `#display` (ogni stringa deve essere unica).

```js
const collection = new Set();
const disp = document.querySelector('#display');
const pads = document.querySelectorAll('.drum-pad');
assert.lengthOf(pads, 9);
pads.forEach(el => {
  const audio = el.querySelector('audio');
  el.click();
  collection.add(disp.innerText.trim());
  audio.pause();
});
assert.lengthOf(collection, 9);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Drum Machine</title>
  </head>
  <body></body>
</html>
```

```css

```

```js

```

# --solutions--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Drum Machine</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div class="inner-container" id="drum-machine">
      <div id="pad-bank"></div>

      <div class="controls-container">
        <div class="control power-control">
          <label class="select">
            <span>Power: <span id="power-status">On</span></span>
            <span class="toggle-container">
              <input type="checkbox" id="power-toggle" checked />
              <span class="slider"></span>
            </span>
          </label>
        </div>

        <!-- Display for Drum Pad Name and Volume -->
        <p id="display"></p>
        
        <div class="volume-slider">
          <label for="volume-control">Volume</label>
          <input type="range" id="volume-control" min="0" max="1" step="0.01" value="0.3" />
        </div>

        <div class="control bank-control">
          <label class="select">
            <span>Bank: <span id="current-bank">Heater Kit</span></span>
            <span class="toggle-container">
              <input type="checkbox" id="bank-toggle" checked />
              <span class="slider"></span>
            </span>
          </label>
        </div>
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
/* General Reset */
* {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

body {
  font-family: 'Arial', sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
  background-color: #282c34;
  color: white;
}

.inner-container {
  display: flex;
  flex-direction: column;
  align-items: center;
  background: #3b3b5f;
  border-radius: 8px;
  padding: 20px;
  width: 300px;
}

/* Drum Pad Styles */
#pad-bank {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 10px;
  margin-bottom: 15px;
}

.drum-pad {
  display: flex;
  align-items: center;
  justify-content: center;
  width: 70px;
  height: 70px;
  background: #666;
  border: none;
  border-radius: 5px;
  font-size: 1.2em;
  color: #fff;
  cursor: pointer;
  font-family: inherit;
}

.drum-pad:hover {
  background: #777;
}

.drum-pad:focus-visible {
  outline: 2px solid #0d6efd;
  outline-offset: 2px;
}

.drum-pad.active {
  background: #ffaa33;
  color: #000;
}

.controls-container {
  display: flex;
  flex-direction: column;
  align-items: center;
  width: 100%;
}

.control {
  display: flex;
  justify-content: space-between;
  width: 100%;
  margin: 10px 0;
}

.select {
  display: flex;
  align-items: center;
  justify-content: space-between;
  gap: 12px;
  width: 100%;
}

.toggle-container {
  position: relative;
  display: inline-flex;
  align-items: center;
  width: 50px;
  height: 24px;
  flex: 0 0 auto;
}

.toggle-container input[type="checkbox"] {
  opacity: 0;
  width: 0;
  height: 0;
  position: absolute;
}

.toggle-container .slider {
  position: absolute;
  cursor: pointer;
  top: 0;
  left: 0;
  right: 0;
  bottom: 0;
  background: #6c757d; /* default gray for inactive state */
  border-radius: 12px;
  transition: background 0.25s ease, box-shadow 0.25s ease;
}

.toggle-container .slider::before {
  position: absolute;
  content: "";
  width: 18px;
  height: 18px;
  left: 1px;
  bottom: 3px;
  background: #fff;
  border-radius: 50%;
  transition: transform 0.25s ease, background 0.25s ease;
  box-shadow: 0 1px 2px rgba(0,0,0,0.2);
}

/* move knob and set green background */
.toggle-container input:checked + .slider {
  background: #28a745;
}

.toggle-container input:checked + .slider::before {
  transform: translateX(26px);
}

/* add glow when active */
.toggle-container input#power-toggle:checked + .slider {
  box-shadow: 0 0 0 4px rgba(40,167,69,0.12);
}

/* Bank toggle: always show green and glow in both states */
.toggle-container input#bank-toggle + .slider {
  background: #28a745;
  box-shadow: 0 0 0 4px rgba(40,167,69,0.12);
}

.toggle-container input:focus-visible + .slider {
  outline: 2px solid #0d6efd;
  outline-offset: 2px;
}

#display {
  margin: 15px 0;
  text-align: center;
  font-size: 1em;
  padding: 20px;
  background: #444;
  border-radius: 5px;
  width: 100%;
}

.volume-slider {
  width: 100%;
}

.volume-slider label {
  display: block;
  margin-bottom: 5px;
  font-size: 0.9em;
}

.volume-slider input[type='range'] {
  width: 100%;
  -webkit-appearance: none;
  appearance: none;
  height: 4px;
  background: #ffaa33;
  border-radius: 2px;
  margin-bottom: 10px;
}

.volume-slider input[type='range']::-webkit-slider-thumb {
  -webkit-appearance: none;
  width: 12px;
  height: 12px;
  background: #fff;
  border-radius: 50%;
  cursor: pointer;
}

.volume-slider input[type='range']:focus-visible {
  outline: 2px solid #0d6efd;
  outline-offset: 2px;
}
```

```js
const projectName = 'drum-machine';

const bankOne = [
  {
    keyCode: 81,
    keyTrigger: 'Q',
    id: 'Heater-1',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-1.mp3'
  },
  {
    keyCode: 87,
    keyTrigger: 'W',
    id: 'Heater-2',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-2.mp3'
  },
  {
    keyCode: 69,
    keyTrigger: 'E',
    id: 'Heater-3',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-3.mp3'
  },
  {
    keyCode: 65,
    keyTrigger: 'A',
    id: 'Heater-4',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-4_1.mp3'
  },
  {
    keyCode: 83,
    keyTrigger: 'S',
    id: 'Clap',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Heater-6.mp3'
  },
  {
    keyCode: 68,
    keyTrigger: 'D',
    id: 'Open-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Dsc_Oh.mp3'
  },
  {
    keyCode: 90,
    keyTrigger: 'Z',
    id: "Kick-n'-Hat",
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Kick_n_Hat.mp3'
  },
  {
    keyCode: 88,
    keyTrigger: 'X',
    id: 'Kick',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/RP4_KICK_1.mp3'
  },
  {
    keyCode: 67,
    keyTrigger: 'C',
    id: 'Closed-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Cev_H2.mp3'
  }
];

const bankTwo = [
  {
    keyCode: 81,
    keyTrigger: 'Q',
    id: 'Chord-1',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Chord_1.mp3'
  },
  {
    keyCode: 87,
    keyTrigger: 'W',
    id: 'Chord-2',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Chord_2.mp3'
  },
  {
    keyCode: 69,
    keyTrigger: 'E',
    id: 'Chord-3',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Chord_3.mp3'
  },
  {
    keyCode: 65,
    keyTrigger: 'A',
    id: 'Shaker',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Give_us_a_light.mp3'
  },
  {
    keyCode: 83,
    keyTrigger: 'S',
    id: 'Open-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Dry_Ohh.mp3'
  },
  {
    keyCode: 68,
    keyTrigger: 'D',
    id: 'Closed-HH',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Bld_H1.mp3'
  },
  {
    keyCode: 90,
    keyTrigger: 'Z',
    id: 'Punchy-Kick',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/punchy_kick_1.mp3'
  },
  {
    keyCode: 88,
    keyTrigger: 'X',
    id: 'Side-Stick',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/side_stick_1.mp3'
  },
  {
    keyCode: 67,
    keyTrigger: 'C',
    id: 'Snare',
    url: 'https://cdn.freecodecamp.org/curriculum/drum/Brk_Snr.mp3'
  }
];

const state = {
  power: true,
  display: ' ',
  currentPadBank: bankOne,
  currentPadBankId: 'Heater Kit',
  sliderVal: 0.3
};

const display = document.getElementById('display');

const setDisplay = text => {
  state.display = text;
  display.innerText = text;
};

const adjustVolume = e => {
  if (state.power) {
    state.sliderVal = e.target.value;
    setDisplay(`Volume: ${Math.round(e.target.value * 100)}`);
    setTimeout(() => setDisplay(' '), 1000);
    document.querySelectorAll('.clip').forEach(sound => {
      sound.volume = state.sliderVal;
    });
  }
};

const activatePad = pad => {
  pad.classList.add('active');
  setTimeout(() => pad.classList.remove('active'), 100);
};

const playSound = (id, name) => {
  if (!state.power) return;

  const sound = document.getElementById(id);
  sound.currentTime = 0;
  sound.play();
  activatePad(sound.parentNode);
  setDisplay(name);
};

const renderPadBank = () => {
  const padBank = document.querySelector('#pad-bank');
  padBank.innerHTML = '';

  state.currentPadBank.forEach(drum => {
    const pad = document.createElement('button');
    pad.classList.add('drum-pad');
    pad.id = drum.id;
    pad.innerText = drum.keyTrigger;

    const audio = document.createElement('audio');
    audio.classList.add('clip');
    audio.id = drum.keyTrigger;
    audio.src = drum.url;
    pad.appendChild(audio);

    pad.addEventListener('click', () =>
      playSound(drum.keyTrigger, drum.id.replace(/-/g, ' '))
    );
    document.addEventListener('keydown', e => {
      if (e.key === drum.keyTrigger.toLowerCase())
        playSound(drum.keyTrigger, drum.id.replace(/-/g, ' '));
    });

    padBank.appendChild(pad);
  });
};

const powerControl = () => {
  state.power = !state.power;
  document.getElementById('power-status').textContent = state.power
    ? 'On'
    : 'Off';

  setDisplay('');
};

const selectBank = () => {
  if (!state.power) return;
  setDisplay(' ');
  state.currentPadBank =
    state.currentPadBankId === 'Heater Kit' ? bankTwo : bankOne;
  state.currentPadBankId =
    state.currentPadBankId === 'Heater Kit' ? 'Smooth Piano Kit' : 'Heater Kit';

  renderPadBank();
  document.getElementById('current-bank').textContent =
    state.currentPadBankId === 'Heater Kit' ? 'Heater Kit' : 'Smooth Piano Kit';
};

document
  .getElementById('power-toggle')
  .addEventListener('change', powerControl);
document
  .getElementById('bank-toggle')
  .addEventListener('change', selectBank);
document
  .getElementById('volume-control')
  .addEventListener('input', adjustVolume);

renderPadBank();
```
