---
id: 615f522dea4f776f64dc3e91
title: Step 31
challengeType: 0
dashedName: step-31
---

# --description--

`rem` 단위는 `root em`를 나타내며 `html` 요소의 폰트 크기를 기준으로 합니다.

`.small-text` 선택자를 만들고 `font-size`를 `0.85rem`로 설정하세요. 이는 대략 `13.6px`에 해당합니다 (`html`의 `font-size`를 `16px`로 설정한 것을 기억하세요).

# --hints--

`.small-text` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.small-text'));
```

`.small-text` 선택자의 `font-size` 속성이 `0.85rem`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.small-text')?.fontSize === '0.85rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
    <div class="calories-info">
      <div class="left-container">
        <h2 class="bold small-text">Amount per serving</h2>
      </div>
    </div>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
