---
id: 615f2d4150fe0d4cbd0f2628
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

`h1` 요소 아래에 `p` 요소를 추가하고 텍스트는 `8 servings per container`로 하세요.

# --hints--

새로운 `p` 요소를 추가해야 합니다.

```js
assert.exists(document.querySelector('p'));
```

`p` 요소는 `body` 요소 안에 있어야 합니다.

```js
assert.equal(document.querySelector('p')?.parentElement?.localName, 'body');
```

`p` 요소는 `h1` 요소 다음에 와야 합니다.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

`p` 요소는 텍스트가 `8 servings per container`이어야 합니다.

```js
assert.equal(document.querySelector('p')?.innerText, '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
    
</body>
</html>
--fcc-editable-region--
```

```css

```
