---
id: 615f3e1b7233ee5c7595771f
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

`h1`의 위치를 미세 조정하려면 위아래 여백을 `-4px`로, 좌우 여백을 `0`로 설정하세요.

# --hints--

`h1` 선택자에 `margin` 속성을 `-4px 0`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.margin === '-4px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--
h1 {
  font-weight: 800;
  text-align: center;
}
--fcc-editable-region--
```
