---
id: 615f51e4e5b24a6e80eccce1
title: Step 30
challengeType: 0
dashedName: step-30
---

# --description--

`.calories-info` 요소 안에 `div` 요소를 만드세요. 그 `div` 요소에 `class` 속성을 `left-container`로 설정하세요. 새로 만든 `div` 요소 안에 텍스트 `h2`을 가진 `Amount per serving` 요소를 만드세요. `h2` 요소에 `class` 속성을 `bold small-text`으로 설정하세요.

# --hints--

`div` 요소 안에 새 `.calories-info` 요소를 만들어야 합니다.

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.localName === 'div');
```

새 `div` 요소에 `class` 속성을 `left-container`로 설정해야 합니다.

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.classList.contains('left-container'));
```

`h2` 요소 안에 새 `.left-container` 요소를 만들어야 합니다.

```js
assert(document.querySelector('.left-container')?.children?.[0]?.localName === 'h2');
```

새 `h2` 요소에 `class` 속성을 `bold small-text`으로 설정해야 합니다.

```js
assert(document.querySelector('.left-container')?.children?.[0]?.classList.contains('bold'));
assert(document.querySelector('.left-container')?.children?.[0]?.classList.contains('small-text'));
```

새 `h2` 요소에 텍스트 `Amount per serving`가 있어야 합니다.

```js
assert(document.querySelector('.left-container')?.children?.[0]?.innerText === 'Amount per serving');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
--fcc-editable-region--
    <div class="calories-info">

    </div>
--fcc-editable-region--
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}
```
