---
id: 615f38279e5c3d53692ea441
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

`h1`와 `p` 요소를 `div` 요소로 감싸세요. 그 `div`에 `class` 속성을 `label`로 설정하세요.

# --hints--

새로운 `div` 요소를 만들어야 합니다.

```js
assert(document.querySelector('div'));
```

새로운 `div` 요소에 `class` 속성을 `label`로 설정해야 합니다.

```js
assert(document.querySelector('div')?.classList?.contains('label'));
```

`h1`와 `p` 요소는 새로운 `.label` 요소 안에 있어야 합니다.

```js
const children = [...document.querySelectorAll('h1'), ...document.querySelectorAll('p')];
assert(children?.every(child => child?.parentElement?.classList?.contains('label')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

--fcc-editable-region--
<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
--fcc-editable-region--
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}
```
