---
id: 615f3949f58e12577dcefb00
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

적절한 여백 사용은 페이지의 중요한 요소에 포커스를 주고, 사용자의 시선을 텍스트 전체로 자연스럽게 안내하는 데 도움이 됩니다.

`.label` 선택자에 `width` 속성을 `270px`로 설정하세요.

# --hints--

`.label` 선택자에 `width` 속성을 `270px`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.width === '270px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--
.label {
  border: 2px solid black;
}
--fcc-editable-region--
```
