---
id: 615f38cabc64e3556f98cc1a
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

테두리는 콘텐츠를 그룹화하고 우선순위를 정하는 데 사용할 수 있습니다.

`.label` 선택자를 만들고 `border`를 `2px solid black`로 설정하세요.

# --hints--

`.label` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label'));
```

`.label` 선택자에 `border` 속성을 `2px solid black`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.border === '2px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--

--fcc-editable-region--
```
