---
id: 615f405b89a7ec5f8e2d11f4
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

새로운 `.divider`에 대한 선택자를 만들고 `border-bottom` 속성을 `1px solid #888989`로 설정하세요. 또한 위쪽과 아래쪽 여백을 `2px`로 지정하세요. 좌우 여백은 없어야 합니다.

# --hints--

`.divider` 선택자를 만들어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.divider'));
```

`.divider` 선택자에 `border-bottom` 속성을 `1px solid #888989`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.divider')?.borderBottom === '1px solid rgb(136, 137, 137)');
```

`.divider` 선택자에 `margin` 속성을 `2px 0`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.divider')?.margin === '2px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
}

p {
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
