---
id: 615f41c979787462e76dab90
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

새로운 클래스에는 아직 스타일이 없습니다. `.bold` 선택자를 만들고 텍스트를 굵게 만들기 위해 `font-weight` 속성을 `800`로 설정하세요.

또한 `font-weight` 선택자에서 `h1` 속성을 제거하세요.

# --hints--

`.bold` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold'));
```

`.bold` 선택자에는 `font-weight` 속성이 `800`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold')?.fontWeight === '800');
```

`h1` 선택자에는 `font-weight` 속성이 없어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === "");
```

`h1` 선택자를 제거하면 안 됩니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p class="bold">Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--
h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

--fcc-editable-region--
```
