---
id: 615f378014c2da526a109c81
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

폰트가 조금 작습니다. `html` 선택자를 만들고 폰트 크기를 `16px`로 설정하세요.

# --hints--

`html` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

`html` 선택자에 `font-size` 속성을 `16px`로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.fontSize === '16px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

body {
  font-family: 'Open Sans', sans-serif;
}
```
