---
id: 615f4172e9eec061d6456221
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

영양 레이블에는 중요한 정보를 강조하기 위해 굵은 텍스트가 많이 사용됩니다. 굵게 표시해야 하는 각 요소를 개별적으로 지정하기보다는, 클래스를 사용해 모든 요소에 굵은 스타일을 적용하는 것이 더 효율적입니다.

두 번째 `p` 요소에 `class` 속성을 `bold`로 설정하세요.

# --hints--

두 번째 `p` 요소에 `class` 속성을 `bold`로 설정해야 합니다.

```js
assert(document.querySelectorAll('p')?.[1]?.classList?.contains('bold'));
```

첫 번째 `p` 요소에는 `class` 속성을 `bold`로 설정하지 마세요.

```js
assert(!document.querySelector('p')?.classList?.contains('bold'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}
```
