---
id: 615f2abbe7d18d49a1e0e1c8
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

기본 HTML 보일러플레이트가 제공되었습니다.

`h1` 요소 안에 `body` 요소를 만들고 텍스트를 `Nutrition Facts`로 지정하세요.

# --hints--

새로운 `h1` 요소를 추가하세요.

```js
assert.exists(document.querySelector('h1'));
```

`h1` 요소 안에 `body` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector('h1')?.parentElement?.localName, 'body');
```

`h1` 요소에 텍스트 `Nutrition Facts`가 있어야 합니다.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
