---
id: 615f34948891834dd77655a6
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

텍스트가 `p`인 두 번째 `Serving size 2/3 cup (55g)` 요소를 추가하세요.

# --hints--

두 번째 `p` 요소가 있어야 합니다.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

두 번째 `p` 요소는 `body` 요소 안에 있어야 합니다.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

두 번째 `p` 요소는 기존 `p` 요소 뒤에 와야 합니다.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

두 번째 `p` 요소는 텍스트가 `Serving size 2/3 cup (55g)` 이어야 합니다.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
    
</body>
</html>
--fcc-editable-region--
```

```css

```
