---
id: 68e43e884ef09d4d2ab05631
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

现在是时候查看使用 `replace()` 方法的结果了。

首先为你的代码添加一个带有值 `"After using the replace() method:"` 的 `console.log()`。

然后在那个 `console.log()` 下面，添加另一个带有 `replacedString` 变量的 `console.log()`。

查看控制台，你应该看到新句子显示为 `"I love dogs."`，而不是原来的 `"I love cats."`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将 `"After using the replace() method:"` 记录到控制台日志。

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

你应该将 `replacedString` 变量记录到控制台日志中。 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
