---
id: 68e448ec4ac4205968c0a805
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

如果在句子 `"I love learning!"` 中 `"love"` 这个词重复三次，那就太好了。 

正如你在之前的课程中回顾的，你可以使用 `repeat()` 方法重复一个字串指定的次数。以下是一个示例：

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

由于字串是不可变的，此方法不会修改原始字串。它将返回一个包含重复内容的新建字串。 

创建一个名为 `repeatedLove` 的变量并赋值为 `"love ".repeat(3)`。然后将 `repeatedLove` 变量日志到控制台，以便你看到结果。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该有一个名为 `repeatedLove` 的变量。

```js
assert.exists(repeatedLove);
```

你的 `repeatedLove` 变量应是一个字串。

```js
assert.isString(repeatedLove);
```

你应该将 `"love ".repeat(3)` 赋值给 `repeatedLove` 变量。

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

你应该将 `repeatedLove` 变量记录到控制台日志。

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
