---
id: 68e42b62d7093345c6975f22
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

在你的 `originalString` 变量下方，添加一个带有字串 `"Original string:"` 的 `console.log()`。 

然后在那个 `console.log()` 下面，添加另一个带有 `originalString` 变量的 `console.log()`。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将字串 `"Original string:"` 记录到控制台日志。

```js
assert.equal(getLogs()[0], "Original string:");
```

你应该将 `originalString` 变量记录到控制台日志。 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
