---
id: 68e41084bb275539f72de9fe
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将通过创建一个字串转换器应用来练习使用 `replace()`、`replaceAll()` 和 `repeat()` 方法。 

首先，创建一个名为 `originalString` 的变量，并赋值为 `"I love cats."`。

# --hints--

你应该有一个名为 `originalString` 的变量。

```js
assert.exists(originalString);
```

你的 `originalString` 变量应是一个字串。

```js
assert.isString(originalString);
```

你应该将字串 `"I love cats."` 赋值给 `originalString` 变量。

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
