---
id: 68e4417b7317835208145323
title: 步骤6
challengeType: 1
dashedName: step-6
---

# --description--

接下来，创建一个名为 `dogsOnlySentence` 的新变量，并赋值为 `exampleSentence.replace("cats", "dogs")`。

然后在该变量下方，向控制台日志 `"Replacing all occurrences of cats with dogs:"`。在该 `console.log()` 下方，为 `dogsOnlySentence` 变量添加另一个 `console.log()`。 

你现在应该看到文本 `"I love dogs and cats are so much fun!"` 被记录到控制台。但这不是期望的结果。在下一步，你将学习如何修复它，使文本显示为 `"I love dogs and dogs are so much fun!"`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该有一个 `dogsOnlySentence` 变量。

```js
assert.exists(dogsOnlySentence);
```

你的 `dogsOnlySentence` 变量应是一个字串。

```js
assert.isString(dogsOnlySentence);
```

你应该将 `exampleSentence.replace("cats", "dogs")` 赋值给你的 `dogsOnlySentence` 变量。

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

你应该将 `"Replacing all occurrences of cats with dogs:"` 记录到控制台日志。

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

你应该将 `dogsOnlySentence` 变量记录到控制台日志。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
