---
id: 68e42de6396f5c49d90d1951
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

在之前的课程中，你学习了用于定位字串中子字串并将其替换为另一个值的 `replace()` 方法。

这是一个示例：

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

请记住，字串是不可变的，这意味着在该示例中原始的 `text` 不会被修改。相反，会创建一个新建的字串。 

创建一个名为 `replacedString` 的新变量，并赋值为 `originalString.replace("cats", "dogs")`。 

# --hints--

你应该有一个名为 `replacedString` 的变量。

```js
assert.exists(replacedString);
```

你的 `replacedString` 变量应是一个字串。

```js
assert.isString(replacedString);
```

你应该将 `originalString.replace("cats", "dogs")` 赋值给 `replacedString` 变量。

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
