---
id: 68e43f94365a964f26c88f78
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

现在是时候使用一个新建的字串方法了。 

首先新建一个名为 `exampleSentence` 的变量，并赋值为字串 `"I love cats and cats are so much fun!"`。

然后在该变量下面，向控制台日志 `"Original sentence:"`。然后在该 `console.log()` 下面，添加另一个带有 `exampleSentence` 变量的 `console.log()`。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该有一个 `exampleSentence` 变量。

```js
assert.exists(exampleSentence);
```

你的 `exampleSentence` 变量应是一个字串。

```js
assert.isString(exampleSentence);
```

你应该将字串 `"I love cats and cats are so much fun!"` 赋值给 `exampleSentence` 变量。

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

你应该将 `"Original sentence:"` 记录到控制台日志。

```js
assert.equal(getLogs()[4], "Original sentence:");
```

你应该将 `exampleSentence` 变量记录到控制台日志中。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
