---
id: 68e44abea678a65b877864dc
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

看到那个重复的词出现在句子中会很不错。创建一个名为 `newSentence` 的新变量，并将 `"I "` 字串与 `repeatedLove` 变量串接的结果赋值给它，随后再串接字串 `" learning."`。然后将 `newSentence` 变量记录到控制台。你可以选择使用 `+` 操作符或模板字面值来进行字串串接。 

现在你应该看到文本 `"I love love love  learning."`。 

*注意*：最后的 `"love"` 后面有一个额外的空格，这是正确的。在最后一步，你将修剪那个额外的空格。 

# --hints--

你应该有一个名为 `newSentence` 的变量。

```js
assert.exists(newSentence);
```

你的 `newSentence` 变量应该是一个字串。

```js
assert.isString(newSentence);
```

你应该将字串 `"I "` 与变量 `repeatedLove` 连接，然后再连接字串 `" learning."`。这个结果应赋值给变量 `newSentence`。注意空格。 

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

你不应该将字面字串 `"I love love love  learning."` 赋值给 `newSentence` 变量。请重新阅读说明以获得正确的方法。

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
