---
id: 68e44519c769bf545bb3f22a
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

要获得期望的输出 `"I love dogs and dogs are so much fun!"`，你需要使用 `replaceAll()` 方法而不是 `replace()` 方法。此方法为所有匹配子串返回一个新建字串。 

将 `exampleSentence.replace("cats", "dogs")` 更新为使用 `replaceAll()` 方法。现在你应该能在控制台看到正确的文本。

# --hints--

你应该更新 `exampleSentence.replace("cats", "dogs")` 这行代码以使用 `replaceAll()` 方法。

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
