---
id: 68e4465fba0bd256868b7deb
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

现在是时候使用一个新建的字串方法了。

首先创建一个名为 `learningSentence` 的变量，并赋值字串 `"I love learning!"`。 

然后在该变量下面，向控制台日志 `"Original learning sentence:"`。在该 `console.log()` 下面，为 `learningSentence` 变量添加另一个 `console.log()`。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该有一个 `learningSentence` 变量。

```js
assert.exists(learningSentence);
```

你的 `learningSentence` 变量应是一个字串。

```js
assert.isString(learningSentence);
```

你应该将字串 `"I love learning!"` 赋值给 `learningSentence` 变量。

```js
assert.equal(learningSentence, "I love learning!");
```

你应该将 `"Original learning sentence:"` 记录到控制台日志。

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

你应该将 `learningSentence` 变量记录到控制台日志。

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
