---
id: 68551b677a5f6898af364b85
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafundishwa zaidi kuhusu Tailwind CSS kwa kutumia madarasa yake ya utiliti ya flexbox kujenga sehemu inayojibadilisha kulingana na kifaa ya wito wa kuchukua hatua (CTA).

Muundo wa msingi wa HTML umewekwa tayari kwako, hivyo unahitaji kwanza kujenga vipengele vya HTML.

Anza kwa kuunda kipengele cha `div`. Ndani ya `div` hiyo, tengeneza vipengele viwili tofauti vya `div`.

# --hints--

Unapaswa kuunda kipengele cha `div`.

```js
assert.exists(document.querySelector("div"));
```

Unapaswa kuwa na vipengele viwili zaidi vya `div` ndani ya kipengele cha `div`.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
