---
id: 68552464ef378e1051a55c28
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Ndani ya `div` ya pili, tengeneza vipengele viwili vya nanga. Kwanza iwe na maandishi `Learn more`, na ya pili iwe na maandishi `Start listening`.

Kwa sifa ya `href` ya vipengele viwili vya nanga, weka zote kuwa `#`.

# --hints--

Unapaswa kutengeneza vipengele viwili vya `a` ndani ya kipengele cha `div` cha pili

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

Kipengele chako cha kwanza cha `a` kinapaswa kuwa na maandishi `Learn more` na sifa ya `href` ya `#` ndani ya kipengele cha `div` cha pili.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

Kipengele chako cha pili cha `a` kinapaswa kuwa na maandishi `Start listening` na sifa ya `href` ya `#` ndani ya kipengele cha `div` cha pili.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
