---
id: 685523c05fb41d0848c4b52f
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Ndani ya kipengele cha kwanza cha `div`, tengeneza kipengele cha `span` chenye maandishi ya `Soundflow`, kipengele cha `h1` chenye maandishi ya `Discover New Music`, na kipengele cha `p` chenye maandishi ya `Stream your favorite tracks and discover new artists.`

# --hints--

Unapaswa kutengeneza kipengele cha `span` chenye maandishi ya `Soundflow` ndani ya kipengele cha kwanza cha `div`.

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

Unapaswa kutengeneza kipengele cha `h1` chenye maandishi ya `Discover New Music` ndani ya kipengele cha kwanza cha `div`.

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

Unapaswa kutengeneza kipengele cha `p` chenye maandishi ya `Stream your favorite tracks and discover new artists.` ndani ya kipengele cha kwanza cha `div`.

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
