---
id: 6858eadd3144c89d45e02bc7
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Ikiwa utaona, `div` ya nje imegandamana sana juu na yaliyomo ndani yake yamekusanyika sana kushoto. Kwa hivyo, unapaswa sasa kufanya kazi kwenye nafasi kidogo.

Ongeza darasa la vitu `mt-8` kwenye `div` ili kuisukuma chini kidogo na `p-4` ili kusogeza maandishi mbali na kushoto.

# --hints--

Kipengele chako cha nje cha `div` kinapaswa kuwa na darasa la vitu la huduma `mt-8`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mt-8"));
```

Kipengele chako cha nje cha `div` kinapaswa kuwa na darasa la vitu la huduma `p-4`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
