---
id: 6858f3a67aa8580ff37016c7
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Ikiwa utarekebisha upana wa kichupo cha onyesho kuwa sawa na upana wa kifaa cha mkononi tena, utaona maudhui daima yakiwa pembeni kwa pembeni. Hii siyo unachotaka kwenye vifaa vya mkononi.

Ili kurekebisha hilo, ongeza darasa la vitu `flex-col` kwenye `div`. Hii itabadilisha mwelekeo wa flex kuwa safu, hivyo vitu vitawekwa juu ya kila mmoja.

Hapa ni jinsi ya kutumia darasa la huduma la `flex-col`:

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

Kipengele chako cha ufunguzi cha `div` kinapaswa kuwa na darasa la huduma la `flex-col`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
