---
id: 6858ebf2552046ab3de9f17c
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Kwa sasa, utaona kuna nafasi nyingi sana upande wa kulia wa maudhui katika `div`. Unaweza kupunguza hii kwa kuongeza upana.

Ongeza darasa la vitu `md:w-1/2` kwenye `div` ili kupunguza upana hadi 50% kwenye skrini za wastani na kubwa zaidi. Hiyo ni, skrini zilizo na pikseli 768 na zaidi.

Ili kuona athari hiyo, panua kichupo cha onyesho awali.

# --hints--

Kipengele chako cha nje cha `div` kinapaswa kuwa na darasa la huduma la `md:w-1/2`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
--fcc-editable-region--
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
