---
id: 6858ee73ed8106ccb214942a
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Sasa kila kitu kimekwama upande wa kushoto, unahitaji kuweka maudhui yote katikati kwa kuongeza ukingo wa auto kwenye mhimili wa x.

Ongeza darasa la Tailwind linalolingana na hilo kwenye kipengele cha nje cha `div`. Ili kuona athari hiyo, rekebisha kichupo cha onyesho hadi maudhui yote yawe katikati.

# --hints--

Kipengele chako cha nje cha `div` kinapaswa kuwa na darasa la huduma la `mx-auto`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--    
  </body>
</html>
```
