---
id: 6858ef974c8ad5dcd7abb9f1
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Yaliyomo ndani ya `div` yanaweza kupangwa na kupangwa kwa nafasi bora zaidi. Kwa nafasi na upangaji wa mabaki, unahitaji kutumia flexbox.

Ongeza darasa la vitu `flex` kwa `div` ili vipengele viwili vya mtoto `div` vikae pembeni kwa pembeni.

Pia, ongeza `justify-around` ili nafasi inayozunguka na kati ya yaliyomo isambazwe kwa usawa, na `items-center` ili yaliyomo yote yawe katikati kabisa ndani ya `div`.

# --hints--

Kipengele chako cha nje cha `div` kinapaswa kuwa na darasa la huduma la `flex`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

Kipengele chako cha nje cha `div` kinapaswa kuwa na darasa la huduma la `justify-around`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

Kipengele chako cha nje cha `div` kinapaswa kuwa na darasa la huduma la `items-center`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
