---
id: 6858e6408a0999603b2a4fee
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Kwa `h1`, ongeza madarasa la vitu `font-bold` na `text-4xl`. Hii itafanya maandishi kuwa herufi nzito, makubwa zaidi, na pia itatoa nafasi zaidi kati ya mistari.

Pia, ongeza darasa la vitu `my-4` kuongeza ukingo wa juu na chini wa `h1` kwa pikseli 16, au `1rem`.

# --hints--

Kipengele chako cha `h1` kinapaswa kuwa na darasa la huduma `font-bold`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("font-bold"));
```

Kipengele chako cha `h1` kinapaswa kuwa na darasa la huduma `text-4xl`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("text-4xl"));
```

Kipengele chako cha `h1` kinapaswa kuwa na darasa la huduma `my-4`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("my-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
         --fcc-editable-region--
        <h1>Discover New Music</h1>
         --fcc-editable-region--
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
