---
id: 68597114730d7c6b8838d204
title: Hatua 15
challengeType: 0
dashedName: step-15
---

# --description--

Aya chini ya `h1` inaonekana nzuri, kwa hivyo haihitaji mtindo zaidi. Hii inamaanisha unaweza kuendelea na vipengele vya nanga ndani ya kipengele cha `div` cha pili.

Toa darasa la vitu `bg-white`, `hover:bg-gray-100`, na `text-indigo-600` kwa kipengele cha nanga cha kwanza. Hii itampa msingi mweupe, rangi ya maandishi ya indigo, na msingi kidogo mwepesi wakati unabofya juu yake.

Kwa kipengele cha nanga cha pili, mpa darasa la vitu `bg-pink-500`, `hover:bg-pink-600`, na `text-white`.

# --hints--

Kipengele chako cha nanga cha kwanza kinapaswa kuwa na darasa la huduma la `bg-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

Kipengele chako cha nanga cha kwanza kinapaswa kuwa na darasa la huduma la `hover:bg-gray-100`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

Kipengele chako cha nanga cha kwanza kinapaswa kuwa na darasa la huduma la `text-indigo-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

Kipengele chako cha nanga cha pili kinapaswa kuwa na darasa la huduma la `bg-pink-500`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

Kipengele chako cha nanga cha pili kinapaswa kuwa na darasa la huduma la `hover:bg-pink-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

Kipengele chako cha nanga cha pili kinapaswa kuwa na darasa la huduma la `text-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
        --fcc-editable-region--
          Start listening
        </a>
      </div>
    </div>
  </body>
</html>
```
