---
id: 695774002591bbc5f8cf3e53
title: Construye un Calculador de Descuentos
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

En este laboratorio escribirás una función que calcula el precio final de un artículo después de aplicar un descuento porcentual.

Por ejemplo, si el precio de un artículo es `50` y se aplica un descuento de `20`, el monto del descuento es `10` y el precio final es `40`.

**Objetivo:** Cumple las historias de usuario a continuación para completar el laboratorio.

**Historias de usuario:**

1. Debes definir una función llamada `apply_discount`.

2. La función `apply_discount` debe tomar exactamente dos parámetros: `price` y `discount`.

3. Si `price` no es un número (`int` o `float`), la función debe devolver la cadena `The price should be a number`.

4. Si `discount` no es un número (`int` o `float`), la función debe devolver la cadena `The discount should be a number`.

5. Si `price` es menor o igual a `0`, la función debería devolver la cadena `The price should be greater than 0`.

6. Si `discount` es menor que `0` o mayor que `100`, la función debe devolver la cadena `The discount should be between 0 and 100`.

7. Si ambas entradas son válidas, la función debe calcular el descuento como un porcentaje del precio.

8. La función debe devolver el precio final después de aplicar el descuento.

# --hints--

Debes tener una función llamada `apply_discount`.

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

Tu función `apply_discount` debe tomar dos parámetros: `price` y `discount`.

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

Cuando se llama a `apply_discount` con un `price` (primer argumento) que no es un número (`int` o `float`), debería devolver `The price should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

Cuando se llama a `apply_discount` con un `discount` (segundo argumento) que no es un número (`int` o `float`), debe devolver `The discount should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

Cuando se llama a `apply_discount` con un `price` menor o igual a `0`, debería devolver `The price should be greater than 0`.

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

Cuando se llama a `apply_discount` con un `discount` menor que `0` o mayor que `100`, debe devolver `The discount should be between 0 and 100`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` debería devolver `80`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` debería devolver `100`.

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` debería devolver `50`.

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

Cuando se llama a `apply_discount` con un descuento de `100`, debería devolver `0`.

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` debería devolver `59.6`.

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
