---
id: 69708d59833ceded7d00541f
title: implementare operazioni su una lista concatenata
challengeType: 26
dashedName: implement-linked-list-operations
---

# --description--

In questo laboratorio, implementerai operazioni aggiuntive per una struttura dati lista concatenata, basandoti sulla funzionalità base di aggiungere e rimuovere nodi.

**Obiettivo**: Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente**

1. Dovresti avere una funzione `contains` che accetta una lista concatenata e un elemento. Deve restituire `true` se l'elemento specificato esiste nella lista concatenata, e `false` altrimenti.
1. Dovresti avere una funzione `getAt` che accetta una lista concatenata e un indice. Deve restituire l'elemento all'indice dato nella lista concatenata. Se l'indice è fuori dai limiti, deve restituire `undefined`.
1. Dovresti avere una funzione `insertAt` che accetta una lista concatenata, un indice e un elemento. Deve inserire l'elemento dato nella posizione specificata nella lista concatenata. Se l'indice è fuori dai limiti, non deve modificare la lista.
1. Dovresti avere una funzione `removeAt` che accetta una lista concatenata e un indice. Deve rimuovere il nodo all'indice dato nella lista concatenata. Se l'indice è fuori dai limiti, non deve modificare la lista.
1. Dovresti avere una funzione `clear` che accetta una lista concatenata. Deve rimuovere tutti gli elementi dalla lista concatenata, resettandola effettivamente a uno stato vuoto.

**Nota:** Alcuni test successivi dipendono da metodi precedenti. Per esempio, se `getAt` non è implementata correttamente, i test per funzioni come `insertAt` e `removeAt` potrebbero fallire anche se quelle funzioni sono quasi corrette.

# --hints--

`initList` dovrebbe restituire un oggetto con `head` impostato su `null` e `length` impostato su `0`.

```js
const list = initList();
assert.isNull(list.head);
assert.strictEqual(list.length, 0);
```

`add` dovrebbe aumentare la lunghezza della lista di uno.

```js
const list = initList();
add(list, 'apple');
assert.strictEqual(list.length, 1);
```

`remove` dovrebbe diminuire la lunghezza della lista di uno.

```js
const list = initList();  
add(list, 'apple');  
add(list, 'banana');  
remove(list, 'apple');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(list.head.element, 'banana');  
assert.isNull(list.head.next);
```

Dovresti avere una funzione `contains`.

```js
assert.isFunction(contains);
```

`contains` dovrebbe restituire `true` se l'elemento esiste nella lista.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

`contains` dovrebbe restituire `false` se l'elemento non esiste nella lista.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

`contains` dovrebbe restituire `false` per una lista vuota.

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

Dovresti avere una funzione `getAt`.

```js
assert.isFunction(getAt);
```

`getAt` dovrebbe restituire l'elemento a un dato indice.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

`getAt` dovrebbe restituire `undefined` per indici non validi.

```js
const list = initList();  
add(list, 'apple');  
assert.isUndefined(getAt(list, -1));  
assert.isUndefined(getAt(list, 1));  
assert.isUndefined(getAt(list, 5));
```

`getAt` dovrebbe restituire `undefined` per una lista vuota.

```js
const list = initList();  
assert.isUndefined(getAt(list, 0));
```

Dovresti avere una funzione `insertAt`.

```js
assert.isFunction(insertAt);
```

`insertAt` dovrebbe inserire un elemento all'inizio della lista.

```js
const list = initList();  
add(list, 'second');  
add(list, 'third');  
insertAt(list, 0, 'first');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` dovrebbe inserire un elemento nel mezzo della lista.

```js
const list = initList();  
add(list, 'first');  
add(list, 'third');  
insertAt(list, 1, 'second');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` dovrebbe inserire un elemento alla fine della lista.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` dovrebbe inserire in una lista vuota all'indice `0`.

```js  
const list = initList();  
insertAt(list, 0, 'first');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.isNull(list.head.next); 
```

`insertAt` non dovrebbe modificare la lista per indici non validi.

```js
const list = initList();  
add(list, 'apple');  
insertAt(list, -1, 'invalid');  
insertAt(list, 10, 'invalid');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
assert.isUndefined(getAt(list, 1));
```

Dovresti avere una funzione `removeAt`.

```js
assert.isFunction(removeAt);
```

`removeAt` dovrebbe rimuovere l'elemento a un indice specificato.

```js
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'third');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` dovrebbe rimuovere correttamente il primo elemento.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
assert.isNull(list.head.next);
```

`removeAt` dovrebbe rimuovere correttamente l'ultimo elemento.

```js  
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 2);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` dovrebbe resettare la lista quando rimuove l'unico elemento.

```js
const list = initList();  
add(list, 'only');  
removeAt(list, 0);  
assert.strictEqual(list.length, 0);  
assert.isNull(list.head);  
assert.isTrue(isEmpty(list));
```

`removeAt` non dovrebbe modificare la lista per indici non validi.

```js
const list = initList();  
add(list, 'apple');  
removeAt(list, -1);  
removeAt(list, 10);  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
```

Dovresti avere una funzione `clear`.

```js
assert.isFunction(clear);
```

`clear` dovrebbe rimuovere tutti gli elementi dalla lista.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

`clear` dovrebbe permettere di riutilizzare la lista.

```js  
const list = initList();  
add(list, 'first');  
clear(list);  
clear(list);  
add(list, 'second');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'second');  
assert.isNull(list.head.next); 
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  
}

function getAt(list, index) {
 
}

function insertAt(list, index, element) {
  
}

function removeAt(list, index) {
 
}

function clear(list) {

}
```

# --solutions--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  let current = list.head;
  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }
  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;
  let current = list.head;
  let i = 0;
  while (i < index) {
    current = current.next;
    i++;
  }
  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;
  const node = { element, next: null };
  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;
  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
