---
id: 69860157f7fce39f9df399a3
title: Implementa el Algoritmo de las N-Reinas
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

El problema de las N-Reinas te pide colocar N reinas en un tablero de ajedrez de N×N de manera que ninguna de las dos reinas se ataque entre sí (ninguna comparta fila, columna o diagonal). 

Por ejemplo, si hay un tablero 4x4, una estructura válida es:

```md
[1, 3, 0, 2]
```

Eso significa que en la fila 0, la reina está colocada en la columna 1; en la fila 1, la reina está colocada en la columna 3; en la fila 2, la reina está colocada en la columna 0; y en la fila 3, la reina está colocada en la columna 2.

Visualmente, esta estructura se ve así:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Donde `Q` representa una reina y `.` representa una casilla vacía.

En este laboratorio, implementarás el solucionador del problema de las N-Reinas usando el enfoque de búsqueda en profundidad.


**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función llamada `dfsNQueens`.
2. La función debe aceptar exactamente un argumento: un entero `n`.
3. Si `n` es menor que `1`, la función debe devolver un arreglo vacío (`[]`).
4. La función debe devolver una lista de soluciones; cada solución es a su vez una lista de longitud `n`, donde el elemento en el índice `i` es el índice de columna (basado en 0) de la reina en la fila `i`.

# --hints--

Debes tener una función llamada `dfsNQueens` que tome un argumento.

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

Si `n` es menor que `1`, la función debe devolver un arreglo vacío.

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

La función debe devolver un arreglo de soluciones, donde cada solución es un arreglo de longitud `n`.

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)` debería devolver `[[0]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)` debería devolver `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)` debería devolver `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)` debe devolver `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)` debe devolver `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length` debería ser `10`.

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length` debería ser `92`.

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens` debe devolver el resultado correcto.

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
