---
id: 594810f028c0303b75339ad3
title: 벡터 내적
challengeType: 1
forumTopicId: 302343
dashedName: vector-dot-product
---

# --description--

벡터는 순서가 있는 모음으로 하나 이상의 값을 나타낼 수 있습니다. 예를 들어 (x), (x, y), 또는 (x, y, z)가 있습니다.

# --instructions--

두 벡터(1차원 배열로 나타낸)를 입력으로 받아 내적을 계산하는 함수를 작성하세요. 벡터 길이가 다르거나 두 벡터 이외의 값이 입력되면 `null`를 반환해야 합니다.

# --hints--

`dotProduct`는 함수여야 합니다.

```js
assert.equal(typeof dotProduct, 'function');
```

`dotProduct()`는 `null`를 반환해야 합니다.

```js
assert.equal(dotProduct(), null);
```

`dotProduct([1], [1])`는 `1`를 반환해야 합니다.

```js
assert.equal(dotProduct([1], [1]), 1);
```

`dotProduct([1], [1, 2])`는 `null`를 반환해야 합니다.

```js
assert.equal(dotProduct([1], [1, 2]), null);
```

`dotProduct([1, 3, -5], [4, -2, -1])`는 `3`를 반환해야 합니다.

```js
assert.equal(dotProduct([1, 3, -5], [4, -2, -1]), 3);
```

`dotProduct([3, 2, 1], [2, 4, 2], [5, 3, 1])`는 `null`를 반환해야 합니다.

```js
assert.equal(dotProduct([3, 2, 1], [2, 4, 2], [5, 3, 1]), null);
```

`dotProduct([ 0, 3, 6, 9, 12 ], [ 0, 4, 8, 12, 16 ])`는 `360`를 반환해야 합니다.

```js
assert.equal(dotProduct([ 0, 3, 6, 9, 12 ], [ 0, 4, 8, 12, 16 ]), 360);
```

# --seed--

## --seed-contents--

```js
function dotProduct(...vectors) {

}
```

# --solutions--

```js
function dotProduct(...vectors) {
  if (!vectors || !vectors.length || vectors.length > 2 || vectors[0].length !== vectors[1].length) {
    return null;
  }
  const vectorLen = vectors[0].length;

  let prod = 0;
  let sum = 0;
  let j = vectorLen;
  let i = 2;
  // Sum terms
  while (j--) {
    i = 2;
    prod = 1;

    while (i--) {
      prod *= vectors[i][j];
    }
    sum += prod;
  }
  return sum;
}
```
