---
id: 594810f028c0303b75339ad6
title: 젝켄도르프 수 표현
challengeType: 1
forumTopicId: 302346
dashedName: zeckendorf-number-representation
---

# --description--

숫자를 10의 거듭제곱(십진법)이나 2의 거듭제곱(이진법)의 배수 합으로 위치 표기법으로 나타낼 수 있듯이, 모든 양의 정수는 피보나치 수열의 서로 다른 항 각각에 0 또는 1을 곱한 값의 합으로 나타낼 수 있습니다. 처음 여섯 개의 서로 다른 피보나치 수는 `1, 2, 3, 5, 8, 13`입니다.

십진수 11은 위치 표기법에서 `0*13 + 1*8 + 0*5 + 1*3 + 0*2 + 0*1` 또는 `010100`로 쓸 수 있는데, 여기서 각 자리수는 수열의 특정 항과 곱하는 것을 나타냅니다. 선행 0은 생략되어 11 십진수는 `10100`가 됩니다. 그러나 10100만이 피보나치 수로 11을 만드는 유일한 방법은 아니며 `0*13 + 1*8 + 0*5 + 0*3 + 1*2 + 1*1` 또는 010011도 십진수 11을 나타냅니다. 진정한 젝켄도르프 수는 *두 개의 연속된 피보나치 수를 사용할 수 없다*는 추가 제한이 있어 앞의 유일한 해를 만듭니다.

# --instructions--

`n`의 젝켄도르프 수 표현을 생성하여 반환하는 함수를 작성하세요.

# --hints--

`zeckendorf`는 함수여야 합니다.

```js
assert.equal(typeof zeckendorf, 'function');
```

`zeckendorf(0)`는 `0`를 반환해야 합니다.

```js
assert.equal(zeckendorf(0), 0);

```

`zeckendorf(1)`는 `1`를 반환해야 합니다.

```js
assert.equal(zeckendorf(1), 1);
```

`zeckendorf(2)`는 `10`를 반환해야 합니다.

```js
assert.equal(zeckendorf(2), 10);
```

`zeckendorf(3)`는 `100`를 반환해야 합니다.

```js
assert.equal(zeckendorf(3), 100);
```

`zeckendorf(4)`는 `101`를 반환해야 합니다.

```js
assert.equal(zeckendorf(4), 101);
```

`zeckendorf(5)`는 `1000`를 반환해야 합니다.

```js
assert.equal(zeckendorf(5), 1000);
```

`zeckendorf(6)`는 `1001`를 반환해야 합니다.

```js
assert.equal(zeckendorf(6), 1001);
```

`zeckendorf(7)`는 `1010`를 반환해야 합니다.

```js
assert.equal(zeckendorf(7), 1010);
```

`zeckendorf(8)`는 `10000`를 반환해야 합니다.

```js
assert.equal(zeckendorf(8), 10000);
```

`zeckendorf(9)`는 `10001`를 반환해야 합니다.

```js
assert.equal(zeckendorf(9), 10001);
```

`zeckendorf(10)`는 `10010`를 반환해야 합니다.

```js
assert.equal(zeckendorf(10), 10010);
```

`zeckendorf(11)`는 `10100`를 반환해야 합니다.

```js
assert.equal(zeckendorf(11), 10100);
```

`zeckendorf(12)`는 `10101`를 반환해야 합니다.

```js
assert.equal(zeckendorf(12), 10101);
```

`zeckendorf(13)`는 `100000`를 반환해야 합니다.

```js
assert.equal(zeckendorf(13), 100000);
```

`zeckendorf(14)`는 `100001`를 반환해야 합니다.

```js
assert.equal(zeckendorf(14), 100001);
```

`zeckendorf(15)`는 `100010`를 반환해야 합니다.

```js
assert.equal(zeckendorf(15), 100010);
```

`zeckendorf(16)`는 `100100`를 반환해야 합니다.

```js
assert.equal(zeckendorf(16), 100100);
```

`zeckendorf(17)`는 `100101`를 반환해야 합니다.

```js
assert.equal(zeckendorf(17), 100101);
```

`zeckendorf(18)`는 `101000`를 반환해야 합니다.

```js
assert.equal(zeckendorf(18), 101000);
```

`zeckendorf(19)`는 `101001`를 반환해야 합니다.

```js
assert.equal(zeckendorf(19), 101001);
```

`zeckendorf(20)`는 `101010`를 반환해야 합니다.

```js
assert.equal(zeckendorf(20), 101010);
```

# --seed--

## --seed-contents--

```js
function zeckendorf(n) {

}
```

# --solutions--

```js
// zeckendorf :: Int -> Int
function zeckendorf(n) {
  const f = (m, x) => (m < x ? [m, 0] : [m - x, 1]);
  return parseInt((n === 0 ? ([0]) :
    mapAccumL(f, n, reverse(
      tail(fibUntil(n))
    ))[1]).join(''));
}

// fibUntil :: Int -> [Int]
let fibUntil = n => {
  const xs = [];
  until(
      ([a]) => a > n,
      ([a, b]) => (xs.push(a), [b, a + b]), [1, 1]
  );
  return xs;
};

let mapAccumL = (f, acc, xs) => (
  xs.reduce((a, x) => {
    const pair = f(a[0], x);

    return [pair[0], a[1].concat(pair[1])];
  }, [acc, []])
);

let until = (p, f, x) => {
  let v = x;
  while (!p(v)) v = f(v);
  return v;
};

const tail = xs => (
   xs.length ? xs.slice(1) : undefined
);

const reverse = xs => xs.slice(0).reverse();
```
