---
id: 594810f028c0303b75339ad4
title: 단어 줄 바꿈
challengeType: 1
forumTopicId: 302344
dashedName: word-wrap
---

# --description--

오늘날에도 비례 폰트와 복잡한 레이아웃이 있지만, 지정된 열에서 텍스트를 줄 바꿈해야 하는 경우가 여전히 있습니다. 기본 작업은 텍스트 단락을 간단한 방식으로 줄 바꿈하는 것입니다.

# --instructions--

이 텍스트를 원하는 문자 수에 맞게 줄 바꿈할 수 있는 함수를 작성하세요. 입력 텍스트에는 이미 줄 바꿈이 포함되어 있으니, 함수가 이를 적절히 처리해야 합니다. 예를 들어, 80자에 맞게 줄 바꿈한 텍스트는 다음과 같아야 합니다:

<pre>
Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX
algorithm. If your language provides this, you get easy extra credit, but you
must reference documentation indicating that the algorithm is something better
than a simple minimum length algorithm.
</pre>

# --before-each--

```js
const text =
`Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX algorithm.
If your language provides this, you get easy extra credit,
but you ''must reference documentation'' indicating that the algorithm
is something better than a simple minimum length algorithm.`;

const wrapped80 = wrap(text, 80);
const wrapped42 = wrap(text, 42);

const firstRow80 =
    'Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX';

const firstRow42 = 'Wrap text using a more sophisticated';
```

# --hints--

`wrap`는 함수여야 합니다.

```js
assert.equal(typeof wrap, 'function');
```

`wrap`는 문자열을 반환해야 합니다.

```js
assert.equal(typeof wrap('abc', 10), 'string');
```

`wrap(text,80)`는 4줄을 반환해야 합니다.

```js
assert(wrapped80.split('\n').length === 4);
```

`wrap` 함수는 예상한 텍스트를 반환해야 합니다.

```js
assert.equal(wrapped80.split('\n')[0], firstRow80);
```

`wrap(text,42)`는 7줄을 반환해야 합니다.

```js
assert(wrapped42.split('\n').length === 7);
```

`wrap` 함수는 예상한 텍스트를 반환해야 합니다.

```js
assert.equal(wrapped42.split('\n')[0], firstRow42);
```

# --seed--

## --seed-contents--

```js
function wrap(text, limit) {
  return text;
}
```

# --solutions--

```js
function wrap(text, limit) {
  const noNewlines = text.replace('\n', '');
  if (noNewlines.length > limit) {
    // find the last space within limit
    const edge = noNewlines.slice(0, limit).lastIndexOf(' ');
    if (edge > 0) {
      const line = noNewlines.slice(0, edge);
      const remainder = noNewlines.slice(edge + 1);
      return line + '\n' + wrap(remainder, limit);
    }
  }
  return text;
}
```
