---
id: 594810f028c0303b75339ad7
title: 장-수엔 세밀화 알고리즘
challengeType: 1
forumTopicId: 302347
dashedName: zhang-suen-thinning-algorithm
---

# --description--

이 알고리즘은 흑백, 즉 픽셀당 1비트 이미지의 세밀화를 위해 사용됩니다. 예를 들어, 다음과 같은 입력 이미지가 주어지면:

```js
const testImage1 = [
 '                               ',
 '#########       ########       ',
 '###   ####     ####  ####      ',
 '###    ###     ###    ###      ',
 '###   ####     ###             ',
 '#########      ###             ',
 '### ####       ###    ###      ',
 '###  ####  ### ####  #### ###  ',
 '###   #### ###  ########  ###  ',
 '                               '
];
```

다음과 같이 세밀화된 출력이 생성됩니다:

```js
[ '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               ' ];
```

## 알고리즘

검은 픽셀을 1, 흰 픽셀을 0으로 가정하고, 입력 이미지는 1과 0으로 이루어진 N×M 직사각형 배열이라고 가정합니다. 알고리즘은 8개의 이웃을 가질 수 있는 모든 검은 픽셀 P1에 대해 작동합니다. 이웃 픽셀들은 다음과 같이 순서대로 배열됩니다:

$$\begin{array}{|c|c|c|}
  \\hline
  P9 & P2              & P3\\\\ \\hline
  P8 & \boldsymbol{P1} & P4\\\\ \\hline
  P7 & P6              & P5\\\\ \\hline
\end{array}$$

이미지의 경계 픽셀은 8개의 이웃을 모두 가질 수 없습니다.

- $A(P1)$ = 수열 P2, P3, P4, P5, P6, P7, P8, P9, P2에서 흰색에서 검은색으로 전환되는 횟수 ($0 \to 1$)를 정의합니다. (마지막에 추가된 P2는 순환을 의미합니다)
- $B(P1)$ = P1의 검은색 이웃 픽셀 수 ($= \\sum(P2 \ldots P9)$)

**1단계:**

모든 픽셀을 검사하고 다음 조건을 모두 만족하는 픽셀을 이 단계에서 기록합니다.

1. 픽셀이 검은색이고 8개의 이웃을 가짐
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. $P2$, $P4$, $P6$ 중 적어도 하나가 흰색
5. $P4$, $P6$, $P8$ 중 적어도 하나가 흰색

이미지를 반복 검사하여 1단계 조건을 모두 만족하는 픽셀을 모두 흰색으로 바꿉니다.

**2단계:**

다시 모든 픽셀을 검사하고 다음 조건을 모두 만족하는 픽셀을 이 단계에서 기록합니다.

1. 픽셀이 검은색이고 8개의 이웃을 가짐
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. $P2$, $P4$, $P8$ 중 적어도 하나가 흰색
5. $P2$, $P6$, $P8$ 중 적어도 하나가 흰색

이미지를 반복 검사하여 2단계 조건을 모두 만족하는 픽셀을 모두 흰색으로 바꿉니다.

**반복:**

1단계 또는 2단계에서 픽셀이 변경된 경우, 더 이상 변경이 없을 때까지 모든 단계를 반복합니다.

# --instructions--

제공된 `image` 배열(각 문자열이 이미지의 한 줄을 나타냄)에 대해 장-수엔 세밀화를 수행하는 함수를 작성하세요. 문자열에서 `#`은 검은 픽셀을 나타내고, 공백은 흰 픽셀을 나타냅니다. 함수는 동일한 표현 방식을 사용하여 세밀화된 이미지를 반환해야 합니다.

# --before-each--

```js
const _testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
const expected1 = [
  '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               '
];
const _testImage2 = [
  '                                                          ',
  ' #################                   #############        ',
  ' ##################               ################        ',
  ' ###################            ##################        ',
  ' ########     #######          ###################        ',
  '   ######     #######         #######       ######        ',
  '   ######     #######        #######                      ',
  '   #################         #######                      ',
  '   ################          #######                      ',
  '   #################         #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######         #######       ######        ',
  ' ########     #######          ###################        ',
  ' ########     ####### ######    ################## ###### ',
  ' ########     ####### ######      ################ ###### ',
  ' ########     ####### ######         ############# ###### ',
  '                                                          '
];
const expected2 = [
  '                                                          ',
  '                                                          ',
  '    # ##########                       #######            ',
  '     ##        #                   ####       #           ',
  '     #          #                 ##                      ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     ############               #                         ',
  '     #          #               #                         ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #                            ##                      ',
  '     #                             ############           ',
  '                       ###                          ###   ',
  '                                                          ',
  '                                                          '
];
```

# --hints--

`thinImage`는 함수여야 합니다.

```js
assert.equal(typeof thinImage, 'function');
```

`thinImage`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(thinImage(_testImage1)));
```

`thinImage`는 문자열 배열을 반환해야 합니다.

```js
assert.equal(typeof thinImage(_testImage1)[0], 'string');
```

`thinImage(testImage1)`는 예시와 같은 세밀화된 이미지를 반환해야 합니다.

```js
assert.deepEqual(thinImage(_testImage1), expected1);
```

`thinImage(testImage2)`는 세밀화된 이미지를 반환해야 합니다.

```js
assert.deepEqual(thinImage(_testImage2), expected2);
```

# --seed--

## --seed-contents--

```js
function thinImage(image) {

}

const testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
```

# --solutions--

```js
function Point(x, y) {
  this.x = x;
  this.y = y;
}

const ZhangSuen = (function () {
  function ZhangSuen() {
  }

  ZhangSuen.nbrs = [[0, -1], [1, -1], [1, 0], [1, 1], [0, 1], [-1, 1], [-1, 0], [-1, -1], [0, -1]];

  ZhangSuen.nbrGroups = [[[0, 2, 4], [2, 4, 6]], [[0, 2, 6], [0, 4, 6]]];

  ZhangSuen.toWhite = [];

  ZhangSuen.main = function (image) {
    ZhangSuen.grid = new Array(image);
    for (let r = 0; r < image.length; r++) {
      ZhangSuen.grid[r] = image[r].split('');
    }
    ZhangSuen.thinImage();
    return ZhangSuen.getResult();
  };

  ZhangSuen.thinImage = function () {
    let firstStep = false;
    let hasChanged;
    do {
      hasChanged = false;
      firstStep = !firstStep;
      for (let r = 1; r < ZhangSuen.grid.length - 1; r++) {
        for (let c = 1; c < ZhangSuen.grid[0].length - 1; c++) {
          if (ZhangSuen.grid[r][c] !== '#') {
            continue;
          }
          const nn = ZhangSuen.numNeighbors(r, c);
          if (nn < 2 || nn > 6) {
            continue;
          }
          if (ZhangSuen.numTransitions(r, c) !== 1) {
            continue;
          }
          if (!ZhangSuen.atLeastOneIsWhite(r, c, firstStep ? 0 : 1)) {
            continue;
          }
          ZhangSuen.toWhite.push(new Point(c, r));
          hasChanged = true;
        }
      }
      for (let i = 0; i < ZhangSuen.toWhite.length; i++) {
        const p = ZhangSuen.toWhite[i];
        ZhangSuen.grid[p.y][p.x] = ' ';
      }
      ZhangSuen.toWhite = [];
    } while ((firstStep || hasChanged));
  };

  ZhangSuen.numNeighbors = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === '#') {
        count++;
      }
    }
    return count;
  };

  ZhangSuen.numTransitions = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === ' ') {
        if (ZhangSuen.grid[r + ZhangSuen.nbrs[i + 1][1]][c + ZhangSuen.nbrs[i + 1][0]] === '#') {
          count++;
        }
      }
    }
    return count;
  };

  ZhangSuen.atLeastOneIsWhite = function (r, c, step) {
    let count = 0;
    const group = ZhangSuen.nbrGroups[step];
    for (let i = 0; i < 2; i++) {
      for (let j = 0; j < group[i].length; j++) {
        const nbr = ZhangSuen.nbrs[group[i][j]];
        if (ZhangSuen.grid[r + nbr[1]][c + nbr[0]] === ' ') {
          count++;
          break;
        }
      }
    }
    return count > 1;
  };

  ZhangSuen.getResult = function () {
    const result = [];
    for (let i = 0; i < ZhangSuen.grid.length; i++) {
      const row = ZhangSuen.grid[i].join('');
      result.push(row);
    }
    return result;
  };
  return ZhangSuen;
}());

function thinImage(image) {
  return ZhangSuen.main(image);
}
```
