---
id: 594810f028c0303b75339ad1
title: 행복한 수
challengeType: 1
forumTopicId: 302280
dashedName: happy-numbers
---

# --description--

행복한 수는 다음 과정을 통해 정의됩니다:

임의의 양의 정수에서 시작해, 그 수를 각 자리 숫자의 제곱의 합으로 바꾸고 이 과정을 반복합니다. 숫자가 `1`가 되어 그 상태로 유지되거나, `1`를 포함하지 않는 순환 고리에 빠져 무한히 반복될 때까지 계속합니다. 이 과정이 `1`에서 끝나는 수를 행복한 수라고 하며, `1`에서 끝나지 않는 수는 불행한 수입니다.

# --instructions--

숫자가 행복한 수이면 true를, 그렇지 않으면 false를 반환하는 함수를 구현하세요.

# --hints--

`happy`는 함수여야 합니다.

```js
assert(typeof happy === 'function');
```

`happy(1)`는 불리언을 반환해야 합니다.

```js
assert(typeof happy(1) === 'boolean');
```

`happy(1)`는 `true`를 반환해야 합니다.

```js
assert(happy(1));
```

`happy(2)`는 `false`를 반환해야 합니다.

```js
assert(!happy(2));
```

`happy(7)`는 `true`를 반환해야 합니다.

```js
assert(happy(7));
```

`happy(10)`는 `true`를 반환해야 합니다.

```js
assert(happy(10));
```

`happy(13)`는 `true`를 반환해야 합니다.

```js
assert(happy(13));
```

`happy(19)`는 `true`를 반환해야 합니다.

```js
assert(happy(19));
```

`happy(23)`는 `true`를 반환해야 합니다.

```js
assert(happy(23));
```

`happy(28)`는 `true`를 반환해야 합니다.

```js
assert(happy(28));
```

`happy(31)`는 `true`를 반환해야 합니다.

```js
assert(happy(31));
```

`happy(32)`는 `true`를 반환해야 합니다.

```js
assert(happy(32));
```

`happy(33)`는 `false`를 반환해야 합니다.

```js
assert(!happy(33));
```

# --seed--

## --seed-contents--

```js
function happy(number) {

}
```

# --solutions--

```js
function happy (number) {
  let m;
  let digit;
  const cycle = [];

  while (number !== 1 && cycle[number] !== true) {
    cycle[number] = true;
    m = 0;
    while (number > 0) {
      digit = number % 10;
      m += Math.pow(digit, 2);
      number = (number - digit) / 10;
    }
    number = m;
  }
  return (number === 1);
}
```
