---
id: 594810f028c0303b75339acd
title: '풍부수, 부족수, 완전수 분류'
challengeType: 1
forumTopicId: 302221
dashedName: abundant-deficient-and-perfect-number-classifications
---

# --description--

이들은 진약수를 기준으로 양의 정수를 세 가지 분류로 정의합니다.

$P(n)$을 `n`의 진약수 합으로 정의하며, 진약수는 `n` 자신을 제외한 모든 양의 정수 `n`입니다.

만약 `P(n) < n`라면 `n`는 `deficient`로 분류됩니다.

만약 `P(n) === n`라면 `n`는 `perfect`로 분류됩니다.

만약 `P(n) > n`라면 `n`는 `abundant`로 분류됩니다.

**예시**: `6`의 진약수는 `1`, `2`, 그리고 `3`입니다. `1 + 2 + 3 = 6`이므로 `6`는 완전수로 분류됩니다.

# --instructions--

`1`부터 `num`까지(포함) 정수 중 세 가지 분류에 속하는 수가 각각 몇 개인지 계산하는 함수를 구현하세요. 결과는 `[deficient, perfect, abundant]` 형식의 배열로 출력하세요.

# --hints--

`getDPA`는 함수여야 합니다.

```js
assert(typeof getDPA === 'function');
```

`getDPA(5000)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(getDPA(5000)));
```

`getDPA(5000)` 반환 배열의 길이는 `3`여야 합니다.

```js
assert(getDPA(5000).length === 3);
```

`getDPA(5000)`는 `[3758, 3, 1239]`를 반환해야 합니다.

```js
assert.deepEqual(getDPA(5000), [3758, 3, 1239]);
```

`getDPA(10000)`는 `[7508, 4, 2488]`를 반환해야 합니다.

```js
assert.deepEqual(getDPA(10000), [7508, 4, 2488]);
```

`getDPA(20000)`는 `[15043, 4, 4953]`를 반환해야 합니다.

```js
assert.deepEqual(getDPA(20000), [15043, 4, 4953]);
```

# --seed--

## --seed-contents--

```js
function getDPA(num) {

}
```

# --solutions--

```js
function getDPA(num) {
  const dpa = [1, 0, 0];
  for (let n = 2; n <= num; n += 1) {
    let ds = 1;
    const e = Math.sqrt(n);
    for (let d = 2; d < e; d += 1) {
      if (n % d === 0) {
        ds += d + (n / d);
      }
    }
    if (n % e === 0) {
      ds += e;
    }
    dpa[ds < n ? 0 : ds === n ? 1 : 2] += 1;
  }
  return dpa;
}
```
