---
id: 594810f028c0303b75339ad8
title: 지그재그 행렬
challengeType: 1
forumTopicId: 302348
dashedName: zig-zag-matrix
---

# --description--

'지그재그' 배열은 처음 $N^2$ 정수를 정사각형 배열로 배치한 것으로, 배열의 반대 대각선을 따라 지그재그로 이동하면서 숫자가 순차적으로 증가합니다.

예를 들어, 입력이 `5`일 때 다음과 같은 결과가 나와야 합니다:

<pre>
 0  1  5  6 14
 2  4  7 13 15
 3  8 12 16 21
 9 11 17 20 22
10 18 19 23 24
</pre>

# --instructions--

지그재그 행렬의 크기를 받아 해당 행렬을 2차원 배열로 반환하는 함수를 작성하세요.

# --before-each--

```js
const zm1 = [[0]];
const zm2 = [[0, 1], [2, 3]];
const zm5 = [
  [0, 1, 5, 6, 14],
  [2, 4, 7, 13, 15],
  [3, 8, 12, 16, 21],
  [9, 11, 17, 20, 22],
  [10, 18, 19, 23, 24]
];
```

# --hints--

ZigZagMatrix는 함수여야 합니다.

```js
assert.equal(typeof ZigZagMatrix, 'function');
```

ZigZagMatrix는 배열을 반환해야 합니다.

```js
assert.equal(typeof ZigZagMatrix(1), 'object');
```

ZigZagMatrix는 중첩된 배열의 배열을 반환해야 합니다.

```js
assert.equal(typeof ZigZagMatrix(1)[0], 'object');
```

ZigZagMatrix(1)은 \[[0]]을 반환해야 합니다.

```js
assert.deepEqual(ZigZagMatrix(1), zm1);
```

ZigZagMatrix(2)는 \[[0, 1], [2, 3]]을 반환해야 합니다.

```js
assert.deepEqual(ZigZagMatrix(2), zm2);
```

ZigZagMatrix(5)는 지정된 행렬을 반환해야 합니다.

```js
assert.deepEqual(ZigZagMatrix(5), zm5);
```

# --seed--

## --seed-contents--

```js
function ZigZagMatrix(n) {

  return [[], []];
}
```

# --solutions--

```js
function ZigZagMatrix(n) {
  const mtx = [];
  for (let i = 0; i < n; i++) {
    mtx[i] = [];
  }

  let i = 1;
  let j = 1;
  for (let e = 0; e < n * n; e++) {
    mtx[i - 1][j - 1] = e;
    if ((i + j) % 2 === 0) {
      // Even stripes
      if (j < n) j++;
      else i += 2;
      if (i > 1) i--;
    } else {
      // Odd stripes
      if (i < n) i++;
      else j += 2;
      if (j > 1) j--;
    }
  }
  return mtx;
}
```
