---
id: 594810f028c0303b75339ad2
title: 벡터 외적
challengeType: 1
forumTopicId: 302342
dashedName: vector-cross-product
---

# --description--

벡터는 세 차원을 가지며 세 개의 숫자 (X, Y, Z)로 이루어진 순서 있는 모음으로 나타냅니다.

# --instructions--

두 벡터(배열)를 입력으로 받아 외적을 계산하는 함수를 작성하세요. 벡터 길이가 다를 경우와 같은 잘못된 입력에는 `null`를 반환해야 합니다.

# --hints--

`crossProduct`는 함수여야 합니다.

```js
assert.equal(typeof crossProduct, 'function');
```

`crossProduct()`는 null을 반환해야 합니다.

```js
assert.equal(crossProduct(), null);
```

`crossProduct([1, 2, 3], [4, 5, 6])`는 `[-3, 6, -3]`를 반환해야 합니다.

```js
assert.deepEqual(res12, exp12);
```

# --before-each--

```js
const tv1 = [1, 2, 3];
const tv2 = [4, 5, 6];
const res12 = crossProduct(tv1, tv2);
const exp12 = [-3, 6, -3];
```

# --seed--

## --seed-contents--

```js
function crossProduct(a, b) {

}
```

# --solutions--

```js
function crossProduct(a, b) {
  if (!a || !b) {
    return null;
  }

  // Check lengths
  if (a.length !== 3 || b.length !== 3) {
    return null;
  }

  return [
    (a[1] * b[2]) - (a[2] * b[1]),
    (a[2] * b[0]) - (a[0] * b[2]),
    (a[0] * b[1]) - (a[1] * b[0])
  ];
}
```
