---
id: 594810f028c0303b75339ad0
title: 열 정렬
challengeType: 1
forumTopicId: 302224
dashedName: align-columns
---

# --description--

한 줄 안의 필드들이 단일 `$` 문자로 구분된 여러 줄의 배열이 주어졌을 때, 각 열의 단어들이 최소 한 칸 이상의 공백으로 구분되도록 하여 각 열을 정렬하는 프로그램을 작성하세요. 또한 각 열의 단어들이 해당 열 안에서 왼쪽 정렬, 오른쪽 정렬, 또는 가운데 정렬되도록 하세요.

# --instructions--

다음 텍스트를 사용해 프로그램을 테스트하세요:

```js
const testText = [
  'Given$a$text$file$of$many$lines',
  'where$fields$within$a$line$',
  'are$delineated$by$a$single$"dollar"$character',
  'write$a$program',
  'that$aligns$each$column$of$fields',
  'by$ensuring$that$words$in$each$',
  'column$are$separated$by$at$least$one$space.',
  'Further,$allow$for$each$word$in$a$column$to$be$either$left$',
  'justified,$right$justified',
  'or$center$justified$within$its$column.'
];
```

**다음 사항을 참고하세요:**

- 예시 입력 텍스트의 줄 끝에 달러 기호가 있을 수도 있고 없을 수도 있습니다.
- 모든 열은 동일한 정렬 방식을 공유해야 합니다.
- 줄 끝에 인접해 생성된 연속된 공백 문자는 이 과제 목적상 중요하지 않습니다.
- 출력 텍스트는 고정폭 폰트가 적용된 일반 텍스트 편집기나 기본 터미널에서 볼 것입니다. 줄은 새 줄 문자(`\n`)로 연결해야 합니다.
- 열 사이의 최솟값 공백은 텍스트에서 계산해야 하며 하드코딩하면 안 됩니다.
- 열 사이 또는 주변에 구분 문자를 추가하는 것은 요구사항이 아닙니다.

예를 들어, `testText`의 한 줄을 오른쪽, 왼쪽, 가운데 정렬한 결과는 각각 다음과 같습니다:

```js
'    column        are separated     by     at    least       one space.\n'
'column     are        separated by     at     least    one       space.\n'
'  column      are     separated   by     at    least      one    space.\n'
```

# --before-each--

```js
const _testText = [
  'Given$a$text$file$of$many$lines',
  'where$fields$within$a$line$',
  'are$delineated$by$a$single$"dollar"$character',
  'write$a$program',
  'that$aligns$each$column$of$fields$',
  'by$ensuring$that$words$in$each$',
  'column$are$separated$by$at$least$one$space.',
  'Further,$allow$for$each$word$in$a$column$to$be$either$left$',
  'justified,$right$justified',
  'or$center$justified$within$its$column.'
];

const rightAligned = '     Given          a      text   file     of     many     lines\n' +
'     where     fields    within      a   line \n' +
'       are delineated        by      a single "dollar" character\n' +
'     write          a   program\n' +
'      that     aligns      each column     of   fields \n' +
'        by   ensuring      that  words     in     each \n' +
'    column        are separated     by     at    least       one space.\n' +
'  Further,      allow       for   each   word       in         a column to be either left \n' +
'justified,      right justified\n' +
'        or     center justified within    its  column.';

const leftAligned = 'Given      a          text      file   of     many     lines    \n' +
'where      fields     within    a      line   \n' +
'are        delineated by        a      single "dollar" character\n' +
'write      a          program  \n' +
'that       aligns     each      column of     fields   \n' +
'by         ensuring   that      words  in     each     \n' +
'column     are        separated by     at     least    one       space.\n' +
'Further,   allow      for       each   word   in       a         column to be either left \n' +
'justified, right      justified\n' +
'or         center     justified within its    column. ';

const centerAligned = '  Given        a        text     file    of     many     lines  \n' +
'  where      fields    within     a     line  \n' +
'   are     delineated    by       a    single \"dollar\" character\n' +
'  write        a       program \n' +
'   that      aligns     each    column   of    fields  \n' +
'    by      ensuring    that    words    in     each   \n' +
'  column      are     separated   by     at    least      one    space.\n' +
' Further,    allow       for     each   word     in        a     column to be either left \n' +
'justified,   right    justified\n' +
'    or       center   justified within  its   column. ';
```

# --hints--

`formatText`는 함수여야 합니다.

```js
assert(typeof formatText === 'function');
```

`formatText(testText, 'right')`는 오른쪽 정렬된 열을 가진 텍스트를 생성해야 합니다.

```js
assert.strictEqual(formatText(_testText, 'right'), rightAligned);
```

`formatText(testText, 'left')`는 왼쪽 정렬된 열을 가진 텍스트를 생성해야 합니다.

```js
assert.strictEqual(formatText(_testText, 'left'), leftAligned);
```

`formatText(testText, 'center')`는 가운데 정렬된 열을 가진 텍스트를 생성해야 합니다.

```js
assert.strictEqual(formatText(_testText, 'center'), centerAligned);
```

# --seed--

## --seed-contents--

```js
function formatText(input, justification) {

}

const testText = [
  'Given$a$text$file$of$many$lines',
  'where$fields$within$a$line$',
  'are$delineated$by$a$single$"dollar"$character',
  'write$a$program',
  'that$aligns$each$column$of$fields$',
  'by$ensuring$that$words$in$each$',
  'column$are$separated$by$at$least$one$space.',
  'Further,$allow$for$each$word$in$a$column$to$be$either$left$',
  'justified,$right$justified',
  'or$center$justified$within$its$column.'
];
```

# --solutions--

```js
String.prototype.repeat = function (n) { return new Array(1 + parseInt(n)).join(this); };

function formatText(input, justification) {
  let x, y, max, cols = 0, diff, left, right;
  for (x = 0; x < input.length; x++) {
    input[x] = input[x].split('$');
    if (input[x].length > cols) {
      cols = input[x].length;
    }
  }
  for (x = 0; x < cols; x++) {
    max = 0;
    for (y = 0; y < input.length; y++) {
      if (input[y][x] && max < input[y][x].length) {
        max = input[y][x].length;
      }
    }
    for (y = 0; y < input.length; y++) {
      if (input[y][x]) {
        diff = (max - input[y][x].length) / 2;
        left = ' '.repeat(Math.floor(diff));
        right = ' '.repeat(Math.ceil(diff));
        if (justification === 'left') {
          right += left; left = '';
        }
        if (justification === 'right') {
          left += right; right = '';
        }
        input[y][x] = left + input[y][x] + right;
      }
    }
  }
  for (x = 0; x < input.length; x++) {
    input[x] = input[x].join(' ');
  }
  input = input.join('\n');
  return input;
}
```
