---
id: 594810f028c0303b75339ad5
title: Y 결합자
challengeType: 1
forumTopicId: 302345
dashedName: y-combinator
---

# --description--

엄격한 <a href="https://www.freecodecamp.org/news/the-principles-of-functional-programming/" target="_blank" rel="noopener noreferrer nofollow">함수형 프로그래밍</a>과 람다 계산법에서 함수(람다 표현식)는 상태를 가지지 않으며 오직 감싸고 있는 함수의 인수만 참조할 수 있습니다. 이는 함수가 변수의 상태와 연관되고 이 변수의 상태를 함수 본문에서 사용하는 일반적인 재귀 함수 정의를 배제합니다.

Y 결합자는 무상태 함수로, 다른 무상태 함수에 적용하면 그 함수의 재귀 버전을 반환합니다. Y 결합자는 고정점 결합자라고 불리는 이러한 함수 클래스 중 가장 단순한 함수입니다.

# --instructions--

무상태 Y 결합자 함수를 정의하고 이를 사용해 팩토리얼을 계산하세요. `factorial(N)` 함수는 이미 제공되어 있습니다.

# --hints--

Y는 함수를 반환해야 합니다.

```js
assert.equal(typeof Y((f) => (n) => n), 'function');
```

factorial(1)은 1을 반환해야 합니다.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(1), 1);
```

factorial(2)는 2를 반환해야 합니다.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(2), 2);
```

factorial(3)은 6을 반환해야 합니다.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(3), 6);
```

factorial(4)는 24를 반환해야 합니다.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(4), 24);
```

factorial(10)은 3628800을 반환해야 합니다.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(10), 3628800);
```

# --seed--

## --seed-contents--

```js
function Y(f) {
  return function() {

  };
}

var factorial = Y(function(f) {
  return function (n) {
    return n > 1 ? n * f(n - 1) : 1;
  };
});
```

# --solutions--

```js
var Y = f => (x => x(x))(y => f(x => y(y)(x)));
```
