---
id: 594810f028c0303b75339acc
title: ABC 문제
challengeType: 1
forumTopicId: 302220
dashedName: abc-problem
---

# --description--

ABC 블록 모음(예: 어린 시절 알파벳 블록)이 주어집니다. 각 블록에는 두 글자가 적혀 있으며 총 20개의 블록이 있습니다. 모든 블록 면을 합치면 완전한 알파벳이 보장됩니다. 블록 샘플 모음:

<pre>(B O)
(X K)
(D Q)
(C P)
(N A)
(G T)
(R E)
(T G)
(Q D)
(F S)
(J W)
(H U)
(V I)
(A N)
(O B)
(E R)
(F S)
(L Y)
(P C)
(Z M)
</pre>

# --instructions--

문자열(단어)을 입력받아 주어진 블록 모음으로 단어를 만들 수 있는지 판단하는 함수를 구현하세요.

다음 규칙을 기억하세요:

<ul>
  <li>블록의 한 글자를 사용하면 그 블록은 다시 사용할 수 없습니다.</li>
  <li>함수는 대소문자를 구분하지 않아야 합니다.</li>
</ul>

# --before-each--

```js
const words = ['bark', 'BooK', 'TReAT', 'COMMON', 'squAD', 'conFUSE'];
```

# --hints--

`canMakeWord`는 함수여야 합니다.

```js
assert(typeof canMakeWord === 'function');
```

`canMakeWord`는 불리언 값을 반환해야 합니다.

```js
assert(typeof canMakeWord('hi') === 'boolean');
```

`canMakeWord("bark")`는 true를 반환해야 합니다.

```js
assert(canMakeWord(words[0]));
```

`canMakeWord("BooK")`는 false를 반환해야 합니다.

```js
assert(!canMakeWord(words[1]));
```

`canMakeWord("TReAT")`는 true를 반환해야 합니다.

```js
assert(canMakeWord(words[2]));
```

`canMakeWord("COMMON")`는 false를 반환해야 합니다.

```js
assert(!canMakeWord(words[3]));
```

`canMakeWord("squAD")`는 true를 반환해야 합니다.

```js
assert(canMakeWord(words[4]));
```

`canMakeWord("conFUSE")`는 true를 반환해야 합니다.

```js
assert(canMakeWord(words[5]));
```

# --seed--

## --seed-contents--

```js
function canMakeWord(word) {

}
```

# --solutions--

```js
function canMakeWord(word) {
  const characters = 'BO XK DQ CP NA GT RE TG QD FS JW HU VI AN OB ER FS LY PC ZM';
  const blocks = characters.split(' ').map(pair => pair.split(''));

  const letters = [...word.toUpperCase()];
  let length = letters.length;
  const copy = new Set(blocks);

  letters.forEach(letter => {
    for (let block of copy) {
      const index = block.indexOf(letter);

      if (index !== -1) {
        length--;
        copy.delete(block);
        break;
      }
    }
  });
  return !length;
}
```
