---
id: 594810f028c0303b75339acb
title: 100개의 문
challengeType: 1
forumTopicId: 302217
dashedName: 100-doors
---

# --description--

처음에는 모두 닫혀 있는 100개의 문이 일렬로 있습니다. 문들을 100번 지나갑니다. 첫 번째에는 모든 문을 방문해서 문을 '토글'하세요(문이 닫혀 있으면 열고, 열려 있으면 닫으세요). 두 번째에는 2번째 문, 4번째 문, 6번째 문 등 2번째마다 방문해서 토글하세요. 세 번째에는 3번째 문, 6번째 문, 9번째 문 등 3번째마다 방문해서 토글하세요. 이런 식으로 100번째 문만 방문할 때까지 계속하세요.

# --instructions--

마지막 통과 후 문들의 상태를 결정하는 함수를 구현하세요. 열린 문 번호만 배열에 포함하여 최종 결과를 배열로 반환하세요.

# --hints--

`getFinalOpenedDoors` 는 함수여야 합니다.

```js
assert(typeof getFinalOpenedDoors === 'function');
```

`getFinalOpenedDoors` 는 배열을 반환해야 합니다.

```js
assert(Array.isArray(getFinalOpenedDoors(100)));
```

`getFinalOpenedDoors` 는 올바른 결과를 만들어야 합니다.

```js
const solution = [1, 4, 9, 16, 25, 36, 49, 64, 81, 100];
assert.deepEqual(getFinalOpenedDoors(100), solution);
```

# --seed--

## --seed-contents--

```js
function getFinalOpenedDoors(numDoors) {

}
```

# --solutions--

```js
function getFinalOpenedDoors(numDoors) {
  // this is the final pattern (always squares).
  // thus, the most efficient solution simply returns an array of squares up to numDoors).
  const finalState = [];
  let i = 1;
  while (Math.pow(i, 2) <= numDoors) {
    finalState.push(Math.pow(i, 2));
    i++;
  }
  return finalState;
}
```
