---
id: 594810f028c0303b75339acf
title: 아커만 함수
challengeType: 1
forumTopicId: 302223
dashedName: ackermann-function
---

# --description--

아커만 함수는 재귀 함수의 고전적인 예로, 특히 원시 재귀 함수가 아니라는 점에서 주목받습니다. 이 함수는 값이 매우 빠르게 증가하며, 호출 트리의 크기도 마찬가지로 빠르게 커집니다.

아커만 함수는 보통 다음과 같이 정의됩니다:

$A(m, n) = \\begin{cases} n+1 & \\mbox{if } m = 0 \\\\ A(m-1, 1) & \\mbox{if } m > 0 \\mbox{ and } n = 0 \\\\ A(m-1, A(m, n-1)) & \\mbox{if } m > 0 \\mbox{ and } n > 0. \\end{cases}$

인수는 절대 음수가 아니며 항상 종료됩니다.

# --instructions--

$A(m, n)$의 값을 반환하는 함수를 작성하세요. 함수가 매우 빠르게 증가하기 때문에 임의 정밀도가 선호되지만 필수는 아닙니다.

# --hints--

`ack`는 함수여야 합니다.

```js
assert(typeof ack === 'function');
```

`ack(0, 0)`는 1을 반환해야 합니다.

```js
assert(ack(0, 0) === 1);
```

`ack(1, 1)`는 3을 반환해야 합니다.

```js
assert(ack(1, 1) === 3);
```

`ack(2, 5)`는 13을 반환해야 합니다.

```js
assert(ack(2, 5) === 13);
```

`ack(3, 3)`는 61을 반환해야 합니다.

```js
assert(ack(3, 3) === 61);
```

# --seed--

## --seed-contents--

```js
function ack(m, n) {

}
```

# --solutions--

```js
function ack(m, n) {
  return m === 0 ? n + 1 : ack(m - 1, n === 0 ? 1 : ack(m, n - 1));
}
```
