---
id: 594810f028c0303b75339ace
title: 누적기 생성 함수
challengeType: 1
forumTopicId: 302222
dashedName: accumulator-factory
---

# --description--

Paul Graham이 제시한 문제는 단일(숫자) 인수를 받아 누적기 역할을 하는 또 다른 함수를 반환하는 함수를 만드는 것입니다. 반환된 누적기 함수도 단일 숫자 인수를 받아 지금까지 그 누적기에 전달된 모든 숫자 값의 합을 반환합니다(누적기가 생성될 때 전달된 초기 값 포함).

# --instructions--

숫자 $n$을 받아 그 숫자와 지금까지 전달된 모든 숫자의 합을 반환하는 누적기 함수를 생성하는 함수를 만드세요.

**규칙:**

전역 변수를 사용하지 마세요.

**힌트:**

클로저는 외부 상태를 저장합니다.

# --hints--

`accumulator`는 함수여야 합니다.

```js
assert(typeof accumulator === 'function');
```

`accumulator(0)`는 함수를 반환해야 합니다.

```js
assert(typeof accumulator(0) === 'function');
```

`accumulator(0)(2)`는 숫자를 반환해야 합니다.

```js
assert(typeof accumulator(0)(2) === 'number');
```

3, -4, 1.5, 5를 차례로 전달하면 5.5를 반환해야 합니다.

```js
assert(testFn && testFn(5) === 5.5);
```

# --before-each--

```js
const testFn = typeof accumulator === 'function' && accumulator(3);
if (testFn) {
  testFn(-4);
  testFn(1.5);
}
```

# --seed--

## --seed-contents--

```js
function accumulator(sum) {

}
```

# --solutions--

```js
function accumulator(sum) {
  return function(n) {
    return sum += n;
  };
}
```
