---
id: 5949b579404977fbaefcd736
title: 9억 개의 신의 이름 정수
challengeType: 1
forumTopicId: 302219
dashedName: 9-billion-names-of-god-the-integer
---

# --description--

이 도전 과제는 Arthur C. Clarke의 단편 소설을 변형한 것입니다.

(도전 과제를 완료할 때 발생할 수 있는 결과를 인지하세요.)

자세히 말하면, "이름"이 의미하는 바는 다음과 같습니다:

<ul>
  <li>정수 1은 이름이 1개, "1"입니다.</li>
  <li>정수 2는 이름이 2개, "1+1"과 "2"입니다.</li>
  <li>정수 3은 이름이 3개, "1+1+1", "2+1", 그리고 "3"입니다.</li>
  <li>정수 4는 이름이 5개, "1+1+1+1", "2+1+1", "2+2", "3+1", "4"입니다.</li>
  <li>정수 5는 이름이 7개, "1+1+1+1+1", "2+1+1+1", "2+2+1", "3+1+1", "3+2", "4+1", "5"입니다.</li>
</ul>

이를 다음과 같은 형태로 시각화할 수 있습니다:

<pre>          1
        1   1
      1   1   1
    1   2   1   1
  1   2   2   1   1
1   3   3   2   1   1
</pre>

여기서 $n$번째 행은 정수 $n$에 해당하며, $m$번째 행의 각 열 $C$는 왼쪽에서 오른쪽으로 $C$로 시작하는 이름의 개수를 나타냅니다.

선택적으로, $n$번째 행의 합 $P(n)$은 정수 분할 함수임을 주목하세요.

# --instructions--

$n$번째 행의 합을 반환하는 함수를 구현하세요.

# --hints--

`numberOfNames`는 함수여야 합니다.

```js
assert(typeof numberOfNames === 'function');
```

`numberOfNames(5)`은 7과 같아야 합니다.

```js
assert.equal(numberOfNames(5), 7);
```

`numberOfNames(12)`는 77과 같아야 합니다.

```js
assert.equal(numberOfNames(12), 77);
```

`numberOfNames(18)`는 385와 같아야 합니다.

```js
assert.equal(numberOfNames(18), 385);
```

`numberOfNames(23)`는 1255와 같아야 합니다.

```js
assert.equal(numberOfNames(23), 1255);
```

`numberOfNames(42)`는 53174와 같아야 합니다.

```js
assert.equal(numberOfNames(42), 53174);
```

`numberOfNames(123)`는 2552338241과 같아야 합니다.

```js
assert.equal(numberOfNames(123), 2552338241);
```

# --seed--

## --seed-contents--

```js
function numberOfNames(num) {

  return true;
}
```

# --solutions--

```js
function numberOfNames(num) {
  const cache = [
    [1]
  ];
  for (let l = cache.length; l < num + 1; l++) {
    let Aa;
    let Mi;
    const r = [0];
    for (let x = 1; x < l + 1; x++) {
      r.push(r[r.length - 1] + (Aa = cache[l - x < 0 ? cache.length - (l - x) : l - x])[(Mi = Math.min(x, l - x)) < 0 ? Aa.length - Mi : Mi]);
    }
    cache.push(r);
  }
  return cache[num][cache[num].length - 1];
}
```
