---
id: 5949b579404977fbaefcd737
title: 우호수 쌍
challengeType: 1
forumTopicId: 302225
dashedName: amicable-pairs
---

# --description--

두 정수 $N$과 $M$이 $N \\neq M$이고 $N$의 진약수 합($\\mathrm{sum}(\\mathrm{propDivs}(N))$)이 $M$이며 동시에 $M$의 진약수 합($\\mathrm{sum}(\\mathrm{propDivs}(M))$)이 $N$일 때, 이 두 수를 우호수 쌍이라고 합니다.

**예시:**

**1184**와 **1210**은 우호수 쌍이며, 진약수는 다음과 같습니다:

<ul>
  <li>1, 2, 4, 8, 16, 32, 37, 74, 148, 296, 592 그리고</li>
  <li>1, 2, 5, 10, 11, 22, 55, 110, 121, 242, 605 각각입니다.</li>
</ul>

첫 번째 값인 **1184**의 약수 합은 **1210**이고, 두 번째 값인 **1210**의 약수 합은 **1184**입니다.

# --instructions--

20,000 미만의 우호수 쌍을 계산하여 여기 표시하세요(총 여덟 쌍이 있습니다).

# --before-each--

```js
const answer300 = [[220, 284]];
const answer3000 = [
  [220, 284],
  [1184, 1210],
  [2620, 2924]
];
const answer20000 = [
  [220, 284],
  [1184, 1210],
  [2620, 2924],
  [5020, 5564],
  [6232, 6368],
  [10744, 10856],
  [12285, 14595],
  [17296, 18416]
];
```

# --hints--

`amicablePairsUpTo`는 함수여야 합니다.

```js
assert(typeof amicablePairsUpTo === 'function');
```

`amicablePairsUpTo(300)`는 `[[220,284]]`를 반환해야 합니다.

```js
assert.deepEqual(amicablePairsUpTo(300), answer300);
```

`amicablePairsUpTo(3000)`는 `[[220,284],[1184,1210],[2620,2924]]`를 반환해야 합니다.

```js
assert.deepEqual(amicablePairsUpTo(3000), answer3000);
```

`amicablePairsUpTo(20000)`는 `[[220,284],[1184,1210],[2620,2924],[5020,5564],[6232,6368],[10744,10856],[12285,14595],[17296,18416]]`를 반환해야 합니다.

```js
assert.deepEqual(amicablePairsUpTo(20000), answer20000);
```

# --seed--

## --seed-contents--

```js
function amicablePairsUpTo(maxNum) {

  return true;
}
```

# --solutions--

```js
// amicablePairsUpTo :: Int -> [(Int, Int)]
function amicablePairsUpTo(maxNum) {
  return range(1, maxNum)
    .map(x => properDivisors(x)
      .reduce((a, b) => a + b, 0))
    .reduce((a, m, i, lst) => {
      const n = i + 1;

      return (m > n) && lst[m - 1] === n ?
        a.concat([
          [n, m]
        ]) : a;
    }, []);
}

// properDivisors :: Int -> [Int]
function properDivisors(n) {
  if (n < 2) return [];

  const rRoot = Math.sqrt(n);
  const intRoot = Math.floor(rRoot);
  const blnPerfectSquare = rRoot === intRoot;
  const lows = range(1, intRoot)
  .filter(x => (n % x) === 0);

  return lows.concat(lows.slice(1)
    .map(x => n / x)
    .reverse()
    .slice(blnPerfectSquare | 0));
}

// Int -> Int -> Maybe Int -> [Int]
function range(m, n, step) {
  const d = (step || 1) * (n >= m ? 1 : -1);

  return Array.from({
    length: Math.floor((n - m) / d) + 1
  }, (_, i) => m + (i * d));
}
```
