---
id: 594d8d0ab97724821379b1e6
title: 평균/최빈값
challengeType: 1
forumTopicId: 302226
dashedName: averagesmode
---

# --description--

배열에서 가장 자주 나타나는 값을 찾는 함수 `mode`를 작성하세요.

컬렉션이 비어 있는 경우는 무시할 수 있습니다. 최빈값이 유일하지 않은 경우를 처리하는 데 주의해야 합니다.

일반 컬렉션을 지원하는 것이 적절하지 않거나 불가능한 경우, 가능하면 벡터(배열)를 사용하세요. 지정되지 않은 값 유형을 지원하는 것이 적절하지 않거나 불가능한 경우, 정수를 사용하세요.

# --before-each--

```js
const arr1 = [1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17];
const arr2 = [1, 2, 4, 4, 1];
```

# --hints--

`mode`는 함수여야 합니다.

```js
assert(typeof mode === 'function');
```

`mode([1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17])`는 `[6]`와 같아야 합니다.

```js
assert.deepEqual(mode(arr1), [6]);
```

`mode([1, 2, 4, 4, 1])`는 `[1, 4]`와 같아야 합니다.

```js
assert.deepEqual(mode(arr2).sort(), [1, 4]);
```

# --seed--

## --seed-contents--

```js
function mode(arr) {

  return true;
}
```

# --solutions--

```js
function mode(arr) {
  const counter = {};
  let result = [];
  let max = 0;
  // for (const i in arr) {
  arr.forEach(el => {
    if (!(el in counter)) {
      counter[el] = 0;
    }
    counter[el]++;

    if (counter[el] === max) {
      result.push(el);
    }
    else if (counter[el] > max) {
      max = counter[el];
      result = [el];
    }
  });
  return result;
}
```
