---
id: 594da033de4190850b893874
title: 평균/제곱근 평균 제곱값
challengeType: 1
forumTopicId: 302228
dashedName: averagesroot-mean-square
---

# --description--

1부터 10까지의 숫자에 대한 제곱근 평균 제곱값(RMS)을 계산하세요.

<abbr title="Root mean square">RMS</abbr>는 숫자들의 제곱의 평균에 제곱근을 취해 계산하며, 다음 식으로 표현됩니다:

$$x\_{\\mathrm{rms}} = \\sqrt {{{x_1}^2 + {x_2}^2 + \\cdots + {x_n}^2} \\over n}. $$

# --hints--

`rms`는 함수여야 합니다.

```js
assert(typeof rms === 'function');
```

`rms([1, 2, 3, 4, 5, 6, 7, 8, 9, 10])`는 `6.2048368229954285`와 같아야 합니다.

```js
const arr1 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const answer1 = 6.2048368229954285;
assert.equal(rms(arr1), answer1);
```

# --seed--

## --seed-contents--

```js
function rms(arr) {

}
```

# --solutions--

```js
function rms(arr) {
  const sumOfSquares = arr.reduce((s, x) => s + x * x, 0);
  return Math.sqrt(sumOfSquares / arr.length);
}
```
