---
id: 594dc6c729e5700999302b45
title: 균형 잡힌 괄호
challengeType: 1
forumTopicId: 302230
dashedName: balanced-brackets
---

# --description--

생성된 괄호 문자열이 균형 잡혔는지 판단하세요. 즉, 모든 괄호가 올바른 순서로 열린 괄호와 닫힌 괄호 쌍으로 이루어져 있으며, 잘못 중첩된 부분이 없는지 확인하세요.

**예시:**
| 입력                      | 출력   |
| ------------------------- | ------ |
| <code>\[]</code>          | true   |
| <code>]\[</code>          | false  |
| <code>[][]</code>         | true   |
| <code>]\[]</code>         | false  |
| <code>\[]]\[\[]</code>    | false  |
| <code>\[\[\[\[]]]]</code> | true   |

# --before-each--

```js
const testCases = [
  '[]',
  ']][[[][][][]][',
  '[][[[[][][[[]]]]]]',
  '][',
  '[[[]]]][[]',
  '][[]',
  '][[][]][[[]]',
  '[[][]]][',
  '[[[]]][[]]]][][[',
  '[]][[]]][[[[][]]',
  '][]][[][',
  '[[]][[][]]',
  '[[]]',
  ']][]][[]][[[',
  '][]][][[',
  '][][',
  '[]]]',
  ''
];
```

# --hints--

`isBalanced`는 함수여야 합니다.

```js
assert(typeof isBalanced === 'function');
```

`isBalanced("[]")`은 true를 반환해야 합니다.

```js
assert(isBalanced(testCases[0]));
```

`isBalanced("]][[[][][][]][")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[1]));
```

`isBalanced("[][[[[][][[[]]]]]]")`은 true를 반환해야 합니다.

```js
assert(isBalanced(testCases[2]));
```

`isBalanced("][")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[3]));
```

`isBalanced("[[[]]]][[]")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[4]));
```

`isBalanced("][[]")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[5]));
```

`isBalanced("][[][]][[[]]")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[6]));
```

`isBalanced("[[][]]][")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[7]));
```

`isBalanced("[[[]]][[]]]][][[")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[8]));
```

`isBalanced("[]][[]]][[[[][]]")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[9]));
```

`isBalanced("][]][[][")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[10]));
```

`isBalanced("[[]][[][]]")`는 true를 반환해야 합니다.

```js
assert(isBalanced(testCases[11]));
```

`isBalanced("[[]]")`은 true를 반환해야 합니다.

```js
assert(isBalanced(testCases[12]));
```

`isBalanced("]][]][[]][[[")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[13]));
```

`isBalanced("][]][][[")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[14]));
```

`isBalanced("][][")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[15]));
```

`isBalanced("[]]]")`는 false를 반환해야 합니다.

```js
assert(!isBalanced(testCases[16]));
```

`isBalanced("")`는 true를 반환해야 합니다.

```js
assert(isBalanced(testCases[17]));
```

# --seed--

## --seed-contents--

```js
function isBalanced(str) {

  return true;
}
```

# --solutions--

```js
function isBalanced(str) {
  if (str === '') return true;
  let a = str;
  let b;
  do {
    b = a;
    a = a.replace(/\[\]/g, '');
  } while (a !== b);
  return !a;
}
```
