---
id: 594ecc0d9a8cf816e3340187
title: 택시캡 수
challengeType: 1
forumTopicId: 302337
dashedName: taxicab-numbers
---

# --description--

택시캡 수(여기서 사용되는 정의)는 두 개의 양의 세제곱수 합으로 여러 가지 방법으로 표현할 수 있는 양의 정수입니다.

첫 번째 택시캡 수는 `1729`이며, 이는 다음과 같습니다:

1<sup>3</sup> + 12<sup>3</sup> 그리고

9<sup>3</sup> + 10<sup>3</sup>.

택시캡 수는 다음 이름으로도 알려져 있습니다:

<ul>
  <li>택시 수</li>
  <li>택시캡 수</li>
  <li>택시 캡 수</li>
  <li>하디-라마누잔 수</li>
</ul>

# --instructions--

가장 낮은 `n` 택시캡 수를 반환하는 함수를 작성하세요. 각 택시캡 수에 대해 그 수와 구성하는 세제곱수를 함께 보여주세요.

# --before-each--

```js
const res4 = [1729, 4104, 13832, 20683];
const res25 = [
  1729, 4104, 13832, 20683, 32832, 39312, 40033, 46683, 64232, 65728, 110656,
  110808, 134379, 149389, 165464, 171288, 195841, 216027, 216125, 262656, 314496, 320264, 327763,
  373464, 402597
];

const res39From20To29 = [314496, 320264, 327763, 373464, 402597, 439101, 443889, 513000, 513856];
```

# --hints--

`taxicabNumbers`는 함수여야 합니다.

```js
assert(typeof taxicabNumbers === 'function');
```

`taxicabNumbers`는 배열을 반환해야 합니다.

```js
assert(typeof taxicabNumbers(2) === 'object');
```

`taxicabNumbers`는 숫자 배열을 반환해야 합니다.

```js
assert(typeof taxicabNumbers(100)[0] === 'number');
```

`taxicabNumbers(4)`는 [1729, 4104, 13832, 20683]를 반환해야 합니다.

```js
assert.deepEqual(taxicabNumbers(4), res4);
```

`taxicabNumbers(25)`는 [1729, 4104, 13832, 20683, 32832, 39312, 40033, 46683, 64232, 65728, 110656, 110808, 134379, 149389, 165464, 171288, 195841, 216027, 216125, 262656, 314496, 320264, 327763, 373464, 402597]를 반환해야 합니다.

```js
assert.deepEqual(taxicabNumbers(25), res25);
```

`taxicabNumbers(39)`의 20~29번째 결과 수는 [314496,320264,327763,373464,402597,439101,443889,513000,513856]이어야 합니다.

```js
assert.deepEqual(taxicabNumbers(39).slice(20, 29), res39From20To29);
```

# --seed--

## --seed-contents--

```js
function taxicabNumbers(n) {

  return true;
}
```

# --solutions--

```js
function taxicabNumbers(nNumbers) {
  const cubeN = [];
  const s3s = {};

  const e = 100;
  for (let n = 1; n < e; n += 1) {
    cubeN[n] = n * n * n;
  }

  for (let a = 1; a < e - 1; a += 1) {
    const a3 = cubeN[a];
    for (let b = a; b < e; b += 1) {
      const b3 = cubeN[b];
      const s3 = a3 + b3;

      let abs = s3s[s3];
      if (!abs) {
        s3s[s3] = abs = [];
      }
      abs.push([a, b]);
    }
  }

  let i = 0;
  const res = [];
  Object.keys(s3s).forEach(s3 => {
    const abs = s3s[s3];
    if (abs.length >= 2) { // No two cube pairs found
      i += 1;
      if (i <= nNumbers) {
        res.push(s3);
      }
    }
  });
  return res.map(item => parseInt(item, 10));
}
```
