---
id: 594faaab4e2a8626833e9c3d
title: 이스케이프 처리된 문자열 토큰화
challengeType: 1
forumTopicId: 302338
dashedName: tokenize-a-string-with-escaping
---

# --description--

이스케이프 처리되지 않은 구분자 문자가 나올 때마다 문자열을 분할할 수 있는 함수나 프로그램을 작성하세요.

세 가지 입력 매개변수를 받아야 합니다:

<ul>
  <li><strong>문자열</strong></li>
  <li><strong>구분자 문자</strong></li>
  <li><strong>이스케이프 문자</strong></li>
</ul>

출력은 문자열 목록이어야 합니다.

분할 규칙:

<ul>
  <li>구분자로 나뉜 필드들이 출력 목록의 요소가 됩니다.</li>
  <li>빈 필드도 시작과 끝에 있더라도 그대로 유지해야 합니다.</li>
</ul>

이스케이프 규칙:

<ul>
  <li>"이스케이프됨"은 이미 이스케이프되지 않은 이스케이프 문자가 앞에 있는 경우를 의미합니다.</li>
  <li>이스케이프 문자가 특별한 의미가 없는 문자를 앞에 두어도 이스케이프로 간주하지만, 특별한 동작은 하지 않습니다.</li>
  <li>무언가를 이스케이프하는 데 사용된 이스케이프 문자는 출력에 포함되지 않아야 합니다.</li>
</ul>

다음 테스트 케이스를 만족하는지 함수가 작동하는지 보여주세요:

문자열이

<pre>one^|uno||three^^^^|four^^^|^cuatro|</pre>

이고 `|`를 구분자로, `^`를 이스케이프 문자로 사용할 때, 함수는 다음 배열을 출력해야 합니다:

<pre>  ['one|uno', '', 'three^^', 'four^|cuatro', '']
</pre>

# --before-each--

```js
const testStr1 = 'one^|uno||three^^^^|four^^^|^cuatro|';
const res1 = ['one|uno', '', 'three^^', 'four^|cuatro', ''];

const testStr2 = 'a@&bcd&ef&&@@hi';
const res2 = ['a&bcd', 'ef', '', '@hi'];

const testStr3 = 'hello^|world|how^are^you^|';
const res3 = ['hello|world', 'howareyou|'];
```

# --hints--

`tokenize`는 함수여야 합니다.

```js
assert(typeof tokenize === 'function');
```

`tokenize`는 배열을 반환해야 합니다.

```js
assert(typeof tokenize('a', 'b', 'c') === 'object');
```

`tokenize('one^|uno||three^^^^|four^^^|^cuatro|', '|', '^')`는 `['one|uno', '', 'three^^', 'four^|cuatro', '']`를 반환해야 합니다.

```js
assert.deepEqual(tokenize(testStr1, '|', '^'), res1);
```

`tokenize('a@&bcd&ef&&@@hi', '&', '@')`는 `['a&bcd', 'ef', '', '@hi']`를 반환해야 합니다.

```js
assert.deepEqual(tokenize(testStr2, '&', '@'), res2);
```

`tokenize('hello^|world|how^are^you^|', '|', '^')`는 `['hello|world', 'howareyou|']`를 반환해야 합니다.

```js
assert.deepEqual(tokenize(testStr3, '|', '^'), res3);
```

# --seed--

## --seed-contents--

```js
function tokenize(str, sep, esc) {
  return true;
}
```

# --solutions--

```js
// tokenize :: String -> Character -> Character -> [String]
function tokenize(str, charDelim, charEsc) {
  const dctParse = str.split('')
    .reduce((a, x) => {
      const blnEsc = a.esc;
      const blnBreak = !blnEsc && x === charDelim;
      const blnEscChar = !blnEsc && x === charEsc;

      return {
        esc: blnEscChar,
        token: blnBreak ? '' : (
          a.token + (blnEscChar ? '' : x)
        ),
        list: a.list.concat(blnBreak ? a.token : [])
      };
    }, {
      esc: false,
      token: '',
      list: []
    });

  return dctParse.list.concat(
    dctParse.token
  );
}
```
