---
id: 594fa2746886f41f7d8bf225
title: 위상 정렬
challengeType: 1
forumTopicId: 302340
dashedName: topological-sort
---

# --description--

항목과 그 항목이 의존하는 항목 간의 매핑이 주어졌을 때, 위상 정렬은 어떤 항목도 자신이 의존하는 항목보다 먼저 오지 않도록 항목의 순서를 정합니다. 위상 정렬에는 두 가지 인기 있는 알고리즘이 있습니다: Kahn의 위상 정렬(1962)과 깊이 우선 탐색입니다.

# --instructions--

의존성으로부터 라이브러리의 유효한 컴파일 순서를 반환하는 함수를 작성하세요.

- 라이브러리 이름은 한 단어라고 가정하세요.
- 의존하는 항목으로만 언급된 항목은 자체 의존 항목이 없지만, 컴파일 순서는 반드시 지정되어야 합니다.
- 자기 자신에 대한 의존성은 무시하세요.
- 순서 지정이 불가능한 의존성은 무시하세요.

다음 데이터를 예시로 사용하세요:

<pre>
LIBRARY          LIBRARY DEPENDENCIES
=======          ====================
des_system_lib   std synopsys std_cell_lib des_system_lib dw02 dw01 ramlib ieee
dw01             ieee dw01 dware gtech
dw02             ieee dw02 dware
dw03             std synopsys dware dw03 dw02 dw01 ieee gtech
dw04             dw04 ieee dw01 dware gtech
dw05             dw05 ieee dware
dw06             dw06 ieee dware
dw07             ieee dware
dware            ieee dware
gtech            ieee gtech
ramlib           std ieee
std_cell_lib     ieee std_cell_lib
synopsys
</pre>

VHDL 언어에서 라이브러리를 컴파일할 때는, 라이브러리가 의존하는 모든 라이브러리 이후에 컴파일되어야 한다는 제약이 있습니다. 예를 들어, `dw04`가 `dw01`의 의존성 목록에 추가되면 위 데이터는 순서 지정이 불가능해집니다.

함수의 입력은 여러 줄 문자열이며, 각 줄은 라이브러리 이름과 그 의존성(존재하는 경우)으로 구성됩니다.

예를 들어:

```js
const libsSimple =
  `aaa bbb
  bbb`;
```

# --before-each--

```js
const libsSimple =
  `aaa bbb
  bbb`;

const libsVHDL =
  `des_system_lib   std synopsys std_cell_lib des_system_lib dw02 dw01 ramlib ieee
  dw01             ieee dw01 dware gtech
  dw02             ieee dw02 dware
  dw03             std synopsys dware dw03 dw02 dw01 ieee gtech
  dw04             dw04 ieee dw01 dware gtech
  dw05             dw05 ieee dware
  dw06             dw06 ieee dware
  dw07             ieee dware
  dware            ieee dware
  gtech            ieee gtech
  ramlib           std ieee
  std_cell_lib     ieee std_cell_lib
  synopsys`;

const libsCustom =
  `a b c d
  b c d
  d c
  c base
  base`;

const libsUnorderable =
  `TestLib Base MainLib
  MainLib TestLib
  Base`;
```

# --hints--

`topologicalSort`는 함수여야 합니다.

```js
assert(typeof topologicalSort === 'function');
```

`topologicalSort(libsSimple)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(topologicalSort(libsSimple)));
```

`topologicalSort(libsSimple)`는 `['bbb', 'aaa']`를 반환해야 합니다.

```js
assert.deepEqual(topologicalSort(libsSimple), ['bbb', 'aaa']);
```

`topologicalSort(libsVHDL)`는 `['ieee', 'std_cell_lib', 'gtech', 'dware', 'dw07', 'dw06', 'dw05', 'dw02', 'dw01', 'dw04', 'std', 'ramlib', 'synopsys', 'dw03', 'des_system_lib']`를 반환해야 합니다.

```js
assert.deepEqual(topologicalSort(libsVHDL), ['ieee', 'std_cell_lib', 'gtech', 'dware', 'dw07', 'dw06', 'dw05', 'dw02', 'dw01', 'dw04', 'std', 'ramlib', 'synopsys', 'dw03', 'des_system_lib']);
```

`topologicalSort(libsCustom)`는 `['base', 'c', 'd', 'b', 'a']`를 반환해야 합니다.

```js
assert.deepEqual(topologicalSort(libsCustom), ['base', 'c', 'd', 'b', 'a']);
```

`topologicalSort`는 순서 지정이 불가능한 의존성을 무시해야 합니다.

```js
assert.deepEqual(topologicalSort(libsUnorderable), ['Base']);
```

# --seed--

## --seed-contents--

```js
function topologicalSort(libs) {

  return true;
}
```

# --solutions--

```js
function topologicalSort(libs) {
  // A map of the input data, with the keys as the packages, and the values as
  // and array of packages on which it depends.
  const D = libs
    .split('\n')
    .map(e => e.split(' ').filter(ep => ep !== ''))
    .reduce((p, c) =>
      p.set(c[0], c.filter((e, i) => (i > 0 && e !== c[0] ? e : null))), new Map());
  [].concat(...D.values()).forEach(e => {
    D.set(e, D.get(e) || []);
  });

  // The above map rotated so that it represents a DAG of the form
  // Map {
  //    A => [ A, B, C],
  //    B => [C],
  //    C => []
  // }
  // where each key represents a node, and the array contains the edges.
  const G = [...D.keys()].reduce((p, c) =>
    p.set(
      c,
      [...D.keys()].filter(e => D.get(e).includes(c))),
    new Map()
  );

  // An array of leaf nodes; nodes with 0 in degrees.
  const Q = [...D.keys()].filter(e => D.get(e).length === 0);

  // The result array.
  const S = [];
  while (Q.length) {
    const u = Q.pop();
    S.push(u);
    G.get(u).forEach(v => {
      D.set(v, D.get(v).filter(e => e !== u));
      if (D.get(v).length === 0) {
        Q.push(v);
      }
    });
  }

  return S;
}
```
