---
id: 595011cba5a81735713873bd
title: 그룹별 최고 순위
challengeType: 1
forumTopicId: 302339
dashedName: top-rank-per-group
---

# --description--

`n`이 매개변수로 주어질 때, 각 그룹에서 상위 `n` 순위 데이터를 찾으세요. 순위 이름과 그룹 이름도 매개변수로 제공됩니다.

다음 데이터가 주어졌을 때:

```js
testData1 = [
  { name: 'Tyler Bennett', id: 'E10297', salary: 32000, dept: 'D101' },
  { name: 'John Rappl', id: 'E21437', salary: 47000, dept: 'D050' },
  { name: 'George Woltman', id: 'E00127', salary: 53500, dept: 'D101' },
  { name: 'Adam Smith', id: 'E63535', salary: 18000, dept: 'D202' },
  { name: 'Claire Buckman', id: 'E39876', salary: 27800, dept: 'D202' },
  { name: 'David McClellan', id: 'E04242', salary: 41500, dept: 'D101' },
  { name: 'Rich Holcomb', id: 'E01234', salary: 49500, dept: 'D202' },
  { name: 'Nathan Adams', id: 'E41298', salary: 21900, dept: 'D050' },
  { name: 'Richard Potter', id: 'E43128', salary: 15900, dept: 'D101' },
  { name: 'David Motsinger', id: 'E27002', salary: 19250, dept: 'D202' },
  { name: 'Tim Sampair', id: 'E03033', salary: 27000, dept: 'D101' },
  { name: 'Kim Arlich', id: 'E10001', salary: 57000, dept: 'D190' },
  { name: 'Timothy Grove', id: 'E16398', salary: 29900, dept: 'D190' }
];
```

다음과 같이 호출하여 각 부서별 상위 10명의 직원을 순위 매길 수 있습니다:

```js
topRankPerGroup(10, testData1, 'dept', 'salary')
```

다음 데이터가 주어졌을 때:

```js
testData2 = [
  { name: 'Friday 13th', genre: 'horror', rating: 9.9 },
  { name: "Nightmare on Elm's Street", genre: 'horror', rating: 5.7 },
  { name: 'Titanic', genre: 'drama', rating: 7.3 },
  { name: 'Maze Runner', genre: 'scifi', rating: 7.1 },
  { name: 'Blade runner', genre: 'scifi', rating: 8.9 }
];
```

다음과 같이 호출하여 각 장르별 최고 평점 영화를 순위 매길 수 있습니다:

```js
topRankPerGroup(1, testData2, 'genre', 'rating')
```

함수는 각 그룹별로 상위 `n` 객체를 포함하는 배열을 반환해야 합니다.

예를 들어, 다음 데이터가 주어졌을 때:

```js
[
  { name: 'Claire Buckman', id: 'E39876', salary: 27800, dept: 'D101' },
  { name: 'Rich Holcomb', id: 'E01234', salary: 49500, dept: 'D050' },
  { name: 'David Motsinger', id: 'E27002', salary: 19250, dept: 'D050' },
  { name: 'Tim Sampair', id: 'E03033', salary: 27000, dept: 'D101' },
  { name: 'Kim Arlich', id: 'E10001', salary: 57000, dept: 'D050' },
  { name: 'Timothy Grove', id: 'E16398', salary: 29900, dept: 'D101' }
];
```

각 부서별 급여 상위 두 명의 직원은 다음과 같습니다:

```js
[ [ { name: 'Kim Arlich', id: 'E10001', salary: 57000, dept: 'D050' },
    { name: 'Rich Holcomb', id: 'E01234', salary: 49500, dept: 'D050' } ],
  [ { name: 'Timothy Grove', id: 'E16398', salary: 29900, dept: 'D101' },
    { name: 'Claire Buckman', id: 'E39876', salary: 27800, dept: 'D101' } ] ]
```

# --before-each--

```js
const testData1 = [
  { name: 'Tyler Bennett', id: 'E10297', salary: 32000, dept: 'D101' },
  { name: 'John Rappl', id: 'E21437', salary: 47000, dept: 'D050' },
  { name: 'George Woltman', id: 'E00127', salary: 53500, dept: 'D101' },
  { name: 'Adam Smith', id: 'E63535', salary: 18000, dept: 'D202' },
  { name: 'Claire Buckman', id: 'E39876', salary: 27800, dept: 'D202' },
  { name: 'David McClellan', id: 'E04242', salary: 41500, dept: 'D101' },
  { name: 'Rich Holcomb', id: 'E01234', salary: 49500, dept: 'D202' },
  { name: 'Nathan Adams', id: 'E41298', salary: 21900, dept: 'D050' },
  { name: 'Richard Potter', id: 'E43128', salary: 15900, dept: 'D101' },
  { name: 'David Motsinger', id: 'E27002', salary: 19250, dept: 'D202' },
  { name: 'Tim Sampair', id: 'E03033', salary: 27000, dept: 'D101' },
  { name: 'Kim Arlich', id: 'E10001', salary: 57000, dept: 'D190' },
  { name: 'Timothy Grove', id: 'E16398', salary: 29900, dept: 'D190' }
];

const testData2 = [
  { name: 'Friday 13th', genre: 'horror', rating: 9.9 },
  { name: "Nightmare on Elm's Street", genre: 'horror', rating: 5.7 },
  { name: 'Titanic', genre: 'drama', rating: 7.3 },
  { name: 'Maze Runner', genre: 'scifi', rating: 7.1 },
  { name: 'Blade runner', genre: 'scifi', rating: 8.9 }
];

function getResults_() {
  return {
    res1: topRankPerGroup(10, testData1, 'dept', 'salary'),
    res2: topRankPerGroup(1, testData2, 'genre', 'rating'),
    res3: topRankPerGroup(2, testData2, 'genre', 'rating')
  };
}
```

# --hints--

`topRankPerGroup`는 함수여야 합니다.

```js
assert(typeof topRankPerGroup === 'function');
```

`topRankPerGroup`는 음수 n 값에 대해 undefined를 반환해야 합니다.

```js
assert(typeof topRankPerGroup(-1, []) === 'undefined');
```

`topRankPerGroup(10, testData1, 'dept', 'salary')`의 경우, 첫 번째 그룹의 첫 번째 결과는 `{ name: 'John Rappl', id: 'E21437', salary: 47000, dept: 'D050'}`이어야 합니다.

```js
assert.deepEqual(getResults_().res1[0][0], { name: 'John Rappl', id: 'E21437', salary: 47000, dept: 'D050'});
```

`topRankPerGroup(10, testData1, 'dept', 'salary')`의 경우, 마지막 그룹의 마지막 결과는 `{ name: 'Adam Smith', id: 'E63535', salary: 18000, dept: 'D202' }`이어야 합니다.

```js
assert.deepEqual(getResults_().res1[3][3], { name: 'Adam Smith', id: 'E63535', salary: 18000, dept: 'D202' });
```

`topRankPerGroup(1, ...)`는 각 그룹별로 상위 한 개의 결과만 반환해야 합니다.

```js
const { res2 } = getResults_();
assert.equal(res2[2].length, 1);
```

`topRankPerGroup(2, ...)`는 각 그룹별로 두 개의 결과를 반환해야 합니다.

```js
const { res3 } = getResults_();
assert.equal(res3[2][1].name, 'Maze Runner');
```

# --seed--

## --seed-contents--

```js
function topRankPerGroup(n, data, groupName, rankName) {

  return true;
}
```

# --solutions--

```js
const collectDept = function (arrOfObj, groupName) {
  const collect = arrOfObj.reduce((rtnObj, obj) => {
    if (rtnObj[obj[groupName]] === undefined) {
      rtnObj[obj[groupName]] = [];
    }
    rtnObj[obj[groupName]].push(obj);
    return rtnObj;
  }, {} // initial value to reduce
  );

  return Object.keys(collect).sort().map(key => collect[key]);
};

const sortRank = function (arrOfRankArrs, rankName) {
  return arrOfRankArrs.map(item => item.sort((a, b) => {
    if (a[rankName] > b[rankName]) { return -1; }
    if (a[rankName] < b[rankName]) { return 1; }
    return 0;
  }));
};

function topRankPerGroup(n, data, groupName, rankName) {
  if (n < 0) { return; }
  return sortRank(collectDept(data, groupName),
    rankName).map(list => list.slice(0, n));
}
```
