---
id: 5951815dd895584b06884620
title: 두 점을 지나는 주어진 반지름의 원
challengeType: 1
forumTopicId: 302231
dashedName: circles-of-given-radius-through-two-points
---

# --description--

평면 위의 두 점과 반지름이 주어지면, 보통 두 점을 지나는 주어진 반지름의 원 두 개를 그릴 수 있습니다.

**예외 사항:**

<ul>
  <li>반지름이 0인 경우는 원을 나타내지 않는 것으로 처리해야 합니다(단, 두 점이 일치하는 경우는 제외).</li>
  <li>두 점이 일치하면, 반지름이 0이 아닌 한 무한히 많은 원이 그 점을 원주 위에 가질 수 있습니다. 반지름도 0이면 원이 점으로 축소됩니다.</li>
  <li>두 점이 지름을 이루면 하나의 원만 반환합니다.</li>
  <li>두 점이 너무 멀리 떨어져 있으면 원을 그릴 수 없습니다.</li>
</ul>

# --instructions--

두 점과 반지름을 입력받아 그 점들을 지나는 두 원을 반환하는 함수를 구현하세요. 각 원에 대해 중심 좌표를 소수점 넷째 자리까지 반올림하여 제공하세요. 각 좌표는 배열로, 좌표들의 모음도 배열로 반환하세요.

**특수한 경우에는 다음을 반환하세요:**

<ul>
  <li>두 점이 지름 위에 있으면 하나의 점을 반환하세요. 단, 반지름도 0이면 <code>"Radius Zero"</code>를 반환하세요.</li>
  <li>두 점이 일치하면 <code>"Coincident point. Infinite solutions"</code>를 반환하세요.</li>
  <li>두 점이 지름보다 멀리 떨어져 있으면 <code>"No intersection. Points further apart than circle diameter"</code>를 반환하세요.</li>
</ul>

**입력 예시:**

<pre>      p1                p2           r
0.1234, 0.9876    0.8765, 0.2345    2.0
0.0000, 2.0000    0.0000, 0.0000    1.0
0.1234, 0.9876    0.1234, 0.9876    2.0
0.1234, 0.9876    0.8765, 0.2345    0.5
0.1234, 0.9876    0.1234, 0.9876    0.0
</pre>

# --before-each--

```js
const testCases = [
  [[0.1234, 0.9876], [0.8765, 0.2345], 2.0],
  [[0.0000, 2.0000], [0.0000, 0.0000], 1.0],
  [[0.1234, 0.9876], [0.1234, 0.9876], 2.0],
  [[0.1234, 0.9876], [0.8765, 0.2345], 0.5],
  [[0.1234, 0.9876], [0.1234, 0.9876], 0.0]
];
const answers = [
  [[1.8631, 1.9742], [-0.8632, -0.7521]],
  [0, 1],
  'Coincident point. Infinite solutions',
  'No intersection. Points further apart than circle diameter',
  'Radius Zero'
];
```

# --hints--

`getCircles` should be a function.

```js
assert(typeof getCircles === 'function');
```

`getCircles([0.1234, 0.9876], [0.8765, 0.2345], 2.0)` should return `[[1.8631, 1.9742], [-0.8632, -0.7521]]`.

```js
assert.deepEqual(getCircles(...testCases[0]), answers[0]);
```

`getCircles([0.0000, 2.0000], [0.0000, 0.0000], 1.0)` should return `[0, 1]`

```js
assert.deepEqual(getCircles(...testCases[1]), answers[1]);
```

`getCircles([0.1234, 0.9876], [0.1234, 0.9876], 2.0)` should return `Coincident point. Infinite solutions`

```js
assert.deepEqual(getCircles(...testCases[2]), answers[2]);
```

`getCircles([0.1234, 0.9876], [0.8765, 0.2345], 0.5)` should return `No intersection. Points further apart than circle diameter`

```js
assert.deepEqual(getCircles(...testCases[3]), answers[3]);
```

`getCircles([0.1234, 0.9876], [0.1234, 0.9876], 0.0)` should return `Radius Zero`

```js
assert.deepEqual(getCircles(...testCases[4]), answers[4]);
```

# --seed--

## --seed-contents--

```js
function getCircles(...args) {

  return true;
}
```

# --solutions--

```js
const hDist = (p1, p2) => Math.hypot(...p1.map((e, i) => e - p2[i])) / 2;
const pAng = (p1, p2) => Math.atan(p1.map((e, i) => e - p2[i]).reduce((p, c) => c / p, 1));
const solveF = (p, r) => t => [parseFloat((r * Math.cos(t) + p[0]).toFixed(4)), parseFloat((r * Math.sin(t) + p[1]).toFixed(4))];
const diamPoints = (p1, p2) => p1.map((e, i) => parseFloat((e + (p2[i] - e) / 2).toFixed(4)));

function getCircles(...args) {
  const [p1, p2, s] = args;
  const solve = solveF(p1, s);
  const halfDist = hDist(p1, p2);

  let msg = [];
  switch (Math.sign(s - halfDist)) {
    case 0:
      msg = s ? diamPoints(p1, p2) :
        'Radius Zero';
      break;
    case 1:
      if (!halfDist) {
        msg = 'Coincident point. Infinite solutions';
      }
      else {
        const theta = pAng(p1, p2);
        const theta2 = Math.acos(halfDist / s);
        [1, -1].map(e => solve(theta + e * theta2)).forEach(
          e => msg.push(e));
      }
      break;
    case -1:
      msg = 'No intersection. Points further apart than circle diameter';
      break;
    default:
      msg = 'Reached the default';
  }
  return msg;
}
```
