---
id: 5951e88f64ebf159166a1176
title: 24 게임
challengeType: 1
forumTopicId: 302218
dashedName: 24-game
---

# --description--

24 게임은 사람의 정신 산술 능력을 테스트합니다.

게임의 목표는 네 개의 숫자를 평가했을 때 결과가 24가 되도록 배열하는 것입니다.

# --instructions--

네 자리 숫자 문자열을 인수로 받아, 각 숫자는 1부터 9까지(중복 허용)이며, 결과가 24가 되는 산술 표현식을 반환하는 함수를 구현하세요. 그런 해가 없으면 "no solution exists"를 반환하세요.

**규칙:**
<ul>
  <li> 허용되는 연산자/함수는 곱셈, 나눗셈, 덧셈, 뺄셈만 가능합니다. </li>
  <li> 나눗셈은 나머지를 보존하기 위해 부동소수점 또는 유리수 산술 등을 사용해야 합니다. </li>
  <li> 주어진 숫자들로 여러 자리 숫자를 만드는 것은 허용되지 않습니다. (예: 1, 2, 2, 1이 주어졌을 때 12+12는 잘못된 답입니다.) </li>
  <li> 주어진 숫자의 순서를 유지할 필요는 없습니다. </li>
</ul>

| 입력 예시                     | 출력 예시                 |
| ----------------------------- | ------------------------- |
| <code>solve24("4878");</code> | <code>(7-8/8)*4</code>    |
| <code>solve24("1234");</code> | <code>3*1*4*2</code>      |
| <code>solve24("6789");</code> | <code>(6*8)/(9-7)</code>  |
| <code>solve24("1127");</code> | <code>(1+7)*(2+1)</code>  |

# --before-each--

```js
const testCases_ = [
  '4878',
  '1234',
  '6789',
  '1127'
];

const OPERATORS_ = {
  '+': (a, b) => a + b,
  '-': (a, b) => a - b,
  '*': (a, b) => a * b,
  '/': (a, b) => a / b
};

const PRECEDENCE_ = {
  '+': 1,
  '-': 1,
  '*': 2,
  '/': 2
};

function evaluate_(expression) {
  expression = expression.replace('/\\s+/g', '');
  const stack = [];
  let postfix = '';

  // Convert from infix to postfix
  let head = 0;
  while (head < expression.length) {
    let c = expression[head];
    switch (c) {
      case '(':
        stack.push(c);
        break;
      case ')':
        let last = stack.pop();
        while (last !== '(') {
          postfix += last;
          last = stack.pop();
        }
        break;
      case '+':
      case '-':
      case '*':
      case '/':
        while (stack.length && PRECEDENCE_[c] <= PRECEDENCE_[stack[stack.length - 1]]) {
          postfix += stack.pop();
        }
        stack.push(c);
        break;
      case '0':
      case '1':
      case '2':
      case '3':
      case '4':
      case '5':
      case '6':
      case '7':
      case '8':
      case '9':
        postfix += c;
        break;
      default:
        return false;
    }
    head++;
  }

  // Clear out stack
  while (stack.length) {
    postfix += stack.pop();
  }

  // Evaluate postfix
  for (let c of postfix) {
    switch (c) {
      case '+':
      case '-':
      case '*':
      case '/':
        const b = +stack.pop();
        const a = +stack.pop();
        stack.push(OPERATORS_[c](a, b));
        break;
      default:
        stack.push(c);
    }
  }
  return stack.pop();
}

function isValidSolution_(userSolution) {
  return evaluate_(userSolution) === 24;
}
```

# --hints--

`solve24`는 함수여야 합니다.

```js
assert(typeof solve24 === 'function');
```

`solve24("4878")`는 `(7-8/8)*4`, `4*(7-8/8)` 또는 유사한 유효한 문자열을 반환해야 합니다.

```js
assert(isValidSolution_(solve24(testCases_[0])));
```

`solve24("1234")`는 `1*2*3*4` 또는 유사한 유효한 문자열을 반환해야 합니다.

```js
assert(isValidSolution_(solve24(testCases_[1])));
```

`solve24("6789")`는 `(6*8)/(9-7)`, `(8*6)/(9-7)` 또는 유사한 유효한 문자열을 반환해야 합니다.

```js
assert(isValidSolution_(solve24(testCases_[2])));
```

`solve24("1127")`는 `(1+7)*(1+2)` 또는 유사한 유효한 문자열을 반환해야 합니다.

```js
assert(isValidSolution_(solve24(testCases_[3])));
```

# --seed--

## --seed-contents--

```js
function solve24 (numStr) {

  return true;
}
```

# --solutions--

```js
function solve24(numStr) {
  const digitsArr = numStr.split('');
  const answers = [];

  const digitPermutations = [];
  const operatorPermutations = [];

  function generateDigitPermutations (digits, permutations = []) {
    if (digits.length === 0) {
      digitPermutations.push(permutations);
    }
    else {
      for (let i = 0; i < digits.length; i++) {
        const curr = digits.slice();
        const next = curr.splice(i, 1);
        generateDigitPermutations(curr.slice(), permutations.concat(next));
      }
    }
  }

  function generateOperatorPermutations (permutations = []) {
    const operators = ['+', '-', '*', '/'];
    if (permutations.length === 3) {
      operatorPermutations.push(permutations);
    }
    else {
      for (let i = 0; i < operators.length; i++) {
        const curr = permutations.slice();
        curr.push(operators[i]);
        generateOperatorPermutations(curr);
      }
    }
  }

  generateDigitPermutations(digitsArr);
  generateOperatorPermutations();

  interleave();

  return answers[0];

  function interleave () {
    for (let i = 0; i < digitPermutations.length; i++) {
      for (let j = 0; j < operatorPermutations.length; j++) {
        const d = digitPermutations[i];
        const o = operatorPermutations[j];
        const perm = [
          `${d[0]}${o[0]}${d[1]}${o[1]}${d[2]}${o[2]}${d[3]}`,
          `(${d[0]}${o[0]}${d[1]})${o[1]}${d[2]}${o[2]}${d[3]}`,
          `${d[0]}${o[0]}(${d[1]}${o[1]}${d[2]})${o[2]}${d[3]}`,
          `${d[0]}${o[0]}${d[1]}${o[1]}(${d[2]}${o[2]}${d[3]})`,
          `${d[0]}${o[0]}(${d[1]}${o[1]}${d[2]}${o[2]}${d[3]})`,
          `(${d[0]}${o[0]}${d[1]}${o[1]}${d[2]})${o[2]}${d[3]}`,
          `(${d[0]}${o[0]}${d[1]})${o[1]}(${d[2]}${o[2]}${d[3]})`
        ];

        perm.forEach(combination => {
          const res = eval(combination);

          if (res === 24) {
            return answers.push(combination);
          }
        });
      }
    }
  }
}
```
