---
id: 5951ed8945deab770972ae56
title: 하노이 탑
challengeType: 1
forumTopicId: 302341
dashedName: towers-of-hanoi
---

# --description--

하노이 탑 문제를 해결하세요. 물체의 개수는 첫 번째 매개변수로 주어지고, 이어서 각 물체 더미를 식별하는 데 사용되는 문자열이 옵니다. 이동 목록을 포함하는 중첩된 배열을 만드세요, `["source", "destination"]`.
 
예를 들어, 매개변수 `(4, 'A', 'B', 'C')`는 이동의 중첩 배열 `[['A', 'C'], ['B', 'A']]`를 생성하는데, 이는 첫 번째 이동이 더미 `A`에서 `C`로 물체를 옮긴 것이고 두 번째 이동이 더미 `B`에서 `A`로 물체를 옮긴 것임을 나타냅니다.

# --instructions--

물체를 쌓기 위한 이동을 중첩 배열로 반환하는 함수를 작성하세요.

# --before-each--

```js
const res3Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B']];
const res7First10Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B'], ['A', 'C'], ['B', 'C'], ['B', 'A']];

function getTowerResults_() {
  return {
    res3: towerOfHanoi(3, 'A', 'B', 'C'),
    res5: towerOfHanoi(5, 'X', 'Y', 'Z'),
    res7: towerOfHanoi(7, 'A', 'B', 'C')
  };
}
```

# --hints--

`towerOfHanoi`는 함수여야 합니다.

```js
assert(typeof towerOfHanoi === 'function');
```

`towerOfHanoi(3, ...)`는 7번의 이동을 반환해야 합니다.

```js
assert(getTowerResults_().res3.length === 7);
```

`towerOfHanoi(3, 'A', 'B', 'C')`는 `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B']]`를 반환해야 합니다.

```js
assert.deepEqual(getTowerResults_().res3, res3Moves);
```

`towerOfHanoi(5, "X", "Y", "Z")`의 10번째 이동은 Y -> X여야 합니다.

```js
const { res5 } = getTowerResults_();
assert.deepEqual(res5[9], ['Y', 'X']);
```

`towerOfHanoi(7, 'A', 'B', 'C')`의 처음 10번 이동은 `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B'], ['A','C'], ['B','C'], ['B','A']]`이어야 합니다.

```js
const { res7 } = getTowerResults_();
assert.deepEqual(res7.slice(0, 10), res7First10Moves);
```

# --seed--

## --seed-contents--

```js
function towerOfHanoi(n, a, b, c) {

  return [[]];
}
```

# --solutions--

```js
function towerOfHanoi(n, a, b, c) {
  const res = [];
  towerOfHanoiHelper(n, a, c, b, res);
  return res;
}

function towerOfHanoiHelper(n, a, b, c, res) {
  if (n > 0) {
    towerOfHanoiHelper(n - 1, a, c, b, res);
    res.push([a, c]);
    towerOfHanoiHelper(n - 1, b, a, c, res);
  }
}
```
