---
id: 595608ff8bcd7a50bd490181
title: 헤일스톤 수열
challengeType: 1
forumTopicId: 302279
dashedName: hailstone-sequence
---

# --description--

헤일스톤 수열은 시작하는 양의 정수 `n`에서 다음과 같이 생성할 수 있습니다:

- `n`이 `1`이면 수열이 종료됩니다
- `n`이 `even`이면 수열의 다음 `n`는 `= n/2`입니다
- `n`이 `odd`이면 수열의 다음 `n`는 `= (3 * n) + 1`입니다

(증명되지 않은) 콜라츠 추측은 어떤 시작 숫자에 대해서도 헤일스톤 수열이 항상 종료된다는 것입니다.

헤일스톤 수열은 보통 구름 속의 우박처럼 여러 번 내려갔다 올라갔다 하는 값이기 때문에 헤일스톤 수라고도 하며, 콜라츠 수열이라고도 합니다.

# --instructions--

1. 숫자에 대한 헤일스톤 수열을 생성하는 루틴을 만드세요
2. 함수는 `limit`보다 작은 숫자 중 가장 긴 헤일스톤 수열을 가진 숫자와 그 수열의 길이를 배열로 반환해야 합니다. (실제 수열은 보여주지 마세요!)

# --hints--

`hailstoneSequence`는 함수여야 합니다.

```js
assert(typeof hailstoneSequence === 'function');
```

`hailstoneSequence(30)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(hailstoneSequence(30)));
```

`hailstoneSequence(30)`는 `[27, 112]`를 반환해야 합니다.

```js
assert.deepEqual(hailstoneSequence(30), [27, 112]);
```

`hailstoneSequence(50000)`는 `[35655, 324]`를 반환해야 합니다.

```js
assert.deepEqual(hailstoneSequence(50000), [35655, 324]);
```

`hailstoneSequence(100000)`는 `[77031, 351]`를 반환해야 합니다.

```js
assert.deepEqual(hailstoneSequence(100000), [77031, 351]);
```

# --seed--

## --seed-contents--

```js
function hailstoneSequence(limit) {
  const res = [];


  return res;
}
```

# --solutions--

```js
function hailstoneSequence (limit) {
  function hailstone(n) {
    const seq = [n];
    while (n > 1) {
      n = n % 2 ? 3 * n + 1 : n / 2;
      seq.push(n);
    }
    return seq;
  }

  let n = 0;
  let max = 0;
  for (let i = limit; --i;) {
    const seq = hailstone(i);
    const sLen = seq.length;

    if (sLen > max) {
      n = i;
      max = sLen;
    }
  }

  return [n, max];
}
```
