---
id: 595671d4d2cdc305f0d5b36f
title: 두 배열로 해시 만들기
challengeType: 1
forumTopicId: 302283
dashedName: hash-from-two-arrays
---

# --description--

두 개의 배열이 주어졌을 때, 한 배열의 요소들(키)이 다른 배열의 요소들(값)과 연결된 해시 객체를 만드세요.

예시:

```js
const firstArr = [1, 2, 3];
const secondArr = ["a", "b", "c"];
const hashObj = arrToObj(firstArr, secondArr);
```

위 배열들로부터 만들어진 해시 객체는 `{1: "a", 2: "b", 3: "c"}`입니다.

# --before-each--

```js
const testCases = [
  [[1, 2, 3, 4, 5], ['a', 'b', 'c', 'd', 'e']],
  [[1, 2, 3, 4, 5], ['a', 'b', 'c', 'd']],
  [[1, 2, 3], ['a', 'b', 'c', 'd', 'e']],
  [['a', 'b', 'c', 'd', 'e'], [1, 2, 3, 4, 5]],
  [['a', 'b', 'c', 'd', 'e'], [1, 2, 3, 4]],
  [['a', 'b', 'c'], [1, 2, 3, 4, 5]]
];

const res = [
  { 1: 'a', 2: 'b', 3: 'c', 4: 'd', 5: 'e' },
  { 1: 'a', 2: 'b', 3: 'c', 4: 'd', 5: undefined },
  { 1: 'a', 2: 'b', 3: 'c' },
  { a: 1, b: 2, c: 3, d: 4, e: 5 },
  { a: 1, b: 2, c: 3, d: 4, e: undefined },
  { a: 1, b: 2, c: 3 }
];
```

# --hints--

`arrToObj`은 함수여야 합니다.

```js
assert(typeof arrToObj === 'function');
```

`arrToObj([1, 2, 3, 4, 5], ["a", "b", "c", "d", "e"])`는 `{ 1: "a", 2: "b", 3: "c", 4: "d", 5: "e" }`을 반환해야 합니다.

```js
assert.deepEqual(arrToObj(...testCases[0]), res[0]);
```

`arrToObj([1, 2, 3, 4, 5], ["a", "b", "c", "d"])`는 `{ 1: "a", 2: "b", 3: "c", 4: "d", 5: undefined }`를 반환해야 합니다.

```js
assert.deepEqual(arrToObj(...testCases[1]), res[1]);
```

`arrToObj([1, 2, 3], ["a", "b", "c", "d", "e"])`는 `{ 1: "a", 2: "b", 3: "c" }`를 반환해야 합니다.

```js
assert.deepEqual(arrToObj(...testCases[2]), res[2]);
```

`arrToObj(["a", "b", "c", "d", "e"], [1, 2, 3, 4, 5])`는 `{ "a": 1, "b": 2, "c": 3 , "d": 4, "e": 5 }`를 반환해야 합니다.

```js
assert.deepEqual(arrToObj(...testCases[3]), res[3]);
```

`arrToObj(["a", "b", "c", "d", "e"], [1, 2, 3, 4])`는 `{ "a": 1, "b": 2, "c": 3 , "d": 4, "e": undefined }`를 반환해야 합니다.

```js
assert.deepEqual(arrToObj(...testCases[4]), res[4]);
```

`arrToObj(["a", "b", "c"], [1, 2, 3, 4, 5])`는 `{ "a": 1, "b": 2, "c": 3 }`를 반환해야 합니다.

```js
assert.deepEqual(arrToObj(...testCases[5]), res[5]);
```

# --seed--

## --seed-contents--

```js
function arrToObj (keys, vals) {

  return true;
}
```

# --solutions--

```js
function arrToObj (keys, vals) {
  return keys.reduce((map, key, index) => {
    map[key] = vals[index];
    return map;
  }, {});
}
```
