---
id: 595668ca4cfe1af2fb9818d4
title: 하샤드 또는 니븐 수열
challengeType: 1
forumTopicId: 302281
dashedName: harshad-or-niven-series
---

# --description--

하샤드 또는 니븐 수는 각 자리 숫자의 합으로 나누어지는 1 이상의 양의 정수입니다.

예를 들어, `42`는 `42`가 `(4 + 2)`로 나누어떨어지므로 하샤드 수입니다.

수열은 오름차순으로 정의된다고 가정하세요.

# --instructions--

하샤드 수열의 연속된 항을 생성하는 함수를 구현하세요.

이를 사용해 `n`보다 큰 첫 번째 하샤드 수부터 시작하는 수열의 10개 항을 배열로 반환하세요.

# --hints--

`isHarshadOrNiven`는 함수여야 합니다.

```js
assert(typeof isHarshadOrNiven === 'function');
```

`isHarshadOrNiven(10)`는 `[12, 18, 20, 21, 24, 27, 30, 36, 40, 42]`를 반환해야 합니다.

```js
assert.deepEqual(isHarshadOrNiven(10), [12, 18, 20, 21, 24, 27, 30, 36, 40, 42]);
```

`isHarshadOrNiven(400)`는 `[402, 405, 407, 408, 410, 414, 420, 423, 432, 440]`를 반환해야 합니다.

```js
assert.deepEqual(isHarshadOrNiven(400), [402, 405, 407, 408, 410, 414, 420, 423, 432, 440]);
```

`isHarshadOrNiven(1000)`는 `[1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]`를 반환해야 합니다.

```js
assert.deepEqual(isHarshadOrNiven(1000), [1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]);
```

# --seed--

## --seed-contents--

```js
function isHarshadOrNiven(n) {
  const res = [];

  return res;
}
```

# --solutions--

```js
function isHarshadOrNiven(n) {
  function isHarshad(num) {
    let s = 0;
    const nStr = num.toString();
    for (let i = 0; i < nStr.length; ++i) {
      s += parseInt(nStr.charAt(i), 10);
    }
    return n % s === 0;
  }

  const res = [];
  let count = 0;

  while (count < 10) {
    n++;
    if (isHarshad(n)) {
      count++;
      res.push(n);
    }
  }

  return res;
}
```
