---
id: 59622f89e4e137560018a40e
title: 호프스태터 그림-그림 수열
challengeType: 1
forumTopicId: 302286
dashedName: hofstadter-figure-figure-sequences
---

# --description--

호프스태터 그림-그림 수열 $R_n$과 $S_n$은 다음과 같이 정의됩니다.

$R_1 = 1\\ ;\\ S_1 = 2 \\\\R_n = R_{n-1} + S_{n-1}, \\quad n>1.$

특히, 수열 $R_n$은 다음 값을 포함합니다.

<pre>1, 3, 7, 12, 18, ...</pre>

그리고 수열 $S_n$은 다음 값을 포함합니다.

<pre>2, 4, 5, 6, 8, ...</pre>

수열 $R_n$은 점화식 $R_n = R_{n-1} + S_{n-1}$으로 정의되고, $S_n$은 수열 $R_n$에 포함되지 않는 양의 정수로 이루어진 수열입니다.

# --instructions--

`ffr`와 `ffs`이라는 이름의 두 함수를 만들어서, 임의의 인덱스 `R(n)`에 대해 각각 `S(n)` 또는 `n`를 반환하세요. 호프스태터 그림-그림 수열은 1부터 시작하며, $R_1 = 1$과 $S_1 = 2$임을 유의하세요.

`n`에 대한 최대값은 가정하지 마세요.

**참조**

 <p>Rosetta: <a href='https://rosettacode.org/wiki/Hofstadter_Figure-Figure_sequences' target='_blank'>Hofstadter Figure-Figure sequences</a></p>.


# --before-each--

```js
const ffrParamRes = [[10, 69], [50, 1509], [100, 5764], [1000, 526334]];
const ffsParamRes = [[10, 14], [50, 59], [100, 112], [1000, 1041]];
```

# --hints--

`ffr`는 함수여야 합니다.

```js
assert(typeof ffr === 'function');
```

`ffs`는 함수여야 합니다.

```js
assert(typeof ffs === 'function');
```

`ffr`는 정수를 반환해야 합니다.

```js
assert(Number.isInteger(ffr(1)));
```

`ffs`는 정수를 반환해야 합니다.

```js
assert(Number.isInteger(ffs(1)));
```

`ffr(10)`는 `69`를 반환해야 합니다.

```js
assert.equal(ffr(ffrParamRes[0][0]), ffrParamRes[0][1]);
```

`ffr(50)`는 `1509`를 반환해야 합니다.

```js
assert.equal(ffr(ffrParamRes[1][0]), ffrParamRes[1][1]);
```

`ffr(100)`는 `5764`를 반환해야 합니다.

```js
assert.equal(ffr(ffrParamRes[2][0]), ffrParamRes[2][1]);
```

`ffr(1000)`는 `526334`를 반환해야 합니다.

```js
assert.equal(ffr(ffrParamRes[3][0]), ffrParamRes[3][1]);
```

`ffs(10)`는 `14`를 반환해야 합니다.

```js
assert.equal(ffs(ffsParamRes[0][0]), ffsParamRes[0][1]);
```

`ffs(50)`는 `59`를 반환해야 합니다.

```js
assert.equal(ffs(ffsParamRes[1][0]), ffsParamRes[1][1]);
```

`ffs(100)`는 `112`를 반환해야 합니다.

```js
assert.equal(ffs(ffsParamRes[2][0]), ffsParamRes[2][1]);
```

`ffs(1000)`는 `1041`를 반환해야 합니다.

```js
assert.equal(ffs(ffsParamRes[3][0]), ffsParamRes[3][1]);
```

# --seed--

## --seed-contents--

```js
function ffr(n) {
  return n;
}

function ffs(n) {
  return n;
}
```

# --solutions--

```js
const R = [null, 1];
const S = [null, 2];

function extendSequences (n) {
  let current = Math.max(R[R.length - 1], S[S.length - 1]);
  let i;
  while (R.length <= n || S.length <= n) {
    i = Math.min(R.length, S.length) - 1;
    current += 1;
    if (current === R[i] + S[i]) {
      R.push(current);
    } else {
      S.push(current);
    }
  }
}

function ffr (n) {
  extendSequences(n);
  return R[n];
}

function ffs (n) {
  extendSequences(n);
  return S[n];
}
```
