---
id: 59637c4d89f6786115efd814
title: 호프스태터 Q 수열
challengeType: 1
forumTopicId: 302287
dashedName: hofstadter-q-sequence
---

# --description--

호프스태터 Q 수열은 다음과 같이 정의됩니다:

$Q(1)=Q(2)=1, \\\\ Q(n)=Q\\big(n-Q(n-1)\\big)+Q\\big(n-Q(n-2)), \\quad n>2.$

피보나치 수열처럼 정의되지만, 피보나치 수열에서 다음 항은 이전 두 항의 합인 반면, Q 수열에서는 이전 두 항이 Q 수열에서 얼마나 뒤로 가야 하는지를 알려주어 그 두 수를 더해 다음 항을 만듭니다.

# --instructions--

호프스태터 Q 수열 공식을 함수로 구현하세요. 함수는 숫자 `n`를 입력받아 정수를 반환해야 합니다.

# --before-each--

```js
const testCase = [1000, 1500, 2000, 2500];
const res = [502, 755, 1005, 1261];
```

# --hints--

`hofstadterQ`은 함수여야 합니다.

```js
assert(typeof hofstadterQ === 'function');
```

`hofstadterQ()`는 `integer`를 반환해야 합니다.

```js
assert(Number.isInteger(hofstadterQ(1000)));
```

`hofstadterQ(1000)`는 `502`를 반환해야 합니다.

```js
assert.equal(hofstadterQ(testCase[0]), res[0]);
```

`hofstadterQ(1500)`는 `755`를 반환해야 합니다.

```js
assert.equal(hofstadterQ(testCase[1]), res[1]);
```

`hofstadterQ(2000)`는 `1005`를 반환해야 합니다.

```js
assert.equal(hofstadterQ(testCase[2]), res[2]);
```

`hofstadterQ(2500)`는 `1261`를 반환해야 합니다.

```js
assert.equal(hofstadterQ(testCase[3]), res[3]);
```

# --seed--

## --seed-contents--

```js
function hofstadterQ(n) {

  return n;
}
```

# --solutions--

```js
function hofstadterQ (n) {
  const memo = [1, 1, 1];
  const Q = function (i) {
    let result = memo[i];
    if (typeof result !== 'number') {
      result = Q(i - Q(i - 1)) + Q(i - Q(i - 2));
      memo[i] = result;
    }
    return result;
  };
  return Q(n);
}
```
