---
id: 59669d08d75b60482359409f
title: 날짜 형식
challengeType: 1
forumTopicId: 302243
dashedName: date-format
---

# --description--

다음 조건에 맞는 현재 날짜의 두 개 날짜 문자열을 포함하는 배열을 반환하세요:

- 첫 번째 문자열의 날짜 순서는 연도, 월, 일이 대시(`-`)로 구분되어야 합니다.
- 첫 번째 문자열의 연도는 네 자리 숫자여야 합니다.
- 첫 번째 문자열의 월과 일은 앞에 0이 없어야 합니다.
- 두 번째 문자열의 요일과 월 이름은 축약하지 않아야 합니다.
- 두 번째 문자열의 일은 앞에 0이 없어야 합니다.

예시 출력:

```js
['2007-11-23', 'Friday, November 23, 2007']
['2021-3-2', 'Tuesday, March 2, 2021']
```

# --hints--

`getDateFormats`은 함수여야 합니다.

```js
assert(typeof getDateFormats === 'function');
```

`getDateFormats`는 객체를 반환해야 합니다.

```js
assert(typeof getDateFormats() === 'object');
```

`getDateFormats`는 2개의 요소를 가진 배열을 반환해야 합니다.

```js
assert(getDateFormats().length === 2);
```

`getDateFormats`는 올바른 형식의 정확한 날짜를 반환해야 합니다.

```js
const getDateSolution = () => {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
};

const dates = getDateSolution();
const equalsMessage = `message: <code>getDataFormats()</code> should return <code>["${dates[0]}", "${dates[1]}"]</code>.`;
assert.deepEqual(getDateFormats(), dates, equalsMessage);
```

# --seed--

## --seed-contents--

```js
function getDateFormats() {

  return true;
}
```

# --solutions--

```js
function getDateFormats() {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
}
```
