---
id: 59667989bf71cf555dd5d2ff
title: S-표현식
challengeType: 1
forumTopicId: 302303
dashedName: s-expressions
---

# --description--

<a href="https://rosettacode.org/wiki/S-expressions" target="_blank" rel="noopener noreferrer nofollow">S-표현식</a>은 데이터를 파싱하고 저장하는 편리한 방법 중 하나입니다.

# --instructions--

따옴표가 있거나 없는 문자열, 정수 및 부동소수점을 처리하는 간단한 S-표현식 리더/파서를 작성하세요.

이 함수는 문자열에서 단일이지만 중첩된 S-표현식을 읽고 (중첩된) 배열로 반환해야 합니다.

줄바꿈과 기타 공백은 따옴표로 묶인 문자열 안에 포함되어 있지 않으면 무시할 수 있습니다.

따옴표로 묶인 문자열 안의 "`()`"는 해석되지 않고 문자열의 일부로 처리됩니다.

문자열 안에서 이스케이프된 따옴표를 처리하는 것은 선택 사항입니다; 따라서 "`(foo"bar)`"는 문자열 "`foo"bar`"로 처리하거나 오류로 처리할 수 있습니다.

이를 위해 리더는 `\`를 이스케이프 용도로 인식할 필요는 없지만, 언어에 적절한 자료형이 있다면 숫자를 인식해야 합니다.

`()"` (이스케이프가 지원된다면 `\`)와 공백을 제외하면 특별한 문자가 없다는 점에 유의하세요. 그 외의 모든 것은 따옴표 없이 허용됩니다.

리더는 다음 입력을 읽을 수 있어야 합니다.

<pre>((data "quoted data" 123 4.5)
(data (!@# (4.5) "(more" "data)")))
</pre>

그리고 이를 네이티브 자료 구조로 변환해야 합니다.

# --before-each--

```js
const simpleSExpr = '(data1 data2 data3)';
const simpleSolution = ['data1', 'data2', 'data3'];

const basicSExpr = '((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))';
const basicSolution = [["data","\"quoted data\"",123,4.5],["data",["!@#",[4.5],"\"(more\"","\"data)\""]]];
```

# --hints--

`parseSexpr`는 함수여야 합니다.

```js
assert(typeof parseSexpr === 'function');
```

`parseSexpr('(data1 data2 data3)')`는 `['data1', 'data2', 'data3']`를 반환해야 합니다.

```js
assert.deepEqual(parseSexpr(simpleSExpr), simpleSolution);
```

`parseSexpr('((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))')`는 `[['data', '"quoted data"', 123, 4.5], ['data', ['!@#', [4.5], '"(more"', '"data)"']]]`를 반환해야 합니다.

```js
assert.deepEqual(parseSexpr(basicSExpr), basicSolution);
```

# --seed--

## --seed-contents--

```js
function parseSexpr(str) {

  return true;
}
```

# --solutions--

```js
function parseSexpr(str) {
  const t = str.match(/\s*("[^"]*"|\(|\)|"|[^\s()"]+)/g);
  for (var o, c = 0, i = t.length - 1; i >= 0; i--) {
    var n,
      ti = t[i].trim();
    if (ti == '"') return;
    else if (ti == '(') t[i] = '[', c += 1;
    else if (ti == ')') t[i] = ']', c -= 1;
    else if ((n = +ti) == ti) t[i] = n;
    else t[i] = `'${ti.replace('\'', '\\\'')}'`;
    if (i > 0 && ti != ']' && t[i - 1].trim() != '(') t.splice(i, 0, ',');
    if (!c) if (!o) o = true; else return;
  }
  return c ? undefined : eval(t.join(''));
}
```
