---
id: 5966f99c45e8976909a85575
title: 요일
challengeType: 1
forumTopicId: 302245
dashedName: day-of-the-week
---

# --description--

한 회사는 크리스마스가 일요일인 경우, 모든 직원에게 추가 유급 휴가를 주어 공휴일과 합쳐서 12월 25일부터 1월 1일 사이에 다음 주에 근무하지 않도록 하기로 결정했습니다.

# --instructions--

시작 연도와 종료 연도를 입력받아 12월 25일이 일요일인 모든 연도를 배열로 반환하는 함수를 작성하세요.

# --before-each--

```js
const firstSolution = [1977, 1983, 1988, 1994, 2005, 2011, 2016];
const secondSolution = [2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118];
```

# --hints--

`findXmasSunday`는 함수여야 합니다.

```js
assert(typeof findXmasSunday === 'function');
```

`findXmasSunday(2000, 2100)`는 배열을 반환해야 합니다.

```js
assert(typeof findXmasSunday(2000, 2100) === 'object');
```

`findXmasSunday(1970, 2017)`는 `[1977, 1983, 1988, 1994, 2005, 2011, 2016]`를 반환해야 합니다.

```js
assert.deepEqual(findXmasSunday(1970, 2017), firstSolution);
```

`findXmasSunday(2008, 2121)`는 `[2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118]`를 반환해야 합니다.

```js
assert.deepEqual(findXmasSunday(2008, 2121), secondSolution);
```

# --seed--

## --seed-contents--

```js
function findXmasSunday(start, end) {

  return true;
}
```

# --solutions--

```js
function findXmasSunday(start, end) {
  const xmasSunday = [];
  for (let year = start; year <= end; year++) {
    const xmas = new Date(year, 11, 25);
    if (xmas.getDay() === 0) {
      xmasSunday.push(year);
    }
  }
  return xmasSunday;
}
```
