---
id: 59694356a6e7011f7f1c5f4e
title: FreeCell 카드 나누기
challengeType: 1
forumTopicId: 302246
dashedName: deal-cards-for-freecell
---

# --description--

*FreeCell*은 1978년 Paul Alfille이 PLATO 시스템에 도입한 솔리테어 카드 게임입니다. Microsoft의 Jim Horne이 이름을 FreeCell로 바꾸고 DOS와 Windows용으로 게임을 다시 구현했습니다. 이 버전은 32000개의 번호가 매겨진 딜을 도입했습니다.

게임이 인기를 얻으면서 Jim Horne은 알고리즘을 공개했고, 다른 FreeCell 구현체들이 Microsoft 딜을 재현하기 시작했습니다. 이 딜들은 1부터 32000까지 번호가 매겨져 있습니다. Microsoft의 최신 버전은 1백만 개의 딜을 가지고 있으며, 1부터 1000000까지 번호가 매겨져 있습니다; 일부 구현체는 이 범위를 벗어난 번호도 허용합니다.

알고리즘은 Microsoft C의 선형 합동 생성기를 사용합니다:

<ul>
  <li>$state_{n + 1} \equiv 214013 \times state_n + 2531011 \pmod{2^{31}}$</li>
  <li>$rand_n = state_n \div 2^{16}$</li>
  <li>$rand_n$는 0부터 32767 범위에 있습니다.</li>
</ul>

알고리즘은 다음과 같습니다:

<ol>
  <li>딜 번호로 RNG를 시드합니다.
  </li><li>52장의 카드 배열을 만듭니다: 클럽 에이스, 다이아몬드 에이스, 하트 에이스, 스페이드 에이스, 클럽 2, 다이아몬드 2, 그리고 순서대로 랭크: 에이스, 2, 3, 4, 5, 6, 7, 8, 9, 10, 잭, 퀸, 킹. 배열 인덱스는 0부터 51까지이며, 클럽 에이스가 0, 스페이드 킹이 51입니다.</li>
  <li>배열이 빌 때까지:</li>
    <ul>
      <li>무작위 카드의 <i>인덱스</i> ≡ <i>다음 무작위 수</i> (mod <i>배열 길이</i>)를 선택합니다.</li>
      <li>이 무작위 카드를 배열의 마지막 카드와 교환합니다.</li>
      <li>이 무작위 카드를 배열에서 제거합니다. (배열 길이가 1 줄어듭니다.)</li>
      <li>이 무작위 카드를 나눕니다.</li>
    </ul>
  <li>52장 카드를 모두 앞면이 보이도록 8열에 나눕니다. 처음 8장은 8열에 각각 놓이고, 다음 8장은 첫 8장 위에 놓이며, 이런 식으로 계속됩니다.</li>
</ol>

**예시:**

**카드 나누기 순서**

<pre> 1  2  3  4  5  6  7  8
 9 10 11 12 13 14 15 16
17 18 19 20 21 22 23 24
25 26 27 28 29 30 31 32
33 34 35 36 37 38 39 40
41 42 43 44 45 46 47 48
49 50 51 52</pre>

**게임 #1**

```js
[
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
]
```

**게임 #617**

```js
[
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
]
```

# --instructions--

딜 번호를 받아 이 알고리즘과 같은 순서로 카드를 나누는 함수를 작성하세요. 함수는 FreeCell 보드를 나타내는 2차원 배열을 반환해야 합니다.

# --before-each--

```js
const replaceThis = 3;
const game1 = [
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
];
const game617 = [
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
];
```

# --hints--

`dealFreeCell`는 함수여야 합니다.

```js
assert(typeof dealFreeCell === 'function');
```

`dealFreeCell(seed)`는 객체를 반환해야 합니다.

```js
assert(typeof dealFreeCell(1) === 'object');
```

`dealFreeCell(seed)`는 길이 7인 배열을 반환해야 합니다.

```js
assert(dealFreeCell(1).length === 7);
```

`dealFreeCell(1)`는 예시 "게임 #1"과 동일한 배열을 반환해야 합니다.

```js
assert.deepEqual(dealFreeCell(1), game1);
```

`dealFreeCell(617)`는 예시 "게임 #617"과 동일한 배열을 반환해야 합니다.

```js
assert.deepEqual(dealFreeCell(617), game617);
```

# --seed--

## --seed-contents--

```js
function dealFreeCell(seed) {

  return true;
}
```

# --solutions--

```js
// RNG
function FreeCellRNG(seed) {
  return {
    lastNum: seed,
    next() {
      this.lastNum = ((214013 * this.lastNum) + 2531011) % (Math.pow(2, 31));
      return this.lastNum >> 16;
    }
  };
}
// Get cards
function getDeck() {
  const ranks = ['A', '2', '3', '4', '5', '6', '7', '8', '9', 'T', 'J', 'Q', 'K'];
  const suits = ['C', 'D', 'H', 'S'];
  const cards = [];
  for (let i = 0; i < ranks.length; i += 1) {
    for (let j = 0; j < suits.length; j += 1) {
      cards.push(`${ranks[i]}${suits[j]}`);
    }
  }
  return cards;
}
function dealFreeCell(seed) {
  const rng = FreeCellRNG(seed);
  const deck = getDeck();

  const dealtCards = [[], [], [], [], [], [], []];
  let currentColumn = 0;
  let currentRow = 0;

  let rand;
  let temp;
  let card;
  while (deck.length > 0) {
    // Choose a random card
    rand = rng.next() % deck.length;

    // Swap this random card with the last card in the array
    temp = deck[deck.length - 1];
    deck[deck.length - 1] = deck[rand];
    deck[rand] = temp;

    // Remove this card from the array
    card = deck.pop();

    // Deal this card
    dealtCards[currentRow].push(card);
    currentColumn += 1;
    if (currentColumn === 8) {
      currentColumn = 0;
      currentRow += 1;
    }
  }

  return dealtCards;
}
```
