---
id: 596a8888ab7c01048de257d5
title: 깊은 복사
challengeType: 1
forumTopicId: 302247
dashedName: deepcopy
---

# --description--

주어진 객체의 깊은 복사본을 반환하는 함수를 작성하세요. 복사본은 주어진 객체와 동일한 객체가 아니어야 합니다.

이 도전 과제는 다음을 테스트하지 않습니다:

<ul>
  <li>프로퍼티가 함수인 객체</li>
  <li>날짜 객체 또는 프로퍼티가 날짜 객체인 객체</li>
  <li>정규식 또는 프로퍼티가 정규식 객체인 객체</li>
  <li>프로토타입 복사</li>
</ul>

# --before-each--

```js
const obj1 = { test: 'test' };
const obj2 = {
  t: 'test',
  a: ['an', 'array']
};
const obj3 = {
  t: 'try',
  o: obj2
};
```

# --hints--

`deepcopy`는 함수여야 합니다.

```js
assert(typeof deepcopy === 'function');
```

`deepcopy({test: "test"})`는 객체를 반환해야 합니다.

```js
assert(typeof deepcopy(obj1) === 'object');
```

`deepcopy`는 제공된 객체와 동일한 객체를 반환하면 안 됩니다.

```js
assert(deepcopy(obj2) != obj2);
```

배열을 포함하는 객체가 전달되면 `deepcopy`는 객체의 깊은 복사본을 반환해야 합니다.

```js
assert.deepEqual(deepcopy(obj2), obj2);
```

다른 객체를 포함하는 객체가 전달되면 `deepcopy`는 객체의 깊은 복사본을 반환해야 합니다.

```js
assert.deepEqual(deepcopy(obj3), obj3);
```

# --seed--

## --seed-contents--

```js
function deepcopy(obj) {

  return true;
}
```

# --solutions--

```js
function deepcopy(obj) {
  return JSON.parse(JSON.stringify(obj));
}
```
