---
id: 596e414344c3b2872167f0fe
title: 쉼표 다툼
challengeType: 1
forumTopicId: 302234
dashedName: comma-quibbling
---

# --description--

<a href="https://rosettacode.org/wiki/Comma_quibbling" target="_blank" rel="noopener noreferrer nofollow">쉼표 다툼</a>은 Eric Lippert가 자신의 블로그에서 처음 제시한 도전 과제입니다.

# --instructions--

다음 조건에 따라 리스트/수열에서 입력된 단어들을 이어 붙인 문자열 출력을 생성하는 함수를 작성하세요:

<ol>
  <li>단어가 하나도 없으면 출력 문자열은 중괄호 두 개만 포함합니다 (<code>"{}"</code>)</li>
  <li>단어가 하나만 있으면, 예를 들어 <code>["ABC"]</code>라면, 출력 문자열은 중괄호 안에 그 단어가 들어갑니다, 예: <code>"{ABC}"</code></li>
  <li>단어가 두 개면, 예를 들어 <code>["ABC", "DEF"]</code>라면, 출력 문자열은 중괄호 안에 두 단어가 <code>" and "</code>로 구분되어 들어갑니다, 예: <code>"{ABC and DEF}"</code></li>
  <li>단어가 세 개 이상이면, 예를 들어 <code>["ABC", "DEF", "G", "H"]</code>라면, 출력 문자열은 마지막 단어를 제외한 모든 단어가 <code>", "</code>로 구분되고 마지막 단어는 <code>" and "</code>로 구분되어 중괄호 안에 들어갑니다; 예: <code>"{ABC, DEF, G and H}"</code></li>
</ol>

다음 입력 시리즈로 함수를 테스트하고 이 페이지에 출력 결과를 보여주세요:

<ul>
  <li>[] # (단어가 없는 경우)</li>
  <li>["ABC"]</li>
  <li>["ABC", "DEF"]</li>
  <li>["ABC", "DEF", "G", "H"]</li>
</ul>

**참고:** 이 도전 과제에서는 단어가 모두 대문자 문자열이며 빈 문자열이 아니라고 가정하세요.

# --before-each--

```js
const testCases = [[], ["ABC"], ["ABC", "DEF"], ["ABC", "DEF", "G", "H"]];
const results = ["{}", "{ABC}", "{ABC and DEF}", "{ABC, DEF, G and H}"];
```

# --hints--

`quibble`는 함수여야 합니다.

```js
assert(typeof quibble === 'function');
```

`quibble(["ABC"])`는 문자열을 반환해야 합니다.

```js
assert(typeof quibble(['ABC']) === 'string');
```

`quibble([])`는 "{}"를 반환해야 합니다.

```js
assert.equal(quibble(testCases[0]), results[0]);
```

`quibble(["ABC"])`는 `"{ABC}"`를 반환해야 합니다.

```js
assert.equal(quibble(testCases[1]), results[1]);
```

`quibble(["ABC", "DEF"])`는 `"{ABC and DEF}"`를 반환해야 합니다.

```js
assert.equal(quibble(testCases[2]), results[2]);
```

`quibble(["ABC", "DEF", "G", "H"])`는 `"{ABC, DEF, G and H}"`를 반환해야 합니다.

```js
assert.equal(quibble(testCases[3]), results[3]);
```

# --seed--

## --seed-contents--

```js
function quibble(words) {

  return true;
}
```

# --solutions--

```js
function quibble(words) {
  return "{" +
    words.slice(0, words.length - 1).join(", ") +
   (words.length > 1 ? " and " : "") +
   (words[words.length - 1] || '') +
  "}";
}
```
