---
id: 596e457071c35c882915b3e4
title: 문자열 목록 비교
challengeType: 1
forumTopicId: 302235
dashedName: compare-a-list-of-strings
---

# --description--

목록은 중복을 포함할 수 있는 순서가 있는 값들의 집합입니다. 예시는 다음과 같습니다:

```js
const list = [['AA',  'BB', 'CC'], ['AA', 'ACB', 'AA'], [], ['AA']];
````

임의 개수의 문자열 목록이 주어졌을 때, 다음 조건 각각에 대해 함수를 구현하세요:

<ul>
  <li>모든 문자열이 사전식으로 같은지 테스트하세요</li>
  <li>각 문자열이 그 다음 문자열보다 사전식으로 작은지 테스트하세요 (즉, 목록이 엄격한 오름차순인지 확인)</li>
</ul>

# --before-each--

```js
const testCases = [['AA', 'AA', 'AA', 'AA'], ['AA', 'ACB', 'BB', 'CC'], [], ['AA'], ['BB', 'AA']];
```

# --hints--

`allEqual`는 함수여야 합니다.

```js
assert(typeof allEqual === 'function');
```

`azSorted`는 함수여야 합니다.

```js
assert(typeof azSorted === 'function');
```

`allEqual(["AA", "AA", "AA", "AA"])`는 true를 반환해야 합니다.

```js
assert(allEqual(testCases[0]));
```

`azSorted(["AA", "AA", "AA", "AA"])`는 false를 반환해야 합니다.

```js
assert(!azSorted(testCases[0]));
```

`allEqual(["AA", "ACB", "BB", "CC"])`는 false를 반환해야 합니다.

```js
assert(!allEqual(testCases[1]));
```

`azSorted(["AA", "ACB", "BB", "CC"])`는 true를 반환해야 합니다.

```js
assert(azSorted(testCases[1]));
```

`allEqual([])`는 true를 반환해야 합니다.

```js
assert(allEqual(testCases[2]));
```

`azSorted([])`는 true를 반환해야 합니다.

```js
assert(azSorted(testCases[2]));
```

`allEqual(["AA"])`는 true를 반환해야 합니다.

```js
assert(allEqual(testCases[3]));
```

`azSorted(["AA"])`는 true를 반환해야 합니다.

```js
assert(azSorted(testCases[3]));
```

`allEqual(["BB", "AA"])`는 false를 반환해야 합니다.

```js
assert(!allEqual(testCases[4]));
```

`azSorted(["BB", "AA"])`는 false를 반환해야 합니다.

```js
assert(!azSorted(testCases[4]));
```

# --seed--

## --seed-contents--

```js
function allEqual(arr) {

  return true;
}

function azSorted(arr) {

  return true;
}
```

# --solutions--

```js
function allEqual(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] === a[i]);
  } return out;
}

function azSorted(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] < a[i]);
  } return out;
}
```
