---
id: 596fda99c69f779975a1b67d
title: 부분 문자열의 발생 횟수 세기
challengeType: 1
forumTopicId: 302237
dashedName: count-occurrences-of-a-substring
---

# --description--

부분 문자열이 문자열 안에 겹치지 않고 나타나는 횟수를 세는 함수를 작성하거나 내장 함수를 보여주세요.

함수는 두 개의 인수를 받아야 합니다:

<ul>
  <li>첫 번째 인수는 검색할 문자열이고,</li>
  <li>두 번째는 검색할 부분 문자열입니다.</li>
</ul>

정수형 개수를 반환해야 합니다.

일치 결과는 겹치지 않는 최대 개수의 매치를 반환해야 합니다.

일반적으로 이는 왼쪽에서 오른쪽 또는 오른쪽에서 왼쪽으로 매칭하는 것을 의미합니다.

# --before-each--

```js
const testCases = ['the three truths', 'ababababab', 'abaabba*bbaba*bbab'];
const searchString = ['th', 'abab', 'a*b'];
const results = [3, 2, 2];
```

# --hints--

`countSubstring`는 함수여야 합니다.

```js
assert(typeof countSubstring === 'function');
```

`countSubstring("the three truths", "th")`는 `3`를 반환해야 합니다.

```js
assert.equal(countSubstring(testCases[0], searchString[0]), results[0]);
```

`countSubstring("ababababab", "abab")`는 `2`를 반환해야 합니다.

```js
assert.equal(countSubstring(testCases[1], searchString[1]), results[1]);
```

`countSubstring("abaabba*bbaba*bbab", "a*b")`는 `2`를 반환해야 합니다.

```js
assert.equal(countSubstring(testCases[2], searchString[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function countSubstring(str, subStr) {

  return true;
}
```

# --solutions--

```js
function countSubstring(str, subStr) {
  const escapedSubStr = subStr.replace(/[.+*?^$[\]{}()|/]/g, '\\$&');
  const matches = str.match(new RegExp(escapedSubStr, 'g'));
  return matches ? matches.length : 0;
}
```
