---
id: 596fd036dc1ab896c5db98b1
title: 초를 복합 기간으로 변환하기
challengeType: 1
forumTopicId: 302236
dashedName: convert-seconds-to-compound-duration
---

# --description--

다음 기능을 구현하세요:

<ul>
  <li>양의 정수로 된 초 단위 기간을 입력으로 받고(예: <code>100</code>),</li>
  <li>같은 기간을 주, 일, 시간, 분, 초 단위로 분해하여 문자열로 반환합니다(예: <code>1 min, 40 sec</code>).</li>
</ul>

다음 세 가지 테스트 케이스를 통과하는지 확인하세요:

<div style='font-size:115%; font-weight: bold;'>테스트 케이스</div>

| 입력 숫자    | 출력 문자열                          |
| ------------ | ------------------------------------- |
| 7259         | <code>2 hr, 59 sec</code>             |
| 86400        | <code>1 d</code>                      |
| 6000000      | <code>9 wk, 6 d, 10 hr, 40 min</code> |

<div style="font-size:115%; font-weight: bold;">세부 사항</div>
<ul>
  <li>
    다음 다섯 단위를 사용하세요:

| 단위   | 출력에 사용되는 접미사 | 변환 기준              |
| ------ | --------------------- | --------------------- |
| 주     | <code>wk</code>       | 1주 = 7일             |
| 일     | <code>d</code>        | 1일 = 24시간          |
| 시간   | <code>hr</code>       | 1시간 = 60분          |
| 분     | <code>min</code>      | 1분 = 60초            |
| 초     | <code>sec</code>      | ---                   |

  </li>
  <li>
    단, 출력에는 0이 아닌 값만 포함하세요(예: <code>1 d</code>는 반환하지만 <code>0 wk, 1 d, 0 hr, 0 min, 0 sec</code>는 반환하지 마세요).
  </li>
  <li>
    가능한 한 큰 단위를 우선으로 하세요(예: <code>2 min, 10 sec</code>는 반환하지만 <code>1 min, 70 sec</code>나 <code>130 sec</code>는 반환하지 마세요).
  </li>
  <li>
    테스트 케이스에 나온 형식을 따라 하세요(수량은 큰 단위부터 작은 단위 순으로 정렬하고, 각 단위는 쉼표+공백으로 구분하며, 값과 단위는 공백으로 구분).
  </li>
</ul>

# --before-each--

```js
const testCases = [7259, 86400, 6000000];
const results = ['2 hr, 59 sec', '1 d', '9 wk, 6 d, 10 hr, 40 min'];
```

# --hints--

`convertSeconds`는 함수여야 합니다.

```js
assert(typeof convertSeconds === 'function');
```

`convertSeconds(7259)`는 `2 hr, 59 sec`를 반환해야 합니다.

```js
assert.equal(convertSeconds(testCases[0]), results[0]);
```

`convertSeconds(86400)`는 `1 d`를 반환해야 합니다.

```js
assert.equal(convertSeconds(testCases[1]), results[1]);
```

`convertSeconds(6000000)`는 `9 wk, 6 d, 10 hr, 40 min`를 반환해야 합니다.

```js
assert.equal(convertSeconds(testCases[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function convertSeconds(sec) {

  return true;
}
```

# --solutions--

```js
function convertSeconds(sec) {
  const localNames = ['wk', 'd', 'hr', 'min', 'sec'];
  // compoundDuration :: [String] -> Int -> String
  const compoundDuration = (labels, intSeconds) =>
    weekParts(intSeconds)
    .map((v, i) => [v, labels[i]])
    .reduce((a, x) =>
      a.concat(x[0] ? [`${x[0]} ${x[1] || '?'}`] : []), []
    )
    .join(', ');

    // weekParts :: Int -> [Int]
  const weekParts = intSeconds => [0, 7, 24, 60, 60]
    .reduceRight((a, x) => {
      const r = a.rem;
      const mod = x !== 0 ? r % x : r;

      return {
        rem: (r - mod) / (x || 1),
        parts: [mod].concat(a.parts)
      };
    }, {
      rem: intSeconds,
      parts: []
    })
    .parts;

  return compoundDuration(localNames, sec);
}
```
