---
id: 597089c87eec450c68aa1643
title: 원시 데이터 타입 정의하기
challengeType: 1
forumTopicId: 302248
dashedName: define-a-primitive-data-type
---

# --description--

정수처럼 동작하지만 유효한 최솟값이 1이고 최댓값이 10인 타입을 정의하세요.

오류 처리:

<ul>
  <li><code>Num</code>을 1에서 10 범위를 벗어난 값으로 인스턴스화하려 하면, <code>'Out of range'</code>라는 오류 메시지를 가진 <code>TypeError</code>를 발생시켜야 합니다.</li>
  <li><code>Num</code>을 숫자가 아닌 값으로 인스턴스화하려 하면, <code>'Not a Number'</code>라는 오류 메시지를 가진 <code>TypeError</code>를 발생시켜야 합니다.</li>
</ul>

# --hints--

`Num`는 함수여야 합니다.

```js
assert(typeof Num === 'function');
```

`new Num(4)`은 객체를 반환해야 합니다.

```js
assert(typeof new Num(4) === 'object');
```

`new Num('test')`는 'Not a Number' 메시지를 가진 TypeError를 발생시켜야 합니다.

```js
assert.throws(() => new Num('test'), TypeError);
```

`new Num(0)`는 'Out of range' 메시지를 가진 TypeError를 발생시켜야 합니다.

```js
assert.throws(() => new Num(0), TypeError);
```

`new Num(-5)`는 'Out of range' 메시지를 가진 TypeError를 발생시켜야 합니다.

```js
assert.throws(() => new Num(-5), TypeError);
```

`new Num(10)`는 'Out of range' 메시지를 가진 TypeError를 발생시켜야 합니다.

```js
assert.throws(() => new Num(11), TypeError);
```

`new Num(20)`는 'Out of range' 메시지를 가진 TypeError를 발생시켜야 합니다.

```js
assert.throws(() => new Num(20), TypeError);
```

`new Num(3) + new Num(4)`는 7과 같아야 합니다.

```js
assert.equal(new Num(3) + new Num(4), 7);
```

`new Num(3) - new Num(4)`는 -1과 같아야 합니다.

```js
assert.equal(new Num(3) - new Num(4), -1);
```

`new Num(3) * new Num(4)`는 12와 같아야 합니다.

```js
assert.equal(new Num(3) * new Num(4), 12);
```

`new Num(3) / new Num(4)`는 0.75와 같아야 합니다.

```js
assert.equal(new Num(3) / new Num(4), 0.75);
```

`new Num(3) < new Num(4)`는 true여야 합니다.

```js
assert(new Num(3) < new Num(4));
```

`new Num(3) > new Num(4)`는 false여야 합니다.

```js
assert(!(new Num(3) > new Num(4)));
```

`(new Num(5)).toString()`는 '5'를 반환해야 합니다.

```js
assert.equal(new Num(5).toString(), '5');
```

# --seed--

## --seed-contents--

```js
function Num(n) {

  return n;
}
```

# --solutions--

```js
function Num(n) {
  if (isNaN(n)) {
    throw new TypeError('Not a Number');
  }
  if (n < 1 || n > 10) {
    throw new TypeError('Out of range');
  }

  this._value = +n;
}
Num.prototype.valueOf = function() { return this._value; };
Num.prototype.toString = function () { return this._value.toString(); };
```
