---
id: 59713bd26bdeb8a594fb9413
title: 동전 개수 세기
challengeType: 1
forumTopicId: 302238
dashedName: count-the-coins
---

# --description--

미국 화폐에는 네 가지 종류의 일반 동전이 있습니다:

<ul>
  <li>쿼터(25센트)</li>
  <li>다임(10센트)</li>
  <li>니켈(5센트), 그리고</li>
  <li>페니(1센트)</li>
</ul>

<p>15센트를 거슬러 주는 방법은 여섯 가지가 있습니다:</p>

<ul>
  <li>다임 1개와 니켈 1개</li>
  <li>다임 1개와 페니 5개</li>
  <li>니켈 3개</li>
  <li>니켈 2개와 페니 5개</li>
  <li>니켈 1개와 페니 10개</li>
  <li>페니 15개</li>
</ul>

# --instructions--

주어진 입력 `cents`이 미국 페니 단위 금액을 나타낼 때, 이 일반 동전들을 사용해 거슬러 줄 수 있는 방법의 수를 결정하는 함수를 구현하세요.

# --hints--

`countCoins`은 함수여야 합니다.

```js
assert(typeof countCoins === 'function');
```

`countCoins(15)`는 `6`을 반환해야 합니다.

```js
assert.equal(countCoins(15), 6);
```

`countCoins(85)`는 `163`를 반환해야 합니다.

```js
assert.equal(countCoins(85), 163);
```

`countCoins(100)`는 `242`을 반환해야 합니다.

```js
assert.equal(countCoins(100), 242);
```

# --seed--

## --seed-contents--

```js
function countCoins(cents) {

  return true;
}
```

# --solutions--

```js
function countCoins(cents) {
  const operands = [1, 5, 10, 25];
  const targetsLength = cents + 1;
  const operandsLength = operands.length;
  const t = [1];

  for (let a = 0; a < operandsLength; a++) {
    for (let b = 1; b < targetsLength; b++) {
      // initialise undefined target
      t[b] = t[b] ? t[b] : 0;

      // accumulate target + operand ways
      t[b] += (b < operands[a]) ? 0 : t[b - operands[a]];
    }
  }

  return t[targetsLength - 1];
}
```
