---
id: 597b2b2a2702b44414742771
title: 팩토리얼
challengeType: 1
forumTopicId: 302263
dashedName: factorial
---

# --description--

숫자의 팩토리얼을 반환하는 함수를 작성하세요.

숫자의 팩토리얼은 다음과 같이 정의됩니다:

<pre><big>n! = n * (n-1) * (n-2) * ..... * 1</big>
</pre>

예를 들어:

<ul>
  <li><code>3! = 3 * 2 * 1 = 6</code></li>
  <li><code>4! = 4 * 3 * 2 * 1 = 24</code></li>
</ul>

**참고:** `0! = 1`

# --hints--

`factorial`는 함수여야 합니다.

```js
assert(typeof factorial === 'function');
```

`factorial(2)`는 숫자를 반환해야 합니다.

```js
assert(typeof factorial(2) === 'number');
```

`factorial(3)`는 6을 반환해야 합니다.

```js
assert.equal(factorial(3), 6);
```

`factorial(5)`는 120을 반환해야 합니다.

```js
assert.equal(factorial(5), 120);
```

`factorial(10)`는 3,628,800을 반환해야 합니다.

```js
assert.equal(factorial(10), 3628800);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

}
```

# --solutions--

```js
function factorial(n) {
  let sum = 1;
  while (n > 1) {
    sum *= n;
    n--;
  }
  return sum;
}
```
