---
id: 597f1e7fbc206f0e9ba95dc4
title: 정수의 약수
challengeType: 1
forumTopicId: 302265
dashedName: factors-of-an-integer
---

# --description--

양의 정수의 약수를 배열로 반환하는 함수를 작성하세요.

이 약수들은 나누어지는 수를 나누었을 때 양의 정수 결과가 나오는 양의 정수입니다.

# --before-each--

```js
const ans=[[1,3,5,9,15,45],[1,53],[1,2,4,8,16,32,64]];
```

# --hints--

`factors`는 함수여야 합니다.

```js
assert(typeof factors === 'function');
```

`factors(45)`는 `[1,3,5,9,15,45]`를 반환해야 합니다.

```js
assert.deepEqual(factors(45), ans[0]);
```

`factors(53)`는 `[1,53]`를 반환해야 합니다.

```js
assert.deepEqual(factors(53), ans[1]);
```

`factors(64)`는 `[1,2,4,8,16,32,64]`를 반환해야 합니다.

```js
assert.deepEqual(factors(64), ans[2]);
```

# --seed--

## --seed-contents--

```js
function factors(num) {

}
```

# --solutions--

```js
function factors(num)
{
 let n_factors = [], i, sqr=Math.floor(Math.sqrt(num));

 for (i = 1; i <=sqr ; i += 1)
  if (num % i === 0)
  {
   n_factors.push(i);
   if (num / i !== i)
    n_factors.push(num / i);
  }
 n_factors.sort(function(a, b){return a - b;});
 return n_factors;
}
```
