---
id: 597f24c1dda4e70f53c79c81
title: 피보나치 수열
challengeType: 1
forumTopicId: 302268
dashedName: fibonacci-sequence
---

# --description--

<code>n<sup>번째</sup></code> 피보나치 수를 생성하는 함수를 작성하세요.

<code>n<sup>번째</sup></code> 피보나치 수는 다음과 같이 정의됩니다:

<code>F<sub>n</sub> = F<sub>n-1</sub> + F<sub>n-2</sub></code>

수열의 처음 두 항은 0과 1입니다.

따라서 수열은 다음과 같습니다: 0, 1, 1, 2, 3, 5, 8, 13...

# --hints--

`fibonacci`는 함수여야 합니다.

```js
assert(typeof fibonacci === 'function');
```

`fibonacci(2)`는 숫자를 반환해야 합니다.

```js
assert(typeof fibonacci(2) == 'number');
```

`fibonacci(3)`는 2를 반환해야 합니다.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)`는 5를 반환해야 합니다.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)`는 55를 반환해야 합니다.

```js
assert.equal(fibonacci(10), 55);
```

# --seed--

## --seed-contents--

```js
function fibonacci(n) {

}
```

# --solutions--

```js
function fibonacci(n) {
  let a = 0, b = 1, t;
  while (--n >= 0) {
    t = a;
    a = b;
    b += t;
  }
  return a;
}
```
