---
id: 5987fd532b954e0f21b5d3f6
title: 평형 인덱스
challengeType: 1
forumTopicId: 302255
dashedName: equilibrium-index
---

# --description--

수열의 평형 인덱스는 해당 인덱스보다 낮은 인덱스에 있는 원소들의 합이 높은 인덱스에 있는 원소들의 합과 같은 수열의 인덱스입니다.

예를 들어, 수열 $A$에서:

<ul style='list-style: none;'>
  <li><big>$A_0 = -7$</big></li>
  <li><big>$A_1 =  1$</big></li>
  <li><big>$A_2 =  5$</big></li>
  <li><big>$A_3 =  2$</big></li>
  <li><big>$A_4 = -4$</big></li>
  <li><big>$A_5 =  3$</big></li>
  <li><big>$A_6 =  0$</big></li>
</ul>

`3`는 평형 인덱스입니다. 왜냐하면:

<ul style='list-style: none;'>
  <li><big>$A_0 + A_1 + A_2 = A_4 + A_5 + A_6$</big></li>
</ul>

`6`도 평형 인덱스입니다. 왜냐하면:

<ul style='list-style: none;'>
  <li><big>$A_0 + A_1 + A_2 + A_3 + A_4 + A_5 = 0$</big></li>
</ul>

(원소가 0개인 합은 0입니다)

`7`는 평형 인덱스가 아닙니다. 수열 $A$의 유효한 인덱스가 아니기 때문입니다.

# --instructions--

수열이 주어졌을 때, 평형 인덱스들을 반환하는 함수를 작성하세요(있다면).

수열이 매우 길 수 있다고 가정하세요.

# --hints--

`equilibrium`는 함수여야 합니다.

```js
assert(typeof equilibrium === 'function');
```

`equilibrium([-7, 1, 5, 2, -4, 3, 0])`는 `[3,6]`를 반환해야 합니다.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[0]),
  globalThis.ans[0]
);
```

`equilibrium([2, 4, 6])`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[1]),
  globalThis.ans[1]
);
```

`equilibrium([2, 9, 2])`는 `[1]`를 반환해야 합니다.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[2]),
  globalThis.ans[2]
);
```

`equilibrium([1, -1, 1, -1, 1, -1, 1])`는 `[0,1,2,3,4,5,6]`를 반환해야 합니다.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[3]),
  globalThis.ans[3]
);
```

`equilibrium([1])`는 `[0]`를 반환해야 합니다.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[4]),
  globalThis.ans[4]
);
```

`equilibrium([])`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[5]),
  globalThis.ans[5]
);
```

# --before-all--

```js
globalThis.equilibriumTests =
  [[-7, 1, 5, 2, -4, 3, 0], // 3, 6
  [2, 4, 6], // empty
  [2, 9, 2], // 1
  [1, -1, 1, -1, 1, -1, 1], // 0,1,2,3,4,5,6
  [1], // 0
  [] // empty
  ];
globalThis.ans = [[3, 6], [], [1], [0, 1, 2, 3, 4, 5, 6], [0], []];
```

# --seed--

## --seed-contents--

```js
function equilibrium(a) {

}
```

# --solutions--

```js
function equilibrium(a) {
  let N = a.length,
    i,
    l = [],
    r = [],
    e = [];
  for (l[0] = a[0], r[N - 1] = a[N - 1], i = 1; i < N; i++)
    { l[i] = l[i - 1] + a[i], r[N - i - 1] = r[N - i] + a[N - i - 1]; }
  for (i = 0; i < N; i++)
    { if (l[i] === r[i]) e.push(i); }
  return e;
}
```
