---
id: 59880443fb36441083c6c20e
title: 오일러 방법
challengeType: 1
forumTopicId: 302258
dashedName: euler-method
---

# --description--

오일러 방법은 주어진 초기값을 가진 1차 상미분방정식(ODE)의 해를 수치적으로 근사합니다. 이는 초기값 문제(IVP)를 푸는 명시적 방법으로, <a href="https://www.freecodecamp.org/news/eulers-method-explained-with-examples/" title="Euler's Method Explained with Examples" target="_blank" rel="noopener noreferrer nofollow">이 글</a>에서 설명합니다.

ODE는 다음 형태로 주어져야 합니다:

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt} = f(t,y(t))$</big></li>
</ul>

초기값은

<ul style='list-style: none;'>
  <li><big>$y(t_0) = y_0$</big></li>
</ul>

입니다.

수치해를 구하기 위해 좌변의 미분을 유한차분 근사로 대체합니다:

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt}  \approx \frac{y(t+h)-y(t)}{h}$</big></li>
</ul>

그다음 $y(t+h)$에 대해 풀면:

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, \frac{dy(t)}{dt}$</big></li>
</ul>

이는 다음과 같습니다:

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, f(t,y(t))$</big></li>
</ul>

반복해 구하는 해 규칙은 다음과 같습니다:

<ul style='list-style: none;'>
  <li><big>$y_{n+1} = y_n + h \, f(t_n, y_n)$</big></li>
</ul>

여기서 $h$는 단계 크기로, 해의 정확도에 가장 중요한 매개변수입니다. 단계 크기가 작을수록 정확도는 높아지지만 계산 비용도 증가하므로 문제에 맞게 직접 선택해야 합니다.

**예시: 뉴턴의 냉각 법칙**

뉴턴의 냉각 법칙은 초기 온도 $T(t_0) = T_0$인 물체가 온도 $T_R$인 환경에서 어떻게 냉각되는지 설명합니다:

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, \Delta T$</big></li>
</ul>

또는

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, (T(t) - T_R)$</big></li>
</ul>

이는 물체의 냉각 속도 $\\frac{dT(t)}{dt}$가 현재 온도 차이 $\\Delta T = (T(t) - T_R)$에 비례함을 의미합니다.

우리가 수치 근사와 비교할 해석적 해는 다음과 같습니다:

<ul style='list-style: none;'>
  <li><big>$T(t) = T_R + (T_0 - T_R) \; e^{-k t}$</big></li>
</ul>

# --instructions--

오일러 방법 루틴을 구현한 뒤, 주어진 뉴턴 냉각 법칙 예제를 다음 세 가지 단계 크기로 풀고 해석적 해와 비교하세요:

<ul>
  <li><code>2 s</code></li>
  <li><code>5 s</code></li>
  <li><code>10 s</code></li>
</ul>

**초기값:**

<ul>
  <li>초기 온도 <big>$T_0$</big>는 <code>100 °C</code>로 하세요</li>
  <li>실내 온도 <big>$T_R$</big>는 <code>20 °C</code>로 하세요</li>
  <li>냉각 상수 <big>$k$</big>는 <code>0.07</code>로 하세요</li>
  <li>계산할 시간 구간은 <code>0 s</code>부터 <code>100 s</code>까지로 하세요</li>
</ul>  

함수의 첫 번째 매개변수는 초기 시간, 두 번째는 초기 온도, 세 번째는 경과 시간, 네 번째는 단계 크기입니다.

# --hints--

`eulersMethod` should be a function.

```js
assert(typeof eulersMethod === 'function');
```

`eulersMethod(0, 100, 100, 2)` should return a number.

```js
assert(typeof eulersMethod(0, 100, 100, 2) === 'number');
```

`eulersMethod(0, 100, 100, 2)` should return 20.0424631833732.

```js
assert.equal(eulersMethod(0, 100, 100, 2), 20.0424631833732);
```

`eulersMethod(0, 100, 100, 5)` should return 20.01449963666907.

```js
assert.equal(eulersMethod(0, 100, 100, 5), 20.01449963666907);
```

`eulersMethod(0, 100, 100, 10)` should return 20.000472392.

```js
assert.equal(eulersMethod(0, 100, 100, 10), 20.000472392);
```

# --seed--

## --seed-contents--

```js
function eulersMethod(x1, y1, x2, h) {

}
```

# --solutions--

```js
function eulersMethod(x1, y1, x2, h) {
  let x = x1;
  let y = y1;

  while ((x < x2 && x1 < x2) || (x > x2 && x1 > x2)) {
    y += h * (-0.07 * (y - 20));
    x += h;
  }

  return y;
}
```
