---
id: 598de241872ef8353c58a7a2
title: 이항계수 계산하기
challengeType: 1
forumTopicId: 302259
dashedName: evaluate-binomial-coefficients
---

# --description--

주어진 n과 k 값에 대해 이항계수를 계산하는 함수를 작성하세요.

추천하는 공식은 다음과 같습니다:

$\\binom{n}{k} = \\frac{n!}{(n-k)!k!} = \\frac{n(n-1)(n-2)\\ldots(n-k+1)}{k(k-1)(k-2)\\ldots 1}$

# --hints--

`binom` 는 함수여야 합니다.

```js
assert(typeof binom === 'function');
```

`binom(5,3)` 은 10을 반환해야 합니다.

```js
assert.equal(binom(5, 3), 10);
```

`binom(7,2)` 은 21을 반환해야 합니다.

```js
assert.equal(binom(7, 2), 21);
```

`binom(10,4)` 은 210을 반환해야 합니다.

```js
assert.equal(binom(10, 4), 210);
```

`binom(6,1)` 은 6을 반환해야 합니다.

```js
assert.equal(binom(6, 1), 6);
```

`binom(12,8)` 은 495를 반환해야 합니다.

```js
assert.equal(binom(12, 8), 495);
```

# --seed--

## --seed-contents--

```js
function binom(n, k) {

}
```

# --solutions--

```js
function binom(n, k) {
  let coeff = 1;
  for (let i = n - k + 1; i <= n; i++) coeff *= i;
  for (let i = 1; i <= k; i++) coeff /= i;
  return coeff;
}
```
