---
id: 598ee8b91b410510ae82efef
title: 확장 가능한 소수 생성기
challengeType: 1
forumTopicId: 302262
dashedName: extensible-prime-generator
---

# --description--

어떤 합리적으로 큰 소수 생성도 자동으로 조정할 수 있는, 순서대로 소수를 생성하는 생성기를 작성하세요.

생성기는 다음 기능을 할 수 있어야 합니다:

<ul>
  <li>처음 <code>n</code>개의 소수를 보여주기</li>
  <li>범위 내 소수를 보여주기</li>
  <li>범위 내 소수의 개수를 보여주기</li>
  <li><code>n<sup>번째</sup></code> 소수를 보여주기</li>
</ul>

함수는 두 개의 매개변수를 가져야 합니다. 첫 번째는 `n` 또는 범위를 배열로 받습니다. 두 번째는 불리언을 받아 함수가 소수를 배열로 반환할지, 아니면 단일 숫자(범위 내 소수 개수 또는 <code>n<sup>번째</sup></code> 소수)를 반환할지 지정합니다. 매개변수에 따라 함수는 배열을 반환해야 합니다.

# --hints--

`primeGenerator`은 함수여야 합니다.

```js
assert(typeof primeGenerator === 'function');
```

`primeGenerator(20, true)`는 `[2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47, 53, 59, 61, 67, 71]`를 반환해야 합니다.

```js
assert.deepEqual(primeGenerator(20, true), [
  2,
  3,
  5,
  7,
  11,
  13,
  17,
  19,
  23,
  29,
  31,
  37,
  41,
  43,
  47,
  53,
  59,
  61,
  67,
  71
]);
```

`primeGenerator([100, 150], true)`는 `[101, 103, 107, 109, 113, 127, 131, 137, 139, 149]`를 반환해야 합니다.

```js
assert.deepEqual(primeGenerator([100, 150], true), [
  101,
  103,
  107,
  109,
  113,
  127,
  131,
  137,
  139,
  149
]);
```

`primeGenerator([7700, 8000], false)`는 `30`를 반환해야 합니다.

```js
assert.equal(primeGenerator([7700, 8000], false), 30);
```

`primeGenerator(10000, false)`는 `104729`를 반환해야 합니다.

```js
assert.equal(primeGenerator(10000, false), 104729);
```

# --seed--

## --seed-contents--

```js
function primeGenerator(num, showPrimes) {

}
```

# --solutions--

```js
function primeGenerator(num, showPrimes) {
  let i,
    arr = [];

  function isPrime(num) {
    // try primes <= 16
    if (num <= 16) { return (
      num == 2 || num == 3 || num == 5 || num == 7 || num == 11 || num == 13
    ); }
    // cull multiples of 2, 3, 5 or 7
    if (num % 2 == 0 || num % 3 == 0 || num % 5 == 0 || num % 7 == 0)
      { return false; }
    // cull square numbers ending in 1, 3, 7 or 9
    for (let i = 10; i * i <= num; i += 10) {
      if (num % (i + 1) == 0) return false;
      if (num % (i + 3) == 0) return false;
      if (num % (i + 7) == 0) return false;
      if (num % (i + 9) == 0) return false;
    }
    return true;
  }

  if (typeof num === 'number') {
    for (i = 0; arr.length < num; i++) if (isPrime(i)) arr.push(i);
    // first x primes
    if (showPrimes) return arr;
    // xth prime
    return arr.pop();
  }

  if (Array.isArray(num)) {
    for (i = num[0]; i <= num[1]; i++) if (isPrime(i)) arr.push(i);
    // primes between x .. y
    if (showPrimes) return arr;
    // number of primes between x .. y
    return arr.length;
  }
}
```
