---
id: 598eef80ba501f1268170e1e
title: 피보나치 n단계 수열
challengeType: 1
forumTopicId: 302267
dashedName: fibonacci-n-step-number-sequences
---

# --description--

이 수열들은 일반 피보나치 수열을 확장한 것으로 다음과 같습니다:

<ol>
  <li>$n = 2$일 때는 피보나치 수열이며, 초기값은 $[1, 1]$이고 $F_k^2 = F_{k-1}^2 + F_{k-2}^2$입니다</li>
  <li>$n = 3$일 때는 트리보나치 수열이며, 초기값은 $[1, 1, 2]$이고 $F_k^3 = F_{k-1}^3 + F_{k-2}^3 + F_{k-3}^3$입니다</li>
  <li>$n = 4$일 때는 테트라나치 수열이며, 초기값은 $[1, 1, 2, 4]$이고 $F_k^4 = F_{k-1}^4 + F_{k-2}^4 + F_{k-3}^4 + F_{k-4}^4$입니다...</li>
  <li>일반적으로 $n>2$일 때는 피보나치 $n$단계 수열 $F_k^n$이며, 초기값은 $(n-1)$번째 피보나치 $n$단계 수열 $F_k^{n-1}$의 처음 $n$개 값이고, 이 $n$번째 수열의 $k$번째 값은 $F_k^n = \sum_{i=1}^{(n)} {F_{k-i}^{(n)}}$입니다</li>
</ol>

작은 $n$ 값에 대해서는 그리스어 숫자 접두사를 사용해 각각의 수열 이름을 부르기도 합니다.

피보나치 $n$단계 수열:

| $n$ | 수열 이름    | 값                                                         |
| --- | ------------ | ---------------------------------------------------------- |
| 2   | fibonacci    | 1 1 2 3 5 8 13 21 34 55 89 144 233 377 610 ...             |
| 3   | tribonacci   | 1 1 2 4 7 13 24 44 81 149 274 504 927 1705 3136 ...        |
| 4   | tetranacci   | 1 1 2 4 8 15 29 56 108 208 401 773 1490 2872 5536 ...       |
| 5   | pentanacci   | 1 1 2 4 8 16 31 61 120 236 464 912 1793 3525 6930 ...       |
| 6   | hexanacci    | 1 1 2 4 8 16 32 63 125 248 492 976 1936 3840 7617 ...       |
| 7   | heptanacci   | 1 1 2 4 8 16 32 64 127 253 504 1004 2000 3984 7936 ...      |
| 8   | octonacci    | 1 1 2 4 8 16 32 64 128 255 509 1016 2028 4048 8080 ...      |
| 9   | nonanacci    | 1 1 2 4 8 16 32 64 128 256 511 1021 2040 4076 8144 ...      |
| 10  | decanacci    | 1 1 2 4 8 16 32 64 128 256 512 1023 2045 4088 8172 ...      |

초기값을 바꿔 생성할 수 있는 관련 수열도 있습니다: 루카스 수열은 $n=2$일 때 피보나치 수열처럼 두 이전 값을 더하지만 초기값으로 $\[2, 1]$을 사용합니다.

# --instructions--

피보나치 $n$단계 수열과 루카스 수열을 생성하는 함수를 작성하세요. 첫 번째 매개변수는 $n$입니다. 두 번째 매개변수는 반환할 요소 개수입니다. 세 번째 매개변수는 피보나치 수열을 출력할지 루카스 수열을 출력할지 지정합니다. 매개변수가 `"f"`이면 피보나치 수열을 반환하고, `"l"`이면 루카스 수열을 반환하세요. 수열은 배열로 반환해야 합니다.

# --hints--

`fib_luc`는 함수여야 합니다.

```js
assert(typeof fib_luc === 'function');
```

`fib_luc(2,10,"f")`는 `[1,1,2,3,5,8,13,21,34,55]`를 반환해야 합니다.

```js
assert.deepEqual(fib_luc(2, 10, 'f'), globalThis.ans[0]);
```

`fib_luc(3,15,"f")`는 `[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136]`를 반환해야 합니다.

```js
assert.deepEqual(fib_luc(3, 15, 'f'), globalThis.ans[1]);
```

`fib_luc(4,15,"f")`는 `[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536]`를 반환해야 합니다.

```js
assert.deepEqual(fib_luc(4, 15, 'f'), globalThis.ans[2]);
```

`fib_luc(2,10,"l")`는 `[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76]`를 반환해야 합니다.

```js
assert.deepEqual(fib_luc(2, 10, 'l'), globalThis.ans[3]);
```

`fib_luc(3,15,"l")`는 `[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ]`를 반환해야 합니다.

```js
assert.deepEqual(fib_luc(3, 15, 'l'), globalThis.ans[4]);
```

`fib_luc(4,15,"l")`는 `[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ]`를 반환해야 합니다.

```js
assert.deepEqual(fib_luc(4, 15, 'l'), globalThis.ans[5]);
```

`fib_luc(5,15,"l")`는 `[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]`를 반환해야 합니다.

```js
assert.deepEqual(fib_luc(5, 15, 'l'), globalThis.ans[6]);
```

# --before-all--

```js
globalThis.ans = [[1,1,2,3,5,8,13,21,34,55],
[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136],
[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536],
[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76],
[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ],
[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ],
[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]];
```

# --seed--

## --seed-contents--

```js
function fib_luc(n, len, w) {

}
```

# --solutions--

```js
function fib_luc(n, len, w) {
    function nacci(a, n, len) {
        while (a.length < len) {
            let sum = 0;
            for (let i = Math.max(0, a.length - n); i < a.length; i++)
                sum += a[i];
            a.push(sum);
        }
        return a;
    }
    if(w=="f"){
        return nacci(nacci([1,1], n, n), n, len);
    }else{
        return nacci(nacci([2,1], n, n), n, len);
    }
}
```
