---
id: 598eea87e5cf4b116c3ff81a
title: 머센수의 약수
challengeType: 1
forumTopicId: 302264
dashedName: factors-of-a-mersenne-number
---

# --description--

머센수는 <code>2<sup>P</sup>-1</code> 형태의 수입니다.

만약 `P`가 소수라면, 머센수는 머센 소수일 수 있습니다. (`P`가 소수가 아니면 머센수도 소수가 아닙니다.)

머센 소수를 찾을 때, 잠재적으로 긴 <a href="https://rosettacode.org/wiki/Lucas-Lehmer test" target="_blank" rel="noopener noreferrer nofollow">Lucas-Lehmer 시험</a>을 시작하기 전에 작은 약수를 찾아 지수를 제거하는 것이 유리합니다.

<code>2<sup>P</sup>-1</code>을 나누는 수를 결정하는 매우 효율적인 알고리즘들이 있습니다(또는 동등하게, <code>2<sup>P</sup> mod (그 수) = 1</code>인지 확인하는 방법).

몇몇 언어는 이미 이 지수와 나머지 연산을 내장 구현(modPow 또는 유사한 이름)으로 제공합니다.

다음은 이 modPow를 직접 구현하는 방법입니다:

예를 들어, <code>2<sup>23</sup> mod 47</code>을 계산해 봅시다.

지수 23을 이진수로 변환하면 10111이 됩니다. <code><tt>square</tt> = 1</code>에서 시작해 반복해서 제곱하세요.

지수의 최상위 비트를 제거하고, 만약 1이면 `square`에 지수의 밑수(2)를 곱한 뒤 <code><tt>square</tt> modulo 47</code>을 계산하세요.

마지막 단계에서 나온 나머지 결과를 다음 단계의 `square` 초기값으로 사용하세요:

<pre>제거   선택적
square    최상위 비트  2 곱하기  47로 나눈 나머지
------------  -------  -------------  ------
1*1 = 1       1  0111  1*2 = 2           2
2*2 = 4       0   111     없음             4
4*4 = 16      1    11  16*2 = 32        32
32*32 = 1024  1     1  1024*2 = 2048    27
27*27 = 729   1        729*2 = 1458      1
</pre>

<code>2<sup>23</sup> mod 47 = 1</code>이므로 47은 <code>2<sup>P</sup>-1</code>의 약수입니다.

(이를 확인하려면 양쪽에서 1을 빼면 됩니다: <code>2<sup>23</sup>-1 = 0 mod 47</code>.)

47이 약수임을 보였으므로 <code>2<sup>23</sup>-1</code>은 소수가 아닙니다.

머센수의 추가 성질을 이용해 이 과정을 더 정교하게 할 수 있습니다.

<code>2<sup>P</sup>-1</code>의 약수 `q`는 `2kP+1` 형태여야 하며, `k`는 0 또는 양의 정수입니다. 또한 `q`는 `1` 또는 `7 mod 8`여야 합니다.

마지막으로, 잠재적 약수 `q`는 <a href="https://rosettacode.org/wiki/Primality_by_trial_division" target="_blank" rel="noopener noreferrer nofollow">소수</a>여야 합니다.

다른 나눗셈 시험 알고리즘과 마찬가지로, 알고리즘은 `2kP+1 > sqrt(N)`일 때 정지합니다. 이 테스트들은 주로 `P`가 소수인 머센수에만 적용됩니다. 예를 들어, <code>M<sub>4</sub>=15</code>는 이 기법으로 약수를 찾지 못하지만 3과 5로 인수분해되며, 둘 다 `2kP+1`에 맞지 않습니다.

# --instructions--

위 방법을 사용해 <code>2<sup>p</sup>-1</code>의 약수를 찾으세요.

# --hints--

`check_mersenne`는 함수여야 합니다.

```js
assert(typeof check_mersenne === 'function');
```

`check_mersenne(3)`는 문자열을 반환해야 합니다.

```js
assert(typeof check_mersenne(3) == 'string');
```

`check_mersenne(3)`는 문자열 `M3 = 2^3-1 is prime`를 반환해야 합니다.

```js
assert.equal(check_mersenne(3), 'M3 = 2^3-1 is prime');
```

`check_mersenne(23)`는 문자열 `M23 = 2^23-1 is composite with factor 47`를 반환해야 합니다.

```js
assert.equal(check_mersenne(23), 'M23 = 2^23-1 is composite with factor 47');
```

`check_mersenne(929)`는 문자열 `M929 = 2^929-1 is composite with factor 13007`를 반환해야 합니다.

```js
assert.equal(
  check_mersenne(929),
  'M929 = 2^929-1 is composite with factor 13007'
);
```

# --seed--

## --seed-contents--

```js
function check_mersenne(p) {

}
```

# --solutions--

```js
function check_mersenne(p){
    function isPrime(value){
      for (let i=2; i < value; i++){
        if (value % i == 0){
          return false;
        }
        if (value % i != 0){
          return true;
         }
      }
    }

    function trial_factor(base, exp, mod){
      let square, bits;
      square = 1;
      bits = exp.toString(2).split('');
      for (let i=0,ln=bits.length; i<ln; i++){
        square = Math.pow(square, 2) * (bits[i] == 1 ? base : 1) % mod;
      }
      return (square == 1);
    }

    function mersenne_factor(p){
      let limit, k, q;
      limit = Math.sqrt(Math.pow(2,p) - 1);
      k = 1;
      while ((2*k*p - 1) < limit){
        q = 2*k*p + 1;
        if (isPrime(q) && (q % 8 == 1 || q % 8 == 7) && trial_factor(2,p,q)){
          return q; // q is a factor of 2**p-1
        }
        k++;
      }
      return null;
    }
  let f, result;
  result="M"+p+" = 2^"+p+"-1 is ";
  f = mersenne_factor(p);
  result+=f == null ? "prime" : "composite with factor "+f;
  return result;
}
```
