---
id: 5992e222d397f00d21122931
title: 피보나치 단어
challengeType: 1
forumTopicId: 302269
dashedName: fibonacci-word
---

# --description--

피보나치 단어 수열은 피보나치 수열과 유사한 방식으로 생성할 수 있지만, 연결 반복에 중점을 둡니다.

<pre>F_Word<sub>1</sub> 을 <strong>1</strong>로 정의하세요
F_Word<sub>2</sub> 을 <strong>0</strong>으로 정의하세요
F_Word<sub>3</sub> 을 F_Word<sub>2</sub> 와 F_Word<sub>1</sub> 을 연결하여 만드세요, 즉: <strong>01</strong>
F_Word<sub>n</sub> 은 F_Word<sub>n-1</sub> 과 F_word<sub>n-2</sub> 를 연결하여 만드세요
</pre>

이 도전 과제에서는 엔트로피 계산이 필요합니다, <a href="https://www.freecodecamp.org/learn/coding-interview-prep/rosetta-code/entropy" target="_blank" rel="noopener noreferrer nofollow">이 Rosetta Code 도전 과제에서 보여준 것처럼</a>

# --instructions--

처음 `n`개의 피보나치 단어를 반환하는 함수를 작성하세요. `n`의 수가 함수의 매개변수로 제공됩니다. 함수는 객체 배열을 반환해야 합니다. 객체는 `{ N: 1, Length: 1, Entropy: 0, Word: '1' }` 형식이어야 합니다. `Entropy`은 문자열 `Word`에 대해 계산되며 소수점 8자리까지 반올림됩니다. 이 수열의 인덱스는 `1`부터 시작한다는 점에 유의하세요.

# --hints--

`fibWord`는 함수여야 합니다.

```js
assert(typeof fibWord === 'function');
```

`fibWord(5)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(fibWord(5)));
```

`fibWord(5)`는 `[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }]`를 반환해야 합니다.

```js
assert.deepEqual(fibWord(5), globalThis.words5);
```

`fibWord(7)`는 `[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }, { N:6, Length:8, Entropy:0.954434, Word:'01001010' }, { N:7, Length:13, Entropy:0.9612366, Word:'0100101001001' }]`를 반환해야 합니다.

```js
assert.deepEqual(fibWord(7), globalThis.words7);
```

# --before-all--

```js
globalThis.words5 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' }
];

globalThis.words7 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' },
  { N: 6, Length: 8, Entropy: 0.954434, Word: '01001010' },
  { N: 7, Length: 13, Entropy: 0.9612366, Word: '0100101001001' }
];
```

# --seed--

## --seed-contents--

```js
function fibWord(n) {

}
```

# --solutions--

```js
// Round to digits
function roundFloat(num, digits) {
  return Math.round(num * 10.0**digits) / (10.0**digits);
}

// Entropy calculation for string with only 0 and 1
function entropy(word) {
  function digitEntropy(count) {
    return count < 1 ? 0
      : - count / word.length * Math.log2(count / word.length);
  }
  const numZeros = word.split('').filter(e => e === '0').length;
  const numOnes  = word.length - numZeros;
  return roundFloat(digitEntropy(numZeros) + digitEntropy(numOnes), 8);
}

// Compute array of Fibonacci words
function fibWord(n) {
  return [...Array(n).keys()]
    .reduce((words, i) => {
      const word = i === 0 ? "1"
                 : i === 1 ? "0"
                 : words[i - 1].Word + words[i - 2].Word;
      words.push(
        { N: i + 1, Length: word.length, Entropy: entropy(word), Word: word }
      );
      return words;
    }, []);
}
```
