---
id: 599d1566a02b571412643b84
title: 에티오피아 곱셈
challengeType: 1
forumTopicId: 302257
dashedName: ethiopian-multiplication
---

# --description--

에티오피아 곱셈은 덧셈, 두 배로 늘리기, 반으로 나누기만 사용하여 정수를 곱하는 방법입니다.

**방법:**

<ol>
  <li>곱할 두 수를 가져와 두 열의 맨 위에 적으세요</li>
  <li>왼쪽 열에서 마지막 수를 반복해서 반으로 나누고 나머지는 버리며, 결과를 같은 열 아래에 적으세요. 값이 <code>1</code>이 될 때까지 계속하세요</li>
  <li>오른쪽 열에서 마지막 수를 반복해서 두 배로 늘리고 결과를 아래에 적으세요. 왼쪽 열에 <code>1</code>이 있는 행과 같은 행에 결과가 추가되면 멈추세요</li>
  <li>만들어진 표를 살펴보고 왼쪽 열의 값이 짝수인 행은 모두 버리세요</li>
  <li>남은 오른쪽 열의 값을 모두 더해 원래 두 수를 곱한 결과를 만드세요</li>
</ol>

**예를 들어:** `17 × 34`

<pre>17   34
</pre>

첫 번째 열을 반으로 나누기:

<pre>17   34
8
4
2
1
</pre>

두 번째 열을 두 배로 늘리기:

<pre>17   34
8    68
4   136
2   272
1   544
</pre>

첫 번째 셀이 짝수인 행을 취소선 처리:

<pre>17   34
8    <strike>68</strike>
4   <strike>136</strike>
2   <strike>272</strike>
1   544
</pre>

오른쪽 열에 남은 수를 모두 더하기:

<!-- markdownlint-disable MD003 -->

<pre>17   34
8    --
4   ---
2   ---
1   544
   ====
    578
</pre>

<!-- markdownlint-enable MD003 -->

그래서 `17`과 `34`를 에티오피아 방법으로 곱하면 `578`입니다.

# --instructions--

세 개의 이름 있는 함수/메서드/프로시저/서브루틴을 정의하는 것이 과제입니다:

<ol>
  <li>정수를 반으로 나누는 함수,</li>
  <li>정수를 두 배로 늘리는 함수,</li>
  <li>정수가 짝수인지 판별하는 함수</li>
</ol>

이 함수들을 사용해 에티오피아 곱셈을 수행하는 함수를 만드세요.

<!-- markdownlint-disable MD046-->

# --hints--

`eth_mult`는 함수여야 합니다.

```js
assert(typeof eth_mult === 'function');
```

`eth_mult(17,34)`는 `578`를 반환해야 합니다.

```js
assert.equal(eth_mult(17, 34), 578);
```

`eth_mult(23,46)`는 `1058`를 반환해야 합니다.

```js
assert.equal(eth_mult(23, 46), 1058);
```

`eth_mult(12,27)`는 `324`를 반환해야 합니다.

```js
assert.equal(eth_mult(12, 27), 324);
```

`eth_mult(56,98)`는 `5488`를 반환해야 합니다.

```js
assert.equal(eth_mult(56, 98), 5488);
```

`eth_mult(63,74)`는 `4662`를 반환해야 합니다.

```js
assert.equal(eth_mult(63, 74), 4662);
```

# --seed--

## --seed-contents--

```js
function eth_mult(a, b) {

}
```

# --solutions--

```js
function eth_mult(a, b) {
  let sum = 0; a = [a]; b = [b];

  let half = a => a / 2,
    double = a => a * 2,
    is_even = a => a % 2 == 0;

  while (a[0] !== 1) {
    a.unshift(Math.floor(half(a[0])));
    b.unshift(double(b[0]));
  }

  for (let i = a.length - 1; i > 0; i -= 1) {
    if (!is_even(a[i])) {
      sum += b[i];
    }
  }
  return sum + b[0];
}
```
