---
id: 599d15309e88c813a40baf58
title: 엔트로피
challengeType: 1
forumTopicId: 302254
dashedName: entropy
---

# --description--

주어진 입력 문자열의 샤논 엔트로피 H를 계산하세요.

이산 확률 변수 $X$가 $n$개의 서로 다른 문자(이진의 경우 n=2)로 구성된 $N$개의 "기호"(전체 문자) 문자열일 때, $X$의 샤논 엔트로피는 비트/기호 단위로 다음과 같습니다:

$H_2(X) = -\\sum\_{i=1}^n \\frac{count_i}{N} \\log_2 \\left(\\frac{count_i}{N}\\right)$

여기서 $count_i$는 문자 $n_i$의 개수입니다.

# --hints--

`entropy`는 함수여야 합니다.

```js
assert(typeof entropy === 'function');
```

`entropy("0")`는 `0`를 반환해야 합니다.

```js
assert.equal(entropy('0'), 0);
```

`entropy("01")`는 `1`를 반환해야 합니다.

```js
assert.equal(entropy('01'), 1);
```

`entropy("0123")`는 `2`를 반환해야 합니다.

```js
assert.equal(entropy('0123'), 2);
```

`entropy("01234567")`는 `3`를 반환해야 합니다.

```js
assert.equal(entropy('01234567'), 3);
```

`entropy("0123456789abcdef")`는 `4`를 반환해야 합니다.

```js
assert.equal(entropy('0123456789abcdef'), 4);
```

`entropy("1223334444")`는 `1.8464393446710154`를 반환해야 합니다.

```js
assert.equal(entropy('1223334444'), 1.8464393446710154);
```

# --seed--

## --seed-contents--

```js
function entropy(s) {

}
```

# --solutions--

```js
function entropy(s) {
    // Create a dictionary of character frequencies and iterate over it.
  function process(s, evaluator) {
    let h = Object.create(null),
      k;
    s.split('').forEach(c => {
      h[c] && h[c]++ || (h[c] = 1); });
    if (evaluator) for (k in h) evaluator(k, h[k]);
    return h;
  }
    // Measure the entropy of a string in bits per symbol.

  let sum = 0,
    len = s.length;
  process(s, (k, f) => {
    const p = f / len;
    sum -= p * Math.log(p) / Math.log(2);
  });
  return sum;
}
```
