---
id: 599d0ba974141b0f508b37d5
title: 에미르프 소수
challengeType: 1
forumTopicId: 302253
dashedName: emirp-primes
---

# --description--

에미르프는 (**소수**를 거꾸로 쓴 것) 10진수 표기에서 뒤집었을 때 다른 소수가 되는 소수입니다.

# --instructions--

다음 기능을 하는 함수를 작성하세요:

<ul>
  <li>처음 <code>n</code>개의 에미르프 수를 보여줍니다.</li>
  <li>범위 내의 에미르프 수를 보여줍니다.</li>
  <li>범위 내 에미르프 수의 개수를 보여줍니다.</li>
  <li><code>n<sup>번째</sup></code> 에미르프 수를 보여줍니다.</li>
</ul>

함수는 두 개의 매개변수를 받아야 합니다. 첫 번째 매개변수는 `n` 또는 배열 형태의 범위를 받습니다. 두 번째 매개변수는 불리언으로, 함수가 에미르프 수를 배열로 반환할지 아니면 단일 숫자(범위 내 소수 개수 또는 <code>n<sup>번째</sup></code> 소수)로 반환할지 지정합니다. 매개변수에 따라 함수는 배열 또는 숫자를 반환해야 합니다.

# --hints--

`emirps`은 함수여야 합니다.

```js
assert(typeof emirps === 'function');
```

`emirps(20,true)`는 `[13,17,31,37,71,73,79,97,107,113,149,157,167,179,199,311,337,347,359,389]`를 반환해야 합니다.

```js
assert.deepEqual(emirps(20, true), [
  13,
  17,
  31,
  37,
  71,
  73,
  79,
  97,
  107,
  113,
  149,
  157,
  167,
  179,
  199,
  311,
  337,
  347,
  359,
  389
]);
```

`emirps(1000)`는 `70529`를 반환해야 합니다.

```js
assert.deepEqual(emirps(1000), 70529);
```

`emirps([7700,8000],true)`는 `[7717,7757,7817,7841,7867,7879,7901,7927,7949,7951,7963]`를 반환해야 합니다.

```js
assert.deepEqual(emirps([7700, 8000], true), [
  7717,
  7757,
  7817,
  7841,
  7867,
  7879,
  7901,
  7927,
  7949,
  7951,
  7963
]);
```

`emirps([7700,8000],false)`는 `11`를 반환해야 합니다.

```js
assert.deepEqual(emirps([7700, 8000], false), 11);
```

# --seed--

## --seed-contents--

```js
function emirps(n) {

}
```

# --solutions--

```js
function emirps(num, showEmirps)
{
  const is_prime = function(n)
    {
    if (!(n % 2) || !(n % 3)) return false;
    let p = 1;
    while (p * p < n)
                    { if (n % (p += 4) == 0 || n % (p += 2) == 0)
                            { return false; } }
    return true;
  };
  const is_emirp = function(n) {
    const r = parseInt(n.toString().split('').reverse().join(''));
    return r != n && is_prime(n) && is_prime(r);
  };

  let i,
    arr = [];
  if (typeof num === 'number') {
    for (i = 0; arr.length < num; i++) if (is_emirp(i)) arr.push(i);
    // first x emirps
    if (showEmirps) return arr;
    // xth emirp
    return arr.pop();
  }

  if (Array.isArray(num)) {
    for (i = num[0]; i <= num[1]; i++) if (is_emirp(i)) arr.push(i);
    // emirps between x .. y
    if (showEmirps) return arr;
    // number of emirps between x .. y
    return arr.length;
  }
}
```
