---
id: 59c3ec9f15068017c96eb8a3
title: 파레이 수열
challengeType: 1
forumTopicId: 302266
dashedName: farey-sequence
---

# --description--

차수 `n`의 파레이 수열 <code>F<sub>n</sub></code>은 `0`과 `1` 사이의 완전히 기약분수로, 기약분수일 때 분모가 `n` 이하인 분수들을 크기 순서대로 나열한 수열입니다.

*파레이 수열*은 때때로 잘못 *파레이 급수*라고 불리기도 합니다.

각 파레이 수열은:

<ul>
  <li>값 0으로 시작하며, 이는 분수 $ \frac{0}{1} $로 나타냅니다</li>
  <li>값 1로 끝나며, 이는 분수 $ \frac{1}{1}$로 나타냅니다.</li>
</ul>

차수 `1`부터 `5`까지의 파레이 수열은 다음과 같습니다:

<ul>
  <li style='list-style: none;'>${\bf\it{F}}_1 = \frac{0}{1}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_2 = \frac{0}{1}, \frac{1}{2}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_3 = \frac{0}{1}, \frac{1}{3}, \frac{1}{2}, \frac{2}{3}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_4 = \frac{0}{1}, \frac{1}{4}, \frac{1}{3}, \frac{1}{2}, \frac{2}{3}, \frac{3}{4}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_5 = \frac{0}{1}, \frac{1}{5}, \frac{1}{4}, \frac{1}{3}, \frac{2}{5}, \frac{1}{2}, \frac{3}{5}, \frac{2}{3}, \frac{3}{4}, \frac{4}{5}, \frac{1}{1}$</li>
</ul>

# --instructions--

차수 `n`의 파레이 수열을 반환하는 함수를 작성하세요. 함수는 `n` 하나의 매개변수를 가져야 합니다. 수열은 배열로 반환해야 합니다.

# --hints--

`farey`는 함수여야 합니다.

```js
assert(typeof farey === 'function');
```

`farey(3)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(farey(3)));
```

`farey(3)`는 `['0/1','1/3','1/2','2/3','1/1']`를 반환해야 합니다.

```js
assert.deepEqual(farey(3),['0/1', '1/3', '1/2', '2/3', '1/1']);
```

`farey(4)`는 `['0/1','1/4','1/3','1/2','2/3','3/4','1/1']`를 반환해야 합니다.

```js
assert.deepEqual(farey(4), ['0/1', '1/4', '1/3', '1/2', '2/3', '3/4', '1/1']);
```

`farey(5)`는 `['0/1','1/5','1/4','1/3','2/5','1/2','3/5','2/3','3/4','4/5','1/1']`를 반환해야 합니다.

```js
assert.deepEqual(farey(5), [
  '0/1',
  '1/5',
  '1/4',
  '1/3',
  '2/5',
  '1/2',
  '3/5',
  '2/3',
  '3/4',
  '4/5',
  '1/1'
]);
```

# --seed--

## --seed-contents--

```js
function farey(n) {

}
```

# --solutions--

```js
function farey(n) {
  const sequence = [{ string: "0/1", float: 0.0 }];
  for (let i = 1; i < n; i++) {
    for (let j = n; j >= i; j--) {
      if (i === 1 || j % i > 0) {
        sequence.push({ string: `${i}/${j}`, float: i / j });
      }
    }
  }
  return sequence
    .sort((a, b) => a.float - b.float)
    .map(e => e.string)
}
```
