---
id: 59d9c6bc214c613ba73ff012
title: SEDOL 목록
challengeType: 1
forumTopicId: 302305
dashedName: sedols
---

# --description--

<abbr title="Stock Exchange Daily Official List">SEDOL</abbr>은 런던 증권거래소에서 발행하는 증권 식별 번호 목록입니다.

6자리 <abbr title="Stock Exchange Daily Official List">SEDOL</abbr> 번호 목록 각각에 대해 체크섬 숫자를 계산하여 덧붙이세요. 즉, 왼쪽에 주어진 입력 문자열에 대해 함수가 오른쪽에 해당하는 문자열을 반환해야 합니다:

<pre>
710889 => 7108899
B0YBKJ => B0YBKJ7
406566 => 4065663
B0YBLH => B0YBLH2
228276 => 2282765
B0YBKL => B0YBKL9
557910 => 5579107
B0YBKR => B0YBKR5
585284 => 5852842
B0YBKT => B0YBKT7
B00030 => B000300
</pre>

각 입력이 올바르게 형성되었는지, 특히 SEDOL 문자열에 허용되는 유효한 문자인지 확인하세요. 잘못된 입력에 대해서는 함수가 `null`를 반환해야 합니다.

# --hints--

`sedol`는 함수여야 합니다.

```js
assert(typeof sedol === 'function');
```

`sedol('a')`는 null을 반환해야 합니다.

```js
assert(sedol('a') === null);
```

`sedol('710889')`는 '7108899'를 반환해야 합니다.

```js
assert(sedol('710889') === '7108899');
```

`sedol('BOATER')`는 null을 반환해야 합니다.

```js
assert(sedol('BOATER') === null);
```

`sedol('228276')`는 '2282765'를 반환해야 합니다.

```js
assert(sedol('228276') === '2282765');
```

# --seed--

## --seed-contents--

```js
function sedol(input) {
  const checkSum = 0
  
  return checkSum;
}
```

# --solutions--

```js
function sedol(input) {
  const checkDigit = sedolCheckDigit(input);
  if (checkDigit !== null) {
    return input + checkDigit;
  }
  return null;
}

const weight = [1, 3, 1, 7, 3, 9, 1];
function sedolCheckDigit(char6) {
  if (char6.search(/^[0-9BCDFGHJKLMNPQRSTVWXYZ]{6}$/) === -1) {
    return null;
  }

  let sum = 0;
  for (let i = 0; i < char6.length; i++) {
    sum += weight[i] * parseInt(char6.charAt(i), 36);
  }
  const checkSum = (10 - (sum % 10)) % 10;
  return checkSum.toString();
}
```
