---
id: 59da22823d04c95919d46269
title: '선원, 코코넛 그리고 원숭이 문제'
challengeType: 1
forumTopicId: 302304
dashedName: sailors-coconuts-and-a-monkey-problem
---

# --description--

다섯 명의 선원이 난파되어 섬에 갇히고 낮 동안 많은 코코넛을 모읍니다.

그날 밤 첫 번째 선원이 깨어나 먼저 자신의 몫을 가져가기로 결심하고 코코넛 더미를 다섯 개로 똑같이 나누려고 하지만 코코넛이 하나 남아 원숭이에게 던져줍니다. 그리고 다섯 개로 나눈 코코넛 더미 중 하나를 숨기고 나머지 네 개 더미를 다시 합쳐 하나의 보이는 코코넛 더미로 만든 뒤 잠자리에 듭니다.

간단히 말해, 밤새 각 선원이 차례로 일어나 코코넛 더미를 다섯 개로 나누고, 하나의 코코넛이 남아 그 나머지 코코넛을 원숭이에게 주고, 나머지 네 개 더미를 다시 합쳐 하나의 더미를 만듭니다.

아침이 되면 (밤새 각 선원이 몰래 따로 한 행동 후) 남은 코코넛을 다섯 개의 똑같은 더미로 나누는데, 이때 코코넛 더미가 나머지 없이 선원들 사이에 정확히 나누어집니다. (아침에는 원숭이에게 주는 코코넛이 없습니다.)

# --instructions--

`N` 명의 선원을 위해 낮 동안 모은 초기 코코넛 더미의 최솟값을 반환하는 함수를 만드세요. **참고:** 이야기 속 시간 안에 코코넛을 아무 수만큼 모으고 여러 번 나누는 일이 가능하다고 가정하여 수학적 계산에 영향을 주지 않습니다.

# --hints--

`splitCoconuts` 은 함수여야 합니다.

```js
assert(typeof splitCoconuts === 'function');
```

`splitCoconuts(5)` 는 3121을 반환해야 합니다.

```js
assert(splitCoconuts(5) === 3121);
```

`splitCoconuts(6)` 는 233275를 반환해야 합니다.

```js
assert(splitCoconuts(6) === 233275);
```

`splitCoconuts(7)` 는 823537을 반환해야 합니다.

```js
assert(splitCoconuts(7) === 823537);
```

# --seed--

## --seed-contents--

```js
function splitCoconuts(intSailors) {

  return true;
}
```

# --solutions--

```js
function splitCoconuts(intSailors) {
  let intNuts = intSailors;
  let result = splitCoconutsHelper(intNuts, intSailors);
  while (!result) {
    intNuts += 1;
    result = splitCoconutsHelper(intNuts, intSailors);
  }

  return intNuts;
}

function splitCoconutsHelper(intNuts, intSailors, intDepth) {
  const nDepth = intDepth !== undefined ? intDepth : intSailors;
  const portion = Math.floor(intNuts / intSailors);
  const remain = intNuts % intSailors;

  if (portion <= 0 || remain !== (nDepth ? 1 : 0)) {
    return null;
  }

  if (nDepth) {
    return splitCoconutsHelper(
      intNuts - portion - remain, intSailors, nDepth - 1
    );
  }

  return intNuts;
}
```
