---
id: 59e0a8df964e4540d5abe599
title: Brain**** 인터프리터 구현하기
challengeType: 1
forumTopicId: 302261
dashedName: execute-brain
---

# --description--

Brain**** 인터프리터를 구현하는 함수를 작성하세요. 이 함수는 문자열을 매개변수로 받고, 출력으로 문자열을 반환해야 합니다. 자세한 내용은 아래에 나와 있습니다.

RCBF는 다양한 언어로 작성된 Rosetta Code용 <a href="https://rosettacode.org/wiki/Brainf***" target="_blank" rel="noopener noreferrer nofollow">Brainf***</a> 컴파일러와 인터프리터 모음입니다.

아래는 RCBF의 각 버전에 대한 링크입니다.

구현은 다음 명령어만 정확히 구현하면 됩니다:

| 명령어           | 설명                                                                        |
| ----------------- | --------------------------------------------------------------------------- |
| <code>></code>    | 포인터를 오른쪽으로 이동                                                    |
| <code>&lt;</code> | 포인터를 왼쪽으로 이동                                                     |
| <code>+</code>    | 포인터가 가리키는 메모리 셀 값을 1 증가                                   |
| <code>-</code>    | 포인터가 가리키는 메모리 셀 값을 1 감소                                   |
| <code>.</code>    | 포인터가 가리키는 셀의 값을 문자로 출력                                   |
| <code>,</code>    | 문자를 입력받아 포인터가 가리키는 셀에 저장                              |
| <code>\[</code>   | 포인터가 가리키는 셀 값이 0이면 대응하는 <code>]</code> 다음으로 점프     |
| <code>]</code>    | 포인터가 가리키는 셀 값이 0이 아니면 대응하는 <code>\[</code>로 점프      |

셀 크기는 자유롭게 설정할 수 있으며, EOF(파일 끝) 지원은 선택 사항입니다. 메모리가 유한한지 무한한지도 선택 사항입니다.

# --hints--

`brain(globalThis.hello)`는 문자열을 반환해야 합니다

```js
assert(typeof brain(globalThis.hello) === 'string');
```

`brain("++++++[>++++++++++<-]>+++++.")`는 "A"를 반환해야 합니다

```js
assert.equal(brain('++++++[>++++++++++<-]>+++++.'), 'A');
```

`brain(globalThis.hello)`는 `Hello World!\n`를 반환해야 합니다

```js
assert.equal(brain(globalThis.hello), 'Hello World!\n');
```

`brain(globalThis.fib)`는 `1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89`를 반환해야 합니다

```js
assert.equal(brain(globalThis.fib), '1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89');
```

# --before-all--

```js
globalThis.fib=`+

++

+++

++++

+>+>>

>>++++

+++++++

++++++++

+++++++++

++++++++++

++++++>++++

++++++++++++

+++++++++++++

+++<<<<<<[>[>>

>>>>+>+<<<<<<<-

]>>>>>>>[<<<<<<<

+>>>>>>>-]<[>++++

++++++[-<-[>>+>+<<

<-]>>>[<<<+>>>-]+<[

>[-]<[-]]>[<<[>>>+<<

<-]>>[-]]<<]>>>[>>+>+

<<<-]>>>[<<<+>>>-]+<[>

[-]<[-]]>[<<+>>[-]]<<<<

<<<]>>>>>[++++++++++++++

+++++++++++++++++++++++++

+++++++++.[-]]++++++++++<[

->-<]>+++++++++++++++++++++

+++++++++++++++++++++++++++.

[-]<<<<<<<<<<<<[>>>+>+<<<<-]>

>>>[<<<<+>>>>-]<-[>>.>.<<<[-]]

<<[>>+>+<<<-]>>>[<<<+>>>-]<<[<+

>-]>[<+>-]<<<-]`;
globalThis.hello='++++++++[>++++[>++>+++>+++>+<<<<-]>+>+>->>+[<]<-]>>.>---.+++++++..+++.>>.<-.<.+++.------.--------.>>+.>++.'
```

# --seed--

## --seed-contents--

```js
function brain(prog) {

}
```

# --solutions--

```js
function brain(prog){
  var output="";
    var code; // formatted code
  var ip = 0; // current instruction within code
  var nest = 0; // current bracket nesting (for Out button)
  var ahead = []; // locations of matching brackets

  var data = [0]; // data array (mod by +, -)
  var dp = 0; // index into data (mod by <, >)

  var inp = 0; // current input character (fetch with ,)
  var quit = 0;
    var commands = {
    '>':function() { if (++dp >= data.length) data[dp]=0 },
    '<':function() { if (--dp < 0) quit++ },
    '+':function() { ++data[dp] },
    '-':function() { --data[dp] },
    '[':function() { if (!data[dp]) ip = ahead[ip]; else ++nest },
    ']':function() { if ( data[dp]) ip = ahead[ip]; else --nest },
    ',':function() {
        var c = document.getElementById("input").value.charCodeAt(inp++);
        data[dp] = isNaN(c) ? 0 : c; // EOF: other options are -1 or no change
    },
    '.':function() {
            output+=String.fromCharCode(data[dp]);
            /*var s = document.getElementById("output").innerHTML)
             + String.fromCharCode(data[dp]);
            s = s.replace(/\n/g,"<br>").replace(/ /g,"&amp;nbsp;");
            document.getElementById("output").innerHTML = s;*/
        },
    };

    let ar=prog.split('');
    var st = [], back, error = -1;
    for (ip=0; ip<ar.length; ip++) {
        switch(ar[ip]) {
        case '[':
            st.push(ip);
            break;
        case ']':
            if (st.length == 0) error = ip;
            back = st.pop();
            ahead[ip] = back;
            ahead[back] = ip;
            break;
        }
    }

    for(ip=0;ip<ar.length;ip++){
    if(commands.hasOwnProperty(ar[ip]))
          commands[ar[ip]]();
    }

    return output;
}
```
