---
id: 59f40b17e79dbf1ab720ed7a
title: 부서 번호
challengeType: 1
forumTopicId: 302249
dashedName: department-numbers
---

# --description--

매우 체계적인 도시가 각 부서에 번호를 할당하기로 결정했습니다:

<ul>
  <li>경찰서</li>
  <li>위생부서</li>
  <li>소방서</li>
</ul>

각 부서는 1부터 7까지(포함) 번호를 가질 수 있습니다.

세 부서 번호는 서로 달라야 하며 합이 12가 되어야 합니다.

경찰서장은 홀수를 싫어해서 자신의 부서 번호는 짝수여야 합니다.

# --instructions--

모든 유효한 조합을 배열로 출력하는 프로그램을 작성하세요.

```js
[2, 3, 7] [2, 4, 6] [2, 6, 4]
[2, 7, 3] [4, 1, 7] [4, 2, 6]
[4, 3, 5] [4, 5, 3] [4, 6, 2]
[4, 7, 1] [6, 1, 5] [6, 2, 4]
[6, 4, 2] [6, 5, 1]
```

# --before-each--

```js
const nums = [1, 2, 3, 4, 5, 6, 7];
const total = 12;
const len = 14;
const result = [
  [2, 3, 7],
  [2, 4, 6],
  [2, 6, 4],
  [2, 7, 3],
  [4, 1, 7],
  [4, 2, 6],
  [4, 3, 5],
  [4, 5, 3],
  [4, 6, 2],
  [4, 7, 1],
  [6, 1, 5],
  [6, 2, 4],
  [6, 4, 2],
  [6, 5, 1]
];
```

# --hints--

`combinations`는 함수여야 합니다.

```js
assert(typeof combinations === 'function');
```

`combinations([1, 2, 3], 6)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(combinations([1, 2, 3], 6)));
```

`combinations([1, 2, 3, 4, 5, 6, 7], 12)`는 길이 14인 배열을 반환해야 합니다.

```js
assert(combinations(nums, total).length === len);
```

`combinations([1, 2, 3, 4, 5, 6, 7], 12)`는 모든 유효한 조합을 반환해야 합니다.

```js
assert.deepEqual(combinations(nums, total), result);
```

# --seed--

## --seed-contents--

```js
function combinations(possibleNumbers, total) {

  return true;
}
```

# --solutions--

```js
function combinations(possibleNumbers, total) {
  let firstNumber;
  let secondNumber;
  let thirdNumber;
  const allCombinations = [];

  for (let i = 0; i < possibleNumbers.length; i += 1) {
    firstNumber = possibleNumbers[i];

    if (firstNumber % 2 === 0) {
      for (let j = 0; j < possibleNumbers.length; j += 1) {
        secondNumber = possibleNumbers[j];

        if (j !== i && firstNumber + secondNumber <= total) {
          thirdNumber = total - firstNumber - secondNumber;

          if (thirdNumber !== firstNumber && thirdNumber !== secondNumber && possibleNumbers.includes(thirdNumber)) {
            allCombinations.push([firstNumber, secondNumber, thirdNumber]);
          }
        }
      }
    }
  }
  return allCombinations;
}
```
