---
id: 59f4eafba0343628bb682785
title: 디스코디안 날짜
challengeType: 1
forumTopicId: 302250
dashedName: discordian-date
---

# --description--
그레고리력은 12개월로 구성된 태양력으로, 각 달은 28~31일입니다. 한 주는 7일로 이루어져 있고, 1년은 52주와 1일로 구성됩니다. 1년은 365일이며, 윤년마다 2월에 하루가 추가됩니다. 윤년은 4년마다 발생하지만, 400년마다 처음 세 세기년은 제외됩니다.

<a href="http://www.rosettacode.org/wiki/Discordian_date" target="_blank" rel="noopener noreferrer nofollow">디스코디안 달력</a>은 그레고리력과 일치하며 1월 1일부터 시작합니다. 한 주는 5일로 구성되고, 1년은 73주입니다. 1년은 365일이며, 윤년마다 Chaos 59와 Chaos 60 사이에 하루가 추가됩니다.

디스코디안 달력의 계절, 요일, 사도 이벤트, 휴일은 다음과 같습니다:

계절: `'Chaos', 'Discord', 'Confusion', 'Bureaucracy', 'The Aftermath'`.
  
요일: `'Sweetmorn', 'Boomtime', 'Pungenday', 'Prickle-Prickle', 'Setting Orange'`.

사도: `'Mungday', 'Mojoday', 'Syaday', 'Zaraday', 'Maladay'`.

휴일: `'Chaoflux', 'Discoflux', 'Confuflux', 'Bureflux', 'Afflux'`.

# --instructions--

주어진 날짜를 그레고리력에서 디스코디안 달력으로 변환하세요.

디스코디안 달력에서 Chaos 1, 3188 YOLD는 그레고리력으로 2022년 1월 1일임을 참고하세요.

# --hints--

`discordianDate`는 함수여야 합니다.

```js
assert(typeof discordianDate === 'function');
```

`discordianDate(new Date(2010, 6, 22))`는 `"Pungenday, the 57th day of Confusion in the YOLD 3176"`를 반환해야 합니다.

```js
assert(
  discordianDate(new Date(2010, 6, 22)) ===
    'Pungenday, the 57th day of Confusion in the YOLD 3176'
);
```

`discordianDate(new Date(2012, 1, 28))`는 `"Prickle-Prickle, the 59th day of Chaos in the YOLD 3178"`를 반환해야 합니다.

```js
assert(
  discordianDate(new Date(2012, 1, 28)) ===
    'Prickle-Prickle, the 59th day of Chaos in the YOLD 3178'
);
```

`discordianDate(new Date(2012, 1, 29))`는 `"Setting Orange, the 60th day of Chaos in the YOLD 3178. Celebrate St. Tib\'s Day!"`를 반환해야 합니다.

```js
assert(
  discordianDate(new Date(2012, 1, 29)) ===
    "Setting Orange, the 60th day of Chaos in the YOLD 3178. Celebrate St. Tib's Day!"
);
```

`discordianDate(new Date(2012, 2, 1))`는 `"Setting Orange, the 60th day of Chaos in the YOLD 3178"`를 반환해야 합니다.

```js
assert(
  discordianDate(new Date(2012, 2, 1)) ===
    'Setting Orange, the 60th day of Chaos in the YOLD 3178'
);
```

`discordianDate(new Date(2010, 0, 5))`는 `"Setting Orange, the 5th day of Chaos in the YOLD 3176. Celebrate Mungday!"`를 반환해야 합니다.

```js
assert(
  discordianDate(new Date(2010, 0, 5)) ===
    'Setting Orange, the 5th day of Chaos in the YOLD 3176. Celebrate Mungday!'
);
```

`discordianDate(new Date(2011, 4, 3))`는 `"Pungenday, the 50th day of Discord in the YOLD 3177. Celebrate Discoflux!"`를 반환해야 합니다.

```js
assert(
  discordianDate(new Date(2011, 4, 3)) ===
    'Pungenday, the 50th day of Discord in the YOLD 3177. Celebrate Discoflux!'
);
```

`discordianDate(new Date(2015, 9, 19))`는 `"Boomtime, the 73rd day of Bureaucracy in the YOLD 3181"`를 반환해야 합니다.

```js
assert(
  discordianDate(new Date(2015, 9, 19)) ===
    'Boomtime, the 73rd day of Bureaucracy in the YOLD 3181'
);
```

# --seed--

## --seed-contents--

```js
function discordianDate(date) {

  return true;
}
```

# --solutions--

```js
/**
 * All Hail Discordia! - this script prints Discordian date using system date.
 *
 * lang: JavaScript
 * author: jklu
 * contributors: JamesMcGuigan
 *
 * source: https://rosettacode.org/wiki/Discordian_date#JavaScript
 */
const seasons = [
  'Chaos', 'Discord', 'Confusion',
  'Bureaucracy', 'The Aftermath'
];
const weekday = [
  'Sweetmorn', 'Boomtime', 'Pungenday',
  'Prickle-Prickle', 'Setting Orange'
];

const apostle = [
  'Mungday', 'Mojoday', 'Syaday',
  'Zaraday', 'Maladay'
];

const holiday = [
  'Chaoflux', 'Discoflux', 'Confuflux',
  'Bureflux', 'Afflux'
];


Date.prototype.isLeapYear = function() {
  const year = this.getFullYear();
  if ((year & 3) !== 0) { return false; }
  return ((year % 100) !== 0 || (year % 400) === 0);
};

// Get Day of Year
Date.prototype.getDOY = function() {
  const dayCount = [0, 31, 59, 90, 120, 151, 181, 212, 243, 273, 304, 334];
  const mn = this.getMonth();
  const dn = this.getDate();
  let dayOfYear = dayCount[mn] + dn;
  if (mn > 1 && this.isLeapYear()) { dayOfYear += 1; }
  return dayOfYear;
};

Date.prototype.isToday = function() {
  const today = new Date();
  return this.getDate() === today.getDate()
      && this.getMonth() === today.getMonth()
      && this.getFullYear() === today.getFullYear()
  ;
};

function discordianDate(date) {
  if (!date) { date = new Date(); }

  const y = date.getFullYear();
  const yold = y + 1166;
  let dayOfYear = date.getDOY();
  let celebrateHoliday = null;

  if (date.isLeapYear()) {
    if (dayOfYear === 60) {
      celebrateHoliday = 'St. Tib\'s Day';
    }
    else if (dayOfYear > 60) {
      dayOfYear--;
    }
  }
  dayOfYear--;

  const divDay = Math.floor(dayOfYear / 73);

  const seasonDay = (dayOfYear % 73) + 1;
  if (seasonDay === 5) {
    celebrateHoliday = apostle[divDay];
  }
  if (seasonDay === 50) {
    celebrateHoliday = holiday[divDay];
  }

  const season = seasons[divDay];
  const dayOfWeek = weekday[dayOfYear % 5];

  const nth = (seasonDay % 10 === 1) ? 'st'
          : (seasonDay % 10 === 2) ? 'nd'
          : (seasonDay % 10 === 3) ? 'rd'
                                  : 'th';

  return ''
         + dayOfWeek
         + ', the ' + seasonDay + nth
         + ' day of ' + season
         + ' in the YOLD ' + yold
         + (celebrateHoliday ? '. Celebrate ' + celebrateHoliday + '!' : '')
    ;
}
```
