---
id: 5a23c84252665b21eecc7e84
title: 최대 부분열 합
challengeType: 1
forumTopicId: 302278
dashedName: greatest-subsequential-sum
---

# --description--

정수 수열이 주어졌을 때, 그 요소들의 합이 최대가 되는 연속 부분열을 찾으세요. 즉, 다른 어떤 단일 부분열의 합보다 더 큰 값을 가지는 부분열을 찾아야 합니다.

빈 부분열은 합이 \\( 0 \\)으로 간주합니다. 따라서 모든 요소가 음수일 경우 결과는 빈 부분열이어야 합니다.

# --hints--

`maximumSubsequence`는 함수여야 합니다.

```js
assert(typeof maximumSubsequence == 'function');
```

`maximumSubsequence([ 1, 2, -1, 3, 10, -10 ])`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(maximumSubsequence([1, 2, -1, 3, 10, -10])));
```

`maximumSubsequence([ 1, 2, -1, 3, 10, -10 ])`는 `[ 1, 2, -1, 3, 10 ]`를 반환해야 합니다.

```js
assert.deepEqual(maximumSubsequence([1, 2, -1, 3, 10, -10]), [1, 2, -1, 3, 10]);
```

`maximumSubsequence([ 0, 8, 10, -2, -4, -1, -5, -3 ])`는 `[ 0, 8, 10 ]`를 반환해야 합니다.

```js
assert.deepEqual(maximumSubsequence([0, 8, 10, -2, -4, -1, -5, -3]), [
  0,
  8,
  10
]);
```

`maximumSubsequence([ 9, 9, -10, 1 ])`는 `[ 9, 9 ]`를 반환해야 합니다.

```js
assert.deepEqual(maximumSubsequence([9, 9, -10, 1]), [9, 9]);
```

`maximumSubsequence([ 7, 1, -5, -3, -8, 1 ])`는 `[ 7, 1 ]`를 반환해야 합니다.

```js
assert.deepEqual(maximumSubsequence([7, 1, -5, -3, -8, 1]), [7, 1]);
```

`maximumSubsequence([ -3, 6, -1, 4, -4, -6 ])`는 `[ 6, -1, 4 ]`를 반환해야 합니다.

```js
assert.deepEqual(maximumSubsequence([-3, 6, -1, 4, -4, -6]), [6, -1, 4]);
```

`maximumSubsequence([ -1, -2, 3, 5, 6, -2, -1, 4, -4, 2, -1 ])`는 `[ 3, 5, 6, -2, -1, 4 ]`를 반환해야 합니다.

```js
assert.deepEqual(maximumSubsequence([-1, -2, 3, 5, 6, -2, -1, 4, -4, 2, -1]), [
  3,
  5,
  6,
  -2,
  -1,
  4
]);
```

# --seed--

## --seed-contents--

```js
function maximumSubsequence(population) {

}
```

# --solutions--

```js
function maximumSubsequence(population) {
  function sumValues(arr) {
      var result = 0;
      for (var i = 0, len = arr.length; i < len; i++) {
          result += arr[i];
      }
      return result;
  }
  var greatest;
  var maxValue = 0;

  for (var i = 0, len = population.length; i < len; i++) {
      for (var j = i; j <= len; j++) {
          var subsequence = population.slice(i, j);
          var value = sumValues(subsequence);
          if (value > maxValue) {
              maxValue = value;
              greatest = subsequence;
          };
      }
  }

  return greatest;
}
```
