---
id: 5a23c84252665b21eecc7e80
title: 그레이 코드
challengeType: 1
forumTopicId: 302276
dashedName: gray-code
---

# --description--

그레이 코드는 연속하는 숫자 사이의 변화가 단 하나의 비트만 다른 이진 인코딩 방식입니다.

이 인코딩은 빠르게 변하는 값이나 느린 하드웨어에 입력으로 연결되는 값의 하드웨어 데이터 위험을 줄이는 데 유용합니다.

또한 카르노 맵의 입력을 왼쪽에서 오른쪽 또는 위에서 아래 순서로 생성할 때도 유용합니다.

# --instructions--

숫자를 그레이 코드로 인코딩하고 그레이 코드에서 숫자를 디코딩하는 함수를 만드세요. 함수는 2개의 매개변수를 가집니다.

첫 번째 매개변수는 불리언입니다. true일 때 인코딩하고 false일 때 디코딩해야 합니다. 두 번째 매개변수는 인코딩/디코딩할 숫자입니다.

모든 5비트 이진수(0부터 31까지, 앞의 0은 필요 없음)에 대해 일반 이진 표현, 그레이 코드 표현, 그리고 디코딩된 그레이 코드 값을 출력하세요.

그레이 코드는 여러 가지가 있지만, 다음은 "이진 반사 그레이 코드"라고 불리는 것을 인코딩하는 방법입니다.

인코딩 (MSB가 비트 0, b는 이진, g는 그레이 코드):

<pre>if b[i-1] = 1
  g[i] = not b[i]
else
  g[i] = b[i]
</pre>

또는:

<pre>g = b xor (b logically right shifted 1 time)
</pre>

디코딩 (MSB가 비트 0, b는 이진, g는 그레이 코드):

<pre>b[0] = g[0]<br>
for other bits:
b[i] = g[i] xor b[i-1]
</pre>

# --hints--

`gray`는 함수여야 합니다.

```js
assert(typeof gray == 'function');
```

`gray(true,177)`는 숫자를 반환해야 합니다.

```js
assert(typeof gray(true, 177) == 'number');
```

`gray(true,177)`는 `233`를 반환해야 합니다.

```js
assert.equal(gray(true, 177), 233);
```

`gray(true,425)`는 `381`를 반환해야 합니다.

```js
assert.equal(gray(true, 425), 381);
```

`gray(true,870)`는 `725`를 반환해야 합니다.

```js
assert.equal(gray(true, 870), 725);
```

`gray(false,233)`는 `177`를 반환해야 합니다.

```js
assert.equal(gray(false, 233), 177);
```

`gray(false,381)`는 `425`를 반환해야 합니다.

```js
assert.equal(gray(false, 381), 425);
```

`gray(false,725)`는 `870`를 반환해야 합니다.

```js
assert.equal(gray(false, 725), 870);
```

# --seed--

## --seed-contents--

```js
function gray(enc, number) {

}
```

# --solutions--

```js
function gray(enc, number){
  if(enc){
      return number ^ (number >> 1);
  }else{
      let n = number;

      while (number >>= 1) {
          n ^= number;
      }
      return n;
  }
}
```
