---
id: 5a23c84252665b21eecc801d
title: 문자 변화에 따라 문자열 나누기
challengeType: 1
forumTopicId: 302322
dashedName: split-a-character-string-based-on-change-of-character
---

# --description--

문자열을 문자 변화(왼쪽에서 오른쪽) 기준으로 쉼표(및 공백)로 구분된 문자열들로 나누세요. 공백은 다른 문자와 동일하게 처리해야 합니다(단, 명확하게 표시하기 어렵다는 점 제외). 쉼표도 마찬가지입니다. 예를 들어, 다음 문자열:

<pre>
"gHHH5YY++///\\"
</pre>

는 다음과 같이 나누어져야 합니다:

<pre>
["g", "HHH", "5", "YY", "++", "///", "\\" ];
</pre>

# --hints--

`split`는 함수여야 합니다.

```js
assert(typeof split == 'function');
```

`split("hello")`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(split('hello')));
```

`split("hello")`는 `["h", "e", "ll", "o"]`를 반환해야 합니다.

```js
assert.deepEqual(split('hello'), ['h', 'e', 'll', 'o']);
```

`split("commission")`는 `["c", "o", "mm", "i", "ss", "i", "o", "n"]`를 반환해야 합니다.

```js
assert.deepEqual(split('commission'), [
  'c',
  'o',
  'mm',
  'i',
  'ss',
  'i',
  'o',
  'n'
]);
```

`split("ssss----====llloooo")`는 `["ssss", "----", "====", "lll", "oooo"]`를 반환해야 합니다.

```js
assert.deepEqual(split('ssss----====llloooo'), [
  'ssss',
  '----',
  '====',
  'lll',
  'oooo'
]);
```

`split("sssmmmaaammmaaat")`는 `["sss", "mmm", "aaa", "mmm", "aaa", "t"]`를 반환해야 합니다.

```js
assert.deepEqual(split('sssmmmaaammmaaat'), [
  'sss',
  'mmm',
  'aaa',
  'mmm',
  'aaa',
  't'
]);
```

`split("gHHH5YY++///\\")`는 `["g", "HHH", "5", "YY", "++", "///", "\\"]`를 반환해야 합니다.

```js
assert.deepEqual(split('gHHH5YY++///\\'), [
  'g',
  'HHH',
  '5',
  'YY',
  '++',
  '///',
  '\\'
]);
```

# --seed--

## --seed-contents--

```js
function split(str) {

}
```

# --solutions--

```js
function split(str) {
  const concat = xs =>
    xs.length > 0
      ? (() => {
          const unit = typeof xs[0] === 'string' ? '' : [];
          return unit.concat.apply(unit, xs);
        })()
      : [];

  const group = xs => groupBy((a, b) => a === b, xs);

  const groupBy = (f, xs) => {
    const dct = xs.slice(1).reduce(
      (a, x) => {
        const h = a.active.length > 0 ? a.active[0] : undefined,
          blnGroup = h !== undefined && f(h, x);
        return {
          active: blnGroup ? a.active.concat([x]) : [x],
          sofar: blnGroup ? a.sofar : a.sofar.concat([a.active])
        };
      },
      {
        active: xs.length > 0 ? [xs[0]] : [],
        sofar: []
      }
    );
    return dct.sofar.concat(dct.active.length > 0 ? [dct.active] : []);
  };

  const map = (f, xs) => xs.map(f);

  const stringChars = s => s.split('');

  return map(concat, group(stringChars(str)));
}
```
