---
id: 5a23c84252665b21eecc7ec4
title: JortSort
challengeType: 1
forumTopicId: 302293
dashedName: jortsort
---

# --description--

jortSort는 사용자가 직접 작업하도록 하면서 다시는 정렬할 필요가 없기 때문에 효율성을 보장하는 정렬 도구 모음입니다. 원래는 Jenn "Moneydollars" Schiffer가 권위 있는 JSConf2014에서 발표했습니다.

jortSort는 비교 가능한 객체들의 단일 배열을 인수로 받는 함수여야 합니다. 그런 다음 배열을 오름차순으로 정렬하고 정렬된 배열을 원래 제공된 배열과 비교합니다. 배열이 일치하면(즉, 원래 배열이 이미 정렬되어 있으면) 함수는 true를 반환합니다. 배열이 일치하지 않으면(즉, 원래 배열이 정렬되어 있지 않으면) 함수는 false를 반환합니다.

# --hints--

`jortsort`는 함수여야 합니다.

```js
assert(typeof jortsort == 'function');
```

`jortsort([1,2,3,4,5])`는 불리언을 반환해야 합니다.

```js
assert(typeof jortsort([1, 2, 3, 4, 5]) == 'boolean');
```

`jortsort([1,2,3,4,5])`는 `true`를 반환해야 합니다.

```js
assert.equal(jortsort([1, 2, 3, 4, 5]), true);
```

`jortsort([1,2,13,4,5])`는 `false`를 반환해야 합니다.

```js
assert.equal(jortsort([1, 2, 13, 4, 5]), false);
```

`jortsort([12,4,51,2,4])`는 `false`를 반환해야 합니다.

```js
assert.equal(jortsort([12, 4, 51, 2, 4]), false);
```

`jortsort([1,2])`는 `true`를 반환해야 합니다.

```js
assert.equal(jortsort([1, 2]), true);
```

`jortsort([5,4,3,2,1])`는 `false`를 반환해야 합니다.

```js
assert.equal(jortsort([5, 4, 3, 2, 1]), false);
```

`jortsort([1,1,1,1,1])`는 `true`를 반환해야 합니다.

```js
assert.equal(jortsort([1, 1, 1, 1, 1]), true);
```

# --seed--

## --seed-contents--

```js
function jortsort(array) {

}
```

# --solutions--

```js
function jortsort(array) {
  // sort the array
  var originalArray = array.slice(0);
  array.sort( function(a,b){return a - b} );

  // compare to see if it was originally sorted
  for (var i = 0; i < originalArray.length; ++i) {
    if (originalArray[i] !== array[i]) return false;
  }

  return true;
};
```
