---
id: 5a23c84252665b21eecc7ede
title: 윤년
challengeType: 1
forumTopicId: 302300
dashedName: leap-year
---

# --description--

주어진 연도가 그레고리력에서 윤년인지 판단하세요.

# --hints--

`isLeapYear`는 함수여야 합니다.

```js
assert(typeof isLeapYear == 'function');
```

`isLeapYear()`은 불리언을 반환해야 합니다.

```js
assert(typeof isLeapYear(2018) == 'boolean');
```

`isLeapYear(2018)`는 `false`을 반환해야 합니다.

```js
assert.equal(isLeapYear(2018), false);
```

`isLeapYear(2016)`는 `true`를 반환해야 합니다.

```js
assert.equal(isLeapYear(2016), true);
```

`isLeapYear(2000)`는 `true`을 반환해야 합니다.

```js
assert.equal(isLeapYear(2000), true);
```

`isLeapYear(1900)`는 `false`를 반환해야 합니다.

```js
assert.equal(isLeapYear(1900), false);
```

`isLeapYear(1996)`는 `true`을 반환해야 합니다.

```js
assert.equal(isLeapYear(1996), true);
```

`isLeapYear(1800)`는 `false`을 반환해야 합니다.

```js
assert.equal(isLeapYear(1800), false);
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

}
```

# --solutions--

```js
function isLeapYear(year) {
  return year % 100 === 0 ? year % 400 === 0 : year % 4 === 0;
}
```
