---
id: 5a23c84252665b21eecc8045
title: 서덜랜드-호지먼 다각형 클리핑
challengeType: 1
forumTopicId: 302336
dashedName: sutherland-hodgman-polygon-clipping
---

# --description--

서덜랜드-호지먼 클리핑 알고리즘은 임의의 다각형(“주제 다각형”)과 볼록 다각형(“클립 다각형”) 사이의 교차하는 다각형을 찾습니다. 이 알고리즘은 컴퓨터 그래픽(특히 2D 그래픽)에서 장면의 복잡성을 줄이기 위해 다각형 중 표시할 필요가 없는 부분을 제거하는 데 사용됩니다. 다음 점들로 정의된 닫힌 다각형을 생각해 보세요:

<pre>[(50, 150), (200, 50), (350, 150), (350, 300), (250, 300), (200, 250), (150, 350), (100, 250), (100, 200)]</pre>

그리고 다음 점들로 정의된 사각형으로 클리핑하세요:

<pre>[(100, 100), (300, 100), (300, 300), (100, 300)]</pre>

# --instructions--

두 개의 배열을 매개변수로 받는 함수를 작성하세요. 첫 번째 배열은 주제 다각형의 점들을 포함하고 두 번째 배열은 클리핑 다각형의 점들을 포함합니다. 함수는 클리핑된 다각형의 점들을 포함하는 배열을 반환해야 합니다. 각 숫자는 소수점 셋째 자리까지 반올림하세요.

# --before-each--

```js
const subjectPolygon1 = [
  [50, 150],
  [200, 50],
  [350, 150],
  [350, 300],
  [250, 300],
  [200, 250],
  [150, 350],
  [100, 250],
  [100, 200]
];
const clipPolygon1 = [
  [100, 100],
  [300, 100],
  [300, 300],
  [100, 300]
];
const expected1 = [
  [100, 116.667],
  [125, 100],
  [275, 100],
  [300, 116.667],
  [300, 300],
  [250, 300],
  [200, 250],
  [175, 300],
  [125, 300],
  [100, 250]
];

const subjectPolygon2 = [
  [150, 200],
  [400, 450],
  [30, 50]
];
const clipPolygon2 = [
  [10, 10],
  [300, 200],
  [400, 600],
  [100, 300]
];
const expected2 = [
  [150, 200],
  [350, 400],
  [348.611, 394.444],
  [30, 50]
];

const subjectPolygon3 = [
  [250, 200],
  [100, 450],
  [130, 250]
];
const clipPolygon3 = [
  [50, 60],
  [100, 230],
  [400, 600],
  [100, 300]
];
const expected3 = [
  [129.167, 329.167],
  [119.565, 319.565],
  [121.854, 304.305]
];
```

# --hints--

`clip`는 함수여야 합니다.

```js
assert(typeof clip == 'function');
```

`clip([[50, 150], [200, 50], [350, 150], [350, 300], [250, 300], [200, 250], [150, 350], [100, 250], [100, 200]], [[100, 100], [300, 100], [300, 300], [100, 300]])`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(clip(subjectPolygon1, clipPolygon1)));
```

`clip([[50, 150], [200, 50], [350, 150], [350, 300], [250, 300], [200, 250], [150, 350], [100, 250], [100, 200]], [[100, 100], [300, 100], [300, 300], [100, 300]])`는 `[[100, 116.667], [125, 100], [275, 100], [300, 116.667], [300, 300], [250, 300], [200, 250], [175, 300], [125, 300], [100, 250]]`를 반환해야 합니다.

```js
assert.deepEqual(clip(subjectPolygon1, clipPolygon1), expected1);
```

`clip([[150, 200], [400, 450], [30, 50]], [[10, 10], [300, 200], [400, 600], [100, 300]])`는 `[[150, 200], [350, 400], [348.611, 394.444], [30, 50]]`를 반환해야 합니다.

```js
assert.deepEqual(clip(subjectPolygon2, clipPolygon2), expected2);
```

`clip([[250, 200], [100, 450], [130, 250]], [[50, 60], [100, 230], [400, 600], [100, 300]])`는 `[[129.167, 329.167], [119.565, 319.565], [121.854, 304.305]]`를 반환해야 합니다.

```js
assert.deepEqual(clip(subjectPolygon3, clipPolygon3), expected3);
```

# --seed--

## --seed-contents--

```js
function clip(subjectPolygon, clipPolygon) {

}
```

# --solutions--

```js
function clip(subjectPolygon, clipPolygon) {
  var cp1, cp2, s, e, i, j;
  var inside = function(p) {
    return (
      (cp2[0] - cp1[0]) * (p[1] - cp1[1]) > (cp2[1] - cp1[1]) * (p[0] - cp1[0])
    );
  };
  var intersection = function() {
    var dc = [cp1[0] - cp2[0], cp1[1] - cp2[1]],
      dp = [s[0] - e[0], s[1] - e[1]],
      n1 = cp1[0] * cp2[1] - cp1[1] * cp2[0],
      n2 = s[0] * e[1] - s[1] * e[0],
      n3 = 1.0 / (dc[0] * dp[1] - dc[1] * dp[0]);
    return [(n1 * dp[0] - n2 * dc[0]) * n3, (n1 * dp[1] - n2 * dc[1]) * n3];
  };
  var outputList = subjectPolygon;
  cp1 = clipPolygon[clipPolygon.length - 1];
  for (j in clipPolygon) {
    var cp2 = clipPolygon[j];
    var inputList = outputList;
    outputList = [];
    s = inputList[inputList.length - 1]; //last on the input list
    for (i in inputList) {
      var e = inputList[i];
      if (inside(e)) {
        if (!inside(s)) {
          outputList.push(intersection());
        }
        outputList.push(e);
      } else if (inside(s)) {
        outputList.push(intersection());
      }
      s = e;
    }
    cp1 = cp2;
  }
  return outputList.map(e => e.map(f => Math.round(f * 1000) / 1000));
}
```
