---
id: 5a23c84252665b21eecc7ec5
title: 요세푸스 문제
challengeType: 1
forumTopicId: 302294
dashedName: josephus-problem
---

# --description--

요세푸스 문제는 암울한 설명을 가진 수학 퍼즐입니다: $n$명의 죄수가 원형으로 서 있고, 0부터 $n-1$까지 순서대로 번호가 매겨져 있습니다.

집행관이 원을 따라 걸으며 죄수 0번부터 시작해 매 $k$번째 죄수를 제거하고 처형합니다.

이 과정이 계속되면서 원은 점점 작아지고, 결국 한 명의 죄수만 남게 되며 그 죄수는 풀려납니다.

예를 들어, 죄수가 $n=5$명이고 $k=2$일 때, 죄수가 처형되는 순서(이를 "처형 수열"이라 부르겠습니다)는 1, 3, 0, 4이고, 최종 생존자는 2번입니다.

임의의 $n, k > 0$가 주어졌을 때, 최종 생존하는 죄수가 누구인지 찾아내세요.

한 사건에서 죄수는 41명이었고, 매 3번째 죄수가 처형되었습니다 ($k=3$).

그중에 요세푸스라는 영리한 사람이 문제를 풀어 생존 위치에 섰고, 살아남아 이야기를 전했습니다.

그는 몇 번 죄수였을까요?

# --instructions--

처음 죄수 수와 `k`를 매개변수로 받아 최종 생존하는 죄수 번호를 반환하는 함수를 작성하세요.

# --hints--

`josephus`는 함수여야 합니다.

```js
assert(typeof josephus == 'function');
```

`josephus(30,3)`는 숫자를 반환해야 합니다.

```js
assert(typeof josephus(30, 3) == 'number');
```

`josephus(30,3)`는 `28`를 반환해야 합니다.

```js
assert.equal(josephus(30, 3), 28);
```

`josephus(30,5)`는 `2`를 반환해야 합니다.

```js
assert.equal(josephus(30, 5), 2);
```

`josephus(20,2)`는 `8`를 반환해야 합니다.

```js
assert.equal(josephus(20, 2), 8);
```

`josephus(17,6)`는 `1`를 반환해야 합니다.

```js
assert.equal(josephus(17, 6), 1);
```

`josephus(29,4)`는 `1`를 반환해야 합니다.

```js
assert.equal(josephus(29, 4), 1);
```

# --seed--

## --seed-contents--

```js
function josephus(init, kill) {

}
```

# --solutions--

```js
function josephus(init, kill) {
  const arr = Array.from(Array(init).keys());
  let curr = -1
  while (arr.length > 1) {
    curr = (curr + kill) % arr.length;
    arr.splice(curr, 1);
    curr--;
  }
  return arr[0];
}
```
