---
id: 5a23c84252665b21eecc8038
title: Subleq
challengeType: 1
forumTopicId: 302328
dashedName: subleq
---

# --description--

Subleq는 단일 명령어 집합 컴퓨터(OISC)의 예시입니다.

이 이름은 유일한 명령어인 **SU**btract and **B**ranch if **L**ess than or **EQ**ual to zero에서 따왔습니다.

이 기계를 에뮬레이트하는 인터프리터를 작성하는 것이 과제입니다.

기계의 메모리는 부호 있는 정수 배열로 구성됩니다. 적절한 워드 크기는 자유지만, 메모리는 양수와 음수를 모두 저장할 수 있어야 합니다.

실행은 명령어 포인터가 첫 번째 워드, 즉 주소 0을 가리키는 상태에서 시작합니다. 실행 과정은 다음과 같습니다:

<ol>
  <li>명령어 포인터가 가리키는 위치부터 연속된 세 워드에 저장된 값을 각각 A, B, C라고 합니다.</li>
  <li>명령어 포인터를 3 워드 앞으로 이동시켜 C가 저장된 주소 다음을 가리키게 합니다.</li>
  <li>A가 -1이면 표준 입력에서 문자를 읽고, 그 코드 포인트를 B가 가리키는 주소에 저장합니다. C는 사용하지 않습니다.</li>
  <li>B가 -1이면 A가 가리키는 주소에 저장된 숫자를 코드 포인트로 해석해 해당 문자를 출력합니다. C는 사용하지 않습니다.</li>
  <li>그 외에는 A와 B를 메모리 주소로 취급합니다. A가 가리키는 주소의 숫자를 B가 가리키는 주소의 숫자에서 빼고, 그 결과를 다시 B가 가리키는 주소에 저장합니다. 결과가 0 이하이면 C의 값이 새로운 명령어 포인터가 됩니다.</li>
  <li>명령어 포인터가 음수가 되면 실행을 중단합니다.</li>
</ol>

-1 이외의 다른 음수 주소는 -1과 동일하게 처리하거나 오류를 발생시켜도 됩니다.

작성한 해법은 기계에서 실행할 프로그램과 프로그램에 입력되는 데이터를 분리해서 받아야 합니다.

프로그램은 순수한 subleq "기계어"로, 공백으로 구분된 10진수 숫자들로 구성되어야 하며, 기호 이름이나 어셈블리 수준 확장 없이 주소 0부터 메모리에 로드됩니다. 이 "Hello, world!" 프로그램을 실행했을 때 출력 결과를 보여주세요. (예시는 ASCII 또는 라틴 문자 집합이나 유니코드 같은 상위 집합을 가정합니다. 구현 환경이 ASCII 호환이 아니면 다른 문자 집합으로 번역해도 됩니다.)

<pre>15 17 -1 17 -1 -1 16 1 -1 16 3 -1 15 15 0 0 -1 72 101 108 108 111 44 32 119 111 114 108 100 33 10 0</pre>

가상의 어셈블러 언어로는 다음과 비슷합니다:

<pre>start:
    zero, message, -1
    message, -1, -1
    neg1, start+1, -1
    neg1, start+3, -1
    zero, zero, start
zero: 0
neg1: -1
message: "Hello, world!\n\0"
</pre>

# --instructions--

정수 배열을 매개변수로 받는 함수를 작성하세요. 이 배열은 메모리 요소를 나타냅니다. 함수는 수열을 해석해 출력 문자열을 반환해야 합니다. 이 과제에서는 표준 입력이 없다고 가정하세요.

# --hints--

`Subleq`는 함수여야 합니다.

```js
assert(typeof Subleq == 'function');
```

`Subleq([15, 17, -1, 17, -1, -1, 16, 1, -1, 16, 3, -1, 15, 15, 0, 0, -1, 72, 101, 108, 108, 111, 44, 32, 119, 111, 114, 108, 100, 33, 0])`는 문자열을 반환해야 합니다.

```js
assert(
  typeof Subleq([
    15,
    17,
    -1,
    17,
    -1,
    -1,
    16,
    1,
    -1,
    16,
    3,
    -1,
    15,
    15,
    0,
    0,
    -1,
    72,
    101,
    108,
    108,
    111,
    44,
    32,
    119,
    111,
    114,
    108,
    100,
    33,
    0
  ]) == 'string'
);
```

`Subleq([15, 17, -1, 17, -1, -1, 16, 1, -1, 16, 3, -1, 15, 15, 0, 0, -1, 72, 101, 108, 108, 111, 44, 32, 119, 111, 114, 108, 100, 33, 0])`는 `"Hello, world!"`를 반환해야 합니다.

```js
assert.equal(
  Subleq([
    15,
    17,
    -1,
    17,
    -1,
    -1,
    16,
    1,
    -1,
    16,
    3,
    -1,
    15,
    15,
    0,
    0,
    -1,
    72,
    101,
    108,
    108,
    111,
    44,
    32,
    119,
    111,
    114,
    108,
    100,
    33,
    0
  ]),
  'Hello, world!'
);
```

# --seed--

## --seed-contents--

```js
function Subleq(mem) {

}
```

# --solutions--

```js
function Subleq(mem) {
  var out = '';
  var instructionPointer = 0;
  do {
    var a = mem[instructionPointer];
    var b = mem[instructionPointer + 1];
    if (a === -1) {
    } else if (b === -1) {
      out += String.fromCharCode(mem[a]);
    } else {
      mem[b] -= mem[a];
      if (mem[b] < 1) {
        instructionPointer = mem[instructionPointer + 2];
        continue;
      }
    }
    instructionPointer += 3;
  } while (instructionPointer >= 0);

  return out;
}
```
