---
id: 5a23c84252665b21eecc7edf
title: 최소공배수
challengeType: 1
forumTopicId: 302301
dashedName: least-common-multiple
---

# --description--

12와 18의 최소공배수는 36입니다. 12는 36의 약수입니다(12 × 3 = 36), 18도 36의 약수입니다(18 × 2 = 36). 그리고 36보다 작은 양의 정수 중에서 두 수 모두의 약수가 되는 수는 없습니다. 특별한 경우로, $m$ 또는 $n$이 0이면 최소공배수는 0입니다. 최소공배수를 계산하는 한 가지 방법은 $m$의 배수를 모두 반복해서 확인하다가 $n$의 배수이기도 한 수를 찾는 것입니다. 만약 <a href="https://rosettacode.org/wiki/Greatest_common_divisor" target="_blank" rel="noopener noreferrer nofollow">최대공약수</a>인 $gcd$를 이미 알고 있다면, 이 공식으로 $lcm$을 계산할 수 있습니다.

$$
\\operatorname{lcm}(m, n) = \\frac{|m \\times n|}{\\operatorname{gcd}(m, n)}
$$

# --instructions--

정수 배열의 최소공배수를 계산하세요. *m*과 *n*이 주어지면, 최소공배수는 *m*과 *n* 모두의 약수인 가장 작은 양의 정수입니다.

# --hints--

`LCM`는 함수여야 합니다.

```js
assert(typeof LCM == 'function');
```

`LCM([2, 4, 8])`는 숫자를 반환해야 합니다.

```js
assert(typeof LCM([2, 4, 8]) == 'number');
```

`LCM([2, 4, 8])`는 `8`를 반환해야 합니다.

```js
assert.equal(LCM([2, 4, 8]), 8);
```

`LCM([4, 8, 12])`는 `24`를 반환해야 합니다.

```js
assert.equal(LCM([4, 8, 12]), 24);
```

`LCM([3, 4, 5, 12, 40])`는 `120`를 반환해야 합니다.

```js
assert.equal(LCM([3, 4, 5, 12, 40]), 120);
```

`LCM([11, 33, 90])`는 `990`를 반환해야 합니다.

```js
assert.equal(LCM([11, 33, 90]), 990);
```

`LCM([-50, 25, -45, -18, 90, 447])`는 `67050`를 반환해야 합니다.

```js
assert.equal(LCM([-50, 25, -45, -18, 90, 447]), 67050);
```

# --seed--

## --seed-contents--

```js
function LCM(A) {

}
```

# --solutions--

```js
function LCM(A) {
  var n = A.length,
    a = Math.abs(A[0]);
  for (var i = 1; i < n; i++) {
    var b = Math.abs(A[i]),
      c = a;
    while (a && b) {
      a > b ? (a %= b) : (b %= a);
    }
    a = Math.abs(c * A[i]) / (a + b);
  }
  return a;
}
```
