---
id: 5a23c84252665b21eecc7edc
title: 매달 마지막 금요일
challengeType: 1
forumTopicId: 302299
dashedName: last-friday-of-each-month
---

# --description--

주어진 연도와 월에 대해 그 달의 마지막 금요일 날짜를 반환하는 함수를 작성하세요.

# --hints--

`lastFriday`는 함수여야 합니다.

```js
assert(typeof lastFriday == 'function');
```

`lastFriday(2018, 1)`은 숫자를 반환해야 합니다.

```js
assert(typeof lastFriday(2018, 1) == 'number');
```

`lastFriday(2018, 1)`는 `26`을 반환해야 합니다.

```js
assert.equal(lastFriday(2018, 1), 26);
```

`lastFriday(2017, 2)`는 `24`를 반환해야 합니다.

```js
assert.equal(lastFriday(2017, 2), 24);
```

`lastFriday(2012, 3)`는 `30`을 반환해야 합니다.

```js
assert.equal(lastFriday(2012, 3), 30);
```

`lastFriday(1900, 4)`는 `27`를 반환해야 합니다.

```js
assert.equal(lastFriday(1900, 4), 27);
```

`lastFriday(2000, 5)`는 `26`을 반환해야 합니다.

```js
assert.equal(lastFriday(2000, 5), 26);
```

`lastFriday(2006, 6)`는 `30`을 반환해야 합니다.

```js
assert.equal(lastFriday(2006, 6), 30);
```

`lastFriday(2010, 7)`는 `30`를 반환해야 합니다.

```js
assert.equal(lastFriday(2010, 7), 30);
```

`lastFriday(2005, 8)`는 `26`을 반환해야 합니다.

```js
assert.equal(lastFriday(2005, 8), 26);
```

# --seed--

## --seed-contents--

```js
function lastFriday(year, month) {

}
```

# --solutions--

```js
function lastFriday(year, month) {
  var i, last_day;
  i = 0;
  while (true) {
    last_day = new Date(year, month, i);
    if (last_day.getDay() === 5) {
      return last_day.getDate();
    }
    i -= 1;
  }
}
```
