---
id: 5a23c84252665b21eecc800b
title: 정렬 알고리즘/팬케이크 정렬
challengeType: 1
forumTopicId: 302315
dashedName: sorting-algorithmspancake-sort
---

# --description--

팬케이크 정렬을 사용해 정수 배열(편한 크기)을 오름차순으로 정렬하는 함수를 작성하세요. 함수는 정렬된 배열을 반환해야 합니다.

간단히 말해, 개별 요소를 정렬하는 대신 리스트 한쪽 끝을 뒤집는 "플립" 연산만 허용됩니다. 예를 들어:

<pre>정렬 전:
<b>6 7 8 9</b> 2 5 3 4 1<br>
정렬 후:
<b>9 8 7 6</b> 2 5 3 4 1
</pre>

리스트의 한쪽 끝만 뒤집을 수 있습니다; 보통은 낮은 쪽 끝이어야 하지만, 코딩하기 쉽거나 더 잘 작동한다면 높은 쪽 끝도 괜찮습니다. 단, 전체 해결 과정에서 뒤집는 끝은 반드시 동일해야 합니다. (뒤집는 끝을 임의로 바꿀 수 없습니다.)

# --hints--

`pancakeSort`는 함수여야 합니다.

```js
assert(typeof pancakeSort == 'function');
```

`pancakeSort([25, 32, 12, 7, 20])`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(pancakeSort([25, 32, 12, 7, 20])));
```

`pancakeSort([25, 32, 12, 7, 20])`는 `[7, 12, 20, 25, 32]`를 반환해야 합니다.

```js
assert.deepEqual(pancakeSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`pancakeSort([38, 45, 35, 8, 13])`는 `[8, 13, 35, 38, 45]`를 반환해야 합니다.

```js
assert.deepEqual(pancakeSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`pancakeSort([43, 36, 20, 34, 24])`는 `[20, 24, 34, 36, 43]`를 반환해야 합니다.

```js
assert.deepEqual(pancakeSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`pancakeSort([12, 33, 26, 18, 1, 16, 38])`는 `[1, 12, 16, 18, 26, 33, 38]`를 반환해야 합니다.

```js
assert.deepEqual(pancakeSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`pancakeSort([3, 39, 48, 16, 1, 4, 29])`는 `[1, 3, 4, 16, 29, 39, 48]`를 반환해야 합니다.

```js
assert.deepEqual(pancakeSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function pancakeSort(arr) {

}
```

# --solutions--

```js
function pancakeSort(arr) {
  for (var i = arr.length - 1; i >= 1; i--) {
    // find the index of the largest element not yet sorted
    var max_idx = 0;
    var max = arr[0];
    for (var j = 1; j <= i; j++) {
      if (arr[j] > max) {
        max = arr[j];
        max_idx = j;
      }
    }

    if (max_idx == i) continue; // element already in place

    var new_slice;

    // flip arr max element to index 0
    if (max_idx > 0) {
      new_slice = arr.slice(0, max_idx + 1).reverse();
      for (var j = 0; j <= max_idx; j++) arr[j] = new_slice[j];
    }

    // then flip the max element to its place
    new_slice = arr.slice(0, i + 1).reverse();
    for (var j = 0; j <= i; j++) arr[j] = new_slice[j];
  }
  return arr;
}
```
