---
id: 5a23c84252665b21eecc8040
title: 3과 5의 배수 합계
challengeType: 1
forumTopicId: 302332
dashedName: sum-multiples-of-3-and-5
---

# --description--

목표는 *n* 미만의 모든 3 또는 5의 양의 배수 합계를 찾는 함수를 작성하는 것입니다.

# --hints--

`sumMults`는 함수여야 합니다.

```js
assert(typeof sumMults == 'function');
```

`sumMults(10)`는 숫자를 반환해야 합니다.

```js
assert(typeof sumMults(10) == 'number');
```

`sumMults(10)`는 `23`를 반환해야 합니다.

```js
assert.equal(sumMults(10), 23);
```

`sumMults(100)`는 `2318`를 반환해야 합니다.

```js
assert.equal(sumMults(100), 2318);
```

`sumMults(1000)`는 `233168`를 반환해야 합니다.

```js
assert.equal(sumMults(1000), 233168);
```

`sumMults(10000)`는 `23331668`를 반환해야 합니다.

```js
assert.equal(sumMults(10000), 23331668);
```

`sumMults(100000)`는 `2333316668`를 반환해야 합니다.

```js
assert.equal(sumMults(100000), 2333316668);
```

# --seed--

## --seed-contents--

```js
function sumMults(n) {

}
```

# --solutions--

```js
function sumMults(n) {
  var sum = 0;
  for (var i = 1; i < n; i++) {
    if (i % 3 == 0 || i % 5 == 0) sum += i;
  }
  return sum;
}
```
