---
id: 5a23c84252665b21eecc7e7b
title: 제너레이터/지수
challengeType: 1
forumTopicId: 302275
dashedName: generatorexponential
---

# --description--

제너레이터는 함수나 프로시저처럼 실행 가능한 엔티티로, 한 번에 하나씩 값을 생성하는 코드를 포함하여 제너레이터를 호출할 때마다 수열의 다음 값을 제공합니다.

제너레이터는 내부 객체 상태가 "자연스럽게" 처리되도록 코루틴이나 객체 위에 구축되는 경우가 많습니다.

제너레이터는 수열이 잠재적으로 무한하고, 최소한의 상태만으로 수열의 다음 값을 구성할 수 있는 상황에서 자주 사용됩니다.

# --instructions--

제너레이터를 사용해 제곱수와 세제곱수를 생성하는 함수를 작성하세요. 제곱수 제너레이터에서 모든 세제곱수를 필터링하는 새 제너레이터를 만드세요.

함수는 필터링된 제너레이터의 \\( n^{th} \\) 값을 반환해야 합니다.

예를 들어 \\(n=7\\)일 때, 수열은 4, 9, 16, 25, 36, 49, 81이므로 함수는 81을 반환해야 합니다. 여기서 64는 세제곱수이므로 필터링됩니다.

# --hints--

`exponentialGenerator`는 함수여야 합니다.

```js
assert(typeof exponentialGenerator == 'function');
```

`exponentialGenerator()`은 숫자를 반환해야 합니다.

```js
assert(typeof exponentialGenerator(10) == 'number');
```

`exponentialGenerator(10)`는 `144`을 반환해야 합니다.

```js
assert.equal(exponentialGenerator(10), 144);
```

`exponentialGenerator(12)`는 `196`를 반환해야 합니다.

```js
assert.equal(exponentialGenerator(12), 196);
```

`exponentialGenerator(14)`는 `256`을 반환해야 합니다.

```js
assert.equal(exponentialGenerator(14), 256);
```

`exponentialGenerator(20)`는 `484`를 반환해야 합니다.

```js
assert.equal(exponentialGenerator(20), 484);
```

`exponentialGenerator(25)`는 `784`을 반환해야 합니다.

```js
assert.equal(exponentialGenerator(25), 784);
```

# --seed--

## --seed-contents--

```js
function exponentialGenerator(n) {

}
```

# --solutions--

```js
function exponentialGenerator(n){
  function* PowersGenerator(m) {
    var n=0;
    while(1) {
        yield Math.pow(n, m);
        n += 1;
    }
  }

  function* FilteredGenerator(g, f){
    var value = g.next().value;
    var filter = f.next().value;
    while(1) {
        if( value < filter ) {
            yield value;
            value = g.next().value;
        } else if ( value > filter ) {
            filter = f.next().value;
        } else {
            value = g.next().value;
            filter = f.next().value;
        }
    }
  }

  var squares = PowersGenerator(2);
  var cubes = PowersGenerator(3);

  var filtered = FilteredGenerator(squares, cubes);

  var curr=0;
  for(var i=0;i<n;i++) curr=filtered.next();

  return curr.value;
}
```
