---
id: 5a23c84252665b21eecc8017
title: 사운덱스
challengeType: 1
forumTopicId: 302320
dashedName: soundex
---

# --description--

**사운덱스 알고리즘**은 단어의 *의도*를 다룹니다. 비슷한 발음의 단어를 나타내는 표현을 만듭니다.

이 알고리즘은 <em>이름</em>과 <em>주소</em>를 검색하는 데 사용됩니다. 즉, <em>이름</em>을 입력한 사람이 <em>이름</em>의 철자를 고치기보다 발음에 집중할 수 있다는 뜻입니다.

예를 들어:

처음으로 `Quenci`라는 이름을 듣고 철자를 틀리게 입력해도, **사운덱스** 코드는 `Q520`가 됩니다.

다음에 `Quincy`라는 이름을 정확히 철자하면, 여전히 같은 코드 `Q520`가 나오므로, 여러 이름 발음을 같은 <em>사람</em>으로 연결할 수 있고 모든 철자를 추가할 필요가 없습니다.

규칙은 다음과 같습니다:

<ul>
  <li>모음(A, E, I, O, U)이 같은 사운덱스 코드를 가진 두 자음 사이에 있으면, 모음 오른쪽 자음을 코딩합니다. 예를 들어 Tymczak은 T-522로 코딩됩니다 (T, M은 5, C는 2, Z는 무시됨(위의 "나란히" 규칙 참조), K는 2). 모음 "A"가 Z와 K 사이에 있으므로 K가 코딩됩니다.</li>
  <li>"H"나 "W"가 같은 사운덱스 코드를 가진 두 자음 사이에 있으면, 모음 오른쪽 자음은 코딩하지 않습니다. 예를 들어 Ashcraft는 A-261로 코딩됩니다 (A, S는 2, C는 무시됨, R은 6, F는 1). A-226이 아닙니다.</li>
</ul>

# --instructions--

문자열을 매개변수로 받아 인코딩된 문자열을 반환하는 함수를 작성하세요.

# --hints--

`soundex`는 함수여야 합니다.

```js
assert(typeof soundex == 'function');
```

`soundex("Soundex")`는 문자열을 반환해야 합니다.

```js
assert(typeof soundex('Soundex') == 'string');
```

`soundex("Soundex")`는 `"S532"`를 반환해야 합니다.

```js
assert.equal(soundex('Soundex'), 'S532');
```

`soundex("Example")`는 `"E251"`를 반환해야 합니다.

```js
assert.equal(soundex('Example'), 'E251');
```

`soundex("Sownteks")`는 `"S532"`를 반환해야 합니다.

```js
assert.equal(soundex('Sownteks'), 'S532');
```

`soundex("Ekzampul")`는 `"E251"`를 반환해야 합니다.

```js
assert.equal(soundex('Ekzampul'), 'E251');
```

`soundex("Euler")`는 `"E460"`를 반환해야 합니다.

```js
assert.equal(soundex('Euler'), 'E460');
```

`soundex("Gauss")`는 `"G200"`를 반환해야 합니다.

```js
assert.equal(soundex('Gauss'), 'G200');
```

`soundex("Hilbert")`는 `"H416"`를 반환해야 합니다.

```js
assert.equal(soundex('Hilbert'), 'H416');
```

`soundex("Knuth")`는 `"K530"`를 반환해야 합니다.

```js
assert.equal(soundex('Knuth'), 'K530');
```

`soundex("Lloyd")`는 `"L300"`를 반환해야 합니다.

```js
assert.equal(soundex('Lloyd'), 'L300');
```

`soundex("Lukasiewicz")`는 `"L222"`를 반환해야 합니다.

```js
assert.equal(soundex('Lukasiewicz'), 'L222');
```

# --seed--

## --seed-contents--

```js
function soundex(s) {

}
```

# --solutions--

```js
function soundex(s) {
  var a = s.toLowerCase().split('');
  var f = a.shift(),
    r = '',
    codes = {
      a: '',
      e: '',
      i: '',
      o: '',
      u: '',
      b: 1,
      f: 1,
      p: 1,
      v: 1,
      c: 2,
      g: 2,
      j: 2,
      k: 2,
      q: 2,
      s: 2,
      x: 2,
      z: 2,
      d: 3,
      t: 3,
      l: 4,
      m: 5,
      n: 5,
      r: 6
    };
  r =
    f +
    a
      .map(function(v, i, a) {
        return codes[v];
      })
      .filter(function(v, i, a) {
        return i === 0 ? v !== codes[f] : v !== a[i - 1];
      })
      .join('');

  return (r + '000').slice(0, 4).toUpperCase();
}
```
