---
id: 5a23c84252665b21eecc8028
title: 스턴-브로코트 수열
challengeType: 1
forumTopicId: 302324
dashedName: stern-brocot-sequence
---

# --description--

이 도전 과제에서는 <a href="https://rosettacode.org/wiki/Fibonacci_sequence" target="_blank" rel="noopener noreferrer nofollow">피보나치 수열</a>을 생성하는 데 사용된 알고리즘과 유사한 방식으로 스턴-브로코트 수열을 생성합니다.

<ol>
  <li>수열의 첫 번째와 두 번째 항은 모두 1입니다:</li>
    <ul><li>1, 1</li></ul>
  <li>수열의 두 번째 항부터 시작하세요</li>
  <li>현재 고려 중인 항과 그 이전 항을 더합니다, (1 + 1) = 2, 그리고 수열 끝에 추가하세요:</li>
    <ul><li>1, 1, 2</li></ul>
  <li>현재 고려 중인 항을 수열 끝에 추가하세요:</li>
    <ul><li>1, 1, 2, 1</li></ul>
  <li>다음 항을 고려하세요, (세 번째 항 즉 2)</li>
  <li>3단계로 돌아가세요</li>
    <ul>
      <li></li>
      <li> ─── 반복을 확장하면 다음과 같습니다: ───</li>
      <li></li>
    </ul>
  <li>현재 고려 중인 항과 그 이전 항을 더합니다, (2 + 1) = 3, 그리고 수열 끝에 추가하세요:</li>
    <ul><li>1, 1, 2, 1, 3</li></ul>
  <li>현재 고려 중인 항을 수열 끝에 추가하세요:</li>
    <ul><li>1, 1, 2, 1, 3, 2</li></ul>
  <li>다음 항을 고려하세요, (네 번째 항 즉 1)</li>
</ol>

# --instructions--

위에 설명된 방법으로 생성된 스턴-브로코트 수열에서 $ n $이 처음 나타나는 위치를 반환하는 함수를 만드세요. 이 수열은 1부터 시작하는 인덱스를 사용합니다.

# --hints--

`sternBrocot`는 함수여야 합니다.

```js
assert(typeof sternBrocot == 'function');
```

`sternBrocot(2)`는 숫자를 반환해야 합니다.

```js
assert(typeof sternBrocot(2) == 'number');
```

`sternBrocot(2)`는 `3`를 반환해야 합니다.

```js
assert.equal(sternBrocot(2), 3);
```

`sternBrocot(3)`는 `5`를 반환해야 합니다.

```js
assert.equal(sternBrocot(3), 5);
```

`sternBrocot(5)`는 `11`를 반환해야 합니다.

```js
assert.equal(sternBrocot(5), 11);
```

`sternBrocot(7)`는 `19`를 반환해야 합니다.

```js
assert.equal(sternBrocot(7), 19);
```

`sternBrocot(10)`는 `39`를 반환해야 합니다.

```js
assert.equal(sternBrocot(10), 39);
```

# --seed--

## --seed-contents--

```js
function sternBrocot(num) {

}
```

# --solutions--

```js
function sternBrocot(num) {
  function f(n) {
    return n < 2
      ? n
      : n & 1
      ? f(Math.floor(n / 2)) + f(Math.floor(n / 2 + 1))
      : f(Math.floor(n / 2));
  }

  function gcd(a, b) {
    return a ? (a < b ? gcd(b % a, a) : gcd(a % b, b)) : b;
  }
  var n;
  for (n = 1; f(n) != num; n++);
  return n;
}
```
