---
id: 5a23c84252665b21eecc8014
title: 정렬 안정성
challengeType: 1
forumTopicId: 302308
dashedName: sort-stability
---

# --description--

특정 열이나 필드를 기준으로 표의 레코드를 정렬할 때, <a href="https://www.freecodecamp.org/news/stability-in-sorting-algorithms-a-treatment-of-equality-fa3140a5a539/" target="_blank" rel="noopener noreferrer nofollow">안정 정렬</a>은 동일한 키를 가진 레코드들의 상대적인 순서를 항상 유지합니다.

예를 들어, 이 국가와 도시의 표에서 두 번째 열인 도시를 기준으로 안정 정렬을 하면, 미국 버밍엄이 영국 버밍엄 위에 유지됩니다. (불안정 정렬은 이 경우 미국 버밍엄을 영국 버밍엄 위에 둘 수도 있지만, 안정 정렬은 이를 *보장*합니다).

<pre>UK  London
US  New York
US  Birmingham
UK  Birmingham
</pre>

마찬가지로, 첫 번째 열만 안정 정렬하면 "UK London"이 첫 번째 항목이 되고 "US Birmingham"이 마지막 항목이 됩니다(같은 첫 단어인 "UK" 또는 "US"를 가진 요소들의 순서가 유지되기 때문입니다).

# --instructions--

2차원 배열을 매개변수로 받는 함수를 작성하세요. 각 요소는 위 예시와 비슷하게 2개의 요소를 가집니다. 함수는 앞서 설명한 대로 배열을 정렬하고 정렬된 배열을 반환해야 합니다.

# --hints--

`stableSort`는 함수여야 합니다.

```js
assert(typeof stableSort == 'function');
```

`stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])`는 배열을 반환해야 합니다.

```js
assert(
  Array.isArray(
    stableSort([
      ['UK', 'London'],
      ['US', 'New York'],
      ['US', 'Birmingham'],
      ['UK', 'Birmingham']
    ])
  )
);
```

`stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])`는 `[["US", "Birmingham"], ["UK", "Birmingham"], ["UK", "London"], ["US", "New York"]]`를 반환해야 합니다.

```js
assert.deepEqual(
  stableSort([
    ['UK', 'London'],
    ['US', 'New York'],
    ['US', 'Birmingham'],
    ['UK', 'Birmingham']
  ]),
  [
    ['US', 'Birmingham'],
    ['UK', 'Birmingham'],
    ['UK', 'London'],
    ['US', 'New York']
  ]
);
```

`stableSort([[2, 2], [1, 2], [1, 4], [1, 5]])`는 `[[2, 2], [1, 2], [1, 4], [1, 5]]`를 반환해야 합니다.

```js
assert.deepEqual(
  stableSort([
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]),
  [
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]
);
```

`stableSort([[11, 55], [12, 45], [11, 45], [32, 45]])`는 `[[12, 45], [11, 45], [32, 45], [11, 55]]`를 반환해야 합니다.

```js
assert.deepEqual(
  stableSort([
    [11, 55],
    [12, 45],
    [11, 45],
    [32, 45]
  ]),
  [
    [12, 45],
    [11, 45],
    [32, 45],
    [11, 55]
  ]
);
```

`stableSort([[10, 22], [1, 2], [1, 4], [1, 5], [10, 9]])`는 `[[1, 2], [1, 4], [1, 5], [10, 9], [10, 22]]`를 반환해야 합니다.

```js
assert.deepEqual(
  stableSort([
    [10, 22],
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9]
  ]),
  [
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9],
    [10, 22]
  ]
);
```

`stableSort([[55, 54], [12, 22], [31, 43], [31, 54], [10, 49]])`는 `[[12, 22], [31, 43], [10, 49], [55, 54], [31, 54]]`를 반환해야 합니다.

```js
assert.deepEqual(
  stableSort([
    [55, 54],
    [12, 22],
    [31, 43],
    [31, 54],
    [10, 49]
  ]),
  [
    [12, 22],
    [31, 43],
    [10, 49],
    [55, 54],
    [31, 54]
  ]
);
```

# --seed--

## --seed-contents--

```js
function stableSort(arr) {

}
```

# --solutions--

```js
function stableSort(arr) {
  arr.sort(function(a, b) {
    return a[1] < b[1] ? -1 : a[1] > b[1] ? 1 : 0;
  });
  return arr;
}
```
