---
id: 5a23c84252665b21eecc7e7a
title: 소문자 ASCII 알파벳 생성
challengeType: 1
forumTopicId: 302274
dashedName: generate-lower-case-ascii-alphabet
---

# --description--

주어진 범위에 대해 소문자 ASCII 문자 배열을 생성하는 함수를 작성하세요. 예를 들어, 범위 `['a', 'd']`가 주어지면 함수는 `['a', 'b', 'c', 'd']`을 반환해야 합니다.

# --before-each--

```js
let results=[
  [ 'a', 'b', 'c', 'd' ],
  [ 'c', 'd', 'e', 'f', 'g', 'h', 'i' ],
  [ 'm', 'n', 'o', 'p', 'q' ],
  [ 'k', 'l', 'm', 'n' ],
  [ 't', 'u', 'v', 'w', 'x', 'y', 'z' ]
]
```

# --hints--

`lascii`는 함수여야 합니다.

```js
assert(typeof lascii == 'function');
```

`lascii("a","d")`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(lascii('a', 'd')));
```

`lascii('a','d')`는 `[ 'a', 'b', 'c', 'd' ]`를 반환해야 합니다.

```js
assert.deepEqual(lascii('a', 'd'), results[0]);
```

`lascii('c','i')`는 `[ 'c', 'd', 'e', 'f', 'g', 'h', 'i' ]`를 반환해야 합니다.

```js
assert.deepEqual(lascii('c', 'i'), results[1]);
```

`lascii('m','q')`는 `[ 'm', 'n', 'o', 'p', 'q' ]`를 반환해야 합니다.

```js
assert.deepEqual(lascii('m', 'q'), results[2]);
```

`lascii('k','n')`는 `[ 'k', 'l', 'm', 'n' ]`를 반환해야 합니다.

```js
assert.deepEqual(lascii('k', 'n'), results[3]);
```

`lascii('t','z')`는 `[ 't', 'u', 'v', 'w', 'x', 'y', 'z' ]`를 반환해야 합니다.

```js
assert.deepEqual(lascii('t', 'z'), results[4]);
```

# --seed--

## --seed-contents--

```js
function lascii(cFrom, cTo) {

}
```

# --solutions--

```js
function lascii(cFrom, cTo) {

  function cRange(cFrom, cTo) {
    var iStart = cFrom.charCodeAt(0);

    return Array.apply(
      null, Array(cTo.charCodeAt(0) - iStart + 1)
    ).map(function (_, i) {

      return String.fromCharCode(iStart + i);

    });
  }

  return cRange(cFrom, cTo);

}
```
