---
id: 5a23c84252665b21eecc8010
title: 셸 정렬 알고리즘
challengeType: 1
forumTopicId: 302317
dashedName: sorting-algorithmsshell-sort
---

# --description--

셸 정렬 알고리즘을 사용해 요소 배열을 정렬하는 함수를 작성하세요. 이 알고리즘은 점차 감소하는 간격 정렬입니다. 함수는 정렬된 배열을 반환해야 합니다.

셸 정렬(또는 Shellsort, Shell의 방법)은 1959년에 알고리즘을 발표한 발명가 Donald Shell의 이름을 따서 명명되었습니다.

셸 정렬은 간격 수열을 기반으로 한 삽입 정렬들의 교차 수열입니다. 각 반복마다 간격 크기가 줄어들며, 최종적으로 간격 크기가 1이 될 때까지 진행합니다.

간격 크기가 1일 때는 기본 삽입 정렬이 되지만, 이 시점에서는 데이터가 거의 정렬되어 있어 삽입 정렬의 "최선의 경우"가 됩니다.

간격 수열이 1로 끝나기만 하면 어떤 수열이든 데이터를 정렬할 수 있지만, 일부 수열이 더 효율적입니다.

실험적 연구에 따르면 약 2.2의 비율을 가진 기하급수적 간격 수열이 실제로 잘 작동하는 것으로 나타났습니다.

# --hints--

`shellSort`는 함수여야 합니다.

```js
assert(typeof shellSort == 'function');
```

`shellSort([25, 32, 12, 7, 20])`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(shellSort([25, 32, 12, 7, 20])));
```

`shellSort([25, 32, 12, 7, 20])`는 `[7, 12, 20, 25, 32]`를 반환해야 합니다.

```js
assert.deepEqual(shellSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`shellSort([38, 45, 35, 8, 13])`는 `[8, 13, 35, 38, 45]`를 반환해야 합니다.

```js
assert.deepEqual(shellSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`shellSort([43, 36, 20, 34, 24])`는 `[20, 24, 34, 36, 43]`를 반환해야 합니다.

```js
assert.deepEqual(shellSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`shellSort([12, 33, 26, 18, 1, 16, 38])`는 `[1, 12, 16, 18, 26, 33, 38]`를 반환해야 합니다.

```js
assert.deepEqual(shellSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`shellSort([3, 39, 48, 16, 1, 4, 29])`는 `[1, 3, 4, 16, 29, 39, 48]`를 반환해야 합니다.

```js
assert.deepEqual(shellSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function shellSort(a) {

}
```

# --solutions--

```js
function shellSort(a) {
  for (var h = a.length; h > 0; h = parseInt(h / 2)) {
    for (var i = h; i < a.length; i++) {
      var k = a[i];
      for (var j = i; j >= h && k < a[j - h]; j -= h) a[j] = a[j - h];
      a[j] = k;
    }
  }
  return a;
}
```
