---
id: 5a23c84252665b21eecc7eb1
title: 항등 행렬
challengeType: 1
forumTopicId: 302290
dashedName: identity-matrix
---

# --description--

*항등 행렬*은 크기가 \\( n \times n \\)인 정사각 행렬로, 대각선 요소는 모두 `1`(1)이고 나머지 모든 요소는 `0`(0)입니다.

<ul>
  <li style='list-style: none;'>\(\displaystyle I_{n}=\begin{bmatrix} 1 &#x26; 0 &#x26; 0 \cr 0 &#x26; 1 &#x26; 0 \cr 0 &#x26; 0 &#x26; 1 \cr \end{bmatrix}\)</li>
</ul>

# --instructions--

숫자 `n`를 매개변수로 받아 \\( n \times n \\) 크기의 항등 행렬을 반환하는 함수를 작성하세요.

# --before-each--

```js
let results=[[ [ 1 ] ],
[ [ 1, 0 ], [ 0, 1 ] ],
[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ],
[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]]
```

# --hints--

`idMatrix`는 함수여야 합니다.

```js
assert(typeof idMatrix == 'function');
```

`idMatrix(1)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(idMatrix(1)));
```

`idMatrix(1)`는 `[ [ 1 ] ]`를 반환해야 합니다.

```js
assert.deepEqual(idMatrix(1), results[0]);
```

`idMatrix(2)`는 `[ [ 1, 0 ], [ 0, 1 ] ]`를 반환해야 합니다.

```js
assert.deepEqual(idMatrix(2), results[1]);
```

`idMatrix(3)`는 `[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ]`를 반환해야 합니다.

```js
assert.deepEqual(idMatrix(3), results[2]);
```

`idMatrix(4)`는 `[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]`를 반환해야 합니다.

```js
assert.deepEqual(idMatrix(4), results[3]);
```

# --seed--

## --seed-contents--

```js
function idMatrix(n) {

}
```

# --solutions--

```js
function idMatrix(n) {
    return Array.apply(null, new Array(n)).map(function (x, i, xs) {
        return xs.map(function (_, k) {
            return i === k ? 1 : 0;
        })
    });
}
```
