---
id: 5a23c84252665b21eecc8046
title: 대칭 차집합
challengeType: 1
forumTopicId: 16086
dashedName: symmetric-difference
---

# --description--

두 집합 *A*와 *B*가 주어졌을 때, $(A \\setminus B) \\cup (B \\setminus A)$를 계산하세요. 즉, *A* 또는 *B*에는 속하지만 둘 다에는 속하지 않는 항목들을 나열하세요. 이 집합을 *A*와 *B*의 대칭 차집합이라고 합니다. 다시 말해, $(A \\cup B) \\setminus (A \\cap B)$ (적어도 하나의 집합 *A* 또는 *B*에 속하는 항목들의 집합에서 *A*와 *B* 모두에 속하는 항목들의 집합을 뺀 것)입니다.

예시:

집합 `A = [1, 2, 3]`와 `B = [1, 3, 4]`에 대해, *A*와 *B*의 대칭 차집합은 `[2, 4]`입니다.

# --instructions--

두 개의 배열을 매개변수로 받아 대칭 차집합을 반환하는 함수를 작성하세요. 반환하기 전에 결과 배열을 정렬하세요.

# --hints--

`symmetricDifference`는 함수여야 합니다.

```js
assert(typeof symmetricDifference == 'function');
```

`symmetricDifference(["John", "Bob", "Mary", "Serena"], ["Jim", "Mary", "John", "Bob"])`는 배열을 반환해야 합니다.

```js
assert(
  Array.isArray(
    symmetricDifference(
      ['John', 'Bob', 'Mary', 'Serena'],
      ['Jim', 'Mary', 'John', 'Bob']
    )
  )
);
```

`symmetricDifference(["John", "Bob", "Mary", "Serena"], ["Jim", "Mary", "John", "Bob"])`는 `["Jim", "Serena"]`를 반환해야 합니다.

```js
assert.deepEqual(
  symmetricDifference(
    ['John', 'Bob', 'Mary', 'Serena'],
    ['Jim', 'Mary', 'John', 'Bob']
  ),
  ['Jim', 'Serena']
);
```

`symmetricDifference([1, 2, 3], [3, 4])`는 `[1, 2, 4]`를 반환해야 합니다.

```js
assert.deepEqual(symmetricDifference([1, 2, 3], [3, 4]), [1, 2, 4]);
```

`symmetricDifference([1, 2, 3, 4, 5], [3, 4, 8, 7])`는 `[1, 2, 5, 7, 8]`를 반환해야 합니다.

```js
assert.deepEqual(symmetricDifference([1, 2, 3, 4, 5], [3, 4, 8, 7]), [
  1,
  2,
  5,
  7,
  8
]);
```

`symmetricDifference([1, 2, 3, 4, 5, 6, 7, 8], [1, 3, 5, 6, 7, 8, 9])`는 `[2, 4, 9]`를 반환해야 합니다.

```js
assert.deepEqual(
  symmetricDifference([1, 2, 3, 4, 5, 6, 7, 8], [1, 3, 5, 6, 7, 8, 9]),
  [2, 4, 9]
);
```

`symmetricDifference([1, 2, 4, 7, 9], [2, 3, 7, 8, 9])`는 `[1, 3, 4, 8]`를 반환해야 합니다.

```js
assert.deepEqual(symmetricDifference([1, 2, 4, 7, 9], [2, 3, 7, 8, 9]), [
  1,
  3,
  4,
  8
]);
```

# --seed--

## --seed-contents--

```js
function symmetricDifference(A, B) {

}
```

# --solutions--

```js
function symmetricDifference(A, B) {
  function relative_complement(A, B) {
    return A.filter(function(elem) {
      return B.indexOf(elem) == -1;
    });
  }

  function unique(ary) {
    var u = ary.concat().sort();
    for (var i = 1; i < u.length; ) {
      if (u[i - 1] === u[i]) u.splice(i, 1);
      else i++;
    }
    return u;
  }

  return unique(
    relative_complement(A, B).concat(relative_complement(B, A))
  ).sort();
}
```
