---
id: 5a23c84252665b21eecc8042
title: 제곱합
challengeType: 1
forumTopicId: 302334
dashedName: sum-of-squares
---

# --description--

정수 배열의 제곱합을 구하는 함수를 작성하세요.

# --hints--

`sumsq`는 함수여야 합니다.

```js
assert(typeof sumsq == 'function');
```

`sumsq([1, 2, 3, 4, 5])`은 숫자를 반환해야 합니다.

```js
assert(typeof sumsq([1, 2, 3, 4, 5]) == 'number');
```

`sumsq([1, 2, 3, 4, 5])`는 `55`을 반환해야 합니다.

```js
assert.equal(sumsq([1, 2, 3, 4, 5]), 55);
```

`sumsq([25, 32, 12, 7, 20])`는 `2242`를 반환해야 합니다.

```js
assert.equal(sumsq([25, 32, 12, 7, 20]), 2242);
```

`sumsq([38, 45, 35, 8, 13])`는 `4927`을 반환해야 합니다.

```js
assert.equal(sumsq([38, 45, 35, 8, 13]), 4927);
```

`sumsq([43, 36, 20, 34, 24])`는 `5277`를 반환해야 합니다.

```js
assert.equal(sumsq([43, 36, 20, 34, 24]), 5277);
```

`sumsq([12, 33, 26, 18, 1, 16, 3])`는 `2499`을 반환해야 합니다.

```js
assert.equal(sumsq([12, 33, 26, 18, 1, 16, 3]), 2499);
```

# --seed--

## --seed-contents--

```js
function sumsq(array) {

}
```

# --solutions--

```js
function sumsq(array) {
  var sum = 0;
  var i, iLen;

  for (i = 0, iLen = array.length; i < iLen; i++) {
    sum += array[i] * array[i];
  }
  return sum;
}
```
