---
id: 5a23c84252665b21eecc8041
title: 수열의 합
challengeType: 1
forumTopicId: 302333
dashedName: sum-of-a-series
---

# --description--

<em>수열</em>의 처음 **n**개 항의 합, 즉 <em>수열</em>의 **n**번째 항까지의 합을 계산하세요. 비공식적으로 이 값이나 **n**이 무한대로 갈 때의 극한값을 *수열의 합*이라고도 하며, 이 도전 과제의 제목이기도 합니다. 이 과제에서는 $S_n = \displaystyle\sum_{k=1}^n \frac{1}{k^2}$를 사용하세요.

# --instructions--

$a$와 $b$를 매개변수로 받아 수열의 $a$번째 항부터 $b$번째 항까지의 합을 반환하는 함수를 작성하세요.

# --hints--

`sum`는 함수여야 합니다.

```js
assert(typeof sum == 'function');
```

`sum(1, 100)`는 숫자를 반환해야 합니다.

```js
assert(typeof sum(1, 100) == 'number');
```

`sum(1, 100)`는 `1.6349839001848923`를 반환해야 합니다.

```js
assert.equal(sum(1, 100), 1.6349839001848923);
```

`sum(33, 46)`는 `0.009262256361481223`를 반환해야 합니다.

```js
assert.equal(sum(33, 46), 0.009262256361481223);
```

`sum(21, 213)`는 `0.044086990748706555`를 반환해야 합니다.

```js
assert.equal(sum(21, 213), 0.044086990748706555);
```

`sum(11, 111)`는 `0.08619778593108679`를 반환해야 합니다.

```js
assert.equal(sum(11, 111), 0.08619778593108679);
```

`sum(1, 10)`는 `1.5497677311665408`를 반환해야 합니다.

```js
assert.equal(sum(1, 10), 1.5497677311665408);
```

# --seed--

## --seed-contents--

```js
function sum(a, b) {

}
```

# --solutions--

```js
function sum(a, b) {
  function fn(x) {
    return 1 / (x * x);
  }
  var s = 0;
  for (; a <= b; a++) s += fn(a);
  return s;
}
```
