---
id: 5a23c84252665b21eecc7eb0
title: C 다음이 아니면 I가 E보다 먼저 온다
challengeType: 1
forumTopicId: 302288
dashedName: i-before-e-except-after-c
---

# --description--

**<a href="http://www.rosettacode.org/wiki/I_before_E_except_after_C" target="_blank" rel="noopener noreferrer nofollow">"C 다음이 아니면 I가 E보다 먼저 온다"</a>**는 영어 철자법의 일반 규칙입니다. 어떤 단어가 `ei` 또는 `ie`로 철자되는지 확실하지 않을 때, 이 운율은 앞 글자가 `c`가 아니면 `ie` 순서가 맞고, 만약 앞 글자가 `c`라면 `ei`일 수 있다고 제안합니다.

제공된 단어들을 사용해, 문구의 두 하위 절이 각각 타당한지 확인하세요:

<ol>
  <li>
    <i>"C가 앞에 없으면 I가 E보다 먼저 온다".</i>
  </li>
  <li>
    <i>"C가 앞에 있으면 E가 I보다 먼저 온다".</i>
  </li>
</ol>

두 하위 문구가 모두 타당하다면 원래 문구도 타당하다고 할 수 있습니다.

# --instructions--

단어를 입력받아 이 규칙을 따르는지 확인하는 함수를 작성하세요. 단어가 규칙을 따르면 true를, 그렇지 않으면 false를 반환해야 합니다.

# --hints--

`IBeforeExceptC`는 함수여야 합니다.

```js
assert(typeof IBeforeExceptC == 'function');
```

`IBeforeExceptC("receive")`는 불리언을 반환해야 합니다.

```js
assert(typeof IBeforeExceptC('receive') == 'boolean');
```

`IBeforeExceptC("receive")`는 `true`를 반환해야 합니다.

```js
assert.equal(IBeforeExceptC('receive'), true);
```

`IBeforeExceptC("science")`는 `false`를 반환해야 합니다.

```js
assert.equal(IBeforeExceptC('science'), false);
```

`IBeforeExceptC("imperceivable")`는 `true`를 반환해야 합니다.

```js
assert.equal(IBeforeExceptC('imperceivable'), true);
```

`IBeforeExceptC("inconceivable")`는 `true`를 반환해야 합니다.

```js
assert.equal(IBeforeExceptC('inconceivable'), true);
```

`IBeforeExceptC("insufficient")`는 `false`를 반환해야 합니다.

```js
assert.equal(IBeforeExceptC('insufficient'), false);
```

`IBeforeExceptC("omniscient")`는 `false`를 반환해야 합니다.

```js
assert.equal(IBeforeExceptC('omniscient'), false);
```

# --seed--

## --seed-contents--

```js
function IBeforeExceptC(word) {

}
```

# --solutions--

```js
function IBeforeExceptC(word)
{
    if(word.indexOf("c")==-1 && word.indexOf("ie")!=-1)
        return true;
    else if(word.indexOf("cei")!=-1)
        return true;
    return false;
}
```
