---
id: 5a23c84252665b21eecc7e78
title: 일반화된 FizzBuzz
challengeType: 1
forumTopicId: 302273
dashedName: general-fizzbuzz
---

# --description--

어떤 인수 목록과 그에 대응하는 단어들에 대해 작동하는 일반화된 <a href="https://rosettacode.org/wiki/FizzBuzz" target="_blank" rel="noopener noreferrer nofollow">FizzBuzz</a>를 작성하세요.

이는 기본적으로 게임 규칙을 사용자에게 제공하는 "fizzbuzz" 구현입니다. 이를 구현하는 함수를 만드세요. 함수는 두 개의 매개변수를 받아야 합니다.

첫 번째는 FizzBuzz 규칙이 담긴 배열입니다. 예를 들어: `[ [3, "Fizz"] , [5, "Buzz"] ]`.

이것은 숫자가 3의 배수일 때 `Fizz`을 출력하고, 5의 배수일 때 `Buzz`를 출력해야 함을 나타냅니다. 만약 두 배수 모두에 해당하면 배열에 지정된 순서대로 문자열을 이어 붙여야 합니다. 이 경우, 숫자가 3과 5의 배수일 때 `FizzBuzz`입니다.

두 번째 매개변수는 위에서 설명한 문자열을 반환해야 하는 숫자입니다.

# --hints--

`genFizzBuzz`는 함수여야 합니다.

```js
assert(typeof genFizzBuzz == 'function');
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 6)`는 문자열을 반환해야 합니다.

```js
assert(
  typeof genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    6
  ) == 'string'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 6)`는 `"Fizz"`을 반환해야 합니다.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    6
  ),
  'Fizz'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 10)`는 `"Buzz"`을 반환해야 합니다.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    10
  ),
  'Buzz'
);
```

`genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 12)`는 `"Buzz"`을 반환해야 합니다.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    12
  ),
  'Buzz'
);
```

`genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 13)`는 `"13"`을 반환해야 합니다.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    13
  ),
  '13'
);
```

`genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 15)`는 `"BuzzFizz"`을 반환해야 합니다.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    15
  ),
  'BuzzFizz'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 15)`는 `"FizzBuzz"`을 반환해야 합니다.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    15
  ),
  'FizzBuzz'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"],[7, "Baxx"]], 105)`는 `"FizzBuzzBaxx"`을 반환해야 합니다.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz'],
      [7, 'Baxx']
    ],
    105
  ),
  'FizzBuzzBaxx'
);
```

# --seed--

## --seed-contents--

```js
function genFizzBuzz(rules, num) {

}
```

# --solutions--

```js
function genFizzBuzz(rules, num) {
  let res='';
  rules.forEach(function (e) {
    if(num % e[0] == 0)
      res+=e[1];
  })

  if(res==''){
    res=num.toString();
  }

  return res;
}
```
