---
id: 5a23c84252665b21eecc7e82
title: 최대공약수
challengeType: 1
forumTopicId: 302277
dashedName: greatest-common-divisor
---

# --description--

두 정수의 최대공약수를 반환하는 함수를 작성하세요.

# --hints--

`gcd`는 함수여야 합니다.

```js
assert(typeof gcd == 'function');
```

`gcd(24,36)`는 숫자를 반환해야 합니다.

```js
assert(typeof gcd(24, 36) == 'number');
```

`gcd(24,36)`는 `12`를 반환해야 합니다.

```js
assert.equal(gcd(24, 36), 12);
```

`gcd(30,48)`는 `6`를 반환해야 합니다.

```js
assert.equal(gcd(30, 48), 6);
```

`gcd(10,15)`는 `5`를 반환해야 합니다.

```js
assert.equal(gcd(10, 15), 5);
```

`gcd(100,25)`는 `25`를 반환해야 합니다.

```js
assert.equal(gcd(100, 25), 25);
```

`gcd(13,250)`는 `1`를 반환해야 합니다.

```js
assert.equal(gcd(13, 250), 1);
```

`gcd(1300,250)`는 `50`를 반환해야 합니다.

```js
assert.equal(gcd(1300, 250), 50);
```

# --seed--

## --seed-contents--

```js
function gcd(a, b) {

}
```

# --solutions--

```js
function gcd(a, b) {
  return b==0 ? Math.abs(a):gcd(b, a % b);
}
```
