---
id: 5a23c84252665b21eecc8043
title: 100이 되도록 더하기
challengeType: 1
forumTopicId: 302335
dashedName: sum-to-100
---

# --description--

*100이 되도록 더하기* 퍼즐의 해답을 찾으세요.

십진수 숫자 문자열 **123456789**의 숫자들 앞에 수학 연산자 **+** 또는 **─** (더하기 또는 빼기)를 넣어 결과 수학 표현식의 합이 특정 값(이 상징적인 경우에는 **100**)이 되도록 하세요.

예시:

<pre><b>123 + 4 - 5 + 67 - 89   =   100</b></pre>

# --instructions--

숫자를 매개변수로 받는 함수를 작성하세요. 이 함수는 주어진 숫자에 대한 모든 해답을 포함하는 배열을 반환해야 합니다. 해답은 표현식을 나타내는 문자열이어야 합니다. 예를 들어: "1+23-456+78-9". 반환하기 전에 배열을 정렬하세요.

# --hints--

`sumTo100`는 함수여야 합니다.

```js
assert(typeof sumTo100 == 'function');
```

`sumTo100(199)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(sumTo100(199)));
```

`sumTo100(199)`는 `["-1+2-3+45+67+89", "123-4+5+6+78-9", "123-4+56+7+8+9"]`를 반환해야 합니다.

```js
assert.deepEqual(sumTo100(199), [
  '-1+2-3+45+67+89',
  '123-4+5+6+78-9',
  '123-4+56+7+8+9'
]);
```

`sumTo100(209)`는 `["1+234+56+7-89"]`를 반환해야 합니다.

```js
assert.deepEqual(sumTo100(209), ['1+234+56+7-89']);
```

`sumTo100(243)`는 `["-1-234+567-89", "-12+345+6-7-89", "123+45+6+78-9"]`를 반환해야 합니다.

```js
assert.deepEqual(sumTo100(243), [
  '-1-234+567-89',
  '-12+345+6-7-89',
  '123+45+6+78-9'
]);
```

`sumTo100(197)`는 `["1-2-3+45+67+89", "12+34-5+67+89", "123+4-5+6+78-9"]`를 반환해야 합니다.

```js
assert.deepEqual(sumTo100(197), [
  '1-2-3+45+67+89',
  '12+34-5+67+89',
  '123+4-5+6+78-9'
]);
```

# --seed--

## --seed-contents--

```js
function sumTo100(n) {

}
```

# --solutions--

```js
function sumTo100(n) {
  var permutationsWithRepetition = function(n, as) {
    return as.length > 0
      ? foldl1(curry(cartesianProduct)(as), replicate(n, as))
      : [];
  };

  var cartesianProduct = function(xs, ys) {
    return [].concat.apply(
      [],
      xs.map(function(x) {
        return [].concat.apply(
          [],
          ys.map(function(y) {
            return [[x].concat(y)];
          })
        );
      })
    );
  };

  var curry = function(f) {
    return function(a) {
      return function(b) {
        return f(a, b);
      };
    };
  };

  var flip = function(f) {
    return function(a, b) {
      return f.apply(null, [b, a]);
    };
  };

  var foldl1 = function(f, xs) {
    return xs.length > 0 ? xs.slice(1).reduce(f, xs[0]) : [];
  };

  var replicate = function(n, a) {
    var v = [a],
      o = [];
    if (n < 1) return o;
    while (n > 1) {
      if (n & 1) o = o.concat(v);
      n >>= 1;
      v = v.concat(v);
    }
    return o.concat(v);
  };

  var asSum = function(xs) {
    var dct = xs.reduceRight(
      function(a, sign, i) {
        var d = i + 1; //  zero-based index to [1-9] positions
        if (sign !== 0) {
          // Sum increased, digits cleared
          return {
            digits: [],
            n: a.n + sign * parseInt([d].concat(a.digits).join(''), 10)
          };
        } else
          return {
            // Digits extended, sum unchanged
            digits: [d].concat(a.digits),
            n: a.n
          };
      },
      {
        digits: [],
        n: 0
      }
    );
    return (
      dct.n + (dct.digits.length > 0 ? parseInt(dct.digits.join(''), 10) : 0)
    );
  };

  var asString = function(xs) {
    var ns = xs.reduce(function(a, sign, i) {
      var d = (i + 1).toString();
      return sign === 0 ? a + d : a + (sign > 0 ? '+' : '-') + d;
    }, '');

    return ns[0] === '+' ? tail(ns) : ns;
  };

  var universe = permutationsWithRepetition(9, [0, 1, -1])
    .filter(function(x) {
      return x[0] !== 1 && asSum(x) === n;
    })
    .map(asString);
  return universe.sort();
}
```
