---
id: 5a23c84252665b21eecc7e06
title: 직사각형 자르기
challengeType: 1
forumTopicId: 302242
dashedName: cut-a-rectangle
---

# --description--

주어진 직사각형은 *m* × *n* 개의 정사각형으로 이루어져 있습니다. 만약 *m*과 *n*이 모두 홀수가 아니라면, 정사각형의 모서리를 따라 경로를 잘라 직사각형을 두 개의 연결된 조각으로 나눌 수 있으며, 이 두 조각은 한 조각을 180° 회전시킨 후에 같은 모양이 됩니다. 2 × 2와 4 × 3 직사각형에 대한 모든 가능한 경로가 아래에 표시되어 있습니다.

<img src="https://cdn.freecodecamp.org/curriculum/coding-interview-prep/rosetta-cut-a-rectangle.png" alt="Diagram showing the possible paths for 2 by 2 and 4 by 3 rectangles" />

# --instructions--

*m* × *n* 직사각형을 자르는 서로 다른 방법의 수를 계산하는 함수를 작성하세요.

# --hints--

`cutRectangle`는 함수여야 합니다.

```js
assert(typeof cutRectangle == 'function');
```

`cutRectangle(2, 2)`은 숫자를 반환해야 합니다.

```js
assert(typeof cutRectangle(2, 2) == 'number');
```

`cutRectangle(2, 2)`는 `2`을 반환해야 합니다.

```js
assert.equal(cutRectangle(2, 2), 2);
```

`cutRectangle(4, 3)`는 `9`를 반환해야 합니다.

```js
assert.equal(cutRectangle(4, 3), 9);
```

`cutRectangle(4, 4)`는 `22`을 반환해야 합니다.

```js
assert.equal(cutRectangle(4, 4), 22);
```

`cutRectangle(8, 3)`는 `53`를 반환해야 합니다.

```js
assert.equal(cutRectangle(8, 3), 53);
```

`cutRectangle(7, 4)`는 `151`을 반환해야 합니다.

```js
assert.equal(cutRectangle(7, 4), 151);
```

# --seed--

## --seed-contents--

```js
function cutRectangle(w, h) {

}
```

# --solutions--

```js
function cutRectangle(w, h) {
  if (w % 2 == 1 && h % 2 == 1) return;

  var dirs = [[0, -1], [-1, 0], [0, 1], [1, 0]];

  var grid = new Array(h);
  for (var i = 0; i < grid.length; i++) grid[i] = new Array(w);
  var stack = [];

  var half = Math.floor((w * h) / 2);
  var bits = Math.pow(2, half) - 1;
  var result = 0;
  for (; bits > 0; bits -= 2) {
    for (var i = 0; i < half; i++) {
      var r = Math.floor(i / w);
      var c = i % w;
      grid[r][c] = (bits & (1 << i)) != 0 ? 1 : 0;
      grid[h - r - 1][w - c - 1] = 1 - grid[r][c];
    }

    stack.push(0);
    grid[0][0] = 2;
    var count = 1;
    while (stack.length != 0) {
      var pos = stack.pop();
      var r = Math.floor(pos / w);
      var c = pos % w;

      for (var dir of dirs) {
        var nextR = r + dir[0];
        var nextC = c + dir[1];

        if (nextR >= 0 && nextR < h && nextC >= 0 && nextC < w) {
          if (grid[nextR][nextC] == 1) {
            stack.push(nextR * w + nextC);
            grid[nextR][nextC] = 2;
            count++;
          }
        }
      }
    }

    if (count == half) {
      result++;
    }
  }

  return result;
}
```
