---
id: 5a23c84252665b21eecc8036
title: 문자열에서 제어 코드와 확장 문자를 제거하기
challengeType: 1
forumTopicId: 302327
dashedName: strip-control-codes-and-extended-characters-from-a-string
---

# --description--

문자열에서 제어 코드와 확장 문자를 제거하는 작업입니다. 다음 각 결과를 어떻게 얻는지 보여주는 해결책이어야 합니다: 제어 코드와 확장 문자가 제거된 문자열. ASCII에서 제어 코드는 0부터 31, 그리고 127까지의 십진수 코드입니다. ASCII 기반 시스템에서는 제어 코드를 제거하면 결과 문자열의 모든 문자가 ASCII 표에서 십진수 32부터 126 범위 내에 있게 됩니다. 비 ASCII 기반 시스템에서는 이 작업의 목적상 ASCII 표(십진수 32부터 126 범위) 내에 해당하는 글리프가 없는 문자를 확장 문자로 간주합니다.

# --hints--

`strip`는 함수여야 합니다.

```js
assert(typeof strip == 'function');
```

`strip("abc")`은 문자열을 반환해야 합니다.

```js
assert(typeof strip('abc') == 'string');
```

`strip("\ba\x00b\n\rc\fd\xc3")`는 `"abcd"`을 반환해야 합니다.

```js
assert.equal(strip('\ba\x00b\n\rc\fd\xc3'), 'abcd');
```

`strip("\u0000\n abc\u00E9def\u007F")`는 `" abcdef"`를 반환해야 합니다.

```js
assert.equal(strip('\u0000\n abc\u00E9def\u007F'), ' abcdef');
```

`strip("a\n\tb\u2102d\u2147f")`는 `"abdf"`을 반환해야 합니다.

```js
assert.equal(strip('a\n\tb\u2102d\u2147f'), 'abdf');
```

`strip("Français.")`는 `"Franais."`를 반환해야 합니다.

```js
assert.equal(strip('Français.'), 'Franais.');
```

`strip("123\tabc\u0007DEF\u007F+-*/€æŧðłþ")`은 `"123abcDEF+-*/"`을 반환해야 합니다.

```js
assert.equal(strip('123\tabc\u0007DEF\u007F+-*/€æŧðłþ'), '123abcDEF+-*/');
```

# --seed--

## --seed-contents--

```js
function strip(s) {

}
```

# --solutions--

```js
function strip(s) {
  return s
    .split('')
    .filter(function(x) {
      var n = x.charCodeAt(0);

      return 31 < n && 127 > n;
    })
    .join('');
}
```
