---
id: 5a23c84252665b21eecc7ec1
title: 반복 제곱수 계산
challengeType: 1
forumTopicId: 302291
dashedName: iterated-digits-squaring
---

# --description--

자연수(0보다 큰 정수)의 각 자리 숫자를 제곱한 값을 더하면 항상 1 또는 89로 끝납니다:

<pre>15 -> 26 -> 40 -> 16 -> 37 -> 58 -> 89
7 -> 49 -> 97 -> 130 -> 10 -> 1
</pre>

# --instructions--

숫자를 매개변수로 받아 위 과정을 수행한 후 1 또는 89를 반환하는 함수를 작성하세요.

# --hints--

`iteratedSquare` 는 함수여야 합니다.

```js
assert(typeof iteratedSquare == 'function');
```

`iteratedSquare(4)` 은 숫자를 반환해야 합니다.

```js
assert(typeof iteratedSquare(4) == 'number');
```

`iteratedSquare(4)` 는 `89` 을 반환해야 합니다.

```js
assert.equal(iteratedSquare(4), 89);
```

`iteratedSquare(7)` 는 `1` 을 반환해야 합니다.

```js
assert.equal(iteratedSquare(7), 1);
```

`iteratedSquare(15)` 는 `89` 을 반환해야 합니다.

```js
assert.equal(iteratedSquare(15), 89);
```

`iteratedSquare(20)` 는 `89` 을 반환해야 합니다.

```js
assert.equal(iteratedSquare(20), 89);
```

`iteratedSquare(70)` 은 `1` 을 반환해야 합니다.

```js
assert.equal(iteratedSquare(70), 1);
```

`iteratedSquare(100)` 는 `1` 을 반환해야 합니다.

```js
assert.equal(iteratedSquare(100), 1);
```

# --seed--

## --seed-contents--

```js
function iteratedSquare(n) {

}
```

# --solutions--

```js
function iteratedSquare(n) {
    var total;
    while (n != 89 && n != 1) {
        total = 0;
        while (n > 0) {
            total += Math.pow(n % 10, 2);
            n = Math.floor(n/10);
        }
        n = total;
    }
    return n;
}
```
