---
id: 5a23c84252665b21eecc7eca
title: 카프레카 수
challengeType: 1
forumTopicId: 302296
dashedName: kaprekar-numbers
---

# --description--

양의 정수는 다음 조건을 만족하면 카프레카 수입니다:

<ul>
  <li>1이거나,</li>
  <li>그 제곱의 십진수 표현을 한 번 나누어 두 부분으로 분리할 수 있으며, 이 두 부분이 양의 정수로 이루어져 원래 수의 합과 같아야 합니다.</li>
</ul>

0으로만 이루어진 부분이 생기는 분할은 유효하지 않은데, 0은 양수로 간주되지 않기 때문입니다.

카프레카 수 예시:

<ul>
  <li><code>2223</code>은 카프레카 수입니다. 왜냐하면 <code>2223 * 2223 = 4941729</code>이고, <code>4941729</code>를 <code>494</code>와 <code>1729</code>로 나눌 수 있으며, <code>494 + 1729 = 2223</code>이기 때문입니다.</li>
  <li>카프레카 수의 수열은 A006886으로 알려져 있으며, <code>1, 9, 45, 55, ...</code>로 시작합니다.</li>
</ul>

# --instructions--

숫자 $n$과 진법 $bs$를 입력받아, 주어진 진법에서 $n$이 카프레카 수인지 판단하는 함수를 작성하세요. 카프레카 수라면 true를 반환하고, 그렇지 않으면 false를 반환하세요.

# --hints--

`isKaprekar`는 함수여야 합니다.

```js
assert(typeof isKaprekar == 'function');
```

`isKaprekar(1, 10)`는 불리언 값을 반환해야 합니다.

```js
assert(typeof isKaprekar(1, 10) == 'boolean');
```

`isKaprekar(1, 10)`는 `true`를 반환해야 합니다.

```js
assert.equal(isKaprekar(1, 10), true);
```

`isKaprekar(9, 10)`는 `true`를 반환해야 합니다.

```js
assert.equal(isKaprekar(9, 10), true);
```

`isKaprekar(2223, 10)`는 `true`를 반환해야 합니다.

```js
assert.equal(isKaprekar(2223, 10), true);
```

`isKaprekar(22823, 10)`는 `false`를 반환해야 합니다.

```js
assert.equal(isKaprekar(22823, 10), false);
```

`isKaprekar(9, 17)`는 `false`를 반환해야 합니다.

```js
assert.equal(isKaprekar(9, 17), false);
```

`isKaprekar(225, 17)`는 `true`를 반환해야 합니다.

```js
assert.equal(isKaprekar(225, 17), true);
```

`isKaprekar(999, 17)`는 `false`를 반환해야 합니다.

```js
assert.equal(isKaprekar(999, 17), false);
```

# --seed--

## --seed-contents--

```js
function isKaprekar(n, bs) {

}
```

# --solutions--

```js
function isKaprekar(n, bs) {
  if (n < 1) return false;
  if (n == 1) return true;
  for (var a = n * n, b = 0, s = 1; a; s *= bs) {
    b += (a % bs) * s;
    a = Math.floor(a / bs);
    if (b && a + b == n) return true;
  }
  return false;
}
```
