---
id: 5a23c84252665b21eecc803f
title: 정수의 자리수 합 구하기
challengeType: 1
forumTopicId: 302331
dashedName: sum-digits-of-an-integer
---

# --description--

문자열을 매개변수로 받는 함수를 작성하세요. 이 문자열은 37 미만의 어떤 진법으로 된 숫자를 나타내며, 그 자리수들의 합을 반환해야 합니다.

<ul>
  <li><b>1</b><sub>10</sub>의 자리수 합은 <b>1</b>입니다</li>
  <li><b>1234</b><sub>10</sub>의 자리수 합은 <b>10</b>입니다</li>
  <li><b>fe</b><sub>16</sub>의 자리수 합은 <b>29</b>입니다</li>
  <li><b>f0e</b><sub>16</sub>의 자리수 합은 <b>29</b>입니다</li>
</ul>

# --hints--

`sumDigits`는 함수여야 합니다.

```js
assert(typeof sumDigits == 'function');
```

`sumDigits("1")`은 숫자를 반환해야 합니다.

```js
assert(typeof sumDigits('1') == 'number');
```

`sumDigits("1")`는 `1`을 반환해야 합니다.

```js
assert.equal(sumDigits('1'), 1);
```

`sumDigits("12345")`는 `15`를 반환해야 합니다.

```js
assert.equal(sumDigits('12345'), 15);
```

`sumDigits("254")`는 `11`을 반환해야 합니다.

```js
assert.equal(sumDigits('254'), 11);
```

`sumDigits("fe")`는 `29`를 반환해야 합니다.

```js
assert.equal(sumDigits('fe'), 29);
```

`sumDigits("f0e")`는 `29`을 반환해야 합니다.

```js
assert.equal(sumDigits('f0e'), 29);
```

`sumDigits("999ABCXYZ")`는 `162`을 반환해야 합니다.

```js
assert.equal(sumDigits('999ABCXYZ'), 162);
```

# --seed--

## --seed-contents--

```js
function sumDigits(n) {

}
```

# --solutions--

```js
function sumDigits(n) {
  n += '';
  for (var s = 0, i = 0, e = n.length; i < e; i += 1)
    s += parseInt(n.charAt(i), 36);
  return s;
}
```
