---
id: 5a23c84252665b21eecc7eaf
title: IBAN
challengeType: 1
forumTopicId: 302289
dashedName: iban
---

# --description--

국제 은행 계좌 번호(IBAN)는 국가 간 은행 계좌를 식별하는 국제적으로 합의된 방법으로, 전사 오류가 전파될 위험을 줄입니다.

<abbr title="International Bank Account Number">IBAN</abbr>은 최대 34자의 영숫자로 구성됩니다:

<ul>
  <li>먼저 두 글자의 <abbr title="International Organization for Standardization">ISO</abbr> 3166-1 alpha-2 국가 코드</li>
  <li>그다음 두 자리의 검증 숫자, 그리고</li>
  <li>마지막으로 국가별 기본 은행 계좌 번호(BBAN)</li>
</ul>

검증 숫자는 거래를 제출하기 전에 은행 계좌 번호의 무결성을 확인하는 간단한 검사를 가능하게 합니다.

# --instructions--

IBAN 문자열을 매개변수로 받는 함수를 작성하세요. 유효하면 true를 반환하고, 그렇지 않으면 false를 반환하세요.

# --hints--

`isValid`는 함수여야 합니다.

```js
assert(typeof isValid == 'function');
```

`isValid("GB82 WEST 1234 5698 7654 32")`은 불리언을 반환해야 합니다.

```js
assert(typeof isValid('GB82 WEST 1234 5698 7654 32') == 'boolean');
```

`isValid("GB82 WEST 1234 5698 7654 32")`는 `true`을 반환해야 합니다.

```js
assert.equal(isValid('GB82 WEST 1234 5698 7654 32'), true);
```

`isValid("GB82 WEST 1.34 5698 7654 32")`는 `false`를 반환해야 합니다.

```js
assert.equal(isValid('GB82 WEST 1.34 5698 7654 32'), false);
```

`isValid("GB82 WEST 1234 5698 7654 325")`는 `false`을 반환해야 합니다.

```js
assert.equal(isValid('GB82 WEST 1234 5698 7654 325'), false);
```

`isValid("GB82 TEST 1234 5698 7654 32")`는 `false`를 반환해야 합니다.

```js
assert.equal(isValid('GB82 TEST 1234 5698 7654 32'), false);
```

`isValid("SA03 8000 0000 6080 1016 7519")`는 `true`을 반환해야 합니다.

```js
assert.equal(isValid('SA03 8000 0000 6080 1016 7519'), true);
```

# --seed--

## --seed-contents--

```js
function isValid(iban) {

}
```

# --solutions--

```js
function isValid(iban) {
  var ibanLen = {
    NO:15, BE:16, DK:18, FI:18, FO:18, GL:18, NL:18, MK:19,
    SI:19, AT:20, BA:20, EE:20, KZ:20, LT:20, LU:20, CR:21,
    CH:21, HR:21, LI:21, LV:21, BG:22, BH:22, DE:22, GB:22,
    GE:22, IE:22, ME:22, RS:22, AE:23, GI:23, IL:23, AD:24,
    CZ:24, ES:24, MD:24, PK:24, RO:24, SA:24, SE:24, SK:24,
    VG:24, TN:24, PT:25, IS:26, TR:26, FR:27, GR:27, IT:27,
    MC:27, MR:27, SM:27, AL:28, AZ:28, CY:28, DO:28, GT:28,
    HU:28, LB:28, PL:28, BR:29, PS:29, KW:30, MU:30, MT:31
  }
    iban = __helpers.removeWhiteSpace(iban)
    if (!iban.match(/^[\dA-Z]+$/)) return false
    var len = iban.length
    if (len != ibanLen[iban.substring(0,2)]) return false
    iban = iban.substring(4) + iban.substring(0,4)
    for (var s='', i=0; i<len; i+=1) s+=parseInt(iban.charAt(i),36)
    for (var m=s.substring(0,15)%97, s=s.substring(15); s; s=s.substring(13)) m=(m+s.substring(0,13))%97
    return m == 1
}
```
