---
id: 5a23c84252665b21eecc7ec2
title: 자로 거리
challengeType: 1
forumTopicId: 302292
dashedName: jaro-distance
---

# --description--

자로 거리는 두 문자열 사이의 유사도를 측정하는 방법입니다. 두 문자열의 자로 거리가 높을수록 문자열이 더 비슷하다는 뜻입니다. 점수는 `0`이 유사성이 전혀 없음을, `1`이 완전 일치를 의미하도록 표준화되어 있습니다.

**정의**

두 문자열 \\(s_1\\)과 \\(s_2\\)의 자로 거리 \\( d_j \\)는 다음과 같습니다.

\\begin{align}d_j = \\begin{cases}0& & \\text{if }m=0 \\\\\\\\{\\frac {1}{3}}\\left({\\frac {m}{|s\_{1}|}}+{\\frac {m}{|s\_{2}|}}+{\\frac {m-t}{m}}\\right)& & \\text{otherwise}\\end{cases}\\end{align}

여기서:

<ul>
  <li>\(m\)은 <i>일치하는 문자</i>의 개수입니다;</li>
  <li>\(t\)는 <i>전위</i>의 절반입니다.</li>
</ul>

\\(s_1\\)과 \\(s_2\\)의 각각의 두 문자는 같고, \\(\\left\\lfloor\\frac{\\max(|s_1|,|s_2|)}{2}\\right\\rfloor-1\\)보다 멀리 떨어져 있지 않을 때만 *일치*한다고 간주합니다.

\\(s_1\\)의 각 문자는 \\(s_2\\)의 모든 일치하는 문자와 비교됩니다. 순서가 다르지만 일치하는 문자의 개수를 2로 나눈 값이 *전위*의 개수를 정의합니다.

**예시**

문자열 \\(s_1\\) *DWAYNE*와 \\(s_2\\) *DUANE*가 주어졌을 때:

<ul>
  <li>\(m = 4\)</li>
  <li>\(|s_1| = 6\)</li>
  <li>\(|s_2| = 5\)</li>
  <li>\(t = 0\)</li>
</ul>

자로 점수는 \\(d_j = \\frac{1}{3}\\left(\\frac{4}{6} + \\frac{4}{5} + \\frac{4-0}{4}\\right) = 0.822\\)입니다.

# --instructions--

두 문자열을 매개변수로 받아 해당 자로 거리를 반환하는 함수를 작성하세요.

# --hints--

`jaro`는 함수여야 합니다.

```js
assert(typeof jaro == 'function');
```

`jaro("MARTHA", "MARHTA")`는 숫자를 반환해야 합니다.

```js
assert(typeof jaro('MARTHA', 'MARHTA') == 'number');
```

`jaro("MARTHA", "MARHTA")`는 `0.9444444444444445`를 반환해야 합니다.

```js
assert.equal(jaro('MARTHA', 'MARHTA'), 0.9444444444444445);
```

`jaro("DIXON", "DICKSONX")`는 `0.7666666666666666`를 반환해야 합니다.

```js
assert.equal(jaro('DIXON', 'DICKSONX'), 0.7666666666666666);
```

`jaro("JELLYFISH", "SMELLYFISH")`는 `0.8962962962962964`를 반환해야 합니다.

```js
assert.equal(jaro('JELLYFISH', 'SMELLYFISH'), 0.8962962962962964);
```

`jaro("HELLOS", "CHELLO")`는 `0.888888888888889`를 반환해야 합니다.

```js
assert.equal(jaro('HELLOS', 'CHELLO'), 0.888888888888889);
```

`jaro("ABCD", "BCDA")`는 `0.8333333333333334`를 반환해야 합니다.

```js
assert.equal(jaro('ABCD', 'BCDA'), 0.8333333333333334);
```

# --seed--

## --seed-contents--

```js
function jaro(s, t) {

}
```

# --solutions--

```js
function jaro(s, t) {
  var s_len = s.length;
  var t_len = t.length;

  if (s_len == 0 && t_len == 0) return 1;

  var match_distance = Math.max(s_len, t_len) / 2 - 1;

  var s_matches = new Array(s_len);
  var t_matches = new Array(t_len);

  var matches = 0;
  var transpositions = 0;

  for (var i = 0; i < s_len; i++) {
    var start = Math.max(0, i - match_distance);
    var end = Math.min(i + match_distance + 1, t_len);

    for (var j = start; j < end; j++) {
      if (t_matches[j]) continue;
      if (s.charAt(i) != t.charAt(j)) continue;
      s_matches[i] = true;
      t_matches[j] = true;
      matches++;
      break;
    }
  }

  if (matches == 0) return 0;

  var k = 0;
  for (var i = 0; i < s_len; i++) {
    if (!s_matches[i]) continue;
    while (!t_matches[k]) k++;
    if (s.charAt(i) != t.charAt(k)) transpositions++;
    k++;
  }

  return ((matches / s_len) +
    (matches / t_len) +
    ((matches - transpositions / 2.0) / matches)) / 3.0;
}
```
