---
id: 5a23c84252665b21eecc7ee0
title: 왼쪽 팩토리얼
challengeType: 1
forumTopicId: 302302
dashedName: left-factorials
---

# --description--

**왼쪽 팩토리얼** $ !n $는 *서브팩토리얼* 또는 *팩토리얼 합*을 가리킬 수 있습니다. 같은 표기가 두 가지 다른 정의에 혼동스럽게 사용될 수 있습니다. 때때로, *서브팩토리얼* (또는 *배열*)은 다음 표기 중 어느 것이든 사용할 수 있습니다:

<ul>
  <li>$!n`$</li>
  <li>$!n$</li>
  <li>$n¡$</li>
</ul>

(눈에 잘 띄지 않을 수 있지만, 마지막 예시는 거꾸로 된 느낌표를 사용합니다.) 이 도전 과제에서는 다음 공식을 사용하여 **왼쪽 팩토리얼**을 정의합니다:

$ !n = \\sum\_{k=0}^{n-1} k! $

여기서 $!0 = 0$입니다.

# --instructions--

주어진 수의 왼쪽 팩토리얼을 계산하는 함수를 작성하세요.

# --hints--

`leftFactorial`는 함수여야 합니다.

```js
assert(typeof leftFactorial == 'function');
```

`leftFactorial(0)`는 숫자를 반환해야 합니다.

```js
assert(typeof leftFactorial(0) == 'number');
```

`leftFactorial(0)`는 `0`를 반환해야 합니다.

```js
assert.equal(leftFactorial(0), 0);
```

`leftFactorial(1)`는 `1`를 반환해야 합니다.

```js
assert.equal(leftFactorial(1), 1);
```

`leftFactorial(2)`는 `2`를 반환해야 합니다.

```js
assert.equal(leftFactorial(2), 2);
```

`leftFactorial(3)`는 `4`를 반환해야 합니다.

```js
assert.equal(leftFactorial(3), 4);
```

`leftFactorial(10)`는 `409114`를 반환해야 합니다.

```js
assert.equal(leftFactorial(10), 409114);
```

`leftFactorial(17)`는 `22324392524314`를 반환해야 합니다.

```js
assert.equal(leftFactorial(17), 22324392524314);
```

`leftFactorial(19)`는 `6780385526348314`를 반환해야 합니다.

```js
assert.equal(leftFactorial(19), 6780385526348314);
```

# --seed--

## --seed-contents--

```js
function leftFactorial(n) {

}
```

# --solutions--

```js
function leftFactorial(n) {
  if (n == 0) return 0;
  if (n == 1) return 1;

  // Note: for n>=20, the result may not be correct.
  // This is because JavaScript uses 53 bit integers and
  // for n>=20 result becomes too large.

  let res = 2,
    fact = 2;
  for (var i = 2; i < n; i++) {
    res += fact;
    fact *= i + 1;
  }

  return res;
}
```
