---
id: 5a23c84252665b21eecc7e05
title: CUSIP
challengeType: 1
forumTopicId: 302241
dashedName: cusip
---

# --description--

**CUSIP**는 북미 금융 증권을 식별하는 9자리 영숫자 코드로, 거래의 결제 및 청산을 원활하게 하기 위해 사용됩니다. CUSIP는 Accredited Standards X9.6에 따라 미국 국가 표준으로 채택되었습니다.

# --instructions--

문자열을 매개변수로 받아 해당 문자열이 유효한 CUSIP인지 확인하는 함수를 작성하세요.

# --hints--

`isCusip`는 함수여야 합니다.

```js
assert(typeof isCusip == 'function');
```

`isCusip("037833100")`는 불리언을 반환해야 합니다.

```js
assert(typeof isCusip('037833100') == 'boolean');
```

`isCusip("037833100")`는 `true`를 반환해야 합니다.

```js
assert.equal(isCusip('037833100'), true);
```

`isCusip("17275R102")`는 `true`를 반환해야 합니다.

```js
assert.equal(isCusip('17275R102'), true);
```

`isCusip("38259P50a")`는 `false`를 반환해야 합니다.

```js
assert.equal(isCusip('38259P50a'), false);
```

`isCusip("38259P508")`는 `true`를 반환해야 합니다.

```js
assert.equal(isCusip('38259P508'), true);
```

`isCusip("38259P50#")`는 `false`를 반환해야 합니다.

```js
assert.equal(isCusip('38259P50#'), false);
```

`isCusip("68389X105")`는 `true`를 반환해야 합니다.

```js
assert.equal(isCusip('68389X105'), true);
```

`isCusip("68389X106")`는 `false`를 반환해야 합니다.

```js
assert.equal(isCusip('68389X106'), false);
```

`isCusip("5949181")`는 `false`를 반환해야 합니다.

```js
assert.equal(isCusip('5949181'), false);
```

# --seed--

## --seed-contents--

```js
function isCusip(s) {

}
```

# --solutions--

```js
function isCusip(s) {
  if (s.length != 9) return false;
  var sum = 0;
  var ASCII = x => x.charCodeAt(0);
  for (var i = 0; i < 7; i++) {
    var c = s.charCodeAt(i);

    var v;
    if (c >= ASCII('0') && c <= ASCII('9')) {
      v = c - 48;
    } else if (c >= ASCII('A') && c <= ASCII('Z')) {
      v = c - 64; // lower case letters apparently invalid
    } else if (c == ASCII('*')) {
      v = 36;
    } else if (c == ASCII('@')) {
      v = 37;
    } else if (c == ASCII('#')) {
      v = 38;
    } else {
      return false;
    }
    if (i % 2 == 1) v *= 2; // check if odd as using 0-based indexing
    sum += Math.floor(v / 10) + (v % 10);
  }
  return s.charCodeAt(8) - 48 == (10 - (sum % 10)) % 10;
}
```
