---
id: 5a23c84252665b21eecc8016
title: 사용자 정의 비교 함수로 정렬하기
challengeType: 1
forumTopicId: 302309
dashedName: sort-using-a-custom-comparator
---

# --description--

문자열 배열(또는 리스트)을 길이가 긴 순서대로 내림차순 정렬하고, 길이가 같은 문자열은 사전식 오름차순으로 정렬하는 함수를 작성하세요.

# --hints--

`lengthSorter`는 함수여야 합니다.

```js
assert(typeof lengthSorter == 'function');
```

`lengthSorter(["Here", "are", "some", "sample", "strings", "to", "be", "sorted"])`는 배열을 반환해야 합니다.

```js
assert(
  Array.isArray(
    lengthSorter([
      'Here',
      'are',
      'some',
      'sample',
      'strings',
      'to',
      'be',
      'sorted'
    ])
  )
);
```

`lengthSorter(["Here", "are", "some", "sample", "strings", "to", "be", "sorted"])`는 `["strings", "sample", "sorted", "Here", "some", "are", "be", "to"]`를 반환해야 합니다.

```js
assert.deepEqual(
  lengthSorter([
    'Here',
    'are',
    'some',
    'sample',
    'strings',
    'to',
    'be',
    'sorted'
  ]),
  ['strings', 'sample', 'sorted', 'Here', 'some', 'are', 'be', 'to']
);
```

`lengthSorter(["I", "hope", "your", "day", "is", "going", "good", "?"])`는 `["going", "good", "hope", "your", "day", "is", "?","I"]`를 반환해야 합니다.

```js
assert.deepEqual(
  lengthSorter(['I', 'hope', 'your', 'day', 'is', 'going', 'good', '?']),
  ['going', 'good', 'hope', 'your', 'day', 'is', '?', 'I']
);
```

`lengthSorter(["Mine", "is", "going", "great"])`는 `["going", "great", "Mine", "is"]`를 반환해야 합니다.

```js
assert.deepEqual(lengthSorter(['Mine', 'is', 'going', 'great']), [
  'going',
  'great',
  'Mine',
  'is'
]);
```

`lengthSorter(["Have", "fun", "sorting", "!!"])`는 `["sorting", "Have", "fun", "!!"]`를 반환해야 합니다.

```js
assert.deepEqual(lengthSorter(['Have', 'fun', 'sorting', '!!']), [
  'sorting',
  'Have',
  'fun',
  '!!'
]);
```

`lengthSorter(["Everything", "is", "good", "!!"])`는 `["Everything", "good", "!!", "is"]`를 반환해야 합니다.

```js
assert.deepEqual(lengthSorter(['Everything', 'is', 'good', '!!']), [
  'Everything',
  'good',
  '!!',
  'is'
]);
```

# --seed--

## --seed-contents--

```js
function lengthSorter(arr) {

}
```

# --solutions--

```js
function lengthSorter(arr) {
  arr.sort(function(a, b) {
    var result = b.length - a.length;
    if (result == 0) result = a.localeCompare(b);
    return result;
  });
  return arr;
}
```
