---
id: 5a23c84252665b21eecc8001
title: 정렬 알고리즘/비드 정렬
challengeType: 1
forumTopicId: 302310
dashedName: sorting-algorithmsbead-sort
---

# --description--

*비드 정렬*은 입력 배열에서 가장 큰 요소의 값과 같은 길이의 0으로 이루어진 행렬을 만드는 것부터 시작합니다. 행렬은 현재 요소가 가리키는 인덱스와 0번째 인덱스 사이의 모든 요소에 1을 더하는 방식으로 변환됩니다. 이 과정을 반복하여 행렬을 채웁니다.

행렬을 반복하면서 0보다 큰 요소의 개수를 합산하고, 각 요소의 값을 1씩 감소시키면 정렬된 배열을 얻습니다.

**참고:** 입력 배열의 각 요소는 고유합니다.

비드 정렬 알고리즘을 사용하여 양의 정수 배열을 정렬하세요.

# --hints--

`beadSort`는 함수여야 합니다.

```js
assert(typeof beadSort == 'function');
```

`beadSort([25, 32, 12, 7, 20])`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(beadSort([25, 32, 12, 7, 20])));
```

`beadSort([25, 32, 12, 7, 20])`는 `[7, 12, 20, 25, 32]`를 반환해야 합니다.

```js
assert.deepEqual(beadSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`beadSort([38, 45, 35, 8, 13])`는 `[8, 13, 35, 38, 45]`를 반환해야 합니다.

```js
assert.deepEqual(beadSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`beadSort([43, 36, 20, 34, 24])`는 `[20, 24, 34, 36, 43]`를 반환해야 합니다.

```js
assert.deepEqual(beadSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`beadSort([12, 33, 26, 18, 1, 16, 38])`는 `[1, 12, 16, 18, 26, 33, 38]`를 반환해야 합니다.

```js
assert.deepEqual(beadSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`beadSort([3, 39, 48, 16, 1, 4, 29])`는 `[1, 3, 4, 16, 29, 39, 48]`를 반환해야 합니다.

```js
assert.deepEqual(beadSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function beadSort(arr) {

}
```

# --solutions--

```js
function beadSort(arr) {
  let max = 0;
  for (let i = 0; i < arr.length; i++) if (arr[i] > max) max = arr[i];
  const grid = new Array(arr.length);
  for (let i = 0; i < grid.length; i++) {
    grid[i] = new Array(max);
  }
  const levelcount = new Array(max);
  levelcount.fill(0);
  for (let i = 0; i < max; i++) {
    levelcount[i] = 0;
    for (let j = 0; j < arr.length; j++) grid[j][i] = '_';
  }
  for (let i = 0; i < arr.length; i++) {
    let num = arr[i];
    for (let j = 0; num > 0; j++) {
      grid[levelcount[j]++][j] = '*';
      num--;
    }
  }
  const sorted = new Array(arr.length);
  sorted.fill(0);
  for (let i = 0; i < arr.length; i++) {
    let putt = 0;
    for (
      let j = 0;
      j < max &&
      (function(c) {
        return c.charCodeAt == null ? c : c.charCodeAt(0);
      })(grid[arr.length - 1 - i][j]) == '*'.charCodeAt(0);
      j++
    )
      putt++;
    sorted[i] = putt;
  }
  return sorted;
}
```
