---
id: 5a23c84252665b21eecc801c
title: 나선형 배열
challengeType: 1
forumTopicId: 302321
dashedName: spiral-matrix
---

# --description--

나선형 배열을 만드세요. *나선형 배열*은 처음 N<sup>2</sup> 자연수를 정사각형으로 배열한 것으로, 배열의 가장자리를 따라 나선형으로 안쪽으로 들어가면서 숫자가 순차적으로 증가합니다. 예를 들어, **5**가 주어지면 다음 배열을 만드세요:

<pre>
0  1  2  3  4
15 16 17 18 5
14 23 24 19 6
13 22 21 20 7
12 11 10  9 8
</pre>

# --hints--

`spiralArray`는 함수여야 합니다.

```js
assert(typeof spiralArray == 'function');
```

`spiralArray(3)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(spiralArray(3)));
```

`spiralArray(3)`는 `[[0, 1, 2],[7, 8, 3],[6, 5, 4]]`를 반환해야 합니다.

```js
assert.deepEqual(spiralArray(3), [
  [0, 1, 2],
  [7, 8, 3],
  [6, 5, 4]
]);
```

`spiralArray(4)`는 `[[0, 1, 2, 3],[11, 12, 13, 4],[10, 15, 14, 5],[9, 8, 7, 6]]`를 반환해야 합니다.

```js
assert.deepEqual(spiralArray(4), [
  [0, 1, 2, 3],
  [11, 12, 13, 4],
  [10, 15, 14, 5],
  [9, 8, 7, 6]
]);
```

`spiralArray(5)`는 `[[0, 1, 2, 3, 4],[15, 16, 17, 18, 5],[14, 23, 24, 19, 6],[13, 22, 21, 20, 7],[12, 11, 10, 9, 8]]`를 반환해야 합니다.

```js
assert.deepEqual(spiralArray(5), [
  [0, 1, 2, 3, 4],
  [15, 16, 17, 18, 5],
  [14, 23, 24, 19, 6],
  [13, 22, 21, 20, 7],
  [12, 11, 10, 9, 8]
]);
```

# --seed--

## --seed-contents--

```js
function spiralArray(n) {

}
```

# --solutions--

```js
function spiralArray(n) {
  var arr = Array(n),
    x = 0,
    y = n,
    total = n * n--,
    dx = 1,
    dy = 0,
    i = 0,
    j = 0;
  while (y) arr[--y] = [];
  while (i < total) {
    arr[y][x] = i++;
    x += dx;
    y += dy;
    if (++j == n) {
      if (dy < 0) {
        x++;
        y++;
        n -= 2;
      }
      j = dx;
      dx = -dy;
      dy = j;
      j = 0;
    }
  }
  return arr;
}
```
