---
id: 5a23c84252665b21eecc8005
title: 정렬 알고리즘/콤브 정렬
challengeType: 1
forumTopicId: 302313
dashedName: sorting-algorithmscomb-sort
---

# --description--

*콤브 정렬*을 구현하세요.

**콤브 정렬**은 버블 정렬의 변형입니다.

쉘 정렬처럼 콤브 정렬은 비교와 교환에 사용하는 간격을 증가시킵니다.

간격을 $(1-e^{-\\varphi})^{-1} \\approx 1.247330950103979$로 나누는 것이 가장 효과적이지만, 1.3이 더 실용적일 수 있습니다.

일부 구현은 간격이 일정 값보다 작아지면 삽입 정렬을 사용합니다.

변형:

<ul>
  <li>Combsort11은 간격이 (11, 8, 6, 4, 3, 2, 1)로 끝나도록 하여 다른 두 가지 가능한 끝보다 훨씬 빠릅니다.</li>
  <li>다른 끝을 가진 Combsort는 데이터가 거의 정렬된 경우(간격이 작을 때) 더 효율적인 정렬로 전환합니다. 간격이 낮은 콤브 정렬은 버블 정렬보다 크게 나은 점이 없습니다.</li>
</ul>

의사 코드:

<pre><b>function</b> combsort(<b>array</b> input)
  gap := input<b>.size</b> <i>//간격 크기 초기화</i>
  <b>loop until</b> gap = 1 <b>and</b> swaps = 0
    <i>//다음 콤브를 위한 간격 값 업데이트. 아래는 예시입니다</i>
    gap := int(gap / 1.25)
    <b>if</b> gap &#x3C; 1 
      <i>//최소 간격은 1</i>
      gap := 1
    <b>end if</b>
    i := 0
    swaps := 0 <i>//설명은 <a href='https://rosettacode.org/wiki/Sorting_algorithms/Bubble_sort' target='_blank'>버블 정렬</a> 참고</i>
    <i>//입력 리스트에 대해 한 번의 "콤브" 수행</i>
    <b>loop until</b> i + gap >= input<b>.size</b> <i>//비슷한 아이디어는 <a href='https://rosettacode.org/wiki/Sorting_algorithms/Shell_sort' target='_blank'>쉘 정렬</a> 참고</i>
      <b>if</b> input[i] > input[i+gap]
        <b>swap</b>(input[i], input[i+gap])
        swaps := 1 <i>//교환이 발생했음을 표시, 리스트가 정렬되었다고 보장할 수 없음</i>
      <b>end if</b>
      i := i + 1
    <b>end loop</b>
  <b>end loop</b>
<b>end function</b>
</pre>

# --instructions--

콤브 정렬을 사용해 주어진 배열을 정렬하는 함수를 작성하세요.

# --hints--

`combSort`는 함수여야 합니다.

```js
assert(typeof combSort == 'function');
```

`combSort([25, 32, 12, 7, 20])`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(combSort([25, 32, 12, 7, 20])));
```

`combSort([25, 32, 12, 7, 20])`는 `[7, 12, 20, 25, 32]`를 반환해야 합니다.

```js
assert.deepEqual(combSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`combSort([38, 45, 35, 8, 13])`는 `[8, 13, 35, 38, 45]`를 반환해야 합니다.

```js
assert.deepEqual(combSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`combSort([43, 36, 20, 34, 24])`는 `[20, 24, 34, 36, 43]`를 반환해야 합니다.

```js
assert.deepEqual(combSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`combSort([12, 33, 26, 18, 1, 16, 38])`는 `[1, 12, 16, 18, 26, 33, 38]`를 반환해야 합니다.

```js
assert.deepEqual(combSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`combSort([3, 39, 48, 16, 1, 4, 29])`는 `[1, 3, 4, 16, 29, 39, 48]`를 반환해야 합니다.

```js
assert.deepEqual(combSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function combSort(arr) {

}
```

# --solutions--

```js
function combSort(arr) {
  function is_array_sorted(arr) {
    var sorted = true;
    for (var i = 0; i < arr.length - 1; i++) {
      if (arr[i] > arr[i + 1]) {
        sorted = false;
        break;
      }
    }
    return sorted;
  }
  var iteration_count = 0;
  var gap = arr.length - 2;
  var decrease_factor = 1.25;

  // Until array is not sorted, repeat iterations
  while (!is_array_sorted(arr)) {
    // If not first gap
    if (iteration_count > 0)
      // Calculate gap
      gap = gap == 1 ? gap : Math.floor(gap / decrease_factor);

    // Set front and back elements and increment to a gap
    var front = 0;
    var back = gap;
    while (back <= arr.length - 1) {
      // If elements are not ordered swap them
      if (arr[front] > arr[back]) {
        var temp = arr[front];
        arr[front] = arr[back];
        arr[back] = temp;
      }

      // Increment and re-run swapping
      front += 1;
      back += 1;
    }
    iteration_count += 1;
  }

  return arr;
}
```
