---
id: 5a23c84252665b21eecc8013
title: 정렬 알고리즘/스트랜드 정렬
challengeType: 1
forumTopicId: 302319
dashedName: sorting-algorithmsstrand-sort
---

# --description--

**스트랜드 정렬**은 정렬된 부분집합을 만들어 최종 결과를 병합합니다.

`unsortedArray = [3, 1, 4, 2]`를 생각해 보세요. 첫 번째 항목 `3`을 선택해 별도의 배열에 복사합니다. 이 항목 다음에 더 큰 항목이 있는지 찾으세요. 더 큰 항목을 찾으면, 이 경우 `4`, 별도의 배열 `[3, 4]`에 복사하고 다음 항목들을 이 새로운 값 `4`와 비교합니다.

배열 끝에 도달하면 복사한 항목들 `[3, 4]`을 제거하고 `unsortedArray`에 남아 있는 첫 번째 항목, 이 경우 `1`부터 다시 시작하세요.

이 과정을 따르면 두 개의 정렬된 배열 `[3, 4]`과 `[1, 2]`이 만들어집니다. 이 두 배열을 병합해 `strandSortedArray`을 만드세요.

```js
const unsortedArray = [3, 1, 4, 2];
const strandsortedArray = [1, 2, 3, 4];
```

**스트랜드 정렬**을 사용해 배열을 정렬하는 함수를 작성하세요. 함수는 정렬된 배열을 반환해야 합니다.


# --hints--

`strandSort`은 함수여야 합니다.

```js
assert(typeof strandSort == 'function');
```

`strandSort([25, 32, 12, 7, 20])`은 배열을 반환해야 합니다.

```js
assert(Array.isArray(strandSort([25, 32, 12, 7, 20])));
```

`strandSort([25, 32, 12, 7, 20])`은 `[7, 12, 20, 25, 32]`를 반환해야 합니다.

```js
assert.deepEqual(strandSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`strandSort([38, 45, 35, 8, 13])`은 `[8, 13, 35, 38, 45]`를 반환해야 합니다.

```js
assert.deepEqual(strandSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`strandSort([43, 36, 20, 34, 24])`은 `[20, 24, 34, 36, 43]`를 반환해야 합니다.

```js
assert.deepEqual(strandSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`strandSort([12, 33, 26, 18, 1, 16, 38])`은 `[1, 12, 16, 18, 26, 33, 38]`를 반환해야 합니다.

```js
assert.deepEqual(strandSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`strandSort([3, 39, 48, 16, 1, 4, 29])`은 `[1, 3, 4, 16, 29, 39, 48]`를 반환해야 합니다.

```js
assert.deepEqual(strandSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function strandSort(list) {

}
```

# --solutions--

```js
function strandSort(list) {
  function merge(left, right) {
    var result = [];
    while (left.length != 0 && right.length != 0) {
      if (left[0] <= right[0]) result.push(left.shift());
      else result.push(right.shift());
    }
    result.push.apply(result, left);
    result.push.apply(result, right);
    return result;
  }

  if (list.length <= 1) return list;
  var result = [];
  while (list.length > 0) {
    var sorted = [];
    sorted.push(list.shift());
    var len = list.length;
    for (var i = 1; i < len; i++) {
      var elem = list[i];
      if (sorted[i - 1] <= elem) {
        sorted.push(elem);
        sorted.splice(i, 1);
      }
    }

    result = merge(result, sorted);
  }
  return result;
}
```
