---
id: 5a7dad05be01840e1778a0d1
title: 프랙트란
challengeType: 1
forumTopicId: 302270
dashedName: fractran
---

# --description--

FRACTRAN은 수학자 John Horton Conway가 발명한 튜링 완전 난해 프로그래밍 언어입니다.

FRACTRAN 프로그램은 양의 분수들로 이루어진 순서 있는 목록 $P = (f_1, f_2, \\ldots, f_m)$과 초기 양의 정수 입력 $n$으로 구성됩니다.

프로그램은 다음과 같이 정수 $n$을 업데이트하며 실행합니다:

<ul>
  <li>목록에서 $nf_i$가 정수가 되는 첫 번째 분수 $f_i$를 찾아 $n$을 $nf_i$로 바꿉니다;</li>
  <li>목록의 어떤 분수도 $n$과 곱했을 때 정수가 나오지 않을 때까지 이 규칙을 반복한 후 종료합니다.</li>
</ul>

Conway는 FRACTRAN으로 소수를 구하는 프로그램을 제시했습니다:

$\\dfrac{17}{91}$, $\\dfrac{78}{85}$, $\\dfrac{19}{51}$, $\\dfrac{23}{38}$, $\\dfrac{29}{33}$, $\\dfrac{77}{29}$, $\\dfrac{95}{23}$, $\\dfrac{77}{19}$, $\\dfrac{1}{17}$, $\\dfrac{11}{13}$, $\\dfrac{13}{11}$, $\\dfrac{15}{14}$, $\\dfrac{15}{2}$, $\\dfrac{55}{1}$

$n=2$로 시작하면 이 FRACTRAN 프로그램은 $n$을 $15=2\\times (\\frac{15}{2})$로, 이어서 $825=15\\times (\\frac{55}{1})$로 바꾸며 다음과 같은 정수 수열을 생성합니다:

$2$, $15$, $825$, $725$, $1925$, $2275$, $425$, $390$, $330$, $290$, $770$, $\\ldots$

2 이후 이 수열에는 다음과 같은 2의 거듭제곱들이 포함됩니다:

$2^2=4$, $2^3=8$, $2^5=32$, $2^7=128$, $2^{11}=2048$, $2^{13}=8192$, $2^{17}=131072$, $2^{19}=524288$, $\\ldots$

이는 2의 소수 거듭제곱들입니다.

# --instructions--

프랙트란 프로그램을 문자열 매개변수로 받아 프로그램의 처음 10개 숫자를 배열로 반환하는 함수를 작성하세요. 결과가 10개 미만이면 있는 숫자들을 그대로 반환하세요.

# --hints--

`fractran`는 함수여야 합니다.

```js
assert(typeof fractran == 'function');
```

`fractran("3/2, 1/3")`은 배열을 반환해야 합니다.

```js
assert(Array.isArray(fractran('3/2, 1/3')));
```

`fractran("3/2, 1/3")`는 `[ 2, 3, 1 ]`을 반환해야 합니다.

```js
assert.deepEqual(fractran('3/2, 1/3'), [2, 3, 1]);
```

`fractran("3/2, 5/3, 1/5")`는 `[ 2, 3, 5, 1 ]`를 반환해야 합니다.

```js
assert.deepEqual(fractran('3/2, 5/3, 1/5'), [2, 3, 5, 1]);
```

`fractran("3/2, 6/3")`는 `[ 2, 3, 6, 9, 18, 27, 54, 81, 162, 243 ]`을 반환해야 합니다.

```js
assert.deepEqual(fractran('3/2, 6/3'), [2, 3, 6, 9, 18, 27, 54, 81, 162, 243]);
```

`fractran("2/7, 7/2")`는 `[ 2, 7, 2, 7, 2, 7, 2, 7, 2, 7 ]`를 반환해야 합니다.

```js
assert.deepEqual(fractran('2/7, 7/2'), [2, 7, 2, 7, 2, 7, 2, 7, 2, 7]);
```

`fractran("17/91, 78/85, 19/51, 23/38, 29/33, 77/29, 95/23, 77/19, 1/17, 11/13, 13/11, 15/14, 15/2, 55/1")`는 `[ 2, 15, 825, 725, 1925, 2275, 425, 390, 330, 290 ]`을 반환해야 합니다.

```js
assert.deepEqual(
  fractran(
    '17/91, 78/85, 19/51, 23/38, 29/33, 77/29, 95/23, 77/19, 1/17, 11/13, 13/11, 15/14, 15/2, 55/1'
  ),
  [2, 15, 825, 725, 1925, 2275, 425, 390, 330, 290]
);
```

# --seed--

## --seed-contents--

```js
function fractran(progStr) {

}
```

# --solutions--

```js
function fractran(progStr){
  var num = new Array();
  var den = new Array();
  var val ;
  var out="";
  function compile(prog){
    var regex = /\s*(\d*)\s*\/\s*(\d*)\s*(.*)/m;
    while(regex.test(prog)){
      num.push(regex.exec(prog)[1]);
      den.push(regex.exec(prog)[2]);
      prog = regex.exec(prog)[3];
    }
  }

  function step(val){
    var i=0;
    while(i<den.length && val%den[i] != 0) i++;
    return num[i]*val/den[i];
  }

  var seq=[]

  function exec(val){
    var i = 0;
    while(val && i<limit){
      seq.push(val)
      val = step(val);
      i ++;
    }
  }

  // Main
  compile(progStr);
  var limit = 10;
  exec(2);
  return seq;
}
```
