---
id: 5e4ce2eaac708cc68c1df260
title: 레벤슈타인 거리
challengeType: 1
forumTopicId: 385264
dashedName: levenshtein-distance
---

# --description--

정보 이론과 컴퓨터 과학에서, **레벤슈타인 거리**는 두 수열 간 차이를 측정하는 척도입니다(즉, 편집 거리). 두 문자열 사이의 레벤슈타인 거리는 한 문자열을 다른 문자열로 변환하는 데 필요한 편집의 최솟값으로 정의되며, 허용되는 편집 작업은 단일 문자 삽입, 삭제 또는 대체입니다.

예시:

"**kitten**"과 "**sitting**" 사이의 레벤슈타인 거리는 3입니다. 다음 세 가지 편집으로 한 문자열을 다른 문자열로 바꿀 수 있으며, 세 번 미만의 편집으로는 불가능하기 때문입니다:

<ul>
  <li><strong>k</strong>itten   <strong>s</strong>itten    ('k'를 's'로 대체)</li>
  <li>sitt<strong>e</strong>n   sitt<strong>i</strong>n    ('e'를 'i'로 대체)</li>
  <li>sittin   sittin<strong>g</strong>    (끝에 'g' 삽입).</li>
</ul>

*레벤슈타인 거리는 "**rosettacode**"와 "**raisethysword**" 사이가 **8**입니다.*

*두 문자열의 거리는 두 문자열을 뒤집었을 때도 같습니다.*

# --instructions--

두 문자열을 매개변수로 받아 레벤슈타인 거리를 반환하는 함수를 작성하세요.

# --hints--

`levenshtein`는 함수여야 합니다.

```js
assert(typeof levenshtein == 'function');
```

`levenshtein("mist", "dist")`은 숫자를 반환해야 합니다.

```js
assert(typeof levenshtein('mist', 'dist') == 'number');
```

`levenshtein("mist", "dist")`는 `1`을 반환해야 합니다.

```js
assert.equal(levenshtein('mist', 'dist'), 1);
```

`levenshtein("tier", "tor")`는 `2`를 반환해야 합니다.

```js
assert.equal(levenshtein('tier', 'tor'), 2);
```

`levenshtein("kitten", "sitting")`는 `3`을 반환해야 합니다.

```js
assert.equal(levenshtein('kitten', 'sitting'), 3);
```

`levenshtein("stop", "tops")`는 `2`를 반환해야 합니다.

```js
assert.equal(levenshtein('stop', 'tops'), 2);
```

`levenshtein("rosettacode", "raisethysword")`는 `8`을 반환해야 합니다.

```js
assert.equal(levenshtein('rosettacode', 'raisethysword'), 8);
```

`levenshtein("mississippi", "swiss miss")`는 `8`을 반환해야 합니다.

```js
assert.equal(levenshtein('mississippi', 'swiss miss'), 8);
```

# --seed--

## --seed-contents--

```js
function levenshtein(a, b) {

}
```

# --solutions--

```js
function levenshtein(a, b) {
  var t = [], u, i, j, m = a.length, n = b.length;
  if (!m) { return n; }
  if (!n) { return m; }
  for (j = 0; j <= n; j++) { t[j] = j; }
  for (i = 1; i <= m; i++) {
    for (u = [i], j = 1; j <= n; j++) {
      u[j] = a[i - 1] === b[j - 1] ? t[j - 1] : Math.min(t[j - 1], t[j], u[j - 1]) + 1;
    } t = u;
  } return u[n];
}
```
