---
id: 5e4ce2f5ac708cc68c1df261
title: 선형 합동 생성기
challengeType: 1
forumTopicId: 385266
dashedName: linear-congruential-generator
---

# --description--

선형 합동 생성기(LCG)는 불연속 조각별 선형 방정식으로 계산된 의사 난수 수열을 생성하는 <em>알고리즘</em>입니다. 모든 선형 합동 생성기는 다음 공식을 사용합니다:

$$r_{n + 1} = (a \times r_n + c) \bmod m$$

여기서:

<ul>
<li>$ r_0 $는 시드입니다.</li>
<li>$r_1$, $r_2$, $r_3$, ..., 는 난수입니다.</li>
<li>$a$, $c$, $m$은 상수입니다.</li>
</ul>

$a$, $c$, $m$ 값을 신중하게 선택하면 생성기는 $0$부터 $m - 1$까지 균등 분포된 정수를 생성합니다.

<abbr title="linear congruential generator">LCG</abbr> 숫자는 품질이 좋지 않습니다. $r_n$과 $r\_{n + 1}$은 진짜 난수처럼 독립적이지 않습니다. $r_n$을 아는 사람은 $r\_{n + 1}$을 예측할 수 있으므로 <abbr title="linear congruential generator">LCG</abbr>는 암호학적으로 안전하지 않습니다. 그럼에도 <abbr title="linear congruential generator">LCG</abbr>는 Miller-Rabin 소수성 검사나 FreeCell 카드 배치 같은 간단한 작업에는 충분히 좋습니다. <abbr title="linear congruential generator">LCG</abbr>의 장점 중 하나는 같은 $r_0$에서 수열을 쉽게 재현할 수 있다는 점입니다. 공식이 매우 간단해서 다른 프로그래밍 언어로도 같은 수열을 재현할 수 있습니다.

# --instructions--

$r_0,a,c,m,n$를 매개변수로 받아 $r_n$을 반환하는 함수를 작성하세요.

# --hints--

`linearCongGenerator`는 함수여야 합니다.

```js
assert(typeof linearCongGenerator == 'function');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)`은 숫자를 반환해야 합니다.

```js
assert(typeof linearCongGenerator(324, 1145, 177, 2148, 3) == 'number');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)`는 `855`을 반환해야 합니다.

```js
assert.equal(linearCongGenerator(324, 1145, 177, 2148, 3), 855);
```

`linearCongGenerator(234, 11245, 145, 83648, 4)`는 `1110`를 반환해야 합니다.

```js
assert.equal(linearCongGenerator(234, 11245, 145, 83648, 4), 1110);
```

`linearCongGenerator(85, 11, 1234, 214748, 5)`는 `62217`을 반환해야 합니다.

```js
assert.equal(linearCongGenerator(85, 11, 1234, 214748, 5), 62217);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 1)`는 `12345`를 반환해야 합니다.

```js
assert.equal(linearCongGenerator(0, 1103515245, 12345, 2147483648, 1), 12345);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 2)`는 `1406932606`을 반환해야 합니다.

```js
assert.equal(
  linearCongGenerator(0, 1103515245, 12345, 2147483648, 2),
  1406932606
);
```

# --seed--

## --seed-contents--

```js
function linearCongGenerator(r0, a, c, m, n) {

}
```

# --solutions--

```js
function linearCongGenerator(r0, a, c, m, n) {
    for (let i = 0; i < n; i++) {
        r0 = (a * r0 + c) % m;
    }
    return r0;
}
```
