---
id: 5e4ce2a1ac708cc68c1df25d
title: 긴 곱셈
challengeType: 1
forumTopicId: 385269
dashedName: long-multiplication
---

# --description--

긴 곱셈을 명시적으로 구현하세요.

이것은 임의 정밀도 정수 대수에 대한 한 가지 가능한 접근법입니다.

# --instructions--

두 개의 큰 숫자 문자열을 매개변수로 받는 함수를 작성하세요. 함수는 이 두 큰 숫자의 곱을 문자열로 반환해야 합니다.

**참고:** JavaScript에서는 큰 숫자에 대한 산술 연산이 부정확하므로, 직접 정밀한 곱셈을 구현해야 합니다.

# --hints--

`mult`는 함수여야 합니다.

```js
assert(typeof mult == 'function');
```

`mult("18446744073709551616", "18446744073709551616")`는 문자열을 반환해야 합니다.

```js
assert(typeof mult('18446744073709551616', '18446744073709551616') == 'string');
```

`mult("18446744073709551616", "18446744073709551616")`는 `"340282366920938463463374607431768211456"`를 반환해야 합니다.

```js
assert.equal(
  mult('18446744073709551616', '18446744073709551616'),
  '340282366920938463463374607431768211456'
);
```

`mult("31844674073709551616", "1844674407309551616")`는 `"58743055272886011737990786529368211456"`를 반환해야 합니다.

```js
assert.equal(
  mult('31844674073709551616', '1844674407309551616'),
  '58743055272886011737990786529368211456'
);
```

`mult("1846744073709551616", "44844644073709551616")`는 `"82816580680737279241781007431768211456"`를 반환해야 합니다.

```js
assert.equal(
  mult('1846744073709551616', '44844644073709551616'),
  '82816580680737279241781007431768211456'
);
```

`mult("1844674407370951616", "1844674407709551616")`는 `"3402823669833978308014392742590611456"`를 반환해야 합니다.

```js
assert.equal(
  mult('1844674407370951616', '1844674407709551616'),
  '3402823669833978308014392742590611456'
);
```

`mult("2844674407370951616", "1844674407370955616")`는 `"5247498076580334548376218009219475456"`를 반환해야 합니다.

```js
assert.equal(
  mult('2844674407370951616', '1844674407370955616'),
  '5247498076580334548376218009219475456'
);
```

# --seed--

## --seed-contents--

```js
function mult(strNum1, strNum2) {

}
```

# --solutions--

```js
function mult(strNum1, strNum2) {
    var a1 = strNum1.split("").reverse();
    var a2 = strNum2.toString().split("").reverse();
    var aResult = new Array;

    for ( var iterNum1 = 0; iterNum1 < a1.length; iterNum1++ ) {
        for ( var iterNum2 = 0; iterNum2 < a2.length; iterNum2++ ) {
            var idxIter = iterNum1 + iterNum2;    // Get the current array position.
            aResult[idxIter] = a1[iterNum1] * a2[iterNum2] + ( idxIter >= aResult.length ? 0 : aResult[idxIter] );

            if ( aResult[idxIter] > 9 ) {    // Carrying
                aResult[idxIter + 1] = Math.floor( aResult[idxIter] / 10 ) + ( idxIter + 1 >= aResult.length ? 0 : aResult[idxIter + 1] );
                aResult[idxIter] %= 10;
            }
        }
    }
    return aResult.reverse().join("");
}
```
