---
id: 5e6dd14797f5ce267c2f19d0
title: 룩 앤 세이 수열
challengeType: 1
forumTopicId: 385277
dashedName: look-and-say-sequence
---

# --description--

룩 앤 세이 수열은 재귀적으로 정의된 수열입니다.

수열 정의

<ul><li>십진수를 하나 선택하세요</li>
<li><span>숫자를 보고</span>, 같은 숫자가 연속으로 나오는 부분을 시각적으로 그룹화하세요.</li>
<li><span>숫자를 말하세요</span>, 왼쪽에서 오른쪽으로 그룹별로; 해당 숫자가 몇 개 있는지 말한 다음, 그 숫자를 말하세요.</li></ul><span> 이것이 다음 수열의 숫자가 됩니다.</span>

예시:

<ul><li>숫자 1에서 시작하면, <span>하나의</span> 1이 있어 11이 됩니다</li>
<li>11에서 시작하면, <span>두 개의</span> 1이 있습니다. 즉: 21</li>
<li>21에서 시작하면, <span>하나의</span> 2, 그리고 <span>하나의</span> 1이 있습니다. 즉: (12)(11)로 1211이 됩니다</li>
<li>1211에서 시작하면, <span>하나의</span> 1, <span>하나의</span> 2, 그리고 <span>두 개의</span> 1이 있습니다. 즉: (11)(12)(21)로 111221가 됩니다</li></ul>

# --instructions--

문자열을 매개변수로 받아 처리하고, 결과 문자열을 반환하는 함수를 작성하세요.

# --hints--

`lookAndSay`는 함수여야 합니다.

```js
assert(typeof lookAndSay == 'function');
```

`lookAndSay("1")`는 문자열을 반환해야 합니다.

```js
assert(typeof lookAndSay('1') == 'string');
```

`lookAndSay("1")`는 `"11"`를 반환해야 합니다.

```js
assert.equal(lookAndSay('1'), '11');
```

`lookAndSay("11")`는 `"21"`를 반환해야 합니다.

```js
assert.equal(lookAndSay('11'), '21');
```

`lookAndSay("21")`는 `"1211"`를 반환해야 합니다.

```js
assert.equal(lookAndSay('21'), '1211');
```

`lookAndSay("1211")`는 `"111221"`를 반환해야 합니다.

```js
assert.equal(lookAndSay('1211'), '111221');
```

`lookAndSay("3542")`는 `"13151412"`를 반환해야 합니다.

```js
assert.equal(lookAndSay('3542'), '13151412');
```

# --seed--

## --seed-contents--

```js
function lookAndSay(str) {

}
```

# --solutions--

```js
function lookAndSay(str) {
    return str.replace(/(.)\1*/g, function(seq, p1) {
      return seq.length.toString() + p1;
    });
}
```
