---
id: 5e6dd1278e6ca105cde40ea9
title: 가장 긴 공통 부분열
challengeType: 1
forumTopicId: 385271
dashedName: longest-common-subsequence
---

# --description--

**가장 긴 공통 부분열**(또는 **LCS**)은 두 그룹 A와 B에서 공통으로 나타나고 각 그룹에서 같은 순서로 나오는 가장 긴 요소들의 그룹입니다. 예를 들어, 수열 `1234`와 `1224533324`의 LCS는 `1234`입니다:

<u>1234</u>

<u>12</u>245<u>3</u>332<u>4</u>

문자열 예시로는 수열 `thisisatest`와 `testing123testing`를 생각해 보세요. LCS는 `tsitest`입니다:

<u>t</u>hi<u>si</u>sa<u>test</u>

<u>t</u>e<u>s</u>t<u>i</u>ng123<u>test</u>ing.

코드는 문자열만 다루면 됩니다.

# --instructions--

대소문자를 구분하는 함수를 작성해서 두 문자열의 LCS를 반환하세요. 여러 개의 LCS를 보여줄 필요는 없습니다.

# --hints--

`lcs`는 함수여야 합니다.

```js
assert(typeof lcs == 'function');
```

`lcs("thisisatest", "testing123testing")`는 문자열을 반환해야 합니다.

```js
assert(typeof lcs('thisisatest', 'testing123testing') == 'string');
```

`lcs("thisisatest", "testing123testing")`는 `"tsitest"`를 반환해야 합니다.

```js
assert.equal(lcs('thisisatest', 'testing123testing'), 'tsitest');
```

`lcs("ABCDGH", "AEDFHR")`는 `"ADH"`를 반환해야 합니다.

```js
assert.equal(lcs('ABCDGH', 'AEDFHR'), 'ADH');
```

`lcs("AGGTAB", "GXTXAYB")`는 `"GTAB"`를 반환해야 합니다.

```js
assert.equal(lcs('AGGTAB', 'GXTXAYB'), 'GTAB');
```

`lcs("BDACDB", "BDCB")`는 `"BDCB"`를 반환해야 합니다.

```js
assert.equal(lcs('BDACDB', 'BDCB'), 'BDCB');
```

`lcs("ABAZDC", "BACBAD")`는 `"ABAD"`를 반환해야 합니다.

```js
assert.equal(lcs('ABAZDC', 'BACBAD'), 'ABAD');
```

# --seed--

## --seed-contents--

```js
function lcs(a, b) {

}
```

# --solutions--

```js
function lcs(a, b) {
  var aSub = a.substring(0, a.length - 1);
  var bSub = b.substring(0, b.length - 1);

  if (a.length === 0 || b.length === 0) {
    return '';
  } else if (a.charAt(a.length - 1) === b.charAt(b.length - 1)) {
    return lcs(aSub, bSub) + a.charAt(a.length - 1);
  } else {
    var x = lcs(a, bSub);
    var y = lcs(aSub, b);
    return (x.length > y.length) ? x : y;
  }
}
```
