---
id: 5e6dee7749a0b85a3f1fc7d5
title: 루카스-레머 테스트
challengeType: 1
forumTopicId: 385281
dashedName: lucas-lehmer-test
---

# --description--

루카스-레머 테스트: $p$가 홀수 소수일 때, 메르센 수 $2^p-1$은 $2^p-1$이 $S(p-1)$을 나누는 경우에만 소수입니다. 여기서 $S(n+1)=(S(n))^2-2$이고, $S(1)=4$입니다.

# --instructions--

주어진 메르센 수가 소수인지 아닌지 반환하는 함수를 작성하세요.

# --hints--

`lucasLehmer`는 함수여야 합니다.

```js
assert(typeof lucasLehmer == 'function');
```

`lucasLehmer(11)`는 불리언을 반환해야 합니다.

```js
assert(typeof lucasLehmer(11) == 'boolean');
```

`lucasLehmer(11)`는 `false`를 반환해야 합니다.

```js
assert.equal(lucasLehmer(11), false);
```

`lucasLehmer(15)`는 `false`를 반환해야 합니다.

```js
assert.equal(lucasLehmer(15), false);
```

`lucasLehmer(13)`는 `true`를 반환해야 합니다.

```js
assert.equal(lucasLehmer(13), true);
```

`lucasLehmer(17)`는 `true`를 반환해야 합니다.

```js
assert.equal(lucasLehmer(17), true);
```

`lucasLehmer(19)`는 `true`를 반환해야 합니다.

```js
assert.equal(lucasLehmer(19), true);
```

`lucasLehmer(21)`는 `false`를 반환해야 합니다.

```js
assert.equal(lucasLehmer(21), false);
```

# --seed--

## --seed-contents--

```js
function lucasLehmer(p) {

}
```

# --solutions--

```js
function lucasLehmer(p) {
    function isPrime(p) {
        if (p == 2)
            return true;
        else if (p <= 1 || p % 2 == 0)
            return false;
        else {
            var to = Math.sqrt(p);
            for (var i = 3; i <= to; i += 2)
                if (p % i == 0)
                    return false;
            return true;
        }
    }

    function isMersennePrime(p) {
        if (p == 2)
            return true;
        else {
            var m_p = Math.pow(2, p) - 1
            var s = 4;
            for (var i = 3; i <= p; i++)
                s = (s * s - 2) % m_p
            return s == 0;
        }
    }

    return isPrime(p) && isMersennePrime(p)
}
```
