---
id: 5e6dd139859c290b6ab80292
title: 가장 긴 증가 부분열
challengeType: 1
forumTopicId: 385272
dashedName: longest-increasing-subsequence
---

# --description--

가장 긴 증가 부분열 문제는 주어진 수열에서 부분열의 원소들이 오름차순으로 정렬되어 있고, 그 부분열이 가능한 한 길도록 하는 부분열을 찾는 것입니다. 예를 들어:

다음 배열이 있을 때:

```js
const array = [3, 10, 2, 1, 20];
```

가장 긴 증가 수열은:

$\\{3, 10, 20\\}$

# --instructions--

숫자 배열을 매개변수로 받아 가장 긴 증가 부분열을 반환하는 함수를 작성하세요.

모든 배열에는 반드시 가장 긴 증가 부분열이 존재함이 보장됩니다.

# --hints--

`findSequence`는 함수여야 합니다.

```js
assert(typeof findSequence == 'function');
```

`findSequence([3, 10, 2, 1, 20])`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(findSequence([3, 10, 2, 1, 20])));
```

`findSequence([3, 10, 2, 1, 20])`는 `[3, 10, 20]`를 반환해야 합니다.

```js
assert.deepEqual(findSequence([3, 10, 2, 1, 20]), [3, 10, 20]);
```

`findSequence([2, 7, 3, 5, 8])`는 `[2, 3, 5, 8]`를 반환해야 합니다.

```js
assert.deepEqual(findSequence([2, 7, 3, 5, 8]), [2, 3, 5, 8]);
```

`findSequence([2, 6, 4, 5, 1])`는 `[2, 4, 5]`를 반환해야 합니다.

```js
assert.deepEqual(findSequence([2, 6, 4, 5, 1]), [2, 4, 5]);
```

`findSequence([10, 22, 9, 33, 21, 50, 60, 80])`는 `[10, 22, 33, 50, 60, 80]`를 반환해야 합니다.

```js
assert.deepEqual(findSequence([10, 22, 9, 33, 21, 50, 60, 80]), [
  10,
  22,
  33,
  50,
  60,
  80
]);
```

`findSequence([0, 12, 2, 10, 6, 14, 1, 9, 5, 13, 3, 11, 7, 15])`는 `[0, 2, 6, 9, 11, 15`를 반환해야 합니다.

```js
assert.deepEqual(
  findSequence([0, 12, 2, 10, 6, 14, 1, 9, 5, 13, 3, 11, 7, 15]),
  [0, 2, 6, 9, 11, 15]
);
```

# --seed--

## --seed-contents--

```js
function findSequence(input) {

}
```

# --solutions--

```js
function findSequence(input) {
    var len = input.length;
    var result = []
    for (var i = 0; i < len; i++) result.push(1)

    for (var i = 0; i < len; i++)
        for (var j = i - 1; j >= 0; j--)
            if (input[i] > input[j] && result[j] >= result[i])
                result[i] = result[j] + 1;

    var maxValue = Math.max.apply(null, result);
    var maxIndex = result.indexOf(Math.max.apply(Math, result));
    var output = [];
    output.push(input[maxIndex]);
    for (var i = maxIndex; i >= 0; i--) {
        if (maxValue == 0) break;
        if (input[maxIndex] > input[i] && result[i] == maxValue - 1) {
            output.push(input[i]);
            maxValue--;
        }
    }
    output.reverse();
    return output;
}
```
