---
id: 5e6dd15004c88cf00d2a78b3
title: 여러 배열을 동시에 반복하기
challengeType: 1
forumTopicId: 385279
dashedName: loop-over-multiple-arrays-simultaneously
---

# --description--

여러 배열을 동시에 반복하면서, 주어진 각 배열의 $i^{th}$ 요소를 이어 붙여서 새로운 배열의 $i^{th}$ 요소로 만드세요.

예를 들어, 다음과 같은 배열들의 배열이 주어졌다면:

```js
[ ["a", "b", "c"], ["A", "B", "C"], [1, 2, 3] ]
```

출력은 다음과 같아야 합니다:

```js
["aA1","bB2","cC3"]
```

# --instructions--

배열들의 배열을 매개변수로 받아서, 주어진 설명에 맞는 문자열 배열을 반환하는 함수를 작성하세요.

# --hints--

`loopSimult`는 함수여야 합니다.

```js
assert(typeof loopSimult == 'function');
```

`loopSimult([["a", "b", "c"], ["A", "B", "C"], [1, 2, 3]])`은 배열을 반환해야 합니다.

```js
assert(
  Array.isArray(
    loopSimult([
      ['a', 'b', 'c'],
      ['A', 'B', 'C'],
      [1, 2, 3]
    ])
  )
);
```

`loopSimult([["a", "b", "c"], ["A", "B", "C"], [1, 2, 3]])`는 `["aA1", "bB2", "cC3"]`을 반환해야 합니다.

```js
assert.deepEqual(
  loopSimult([
    ['a', 'b', 'c'],
    ['A', 'B', 'C'],
    [1, 2, 3]
  ]),
  ['aA1', 'bB2', 'cC3']
);
```

`loopSimult([["c", "b", "c"], ["4", "5", "C"], [7, 7, 3]])`는 `["c47", "b57", "cC3"]`를 반환해야 합니다.

```js
assert.deepEqual(
  loopSimult([
    ['c', 'b', 'c'],
    ['4', '5', 'C'],
    [7, 7, 3]
  ]),
  ['c47', 'b57', 'cC3']
);
```

`loopSimult([["a", "b", "c", "d"], ["A", "B", "C", "d"], [1, 2, 3, 4]])`는 `["aA1", "bB2", "cC3", "dd4"]`을 반환해야 합니다.

```js
assert.deepEqual(
  loopSimult([
    ['a', 'b', 'c', 'd'],
    ['A', 'B', 'C', 'd'],
    [1, 2, 3, 4]
  ]),
  ['aA1', 'bB2', 'cC3', 'dd4']
);
```

`loopSimult([["a", "b"], ["A", "B"], [1, 2]])`는 `["aA1", "bB2"]`를 반환해야 합니다.

```js
assert.deepEqual(
  loopSimult([
    ['a', 'b'],
    ['A', 'B'],
    [1, 2]
  ]),
  ['aA1', 'bB2']
);
```

`loopSimult([["b", "c"], ["B", "C"], [2, 3]])`는 `["bB2", "cC3"]`을 반환해야 합니다.

```js
assert.deepEqual(
  loopSimult([
    ['b', 'c'],
    ['B', 'C'],
    [2, 3]
  ]),
  ['bB2', 'cC3']
);
```

# --seed--

## --seed-contents--

```js
function loopSimult(A) {

}
```

# --solutions--

```js
function loopSimult(A) {
    var res = [], output;
    for (var i = 0; i < A[0].length; i += 1) {
        output = "";
        for (var j = 0; j < A.length; j++)
            output += A[j][i];
        res.push(output);
    }
    return res;
}
```
