---
id: 5e6dd14192286d95fc43046e
title: 가장 긴 문자열 도전 과제
challengeType: 1
forumTopicId: 385275
dashedName: longest-string-challenge
---

# --description--

이 도전 과제에서는 주어진 문자열 중 가장 긴 문자열을 찾아야 합니다.

# --instructions--

문자열 배열을 입력받아 가장 긴 길이와 같은 길이를 가진 문자열들을 반환하는 함수를 작성하세요.

# --hints--

`longestString`는 함수여야 합니다.

```js
assert(typeof longestString == 'function');
```

`longestString(["a", "bb", "ccc", "ee", "f", "ggg"])`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(longestString(['a', 'bb', 'ccc', 'ee', 'f', 'ggg'])));
```

`longestString(["a", "bb", "ccc", "ee", "f", "ggg"])`는 `["ccc", "ggg"]`를 반환해야 합니다.

```js
assert.deepEqual(longestString(['a', 'bb', 'ccc', 'ee', 'f', 'ggg']), [
  'ccc',
  'ggg'
]);
```

`longestString(["afedg", "bb", "sdccc", "efdee", "f", "geegg"])`는 `["afedg", "sdccc", "efdee", "geegg"]`를 반환해야 합니다.

```js
assert.deepEqual(
  longestString(['afedg', 'bb', 'sdccc', 'efdee', 'f', 'geegg']),
  ['afedg', 'sdccc', 'efdee', 'geegg']
);
```

`longestString(["a", "bhghgb", "ccc", "efde", "fssdrr", "ggg"])`는 `["bhghgb", "fssdrr"]`를 반환해야 합니다.

```js
assert.deepEqual(
  longestString(['a', 'bhghgb', 'ccc', 'efde', 'fssdrr', 'ggg']),
  ['bhghgb', 'fssdrr']
);
```

`longestString(["ahgfhg", "bdsfsb", "ccc", "ee", "f", "ggdsfg"])`는 `["ahgfhg", "bdsfsb", "ggdsfg"]`를 반환해야 합니다.

```js
assert.deepEqual(
  longestString(['ahgfhg', 'bdsfsb', 'ccc', 'ee', 'f', 'ggdsfg']),
  ['ahgfhg', 'bdsfsb', 'ggdsfg']
);
```

`longestString(["a", "bbdsf", "ccc", "edfe", "gzzzgg"])`는 `["gzzzgg"]`를 반환해야 합니다.

```js
assert.deepEqual(longestString(['a', 'bbdsf', 'ccc', 'edfe', 'gzzzgg']), [
  'gzzzgg'
]);
```

# --seed--

## --seed-contents--

```js
function longestString(strings) {

}
```

# --solutions--

```js
function longestString(strings) {
    var mx = 0;
    var result = []
    strings.forEach(function (e) {
        if (e.length > mx) {
            mx = e.length
            result = [e]
        } else if (e.length == mx)
            result.push(e)
    })

    return result
}
```
