---
id: 5e94a54cc7b022105bf0fd2c
title: 단어 빈도
challengeType: 1
forumTopicId: 393913
dashedName: word-frequency
---

# --description--

문자열과 정수 n이 주어지면, 파일에서 가장 자주 등장하는 n개의 단어와 그 등장 횟수를 빈도 순으로 반환하세요.

# --instructions--

각 단어의 등장 횟수를 세고, 등장 횟수가 많은 순서대로 n개의 단어와 그 횟수를 반환하는 함수를 작성하세요.

함수는 2차원 배열을 반환해야 하며, 각 요소는 다음 형식이어야 합니다: `[word, freq]`. `word`은 단어의 소문자 버전이고 `freq`는 등장 횟수를 나타내는 숫자입니다.

문자열이 제공되지 않으면 함수는 빈 배열을 반환해야 합니다.

함수는 대소문자를 구분하지 않아야 하며, 예를 들어 "Hello"와 "hello"는 같은 단어로 처리해야 합니다.

밑줄, 대시, 아포스트로피, 쉼표 등 특수 문자가 포함된 단어는 서로 다른 단어로 취급할 수 있습니다.

예를 들어, 문자열 "Hello hello goodbye"가 주어지면, 함수는 `[['hello', 2], ['goodbye', 1]]`를 반환해야 합니다.

# --hints--

`wordFrequency`는 함수여야 합니다.

```js
assert.isFunction(wordFrequency);
```

`wordFrequency`는 배열을 반환해야 합니다.

```js
assert.isArray(wordFrequency('test'));
```

`wordFrequency("Hello hello world", 2)`는 `[['hello', 2], ['world', 1]]`를 반환해야 합니다.

```js
const example_1 = 'Hello hello world';
const example_1_solution = [['hello', 2], ['world', 1]];

assert.deepEqual(wordFrequency(example_1, 2), example_1_solution);
```

`wordFrequency("The quick brown fox jumped over the lazy dog", 1)`는 `[['the', 2]]`를 반환해야 합니다.

```js
const example_2 = 'The quick brown fox jumped over the lazy dog';
const example_2_solution = [['the', 2]];

assert.deepEqual(wordFrequency(example_2, 1), example_2_solution);
```

`wordFrequency("Opensource opensource open-source open source", 1)`는 `[['opensource', 2]]`를 반환해야 합니다.

```js
const example_3 = 'Opensource opensource open-source open source';
const example_3_solution = [['opensource', 2]];

assert.deepEqual(wordFrequency(example_3, 1), example_3_solution);
```

`wordFrequency("Apple App apply aPP aPPlE", 3)`는 `[['app', 2], ['apple', 2], ['apply', 1]]` 또는 `[['apple', 2], ['app', 2], ['apply', 1]]`를 반환해야 합니다.

```js
const example_4 = 'Apple App apply aPP aPPlE';
const example_4_solution = [["app", 2],["apple", 2],["apply", 1]];

const arr = wordFrequency(example_4, 3);
assert.includeDeepMembers(arr, example_4_solution);
```

`wordFrequency("c d a d c a b d d c", 4)`는 `[['d', 4], ['c', 3], ['a', 2], ['b', 1]]`를 반환해야 합니다.

```js
const example_5 = 'c d a d c a b d d c';
const example_5_solution = [['d', 4], ['c', 3], ['a', 2], ['b', 1]];

assert.deepEqual(wordFrequency(example_5, 4), example_5_solution);
```

`wordFrequency("", 5)`는 `[]`를 반환해야 합니다.

```js
const example_6 = '';
const example_6_solution = [];

assert.deepEqual(wordFrequency(example_6, 5), example_6_solution);
```

# --seed--

## --seed-contents--

```js
function wordFrequency(txt, n) {

}
```

# --solutions--

```js
function wordFrequency(txt, n) {
  var words = txt.split(/\s+/);
  var wordCount = {};
  words.forEach(word => {
    if (word == '') {
      return;
    }
    const lowerWord = word.toLowerCase();
    wordCount[lowerWord] =
      lowerWord in wordCount ? wordCount[lowerWord] + 1 : 1;
  });

  var wordsArray = [];
  for (let [word, count] of Object.entries(wordCount)) {
    wordsArray.push([word, count]);
  }

  wordsArray.sort((a, b) => {
    if (a[1] !== b[1]) {
      return b[1] - a[1];
    } else if (a[0] !== b[0]) {
      return a[0] < b[0] ? -1 : 1;
    }
    return 0;
  });
  return wordsArray.slice(0, n);
}
```
