---
id: 5e9ddb06ec35240f39657419
title: 피즈버즈
challengeType: 1
forumTopicId: 385370
dashedName: fizzbuzz
---

# --description--

1부터 100까지(포함) 정수 배열을 생성하는 프로그램을 작성하세요. 단:

<ul>
    <li>3의 배수에는 숫자 대신 <code>"Fizz"</code>를 배열에 추가하세요</li>
    <li>5의 배수에는 숫자 대신 <code>"Buzz"</code>를 배열에 추가하세요</li>
    <li>3과 5의 배수에는 숫자 대신 <code>"FizzBuzz"</code>를 배열에 추가하세요</li>
</ul>

# --instructions--

프로그램은 위 규칙에 따라 결과를 담은 배열을 반환해야 합니다.

# --hints--

`fizzBuzz`는 함수여야 합니다.

```js
assert(typeof fizzBuzz == 'function');
```

`fizzBuzz()`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(fizzBuzz()) == true);
```

3으로만 나누어지는 숫자는 `"Fizz"`를 반환해야 합니다.

```js
assert.equal(fizzBuzz()[2], 'Fizz');
```

5로만 나누어지는 숫자는 `"Buzz"`를 반환해야 합니다.

```js
assert.equal(fizzBuzz()[99], 'Buzz');
```

3과 5 모두로 나누어지는 숫자는 `"FizzBuzz"`를 반환해야 합니다.

```js
assert.equal(fizzBuzz()[89], 'FizzBuzz');
```

3이나 5로 나누어지지 않는 숫자는 그 숫자 자체를 반환해야 합니다.

```js
assert.equal(fizzBuzz()[12], 13);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz() {

}
```

# --solutions--

```js
function fizzBuzz() {
    let res=[];
    for (let i =1; i < 101; i++) {
        if (i % 3 === 0  && i % 5 === 0) {
            res.push("FizzBuzz");
        }
        else if (i % 3 === 0) {
            res.push("Fizz");
        }
        else if (i % 5 === 0) {
            res.push("Buzz");
        } 
        else {
            res.push(i);
        }
    }
    return res;
}
```
