---
id: 5ea281203167d2b0bdefca00
title: 루딕 수
challengeType: 1
forumTopicId: 385282
dashedName: ludic-numbers
---

# --description--

루딕 수는 소수와 관련이 있으며, 소수를 생성하는 데 사용되는 에라토스테네스의 체와 매우 비슷한 체를 통해 생성됩니다.

첫 번째 루딕 수는 1입니다.

다음 루딕 수를 생성하려면 2부터 시작하는 증가하는 정수 배열을 만드세요.

<code style='margin-left: 2em;'><span style='color:blue;font-weight:bold'>2</span> 3 4 5 6 7 8 9 10 11 12 13 14 15 16 17 18 19 20 21 22 23 24 25 26 ...</code>

(반복)

<ul>
  <li>결과 배열의 첫 번째 요소를 다음 루딕 수 <span style='color:blue;font-weight:bold'>2</span>로 선택하세요.</li>
  <li>배열에서 <strong>2<sup>번째</sup></strong> 인덱스 항목(첫 번째 포함)을 모두 제거하세요.</li>
  <code style='margin-left: 2em;'><span style='color:blue;font-weight:bold;'><s>2</s></span> 3 <s>4</s> 5 <s>6</s> 7 <s>8</s> 9 <s>10</s> 11 <s>12</s> 13 <s>14</s> 15 <s>16</s> 17 <s>18</s> 19 <s>20</s> 21 <s>22</s> 23 <s>24</s> 25 <s>26</s> ...</code>
</ul>

<ul>
  <li>(몇 번의 반복을 펼치면...)</li>
  <li>결과 배열의 첫 번째 요소를 다음 루딕 수 <span style='color:blue;font-weight:bold'>3</span>로 선택하세요.</li>
  <li>배열에서 <strong>3<sup>번째</sup></strong> 인덱스 항목(첫 번째 포함)을 모두 제거하세요.</li>
  <code style='margin-left: 2em;'><span style='color:blue;font-weight:bold'><s>3</s></span> 5 7 <s>9</s> 11 13 <s>15</s> 17 19 <s>21</s> 23 25 <s>27</s> 29 31 <s>33</s> 35 37 <s>39</s> 41 43 <s>45</s> 47 49 <s>51</s> ...</code>
</ul>

<ul>
  <li>결과 배열의 첫 번째 요소를 다음 루딕 수 <span style='color:blue;font-weight:bold'>5</span>로 선택하세요.</li>
  <li>배열에서 <strong>5<sup>번째</sup></strong> 인덱스 항목(첫 번째 포함)을 모두 제거하세요.</li>
  <code style='margin-left: 2em;'><span style='color:blue;font-weight:bold'><s>5</s></span> 7 11 13 17 <s>19</s> 23 25 29 31 <s>35</s> 37 41 43 47 <s>49</s> 53 55 59 61 <s>65</s> 67 71 73 77 ...</code>
</ul>

<ul>
  <li>결과 배열의 첫 번째 요소를 다음 루딕 수 <span style='color:blue;font-weight:bold'>7</span>로 선택하세요.</li>
  <li>배열에서 <strong>7<sup>번째</sup></strong> 인덱스 항목(첫 번째 포함)을 모두 제거하세요.</li>
  <code style='margin-left: 2em;'><span style='color:blue;font-weight:bold'><s>7</s></span> 11 13 17 23 25 29 <s>31</s> 37 41 43 47 53 55 <s>59</s> 61 67 71 73 77 83 <s>85</s> 89 91 97 ...</code>
</ul>

<ul>
  <li><big><b> ... </b></big></li>
  <li>현재 배열의 첫 번째 요소를 다음 루딕 수 <span style='color:blue;font-weight:bold'>L</span>로 선택하세요.</li>
  <li>배열에서 <strong>L<sup>번째</sup></strong> 인덱스 항목(첫 번째 포함)을 모두 제거하세요.</li>
  <li><big><b> ... </b></big></li>
</ul>

# --instructions--

주어진 숫자보다 작거나 같은 모든 루딕 수를 반환하는 함수를 작성하세요.

# --hints--

`ludic`는 함수여야 합니다.

```js
assert(typeof ludic === 'function', '<code>ludic</code> should be a function.');
```

`ludic(2)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(ludic(2)));
```

`ludic(2)`는 `[1, 2]`를 반환해야 합니다.

```js
assert.deepEqual(ludic(2), [1, 2]);
```

`ludic(3)`는 `[1, 2, 3]`를 반환해야 합니다.

```js
assert.deepEqual(ludic(3), [1, 2, 3]);
```

`ludic(5)`는 `[1, 2, 3, 5]`를 반환해야 합니다.

```js
assert.deepEqual(ludic(5), [1, 2, 3, 5]);
```

`ludic(20)`는 `[1, 2, 3, 5, 7, 11, 13, 17]`를 반환해야 합니다.

```js
assert.deepEqual(ludic(20), [1, 2, 3, 5, 7, 11, 13, 17]);
```

`ludic(26)`는 `[1, 2, 3, 5, 7, 11, 13, 17, 23, 25]`를 반환해야 합니다.

```js
assert.deepEqual(ludic(26), [1, 2, 3, 5, 7, 11, 13, 17, 23, 25]);
```

# --seed--

## --seed-contents--

```js
function ludic(n) {

}
```

# --solutions--

```js
function ludic(n) {
  const makeArr = (s, e) => new Array(e + 1 - s).fill(s).map((e, i) => e + i);

  const filterAtInc = (arr, n) => arr.filter((e, i) => (i + 1) % n);

  const makeLudic = (arr, result) => {
    const iter = arr.shift();
    result.push(iter);
    return arr.length ? makeLudic(filterAtInc(arr, iter), result) : result;
  };

  const ludicResult = makeLudic(makeArr(2, n), [1]);

  return ludicResult;
}
```
