---
id: 5ea28156e79528a9ab248f27
title: 신용카드 번호의 Luhn 검사
challengeType: 1
forumTopicId: 385284
dashedName: luhn-test-of-credit-card-numbers
---

# --description--

Luhn 검사는 일부 신용카드 회사에서 유효한 신용카드 번호를 무작위 숫자 조합과 구별하기 위해 사용합니다.

Luhn 검사를 통해 유효성을 확인할 수 있는 신용카드 번호를 사용하는 회사들은 다음 검사를 통과하는 번호를 가집니다:

<ol>
  <li> 번호의 숫자 순서를 뒤집습니다.</li>
  <li> 뒤집힌 숫자에서 첫 번째, 세 번째, ... 홀수 번째 숫자들을 골라 합산하여 부분합 s1을 만듭니다.</li>
  <li> 뒤집힌 숫자에서 두 번째, 네 번째, ... 짝수 번째 숫자들을 골라:</li>
    <ol>
      <li>각 숫자에 2를 곱하고, 결과가 9보다 크면 각 자리 숫자를 더해 짝수 숫자들의 부분합을 만듭니다.</li>
      <li>짝수 숫자들의 부분합을 합산하여 s2를 만듭니다.</li>
    </ol>
  <li>s1 + s2의 합이 0으로 끝나면 원래 숫자는 Luhn 검사로 확인된 유효한 신용카드 번호 형식입니다.</li>
</ol>

예를 들어, 시험 번호가 49927398716이라면:

```md
Reverse the digits:
  61789372994
Sum the odd digits:
  6 + 7 + 9 + 7 + 9 + 4 = 42 = s1
The even digits:
    1, 8, 3, 2, 9
  Two times each even digit:
    2, 16, 6, 4, 18
  Sum the digits of each multiplication:
    2, 7, 6, 4, 9
  Sum the last:
    2 + 7 + 6 + 4 + 9 = 28 = s2

s1 + s2 = 70 which ends in zero which means that 49927398716 passes the Luhn test.
```

# --instructions--

Luhn 검사를 통해 숫자의 유효성을 확인하는 함수를 작성하세요. 유효한 번호라면 true를 반환하고, 그렇지 않으면 false를 반환하세요.

# --hints--

`luhnTest`는 함수여야 합니다.

```js
assert(typeof luhnTest === 'function');
```

`luhnTest("4111111111111111")`는 불리언을 반환해야 합니다.

```js
assert(typeof luhnTest('4111111111111111') === 'boolean');
```

`luhnTest("4111111111111111")`는 `true`를 반환해야 합니다.

```js
assert.equal(luhnTest('4111111111111111'), true);
```

`luhnTest("4111111111111112")`는 `false`를 반환해야 합니다.

```js
assert.equal(luhnTest('4111111111111112'), false);
```

`luhnTest("49927398716")`는 `true`를 반환해야 합니다.

```js
assert.equal(luhnTest('49927398716'), true);
```

`luhnTest("49927398717")`는 `false`를 반환해야 합니다.

```js
assert.equal(luhnTest('49927398717'), false);
```

`luhnTest("1234567812345678")`는 `false`를 반환해야 합니다.

```js
assert.equal(luhnTest('1234567812345678'), false);
```

`luhnTest("1234567812345670")`는 `true`를 반환해야 합니다.

```js
assert.equal(luhnTest('1234567812345670'), true);
```

# --seed--

## --seed-contents--

```js
function luhnTest(str) {

}
```

# --solutions--

```js
function luhnTest(str) {
  var luhnArr = [0, 2, 4, 6, 8, 1, 3, 5, 7, 9];
  var counter = 0;
  var incNum;
  var odd = false;
  var temp = String(str).replace(/[^\d]/g, '');
  if (temp.length == 0) return false;
  for (var i = temp.length - 1; i >= 0; --i) {
    incNum = parseInt(temp.charAt(i), 10);
    counter += (odd = !odd) ? incNum : luhnArr[incNum];
  }
  return counter % 10 == 0;
}
```
