---
id: 5ea2815a8640bcc6cb7dab3c
title: 라이크렐 수
challengeType: 1
forumTopicId: 385287
dashedName: lychrel-numbers
---

# --description--

<ol>
  <li>0보다 큰 정수 <code>n₀</code>를 선택하세요.</li>
  <li><code>n₀</code>를 뒤집은 수와 <code>n₀</code>를 더해 다음 수 <code>n</code>를 만드세요.</li>
  <li><code>n</code>가 회문, 즉 <code>n</code>의 숫자를 거꾸로 읽은 값이 <code>n</code>와 같아질 때 멈추세요.</li>
</ol>

위 점화식을 대부분의 시작 숫자 `n` = 1, 2, ...에 적용하면 회문이 빠르게 만들어져 종료됩니다.

예를 들어 `n₀` = 12일 때:

```md
12
12 + 21 = 33,  a palindrome!
```

그리고 `n₀` = 55일 때:

```md
55
55 + 55 = 110
110 + 011 = 121,  a palindrome!
```

덧셈 후에 회문인지 확인하는 점에 주목하세요.

일부 시작 숫자는 무한히 계속되는 것처럼 보입니다; 196에 대한 점화식은 수백만 번 반복되어 수백만 자리 숫자가 만들어졌지만 회문이 형성되지 않았습니다. 회문으로 끝나지 않는 이런 숫자를 **라이크렐 수**라고 합니다.

이 도전 과제에서는 500회(또는 그 이상) 반복 내에 회문이 만들어지지 않는 시작 숫자를 라이크렐 수로 간주합니다.

**시드와 관련 라이크렐 수:**

라이크렐 수의 수열에서 생성된 모든 정수도 라이크렐 수입니다.

일반적으로, 한 라이크렐 수에서 시작한 수열은 이전 라이크렐 수 후보의 수열과 합쳐질 *수도* 있습니다; 예를 들어 196과 689의 수열은 다음과 같이 시작합니다:

```md
    196
    196 + 691 = 887
    887 + 788 = 1675
    1675 + 5761 = 7436
    7436 + 6347 = 13783
    13783 + 38731 = 52514
    52514 + 41525 = 94039
    ...
    689
    689 + 986 = 1675
    1675 + 5761 = 7436
    ...
```

따라서 689에서 시작한 수열이 196의 수열과 합쳐져 같은 숫자로 계속됨을 알 수 있습니다.

이 때문에 라이크렐 수를 진짜 **시드** 라이크렐 수 후보와 회문을 만들지 않지만 더 낮은 라이크렐 수의 수열에 포함된 정수를 가진 **관련** 수로 나눌 수 있습니다.

# --instructions--

숫자를 매개변수로 받는 함수를 작성하세요. 숫자가 라이크렐 수이면 true를 반환하세요. 그렇지 않으면 false를 반환하세요. 반복 제한은 500회임을 기억하세요.

# --hints--

`isLychrel`는 함수여야 합니다.

```js
assert(typeof isLychrel === 'function');
```

`isLychrel(12)`는 불리언을 반환해야 합니다.

```js
assert(typeof isLychrel(12) === 'boolean');
```

`isLychrel(12)`는 `false`를 반환해야 합니다.

```js
assert.equal(isLychrel(12), false);
```

`isLychrel(55)`는 `false`를 반환해야 합니다.

```js
assert.equal(isLychrel(55), false);
```

`isLychrel(196)`는 `true`를 반환해야 합니다.

```js
assert.equal(isLychrel(196), true);
```

`isLychrel(879)`는 `true`를 반환해야 합니다.

```js
assert.equal(isLychrel(879), true);
```

`isLychrel(44987)`는 `false`를 반환해야 합니다.

```js
assert.equal(isLychrel(44987), false);
```

`isLychrel(7059)`는 `true`를 반환해야 합니다.

```js
assert.equal(isLychrel(7059), true);
```

# --seed--

## --seed-contents--

```js
function isLychrel(n) {

}
```

# --solutions--

```js
function isLychrel(n) {
  function reverse(num) {
    return parseInt(
      num
        .toString()
        .split('')
        .reverse()
        .join('')
    );
  }

  var i;
  for (i = 0; i < 500; i++) {
    n = n + reverse(n);
    if (n == reverse(n)) break;
  }

  return i == 500;
}
```
