---
id: 5eaf48389ee512d4d103684b
title: 자기 설명 수
challengeType: 1
forumTopicId: 385289
dashedName: self-describing-numbers
---

# --description--

소위 "자기 설명 수" 또는 "자기 설명 정수"라고 불리는 정수가 있습니다.

정수의 자리 위치를 0부터 N-1까지 레이블링했을 때, 각 자리의 숫자가 그 숫자가 해당 정수 내에 나타나는 횟수와 같으면 그 정수를 "자기 설명 수"라고 합니다.

예를 들어, **2020**은 네 자리 자기 설명 수입니다:

<ul>
    <li> 0번 자리의 값은 2이고, 숫자 0이 두 번 나타납니다; </li>
    <li> 1번 자리의 값은 0이고, 숫자 1이 전혀 없습니다; </li>
    <li> 2번 자리의 값은 2이고, 숫자 2가 두 번 나타납니다; </li>
    <li> 3번 자리의 값은 0이고, 숫자 3이 없습니다; </li>
</ul>

100,000,000 미만의 자기 설명 수는 1210, 2020, 21200, 3211000, 42101000입니다.

# --instructions--

양의 정수를 매개변수로 받는 함수를 작성하세요. 만약 그 정수가 자기 설명 수라면 true를 반환하세요. 그렇지 않으면 false를 반환하세요.

# --hints--

`isSelfDescribing`는 함수여야 합니다.

```js
assert(typeof isSelfDescribing == 'function');
```

`isSelfDescribing()`는 불리언을 반환해야 합니다.

```js
assert(typeof isSelfDescribing(2020) == 'boolean');
```

`isSelfDescribing(2020)`는 `true`를 반환해야 합니다.

```js
assert.equal(isSelfDescribing(2020), true);
```

`isSelfDescribing(3021)`는 `false`를 반환해야 합니다.

```js
assert.equal(isSelfDescribing(3021), false);
```

`isSelfDescribing(3211000)`는 `true`를 반환해야 합니다.

```js
assert.equal(isSelfDescribing(3211000), true);
```

# --seed--

## --seed-contents--

```js
function isSelfDescribing(n) {

}
```

# --solutions--

```js
function isSelfDescribing(n) {
    let digits = String(n).split("");
    digits = digits.map(function(e) {return parseInt(e)});
    let count = digits.map((x) => {return 0})
    digits.forEach((d) =>{
        if (d >= count.length) {
            return false
        }
        count[d] += 1;
    });

     if (digits === count) {
        return true;
    }
    if (digits.length != count.length) {
        return false;
    }
    
    for (let i=0; i< digits.length; i++){
      if (digits[i] !== count[i]) {
        return false;
      }
    }
    return true;
}
```
