---
id: 5eb3e4aa847216613aa81983
title: 세미소수
challengeType: 1
forumTopicId: 385318
dashedName: semiprime
---

# --description--

세미소수는 정확히 두 개의 (같을 수도 있는) <a href="https://rosettacode.org/wiki/prime_number" target="_blank" rel="noopener noreferrer nofollow">소수</a>의 곱인 자연수입니다.

<pre>1679  =  23 x 73</pre>

# --instructions--

숫자가 세미소수이면 true를, 그렇지 않으면 false를 반환하는 함수를 작성하세요.

# --hints--

`isSemiPrime`는 함수여야 합니다.

```js
assert(typeof isSemiPrime === 'function');
```

`isSemiPrime(100)`는 불리언을 반환해야 합니다.

```js
assert(typeof isSemiPrime(100) === 'boolean');
```

`isSemiPrime(100)`는 `false`를 반환해야 합니다.

```js
assert.equal(isSemiPrime(100), false);
```

`isSemiPrime(504)`는 `false`를 반환해야 합니다.

```js
assert.equal(isSemiPrime(504), false);
```

`isSemiPrime(4)`는 `true`를 반환해야 합니다.

```js
assert.equal(isSemiPrime(4), true);
```

`isSemiPrime(46)`는 `true`를 반환해야 합니다.

```js
assert.equal(isSemiPrime(46), true);
```

`isSemiPrime(13)`는 `false`를 반환해야 합니다.

```js
assert.equal(isSemiPrime(13), false);
```

`isSemiPrime(74)`는 `true`를 반환해야 합니다.

```js
assert.equal(isSemiPrime(74), true);
```

`isSemiPrime(1679)`는 `true`를 반환해야 합니다.

```js
assert.equal(isSemiPrime(1679), true);
```

`isSemiPrime(2)`는 `false`를 반환해야 합니다.

```js
assert.equal(isSemiPrime(2), false);
```

`isSemiPrime(95)`는 `true`를 반환해야 합니다.

```js
assert.equal(isSemiPrime(95), true);
```

`isSemiPrime(124)`는 `false`를 반환해야 합니다.

```js
assert.equal(isSemiPrime(124), false);
```

# --seed--

## --seed-contents--

```js
function isSemiPrime(n) {

}
```

# --solutions--

```js
function isSemiPrime(n) {
  if (n <= 3) return false;

  var ans = [];
  var done = false;
  while (!done) {
    if (n % 2 === 0) {
      ans.push(2);
      n /= 2;
      continue;
    }
    if (n % 3 === 0) {
      ans.push(3);
      n /= 3;
      continue;
    }
    if (n === 1) return ans.length == 2;
    var sr = Math.sqrt(n);
    done = true;
    // try to divide the checked number by all numbers till its square root.
    for (var i = 6; i <= sr; i += 6) {
      if (n % (i - 1) === 0) {
        // is n divisible by i-1?
        ans.push(i - 1);
        n /= i - 1;
        done = false;
        break;
      }
      if (n % (i + 1) === 0) {
        // is n divisible by i+1?
        ans.push(i + 1);
        n /= i + 1;
        done = false;
        break;
      }
    }
  }
  ans.push(n);
  return ans.length == 2;
}
```
