---
id: 5eb3e4b20aa93c437f9e9717
title: 실수 집합
challengeType: 1
forumTopicId: 385322
dashedName: set-of-real-numbers
---

# --description--

모든 실수는 셀 수 없는 집합 ℝ를 이룹니다. 그 부분집합 중 비교적 단순한 것은 볼록 집합으로, 각각은 *a* ≤ *b*인 두 실수 *a*와 *b* 사이의 구간으로 나타냅니다. "사이"의 의미는 경계가 열린지 닫힌지에 따라 실제로 네 가지 경우가 있습니다:

<ul>
  <li>[<i>a</i>, <i>b</i>]: {<i>x</i> | <i>a</i> ≤ <i>x</i> 그리고 <i>x</i> ≤ <i>b</i> }</li>
  <li>(<i>a</i>, <i>b</i>): {<i>x</i> | <i>a</i> < <i>x</i> 그리고 <i>x</i> < <i>b</i> }</li>
  <li>[<i>a</i>, <i>b</i>): {<i>x</i> | <i>a</i> ≤ <i>x</i> 그리고 <i>x</i> < <i>b</i> }</li>
  <li>(<i>a</i>, <i>b</i>]: {<i>x</i> | <i>a</i> < <i>x</i> 그리고 <i>x</i> ≤ <i>b</i> }</li>
</ul>

*a* = *b*인 경우 네 가지 중 \[*a*, *a*]만 비어 있지 않다는 점에 유의하세요.

**과제**

<ul>
  <li>아래 구현 노트에서 "어떤"의 정의에 따라 임의의 실수 집합을 나타낼 방법을 고안하세요.</li>
  <li>다음과 같은 일반적인 집합 연산에 대한 메서드를 제공하세요 (<i>x</i>는 실수, <i>A</i>와 <i>B</i>는 집합입니다):</li>
  <ul>
    <li>
      <i>x</i> ∈ <i>A</i>: <i>x</i>가 <i>A</i>의 원소인지 판단하기<br>
      예: 1은 [1, 2) 안에 있지만 2, 3, ...은 아닙니다.
    </li>
    <li>
      <i>A</i> ∪ <i>B</i>: <i>A</i>와 <i>B</i>의 합집합, 즉 {<i>x</i> | <i>x</i> ∈ <i>A</i> 또는 <i>x</i> ∈ <i>B</i>}<br>
      예: [0, 2) ∪ (1, 3) = [0, 3); [0, 1) ∪ (2, 3] = 그냥 [0, 1) ∪ (2, 3]
    </li>
    <li>
      <i>A</i> ∩ <i>B</i>: <i>A</i>와 <i>B</i>의 교집합, 즉 {<i>x</i> | <i>x</i> ∈ <i>A</i> 그리고 <i>x</i> ∈ <i>B</i>}<br>
      예: [0, 2) ∩ (1, 3) = (1, 2); [0, 1) ∩ (2, 3] = 공집합
    </li>
    <li>
      <i>A</i> - <i>B</i>: <i>A</i>와 <i>B</i>의 차집합, <i>A</i> \ <i>B</i>로도 쓰며, 즉 {<i>x</i> | <i>x</i> ∈ <i>A</i> 그리고 <i>x</i> ∉ <i>B</i>}<br>
      예: [0, 2) − (1, 3) = [0, 1]
    </li>
  </ul>
</ul>

# --instructions--

두 개의 객체, 문자열과 배열을 매개변수로 받는 함수를 작성하세요. 객체는 집합을 나타내며 속성으로 `low`, `high` 그리고 `rangeType`를 가집니다.

`rangeType`는 `CLOSED`, `BOTH_OPEN`, `LEFT_OPEN`, `RIGHT_OPEN`에 대해 각각 0, 1, 2, 3 값을 가질 수 있습니다. 함수는 이 정보를 사용해 집합을 구현해야 합니다.

문자열은 집합에 수행할 연산을 나타내며 `"union"`, `"intersect"`, `"subtract"` (차집합) 중 하나일 수 있습니다.

연산을 수행한 후, 배열에 있는 값들이 결과 집합에 포함되는지 확인하고 그에 해당하는 불리언 값을 배열에 저장하세요. 함수는 이 배열을 반환해야 합니다.

# --hints--

`realSet`는 함수여야 합니다.

```js
assert(typeof realSet == 'function');
```

`realSet({"low":0, "high":1, "rangeType":2}, {"low":0, "high":2, "rangeType":3}, "union", [1, 2, 3])`는 배열을 반환해야 합니다.

```js
assert(
  Array.isArray(
    realSet(
      { low: 0, high: 1, rangeType: 2 },
      { low: 0, high: 2, rangeType: 3 },
      'union',
      [1, 2, 3]
    )
  )
);
```

`realSet({"low":0, "high":1, "rangeType":2}, {"low":0, "high":2, "rangeType":3}, "union", [1, 2, 3])`는 `[true, false, false]`를 반환해야 합니다.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 1, rangeType: 2 },
    { low: 0, high: 2, rangeType: 3 },
    'union',
    [1, 2, 3]
  ),
  [true, false, false]
);
```

`realSet({"low":0, "high":2, "rangeType":3}, {"low":1, "high":2, "rangeType":2}, "intersect", [0, 1, 2])`는 `[false, false, false]`를 반환해야 합니다.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 2, rangeType: 3 },
    { low: 1, high: 2, rangeType: 2 },
    'intersect',
    [0, 1, 2]
  ),
  [false, false, false]
);
```

`realSet({"low":0, "high":3, "rangeType":3}, {"low":0, "high":1, "rangeType":1}, "subtract", [0, 1, 2])`는 `[true, true, true]`를 반환해야 합니다.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 3, rangeType: 3 },
    { low: 0, high: 1, rangeType: 1 },
    'subtract',
    [0, 1, 2]
  ),
  [true, true, true]
);
```

`realSet({"low":0, "high":3, "rangeType":3}, {"low":0, "high":1, "rangeType":0}, "subtract", [0, 1, 2])`는 `[false, false, true]`를 반환해야 합니다.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 3, rangeType: 3 },
    { low: 0, high: 1, rangeType: 0 },
    'subtract',
    [0, 1, 2]
  ),
  [false, false, true]
);
```

`realSet({"low":0, "high":33, "rangeType":1}, {"low":30, "high":31, "rangeType":0}, "intersect", [30, 31, 32])`는 `[true, true, false]`를 반환해야 합니다.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 33, rangeType: 1 },
    { low: 30, high: 31, rangeType: 0 },
    'intersect',
    [30, 31, 32]
  ),
  [true, true, false]
);
```

# --seed--

## --seed-contents--

```js
function realSet(set1, set2, operation, values) {

}
```

# --solutions--

```js
function realSet(set1, set2, operation, values) {
  const RangeType = {
    CLOSED: 0,
    BOTH_OPEN: 1,
    LEFT_OPEN: 2,
    RIGHT_OPEN: 3
  };

  function Predicate(test) {
    this.test = test;
    this.or = function(other) {
      return new Predicate(t => this.test(t) || other.test(t));
    };
    this.and = function(other) {
      return new Predicate(t => this.test(t) && other.test(t));
    };
    this.negate = function() {
      return new Predicate(t => !this.test(t));
    };
  }

  function RealSet(start, end, rangeType, predF) {
    this.low = start;
    this.high = end;

    if (predF) {
      this.predicate = new Predicate(predF);
    } else {
      this.predicate = new Predicate(d => {
        switch (rangeType) {
          case RangeType.CLOSED:
            return start <= d && d <= end;
          case RangeType.BOTH_OPEN:
            return start < d && d < end;
          case RangeType.LEFT_OPEN:
            return start < d && d <= end;
          case RangeType.RIGHT_OPEN:
            return start <= d && d < end;
        }
      });
    }

    this.contains = function(d) {
      return this.predicate.test(d);
    };

    this.union = function(other) {
      var low2 = Math.min(this.low, other.low);
      var high2 = Math.max(this.high, other.high);
      return new RealSet(low2, high2, null, d =>
        this.predicate.or(other.predicate).test(d)
      );
    };

    this.intersect = function(other) {
      var low2 = Math.min(this.low, other.low);
      var high2 = Math.max(this.high, other.high);
      return new RealSet(low2, high2, null, d =>
        this.predicate.and(other.predicate).test(d)
      );
    };

    this.subtract = function(other) {
      return new RealSet(this.low, this.high, null, d =>
        this.predicate.and(other.predicate.negate()).test(d)
      );
    };
  }
  set1 = new RealSet(set1.low, set1.high, set1.rangeType);
  set2 = new RealSet(set2.low, set2.high, set2.rangeType);
  var result = [];
  values.forEach(function(value) {
    result.push(set1[operation](set2).contains(value));
  });
  return result;
}
```
