---
id: 5eb3e497b8d6d7f63c5517ea
title: 기록 목록 검색
challengeType: 1
forumTopicId: 385315
dashedName: search-a-list-of-records
---

# --description--

기록은 엔터티를 설명하는 속성들로 구성됩니다. 각 속성은 이름과 값을 가집니다. 예를 들어, 사람은 `age`라는 속성을 가지고 값이 25일 수 있습니다. 기록 목록에서 중요한 작업 중 하나는 특정 속성 값과 일치하는 기록을 찾는 것입니다.

# --instructions--

문자열을 매개변수로 받는 함수를 작성하세요. 이 함수는 `list`에서 `name` 속성의 값이 주어진 문자열과 일치하는 항목의 인덱스를 반환해야 합니다.

# --hints--

`searchCity`는 함수여야 합니다.

```js
assert(typeof searchCity === 'function');
```

`searchCity("Dar Es Salaam")`는 숫자를 반환해야 합니다.

```js
assert(typeof searchCity('Dar Es Salaam') === 'number');
```

`searchCity("Dar Es Salaam")`는 `6`를 반환해야 합니다.

```js
assert.equal(searchCity('Dar Es Salaam'), 6);
```

`searchCity("Casablanca")`는 `9`를 반환해야 합니다.

```js
assert.equal(searchCity('Casablanca'), 9);
```

`searchCity("Cairo")`는 `1`를 반환해야 합니다.

```js
assert.equal(searchCity('Cairo'), 1);
```

`searchCity("Mogadishu")`는 `4`를 반환해야 합니다.

```js
assert.equal(searchCity('Mogadishu'), 4);
```

`searchCity("Lagos")`는 `0`를 반환해야 합니다.

```js
assert.equal(searchCity('Lagos'), 0);
```

# --seed--

## --seed-contents--

```js
function searchCity(name) {

}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```

# --solutions--

```js
function searchCity(name) {
  return list.findIndex(item => item.name === name);
}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```
