---
id: 5eb3e4a21f462f409d656c73
title: 자기참조 수열
challengeType: 1
forumTopicId: 385317
dashedName: self-referential-sequence
---

# --description--

자기참조 수열을 생성하는 방법은 여러 가지가 있습니다. 그중 아주 흔한 방법 하나는 (<a href="https://rosettacode.org/wiki/Look-and-say_sequence" target="_blank" rel="noopener noreferrer nofollow">Look-and-say 수열</a>) 양의 정수로 시작해서, 인접한 같은 숫자들의 그룹을 열거하여 이어 붙여 다음 항을 생성하는 것입니다:

<pre>0, 10, 1110, 3110, 132110, 1113122110, 311311222110 ...</pre>

생성된 항들은 길이가 기하급수적으로 늘어나며 수렴하지 않습니다.

또 다른 자기참조 수열 생성 방법은 이전 항을 요약하는 것입니다.

각 같은 숫자가 몇 개 있는지 세고, 정렬된 열거된 숫자마다 합과 숫자를 이어 붙입니다. 처음 다섯 항은 이전 수열과 같습니다.

<pre>0, 10, 1110, 3110, 132110, 13123110, 23124110 ...</pre>

숫자는 큰 것부터 작은 순으로 정렬하세요. 이전 항에 나타나지 않은 숫자의 개수는 포함하지 마세요.

초기값에 따라, 이렇게 생성된 수열은 항상 안정된 값으로 수렴하거나 짧은 주기 패턴으로 수렴합니다. (여기서 수렴한다는 것은 어떤 항이 이전에 본 항과 같아진다는 뜻입니다.) 0을 초기값으로 한 수열은 11회 반복 후 1433223110이라는 안정된 값으로 수렴합니다. 가장 빨리 수렴하는 초기값은 22이며, 첫 번째 항 이후 바로 안정됩니다. (다음 항도 22로 이전에 본 값입니다.)

# --instructions--

초기값을 매개변수로 받아 자기참조 수열을 수렴할 때까지 생성하고, 배열로 반환하는 함수를 작성하세요.

# --hints--

`selfReferential`는 함수여야 합니다.

```js
assert(typeof selfReferential === 'function');
```

`selfReferential(40)`는 배열을 반환해야 합니다.

```js
assert(Array.isArray(selfReferential(40)));
```

`selfReferential(40)`는 `["40", "1410", "142110", "14123110", "1413124110", "2413125110", "151413224110", "152413225110", "251413324110", "152423224110", "152413423110"]`를 반환해야 합니다.

```js
assert.deepEqual(selfReferential(40), [
  '40',
  '1410',
  '142110',
  '14123110',
  '1413124110',
  '2413125110',
  '151413224110',
  '152413225110',
  '251413324110',
  '152423224110',
  '152413423110'
]);
```

`selfReferential(132110)`는 `["132110", "13123110", "23124110", "1413223110", "1423224110", "2413323110", "1433223110"]`를 반환해야 합니다.

```js
assert.deepEqual(selfReferential(132110), [
  '132110',
  '13123110',
  '23124110',
  '1413223110',
  '1423224110',
  '2413323110',
  '1433223110'
]);
```

`selfReferential(132211)`는 `["132211", "132231", "232221", "134211", "14131231", "14231241", "24132231", "14233221"]`를 반환해야 합니다.

```js
assert.deepEqual(selfReferential(132211), [
  '132211',
  '132231',
  '232221',
  '134211',
  '14131231',
  '14231241',
  '24132231',
  '14233221'
]);
```

`selfReferential(1413223110)`는 `["1413223110", "1423224110", "2413323110", "1433223110"]`를 반환해야 합니다.

```js
assert.deepEqual(selfReferential(1413223110), [
  '1413223110',
  '1423224110',
  '2413323110',
  '1433223110'
]);
```

`selfReferential(251413126110)`는 `["251413126110", "16151413225110", "16251413226110", "26151413325110", "16251423225110", "16251413424110", "16153413225110"]`를 반환해야 합니다.

```js
assert.deepEqual(selfReferential(251413126110), [
  '251413126110',
  '16151413225110',
  '16251413226110',
  '26151413325110',
  '16251423225110',
  '16251413424110',
  '16153413225110'
]);
```

# --seed--

## --seed-contents--

```js
function selfReferential(n) {

}
```

# --solutions--

```js
function selfReferential(n) {
  var descending,
    i,
    incr,
    j,
    max_i,
    max_len,
    max_seq,
    seq,
    sequence,
    indexOf =
      [].indexOf ||
      function(item) {
        for (var i = 0, l = this.length; i < l; i++) {
          if (i in this && this[i] === item) return i;
        }
        return -1;
      };

  sequence = function(n) {
    var c, cnt, cnts, d, digit, i, j, l, len, new_cnts, ref, s, seq;
    cnts = {};
    ref = n.toString();
    for (j = 0, len = ref.length; j < len; j++) {
      c = ref[j];
      d = parseInt(c);
      incr(cnts, d);
    }
    seq = [ref];
    while (true) {
      s = '';
      for (i = l = 9; l >= 0; i = --l) {
        if (cnts[i]) {
          s += '' + cnts[i] + i;
        }
      }
      if (indexOf.call(seq, s) >= 0) {
        break;
      }
      seq.push(s);
      new_cnts = {};
      for (digit in cnts) {
        cnt = cnts[digit];
        incr(new_cnts, cnt);
        incr(new_cnts, digit);
      }
      cnts = new_cnts;
    }
    return seq;
  };

  incr = function(h, k) {
    if (h[k] == null) {
      h[k] = 0;
    }
    return (h[k] += 1);
  };

  descending = function(n) {
    var tens;
    if (n < 10) {
      return true;
    }
    tens = n / 10;
    if (n % 10 > tens % 10) {
      return false;
    }
    return descending(tens);
  };

  return sequence(n);
}
```
