---
id: 6720b7f3155a9278d6b6a8b8
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

现在是为你的贺卡设计样式的时候了。

为你的 `body` 元素添加选择器，然后： 

- 将 `font-family` 更改为 `Arial`，后跟通用的 `sans-serif`，
- 在所有边缘添加 `40px` 的内边距，
- 将 `text-align` 属性设置为 `center`。

# --hints--

应该有一个 `body` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

你应该在 `body` 选择器中将 `font-family` 设置为 `"Arial", sans-serif`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

你应该在 `body` 选择器中将 `padding` 设置为 `40px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

你应该在 `body` 选择器中将 `text-align` 设置为 `center`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
