---
id: 6720b3191bbafd269bae0ae8
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

创建一个 `div` 元素，该元素具有 `id` 为 `greeting-card` 和 `class` 为 `card`。

在 `div` 元素内，添加一个带有文本 `Happy Birthday!` 的 `h1`。然后添加一个带有名为 `message` 的 `class` 和文本 `Wishing you all the happiness and joy on your special day!` 的段落元素。

# --hints--

你应该有一个 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

`div` 元素应具有 `id` 为 `greeting-card`。

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

`div` 元素应具有 `card` 的 `class`。

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

你应该在 `div` 元素内有一个 `h1` 元素。

```js
assert.exists(document.querySelector("div > h1"));
```

`h1` 元素应包含文本 `Happy Birthday!`。

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

你应该在 `div` 内有一个 `p` 元素。

```js
assert.exists(document.querySelector("div > p"));
```

`p` 元素应具有 `message` 的 `class`。

```js
assert.exists(document.querySelector("div > p.message"));
```

`p` 元素应包含文本 `Wishing you all the happiness and joy on your special day!`。

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
