---
id: 6720b1f2680ec706ed534acf
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将通过设计贺卡来练习使用不同的伪类和伪元素。已为你提供了 `HTML` 公式化。 

通过链接 `styles.css` 文件开始工作坊。

# --hints--

你的 `link` 元素应该在你的 `head` 元素中。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

你的 `link` 元素应具有一个 `rel` 属性，其值为 `stylesheet`。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

你的 `link` 元素应具有一个 `href` 属性，其值为 `styles.css`。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
