---
id: 6728fda085006772a247e779
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

为 `.card` 添加 `max-width` 为 `400px`，四周的 `padding` 为 `40px`，上下的 `margin` 为 `0`，左右的 `margin` 为 `auto`（使用简写属性）。

# --hints--

`.card` 应该有一个 `max-width` 为 `400px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("max-width"), "400px");
```

`.card` 应该有 `padding` 为 `40px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("padding"), "40px");
```

`.card` 应该有一个 `margin` 为 `0 auto`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("margin"), ["0 auto", "0 auto 0 auto", "0 auto 0", "0px auto", "0px auto 0px auto", "0px auto 0px"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
