---
id: 672b4434dd6d6a1edc405419
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

在 `.card-links` 元素内添加两个 `a` 元素。

第一个应该有一个文本为 `Send Card`、一个 `class` 为 `send-link` 和一个 `href` 为 `#send`。

第二个应该有一个文本为 `Share on Social Media`，一个 `class` 为 `share-link`，以及一个 `href` 为 `#share`。

# --hints--

`.card-links` 内的第一个 `a` 元素的文本应为 `Send Card`。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

第一个 `a` 元素应具有值为 `send-link` 的 `class`。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

第一个 `a` 元素应具有值为 `#send` 的 `href`。

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

`.card-links` 内的第二个 `a` 元素的文本应为 `Share on Social Media`。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

第二个 `a` 元素应具有值为 `share-link` 的 `class`。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

第二个 `a` 元素应具有值为 `#share` 的 `href`。

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
