---
id: 672e44e811e610232fead333
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

为你的 `#send` 添加一个包含文本 `Sending your card...` 的 `h2`，然后添加一个包含文本 `Card successfully sent to your recipient!` 的 `p` 元素。

# --hints--

`#send` 元素应包含一个 `h2` 元素。

```js
assert.exists(document.querySelector("#send > h2"));
```

`h2` 元素应包含文本 `Sending your card...`。

```js
assert.strictEqual(document.querySelector("#send > h2").innerText.trim(), "Sending your card...");
```

`#send` 内的第二个元素应为 `p` 元素。

```js
assert.strictEqual(document.querySelector("#send").children[1].tagName, "P");
```

`p` 元素应包含文本 `Card successfully sent to your recipient!`。

```js
assert.strictEqual(document.querySelector("#send").children[1].innerText.trim(), "Card successfully sent to your recipient!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </section>
    <section id="share"></section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
