---
id: 673332a6b63f9f1e1f81ba9b
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

现在你可以做类似的操作，将表情符号也添加到标题后面。

创建一个选择器，目标为 `h1` 元素的伪元素 `::after`。为它添加一个 `content` 属性，并将其值设置为 `" 🥳"`（注意表情符号前有一个空格）。

# --hints--

你应该有一个 `h1::after` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("h1::after"));
```

`h1::after` 选择器应具有设置为 `" 🥳"` 的 `content` 属性。

```js
assert.strictEqual(
  new __helpers.CSSHelp(document).getStyle("h1::after")?.getPropertyValue("content"),
  `" 🥳"`
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
    </div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

--fcc-editable-region--

--fcc-editable-region--
```
