---
id: 67333082ab35e414eb4ea70d
title: 步骤14
challengeType: 0
dashedName: step-14
---

# --description--

当鼠标悬停在 `a` 元素上时，背景颜色会过渡到另一种颜色。你可以使用 `transition` 属性来调节该过渡的方式：

```css
a {
  transition: color 1s linear;
}
```

`transition` 属性接受的值依次是，应用过渡的属性、过渡的持续时间，然后是时序。

如果有多个属性具有过渡效果，你可以为每个属性写入值，并用逗号分隔：

```css
p {
  transition: property1 0.1s, property2 0.6s linear;
}
```

如果省略了某个值，比如第一个 `property` 的时间，则会应用默认值。

为 `.card` 选择器添加 `transition: transform 0.3s, background-color 0.3s ease`。

试试看，悬停过渡已完成。

# --hints--

`.card` 选择器应具有属性 `transition`，其值设置为 `transform 0.3s, background-color 0.3s ease`。

```js
assert.oneOf(
  new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("transition"),
  [
    "transform 0.3s ease 0s, background-color 0.3s ease 0s", // this is what comes out of `getPropertyValue`
    "transform 0.3s, background-color 0.3s" // different environment got this instead
  ]
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
--fcc-editable-region--
  
--fcc-editable-region--
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

```
