---
id: 673337337f794d3025ded433
title: 第 19 步
challengeType: 0
dashedName: step-19
---

# --description--

定位 `.card-links` 内的 `a` 元素并赋予它们：

- 一个 `text-decoration` 属性设置为 `none`。
- 设置为 `1em` 的 `font-size` 属性
- 一个 `padding` 属性设置为 `10px 20px`
- 设置为 `5px` 的 `border-radius` 属性
- 一个 `color` 属性设置为 `white`
- 设置为 `midnightblue` 的 `background-color` 属性

# --hints--

你应该有一个 `.card-links a` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".card-links a"));
```

`.card-links a` 选择器应具有设置为 `none` 的 `text-decoration` 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("text-decoration"), "none");
```

`.card-links a` 选择器应具有设置为 `1em` 的 `font-size` 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("font-size"), "1em");
```

`.card-links a` 选择器应具有设置为 `10px 20px` 的 `padding` 属性。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("padding"), ["10px 20px", "10px 20px 10px", "10px 20px 10px 20px"]);
```

`.card-links a` 选择器应具有设置为 `5px` 的 `border-radius` 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("border-radius"), "5px");
```

`.card-links a` 选择器应具有设置为 `white` 的 `color` 属性。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("color").toLowerCase(), ["white", "#ffffff", "#fff", "rgb(255, 255, 255)", "hsl(0, 0%, 100%)"]);
```

`.card-links a` 选择器应设置 `background-color` 属性为 `midnightblue`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card-links a")?.getPropertyValue("background-color").toLowerCase(), ["midnightblue", "#191970", "rgb(25, 25, 112)", "hsl(240, 64%, 27%)"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
    </div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

.message {
  font-size: 1.2em;
  margin-bottom: 20px;
}

.card-links {
  margin-top: 20px;
  display: flex;
  justify-content: space-around;
}

--fcc-editable-region--

--fcc-editable-region--
```
