---
id: 6733352ce2a5642827b4f7f0
title: 第 18 步
challengeType: 0
dashedName: step-18
---

# --description--

添加一个 `.card-links` 选择器，并将 `margin-top` 属性设置为 `20px`。

你可以为元素添加 `display: flex` 来设置使用 flexbox，稍后你将在课程中详细学习 flexbox，你可以将此视为一个小预览。

要使两个链接之间的间距相等，为它们添加一个设置为 `flex` 的 `display` 属性，以及一个设置为 `space-around` 的 `justify-content` 属性。

# --hints--

你应该有一个 `.card-links` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".card-links"));
```

`.card-links` 选择器应设置 `margin-top` 属性为 `20px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("margin-top"), "20px");
```

`.card-links` 选择器应具有设置为 `flex` 的 `display` 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("display"), "flex");
```

`.card-links` 选择器应具有设置为 `space-around` 的 `justify-content` 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("justify-content"), "space-around");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
    </div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

.message {
  font-size: 1.2em;
  margin-bottom: 20px;
}

--fcc-editable-region--

--fcc-editable-region--
```
