---
id: 694c84330cbe604821b9deef
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

任务控制需要一种方法来为名单添加新成员。创建一个名为 `addCrewMember` 的空函数，该函数接受两个参数，`crew` 和 `astronaut`。

# --hints--

你应该创建一个名为 `addCrewMember` 的函数。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addCrewMember);
```

你的 `addCrewMember` 函数应具有参数 `crew` 和 `astronaut`。

```js
const explorer = await __helpers.Explorer(code);
const { addCrewMember } = explorer.allFunctions;
assert.lengthOf(addCrewMember?.parameters, 2);
assert.equal(addCrewMember?.parameters?.[0]?.toString(), "crew");
assert.equal(addCrewMember?.parameters?.[1]?.toString(), "astronaut");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
