---
id: 694ca14c8ef00056ad146e51
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

将 `astronaut` 对象添加到 `crew` 数组中，并将 `Added ${astronaut.name} as ${astronaut.role}` 记录到控制台。你可以使用模板字面值或字串串接来记录消息。

# --hints--

你应该将 `astronaut` 推入 `crew` 数组。

```js
const crew = [];
const astronauts = [
  { id: 1, name: "Abraham", role: "Commander", isEVAEligible: true, priority: 1 },
  { id: 2, name: "Boris", role: "Pilot", isEVAEligible: false, priority: 2 },
  { id: 3, name: "Catarina", role: "Engineer", isEVAEligible: true, priority: 3 }
];

astronauts.forEach(astro => addCrewMember(crew, astro));

assert.strictEqual(crew.length, astronauts.length);
astronauts.forEach((astro, idx) => {
  assert.deepEqual(crew[idx], astro);
});
```

你应该以以下形式记录宇航员的 `name`，`Added {astronaut.name} as {astronaut.role}`。

```js
function captureConsoleLogs(fn) {
  const originalLog = console.log;
  const logs = [];

  console.log = function (...args) {
    logs.push(args.join(" "));
  };

  try {
    fn();
  } finally {
    console.log = originalLog;
  }

  return logs;
}

const crewForLogging = [];
const astronautsToLog = [
  { id: 1, name: "Abraham", role: "Commander" },
  { id: 2, name: "Boris", role: "Pilot" },
  { id: 3, name: "Catarina", role: "Engineer" }
];

const logs = captureConsoleLogs(() => {
  astronautsToLog.forEach(astro => addCrewMember(crewForLogging, astro));
});

assert.strictEqual(logs.length, astronautsToLog.length);

astronautsToLog.forEach((astro, idx) => {
  assert.match(
    logs[idx],
    new RegExp(`Added\\s*${astro.name}\\s*as\\s*${astro.role}`, 'i'),
  );
});
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
