---
id: 694c84330cbe604821b9deef
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

任务控制需要一种方法来为名单添加新成员。创建一个名为 `addCrewMember` 的空函数，该函数接受两个参数，`crew` 和 `astronaut`。

# --hints--

你应该创建一个名为 `addCrewMember` 的函数。

```js
assert.isFunction(addCrewMember);
```

你的 `addCrewMember` 函数应具有参数 `crew` 和 `astronaut`。

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
