---
id: 694c84106e9023f289c03b84
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

你的名单从一个指挥官宇航员开始。创建一个名为 `firstAstronaut` 的对象，包含以下属性：

| Key | Value |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

你应该创建一个 `firstAstronaut` 变量。

```js
assert.isDefined(firstAstronaut);
```

你的 `firstAstronaut` 应该是一个对象。

```js
assert.isObject(firstAstronaut);
```

你的 `firstAstronaut` 对象应包含正确的属性值。

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
