---
id: 694c40d86fe6a78d20b6559c
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会中，你将杠杆 JavaScript 来构建一个太空任务的宇航员名单。 

首先，为你的花名册创建一个名为 `squad` 的空数组。

# --hints--

你应该创建一个 `squad` 变量。

```js
assert.isDefined(squad);
```

你的 `squad` 变量应该是一个数组。

```js
assert.isArray(squad);
```

你的 `squad` 数组应为空。

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
