---
id: 694c9c186c089f7bdff0f3c6
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

通常，在函数内部验证你的输入是一个良好的实践。使用 `for` 循环迭代 `crew` 数组，查看是否有任何成员已经拥有与输入的 `astronaut` 相同的 `id`。如果发现副本，记录错误信息 `console.log("Duplicate ID: " + astronaut.id)`，然后调用 `return` 以提前退出函数。 

这是这种常用技术的一个示例：

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

你的函数应使用 `for` 循环扫描现有的 `crew` 成员（任何有效的 `for (...)` 形式都可以）。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

在 `for` 循环内，你应该查看是否有重复的 `id`。如果发现，向控制台记录 `"Duplicate ID: " + astronaut.id"`，然后调用 `return` 退出函数。

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

当给定有效的 `input` 时，你的 `addCrewMember` `function` 不应调用 `console.log()` 误差信息或提前 `return` 退出。

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
