---
id: 694d8ddd106340237efb0c80
title: 步骤 12
challengeType: 1
dashedName: step-12
---

# --description--

`splice()` 方法可以通过在任意位置（包括中间）添加或删除元素来修改数组。由于 `splice()` 返回一个包含被删除元素的数组，你可以使用它在不改变原数组的情况下交换数组中的两个元素。可以使用以下技巧在一行内完成此操作：

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

这是一个示例：

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

该技术应用于第三行，具体如下：

- `splice(3, 1, copyArray[1])` 删除索引为 `3` 的元素（`55`）

- 它将来自 `copyArray[1]`（`97`）的元素插入到索引 `3`。

- `splice()` 返回 `[55]`，一个包含被移除元素的数组

- `[0]` 从该数组中提取 `55`

- 该值（`55`）被赋值回 `copyArray[1]`，完成交换

使用此技术对 `updatedCrew` 数组中位于 `fromIndex` 和 `toIndex` 的宇航员进行交换。

# --hints--

你应该通过调用 `splice()` 方法来执行交换。

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

你应该使用 `updatedCrew` 数组调用 `splice()` 来执行交换。你应该使用单行技术。

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

你的 `swapCrewMembers` 函数不应修改原始输入的 `crew` 数组。

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
