---
id: 698c5229e47d57c6dfaa91bd
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

`controls` 屬性為 `audio` 元素提供播放控制元件，包括暫停／繼續播放、尋找及音量控制。 

`controls` 屬性是一個布林值屬性，且不需要值。

為你的 `audio` 元素添加 `controls` 屬性。

現在你應該會看到頁面上顯示 `audio` 元素。

# --hints--

你的 `audio` 元素應該有 `controls` 屬性。

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
