---
id: 698c58745b73289bb6c89ece
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

最後一首歌曲標題為 `Scratching the Surface`，藝術家是 Quincy Larson，並且要使用的檔案是 `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`。

為你的音樂播放器添加最後一首歌曲以完成，並遵循你對前兩首歌曲使用的相同樣式。

# --hints--

你應該有第三個 `h2` 元素。

```js
assert.exists(document.querySelector('h2:nth-of-type(3)'));
```

你的第三個 `h2` 元素應該包含文字 `Scratching the Surface`。

```js
const h2 = document.querySelector('h2:nth-of-type(3)');
assert.strictEqual(h2?.textContent.trim(), 'Scratching the Surface');
```

你應該有一個 `p` 元素。

```js
assert.exists(document.querySelector('p:nth-of-type(3)'));
```

你的 `p` 元素應該包含文字 `Artist: Quincy Larson`。

```js
const p = document.querySelector('p:nth-of-type(3)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

你應該有第三個 `audio` 元素。

```js
assert.exists(document.querySelector('audio:nth-of-type(3)'));
```

新的 `audio` 元素應該位於 `p` 元素下方。

```js
assert.exists(document.querySelector('p:nth-of-type(3) + audio:nth-of-type(3)'));
```

你的 `audio` 元素應該有一個 `src` 屬性。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('src'));
```

你的 `audio` 元素應該有一個 `src` 屬性，值為 `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

你的 `audio` 元素應該有 `loop` 屬性。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('loop'));
```

你的 `audio` 元素應該有 `controls` 屬性。

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

  <h2>Scratching the Surface</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3" loop controls></audio>
  
</body>
</html>
```
