---
id: 698c5710136539797f296098
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

在 `p` 元素下方，添加一個 `audio` 元素，為它設定一個 `src` 屬性，值為 `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`，並加上 `loop` 和 `controls` 屬性。

# --hints--

你應該有第二個 `audio` 元素。

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

新的 `audio` 元素應該位於 `p` 元素下方。

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

你的 `audio` 元素應該有一個 `src` 屬性。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

你的 `audio` 元素應該有一個 `src` 屬性，值為 `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

你的 `audio` 元素應該有 `loop` 屬性。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

你的 `audio` 元素應該有 `controls` 屬性。

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
