---
id: 698c50535b54b2d75fa985fb
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

`loop` 屬性會在播放完成後重新啟動音訊。想像一個會重複播放的網際網路迷因。省略 `loop` 屬性會讓音訊只播放一次。

`loop` 屬性是一個布林值屬性，且不需要值。

為你的 `audio` 元素添加 `loop` 屬性。

# --hints--

你的 `audio` 元素應該有 `loop` 屬性。

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
