---
id: 698c4fe24f1f1a17be74a23e
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

接著，在 `p` 元素下方創建一個 `audio` 元素。在接下來的幾個步驟中，你將為該音訊元素添加必要的屬性，以使其播放音樂。 

# --hints--

你應該有一個 `audio` 元素。

```js
assert.exists(document.querySelector('audio'));
```

新的 `audio` 元素應該位於 `p` 元素下方。

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
