---
id: 6894c1b715b8d18a18ef3203
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個樂器清單來練習 Python 的類別和物件。

要在 Python 中定義一個類別，請使用 `class` 關鍵字，後接類別名稱和冒號。

```python
class ClassName:
    # class contents go here
```

請記住，在 Python 中，類別名稱通常使用 PascalCase，這表示每個字、單字都以大寫字母開頭，且字與字之間沒有空格。

首先建立一個名為 `MusicalInstrument` 的類別。暫時將 `pass` 作為該類別的主體。

# --hints--

你應該創建一個名為 `MusicalInstrument` 的類別。

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
