---
id: 6895e5337b1914050b0af7bd
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

現在，對 `instrument_1` 實體呼叫 `get_fact()` 方法，並列印結果。

由於 `get_fact()` 傳回的是字串而非直接列印，因此你需要使用 `print()` 來顯示傳回的字串。

# --hints--

你應該在 `instrument_1` 上呼叫 `get_fact()` 並列印結果。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()

--fcc-editable-region--
```
