---
id: 6895e4940e39950288e406c2
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

接著，為你的類別添加另一個名為 `get_fact` 的方法，該方法傳回有關該樂器的事實。

這個方法應該只有 `self` 參數，並傳回一個 f-string，內容為 `The [instrument name] is part of the [instrument type] family of instruments.`。

與直接列印訊息的 `play` 方法不同，這個方法應該傳回一個可以在你的程式碼其他地方使用的字串。

# --hints--

你應該定義一個名為 `get_fact` 且帶有 `self` 參數的方法。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

你的 `get_fact` 方法應該傳回一個包含樂器名稱和型別的 f-string。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
