---
id: 6895e0452b3370f26f365a46
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

在 `__init__` 方法中，你需要將參數指定到實體屬性。實體屬性是屬於類別特定實體的變數。

在你的 `__init__` 方法中，將 `name` 參數指定到 `self.name`，並將 `instrument_type` 參數指定到 `self.instrument_type`。

# --hints--

你應該在 `__init__` 方法中將 `name` 參數指定給 `self.name`。

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

你應該在 `__init__` 方法中將 `instrument_type` 參數指定到 `self.instrument_type`。

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        pass
--fcc-editable-region--
```
