---
id: 6895e2becb7890fad818ad4d
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

既然你已經了解如何存取物件屬性，請移除現有的 `print` 呼叫。

# --hints--

你應該移除 `print(instrument_1.name)` 述語。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

你應該移除 `print(instrument_1.instrument_type)` 述語。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

你應該移除 `print(instrument_2.name)` 述語。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

你應該移除 `print(instrument_2.instrument_type)` 述語。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
print(instrument_1.name)
print(instrument_1.instrument_type)

print(instrument_2.name)
print(instrument_2.instrument_type)

--fcc-editable-region--
```
