---
id: 6895e4e023e07603bcef46f9
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

現在讓我們使用你已經創建的這些方法。 

對你的 `instrument_1` 實體呼叫 `play()` 方法。這將列印有關演奏雙簧管的訊息。

要在實體上呼叫方法，請使用點記法：

```python
instance_name.method_name()
```

# --hints--

你應該在 `instrument_1` 實體上呼叫 `play()` 方法。

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_1.play()")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
