---
id: 6895e235935106f88281b648
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

你可以使用點記法存取你的實體的屬性：

```py
instance_name.attribute_name
```

這讓你能夠擷取在初始化期間儲存的值。

嘗試列印你兩個實體的 `name` 和 `instrument_type` 屬性。你將能看到在創建它們時所指定的值。

# --hints--

你應該列印 `instrument_1` 的 `name` 屬性。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

你應該列印 `instrument_1` 的 `instrument_type` 屬性。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

你應該列印 `instrument_2` 的 `name` 屬性。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

你應該列印 `instrument_2` 的 `instrument_type` 屬性。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
