---
id: 6895e3f91c069e0068a42582
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

在你的 `play` 方法中，使用 `print()` 來顯示有關演奏樂器的訊息。

使用 f-string 在訊息中包含樂器名稱：`The [instrument name] is fun to play!`。`[instrument name]` 應該被取代為該樂器的實際名稱。

# --hints--

你的 `play` 方法應該使用 f-string 來顯示訊息 `The [instrument name] is fun to play!`。使用 `self.name` 來存取樂器名稱。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
