---
id: 6895e197180c9af6988e1e19
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

為一個屬於 `brass` 樂器的 `Trumpet` 創建第二個名為 `instrument_2` 的實體。

注意你如何能從相同的類別模板創建多個實體。每個實體獨立儲存自己的屬性值，彼此互不影響。

# --hints--

你應該生成一個名為 `instrument_2` 的 `MusicalInstrument` 類別實體。

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

你應該在建立 `instrument_2` 時傳入 `Trumpet` 和 `brass` 作為引數。

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')

--fcc-editable-region--
```
