---
id: 6895e0bdb0348bf46952ab67
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

現在是時候創建你的 `MusicalInstrument` 類別的實體。

要創建一個類別的實體，你像呼叫函式一樣呼叫類別名稱，並傳遞 `__init__` 方法所需的引數（除了 `self`，它由 Python 自動控制代碼）。

```python
my_instance = ClassName(arg1, arg2)
```

當你創建這個實體時，Python 會呼叫你的 `__init__` 方法，並將這些值儲存為只屬於此特定實體的實體屬性。

在類別定義式外，為一個屬於 `woodwind` 樂器的 `Oboe` 創建一個名為 `instrument_1` 的實體。

# --hints--

你應該生成一個名為 `instrument_1` 的 `MusicalInstrument` 類別實體。

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

你應該在建立 `instrument_1` 時傳遞 `Oboe` 和 `woodwind` 作為引數。

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
```
