---
id: 6895dffc64e97df0e62e9ef9
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

方法是在類別內定義的函式（程式），可以使用該類別的資料來執行動作。以下是你如何創建方法：

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

在 Python 中，`__init__` 方法是一個特殊的方法，稱為初始化方法，當你創建類別的新實體時會被呼叫。

在你的 `MusicalInstrument` 類別中，為你的 `__init__` 方法添加三個參數：`self`、`name` 和 `instrument_type`。

`self` 參數是取用類別實體的標準方式，並且是類別中所有實體方法的第一個必要參數。

# --hints--

你應該在你的 `MusicalInstrument` 類別中定義一個 `__init__` 方法。

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

你的 `__init__` 方法應該有三個參數：`self`、`name` 和 `instrument_type`。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    pass
--fcc-editable-region--
```
