---
id: 6895e356fd877efd6e3657ee
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--


在 `MusicalInstrument` 類別內，定義另一個名為 `play` 的方法，且僅有 `self` 參數。`self` 參數允許方法存取特定實體的屬性。

# --hints--

你應該定義一個名為 `play` 且帶有 `self` 參數的方法。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
