---
id: 6718d2d59337c822ecb697ff
title: Construir um Programa de Gerenciamento de Conta Bancária
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório. 

**Histórias de Usuário:**

1. Você deve definir uma classe chamada `BankAccount` com um construtor. O construtor deve definir o `balance` padrão como `0` e inicializar um array vazio chamado `transactions` para armazenar registros de transações como objetos.

2. Cada transação armazenada no array `transactions` deve ser um objeto com duas propriedades: `type` e `amount`. A propriedade `type` deve ser `deposit` ou `withdraw`, e a propriedade `amount` deve ser o valor depositado ou retirado.

3. You should define a method named `deposit` that takes the deposit amount as a parameter. When the deposit amount is greater than `0`, it should:

   - Adicione um novo objeto ao array transactions com um tipo `deposit` e o valor depositado.
   - Atualize o saldo.
   - Retorna `"Successfully deposited $[amount]. New balance: $[balance]"`.

4. Se o valor for menor ou igual a `0`, o método `deposit` deve retornar `"O valor do depósito deve ser maior que zero."`.

5. You should define a method named `withdraw` that takes an amount as a parameter. This method should update the current balance according to withdrawals. When the amount to be withdrawn is greater than `0` and less than or equal to the current balance, it should:

   - Adicione um novo objeto ao array transactions com um tipo `withdraw` e o valor retirado.
   - Atualize o saldo.
   - Retorna `"Successfully withdrew $[amount]. New balance: $[balance]"`.

6. Se o valor a ser sacado for menor ou igual a `0` ou maior que o saldo atual, o método `withdraw` deve retornar `"Saldo insuficiente ou valor inválido."`.


7. Você deve definir um método chamado `checkBalance` que retorne o saldo atual no formato `"Saldo atual: R$[balance]"`.

8. Você deve definir um método chamado `listAllDeposits` que itere pelo array de transações e retorne todos os depósitos no formato `"Deposits: amount,amount,..."`.

9. Você deve definir um método chamado `listAllWithdrawals` que itere pelo array de transações e retorne todos os saques no formato `"Withdrawals: amount,amount,..."`.

10. Você deve criar uma nova instância de `BankAccount` chamada `myAccount`.

11. Sua conta bancária `myAccount` deve ter pelo menos cinco transações.

12. Sua conta bancária `myAccount` deve ter pelo menos dois depósitos.

13. Sua conta bancária `myAccount` deve ter pelo menos dois saques.

14. Sua conta bancária `myAccount` deve ter um saldo maior que `$100`.

# --hints--

Você deve definir uma classe chamada `BankAccount`.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

O objeto `BankAccount` deve inicialmente ter um `balance` de `0` e um array vazio `transactions` para armazenar os registros de transações.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

Você deve ter um método `deposit` que recebe o valor do depósito como parâmetro.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

Você deve ter um método `withdraw` que receba o valor do saque como parâmetro.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

Você deve ter um método `checkBalance` que verifica o saldo atual.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

Você deve ter um método `listAllDeposits` que liste todos os depósitos.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

Você deve ter um método `listAllWithdrawals` que liste todas as retiradas.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

"`BankAccount.deposit(100)` deve retornar `"Depósito de $100 realizado com sucesso."` Novo saldo: $100`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` deve retornar `"O valor do depósito deve ser maior que zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` deve retornar `"O valor do depósito deve ser maior que zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

Quando o saldo da conta é `100`, `BankAccount.withdraw(150)` deve retornar `"Saldo insuficiente ou valor inválido."`.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` deve retornar `"Saldo insuficiente ou valor inválido."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` deve retornar `"Saldo insuficiente ou valor inválido."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

Quando o saldo da conta é `200`, `BankAccount.withdraw(150)` deve retornar `"Retirou com sucesso $150."` Novo saldo: $50`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

Quando o saldo da conta é `200`, `BankAccount.checkBalance()` deve retornar `"Saldo atual: R$200"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

Quando você depositar `10`, `35`, `90`, o método `listAllDeposits` deve retornar `"Depósitos: 10,35,90"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

Quando você sacar `20`, `50`, `100`, o método `listAllWithdrawals` deve retornar `"Withdrawals: 20,50,100"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

Você deve ter uma instância de `BankAccount` chamada `myAccount`.

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

Sua conta bancária `myAccount` deve ter pelo menos cinco transações.

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

Sua conta bancária `myAccount` deve ter pelo menos dois depósitos.

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

Sua conta bancária `myAccount` deve ter pelo menos dois saques.

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```


Sua conta bancária `myAccount` deve ter um saldo maior que `$100`.

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
