---
id: 5a9d7286424fe3d0e10cad13
title: 给 CSS 变量设置备用值
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bDNfp'
forumTopicId: 301084
dashedName: attach-a-fallback-value-to-a-css-variable
---

# --description--

When using your variable as a CSS property value, you can attach a fallback value that your browser will revert to if the given variable is invalid.

**注意：** 备用值不是用于增强浏览器的兼容性，它也不适用于 IE 浏览器。 相反，它是用来让浏览器在找不到你的变量时可以显示一种颜色。

下面是操作方式：

```css
background: var(--penguin-skin, black);
```

如果你的变量没有设置，这将会把背景设置为 `black`。 提示：这对调试代码也会很有帮助。

# --instructions--

`.penguin-top` 和 `.penguin-bottom` 类的变量看起来似乎有点问题。 请为 class 为 `penguin-top` 和 `penguin-bottom` 的元素的 `background` 属性设置一个 `black` 的备用色。

# --hints--

class 为 `penguin-top` 的元素的 `background` 属性值应有 `black` 作为备用颜色。

```js
assert.match(__helpers.removeCssComments(code), /.penguin-top\s*?{[\s\S]*background\s*?:\s*?var\(\s*?--pengiun-skin\s*?,\s*?black\s*?\)\s*?;[\s\S]*}[\s\S]*.penguin-bottom\s{/gi);
```

class 为 `penguin-bottom` 的元素的 `background` 属性值应有 `black` 作为备用颜色。

```js
assert.match(__helpers.removeCssComments(code), /.penguin-bottom\s*?{[\s\S]*background\s*?:\s*?var\(\s*?--pengiun-skin\s*?,\s*?black\s*?\)\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
请为 class 为 <code>penguin-top</code> 和 <code>penguin-bottom</code> 的元素的 <code>background</code> 属性设置一个 <code>black</code> 的备用色。
```
 和 penguin-bottom 的元素的 background 属性设置一个 black 的备用色。
</code>

# --solutions--

```html
<style>
  .penguin-top {
    background: var(--pengiun-skin, black);
  }
  .penguin-bottom {
    background: var(--pengiun-skin, black);
  }
</style>
```
