---
id: bad87fee1348bd9aedf08726
title: 使用十六进制编码获得指定颜色
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

你知道在 CSS 里面还有其他方式来代表颜色吗？ 其中一个方法叫十六进制编码，简称 hex。

我们通常使用<dfn>十进制</dfn>，也就是基数为10的数字，每位使用符号0到9。<dfn>十六进制</dfn>（或称<dfn>hex</dfn>）是基数为16的数字。这意味着它使用十六个不同的符号。和十进制一样，符号0-9表示数值零到九。然后A、B、C、D、E、F表示数值十到十五。总的来说，0到F可以表现十六进制中的一个数字，共有16种可能的数值。你可以在<a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">这里</a>找到更多关于十六进制数字的信息。

在 CSS 中，我们可以用6个十六进制数字来表现颜色，每两个数字分别代表红色（R）、绿色（G）和蓝色（B）组件。例如，`#000000` 是黑色，同时也是最低的可能数值。你可以在<a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">这里</a>找到更多关于 RGB 颜色系统的信息。

```css
body {
  color: #000000;
}
```

# --instructions--

将 `body` 元素的背景颜色由 `black` 改为它对应的十六进制编码 `#000000`。

# --hints--

`body` 元素的 `background-color` 应为黑色。

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

应该使用黑色的十六进制代码，而不是 `black` 这个词。

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
