---
id: 68c3144fc395d681180cde15
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

يبدو أن هناك بعض المشاكل في استخدام العناوين على الصفحة. إذا نظرت إلى أول عنصر `div`، ستجد عنصرين `h4` مستخدمين كعناوين فرعية. من الأفضل استخدام عناصر `h3` لهذه العناوين الفرعية لأنها تقع تحت مستوى عنصر `h2`.

غيّر العنصرين `h4` إلى عناصر `h3`.

# --hints--

يجب أن يكون عنصر `h4` الذي يحتوي على النص `Early Challenges` عنصر `h3`.

```js
const post1Subheading = document.querySelector("#post1 + p + h3");
assert.equal(post1Subheading?.tagName, "H3");
assert.equal(post1Subheading?.textContent, "Early Challenges");
```

يجب أن يكون عنصر `h4` الذي يحتوي على النص `Breakthroughs` عنصر `h3`.

```js
const post1SecondSubheading = document.querySelector("#post1 + p + h3 + p + h3");
assert.equal(post1SecondSubheading?.tagName, "H3");
assert.equal(post1SecondSubheading?.textContent, "Breakthroughs");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>

  <h2>Navigation</h2>
  <ul>
    <li><a href="#post1">My Journey</a></li>
    <li><a href="#post2">Accessibility</a></li>
    <li><a href="#post3">Next Steps</a></li>
  </ul>
  <div>
    <h2 id="post1">My Journey Learning to Code</h2>
    <p>I started learning to code a few months ago and it's been a wild ride!</p>
    
    --fcc-editable-region--
    <h4>Early Challenges</h4>
    <p>At first, syntax was really confusing.</p>
    
    <h4>Breakthroughs</h4>
    --fcc-editable-region--
    <p>Eventually things started to click.</p>
  </div>

  <div>
    <h2 id="post2">Accessibility Matters</h2>
    <p>Today I learned that not everyone uses the web the same way I do.</p>

    <h5>Screen Readers</h5>
    <p>These tools help visually impaired users browse websites.</p>
  </div>

  <div>
    <h2 id="post3">What's Next?</h2>
    <p>I'm excited to dive into JavaScript and build interactive features!</p>

    <h3>Coming soon: My first JavaScript project!</h3>
    <p>Stay tuned for some exciting interactive blog features.</p>
  </div>
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
</body>
</html>
```
