---
id: 6716249b5405164036fd0b0d
title: ソーティングビジュアライザーを作成する
challengeType: 25
dashedName: build-a-sorting-visualizer
demoType: onClick
---

# --description--

バブルソートアルゴリズムは、シーケンスの先頭から隣接する要素のペアを比較して整数のシーケンスをソートします。最初の要素が2番目の要素より大きい場合、それらを入れ替えます。その後、次のペアに進みます。シーケンスの最後の要素に到達すると、シーケンスの先頭から新しいサイクルを開始し、要素がソートされるまでこのプロセスを繰り返します。入れ替えが一度も発生しないサイクルが完了するとアルゴリズムは停止します。

このラボでは、すべてのHTMLとCSSが提供されています。JavaScriptを使ってバブルソートビジュアライザーを完成させ、5つの整数の配列をソートするためにバブルソートアルゴリズムが必要とする各ステップを視覚化してください。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `generateElement` という名前の関数があり、`1` から `100` までの範囲でランダムな整数を返す必要があります。
1. `generateArray` という名前の関数があり、`generateElement` 関数を使って5つのランダムな整数を含む配列を返す必要があります。
1. `generateContainer` という名前の関数があり、空の `div` 要素を作成して返す必要があります。
1. `fillArrContainer` という名前の関数があり、最初の引数にHTML要素、2番目の引数に配列を取る必要があります。
1. `fillArrContainer` は最初のパラメータに要素、2番目のパラメータに整数の配列を取り、要素に5つの `span` 要素を追加し、それぞれの `span` が配列の整数の1つを表示する必要があります。
1. `isOrdered` という名前の関数があり、2つの整数を取り、最初の整数が2番目の整数以下かどうかを示すブール値を返す必要があります。
1. `swapElements` という名前の関数があり、整数の配列と数値のインデックスを取ります。
1. `swapElements` 関数は、`isOrdered` が `false` を返す場合に、渡されたインデックスの要素とその次の要素を入れ替えて配列をその場で変更する必要があります。
1. `highlightCurrentEls` という名前の関数があり、HTML要素と数値のインデックスを取ります。
1. `highlightCurrentEls` 関数は、指定された要素の子要素のうち指定されたインデックスとその直後の子要素に対して、`border` の `dashed` スタイル、`red` の色、および任意の幅を設定する必要があります。
1. `#generate-btn` をクリックすると、`fillArrContainer` 関数を使って `#starting-array` に5つの `span` 要素を生成し、それぞれにランダムな数字をテキストとして設定します。もし存在すれば、他の要素は `#array-container` から削除されます。
1. `#sort-btn` をクリックすると、バブルソートアルゴリズムが開始配列をソートするために必要なステップごとに `#array-container` に `div` 要素が含まれ、開始配列を表す `div` とソート済み配列を表す `div` も含まれます。これまでに作成した関数がここで役立ちます。
1. 各 `div` は、現在のソート状態を表す5つの `span` 要素を含む必要があります。
1. `#sort-btn` をクリックすると、`#starting-array` は初期配列と最初の2つの整数がハイライトされた開始ステップを表す必要があります。
1. 各ソートステップで、比較されている2つの数字をハイライトするために `highlightCurrentEls` を使い、次のステップで `swapElements` を使ってそれらを入れ替える必要があります。

# --hints--

`generateElement` という名前の関数を用意してください。

```js
assert.isFunction(generateElement);
```

`generateElement` 関数は `1` から `100` までの範囲でランダムな整数を返す必要があります。

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    assert.strictEqual(generateElement(), 26);
    assert.strictEqual(generateElement(), 9);
} finally {
    randomMocker.restore();
}
```

`generateArray` という名前の関数を用意してください。

```js
assert.isFunction(generateArray)
```

`generateArray` 関数は `generateElement` 関数を利用してください。

```js
let flag = false;
const temp = generateElement;
generateElement = () => flag = true;
try {
    generateArray();
    assert.isTrue(flag);
} finally {
    generateElement = temp;
}
```

`generateArray` 関数は `1` から `100` までの範囲で5つのランダムな整数を含む配列を返す必要があります。

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    assert.deepEqual(generateArray(), [26, 9, 58, 23, 38])
} finally {
    randomMocker.restore();
}
```

`generateContainer` という名前の関数を用意してください。

```js
assert.isFunction(generateContainer);
```

`generateContainer` 関数は空の `div` 要素を返す必要があります。

```js
const div = generateContainer();
assert.equal(div.tagName, "DIV");
assert.isEmpty(div.children);
```

`fillArrContainer` という名前の関数を用意してください。

```js
assert.isFunction(fillArrContainer);
```

`fillArrContainer()` は最初のパラメータに要素、2番目のパラメータに整数の配列を取り、要素に5つの `span` 要素を追加し、それぞれの `span` が配列の整数の1つを表示する必要があります。

```js
const testDiv = document.createElement("div");
const testArr = [15, 98, 17, 5, 63]
fillArrContainer(testDiv, testArr);
const children = testDiv.children;
assert.lengthOf(children, 5);
Array.from(children).forEach((el, i) => {
    assert.equal(el.tagName, "SPAN");
    assert.equal(el.innerText.trim(), testArr[i])
})
```

`isOrdered` という名前の関数を用意してください。

```js
assert.isFunction(isOrdered);
```

`isOrdered` 関数は2つの整数を取り、最初の整数が2番目の整数以下かどうかを示すブール値を返す必要があります。

```js
assert.lengthOf(isOrdered, 2);
assert.isTrue(isOrdered(2, 60));
assert.isFalse(isOrdered(10, 3));
assert.isTrue(isOrdered(5, 5));
```

`swapElements` という名前の関数を用意してください。

```js
assert.isFunction(swapElements);
```

`swapElements` 関数は整数の配列と数値のインデックスを引数に取り、最初の整数が2番目の整数より大きい場合に渡された配列の指定インデックスの要素と次の要素を入れ替えて配列をその場で変更する必要があります。

```js
const testArr = [22, 4, 87, 47, 33];
swapElements(testArr, 0);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 1);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 2);
assert.deepEqual(testArr, [4, 22, 47, 87, 33]);
swapElements(testArr, 3);
assert.deepEqual(testArr, [4, 22, 47, 33, 87]);
```

`highlightCurrentEls` という名前の関数を用意してください。

```js
assert.isFunction(highlightCurrentEls);
```

`highlightCurrentEls` 関数は指定された要素の子孫要素のうち、指定されたインデックスとその次のインデックスにある要素に対して、`dashed`、`red` の `border` を幅は任意で設定する必要があります。

```js
const testDiv = document.createElement("div");
document.querySelector("body").appendChild(testDiv)
for (let i = 0; i < 5; i++) {
    testDiv.appendChild(document.createElement("span"));
}
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|#F00|hsl\(0,\s*100%,\s*50%\))/;
const revertBorder = () => {
    for (const el of children) {
        el.style.border = "revert";
    }
}
const children = testDiv.children;

for (let i = 0; i < 3; i++) {
    highlightCurrentEls(testDiv, i);
    for (let j = 0; j < 5; j++) {
        let b = getComputedStyle(children[j]).border;
        if (j == i || j == i + 1) {            
            assert.match(b, redBorderRegex);
            assert.isAbove(parseFloat(b), 0);
        } else {
            assert.notMatch(b, redBorderRegex);
        }
    }
    revertBorder();
}
testDiv.remove();
```

`#generate-btn` をクリックすると、`#starting-array` に `span` 要素を5つ生成し、それぞれに `1` から `100` までの範囲のランダムな数字をテキストとして設定してください。

```js
const genBtn = document.querySelector("#generate-btn");
genBtn.click();
const children = document.querySelector("#starting-array").querySelectorAll("span");
assert.lengthOf(children, 5);
Array.from(children).forEach(el => {
    assert.equal(el.tagName, "SPAN");
    const num = Number(el.innerText.trim());
    assert.isAtMost(num, 100);
    assert.isAtLeast(num, 1);
})
```

`#starting-array` にすでに生成された配列があるか、`#array-container` にソート済み配列がある場合、`#generate-btn` をクリックすると `#array-container` の他の要素が削除され、新しく生成された数字だけが `#starting-array` に残ります。

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
genBtn.dispatchEvent(new Event("click"));

const prevNumbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));

sortBtn.dispatchEvent(new Event("click"));
genBtn.dispatchEvent(new Event("click"));

const container = document.querySelector("#array-container");
assert.lengthOf(container.children, 1);
const numbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));
assert.lengthOf(numbers, 5);
assert.isTrue(prevNumbers.some((num, index) => num !== numbers[index]))
```

`#sort-btn` をクリックすると、`#array-container` にバブルソートアルゴリズムが開始配列をソートするために必要なステップ数と同じ数の `div` 要素が含まれ、開始配列を表す `div` とソート済み配列を表す `div` も含まれます。

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
// using randomMocker to be sure that the starting array requires 13 steps to be sorted
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.lengthOf(container.children, 13)
    Array.from(container.children).forEach(el => {assert.equal(el.tagName, "DIV")})
} finally {
    randomMocker.restore();
}
```

`#sort-btn` をクリックすると、`#array-container` 内の各 `div` は5つの `span` 要素を含み、それぞれに数字がテキストとして表示され、バブルソートアルゴリズムが開始配列をソートするために必要なステップを表します。

```js
const finalArr = [9, 23, 26, 38, 58];
const arrays = [
    [26, 9, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 23, 58, 38],
    [9, 26, 23, 38, 58],
    [9, 26, 23, 38, 58],
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr
]
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.isNotEmpty(container.children);
    Array.from(container.children).forEach((el, i) => {
        Array.from(el.children).forEach((j, k) => {
            assert.strictEqual(Number(j.innerText.trim()), arrays[i][k])
        })
    })
} finally {
    randomMocker.restore();
}
```

`#sort-btn` をクリックすると、`highlightCurrentEls` 関数を使って各ステップで比較されている要素をハイライトしてください。

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
let flag = false;
const temp = highlightCurrentEls;
highlightCurrentEls = () => flag = true;
try {
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    assert.isTrue(flag);
} finally {
    highlightCurrentEls = temp;
}
```

`#sort-btn` をクリックすると、`#starting-array` は初期配列と最初の2つの整数が `highlightCurrentEls` を使ってハイライトされた開始ステップを表します。

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");   
genBtn.dispatchEvent(new Event("click"));
sortBtn.dispatchEvent(new Event("click"));
const firstContainer = document.querySelector("#starting-array");
const children = firstContainer.children
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|hsl\(0,\s*100%,\s*50%\))/;
assert.match(getComputedStyle(children[0]).border, redBorderRegex);
assert.match(getComputedStyle(children[1]).border, redBorderRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

#sort-btn {
   display: none 
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js
const arrayContainer = document.getElementById("array-container");
const generateArrayBtn = document.getElementById("generate-btn");
const sortArrayBtn = document.getElementById("sort-btn");
const arrayLength = 5;
const minVal = 1;
const maxVal = 100;
const isStart = () => arrayContainer.children.length === 1;
const clearArrayContainer = () => { arrayContainer.innerHTML = '<div id="starting-array"></div>' };
const showSortBtn = () => {sortArrayBtn.style.display = "inline-block"};
const hideSortBtn = () => {sortArrayBtn.style.display = "none"};
const generateElement = () => Math.floor(Math.random() * maxVal + minVal);
const generateArray = () => Array.from({ length: arrayLength }, generateElement);
const fillArrContainer = (container, arr) => {
    container.innerHTML = "";
    arr.forEach(i => {
        const el = document.createElement("span");
        el.innerText = i;
        el.id = `number-${i}`
        container.appendChild(el);
    })
}
const isOrdered = (el1, el2) => el1 <= el2;

const swapElements = (arr, n = 0) => {
    if (n < arr.length - 1 && !isOrdered(arr[n], arr[n + 1])) {
        const temp = arr[n];
        arr[n] = arr[n + 1];
        arr[n + 1] = temp;
    }
}
const getLastChildren = () => arrayContainer.lastElementChild;
const getLastArr = () => {
    const els = Array.from(getLastChildren().children);
    const arr = els.map(el => Number(el.id.replace("number-", "")));
    return arr;
}
const generateContainer = () => {
    const container = document.createElement("div");
    arrayContainer.appendChild(container);
    return container;
}
const highlightCurrentEls = (container, n = 0) => {
    const children = container.children;
    children[n].style.border = "2px dashed red";
    children[n + 1].style.border = "2px dashed red";
}

const highlightSorted = () => {
    getLastChildren().style.border = "4px solid green";
}

const bubbleSort = () => {
    let swapped = true;
    while (swapped) {
        const startingArr = getLastArr();
        startingArr.forEach((_, i) => {
            if (i + 1 < arrayLength) {
                highlightCurrentEls(getLastChildren(), i);
                const arr = getLastArr();
                swapElements(arr, i);
                fillArrContainer(generateContainer(), arr);
            }
        })
        const lastArr = getLastArr()
        if (startingArr.every((el, i) => el === lastArr[i])) swapped = false;
    }
}


generateArrayBtn.addEventListener("click", () => {
    if (!isStart()) {
        clearArrayContainer();
    }
    fillArrContainer(document.getElementById("starting-array"), generateArray());
    showSortBtn();
});

sortArrayBtn.addEventListener("click", () => {
    bubbleSort();
    highlightSorted();
    hideSortBtn();
})
```
