---
id: 657ef2a86d4e545cec9a85fb
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Sia nel camel case che nel pascal case, i caratteri maiuscoli indicano l’inizio di nuove parole. Per convertire la stringa di input in snake case, dovrai verificare se i caratteri nella stringa di input sono maiuscoli.

Puoi usare il metodo `.isupper()` della stringa per controllare se un carattere è maiuscolo. Questo metodo restituisce `True` se il carattere è maiuscolo e `False` se non lo è.

All’interno del loop `for`, aggiungi un’istruzione `if` per verificare se il carattere corrente è maiuscolo. Sposta l’istruzione `pass` all’interno della nuova istruzione `if`.

# --hints--

Dovresti scrivere una nuova istruzione if con char.isupper() come condizione. Ricorda di aggiungere i due punti alla fine e usa la parola chiave pass per riempire il corpo dell’istruzione if.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

Dovresti sostituire l’istruzione `pass` nel corpo del loop con l’istruzione `if`.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
