---
id: 657efa642593c5746acc5c81
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

All'interno del corpo dell'istruzione `if`, devi convertire ogni carattere maiuscolo in minuscolo e anteporre un underscore a questo carattere minuscolo.

Usa il metodo `.lower()` della stringa per convertire i caratteri maiuscoli in minuscoli. Poi, anteponi un underscore al carattere. Assegna i risultati a una variabile chiamata `converted_character`.

# --hints--

Dovresti assegnare il carattere modificato a una variabile chiamata `converted_character`.

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

Non dovresti avere `pass` nella tua istruzione `if`.

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
