---
id: 657efce98e958b75df86b305
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

All'interno del corpo dell'istruzione `if` aggiungerai il carattere convertito alla lista che hai creato in precedenza.

Per questo, verrà usato il metodo `.append()`. Questo metodo aggiunge un oggetto dato alla fine della lista su cui viene invocato.

Usa il metodo `.append()` su `snake_cased_char_list` per aggiungere `converted_character` alla lista.

# --hints--

Dovresti usare il metodo `.append()` per aggiungere il carattere convertito a `snake_cased_char_list`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\s*\(\s*converted_character\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
