---
id: 657ee28cefc4945568287673
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Con la lista vuota pronta, ora puoi iniziare a iterare sulla stringa di input e convertirla in snake case.

All'interno della funzione, sotto la lista che hai appena creato, aggiungi un loop `for` per iterare su `pascal_or_camel_cased_string`. Assicurati di chiamare la variabile obiettivo `char`. Per ora, aggiungi un'istruzione `pass` come segnaposto nel corpo del loop.

# --hints--

Dovresti scrivere un nuovo loop `for` con la variabile obiettivo chiamata `char`. Non dimenticare i due punti alla fine e di usare `pass` per riempire il corpo del loop.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
