---
id: 657effaa2a5e0277d71f9cbe
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

A questo punto, la variabile `snake_cased_char_list` contiene la lista di caratteri convertiti. Per unire questi caratteri in un'unica stringa, puoi usare il metodo `.join()`.

Il metodo `join` funziona concatenando ogni elemento di una lista in una stringa, separati da una stringa designata, chiamata separatore.

```py
result_string = ''.join(characters)
```

L'esempio sopra unisce gli elementi della lista `characters` in un'unica stringa dove ogni elemento è concatenato usando una stringa vuota come separatore.

Ora, subito dopo il loop `for`, usa il metodo `.join()` per unire gli elementi in `snake_cased_char_list` usando una stringa vuota come separatore. Assegna il risultato a una nuova variabile chiamata `snake_cased_string`.


# --hints--

Dovresti unire i caratteri nella lista `snake_cased_char_list` in una stringa usando il metodo `.join()` e assegnare la stringa a una variabile chiamata `snake_cased_string` dopo il tuo loop `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
