---
id: 657e928716b77b2277980276
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo progetto, costruirai un programma che prende come input una stringa formattata `camelCase` o `PascalCase` e la converte in una stringa formattata `snake_case` usando due approcci. Prima userai un for loop e poi la list comprehension per ottenere gli stessi risultati. Vedrai come la list comprehension può rendere il tuo codice più conciso.

Inizia definendo una nuova funzione chiamata `convert_to_snake_case()` che accetta come input una stringa chiamata `pascal_or_camel_cased_string`. Per ora, aggiungi un'istruzione `pass` all'interno della funzione.

# --hints--

Dovresti definire una nuova funzione chiamata `convert_to_snake_case()`. Non dimenticare i due punti alla fine e usa `pass` per completare il corpo della funzione.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

La tua funzione dovrebbe prendere un parametro chiamato `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
