---
id: 657f456223b8c1187b461987
title: Passaggio 19
challengeType: 20
dashedName: step-19
---

# --description--

In Python, una <dfn>list comprehension</dfn> è una struttura che ti permette di generare una nuova lista applicando un’espressione a ogni elemento di un iterabile esistente e, opzionalmente, filtrando gli elementi con una condizione. Oltre a essere più sintetiche, le list comprehension spesso sono più veloci.

Una list comprehension di base consiste in un’espressione seguita da una clausola `for`:

```py
spam = [i * 2 for i in iterable]
```

L’esempio sopra usa la variabile `i` per iterare su `iterable`. Ogni elemento della lista risultante si ottiene valutando l’espressione `i * 2` all’iterazione corrente.

In questo passaggio, devi riempire la lista vuota `snake_cased_char_list` usando la sintassi della list comprehension.

Trasforma la tua lista vuota in una list comprehension che converte ogni carattere in `pascal_or_camel_cased_string` in un carattere minuscolo e gli antepone un underscore (il codice che avevi commentato prima potrebbe aiutarti a scrivere l’espressione). Usa `char` per iterare su `pascal_or_camel_cased_string`. 

# --hints--

Dovresti trasformare `snake_cased_char_list` in una list comprehension che itera su `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

La tua list comprehension dovrebbe usare `char` per iterare su `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

La tua list comprehension dovrebbe valutare `'_' + char.lower()` per ogni `char` in `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
