---
id: 657f0353c9523d7d896873ea
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

All'interno della funzione `main()`, sostituisci l'istruzione `pass` con una chiamata alla funzione `convert_to_snake_case()`, passando come input la stringa `'aLongAndComplexString'`.

Per mostrare l'output, passa la chiamata alla funzione come argomento alla funzione `print()`.

# --hints--

Dovresti chiamare `convert_to_snake_case()` all'interno della funzione `main()` e passare `'aLongAndComplexString'` come input alla funzione.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\s*\(\s*convert_to_snake_case\s*\(\s*("|')aLongAndComplexString\1\s*\)\s*\)/);
    }
})
```

Non dovresti avere `pass` nella tua funzione `main`.

```js
({
    test: () => {
        const commentless_code = __helpers.python.removeComments(code);
        const transformedCode = commentless_code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    pass
--fcc-editable-region--
```
