---
id: 657f465f8e718b19c5105ae5
title: Passaggio 20
challengeType: 20
dashedName: step-20
---

# --description--

Le list comprehension accettano istruzioni condizionali, per valutare l’espressione fornita solo se certe condizioni sono soddisfatte:

```py
spam = [i * 2 for i in iterable if i > 0]
```

Come puoi vedere dall’output, la lista di caratteri generata da `pascal_or_camel_cased_string` è stata unita. Poiché l’espressione all’interno della list comprehension viene valutata per ogni carattere, il risultato è una stringa minuscola con tutti i caratteri separati da un underscore.

Segui l’esempio sopra per aggiungere una clausola `if` alla tua list comprehension in modo che l’espressione venga eseguita solo se il carattere è maiuscolo.

# --hints--

Dovresti aggiungere una clausola `if` con la condizione `char.isupper()` alla tua list comprehension.

```js
({ 
    test: () => assert(runPython(`
    ifs = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_ifs()
    len(ifs) == 1 and ifs[0].is_equivalent("char.isupper()")
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
