---
id: 657f43d341a0dd17120cdb08
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Dovrai convertire i caratteri maiuscoli in minuscoli e aggiungere un underscore prima di essi.

Prima di procedere con la comprensione delle liste, darai alla tua funzione un valore di ritorno. In questo modo potrai controllare l’output.

Usa l’istruzione `return` per restituire la lista `snake_cased_char_list` dalla tua funzione.

# --hints--

Dovresti restituire la lista `snake_cased_char_list`. Assicurati che l’indentazione sia corretta.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*snake_cased_char_list/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--

--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
