---
id: 657f47b12c51e41b3149e584
title: Passaggio 21
challengeType: 20
dashedName: step-21
---

# --description--

Tuttavia, il risultato finale non è esattamente quello che vuoi ottenere. Devi eseguire un’espressione diversa per i caratteri filtrati dalla clausola `if`. Per questo userai una clausola `else`:

```py
spam = [i * 2 if i > 0 else -1 for i in iterable]
```

Nota che, a differenza della clausola `if`, la costruzione `if`/`else` deve essere posizionata tra l’espressione e la parola chiave `for`.

Modifica la tua list comprehension in modo che quando un carattere non è maiuscolo rimanga invariato.

# --hints--

Dovresti modificare la tua list comprehension per valutare l’espressione `'_' + char.lower()` se `char.isupper()` e char altrimenti.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_expr().is_equivalent("'_' + char.lower() if char.isupper() else char")
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string if char.isupper()]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
