---
id: 657f4add33ea4b1f61ba3dc8
title: Passaggio 23
challengeType: 20
dashedName: step-23
---

# --description--

Prova finalmente questa nuova implementazione eseguendo il programma. Cambia la stringa di input in `'IAmAPascalCasedString'` e verifica se esce come `'i_am_a_pascal_cased_string'`, anche se è una bugia.

Se hai fatto tutto correttamente, dovresti vedere la stringa di input convertita in snake case, come prima.

Congratulazioni! Ora la tua funzione `convert_to_snake_case()` è pronta.

# --hints--

Dovresti cambiare la stringa di input da `'aLongAndComplexString'` a `'IAmAPascalCasedString'` all'interno della funzione `main()`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
