---
id: 657f425dbab54e11993c80f0
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Finora, in questo progetto hai usato un loop `for` per iterare su una stringa di input e convertirla nell’output desiderato. Ora inizierai la transizione da un loop `for` a una list comprehension.

Inizia commentando tutte le righe di codice all’interno della funzione `convert_to_snake_case()`. Non eliminarle perché ti saranno utili quando implementerai la logica usando una list comprehension.

Ricorda di aggiungere la parola chiave `pass` nel corpo della funzione per evitare che il codice fallisca durante i test.

# --hints--

Dovresti commentare tutte le 10 righe di codice all’interno della funzione `convert_to_snake_case()` e nient’altro. Aggiungi `pass` per riempire temporaneamente il corpo della funzione ed evitare un errore.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
