---
id: 657f0044be09db790b1eb1c5
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Nel pascal case, le stringhe iniziano con una lettera maiuscola. Dopo aver convertito tutti i caratteri in minuscolo e aver aggiunto un underscore, c'è la possibilità che all'inizio della stringa ci sia un underscore in più.

Il modo più semplice per risolvere questo problema è usare il metodo `.strip()` per le stringhe, che rimuove da una stringa tutti i caratteri iniziali o finali presenti in un insieme di caratteri passato come argomento. Per esempio:

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

Il metodo `strip()` viene applicato a `original_string`. Questo rimuove tutti gli underscore iniziali e finali. Il risultato dell'esempio sopra sarebbe la stringa `'example_string'`.

Dichiara una nuova variabile chiamata `clean_snake_cased_string` e assegnale il risultato del metodo `.strip()` applicato a `snake_cased_string`, passando `'_'` come argomento al metodo.

# --hints--

Dovresti assegnare la stringa risultante da `snake_cased_string.strip('_')` a una variabile chiamata `clean_snake_cased_string`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
