---
id: 663b16e62fee463b4caf46e9
title: Passaggio 18
challengeType: 20
dashedName: step-18
---

# --description--

Dopo aver unito gli elementi della lista `snake_cased_char_list`, dovrai rimuovere eventuali underscore iniziali o finali dalla stringa risultante. Per questo, usa il metodo `strip` con il carattere underscore `_` come argomento.

I metodi possono essere concatenati, il che significa che il risultato di una chiamata a metodo può essere usato come oggetto per un'altra chiamata a metodo.

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

Nell'esempio sopra, il metodo `.upper()` è concatenato a `' '.join(words_list)`, quindi `.upper()` viene chiamato sul risultato della chiamata `.join()`.

Modifica l'istruzione `return` concatenando a `''.join(snake_cased_char_list)` una chiamata al metodo `.strip()` per rimuovere eventuali underscore iniziali o finali.

# --hints--

Dovresti modificare l'istruzione return per includere il metodo `strip` con il carattere underscore `_` come argomento.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
