---
id: 56533eb9ac21ba0edf2244cf
title: 建置記錄群集
challengeType: 26
dashedName: build-a-record-collection
---

# --description--

你正在建立一個函式（程式），用於協助維護音樂專輯群集。該群集被組織為一個物件，包含多個同樣是物件的專輯。每個專輯在群集中以唯一的 `id` 作為屬性名稱來表述。在每個專輯物件中，有各種屬性描述該專輯的資訊。並非所有專輯都有完整的資訊。

`updateRecords` 函式（程式）接受 4 個由以下函式（程式）參數表述的引數：

-   `records`－包含多個獨立專輯的物件
-   `id` - 一個數字，用來表述 `records` 物件中特定的相簿
-   `prop` - 一個字串，表述要更新的專輯屬性名稱
-   `value` - 包含用於更新相簿屬性的資訊的字串

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你的函式必須始終傳回整個 `records` 物件。
2. 如果 `value` 是空字串，則從相簿中刪除指定的 `prop` 屬性。
3. 如果 `prop` 不是 `tracks` 且 `value` 不是空字串，則將該 `value` 指定給該專輯的 `prop`。
4. 如果 `prop` 是 `tracks` 且 `value` 不是空字串，但專輯沒有 `tracks` 屬性，則創建一個空陣列並將 `value` 添加到其中。
5. 如果 prop 是 `tracks` 且 `value` 不是空字串，將 `value` 添加到專輯現有的 `tracks` 陣列末端。

**注意：** 測試中使用了 `recordCollection` 物件的複本。你的函式（程式）不應直接參考 `recordCollection` 物件，只能參考函式（程式）參數。

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

你應該有一個 `updateRecords` 函式（程式）。

```js
assert.isFunction(updateRecords);
```

在執行 `updateRecords(recordCollection, 5439, "artist", "ABBA")` 之後，`artist` 應該是字串 `ABBA`。

```js
assert.equal(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'],
    'ABBA'
);
```

在執行 `updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")` 之後，`tracks` 應該只有字串 `Take a Chance on Me` 作為最後且唯一的元素。

```js
updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me');
assert.lengthOf(_recordCollection[5439]['tracks'], 1);
assert.equal(_recordCollection[5439]['tracks'].pop(), 'Take a Chance on Me');
```

在執行 `updateRecords(recordCollection, 2548, "artist", "")` 之後，`artist` 不應該被設定。

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('artist'));
```

在執行 `updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")` 之後，`tracks` 應該會有字串 `Addicted to Love` 作為最後一個元素。

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop(), 'Addicted to Love'
);
```

在執行 `updateRecords(recordCollection, 2468, "tracks", "Free")` 之後，`tracks` 應該將字串 `1999` 作為第一個元素。

```js
assert.equal(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0], '1999'
);
```

在執行 `updateRecords(recordCollection, 2548, "tracks", "")` 之後，`tracks` 不應該被設定。

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('tracks'));
```

在執行 `updateRecords(recordCollection, 1245, "albumTitle", "Riptide")` 之後，`albumTitle` 應該是字串 `Riptide`。

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ], 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};



```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
