---
id: 68f1f8c248306da204655117
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Sifa ya kipengele `muted` itazima sauti wakati wa kuanzisha uchezaji. Ikiwa unawezeshwa `controls`, mtumizi ataweza kuwasha tena sauti. Kukosa sifa ya kipengele `muted` kutasababisha sauti kuchezwa wakati wa kuanzisha uchezaji.

Sifa ya kipengele `muted` ni sifa ya kipengele ya thamani ya kweli au siyo kweli na haihitaji thamani.

Ongeza sifa ya kipengele `muted` kwenye kipengele cha `video`.

# --hints--

Kipengele chako cha `video` kinapaswa kuwa na sifa ya kipengele `muted`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
