---
id: 68f1f371fefd448b4f70b898
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Katika funzo lililopita, ulijifunza kuhusu sifa mbalimbali za kipengele zinazopatikana kwa kipengele cha `video`. Sifa ya kipengele `width` inaamua upana wa video kwa pikseli.

Ongeza sifa ya kipengele `width` kwa kipengele cha `video` kwa thamani ya `640`.

# --hints--

Kipengele chako cha `video` kinapaswa kuwa na sifa ya kipengele `width` yenye thamani ya `640`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
