---
id: 68f1f6f97af9749a42260a43
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Sifa ya kipengele ya `controls` hutoa vidhibiti vya kucheza ikijumuisha kucheza, kurudisha nyuma, na udhibiti wa sauti kwa kipengele cha `video`. 

Sifa ya kipengele ya `controls` ni sifa ya kipengele ya thamani ya kweli au siyo kweli na haihitaji thamani.

Ongeza sifa ya kipengele ya `controls` kwa kipengele cha `video`.

Sasa unapaswa kuona kipengele cha `video` kikiwa kimeonyeshwa kwenye ukurasa.

# --hints--

Kipengele chako cha `video` kinapaswa kuwa na sifa ya kipengele ya `controls`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
