---
id: 68f1f57723debf9383c75947
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Sifa ya kipengele `loop` itaanzisha upya video mara tu baada ya kuchezwa kabisa. Fikiria meme ya mtandao inayorudia kuchezwa. Kukosa sifa ya kipengele `loop` kutafanya video ichezwe mara moja tu.

Sifa ya kipengele `loop` ni sifa ya kipengele yenye thamani ya kweli au siyo kweli na haihitaji thamani.

Ongeza sifa ya kipengele `loop` kwenye kipengele cha `video`.

# --hints--

Kipengele chako cha `video` kinapaswa kuwa na sifa ya kipengele `loop`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
