---
id: 68f1fac259a4afabb478b960
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Huenda umeona haukufanya unganisho na video halisi. Utayafanya katika awamu inayofuata. Linapokuja suala la aina za faili za video, kuna tofauti katika msaada wa kivinjari. Ili kukidhi hili, unaweza kutumia vipengele vya `source` ndani ya kipengele cha `video` na kivinjari kita-chagua `source` inayolingana ya kwanza.

Hapa kuna mfano wa kipengele cha `source`:

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

Kipengele cha `source` ni kipengele kisicho na maudhui (void element) kwa hivyo hakina lebo ya kufunga.

Ongeza kipengele cha `source` ndani ya kipengele chako cha `video`.


# --hints--

Unapaswa kuwa na kipengele cha `source` ndani ya kipengele chako cha `video`.

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
