---
id: 68f1eb9ad0a7575e4dbc31d5
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Ifuatayo, tengeneza kipengele cha `video` chini ya `h1`. Katika hatua chache zijazo, utaongeza sifa za kipengele zinazohitajika ili kufanya mchezaji wa video ufanye kazi.

# --hints--

Unapaswa kuwa na kipengele cha `video`.

```js
assert.exists(document.querySelector('video'));
```

Kipengele kipya cha `video` kinapaswa kuwa chini ya `h1`.

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
