---
id: 68f1f993460f06a64c033061
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Sifa ya `poster` ni picha ndogo ya video. Fikiria video unazotazama kwenye YouTube. Inaonyeshwa wakati video inapakuliwa. Ikiwa sifa hii itakosekana, fremu ya kwanza ya video itaonyeshwa wakati wa awamu ya upakuaji.

Sasa, ongeza sifa ya `poster` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` kwenye kipengele chako cha `video`.

# --hints--

Kipengele chako cha `video` kinapaswa kuwa na sifa ya `poster` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
