---
id: 68f2701b2160c765f23b60f8
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Aina nyingine ya kawaida ya MIME ni aina ya `video/ogg` MIME.

Ogg ni muundo wa kontena wa vyombo vya habari vya kidijitali ulioundwa kutoa mtiririko mzuri na usimamizi wa vyombo vya habari vya kidijitali. Inasimamiwa na Taasisi ya Xiph.Org na ni huru na wazi, isiyozuiliwa na hati miliki za programu. Jina lake linatokana na "ogging", lugha ya mtaalamu kutoka kwenye mchezo wa kompyuta Netrek.

Chini ya kipengele chako cha `source` cha pili, ongeza kipengele cha tatu cha `source` na umpe sifa ya `src` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` na sifa ya `type` yenye thamani ya `video/ogg`.

# --hints--

Unapaswa kuwa na kipengele cha tatu cha `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

Kipengele chako cha tatu cha `source` kinapaswa kuwa na sifa ya `src`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

Kipengele chako cha tatu cha `source` kinapaswa kuwa na sifa ya `src` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

Kipengele chako cha tatu cha `source` kinapaswa kuwa na sifa ya `type`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

Kipengele chako cha tatu cha `source` kinapaswa kuwa na sifa ya `type` yenye thamani ya `video/ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
