---
id: 68f27139bddf9f6bf2821acf
title: Hatua 13
challengeType: 0
dashedName: step-13
---

# --description--

Kipengele cha mwisho cha `source` utakachoongeza kitakuwa kwa aina ya MIME ya `video/quicktime`.

QuickTime ni usanifu wa vyombo vya habari unaoweza kupanuliwa uliotengenezwa na Apple, unaounga mkono kucheza, kupeleka mtandao, kusimbua, na kubadilisha aina mbalimbali za fomati za vyombo vya habari vya kidijitali. Sio maarufu kama fomati ya MP4, unaweza kuhitaji kwa msaada wa programu ya zamani.

Chini ya kipengele chako cha tatu cha `source`, ongeza kipengele cha nne cha `source` na kipa sifa ya `src` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` na sifa ya `type` yenye thamani ya `video/quicktime`.

Hongera! Umehitimisha warsha ya Mchezaji wa Video wa HTML.

# --hints--

Unapaswa kuwa na kipengele cha nne cha `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

Kipengele chako cha nne cha `source` kinapaswa kuwa na sifa ya `src`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

Kipengele chako cha nne cha `source` kinapaswa kuwa na sifa ya `src` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

Kipengele chako cha nne cha `source` kinapaswa kuwa na sifa ya `type`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

Kipengele chako cha nne cha `source` kinapaswa kuwa na sifa ya `type` yenye thamani ya `video/quicktime`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
