---
id: 68f26d52b13e0f5b6e2d2e09
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Aina nyingine ya kawaida ya MIME ni aina ya MIME ya `video/webm`.

WebM ni muundo wa faili wa vyombo vya habari vya sauti na picha wa chanzo wazi ulioanzishwa na Google, hasa kwa ajili ya maudhui ya vyombo vya habari mtandaoni. Inasaidia codecs za video kama VP8, VP9, na AV1, pamoja na codecs za sauti kama Vorbis na Opus, na hivyo kuwa chaguo maarufu kwa vipengele vya video na sauti vya HTML5.

Chini ya kipengele chako cha kwanza cha `source`, ongeza kipengele kingine cha `source` na kipa sifa ya `src` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` na sifa ya `type` yenye thamani ya `video/webm`.

# --hints--

Unapaswa kuwa na kipengele cha pili cha `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

Kipengele chako cha pili cha `source` kinapaswa kuwa na sifa ya `src`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

Kipengele chako cha pili cha `source` kinapaswa kuwa na sifa ya `src` yenye thamani ya `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

Kipengele chako cha pili cha `source` kinapaswa kuwa na sifa ya `type`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

Kipengele chako cha pili cha `source` kinapaswa kuwa na sifa ya `type` yenye thamani ya `video/webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
