---
id: 68f26a5f79e97c5032c10642
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Umetumia faili la video lenye kiambatisho cha faili la `mp4`, na unahitaji kumwambia kivinjari ili kijue jinsi ya kusoma faili hilo.

Utatumia sifa ya kipengele ya `type` kubainisha aina ya MIME ya `video/mp4`.

MIME (Multipurpose Internet Mail Extensions) ni kiwango cha kuelezea nyaraka katika aina nyingine mbali na maandishi ya ASCII, kwa mfano, sauti, video, na picha.

MP4, inayojulikana rasmi kama MPEG-4 Sehemu ya 14, ni muundo wa kidhibiti wa vyombo vya habari vya kidijitali. Inatumiwa sana kuhifadhi video na sauti, lakini pia inaweza kujumuisha aina nyingine za data kama vichwa vidogo na picha zisizohamishwa. Mafaili ya MP4 yameundwa kwa ajili ya kupelekwa mtandaoni na yanayolingana na vifaa na majukwaa mengi.

Sasa, ongeza sifa ya kipengele ya `type` na thamani `video/mp4`.

# --hints--

Kipengele chako cha `source` kinapaswa kuwa na sifa ya kipengele ya `type`.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

Kipengele chako cha `source` kinapaswa kuwa na sifa ya kipengele ya `type` yenye thamani `video/mp4`.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--
      
--fcc-editable-region--
    >
  </video>
</body>
</html>
```
