---
id: 66fe4f33a2cc9b33f4d5cd9b
title: Побудова генератора перестановок
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

У цій лабораторній роботі ви створите генератор перестановок, який прийматиме рядок і повертатиме всі можливі перестановки символів у цьому рядку. Наприклад, можливі перестановки рядка `cat` — це `cat`, `cta`, `act`, `atc`, `tac` та `tca`.

Рекурсивний спосіб створення перестановок рядка працює шляхом збереження фіксованої початкової частини рядка (префікса) та створення перестановок для залишку.

Наприклад, розглянемо слово `machine`. Перший етап створення перестановок полягає у фіксації `m` як префікса рядка, а потім створенні перестановок для решти рядка, `achine`.

Для решти рядка перестановки продовжуються так само. До префікса додається одна літера, можливо `c`, тож префікс стає `mc`. Потім кожна з перестановок `ahine` конкатенується з префіксом.

Це триває, доки префікс не міститиме всі літери, а решта рядка буде порожньою — це означає, що створена одна перестановка.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити функцію з назвою `permuteString`.
2. Функція `permuteString` має приймати один параметр — рядок, а також два параметри зі значеннями за замовчуванням: значення префікса та порожній масив для збереження і повернення результатів. Значення префікса має бути рядком, який використовується для накопичення символів, щоб сформувати перестановку. Функція викликатиметься з одним аргументом, наприклад `permuteString("cat")`.
3. Усередині функції потрібно перевірити, чи дорівнює довжина переданого рядка `0`. Якщо так, додайте поточний префікс до результатів і поверніть результати.
4. Ітеруйте над кожним символом у вхідному рядку, і для кожної ітерації видаляйте поточний символ із рядка та рекурсивно викликайте функцію `permuteString` з оновленими аргументами, щоб побудувати решту перестановок.
5. Ви повинні повернути фінальний масив результатів.
6. Переконайтеся, що перестановки унікальні, видаляючи дублікати.

# --hints--

Вам потрібна функція `permuteString`.

```js
assert.isFunction(permuteString);
```

У вашій функції `permuteString` слід використовувати рекурсію.

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` має повертати `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`.

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` має повертати `[ "fcc", "cfc", "ccf" ]`.

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` має повертати `[ "p" ]`.

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` має повертати `[""]`.

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")` має повертати `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]`.
`.

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` має повертати правильні результати.

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
