---
id: 68de7360e1832ea180577a8a
title: 建置一個第 N 個費氏數列計算器
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `fibonacci` 的函式（程式）。
2. 你應該在 `fibonacci` 函式（程式）中定義一個名為 `sequence` 的列表，並且它應該以值 `[0, 1]` 初始化。
3. `fibonacci` 函式（程式）應該接受一個參數，一個正整數 `n`。
4. 呼叫 `fibonacci(n)` 應該使用動態程式設計方法來計算並傳回 Fibonacci 數列中的第 n 個數字，其中每個數字是前兩個數字的總和。
5. 費波那契數列中位置為 `n` 的每個計算數字應該儲存在索引為 `n - 1` 的 `sequence` 列表中。

# --hints--

你應該有一個名為 `fibonacci` 的函式（程式）。

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

你的 `fibonacci` 函式應該接受一個單一參數。

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

你應該在 `fibonacci` 函式（程式）中擁有一個名為 `sequence` 的列表，並初始化為 `[0, 1]`。

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

`fibonacci(0)` 應該傳回 `0`。

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

`fibonacci(1)` 應該傳回 `1`。

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

`fibonacci(2)` 應該傳回 `1`。

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

`fibonacci(3)` 應該傳回 `2`。

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

`fibonacci(5)` 應該傳回 `5`。

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

`fibonacci(10)` 應該傳回 `55`。

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

`fibonacci(15)` 應該傳回 `610`。

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

你不應該在你的程式碼中使用遞迴。

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```

