---
id: 56533eb9ac21ba0edf2244b3
title: Celsius in Fahrenheit umrechnen
challengeType: 1
forumTopicId: 16806
dashedName: convert-celsius-to-fahrenheit
---

# --description--

Die Formel zur Umrechnung von Celsius in Fahrenheit ist die Temperatur in Celsius mal `9/5`, plus `32`.

Sie erhalten eine Variable `celsius`, die eine Temperatur in Celsius darstellt. Verwenden Sie die bereits definierte Variable `fahrenheit` und weisen Sie ihr die entsprechende Temperatur in Fahrenheit zu, die der gegebenen Celsius-Temperatur entspricht. Nutzen Sie die oben genannte Formel, um die Celsius-Temperatur in Fahrenheit umzurechnen.

# --hints--

`convertCtoF(0)` sollte eine Zahl zurückgeben

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` sollte einen Wert von `-22` zurückgeben

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` sollte einen Wert von `14` zurückgeben

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` sollte einen Wert von `32` zurückgeben

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` sollte einen Wert von `68` zurückgeben

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` sollte einen Wert von `86` zurückgeben

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js
function convertCtoF(celsius) {
  let fahrenheit;
  return fahrenheit;
}

convertCtoF(30);
```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
