---
id: a24c1a4622e3c05097f71d67
title: Wo gehöre ich hin
challengeType: 1
forumTopicId: 16094
dashedName: where-do-i-belong
---

# --description--

Geben Sie den niedrigsten Index zurück, an dem ein Wert (zweites Argument) in ein Array (erstes Argument) eingefügt werden sollte, nachdem es sortiert wurde. Der zurückgegebene Wert sollte eine Zahl sein.

Zum Beispiel sollte `getIndexToIns([1,2,3,4], 1.5)` `1` zurückgeben, weil es größer ist als `1` (Index 0), aber kleiner als `2` (Index 1).

Ebenso sollte `getIndexToIns([20,3,5], 19)` `2` zurückgeben, weil das Array nach dem Sortieren wie `[3,5,20]` aussehen wird und `19` kleiner als `20` (Index 2) und größer als `5` (Index 1) ist.

# --hints--

`getIndexToIns([10, 20, 30, 40, 50], 35)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` sollte `3` zurückgeben.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 30));
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` sollte `2` zurückgeben.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(getIndexToIns([40, 60], 50));
```

`getIndexToIns([40, 60], 50)` sollte `1` zurückgeben.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(getIndexToIns([3, 10, 5], 3));
```

`getIndexToIns([3, 10, 5], 3)` sollte `0` zurückgeben.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(getIndexToIns([5, 3, 20, 3], 5));
```

`getIndexToIns([5, 3, 20, 3], 5)` sollte `2` zurückgeben.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` sollte `2` zurückgeben.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([2, 20, 10], 19)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(getIndexToIns([2, 20, 10], 19));
```

`getIndexToIns([2, 5, 10], 15)` sollte `3` zurückgeben.

```js
assert.strictEqual(getIndexToIns([2, 5, 10], 15), 3);
```

`getIndexToIns([2, 5, 10], 15)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(getIndexToIns([2, 5, 10], 15));
```

`getIndexToIns([], 1)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(getIndexToIns([], 1));
```

`getIndexToIns([], 1)` sollte `0` zurückgeben.

```js
assert.strictEqual(getIndexToIns([], 1), 0);
```

# --seed--

## --seed-contents--

```js
function getIndexToIns(arr, num) {
  return num;
}

getIndexToIns([40, 60], 50);
```

# --solutions--

```js
function getIndexToIns(arr, num) {
  arr = arr.sort((a, b) => a - b);

  for (let i = 0; i < arr.length; i++) {
    if (arr[i] >= num) {
      return i;
    }
  }

  return arr.length;
}

getIndexToIns([40, 60], 50);
```
