---
id: a6e40f1041b06c996f7b2406
title: Finderlohn
challengeType: 1
forumTopicId: 16016
dashedName: finders-keepers
---

# --description--

Erstellen Sie eine Funktion, die ein Array `arr` durchsucht und das erste Element zurückgibt, das einen „Wahrheitstest“ besteht. Das bedeutet, dass bei einem Element `x` der „Wahrheitstest“ bestanden ist, wenn `func(x)` `true` ist. Wenn kein Element den Test besteht, geben Sie `undefined` zurück.

# --hints--

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` sollte `8` zurückgeben.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` sollte `undefined` zurückgeben.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js
function findElement(arr, func) {
  let num = 0;
  return num;
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
