---
id: a77dbc43c33f39daa4429b4f
title: Boo who
challengeType: 1
forumTopicId: 16000
dashedName: boo-who
---

# --description--

Überprüfen Sie, ob ein Wert als Boolean-Primitiv klassifiziert ist. Geben Sie `true` oder `false` zurück.

Boolean-Primitiven sind `true` und `false`.

# --hints--

`booWho(true)` sollte `true` zurückgeben.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` sollte `true` zurückgeben.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` sollte `false` zurückgeben.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` sollte `false` zurückgeben.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` sollte `false` zurückgeben.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` sollte `false` zurückgeben.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` sollte `false` zurückgeben.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` sollte `false` zurückgeben.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` sollte `false` zurückgeben.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` sollte `false` zurückgeben.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js
function booWho(bool) {
  return bool;
}

booWho(null);
```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

booWho(null);
```
