---
id: ab6137d4e35944e21037b769
title: Einen Satz in Titel-Schreibweise umwandeln
challengeType: 1
forumTopicId: 16088
dashedName: title-case-a-sentence
---

# --description--

Geben Sie den bereitgestellten String zurück, wobei der erste Buchstabe jedes Wortes großgeschrieben ist. Stellen Sie sicher, dass der Rest des Wortes in Kleinbuchstaben geschrieben ist.

Für diese Aufgabe sollten Sie auch verbindende Wörter wie `the` und `of` großschreiben.

# --hints--

`titleCase("I'm a little tea pot")` sollte einen String zurückgeben.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` sollte den String `I'm A Little Tea Pot` zurückgeben.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` sollte den String `Short And Stout` zurückgeben.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` sollte den String `Here Is My Handle Here Is My Spout` zurückgeben.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js
function titleCase(str) {
  return str;
}

titleCase("I'm a little tea pot");
```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}

titleCase("I'm a little tea pot");
```
