---
id: acda2fb1324d9b0fa741e6b5
title: Bestätigen Sie das Ende
challengeType: 1
forumTopicId: 16006
dashedName: confirm-the-ending
---

# --description--

Überprüfen Sie, ob ein String (erstes Argument, `str`) mit dem angegebenen Zielstring (zweites Argument, `target`) endet.

Diese Aufgabe _kann_ mit der `.endsWith()`-Methode gelöst werden, die in ES2015 eingeführt wurde. Für diese Aufgabe möchten wir jedoch, dass Sie stattdessen eine der JavaScript-Substring-Methoden verwenden.

# --hints--

`confirmEnding("Bastian", "n")` sollte `true` zurückgeben.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` sollte `true` zurückgeben.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` sollte `false` zurückgeben.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` sollte `false` zurückgeben.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` sollte `true` zurückgeben.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` sollte `true` zurückgeben.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` sollte `false` zurückgeben.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` sollte `false` zurückgeben.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing", "mountain")` sollte `false` zurückgeben.

```js
assert.isFalse(
  confirmEnding(
    'If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing',
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` sollte `true` zurückgeben.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Ihr Code sollte die eingebaute Methode `.endsWith()` nicht verwenden, um die Aufgabe zu lösen.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js
function confirmEnding(str, target) {
  return str;
}

confirmEnding('Bastian', 'n');
```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
