---
id: afcc8d540bea9ea2669306b6
title: Einen String wiederholen
challengeType: 1
forumTopicId: 16041
dashedName: repeat-a-string-repeat-a-string
---

# --description--

Wiederholen Sie einen gegebenen String `str` (erstes Argument) `num` Mal (zweites Argument). Geben Sie einen leeren String zurück, wenn `num` keine positive Zahl ist. Für diese Aufgabe dürfen Sie die eingebaute `.repeat()`-Methode _nicht_ verwenden.

# --hints--

`repeatStringNumTimes("*", 3)` sollte den String `***` zurückgeben.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` sollte den String `abcabcabc` zurückgeben.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` sollte den String `abcabcabcabc` zurückgeben.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` sollte den String `abc` zurückgeben.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` sollte den String `********` zurückgeben.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` sollte einen leeren String zurückgeben (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

Die eingebaute `repeat()`-Methode darf nicht verwendet werden.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

`repeatStringNumTimes("abc", 0)` sollte `""` zurückgeben.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

# --seed--

## --seed-contents--

```js
function repeatStringNumTimes(str, num) {
  return str;
}

repeatStringNumTimes('abc', 3);
```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
