---
id: 694c84330cbe604821b9deef
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

任務控制元件需要一個方法來為名冊添加新成員。創建一個名為 `addCrewMember` 的空函式，該函式接受兩個參數，`crew` 和 `astronaut`。

# --hints--

你應該創建一個名為 `addCrewMember` 的函式。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addCrewMember);
```

你的 `addCrewMember` 函式（程式）應該有參數 `crew` 和 `astronaut`。

```js
const explorer = await __helpers.Explorer(code);
const { addCrewMember } = explorer.allFunctions;
assert.lengthOf(addCrewMember?.parameters, 2);
assert.equal(addCrewMember?.parameters?.[0]?.toString(), "crew");
assert.equal(addCrewMember?.parameters?.[1]?.toString(), "astronaut");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
