---
id: 694c84106e9023f289c03b84
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

你的名單從一名指揮官太空人開始。建立一個名為 `firstAstronaut` 的物件，並包含以下屬性：

| Key | Value |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

你應該創建一個 `firstAstronaut` 變數。

```js
assert.isDefined(firstAstronaut);
```

你的 `firstAstronaut` 應該是一個物件。

```js
assert.isObject(firstAstronaut);
```

你的 `firstAstronaut` 物件應該包含正確的屬性值。

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
