---
id: 694cabb9f1afa9446f36ffc0
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

現在你可以開始建造你的名單。使用 `addCrewMember()` 將 `firstAstronaut` 添加到你的 `squad`。

# --hints--

你應該使用 `squad` 和 `firstAstronaut` 來呼叫 `addCrewMember()`。

```js
assert.match(__helpers.removeJSComments(code), /addCrewMember\s*\(\s*squad\s*,\s*firstAstronaut\s*\)/);
assert.strictEqual(squad.length, 1);
assert.strictEqual(squad[0], firstAstronaut);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
  console.log(`Added ${astronaut.name} as ${astronaut.role}`);
}

--fcc-editable-region--

--fcc-editable-region--
```
