---
id: 694c40d86fe6a78d20b6559c
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將利用 JavaScript 建置一個太空任務的太空人名單。 

首先，為你的名單創建一個名為 `squad` 的空陣列。

# --hints--

你應該創建一個 `squad` 變數。

```js
assert.isDefined(squad);
```

你的 `squad` 變數應該是一個陣列。

```js
assert.isArray(squad);
```

你的 `squad` 陣列應該是空的。

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
