---
id: 694c9c186c089f7bdff0f3c6
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

通常，在函式內驗證你的輸入是良好的做法。使用 `for` 迴圈來迭代 `crew` 陣列，並檢查是否有任何成員已經有與輸入的 `astronaut` 相同的 `id`。如果發現重複，請記錄錯誤訊息 `console.log("Duplicate ID: " + astronaut.id)`，然後呼叫 `return` 以提前離開函式。 

這裡是一個這種常用技術的範例：

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

你的函式應該使用 `for` 迴圈來掃描現有的 `crew` 成員（任何有效的 `for (...)` 形式皆可）。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

在 `for` 迴圈中，你應該檢查重複的 `id`。如果找到，將 `"Duplicate ID: " + astronaut.id` 記錄到主控台，然後呼叫 `return` 以離開函式。

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

當給定有效輸入時，你的 `addCrewMember` 函式（程式）不應該喚起 `console.log()` 錯誤訊息或提早 `return` 離開。

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
