---
id: 694d8ddd106340237efb0c80
title: 步驟 12
challengeType: 1
dashedName: step-12
---

# --description--

`splice()` 方法可以透過在任何位置（包括中間）添加或移除元素來修改陣列。由於 `splice()` 傳回一個包含被移除元素的陣列，你可以使用它在不改變原始陣列的情況下交換陣列中的兩個元素。這可以使用以下技巧在一行中完成：

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

這裡有一個範例：

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

第三行中應用的技術運作方式如下：

- `splice(3, 1, copyArray[1])` 移除索引為 `3`（`55`）的元素

- 它將來自 `copyArray[1]`（`97`）的元素插入到索引 `3`。

- `splice()` 傳回 `[55]`，一個包含被移除元素的陣列

- `[0]` 從該陣列中擷取 `55`

- 該值（`55`）被指定回 `copyArray[1]`，完成交換。

將此技術用於 `updatedCrew` 陣列，以交換位於 `fromIndex` 和 `toIndex` 的太空人。

# --hints--

你應該透過呼叫 `splice()` 方法來執行交換。

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

你應該使用 `updatedCrew` 陣列來呼叫 `splice()` 以執行交換。你應該使用單行技巧。

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

你的 `swapCrewMembers` 函式不應該變更原始輸入的 `crew` 陣列。

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
