---
id: 694e1931f6509ae16aa1c36f
title: 步驟 18
challengeType: 1
dashedName: step-18
---

# --description--

任務控制已提醒你，EVA 合格太空人的列表也應該依照優先權由高到低排序。排序陣列有幾種方法——也許最基本的是 `bubble sort`。

`Bubble sort` 的運作方式是反覆遍歷一個列表，比較相鄰的條款，並在它們的順序錯誤時交換它們。每次遍歷後，根據你的排序標準應該排在最前面的條款會更接近（或「冒泡」到）它在陣列中的正確位置。以下是如何使用 `bubble sort` 按優先權由高到低排序 `crew`：

```js
// Outer loop: controls how many passes we make
for (let i = 0; i < crew.length - 1; i++) {
  // Inner loop: compares neighboring items
  for (let j = 0; j < crew.length - 1 - i; j++) {
    // If current member has lower priority than next, swap
    if (crew[j].priority < crew[j + 1].priority) {
      // Using a temp variable for the swap
        const temp = crew[j];
        crew[j] = crew[j + 1];
        crew[j + 1] = temp;
    }
  }
}
```

創建一個名為 `sortByPriorityDescending` 的新輔助函式，該函式接受一個 `crew` 參數。此函式應直接對輸入的 `crew` 陣列依優先權由高到低排序，使用兩個巢狀的 `for` 迴圈，且不應傳回任何值。 

# --hints--

你應該有一個名為 `sortByPriorityDescending` 且帶有 `crew` 參數的函式（程式）。

```js
assert.isFunction(sortByPriorityDescending);
const regex = __helpers.functionRegex('sortByPriorityDescending', ['crew']);
const cleaned = __helpers.removeJSComments(sortByPriorityDescending.toString());
assert.match(cleaned, regex);
```

你的 `sortByPriorityDescending` 函式應該使用兩個巢狀的 `for` 迴圈來執行排序。

```js
const loopMatches = sortByPriorityDescending.toString().match(/for\s*\(/g) || [];
assert.isAtLeast(loopMatches.length, 2);
```

你的 `sortByPriorityDescending` 函式（程式）應該將輸入的 `crew` 按 `priority` 以降冪排序。

```js
const crewSample = [
  { name: "A", priority: 2 },
  { name: "B", priority: 5 },
  { name: "C", priority: 3 }
];
sortByPriorityDescending(crewSample);
assert.deepEqual(crewSample.map(a => a.name), ["B", "C", "A"]);
```

你的 `sortByPriorityDescending` 函式應該直接就地排序輸入的 `crew` 陣列，而不是傳回新的陣列。

```js
const crewSample2 = [
  { name: "X", priority: 1 },
  { name: "Y", priority: 2 }
];
const result = sortByPriorityDescending(crewSample2);
assert.strictEqual(result, undefined);
assert.deepEqual(crewSample2.map(a => a.priority), [2, 1]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];

  return updatedCrew; 
}

const updatedSquad = swapCrewMembers(squad, 2, 5);

--fcc-editable-region--

--fcc-editable-region--

function getEVAReadyCrew(crew) {
  const eligible = [];
  for (const astronaut of crew) {
    if (astronaut.isEVAEligible) eligible.push(astronaut);
  }
  
  return eligible;
}
```
