---
id: 587d825a367417b2b2512c89
title: 퀵 정렬 구현하기
challengeType: 1
forumTopicId: 301615
dashedName: implement-quick-sort
---

# --description--

이번에는 중급 정렬 알고리즘인 퀵 정렬로 넘어가겠습니다. 퀵 정렬은 배열을 정렬하는 효율적이고 재귀적인 분할 정복 방법입니다. 우선은 배열에서 피벗 값을 선택합니다. 그런 다음 배열을 피벗 값보다 작은 값과 큰 값으로 이루어진 두 개의 하위 배열로 분할합니다. 이후 두 하위 배열에 대해 퀵 정렬 알고리즘을 재귀적으로 호출한 결과를 합칩니다. 이 과정은 빈 배열이나 한 개의 항목만 있는 배열이라는 기본 사례에 도달할 때까지 계속되며, 이때 배열을 반환합니다. 재귀 호출이 풀리면서 정렬된 배열이 반환됩니다.

퀵 정렬은 평균적으로 *O(nlog(n))* 성능을 제공하는 매우 효율적인 정렬 방법입니다. 또한 구현하기도 비교적 쉽습니다. 이러한 특성 덕분에 널리 사용되고 유용한 정렬 방법입니다.

**지시사항:** 정수 배열을 입력으로 받아 가장 작은 값부터 가장 큰 값 순서로 정렬된 정수 배열을 반환하는 `quickSort` 함수를 작성하세요. 피벗 값 선택이 중요하지만, 여기서는 어떤 피벗을 선택해도 괜찮습니다. 간단하게 첫 번째나 마지막 요소를 사용할 수 있습니다.

# --hints--

`quickSort`은 함수여야 합니다.

```js
assert(typeof quickSort == 'function');
```

`quickSort`는 정렬된 배열(작은 값부터 큰 값 순서)을 반환해야 합니다.

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    quickSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`quickSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])`는 순서만 바뀌었을 뿐 변경되지 않은 배열을 반환해야 합니다.

```js
assert.sameMembers(
  quickSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quickSort`는 내장 `.sort()` 메서드를 사용하지 않아야 합니다.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quickSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function quickSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function quickSort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quickSort(lesser), ...equal, ...quickSort(greater)];
  }
}
```
