---
id: 587d8259367417b2b2512c86
title: 삽입 정렬 구현하기
challengeType: 1
forumTopicId: 301613
dashedName: implement-insertion-sort
---

# --description--

다음으로 살펴볼 정렬 방법은 삽입 정렬입니다. 이 방법은 리스트 앞부분에 정렬된 배열을 점차 만들어 가면서 작동합니다. 먼저 첫 번째 요소로 정렬된 배열을 시작합니다. 그다음 다음 요소를 검사하여 정렬된 배열 안에서 올바른 위치가 될 때까지 뒤로 교환합니다. 리스트 끝에 도달할 때까지 계속해서 새로운 항목을 뒤로 교환하며 정렬된 부분에 삽입합니다. 이 알고리즘은 평균과 최악의 경우에 이차 시간 복잡도를 가집니다.

**지시사항:** 정수 배열을 입력으로 받아 이 정수들을 가장 작은 값부터 가장 큰 값 순으로 정렬된 배열을 반환하는 함수 `insertionSort`를 작성하세요.

# --hints--

`insertionSort`는 함수여야 합니다.

```js
assert.isFunction(insertionSort);
```

`insertionSort`는 정렬된 배열(작은 값에서 큰 값 순)을 반환해야 합니다.

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    insertionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`insertionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])`는 순서만 바뀌었을 뿐 나머지는 변경되지 않은 배열을 반환해야 합니다.

```js
assert.sameMembers(
  insertionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort([5, 4, 33, 2, 8])`는 `[2, 4, 5, 8, 33]`를 반환해야 합니다.

```js
assert.deepEqual(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort`는 내장 `.sort()` 메서드를 사용하지 않아야 합니다.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function insertionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
