---
id: 587d825c367417b2b2512c8f
title: 병합 정렬 구현하기
challengeType: 1
forumTopicId: 301614
dashedName: implement-merge-sort
---

# --description--

또 다른 흔한 중급 정렬 알고리즘은 병합 정렬입니다. 퀵 정렬처럼 병합 정렬도 분할 정복, 재귀 방식을 사용해 배열을 정렬합니다. 두 배열이 각각 정렬되어 있다면 두 배열을 정렬하는 것은 비교적 쉽다는 점을 활용합니다. 하지만 입력으로는 하나의 배열만 주어지므로, 어떻게 두 개의 정렬된 배열로 나누나요? 원본 입력을 재귀적으로 반씩 나누어 한 개의 요소만 가진 배열이 될 때까지 나누면 됩니다. 한 개 요소만 가진 배열은 자연스럽게 정렬된 상태이므로, 그다음에 병합을 시작할 수 있습니다. 이 병합 과정은 원본 배열을 나누었던 재귀 호출을 되돌리며 결국 모든 요소가 정렬된 최종 배열을 만듭니다. 병합 정렬의 단계는 다음과 같습니다:

**1)** 입력 배열을 재귀적으로 반씩 나누어 한 개 요소만 가진 하위 배열이 만들어질 때까지 나눕니다.

**2)** 각 정렬된 하위 배열을 병합하여 최종 정렬된 배열을 만듭니다.

병합 정렬은 시간 복잡도가 *O(nlog(n))*인 효율적인 정렬 방법입니다. 이 알고리즘은 성능이 좋고 구현이 비교적 쉬워 인기가 많습니다.

참고로, 여기서 다루는 마지막 정렬 알고리즘입니다. 하지만 나중에 트리 자료 구조 섹션에서 이진 힙을 사용해 구현하는 또 다른 효율적인 정렬 방법인 힙 정렬을 설명할 예정입니다.

**지시사항:** 정수 배열을 입력으로 받아 이 정수들을 오름차순으로 정렬한 배열을 반환하는 `mergeSort` 함수를 작성하세요. 이를 구현하는 좋은 방법은, 예를 들어 두 개의 정렬된 배열을 병합하는 역할을 하는 `merge` 함수와, 병합에 넣을 한 개 요소 배열을 만드는 재귀를 담당하는 `mergeSort` 함수를 작성하는 것입니다. 행운을 빕니다!

# --hints--

`mergeSort`는 함수여야 합니다.

```js
assert.isFunction(mergeSort);
```

`mergeSort`는 오름차순으로 정렬된 배열을 반환해야 합니다.

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    mergeSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`mergeSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])`는 순서만 바뀌었을 뿐 원본 배열과 같은 요소를 가진 배열을 반환해야 합니다.

```js
assert.sameMembers(
  mergeSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`mergeSort`는 내장 `.sort()` 메서드를 사용하지 않아야 합니다.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    mergeSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function mergeSort(array) {
  if (array.length === 1) {
    return array;
  } else {
    const splitIndex = Math.floor(array.length / 2);
    return merge(
      mergeSort(array.slice(0, splitIndex)),
      mergeSort(array.slice(splitIndex))
    );
  }

  // Merge two sorted arrays
  function merge(array1, array2) {
    let merged = [];
    while (array1.length && array2.length) {
      if (array1[0] < array2[0]) {
        merged.push(array1.shift());
      } else if (array1[0] > array2[0]) {
        merged.push(array2.shift());
      } else {
        merged.push(array1.shift(), array2.shift());
      }
    }

    // After looping ends, one array is empty, and other array contains only
    // values greater than all values in `merged`
    return [...merged, ...array1, ...array2];
  }
}

mergeSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]);
```
