---
id: 587d8259367417b2b2512c85
title: 선택 정렬 구현하기
challengeType: 1
forumTopicId: 301616
dashedName: implement-selection-sort
---

# --description--

여기서는 선택 정렬을 구현할 것입니다. 선택 정렬은 리스트에서 최솟값을 선택해 리스트의 첫 번째 값과 교환하는 방식으로 작동합니다. 그런 다음 두 번째 위치에서 시작해 남은 리스트 중 가장 작은 값을 선택하고 두 번째 요소와 교환합니다. 리스트 끝에 도달할 때까지 계속해서 요소를 반복하며 교환합니다. 이렇게 하면 리스트가 정렬됩니다. 선택 정렬은 모든 경우에 시간 복잡도가 이차 함수입니다.

**지시사항**: 정수 배열을 입력으로 받아 이 정수들을 최솟값부터 최댓값 순으로 정렬한 배열을 반환하는 `selectionSort` 함수를 작성하세요.

# --hints--

`selectionSort`는 함수여야 합니다.

```js
assert.isFunction(selectionSort);
```

`selectionSort`는 정렬된 배열(최솟값부터 최댓값 순)을 반환해야 합니다.

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}

assert.isTrue(
  isSorted(
    selectionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`selectionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])`는 순서만 바뀌었을 뿐 나머지는 변경되지 않은 배열을 반환해야 합니다.

```js
assert.sameMembers(
  selectionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort`는 내장 `.sort()` 메서드를 사용하지 않아야 합니다.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function selectionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
