---
id: 61abc7ebf3029b56226de5b6
title: 이진 검색 구현
challengeType: 1
forumTopicId: 487618
dashedName: implement-binary-search
---

# --description--

이진 검색은 정렬된 배열에서 요소를 찾기 위한 **O(log(n))** 효율 알고리즘입니다. 다음 단계로 진행합니다:

1. 정렬된 배열의 중간 `value`를 찾습니다. 만약 `value == target`라면 `true`를 반환합니다 (값을 찾았고 검색이 완료됨).
1. 중간 `value < target`라면 다음 비교에서 배열의 오른쪽 절반을 검색합니다.
1. 중간 `value > target`라면 다음 비교에서 배열의 왼쪽 절반을 검색합니다.
1. 배열 전체를 검색한 후에도 값이 없으면 `false`를 반환합니다 (배열을 모두 검색했으나 값이 없음).

보시다시피, 배열을 계속 반으로 나누기 때문에 log(n) 효율을 얻습니다. 이 도전 과제에서는 목표 값을 찾기까지의 과정을 보여주길 원합니다... 즉, 거쳐간 경로를 보여주세요!

# --instructions--

이진 검색 알고리즘을 배열에 구현하는 `binarySearch` 함수를 작성하세요. 목표 값을 찾기 위해 거친 경로(각 중간 값 비교)를 반환해야 합니다.

이 함수는 정렬된 정수 배열과 목표 값을 입력으로 받습니다. 목표 값을 찾을 때까지 원래 배열을 반으로 나누면서 찾은 중간 값을 (중위 순서로) 포함하는 배열을 반환합니다. 반환된 배열의 마지막 요소는 목표 값이어야 합니다. 값을 찾지 못하면 문자열 `Value Not Found`를 반환하세요.

예를 들어, `binarySearch([1,2,3,4,5,6,7], 5)`는 `[4,6,5]`를 반환합니다.

이 도전 과제에서는 나눗셈을 할 때 반드시 `Math.floor()`를 사용해야 합니다: `Math.floor(x/2)`. 이렇게 하면 일관되고 테스트 가능한 경로가 나옵니다.

**참고:** 테스트에는 다음 배열이 사용됩니다:

```js
const testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
```

# --hints--

`binarySearch`는 함수여야 합니다.

```js
assert.isFunction(binarySearch);
```

`binarySearch(testArray, 0)`는 `[13, 5, 2, 0]`를 반환해야 합니다.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 0), [13, 5, 2, 0]);
```

`binarySearch(testArray, 1)`는 `[13, 5, 2, 0, 1]`를 반환해야 합니다.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 1), [13, 5, 2, 0, 1]);
```


`binarySearch(testArray, 2)`는 `[13, 5, 2]`를 반환해야 합니다.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 2), [13, 5, 2]);
```

`binarySearch(testArray, 6)`는 문자열 `Value Not Found`를 반환해야 합니다.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.strictEqual(binarySearch(_testArray, 6), 'Value Not Found');
```

`binarySearch(testArray, 11)`는 `[13, 5, 10, 11]`를 반환해야 합니다.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 11), [13, 5, 10, 11])
```

`binarySearch(testArray, 13)`는 `[13]`를 반환해야 합니다.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 13), [13]);
```

`binarySearch(testArray, 70)`는 `[13, 19, 22, 49, 70]`를 반환해야 합니다.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 70), [13, 19, 22, 49, 70]);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let arrayPath = [];
  return arrayPath;
}
```



# --solutions--

```js
let binarySearch = (searchList, value) => {
  let arrayPath = [];

  // set initial L - M - R
  let left = 0;
  let right = searchList.length - 1;
  let middle = Math.floor(right / 2);

  // if first comparison finds value
  if (searchList[middle] == value) {
    arrayPath.push(searchList[middle]);
    return arrayPath;
  }

  while (searchList[middle] !== value) {
    // add to output array
    arrayPath.push(searchList[middle]);

    // not found
    if (right < left) {
      return 'Value Not Found';
    }
    // value is in left or right portion of array
    // update L - M - R
    if (searchList[middle] > value) {
      right = middle - 1;
      middle = left + Math.floor((right - left) / 2);
    } else {
      left = middle + 1;
      middle = left + Math.floor((right - left) / 2);
    }

    // if found update output array and exit
    if (searchList[middle] == value) {
      arrayPath.push(searchList[middle]);

      break;
    }
  }
  return arrayPath;
};
```
