---
id: a3f503de51cf954ede28891d
title: 대칭 차집합 찾기
challengeType: 1
forumTopicId: 301611
dashedName: find-the-symmetric-difference
---

# --description--

수학 용어인 <dfn>대칭 차집합</dfn> (`△` 또는 `⊕`)은 두 집합 중 어느 한쪽에는 속하지만 두 집합 모두에는 속하지 않는 원소들의 집합입니다. 예를 들어, 집합 `A = {1, 2, 3}`와 `B = {2, 3, 4}`에 대해 `A △ B = {1, 4}`입니다.

대칭 차집합은 이항 연산으로, 두 개의 원소에만 작용합니다. 따라서 *세 개*의 원소(`A △ B △ C`)에 대칭 차집합 연산이 포함된 표현식을 평가하려면 한 번에 하나씩 연산을 완료해야 합니다. 따라서 위의 집합 `A`와 `B`, 그리고 `C = {2, 3}`에 대해 `A △ B △ C = (A △ B) △ C = {1, 4} △ {2, 3} = {1, 2, 3, 4}`입니다.

# --instructions--

두 개 이상의 배열을 인수로 받아 그들의 대칭 차집합을 배열로 반환하는 함수를 만드세요. 반환하는 배열에는 중복 없이 고유한 값만 포함되어야 합니다.

# --hints--

`sym([1, 2, 3], [5, 2, 1, 4])`은 `[3, 4, 5]`을 반환해야 합니다.

```js
assert.sameMembers(sym([1, 2, 3], [5, 2, 1, 4]), [3, 4, 5]);
```

`sym([1, 2, 3], [5, 2, 1, 4])`에는 세 개의 원소만 포함되어야 합니다.

```js
assert.equal(sym([1, 2, 3], [5, 2, 1, 4]).length, 3);
```

`sym([1, 2, 3, 3], [5, 2, 1, 4])`은 `[3, 4, 5]`를 반환해야 합니다.

```js
assert.sameMembers(sym([1, 2, 3, 3], [5, 2, 1, 4]), [3, 4, 5]);
```

`sym([1, 2, 3, 3], [5, 2, 1, 4])`에는 세 개의 원소만 포함되어야 합니다.

```js
assert.equal(sym([1, 2, 3, 3], [5, 2, 1, 4]).length, 3);
```

`sym([1, 2, 3], [5, 2, 1, 4, 5])`은 `[3, 4, 5]`을 반환해야 합니다.

```js
assert.sameMembers(sym([1, 2, 3], [5, 2, 1, 4, 5]), [3, 4, 5]);
```

`sym([1, 2, 3], [5, 2, 1, 4, 5])`에는 세 개의 원소만 포함되어야 합니다.

```js
assert.equal(sym([1, 2, 3], [5, 2, 1, 4, 5]).length, 3);
```

`sym([1, 2, 5], [2, 3, 5], [3, 4, 5])`은 `[1, 4, 5]`을 반환해야 합니다.

```js
assert.sameMembers(sym([1, 2, 5], [2, 3, 5], [3, 4, 5]), [1, 4, 5]);
```

`sym([1, 2, 5], [2, 3, 5], [3, 4, 5])`에는 세 개의 원소만 포함되어야 합니다.

```js
assert.equal(sym([1, 2, 5], [2, 3, 5], [3, 4, 5]).length, 3);
```

`sym([1, 1, 2, 5], [2, 2, 3, 5], [3, 4, 5, 5])`은 `[1, 4, 5]`을 반환해야 합니다.

```js
assert.sameMembers(sym([1, 1, 2, 5], [2, 2, 3, 5], [3, 4, 5, 5]), [1, 4, 5]);
```

`sym([1, 1, 2, 5], [2, 2, 3, 5], [3, 4, 5, 5])`에는 세 개의 원소만 포함되어야 합니다.

```js
assert.equal(sym([1, 1, 2, 5], [2, 2, 3, 5], [3, 4, 5, 5]).length, 3);
```

`sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3])`은 `[2, 3, 4, 6, 7]`을 반환해야 합니다.

```js
assert.sameMembers(
  sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3]),
  [2, 3, 4, 6, 7]
);
```

`sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3])`에는 다섯 개의 원소만 포함되어야 합니다.

```js
assert.equal(
  sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3]).length,
  5
);
```

`sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3], [5, 3, 9, 8], [1])`은 `[1, 2, 4, 5, 6, 7, 8, 9]`을 반환해야 합니다.

```js
assert.sameMembers(
  sym(
    [3, 3, 3, 2, 5],
    [2, 1, 5, 7],
    [3, 4, 6, 6],
    [1, 2, 3],
    [5, 3, 9, 8],
    [1]
  ),
  [1, 2, 4, 5, 6, 7, 8, 9]
);
```

`sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3], [5, 3, 9, 8], [1])`에는 여덟 개의 원소만 포함되어야 합니다.

```js
assert.equal(
  sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3], [5, 3, 9, 8], [1])
    .length,
  8
);
```

# --seed--

## --seed-contents--

```js
function sym(args) {
  return args;
}

sym([1, 2, 3], [5, 2, 1, 4]);
```

# --solutions--

```js
function sym() {
  var arrays = [].slice.call(arguments);
  return arrays.reduce(function (symDiff, arr) {
    return symDiff.concat(arr).filter(function (val, idx, theArr) {
      return theArr.indexOf(val) === idx
        && (symDiff.indexOf(val) === -1 || arr.indexOf(val) === -1);
    });
  });
}
sym([1, 2, 3], [5, 2, 1, 4]);
```
