---
id: a3f503de51cfab748ff001aa
title: 쌍으로 묶기
challengeType: 1
forumTopicId: 301617
dashedName: pairwise
---

# --description--

배열 `arr`이 주어졌을 때, 두 번째 인수 `arg`과 합이 같은 요소 쌍을 찾아 그 인덱스 합을 반환하세요.

같은 숫자 요소를 가진 여러 쌍을 사용할 수 있지만 인덱스는 달라야 합니다. 각 쌍은 가능한 가장 낮은 인덱스를 사용해야 합니다. 한 번 사용된 요소는 다른 요소와 쌍을 이루기 위해 다시 사용할 수 없습니다. 예를 들어 `pairwise([1, 1, 2], 3)`는 인덱스 1의 1 대신 인덱스 0의 1을 사용하여 쌍 `[2, 1]`을 만드는데, 그 이유는 0+2가 1+2보다 작기 때문입니다.

예를 들어 `pairwise([7, 9, 11, 13, 15], 20)`는 `6`를 반환합니다. 합이 20인 쌍은 `[7, 13]`와 `[9, 11]`입니다. 배열과 각 요소의 인덱스를 다음과 같이 나타낼 수 있습니다.

<div style='margin-left: 2em;'>

| 인덱스 | 0 | 1 | 2  | 3  | 4  |
| ------ | - | - | -- | -- | -- |
| 값     | 7 | 9 | 11 | 13 | 15 |

</div>

아래에서 해당 인덱스를 더해 보겠습니다.

<div style='margin-left: 2em;'>

7 + 13 = 20 → 인덱스 0 + 3 = 3  
9 + 11 = 20 → 인덱스 1 + 2 = 3  
3 + 3 = 6 → `6` 반환

</div>

# --hints--

`pairwise([1, 4, 2, 3, 0, 5], 7)`는 11을 반환해야 합니다.

```js
assert.deepEqual(pairwise([1, 4, 2, 3, 0, 5], 7), 11);
```

`pairwise([1, 3, 2, 4], 4)`는 1을 반환해야 합니다.

```js
assert.deepEqual(pairwise([1, 3, 2, 4], 4), 1);
```

`pairwise([1, 1, 1], 2)`는 1을 반환해야 합니다.

```js
assert.deepEqual(pairwise([1, 1, 1], 2), 1);
```

`pairwise([0, 0, 0, 0, 1, 1], 1)`는 10을 반환해야 합니다.

```js
assert.deepEqual(pairwise([0, 0, 0, 0, 1, 1], 1), 10);
```

`pairwise([], 100)`는 0을 반환해야 합니다.

```js
assert.deepEqual(pairwise([], 100), 0);
```

# --seed--

## --seed-contents--

```js
function pairwise(arr, arg) {
  return arg;
}

pairwise([1,4,2,3,0,5], 7);
```

# --solutions--

```js
function pairwise(arr, arg) {
  var sum = 0;
  arr.forEach(function(e, i, a) {
    if (e != null) {
      var diff = arg-e;
      a[i] = null;
      var dix = a.indexOf(diff);
      if (dix !== -1) {
        sum += dix;
        sum += i;
        a[dix] = null;
      }
    }
  });
  return sum;
}

pairwise([1,4,2,3,0,5], 7);
```
