---
id: a7bf700cd123b9a54eef01d5
title: 반복 없이 개수 세기
challengeType: 1
forumTopicId: 16037
dashedName: no-repeats-please
---

# --description--

연속해서 반복되는 문자가 없는 제공된 문자열의 모든 순열 개수를 반환하세요. 제공된 문자열의 모든 문자는 각각 고유하다고 가정합니다.

예를 들어 `aab`는 총 6개의 순열(`aab`, `aab`, `aba`, `aba`, `baa`, `baa`)이 있지만, 그중에서 연속된 같은 문자(`aba`)가 반복되지 않는 순열은 2개(`aba`와 `a`)이므로 2를 반환해야 합니다.

# --hints--

`permAlone("aab")`는 숫자를 반환해야 합니다.

```js
assert.isNumber(permAlone('aab'));
```

`permAlone("aab")`는 2를 반환해야 합니다.

```js
assert.strictEqual(permAlone('aab'), 2);
```

`permAlone("aaa")`는 0을 반환해야 합니다.

```js
assert.strictEqual(permAlone('aaa'), 0);
```

`permAlone("aabb")`는 8을 반환해야 합니다.

```js
assert.strictEqual(permAlone('aabb'), 8);
```

`permAlone("abcdefa")`는 3600을 반환해야 합니다.

```js
assert.strictEqual(permAlone('abcdefa'), 3600);
```

`permAlone("abfdefa")`는 2640을 반환해야 합니다.

```js
assert.strictEqual(permAlone('abfdefa'), 2640);
```

`permAlone("zzzzzzzz")`는 0을 반환해야 합니다.

```js
assert.strictEqual(permAlone('zzzzzzzz'), 0);
```

`permAlone("a")`는 1을 반환해야 합니다.

```js
assert.strictEqual(permAlone('a'), 1);
```

`permAlone("aaab")`는 0을 반환해야 합니다.

```js
assert.strictEqual(permAlone('aaab'), 0);
```

`permAlone("aaabb")`는 12를 반환해야 합니다.

```js
assert.strictEqual(permAlone('aaabb'), 12);
```

# --seed--

## --seed-contents--

```js
function permAlone(str) {
  return str;
}

permAlone('aab');
```

# --solutions--

```js
function permAlone(str) {
  return permuter(str).filter(function(perm) {
    return !perm.match(/(.)\1/g);
  }).length;
}

function permuter(str) {
  // http://staff.roguecc.edu/JMiller/JavaScript/permute.html
  //permArr: Global array which holds the list of permutations
  //usedChars: Global utility array which holds a list of "currently-in-use" characters
  var permArr = [], usedChars = [];
  function permute(input) {
    //convert input into a char array (one element for each character)
    var i, ch, chars = input.split("");
    for (i = 0; i < chars.length; i++) {
      //get and remove character at index "i" from char array
      ch = chars.splice(i, 1);
      //add removed character to the end of used characters
      usedChars.push(ch);
      //when there are no more characters left in char array to add, add used chars to list of permutations
      if (chars.length === 0) permArr[permArr.length] = usedChars.join("");
      //send characters (minus the removed one from above) from char array to be permuted
      permute(chars.join(""));
      //add removed character back into char array in original position
      chars.splice(i, 0, ch);
      //remove the last character used off the end of used characters array
      usedChars.pop();
    }
  }
  permute(str);
  return permArr;
}

permAlone('aab');
```
