---
id: 68af273960e009597dd861e2
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Quieres una visualización legible para depuración e impresión. Crea un método `__str__` que devuelva una cadena con el formato `title (year) - duration min, director` (donde `title`, `year`, `duration` y `director` deben ser reemplazados por los valores de los atributos correspondientes) para que imprimir una película dé un buen resumen.

# --hints--

Tu clase `Movie` debería tener un método `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

Tu método `__str__` debe devolver una cadena en el formato `title (year) - duration min, director`, donde `title`, `year`, `duration` y `director` deben ser reemplazados con los valores de los atributos correspondientes.

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    
--fcc-editable-region--
```
