---
id: 68af272f60e009597dd861e1
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, vas a crear un catálogo de medios. Comienza definiendo una clase llamada `Movie`.

Dale a tu clase un método `__init__` con los siguientes parámetros: `self`, `title`, `year`, `director` y `duration`. Dentro del método `__init__`, asigna los parámetros `title`, `year`, `director` y `duration` a atributos de instancia con el mismo nombre.

# --hints--

Debes crear una clase llamada `Movie`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

Tu clase `Movie` debe tener un método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

Tu método `__init__` debe tener los siguientes parámetros: `self`, `title`, `year`, `director` y `duration`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

Debes asignar `title` a `self.title` dentro de tu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

Debes asignar `year` a `self.year` dentro de tu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

Debes asignar `director` a `self.director` dentro de tu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

Debes asignar `duration` a `self.duration` dentro de tu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
