---
id: 68af277560e009597dd861e8
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Ahora elimina tu línea `print(movie1)`. Luego, coloca tu asignación `movie1` dentro de un bloque `try` y crea una cláusula `except` que capture `ValueError as e` y use una f-string para imprimir `Validation Error:` seguido de un espacio y el mensaje de error.

Siéntete libre de cambiar los argumentos pasados para instanciar `movie1` y ver los diferentes mensajes de error.

# --hints--

No debes tener `print(movie1)` en tu código.

```js
({
  test: () => runPython(`assert not _Node(_code).has_call("print(movie1)")`)
})
```

Debes tener un bloque `try` con `movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)` dentro.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().has_stmt("movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)")`)
})
```

Debes tener una cláusula `except` que capture `ValueError as e`.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].has_except("ValueError", "e")`)
})
```

Debes usar una f-string para imprimir `Validation Error:` seguido de un espacio y el mensaje de error dentro del bloque `except`.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_except("ValueError", "e").has_call("print(f'Validation Error: {e}')")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
--fcc-editable-region--
```
