---
id: 68af274560e009597dd861e3
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Ahora, crea una instancia de `Movie` pasando argumentos de tu elección y asígnala a una variable llamada `movie1`. Luego, imprime `movie1` en la consola.

# --hints--

Debes tener una variable llamada `movie1`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

Debes asignar una instancia de `Movie` a la variable `movie1`.

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

Debes imprimir `movie1` en la consola.

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
```
