---
id: 68af274e60e009597dd861e4
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

En los pasos anteriores, creaste un objeto básico de película. Ahora, agregarás algo de validación para asegurarte de que las películas no puedan ser instanciadas con atributos inválidos.

Al comienzo del método `__init__`, crea una estructura `if` que verifique si el `title` está vacío o contiene solo espacios en blanco (puedes usar el método `strip` para construir la condición).

Si es así, lanza un `ValueError` con el mensaje `Title cannot be empty`.

# --hints--

Debes lanzar un `ValueError` con el mensaje `Title cannot be empty` cuando `title` esté vacío o contenga solo espacios en blanco.

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

No debes lanzar un `ValueError` cuando `title` es una cadena válida.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
