---
id: 68b6d5244e6ce639e52fd1e7
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

¿Qué pasa si quieres representar un tipo diferente de elemento multimedia, digamos una serie de televisión? Podrías crear otra clase y hacer que herede de `Movie` para tener acceso a todos sus métodos y atributos.

Crea una clase vacía llamada `TVSeries` y haz que herede de `Movie`. Aquí tienes cómo hacer que una clase `Child` herede de una clase `Parent`.

```py
class Parent:
    ...

class Child(Parent):
    ...
```

# --hints--

Debes tener una clase llamada `TVSeries`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("TVSeries")`)
})
```

Tu clase `TVSeries` debería heredar de `Movie`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").inherits_from("Movie")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
