---
id: 68b702892910ab7b495de081
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Necesitarás una forma adecuada de representar visualmente el catálogo. Así que crea un método `__str__` en tu clase `MediaCatalogue`.

Si la lista `items` está vacía, asegúrate de que el método `__str__` devuelva `Media Catalogue (empty)`.

# --hints--

Debes tener un método `__str__` en tu clase `MediaCatalogue`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").has_function("__str__")`)
})
```

Tu método `__str__` debería devolver `Media Catalogue (empty)` cuando `self.items` esté vacío.

```js
({
  test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    print(MediaCatalogue())
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Media Catalogue (empty)" in output
  `)
})
```

Tu método `__str__` no debe devolver `Media Catalogue (empty)` cuando `self.items` no está vacío.

```js
({
  test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
    _catalogue = MediaCatalogue()
    _catalogue.add(Movie('The Matrix', 1999, 'The Wachowskis', 136))
    try:
      print(_catalogue)
    except TypeError:
      pass
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Media Catalogue (empty)" not in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
