---
id: 68b804f3be853a355f055fef
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

La función `super` se usa para referirse a la clase padre sin nombrarla explícitamente. Por ejemplo, `super().__init__(arg1, arg2, ...)` llama al método `__init__` definido en la clase padre.

Por el momento, has sobrescrito efectivamente el método `__init__` heredado de `Movie`. Esto significa que un objeto `TVSeries` tendrá solo un atributo `seasons` y un atributo `total_episodes`. Esto no es lo que quieres.

Para mantener todo el código definido en el `__init__` de la clase `Movie`, usa `super` para llamar al `__init__` del padre. Ten en cuenta que necesitarás pasar los argumentos requeridos por el método. En este caso, `title`, `year`, `director` y `duration`.

# --hints--

Debes llamar a `super().__init__(title, year, director, duration)` dentro de tu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("super().__init__(title, year, director, duration)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        
        self.seasons = seasons
        self.total_episodes = total_episodes
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
