---
id: 68b95fa5f6e575879105e74f
title: Paso 32
challengeType: 20
dashedName: step-32
---

# --description--

Ahora que la clase `MediaError` está completa, modifica la sentencia `if` en el método `add`. Reemplaza el `TypeError` con un `MediaError` usando el mismo mensaje.

Además, pasa `media_item` como el segundo argumento después del mensaje. Esto te permitirá almacenar el elemento que causa el error en el objeto de excepción.

# --hints--

Debes lanzar un `MediaError` con el mensaje `Only Movie or TVSeries instances can be added` cuando `media_item` no sea una instancia de `Movie` o de sus subclases.

```js
({
  test: () => runPython(`
  _msg = 'Only Movie or TVSeries instances can be added'
  _expected_error = f"Expected to raise 'MediaError: {_msg}'"
  _c = MediaCatalogue()
  try:    
    _c.add("Not a movie")
  except MediaError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    _c.add({})
  except MediaError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Debes pasar `media_item` como el segundo argumento a `MediaError`.

```js
({
  test: () => runPython(`
  _c = MediaCatalogue()
  _obj = "Not a movie"
  try:    
    _c.add(_obj)
  except MediaError as e:
    assert e.obj == _obj
  else:
    assert False
  `)
})
```

No debes generar un `MediaError` cuando `media_item` es un elemento válido.

```js
({
    test: () => runPython(`
    _c = MediaCatalogue()
    try:
      _c.add(Movie('The Matrix', 1999, 'The Wachowskis', 136))
    except MediaError:
      assert False, 'Expected not to raise a MediaError with valid media_item'
    except TypeError:
      pass

    try:
      _c.add(TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182))
    except MediaError:
      assert False, 'Expected not to raise a MediaError with valid media_item'
    except TypeError:
      pass
    `)
})
```

Tu método `add` no debería generar un `TypeError`.

```js
({
    test: () => runPython(`
    _c = MediaCatalogue()
    try:
      _c.add(Movie('The Matrix', 1999, 'The Wachowskis', 136))
    except TypeError:
      assert False, 'Expected not to raise a TypeError from add method'

    try:
      _c.add(TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182))
    except TypeError:
      assert False, 'Expected not to raise a TypeError from add method'

    try:
      _c.add("Not a movie")
    except TypeError:
      assert False, 'Expected not to raise a TypeError from add method'
    except MediaError:
      pass
    `)
})
```

# --seed--

## --seed-contents--

```py
class MediaError(Exception):
    """Custom exception for media-related errors."""

    def __init__(self, message, obj):
        super().__init__(message)
        self.obj = obj

class Movie:
    """Parent class representing a movie."""
    
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    """Child class representing an entire TV series."""

    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)

        if seasons < 1:
            raise ValueError('Seasons must be 1 or greater')
        if total_episodes < 1:
            raise ValueError('Total episodes must be 1 or greater')
        
        self.seasons = seasons
        self.total_episodes = total_episodes

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.seasons} seasons, {self.total_episodes} episodes, {self.duration} min avg, {self.director}'

class MediaCatalogue:
    """A catalogue that can store different types of media items."""

    def __init__(self):
        self.items = []

    def add(self, media_item):
--fcc-editable-region--
        if not isinstance(media_item, Movie):
            raise TypeError('Only Movie or TVSeries instances can be added')
--fcc-editable-region--
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)

    series1 = TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182)
    catalogue.add(series1)
    series2 = TVSeries('Breaking Bad', 2008, 'Vince Gilligan', 47, 5, 62)
    catalogue.add(series2)
    
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
