---
id: 655cd899f8de09431eabb40c
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Projekt lernen Sie Datenstrukturen kennen, indem Sie den Merge-Sort-Algorithmus erstellen.

Dies ist ein Sortieralgorithmus, der das Prinzip „Teile und herrsche“ verwendet, um Datensammlungen zu sortieren. Das bedeutet, er „teilt“ eine Sammlung in kleinere Unterteile auf, „beherrscht“ die Unterteile, indem er sie unabhängig voneinander sortiert, und führt dann die sortierten Unterteile zusammen.

Erstellen Sie eine Funktion namens `merge_sort`. Diese Funktion übernimmt die Aufgabe, eine Liste von Zahlen zu sortieren.

Verwenden Sie das Schlüsselwort `pass` im Funktionskörper.

# --hints--

Sie sollten eine Funktion namens `merge_sort` deklarieren. Vergessen Sie nicht, das Schlüsselwort `pass` im Funktionskörper zu verwenden.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
