---
id: 6564aee9c077774ea49c3faf
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Der Merge-Sort-Algorithmus führt hauptsächlich drei Aktionen aus:

- Teilen einer unsortierten Folge von Elementen in Unterteile
- Sortieren der Elemente in den Unterteilen
- Zusammenführen der sortierten Unterteile

Das oben Genannte geschieht rekursiv, bis die Unterteile zu der vollständig sortierten Folge zusammengeführt sind. Beginnen Sie mit dem Teilen der Folge.

Ersetzen Sie zuerst das Schlüsselwort `pass` durch eine Variable `middle_point`. Verwenden Sie den ganzzahligen Divisionsoperator (`//`), um die Länge der `array`-Liste durch 2 zu teilen, und weisen Sie das Ergebnis Ihrer neuen `middle_point`-Variable zu. Denken Sie daran, Ihren Code einzurücken.

# --hints--

Sie sollten `pass` nicht im Funktionskörper haben.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.notMatch(function_body, /^\s{4}pass/m);
  }
})
```

Sie sollten eine `middle_point`-Variable deklarieren und `len(array) // 2` zuweisen.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
