---
id: 6564a9fe51964c229d5b7f4c
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Sie benötigen einen Parameter, der die zu sortierende Datensammlung bezeichnet. Erstellen Sie einen Parameter namens `array` in der `merge_sort`-Funktion.

# --hints--

Ihre `merge_sort`-Funktion sollte einen einzelnen Parameter annehmen: `array`.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
