---
id: 656639771fed09b5c0e8fe71
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Nachdem Sie die `array`-Liste in zwei separate Listen aufgeteilt haben, teilen Sie jede Liste weiter auf, bis jedes Element in seiner eigenen Liste steht. Eine Liste mit nur einer Zahl ist immer sortiert.

Rufen Sie dazu rekursiv `merge_sort` innerhalb Ihrer Funktion auf. Übergeben Sie vorerst kein Argument.

# --hints--

Sie sollten `merge_sort` am Ende Ihres Funktionskörpers aufrufen.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
