---
id: 656627b47bd2d2a4afbc945d
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Im vorherigen Schritt haben Sie den Mittelpunkt ermittelt. Sie können ihn verwenden, um `array` in zwei Teile zu teilen und jeden Teil neuen Variablen zuzuweisen.

Verwenden Sie die Slice-Syntax, um die linke Hälfte von `array` zu extrahieren und sie einer Variablen namens `left_part` zuzuweisen.


# --hints--

Sie sollten eine Variable namens `left_part` in Ihrer `merge_sort`-Funktion haben.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

Sie sollten `array[:middle_point]` der Variablen `left_part` zuweisen.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
