---
id: 656702f8b4cbd8cbf0a433c6
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Erstellen Sie innerhalb Ihrer Funktion eine `while`-Schleife, die ein Element in `left_part` mit einem Element in `right_part` vergleicht und das kleinere Element in die Hauptliste `array` zusammenführt.

Erstellen Sie zwei Bedingungen für die Schleife: eine, die überprüft, ob `left_array_index` kleiner als die Länge von `left_part` ist, und eine andere, die überprüft, ob `right_array_index` kleiner als die Länge von `right_part` ist.

# --hints--

Ihre `while`-Schleife sollte diese Bedingungen haben: `left_array_index < len(left_part) and right_array_index < len(right_part)`. Vergessen Sie nicht, das Schlüsselwort `pass` im Schleifenkörper zu verwenden.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

Verwenden Sie das Schlüsselwort `pass` im Schleifenkörper

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
