---
id: 656709e50ed928da35260276
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

In einem vorherigen Schritt haben Sie das Element in der `left_part` der `array`-Liste zugewiesen, weil es kleiner war. Aber das ist nicht immer der Fall. In manchen Vergleichsfällen könnte das Element rechts kleiner sein.

Erstellen Sie eine `else`-Klausel, die ausgeführt wird, wenn der `left_part`-Index nicht kleiner als der `right_part`-Index ist.

Weisen Sie innerhalb des `else`-Blocks `right_part[right_array_index]` `array[sorted_index]` zu.

# --hints--

Sie sollten `array[sorted_index] = right_part[right_array_index]` im Körper Ihrer `else`-Klausel haben.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /else:\s*[^}]*array\s*\[\s*sorted_index\s*\]\s*=\s*right_part\s*\[\s*right_array_index\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
