---
id: 65772ef923f922cd720e5008
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Übergeben Sie `left_part` als Argument an den `merge_sort()`-Aufruf, den Sie im letzten Schritt hinzugefügt haben.

# --hints--

Sie sollten `left_part` als Argument an Ihren `merge_sort()`-Aufruf übergeben.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*left_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort()
--fcc-editable-region--
```
