---
id: 657f3a661730891aa64f3568
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Verwenden Sie die Slice-Syntax, um die rechte Hälfte von `array` zu extrahieren und weisen Sie sie einer Variablen namens `right_part` zu.

# --hints--

Sie sollten eine Variable namens `right_part` haben

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_part\s*(?!=)/m);
  }
})
```

Sie sollten `array[middle_point:]` der Variablen `right_part` zuweisen

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_part\s*=\s*array\s*\[\s*middle_point\s*:\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
