---
id: 68e4276483b0c71658aab1d3
title: Paso 14
challengeType: 1
dashedName: step-14
---

# --description--

Como recuerdas de lecciones y talleres anteriores, puedes extraer porciones de una cadena usando el método `slice()`.

Aquí tienes un recordatorio:

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

Cambia el valor de `camelCasedVersion` de una cadena vacía al resultado de usar el método `slice()` en la variable `lowercaseWord`. Pasa los números `0` y `5` que representan los índices de inicio y fin para el método `slice()`.

Ahora deberías ver la palabra `"camel"` en la consola.

# --hints--

Tu variable `camelCasedVersion` no debe ser una cadena vacía.

```js
assert.isNotEmpty(camelCasedVersion);
```

Debes asignar `lowercaseWord.slice(0, 5)` a la variable `camelCasedVersion`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

NO debes asignar la cadena literal `"camel"` a la `camelCasedVersion`. Por favor, relee las instrucciones para el enfoque correcto.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
