---
id: 68e4276483b0c71658aab1ce
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Habrá ocasiones en las que solo querrás eliminar los espacios en blanco al inicio de la cadena. Ahí es donde el método `trimStart()` resulta útil.

Aquí hay un ejemplo:

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

Recuerda que las cadenas son inmutables. Este método creará una nueva cadena con el resultado. No modificará la original.

Crea una variable `trimmedStart` y asígnale el resultado de usar el método `trimStart()` en la variable `userInput`.

# --hints--

Deberías tener una variable `trimmedStart`.

```js
assert.exists(trimmedStart);
```

Tu variable `trimmedStart` debe ser una cadena.

```js
assert.isString(trimmedStart);
```

Debes asignar `userInput.trimStart()` a tu variable `trimmedStart`.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
