---
id: 68e4276483b0c71658aab1d1
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

Si necesitas crear una nueva cadena donde todas las letras estén en minúsculas, puedes usar el método `toLowerCase()`.  

Comienza creando una variable llamada `lowerCaseInput` y asígnale el resultado de usar el método `toLowerCase()` en la variable `cleanedInput`. 

Luego registra la cadena `"Result of using the toLowerCase() method:"`.

Finalmente, agrega otro `console.log()` para la variable `lowerCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes tener una variable `lowerCaseInput`.

```js
assert.exists(lowerCaseInput);
```

Tu variable `lowerCaseInput` debe ser una cadena.

```js
assert.isString(lowerCaseInput);
```

Debes asignar `cleanedInput.toLowerCase()` a tu variable `lowerCaseInput`.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

Debes registrar la cadena `"Result of using the toLowerCase() method:"` en la consola.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

Debes registrar la variable `lowerCaseInput` en la consola.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
