---
id: 68e4276483b0c71658aab1cd
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Ahora es momento de ver el resultado.

Comienza agregando un `console.log()` con la cadena `"Result of trimming whitespace from both ends:"`.

Luego, debajo de ese `console.log()`, agrega otro `console.log()` con la variable `cleanedInput`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes registrar la cadena `"Result of trimming whitespace from both ends:"` en la consola.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

Debes registrar la variable `cleanedInput` en la consola.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
