---
id: 68e4276483b0c71658aab1ca
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, practicarás trabajando con varios métodos de cadenas usados para formatear cadenas.

Comienza creando una variable llamada `userInput` y asígnale el siguiente valor: `"   Hello World!   "`.

# --hints--

Debes tener una variable `userInput`.

```js
assert.exists(userInput);
```

Tu variable `userInput` debe ser una cadena.

```js
assert.isString(userInput);
```

Debes asignar `"   Hello World!   "` a tu variable `userInput`. Asegúrate de que haya tres espacios antes de la `H` y tres espacios después del `!`.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
