---
id: 68e4276483b0c71658aab1d2
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

Habrá ocasiones en las que necesitarás usar múltiples métodos para crear una nueva cadena. En los próximos pasos, usarás algunos de los métodos que has aprendido hasta ahora para crear una nueva cadena que use la convención de nombres camelCase.

Comienza creando una variable llamada `lowercaseWord` y asígnale la cadena `"camelcase"`. 

Debajo de esa variable `lowercaseWord`, crea otra variable llamada `camelCasedVersion` y asígnale una cadena vacía por ahora.

Debajo de esa variable `camelCasedVersion`, agrega un `console.log()` con la cadena `"Camel cased version:"` seguido de otro `console.log()` para la variable `camelCasedVersion`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes tener una variable llamada `lowercaseWord`.

```js
assert.exists(lowercaseWord);
```

Tu `lowercaseWord` debe ser una cadena.

```js
assert.isString(lowercaseWord);
```

Debes asignar la cadena `"camelcase"` a tu variable `lowercaseWord`.

```js
assert.equal(lowercaseWord, "camelcase");
```

Debes tener una variable `camelCasedVersion`.

```js
assert.exists(camelCasedVersion);
```

Tu variable `camelCasedVersion` debería ser una cadena.

```js
assert.isString(camelCasedVersion);
```

Debes asignar una cadena vacía a tu variable `camelCasedVersion`.

```js
assert.isEmpty(camelCasedVersion);
```

Debes registrar la cadena `"Camel cased version:"` en la consola.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

Debes registrar la variable `camelCasedVersion` en la consola.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
