---
id: 68e4276483b0c71658aab1d0
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Habrá ocasiones en las que necesitarás crear una nueva cadena con un formato de mayúsculas y minúsculas diferente. 

Para crear una nueva cadena con todas las letras convertidas a mayúsculas, puedes usar el método `toUpperCase()`.

Aquí hay un ejemplo:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

Crea una variable llamada `upperCaseInput` y asígnale el resultado de usar el método `toUpperCase()` en la variable `cleanedInput`.

# --hints--

Debes tener una variable `upperCaseInput`.

```js
assert.exists(upperCaseInput);
```

Tu variable `upperCaseInput` debe ser una cadena.

```js
assert.isString(upperCaseInput);
```

Debes asignar `cleanedInput.toUpperCase()` a tu variable `upperCaseInput`.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
