---
id: 68e4276483b0c71658aab1cc
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

El valor de tu `userInput` tiene algunos espacios adicionales en ambos extremos de la cadena. Para crear una nueva cadena con esos espacios en blanco eliminados, puedes usar el método `trim()`.  

Aquí hay un ejemplo:

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

Crea una variable llamada `cleanedInput` y asígnale el resultado de usar el método `trim()` en la variable `userInput`.

# --hints--

Debes tener una variable `cleanedInput`.

```js
assert.exists(cleanedInput);
```

Tu variable `cleanedInput` debe ser una cadena.

```js
assert.isString(cleanedInput);
```

Debes asignar `userInput.trim()` a tu variable `cleanedInput`.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
