---
id: 68ed6079fe0e3e98b94a1b75
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

A continuación, agrega un `console.log()` con la cadena `"Original input:"`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes registrar la cadena `"Original input:"` en la consola.

```js
assert.equal(getLogs()[0], "Original input:");
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
--fcc-editable-region--

--fcc-editable-region--
```
