---
id: 68f1179eaea3bb59ba434386
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Para eliminar espacios en blanco al final de una cadena, puedes usar el método `trimEnd()`.

Aquí hay un ejemplo:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

Crea una variable llamada `trimmedEnd` y asígnale el resultado de usar el método `trimEnd()` en la variable `userInput`. 

# --hints--

Deberías tener una variable `trimmedEnd`.

```js
assert.exists(trimmedEnd);
```

Tu variable `trimmedEnd` debe ser una cadena.

```js
assert.isString(trimmedEnd);
```

Debes asignar `userInput.trimEnd()` a tu variable `trimmedEnd`.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
