---
id: 68f3f890010c870e17bbe754
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

Ahora es momento de ver los resultados.

Debajo de tu variable `upperCaseInput`, agrega un `console.log()` con la cadena `"Result of using the toUpperCase() method:"`. 

Luego, debajo de ese `console.log()`, agrega otro `console.log()` con la variable `upperCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes registrar la cadena `"Result of using the toUpperCase() method:"` en la consola.

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

Debes registrar la variable `upperCaseInput` en la consola.

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
