---
id: 68f3f5ca1d2a490c97f14f03
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Ahora es momento de ver los resultados de recortar los espacios en blanco al final de la cadena.

Debajo de tu variable `trimmedEnd`, agrega un `console.log()` con la cadena `"After using the trimEnd() method, trailing spaces removed:"`. 

Luego, debajo de ese `console.log()`, agrega otro `console.log()` con la variable `trimmedEnd`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes registrar en la consola la cadena `"After using the trimEnd() method, trailing spaces removed:"`.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

Debes registrar la variable `trimmedEnd` en la consola.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
