---
id: 68f3eeaa65614107ecd5f1f2
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Debajo de tu primer `console.log()`, agrega otro `console.log()` para registrar la variable `userInput` en la consola.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes registrar la variable `userInput` en la consola.

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
