---
id: 68f3f3d65d51d80ac0b412fc
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Ahora es momento de ver los resultados de recortar los espacios en blanco al inicio de la cadena.

Debajo de tu variable `trimmedStart`, agrega un `console.log()` con la cadena `"After using the trimStart() method, leading spaces removed:"`. 

Luego, debajo de ese `console.log()`, agrega otro `console.log()` con la variable `trimmedStart`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes registrar la cadena `"After using the trimStart() method, leading spaces removed:"` en la consola.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

Debes registrar la variable `trimmedStart` en la consola.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
