---
id: 68f40925d810ae14b2d5a169
title: Paso 16
challengeType: 1
dashedName: step-16
---

# --description--

En este momento, tu variable `camelCasedVersion` contiene el valor de `camelC`. Pero el resultado deseado es `camelCase`. 

Para obtener las letras restantes `ase`, puedes usar el método `slice()` con un valor de índice negativo. Recuerda que usar valores de índice negativos indica al método `slice()` que extraiga desde el final. 

Aquí tienes un ejemplo de recordatorio:

```js
const str = "abcd";
str.slice(-2); // cd
```

Usa el operador `+` para concatenar `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` con el resultado de usar el método `slice()` en `lowercaseWord`. Pasa el número `-3` para el método `slice()`.

¡Y con ese último cambio, tu taller está completo!

# --hints--

Debes concatenar `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` con `lowercaseWord.slice(-3)`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

NO debes asignar la cadena literal `"camelCase"` a la `camelCasedVersion`. Por favor, vuelve a leer las instrucciones para el enfoque correcto.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
