---
id: 68f405c8c9df44132c8fe1bb
title: Paso 15
challengeType: 1
dashedName: step-15
---

# --description--

Como recuerdas de lecciones anteriores, la convención de nombres camelCase comienza con una letra minúscula para la primera palabra, y cada palabra siguiente empieza con una letra mayúscula seguida de letras minúsculas.

Aquí tienes algunos ejemplos de recordatorio: 

```md
freeCodeCamp
camelCase
learningIsFun
```

La segunda palabra en la variable `lowercaseWord` es `"case"`. Para acceder a la `c` en esa palabra, puedes usar `lowercaseWord[5]`. 

Usa el operador `+` para concatenar `lowercaseWord.slice(0, 5)` con el resultado de usar el método correcto para convertir cadenas a mayúsculas en `lowercaseWord[5]`.

Ahora deberías ver `camelC` en la consola.

# --hints--

Debes concatenar `lowercaseWord.slice(0, 5)` con `lowercaseWord[5].toUpperCase()`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

NO debes asignar la cadena literal `"camelC"` a `camelCasedVersion`. Por favor, vuelve a leer las instrucciones para el enfoque correcto.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

