---
id: 66c8ba41a77db20a93f9d7a1
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本工作坊中，你将通过构建一个贷款资格检查应用程序来复习 JavaScript 条件语句。

该应用程序将根据用户的年收入和信用评分，检查其是否有资格申请复式楼房、汽车和公寓贷款。 信用评分是一个数字，代表你管理借贷资金的能力。

首先，创建以下变量和值。

| Variable Name | Value   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

你应该创建一个 `minIncomeForDuplex` 变量，设置为 `60000`。

```js
assert.equal(minIncomeForDuplex, 60000);
```

你应该创建一个 `minCreditScoreForDuplex` 变量，设置为 `700`。

```js
assert.equal(minCreditScoreForDuplex, 700);
```

你应该创建一个 `minIncomeForCar` 变量，设置为 `30000`。

```js
assert.equal(minIncomeForCar, 30000);
```

你应该创建一个 `minCreditScoreForCar` 变量，设置为 `650`。

```js
assert.equal(minCreditScoreForCar, 650);
```

你应该创建一个 `minIncomeForCondo` 变量，设置为 `45000`。

```js
assert.equal(minIncomeForCondo, 45000);
```

你应该创建一个 `minCreditScoreForCondo` 变量，设置为 `680`。

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
