---
id: 66c8ba975ee7230e29f6c4ad
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

要根据 `annualIncome` 和 `creditScore` 检查用户是否符合贷款条件，你必须在 `getLoanMessage` 函数中使用 `if/else if` 语句或三元运算符。 然后，你将在每个条件块中返回相应的信息。

从复式房屋贷款开始，检查 `annualIncome` 是否大于或等于 `minIncomeForDuplex`，且 `creditScore` 是否大于等于 `minCreditScoreForDuplex`。

如果该条件成立，那么申请人就有资格申请复式房屋贷款和其他贷款。 因此，在检查内部，返回字符串 `"You qualify for a duplex, condo, and car loan."`

# --hints--

你的 `getLoanMessage` 函数应返回一个字符串。

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

你的 `getLoanMessage` 函数应返回字符串 `"You qualify for a duplex, condo, and car loan."`。

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

如果申请人的年收入和信用评分不符合复式房屋贷款的要求，你的 `getLoanMessage` 函数应返回 `undefined` 。

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
