---
id: 66c8ba975ee7230e29f6c4ac
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

当用户符合贷款条件时，你需要在控制台中向他们显示一条信息。

为此，你需要建立一个函数，在这个函数中，你需要进行一些检查，以返回申请人符合哪些贷款条件。

创建一个空的 `getLoanMessage` 函数，参数为 `annualIncome` 和 `creditScore` 。

# --hints--

你应创建一个 `getLoanMessage` 函数。

```js
assert.isFunction(getLoanMessage)
```

你的 `getLoanMessage` 函数应将 `annualIncome` 作为 `creditScore` 作为参数。

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` 应该是一个空函数。

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
