---
id: 66c8ba975ee7230e29f6c4af
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

如果申请人的年收入大于或等于 `minIncomeForCondo`，且如果他们的信用分数大于或等于 `minCreditScoreForCondo`，那么他们就有资格申请公寓和汽车贷款。

在 `getLoanMessage` 函数中检查是否为真。 如果是，则返回字符串 `"You qualify for a condo and car loan."`

# --hints--

如果申请人符合公寓和汽车贷款条件，你的 `getLoanMessage` 函数应返回一个字符串。

```js
assert.isString(getLoanMessage(45000, 680));
```

你的 `getLoanMessage` 函数应返回字符串 `"You qualify for a condo and car loan."`。

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

如果申请人的年收入和信用评分不符合公寓贷款的要求，你的 `getLoanMessage` 函数应返回 `undefined` 。

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }

--fcc-editable-region--
}
```
