---
id: 66c8ba975ee7230e29f6c4b0
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

现在，你应该只检查申请人是否有资格申请汽车贷款。 如果是，则返回字符串 `"You qualify for a car loan."`。

# --hints--

如果申请人符合汽车贷款条件，你的 `getLoanMessage` 函数应返回一个字符串。

```js
assert.isString(getLoanMessage(30000, 650));
```

你的 `getLoanMessage` 函数应返回字符串 `"You qualify for a car loan."`。

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

如果申请人的年收入和信用评分不符合汽车贷款的要求，你的 `getLoanMessage` 函数应返回 `undefined` 。

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 

--fcc-editable-region--
}
```
