---
id: 66c8ba975ee7230e29f6c4b2
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

现在，是测试 `getLoanMessage` 函数的时候了。

使用下表创建 `4` 变量及其值：

| 变量名称            | Value                        |
| --------------- | ---------------------------- |
| `duplexLoanMsg` | `getLoanMessage(85000, 850)` |
| `condoLoanMsg`  | `getLoanMessage(65000, 690)` |
| `carLoanMsg`    | `getLoanMessage(45000, 660)` |
| `noLoanMsg`     | `getLoanMessage(25000, 550)` |

然后，将每个变量记录到控制台，查看信息。

至此，你的贷款资格检查项目就完成了！

# --hints--

你应该有一个名为 `duplexLoanMsg` 的变量。

```js
assert.isNotNull(duplexLoanMsg);
```

你的 `duplexLoanMsg` 变量应设置为 `getLoanMessage(85000, 850)`。

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

你应将 `duplexLoanMsg` 变量记录到控制台。

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

你应该有一个名为 `condoLoanMsg` 的变量。

```js
assert.isNotNull(condoLoanMsg);
```

你的 `condoLoanMsg` 变量应设置为 `getLoanMessage(65000, 690)`。

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

你应将 `condoLoanMsg` 变量记录到控制台。

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

你应该有一个名为 `carLoanMsg` 的变量。

```js
assert.isNotNull(carLoanMsg);
```

你的 `carLoanMsg` 变量应设置为 `getLoanMessage(45000, 660)`。

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

你应将 `carLoanMsg` 变量记录到控制台。

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

你应该有一个名为 `noLoanMsg` 的变量。

```js
assert.isNotNull(noLoanMsg);
```

你的 `noLoanMsg` 变量应设置为 `getLoanMessage(25000, 550)`。

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

你应将 `noLoanMsg` 变量记录到控制台。

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
