---
id: 66c06fad3475cd92421b9ac2
title: 윤년 계산기 만들기
challengeType: 26
dashedName: build-a-leap-year-calculator
---

# --description--

윤년은 `4`으로 나누어떨어지는 해이지만, `100`으로 나누어떨어지고 `400`으로 나누어떨어지지 않는 해는 제외됩니다. 예를 들어, `2000`은 윤년이지만 `1900`는 윤년이 아닙니다. 또한 윤년은 2월에 29일이라는 하루가 추가됩니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 숫자를 인수로 받는 `isLeapYear`라는 함수를 정의하세요.
2. 함수 외부에 검사할 연도의 값을 저장하는 `year` 변수를 선언하세요.
3. 함수 내부에서 `if/ else` 문이나 삼항 연산자를 사용해 연도가 윤년인지 확인하세요.
4. 윤년인지 확인하려면 다음 조건을 충족하세요:

   - 연도가 `4`으로 나누어떨어지면 윤년입니다.
   - 단, 연도가 `100`으로도 나누어떨어지면 윤년이 아닙니다.
   - 단, 연도가 `400`으로도 나누어떨어지면 윤년입니다.

5. 연도가 윤년이면 `[year] is a leap year.`을 반환하고, 그렇지 않으면 `[year] is not a leap year.`를 반환하세요. `[year]`은 `isLeapYear` 함수에서 정의한 매개변수로 대체합니다.
6. `isLeapYear` 함수를 `year` 인수로 호출하고 결과를 `result`이라는 변수에 할당하세요.
7. `result` 변수를 `console.log()`를 사용해 콘솔에 출력하세요.

# --hints--

`isLeapYear`라는 함수를 정의해야 합니다.

```js
assert.isFunction(isLeapYear);
```

`isLeapYear` 함수는 숫자를 인수로 받아야 합니다.

```js
assert.match(isLeapYear.toString(), /\s*isLeapYear\(\s*\w+\s*\)/);
```

윤년인지 확인하기 위해 `year` 변수를 선언하고 값을 할당해야 합니다.

```js
assert.isDefined(year);
assert.isNumber(year);
assert.isAtLeast(year, 0);
assert.isTrue(Number.isInteger(year));
```

`year` 변수는 비어 있으면 안 됩니다.

```js
assert.isNotNull(year);
```

`2024` 변수의 값이 `year`일 때, `result`은 `2024 is a leap year.`이어야 합니다.

```js
assert.strictEqual(isLeapYear(2024), '2024 is a leap year.');
```

`2000` 변수의 값이 `year`일 때, `result`은 `2000 is a leap year.`이어야 합니다.

```js
assert.strictEqual(isLeapYear(2000), '2000 is a leap year.');
```

`1900` 변수의 값이 `year`일 때, `result`는 `1900 is not a leap year.`이어야 합니다.

```js
assert.strictEqual(isLeapYear(1900), '1900 is not a leap year.');

```

`isLeapYear` 함수를 호출하고 `year`을 매개변수로 전달해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /isLeapYear\(\s*year\s*\)/);
```

`result` 변수를 선언해야 합니다.

```js
assert.isDefined(result);
```

`isLeapYear` 함수를 호출한 결과를 `result`이라는 변수에 저장해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*isLeapYear\(\s*year\s*\)/);
```

`result`를 `console.log()`를 사용해 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function isLeapYear(year) {
  if ((year % 4 === 0 && year % 100 !== 0) || year % 400 === 0) {
    return(year + " is a leap year.");
  } else {
    return(year + " is not a leap year.");
  }
}

const year = 2024;
const result = isLeapYear(year);
console.log(result);
```
