---
id: 67bf4350777ac6ffdc027745
title: Eine Farbwahl-App erstellen
challengeType: 25
dashedName: build-a-color-picker-app
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um die Aufgabe abzuschließen.

**User Stories:**

1. Sie sollten eine `ColorPicker`-Komponente definieren und exportieren.
1. Sie sollten den `useState`-Hook verwenden.
1. Sie sollten ein `#color-picker-container`-Element mit weißem Hintergrund haben.
1. Sie sollten ein `#color-input`-Element haben, das ein `color`-Input sein sollte.
1. Ihr `#color-input` sollte ein Kind von `#color-picker-container` sein.
1. Wenn `#color-input` geändert wird, sollte `#color-picker-container` seinen Hintergrund auf diesen neuen Wert setzen.

# --hints--

Sie sollten eine `ColorPicker`-Komponente definieren und exportieren.

```js
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const _a = eval(script);

  assert.property(exports, "ColorPicker");
```

Sie sollten den `useState`-Hook verwenden.

```js
    const abuseState = __helpers.spyOn(React, "useState");
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const _b = await __helpers.prepTestComponent(exports.ColorPicker);
    assert.isAtLeast(abuseState.calls.length, 1);
```

Sie sollten den Rückgabewert des `useState`-Hooks verwenden.

```js
assert.match(code, /=\s*(?:React\.)?useState\(/);
```

Sie sollten Destrukturierung verwenden, um den Rückgabewert von `useState` zwei Variablen zuzuweisen.

```js
assert.match(code, /\[.*,.*\]\s*=\s*(?:React\.)?useState\(/);
```

Sie sollten `#ffffff` als Anfangswert an Ihren `useState`-Aufruf übergeben.

```js
    const abuseState = __helpers.spyOn(React, "useState");
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const _b = await __helpers.prepTestComponent(exports.ColorPicker);
    assert.match(abuseState.calls[0], /#ffffff/i);
```

Sie sollten ein Element mit `color-picker-container` als `id` rendern.

```js
assert.exists(document.getElementById("color-picker-container"));
```

Ihr `#color-picker-container` sollte anfangs einen weißen Hintergrund haben.

```js
const el = document.getElementById("color-picker-container");
const style = window.getComputedStyle(el);
assert.strictEqual(style.backgroundColor, "rgb(255, 255, 255)");
```

Sie sollten ein `input`-Element mit einem `id` von `color-input` rendern.

```js
const inputMaybe = document.getElementById("color-input");
assert.exists(inputMaybe);
assert.strictEqual(inputMaybe?.tagName, "INPUT");
```

Ihr `#color-input` sollte ein Kind von `#color-picker-container` sein.

```js
const inputPossibly = document.getElementById("color-input");
const parent = inputPossibly.parentElement;
assert.exists(parent);
assert.strictEqual(parent.getAttribute("id"), "color-picker-container");
```

Ihr `#color-input` sollte ein `color`-Input sein.

```js
const inputDefinitely = document.getElementById("color-input");
assert.equal(inputDefinitely?.type, "color");
```

Ihr `#color-input` sollte ein `onChange`-Attribut haben.

```js
const definitelyAnInput = document.getElementById("color-input");
assert.exists(definitelyAnInput);
const key = Object.keys(definitelyAnInput).find(key => key.startsWith("__reactProps"));
const reactInstance = definitelyAnInput[key]
assert.exists(reactInstance);
assert.isFunction(reactInstance.onChange);
```

Ihr `#color-input` sollte einen Anfangswert von `#ffffff` haben.

```js
const mostCertainlyAnInput = document.getElementById("color-input");
assert.exists(mostCertainlyAnInput);
assert.strictEqual(mostCertainlyAnInput.value, "#ffffff");
```

Der Anfangswert Ihres `#color-input` sollte von Ihrem `useState`-Hook verwaltet werden.

```js
const abuseState = __helpers.spyOn(React, "useState");
const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
const exports = {};
const _a = eval(script);

const rendered = await __helpers.prepTestComponent(exports.ColorPicker);
const inputElement = rendered.querySelector('#color-input');
assert.exists(inputElement, "Could not find #color-input element");

const initialValue = inputElement.value;

const stateIndex = abuseState.returns.findIndex(([stateValue]) => stateValue === initialValue);
assert.isAtLeast(stateIndex, 0, "The #color-input value must be managed by a useState hook");

const [, setterFunction] = abuseState.returns[stateIndex];

const newColor = '#123456';
await React.act(async () => {
  setterFunction(newColor);
});

assert.strictEqual(
  inputElement.value,
  newColor,
  "The #color-input value should update when the useState state changes"
);
```

Wenn sich Ihr `#color-input`-Wert ändert, sollte Ihr `#color-picker-container` seine Hintergrundfarbe ändern, um diesen Wert widerzuspiegeln.

```js
  const inputForSure = document.getElementById("color-input");
  assert.exists(inputForSure);
  const nativeValueSetter = Object.getOwnPropertyDescriptor(window.HTMLInputElement.prototype, "value").set;
  await React.act(async () => {
    nativeValueSetter.call(inputForSure, "#000001");
    inputForSure.dispatchEvent(new Event("input", { bubbles: true}));
  });
  const container = document.getElementById("color-picker-container");
  assert.exists(container);
  const style = window.getComputedStyle(container);
  assert.equal(style?.backgroundColor, "rgb(0, 0, 1)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ColorPicker } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<ColorPicker />);
    </script>
</body>

</html>
```

```css
body,
html {
    margin: 0;
    padding: 0;
    height: 100%;
    font-family: Arial, sans-serif;
}

#color-picker-container {
    height: 100vh;
    display: flex;
    justify-content: center;
    align-items: center;
    flex-direction: column;
    background-color: #ffffff;
}

input[type="color"] {
    position: absolute;
    margin-top: 50px;
    height: 40px;
}
```

```jsx
const { useState } = React;

export const ColorPicker = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ColorPicker } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<ColorPicker />);
    </script>
</body>

</html>
```

```css
body,
html {
    margin: 0;
    padding: 0;
    height: 100%;
    font-family: Arial, sans-serif;
}

#color-picker-container {
    height: 100vh;
    display: flex;
    justify-content: center;
    align-items: center;
    flex-direction: column;
    background-color: #ffffff;
}

#pick-color-btn {
    padding: 15px 30px;
    font-size: 18px;
    color: white;
    background-color: #14183c;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    margin-top: 20px;
}

#pick-color-btn:hover {
    background-color: #08131f;
}

input[type="color"] {
    position: absolute;
    margin-top: 50px;
    height: 40px;
}

.stylep{
    color: #14183c;
    font-size: 15px;
    font-weight: bold;
    margin-bottom: 50px;

}
```

```jsx
const { useState } = React;

export const ColorPicker = () => {
    const [bgColor, setBgColor] = useState('#ffffff');

    const handleColorChange = (event) => {
        setBgColor(event.target.value);
    };

    return (
            <div id="color-picker-container" style={{ backgroundColor: bgColor }}>
            <p className="stylep">
                Choose a color using the color input below:
            </p>

            <input
                id="color-input"
                type="color"
                value={bgColor}
                onChange={handleColorChange}
            />
        </div>
    );
}
```
