---
id: 66b205e6eacba4c4e54ea434
title: 创建一个电子邮件掩码器
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

在本实验室中，您将用星号遮盖电子邮件地址中的用户名部分。 掩码是一个用星号或其他字符隐藏或替换敏感信息的术语。

例如，如果电子邮件地址是 `myEmail@email.com`，那么掩码的电子邮件地址将是 `m*****l@email.com` 。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 创建一个名为 `maskEmail` 的函数， 将 `email` 作为参数。
2. 在函数中，你应对 `email` 进行掩码处理，并将域名附于其后。  记住你可以使用 `slice`、`repeat`、`indexOf` 或 `replace` 等字符串方法来帮忙实现。
3. 在函数外部，声明一个名为 `email` 的变量，用于存储要掩码的电子邮件地址。
4. 调用 `maskEmail` 函数，将 `email` 变量作为参数，并将结果输出到控制台。
5. `maskEmail("apple.pie@example.com")` 应返回 `"a*******e@example.com"`。
6. `maskEmail("freecodecamp@example.com")` 应返回 `"f**********p@example.com"`。
7. `maskEmail("info@test.dev")` 应返回 `"i**o@test.dev"`。
8. `maskEmail("user@domain.org")` 应返回 `"u**r@domain.org"`。

# --hints--

你应定义一个名为 `maskEmail` 的函数。

```js
assert.isFunction(maskEmail);
```

`maskEmail` 函数应接收一个字符串 `email` 作为参数。

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

在函数外部，你应该有一个 `email` 变量。

```js
assert.isDefined(email);
```

你应为 `email` 变量指定一个有效的电子邮件地址。

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` 应返回 `"a*******e@example.com"`。

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` 应返回 `"f**********p@example.com"`。

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` 应返回 `"i**o@test.dev"`。

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` 应返回 `"u**r@domain.org"`。

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

你的 `maskEmail` 应产生正确的结果。

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

你应记录调用 `maskEmail` 的输出结果，其参数为 `email`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
