---
id: a789b3483989747d63b0e427
title: Jenga kipata nambari kubwa zaidi
challengeType: 26
dashedName: build-the-largest-number-finder
---

# --description--

Katika maabara hii, utajenga kitendakazi kinachorudisha safu ya data inayojumuisha nambari kubwa zaidi kutoka kila safu ndogo iliyotolewa.

Kumbuka, unaweza kurudia kupitia safu ya data kwa kutumia mzunguko rahisi wa `for`, na kupata kila mshiriki kwa sintaksia ya safu ya data `arr[i]`.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuunda kitendakazi `largestOfAll` kinachopokea safu ya safu za data kama hoja.
2. Kitendakazi kinapaswa kurudisha safu ya data inayojumuisha nambari kubwa zaidi kutoka kila safu ndogo.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `largestOfAll`.

```js
assert.isFunction(largestOfAll);
```

`largestOfAll` inapaswa kupokea kigezo kimoja.

```js
assert.lengthOf(largestOfAll, 1);
```

`largestOfAll([[4, 5, 1, 3], [13, 27, 18, 26], [32, 35, 37, 39], [1000, 1001, 857, 1]])` inapaswa kurudisha safu ya data.

```js
assert.isArray(
  largestOfAll([
    [4, 5, 1, 3],
    [13, 27, 18, 26],
    [32, 35, 37, 39],
    [1000, 1001, 857, 1]
  ])
);
```

`largestOfAll([[13, 27, 18, 26], [4, 5, 1, 3], [32, 35, 37, 39], [1000, 1001, 857, 1]])` inapaswa kurudisha `[27, 5, 39, 1001]`.

```js
assert.deepEqual(
  largestOfAll([
    [13, 27, 18, 26],
    [4, 5, 1, 3],
    [32, 35, 37, 39],
    [1000, 1001, 857, 1]
  ]),
  [27, 5, 39, 1001]
);
```

`largestOfAll([[4, 9, 1, 3], [13, 35, 18, 26], [32, 35, 97, 39], [1000000, 1001, 857, 1]])` inapaswa kurudisha `[9, 35, 97, 1000000]`.

```js
assert.deepEqual(
  largestOfAll([
    [4, 9, 1, 3],
    [13, 35, 18, 26],
    [32, 35, 97, 39],
    [1000000, 1001, 857, 1]
  ]),
  [9, 35, 97, 1000000]
);
```

`largestOfAll([[17, 23, 25, 12], [25, 7, 34, 48], [4, -10, 18, 21], [-72, -3, -17, -10]])` inapaswa kurudisha `[25, 48, 21, -3]`.

```js
assert.deepEqual(
  largestOfAll([
    [17, 23, 25, 12],
    [25, 7, 34, 48],
    [4, -10, 18, 21],
    [-72, -3, -17, -10]
  ]),
  [25, 48, 21, -3]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function largestOfAll(arr) {
  return arr.map(subArr => Math.max.apply(null, subArr));
}

largestOfAll([
  [4, 5, 1, 3],
  [13, 27, 18, 26],
  [32, 35, 37, 39],
  [1000, 1001, 857, 1]
]);
```
