---
id: 655cd899f8de09431eabb40c
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذا المشروع، ستتعلم هياكل البيانات من خلال بناء خوارزمية دمج الفرز.

هذه خوارزمية فرز تستخدم مبدأ القسمة والتغلب لفرز مجموعات البيانات. أي أنها "تقسم" المجموعة إلى أجزاء فرعية أصغر، و"تتغلب" على الأجزاء الفرعية بفرزها بشكل مستقل، ثم تدمج الأجزاء الفرعية المرتبة.

أنشئ دالة تسمى `merge_sort`. ستتولى هذه الدالة مهمة فرز قائمة من الأرقام.

استخدم الكلمة المفتاحية `pass` داخل جسم الدالة.

# --hints--

يجب إعلان دالة باسم `merge_sort`. لا تنسَ استخدام الكلمة المفتاحية `pass` داخل جسم الدالة.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
