---
id: 6564a9fe51964c229d5b7f4c
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

ستحتاج إلى معلمة تشير إلى مجموعة البيانات التي ستُرتب. أنشئ معلمة باسم `array` في الدالة `merge_sort`.

# --hints--

يجب أن تأخذ دالة `merge_sort` معلمة واحدة فقط: `array`.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
