---
id: 6564aee9c077774ea49c3faf
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

تُنفذ خوارزمية دمج الترتيب ثلاث عمليات رئيسية:

- تقسيم متتالية غير مرتبة من العناصر إلى أجزاء فرعية
- ترتيب العناصر في الأجزاء الفرعية
- دمج الأجزاء الفرعية المرتبة

تحدث العمليات السابقة بشكل متكرر حتى تُدمج الأجزاء الفرعية في المتتالية المرتبة الكاملة. ابدأ بتقسيم المتتالية.

أولًا، استبدل الكلمة المفتاحية `pass` بمتغير `middle_point`. استخدم عامل القسمة الصحيحة (`//`) لقسمة طول قائمة `array` على 2 وعيّن النتيجة إلى متغير `middle_point` الجديد. تذكر أن تُزح شفرتك.

# --hints--

يجب ألا يحتوي جسم الدالة على `pass`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.notMatch(function_body, /^\s{4}pass/m);
  }
})
```

يجب إعلان متغير `middle_point` وتعيين `len(array) // 2` له.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
