---
id: 656680b0fc79f2c38a34d90e
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

حان الوقت الآن لفرز ودمج القائمتين (`left_part` و `right_part`) في الـ `array` الأصلية.

يمكنك فعل ذلك بمقارنة العناصر في كلتا القائمتين، ودمج العنصر الأصغر في القائمة الرئيسية. ستقوم بهذه المقارنة لكل الفهارس في `left_part` و `right_part`.

أنشئ ثلاث متغيرات: `left_array_index` و `right_array_index` و `sorted_index` وضع قيمها على `0`. ستساعدك هذه المتغيرات في تتبع كل فهرس أثناء عملية الفرز.

# --hints--

يجب أن يكون لديك متغير باسم `left_array_index` داخل دالة `merge_sort` الخاصة بك

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index/);
  }
})
```

يجب أن تكون قيمة `left_array_index` هي 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index\s*=\s*0/);
  }
})
```

يجب أن يكون لديك متغير باسم `right_array_index` داخل دالة `merge_sort` الخاصة بك

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index/);
  }
})
```

يجب أن تكون قيمة `right_array_index` هي 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index\s*=\s*0/);
  }
})

```

يجب أن يكون لديك متغير باسم `sorted_index` داخل دالة `merge_sort` الخاصة بك

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index/);
  }
})
```

يجب أن تكون قيمة `sorted_index` هي 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index\s*=\s*0/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)

--fcc-editable-region--

--fcc-editable-region--
```
