---
id: 656639771fed09b5c0e8fe71
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

الآن بعد أن قسمت قائمة `array` إلى قائمتين منفصلتين، استمر في تقسيم كل قائمة حتى يصبح كل عنصر بمفرده في قائمته الخاصة. القائمة التي تحتوي على رقم واحد دائمًا ما تكون مرتبة.

للقيام بذلك، استدعِ دالة `merge_sort` بشكل متكرر داخل دالتك. لا تمرر أي معلمة في الوقت الحالي.

# --hints--

يجب أن تستدعي `merge_sort` في أسفل جسم الدالة.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
