---
id: 656627b47bd2d2a4afbc945d
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

في الخطوة السابقة، حصلت على نقطة المنتصف. يمكنك استخدامها لتقسيم `array` إلى قسمين وتعيين كل جزء إلى متغير جديد.

استخدم صيغة القطع لاستخراج النصف الأيسر من `array` وتعيينه إلى متغير باسم `left_part`.


# --hints--

يجب أن يكون لديك متغير باسم `left_part` في دالة `merge_sort` الخاصة بك.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

يجب تعيين `array[:middle_point]` إلى المتغير `left_part`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
