---
id: 656702f8b4cbd8cbf0a433c6
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

داخل دالتك، أنشئ حلقة `while` تقارن عنصرًا في `left_part` مع عنصر في `right_part`، وتدمج العنصر الأصغر في القائمة الرئيسية `array`.

أنشئ شرطين للحلقة: أحدهما يتحقق مما إذا كان `left_array_index` أقل من طول `left_part`، والآخر يتحقق مما إذا كان `right_array_index` أقل من طول `right_part`.

# --hints--

يجب أن تحتوي حلقة `while` على هذه الشروط: `left_array_index < len(left_part) and right_array_index < len(right_part)`. لا تنسَ استخدام الكلمة المفتاحية `pass` في جسم الحلقة.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

استخدم الكلمة المفتاحية `pass` في جسم الحلقة

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
