---
id: 656707e11e671ed54c96f7ec
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

عندما تُقيّم شرط `if` إلى `True`، فهذا يعني أن العنصر في قائمة `left_part` أصغر من العنصر الذي يُقارن به في قائمة `right_part`.

في هذه الحالة، يمكنك تعيين فهرس `left_part` إلى المصفوفة المرتبة.

داخل كتلة `if`، احذف `pass` وقم بتعيين `left_part[left_array_index]` إلى `array[sorted_index]`.

# --hints--

يجب أن تستبدل `pass` بـ `array[sorted_index] = left_part[left_array_index]` داخل جسم جملة `if`.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
