---
id: 656710d1e0ec62253426db24
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

حلقة `while` التي أنشأتها تقارن عنصرًا واحدًا من `left_part` مع آخر في `right_part`، ثم تضيف العنصر الأصغر إلى القائمة الرئيسية `array`.

ستستمر هذه العملية حتى لا يتبقى عناصر للمقارنة. لكن قد تبقى عناصر في `left_part` بينما لا يوجد في `right_part`، والعكس صحيح.

أنشئ حلقة `while` أخرى لنسخ العناصر المتبقية في `left_part` إلى قائمة `array`. استخدم الشرط `left_array_index < len(left_part)` في حلقة الـ `while`.

# --hints--

يجب أن تحتوي حلقة `while` على هذا الشرط: `left_array_index < len(left_part)`. لا تنسَ استخدام الكلمة المفتاحية `pass` في جسم الحلقة.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):/m);
  }
})
```

يجب أن تحتوي حلقة `while` على الكلمة المفتاحية `pass` في جسمها

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
