---
id: 656706afd65547d22bee0b68
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

داخل حلقة `while`، استبدل `pass` بعبارة `if` تتحقق مما إذا كان فهرس `left_part` أقل من فهرس `right_part`.

استخدم الكلمة المفتاحية `pass` داخل جسم عبارة `if`.

# --hints--

يجب استبدال الكلمة المفتاحية `pass` بعبارة `if` تحتوي على الشرط التالي: `left_part[left_array_index] < right_part[right_array_index]`. أضف الكلمة المفتاحية `pass` داخل جسم عبارة `if`.

```js
({
 test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_body().is_equivalent("if left_part[left_array_index] < right_part[right_array_index]:\\n  pass")`))
 })
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        pass
--fcc-editable-region--
```
