---
id: 6569b743630ee592a65a7e2f
title: الخطوة 26
challengeType: 20
dashedName: step-26
---

# --description--

قبل اختبار الدالة `merge_sort()`، عليك إنشاء حالة أساسية توقف تنفيذ الدالة عندما يكون طول الـ `array` أقل من أو يساوي `1`.

هذه الحالة الأساسية توقف استدعاء التكرار. بدونها، ستستمر عملية الفرز بالدمج في العمل حتى عندما تكون القائمة مرتبة أو لا تحتوي على عناصر.

بعد الإعلان عن الدالة مباشرة، أنشئ جملة `if` بالشرط التالي: `len(array) <= 1`. استخدم الكلمة المفتاحية `pass` في جسم الدالة.

# --hints--

يجب إنشاء جملة `if` للتحقق مما إذا كان `len(array) <= 1`. لا تنسَ استخدام الكلمة المفتاحية `pass` في جسم جملة الـ `if`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):

--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
