---
id: 657f3a661730891aa64f3568
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

استخدم صيغة القطع لاستخراج النصف الأيمن من `array` وتعيينه إلى متغير باسم `right_part`.

# --hints--

يجب أن يكون لديك متغير باسم `right_part`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_part\s*(?!=)/m);
  }
})
```

يجب تعيين `array[middle_point:]` إلى المتغير `right_part`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_part\s*=\s*array\s*\[\s*middle_point\s*:\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
