---
id: 68e51d8f163d3f8b70f9c4e5
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utazoeza kufanya kazi na njia za `includes()` na `slice()`.

Kuanza, tengeneza kigezo kinachoitwa `fccSentence` na kitoe thamani ya mfuatano wa herufi `"freeCodeCamp is a great place to learn web development."`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `fccSentence`.

```js
assert.exists(fccSentence);
```

Kigezo chako cha `fccSentence` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(fccSentence);
```

Kigezo chako cha `fccSentence` kinapaswa kupewa thamani ya mfuatano wa herufi `"freeCodeCamp is a great place to learn web development."`.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
