---
id: 68e6172975008fea742a63cc
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Njia ya `includes()` inaangalia tofauti kati ya herufi kubwa na herufi ndogo, ikimaanisha inatofautisha kati ya herufi kubwa na herufi ndogo.

Tangaza kigezo kinachoitwa `hasLowercaseFCC`. Kisha weka thamani yake kwa matokeo ya kutumia njia ya `includes()` kuangalia kama `fccSentence` ina `"freecodecamp"` (herufi ndogo zote).

Kisha tumia au mfuatano wa herufi wa template au concatenation ya mfuatano wa herufi kuandika ujumbe wa `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` kwenye konsoli. Badilisha `<hasLowercaseFCC>` na thamani halisi ya kigezo.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `hasLowercaseFCC`.

```js
assert.exists(hasLowercaseFCC);
```

Unapaswa kutumia `fccSentence.includes("freecodecamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Unapaswa kuweka thamani ya matokeo ya `fccSentence.includes("freecodecamp")` kwenye kigezo chako cha `hasLowercaseFCC`.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Matokeo ya `fccSentence.includes("freecodecamp")` yanapaswa kuwa `false`. 

```js
assert.isFalse(hasLowercaseFCC);
```

Unapaswa kuandika ujumbe unaojumuisha `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`, ambapo `<hasLowercaseFCC>` inapaswa kubadilishwa na thamani halisi ya kigezo.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
