---
id: 68e6172975008fea742a63ca
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Sasa tumia mfuatano wa herufi wa template au muunganiko wa mfuatano wa herufi kuandika ujumbe `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` kwenye konsoli. Badilisha `<hasFreeCodeCamp>` na thamani halisi ya kigezo.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Unapaswa kuandika ujumbe `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`, ambapo `<hasFreeCodeCamp>` inapaswa kubadilishwa na thamani halisi ya kigezo.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
