---
id: 68e6172875008fea742a63c9
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Kama unavyokumbuka kutoka kwenye mafundisho yaliyopita, njia ya `includes()` huangalia kama mfuatano wa herufi una mfuatano mdogo maalum na hurudisha `true` au `false`.

Hapa kuna mfano ukitumia njia ya `includes()`:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

Tengeneza kigezo kinachoitwa `hasFreeCodeCamp`. Kisha weka thamani yake kuwa matokeo ya kutumia njia ya `includes()` kwenye `fccSentence` ili kuangalia kama ina `"freeCodeCamp"`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `hasFreeCodeCamp`.

```js
assert.exists(hasFreeCodeCamp);
```

Unapaswa kutumia `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Unapaswa kuweka thamani ya matokeo ya `fccSentence.includes("freeCodeCamp")` kwenye kigezo chako cha `hasFreeCodeCamp`.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Matokeo ya `fccSentence.includes("freeCodeCamp")` yanapaswa kuwa `true`. 

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
