---
id: 68e6172975008fea742a63cf
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Sasa ni wakati wa kurekodi matokeo ya kutumia njia ya `slice()`.

Tumia mfuatano wa herufi wa template au muunganiko wa mfuatano wa herufi kurekodi ujumbe `The word "<platform>" was sliced from the message.` kwenye konsoli. Badilisha `<platform>` na thamani halisi ya kigezo.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Unapaswa kurekodi ujumbe `The word "<platform>" was sliced from the message.`, ambapo `<platform>` inapaswa kubadilishwa na thamani halisi ya kigezo.

```js
assert.deepInclude(spy.calls, ['The word "freeCodeCamp" was sliced from the message.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
--fcc-editable-region--

--fcc-editable-region--
```
