---
id: 68e6172975008fea742a63ce
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Njia ya `slice()` hutumia sehemu ya mfuatano wa herufi. Inachukua hoja mbili: kielezo cha kuanzia (kimejumuishwa) na kielezo cha mwisho (hakijajumuishwa).

Hapa kuna mfano:

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

Viashiria vya mfuatano wa herufi huanzia 0. Katika mfuatano wa `message`, herufi iliyoko kwenye kielezo cha 11 ni `f` (mwanzoni mwa "freeCodeCamp").

Tengeneza kigezo kinachoitwa `platform`. Kisha weka thamani yake kuwa matokeo ya kutumia `message.slice(11, 23)` kuchukua `"freeCodeCamp"`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `platform`.

```js
assert.exists(platform);
```

Kigezo chako cha `platform` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(platform);
```

Unapaswa kutumia `message.slice(11, 23)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Unapaswa kuweka thamani ya matokeo ya `message.slice(11, 23)` kwenye kigezo chako cha `platform`.

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Kigezo chako cha `platform` kinapaswa kuwa sawa na `"freeCodeCamp"`.

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
