---
id: 68e6172975008fea742a63d1
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Njia ya `slice()` inaweza kutumia nambari hasi kuhesabu kutoka mwisho wa mfuatano wa herufi. `-1` inarejelea herufi ya mwisho.

Hapa kuna mfano:

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

Tengeneza kigezo kinachoitwa `endPunctuation`. Kisha weka thamani yake kuwa matokeo ya kutumia `message.slice(-1)` kuchukua herufi ya mwisho.

Kisha tumia au mfuatano wa maandishi wa template au muunganiko wa mfuatano wa herufi kuandika ujumbe `The ending punctuation mark is a "<endPunctuation>"` kwenye konsoli. Badilisha `<endPunctuation>` na thamani halisi ya kigezo.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `endPunctuation`.

```js
assert.exists(endPunctuation);
```

Kigezo chako cha `endPunctuation` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(endPunctuation);
```

Unapaswa kutumia `message.slice(-1)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

Unapaswa kuweka thamani ya matokeo ya `message.slice(-1)` kwenye kigezo chako cha `endPunctuation`.

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

Kigezo chako cha `endPunctuation` kinapaswa kuwa sawa na `"!"`.

```js
assert.strictEqual(endPunctuation, "!");
```

Unapaswa kuandika ujumbe `The ending punctuation mark is a "<endPunctuation>"`, ambapo `<endPunctuation>` inapaswa kubadilishwa na thamani halisi ya kigezo.

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
