---
id: 68e6172975008fea742a63cd
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Sasa utaangalia njia ya `slice()` ambayo hutoa sehemu ya mfuatano wa herufi.

Tengeneza kigezo kinachoitwa `message` na ukae thamani ya mfuatano wa herufi `"Welcome to freeCodeCamp!"`.

Kisha tumia `console.log()` kuandika `"Here are some examples of the slice() method:"` kwenye konsoli.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `message`.

```js
assert.exists(message);
```

Kigezo chako cha `message` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(message);
```

Kigezo chako cha `message` kinapaswa kuwa `"Welcome to freeCodeCamp!"`.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

Unapaswa kuandika `"Here are some examples of the slice() method:"` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
