---
id: 68ef6e42fa6e8a9a18a401fa
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Tumia `console.log()` kuandika `"Here are some examples of the includes() method:"` kwenye konsoli.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Unapaswa kuandika `"Here are some examples of the includes() method:"` kwenye konsoli.

```js
assert.deepInclude(spy.calls, ["Here are some examples of the includes() method:"])
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
