---
id: 66bba6fff611169359d9d36a
title: Створіть конфіденційний лист
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Історія користувача:**

1. Ви повинні мати елемент `main` з `id` зі значенням `email`.
1. Елемент `#email` повинен мати `padding` зі значенням `50px`, верхнє поле зі значенням `50px`, `width` зі значенням `500px` та `border` зі значенням `2px`.
1. Загальна ширина елемента `#email` — включно з відступами і кордонами — має дорівнювати `500px`.
1. В межах елемента `#email` має бути два елементи `div`: один з `id` зі значенням `confidential`, а інший — з `id` зі значенням `top-secret`.
1. Елементи `#confidential` та `#top-secret` повинні мати `display` зі значенням `inline-block`.
1. Елементи `#confidential` та `#top-secret` повинні мати `padding`, ліве поле і `border`.
1. Елемент `#confidential` повинен мати текст `CONFIDENTIAL`.
1. Елемент `#top-secret` повинен мати текст `TOP SECRET`.
1. Елементи `#confidential` та `#top-secret` мають бути нахиленими за допомогою CSS-властивості `transform`.
1. В межах елемента `#email` має бути принаймні три абзаци.
1. В межах абзаців має бути принаймні три елементи `span` з класом `blurred`.
1. Ви повинні мати селектор `blurred`, який задає розмиття зі значенням `3px` за допомогою CSS-властивості `filter`.

# --hints--

Ви повинні мати елемент `main` з `id` зі значенням `email`.

```js
assert.exists(document.querySelector('main#email'));
```

Ви повинні мати селектор `#email`, який встановлює властивість `padding` своїх елементів на `50px` (зі всіх сторін).

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

Ви повинні мати селектор `#email`, який встановлює властивість `margin-top` своїх елементів на `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

Ви повинні мати селектор `#email`, який встановлює властивість `width` своїх елементів на `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

Елемент `#email` повинен мати кордон зі значенням `2px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

Елемент `#email` повинен мати `box-sizing` зі значенням `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

В межах `#email` має бути елемент `div` з `id` зі значенням `confidential`.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

В межах `#email` має бути елемент `div` з `id` зі значенням `top-secret`.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

Елемент `#confidential` повинен мати `display` зі значенням `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

Елемент `#top-secret` повинен мати `display` зі значенням `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

Елемент `#confidential` повинен мати властивість `padding` зі всіх сторін.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Елемент `#top-secret` повинен мати властивість `padding` зі всіх сторін.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Елемент `#confidential` повинен мати ліве поле.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

Елемент `#top-secret` повинен мати ліве поле.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

Елемент `#confidential` повинен мати властивість `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

Елемент `#top-secret` повинен мати властивість `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

Текстом елемента `#confidential` має бути `CONFIDENTIAL`. Ви або не написали текст, або маєте друкарську помилку.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

Текстом елемента `#top-secret` має бути `TOP SECRET`. Ви або не написали текст, або маєте друкарську помилку.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

Елемент `#confidential` має використовувати CSS-властивість `transform`, щоб бути нахиленим.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

Елемент `#top-secret` має використовувати CSS-властивість `transform`, щоб бути нахиленим.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

В межах елемента `#email` має бути принаймні три абзаци.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

В межах абзаців має бути принаймні три елементи `span` з класом `blurred`.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

Елементи `.blurred` не можуть бути порожніми.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

Ви повинні мати селектор `.blurred`, який встановлює властивість `filter` своїх елементів на `blur(3px)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
