---
id: 69008059d43d6751385883c1
title: 偵錯有錯誤的應用程式中的遞增和遞減運算子錯誤
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

你的朋友請你幫忙偵錯他們的程式碼。他們正在練習遞增和遞減運算子，但在主控台中得到意外的結果。

你的任務是幫助他們找出錯誤並更新程式碼，使主控台輸出符合預期結果。

你可以回頭參考先前關於遞增和遞減運算子以及前置式和後置式表示法的課程。

**使用者故事：**

1. 你應該有一個名為 `updatedCounter` 的變數，透過使用累加運算子正確產生值 `11`。更新表達式以使用正確的前置式或後置式表示法，而非硬編碼該值。
2. 你應該有一個名為 `finalScore` 的變數，透過使用累加運算子正確產生值 `8`。更新表達式以使用正確的前置式或後置式符號，而非硬編碼該值。
3. 你應該有一個名為 `updatedCoins` 的變數，透過使用遞減運算子正確產生值 `2`。更新表達式以使用正確的前置式或後置式表示法，而非硬編碼該值。
4. 你應該有一個名為 `newHealth` 的變數，透過使用遞減運算子正確產生值 `7`。更新表達式以使用正確的前置式或後置式表示法，而非硬編碼該值。

# --hints--

你應該有一個名為 `updatedCounter` 的變數，使用前置式遞增表示法（`++counter`）。

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` 應該相等於 `11`。

```js
assert.strictEqual(updatedCounter, 11);
```

你應該有一個名為 `finalScore` 的變數，該變數使用後置式遞增表示法（`score++`）。

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` 應該相等於 `8`。

```js
assert.strictEqual(finalScore, 8);
```

你應該有一個名為 `updatedCoins` 的變數，使用前置式遞減符號（`--coins`）。

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` 應該相等於 `2`。

```js
assert.strictEqual(updatedCoins, 2);
```

你應該有一個名為 `newHealth` 的變數，該變數使用後置式遞減表示法（`health--`）。

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` 應該相等於 `7`。

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
