---
id: 6a3ea8a55317ba7a08e21d5b
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Sasa unaweza hatimaye jibu swali la awali: kwa nini idadi ya mayai duniani kote ilibaki `1`? Kitendakazi kinaporudisha thamani, hakibadilishi moja kwa moja vigezo vyako vilivyopo. Lazima uhifadhi thamani hiyo mpya kwa uwazi!

Unatumia kigezo cha `available_eggs` ndani ya kitendakazi cha `make_fried_egg`, lakini si kigezo kile kile kilichotangazwa juu ya msimbo wako. Kigezo kilicho ndani ya kitendakazi ni kigezo cha ndani kinachokuwepo tu wakati kitendakazi kinaendeshwa. Kitendakazi kinapomalizika, kigezo cha ndani kinatupwa, na `available_eggs` cha ulimwengu kinabaki kisibadilike.

Ili kusasisha hesabu halisi ya jikoni lako, badilisha kigezo cha ulimwengu `available_eggs` kwa idadi mpya ya mayai inayorejeshwa na kitendakazi chako cha `make_fried_egg`.

Mwishoni kabisa mwa msimbo wako, fanya wito kwa kitendakazi cha `make_fried_egg`, ukipitisha `available_eggs` kama hoja, na weka thamani inayorejeshwa tena kwenye kigezo cha ulimwengu `available_eggs`.

# --hints--

Unapaswa kuweka thamani mpya kwa kigezo cha `available_eggs` kwa matokeo ya `make_fried_egg(available_eggs)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
