---
id: 6a3e524fc3394ef52ed5b797
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Vitendakazi vina wigo wa ndani wenyewe. Hii inamaanisha vigezo vilivyoundwa ndani ya kitendakazi ni kama viungo vilivyohifadhiwa katika ghala la jikoni hiyo: vinaweza tu kufikiliwa ndani ya kitendakazi hicho maalum.

Hata hivyo, vitendakazi havitengwa kabisa. Vinaweza kusoma vigezo vya kimataifa vilivyotangazwa nje ya wigo wao. Vinaweza kutumia vigezo hivyo kufanya hesabu, kufanya maamuzi, au kuunda vigezo vipya vya ndani.

Ndani ya kitendakazi chako cha `check_kitchen_stock`, badilisha neno la `pass` na kigezo kipya kinachoitwa `total_items` na weka thamani yake kuwa jumla ya `available_eggs`, `available_flour`, na `available_sugar`.

# --hints--

Unapaswa kufafanua kigezo kinachoitwa `total_items` ndani ya kitendakazi chako cha `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

Kigezo chako cha `total_items` kinapaswa kuwa jumla ya `available_eggs`, `available_flour`, na `available_sugar`.

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
