---
id: 6a3e4ae26f2b7e4b5b026cf8
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Ni wakati wa kuunda kitendakazi ili kukagua kile ulicho nacho jikoni mwako. Katika Python, unaeleza vitendakazi kwa kutumia neno kuu `def`, ikifuatiwa na jina la kitendakazi, mabano ya kawaida `()`, na nukta `:`.

Hapa kuna mfano wa jinsi ya kueleza kitendakazi rahisi:

```py
def my_function():
    pass
```

Eleza kitendakazi tupu chenye jina `check_kitchen_stock()`. Tumia neno kuu `pass` ndani ya mwili wa kitendakazi kwa sasa ili Python isitupie kosa la sintaksia.

Kumbuka kuwa Python hutegemea kivutio cha mstari kuainisha sehemu za msimbo. Kwa hivyo hakikisha unavuta mstari mwili wa kitendakazi.

# --hints--

Unapaswa kueleza kitendakazi chenye jina `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
