---
id: 6a3e67fc5ef4045201c28219
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Ndani ya kitendakazi chako cha `use_eggs`, unahitaji kukagua kama una mayai ya kutosha ya kupikia. Unaweza kufanya hivi kwa kutumia usemi wa `if`.

Usemi wa `if` unathibitisha hali fulani. Ikiwa hali hiyo ni kweli, msimbo ndani ya sehemu ya `if` utaendeshwa.

Badilisha neno la `pass` kwa usemi wa `if` unaokagua kama `eggs_to_use` ni kubwa zaidi (`>`) kuliko `available_eggs`. Ndani ya sehemu ya `if`, chapisha mfuatano wa herufi `The kitchen does not have enough eggs.`.

# --hints--

Unapaswa kuwa na usemi wa `if` unaokagua kama `eggs_to_use > available_eggs`.

```js
({
  test: () => runPython(`
  cond = _Node(_code).find_function("use_eggs").find_ifs()[0].find_conditions()[0]
  assert cond.is_equivalent("eggs_to_use > available_eggs") or cond.is_equivalent("available_eggs < eggs_to_use")`)
})
```

Unapaswa kuchapisha `The kitchen does not have enough eggs.` ndani ya usemi wako wa `if`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_call("print('The kitchen does not have enough eggs.')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
