---
id: 6a3e70b94f67d9ee3ecd8486
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Umefanya wito wa `use_eggs(available_eggs, 1)` kutumia yai moja. Hii inamaanisha unapaswa kuwa na mayai `0` yaliyobaki jikoni kwako sasa, sivyo?

Utajua! Mwishoni mwa msimbo wako, tumia kitendakazi cha `print()` kuchapisha thamani ya kigezo cha `available_eggs`.

Unapotekeleza msimbo, utaona jambo la kushangaza: matokeo bado ni `1`! Kwa nini idadi ya mayai haikubadilika? Utachunguza fumbo hili katika hatua chache zijazo.

# --hints--

Unapaswa kuchapisha kigezo cha `available_eggs` mwishoni mwa msimbo wako.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
