---
id: 6a3e73eabddd614f91cc235a
title: Hatua 14
challengeType: 20
dashedName: step-14
---

# --description--

Sasa utaondoa fumbo la awali! Wakati kitendakazi chako cha `use_eggs` kilirudisha idadi iliyobaki ya mayai, hukuhifadhi thamani hiyo mpya mahali popote. Ilitoweka tu kabisa!

Ili kusahihisha hesabu yako, unahitaji kushika thamani hiyo iliyorejeshwa kwa kuweka thamani kwenye kigezo.

Ndani ya kitendakazi chako cha `make_fried_egg`, tumia tamko la `if/else` kukagua kigezo cha `has_enough_eggs`.

- Ikiwa `has_enough_eggs` ni `True`, weka thamani upya ya kigezo chako cha `available_eggs` kwa matokeo ya `use_eggs(available_eggs, 1)` na chapisha ujumbe `Made a fried egg. Yummy!`.

- Vinginevyo, ndani ya sehemu ya `else`, chapisha ujumbe `Could not make a fried egg. Not enough eggs!`.

# --hints--

Unapaswa kuwa na tamko la `if` na utumie `has_enough_eggs` kama sharti lake.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

Ndani ya sehemu yako ya `if`, unapaswa kuweka thamani upya ya `available_eggs` kwa matokeo ya `use_eggs(available_eggs, 1)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

Unapaswa kuchapisha ujumbe `Made a fried egg. Yummy!` wakati `has_enough_eggs` ni `True`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

Unapaswa kuwa na sehemu ya `else` inayochapisha ujumbe `Could not make a fried egg. Not enough eggs!` wakati `has_enough_eggs` ni `False`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
