---
id: 6a3e2b9b7eeca084dd4fb857
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Karibu jikoni! Kama mpishi anavyotegemea mapishi kuandaa chakula, ndivyo utakavyotegemea vitendakazi kujenga programu. Katika warsha hii, utajifunza jinsi ya kuandika vitendakazi vya Python vinavyoweza kutumika tena na kumiliki wigo, ukielewa kwa usahihi viambato gani (vigezo) vinapatikana katika msimbo wako wakati wowote.

Ili kuanza, unapaswa kufafanua data za msingi kwa hesabu ya jikoni yako. Unda vigezo vitatu: `available_eggs`, `available_flour`, na `available_sugar`, na weka thamani zao kuwa `1`, `2`, na `3`, kwa mtiririko huo.

# --hints--

Unapaswa kufafanua kigezo kinachoitwa `available_eggs` na kuweka thamani yake kuwa `1`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

Unapaswa kufafanua kigezo kinachoitwa `available_flour` na kuweka thamani yake kuwa `2`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

Unapaswa kufafanua kigezo kinachoitwa `available_sugar` na kuweka thamani yake kuwa `3`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
