---
id: 6a3e58891e1ff13d07f4ab48
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Python husoma sehemu ya `def` na kukumbuka kitendakazi, lakini haifanyi msimbo uliomo ndani yake mara moja.

Ili kutekeleza msimbo na kuona matokeo yako, unahitaji **kufanya wito** (au kuitisha) kitendakazi. Fanya hivi kwa kuandika jina la kitendakazi kifuatiwe na mabano ya kawaida.

Kwa mfano:

```py
my_function()
```

Mwisho wa msimbo wako, nje ya kitendakazi, fanya wito kwa kitendakazi chako cha `check_kitchen_stock`.

# --hints--

Unapaswa kufanya wito kwa kitendakazi cha `check_kitchen_stock` mwishoni mwa msimbo wako. Hakikisha hakijaandikwa kwa njia ya kuingizwa (indented), ili kifanye kazi kwa kiwango cha jumla.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
