---
id: 6a3e5475c35dbc138d77c6f4
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Sasa umehesabu jumla yako, unapaswa kuonyesha hesabu ya bidhaa. Katika Python, unaweza kutumia kitendakazi cha `print()` pamoja na **f-strings** (mfuatano wa herufi uliopangwa) kwa urahisi kutoa maandishi na vigezo pamoja.

Kwa mfano:

```py
name = 'Alice'
print(f'Hello, {name}!') # Output: Hello, Alice!
```

Ndani ya kitendakazi chako cha `check_kitchen_stock`, chini ya kigezo chako cha `total_items`, ongeza kauli hizi nne za chapisho kama zilivyo:

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

Unapaswa kuwa na kauli ya `print(f'The kitchen has {total_items} total items:')` ndani ya kitendakazi chako.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

Kauli ya `print(f'The kitchen has {total_items} total items:')` inapaswa kuwa wito wa kwanza wa `print()` ndani ya kitendakazi chako.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

Unapaswa kuwa na kauli ya `print(f'- {available_eggs} eggs')` ndani ya kitendakazi chako.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

Unapaswa kuwa na kauli ya `print(f'- {available_flour} flour')` ndani ya kitendakazi chako.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

Unapaswa kuwa na kauli ya `print(f'- {available_sugar} sugar')` ndani ya kitendakazi chako.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
