---
id: 6a3e5ccb99fd6053af9032a1
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Sasa unajua nini kiko jikoni mwako, unaweza kutumia vitu hivyo kupika. Ili kufanya hivyo, tengeneza kitendakazi kipya kinachokubali vigezo. Fikiria vigezo kama viwakilishi vya data unazopitisha ndani ya kitendakazi wakati wa kufanya wito.

Hapa kuna kitendakazi kinachokubali vigezo viwili, `a` na `b`, na kinachochapisha jumla yao kwenye terminali.

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

Chini ya kitendakazi chako cha `check_kitchen_stock`, fafanua kitendakazi tupu kinachoitwa `use_eggs` kinachokubali vigezo viwili: `available_eggs` na `eggs_to_use`, kwa mpangilio huo.

Kumbuka kwamba ukimpa kigezo jina sawa kabisa na kigezo cha kimataifa (kama `available_eggs`), kitendakazi kitatumia kigezo chake cha ndani `available_eggs` badala ya kile cha kimataifa. Hii inajulikana kama kujificha kwa jina kwa sababu kigezo cha ndani "hujificha" kigezo cha kimataifa.

Tumia neno kuu `pass` ndani ya mwili kwa sasa.

# --hints--

Unapaswa kufafanua kitendakazi kinachoitwa `use_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

Kitendakazi chako cha `use_eggs` kinapaswa kuwa na `available_eggs` na `eggs_to_use` kama vigezo vyake.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
