---
id: 6a3e6c0eb5d050ce07781709
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Kama sharti la `if` linatathminiwa kuwa `False` (maana una mayai ya kutosha), kitendakazi kitaepuka sehemu ya `if` na kuendelea kuendesha msimbo uliobaki.

Sasa unahitaji kusindika data yako na kurudisha thamani mpya iliyohesabiwa.

Nje ya sehemu yako ya `if` (kwa kiwango sawa cha kuingizwa kama neno kuu `if`), tumia wito wa `print()` na f-string kuchapisha `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`.

Kisha, kwenye mstari unaofuata, `return` idadi iliyobaki ya mayai kwa kutoa `eggs_to_use` kutoka `available_eggs`.

# --hints--

Unapaswa kuwa na tamko la chapisho linalochapisha `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'` katika kitendakazi chako cha `use_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

Unapaswa kurudisha idadi iliyobaki ya mayai kwa kutoa `eggs_to_use` kutoka `available_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

Unapaswa kurudisha `available_eggs - eggs_to_use` baada ya tamko la `print`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
