---
id: 6a3e71a0a3a44260e675b157
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Ili kutatua fumbo la kwanini hesabu yako ya mayai haikubadilika na pia kupika kifungua kinywa, tengeneza kitendakazi kipya. Fikiria kitendakazi hiki kama mapishi yako, ambacho hufanya wito kwa vitendakazi vingine ili kazi ifanyike.

Chini ya usemi wako wa `print(available_eggs)`, fafanua kitendakazi tupu kinachoitwa `make_fried_egg` kinachopokea `available_eggs` kama kigezo chake pekee.

Tumia neno kuu `pass` ndani ya mwili kwa sasa.

# --hints--

Unapaswa kufafanua kitendakazi kinachoitwa `make_fried_egg`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

Kitendakazi chako cha `make_fried_egg` kinapaswa kuwa na `available_eggs` kama kigezo chake.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
