---
id: 6a3e6fd08125dfa615cb431b
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Sasa kitendakazi chako cha `use_eggs` kimekamilika, unaweza kukijaribu. Unapofanya wito wa kitendakazi kinachohitaji **vigezo**, unapitisha thamani maalum ndani ya mabano ya kawaida. Thamani hizi zinajulikana kama **hoja**.

Hapa kuna mfano wa jinsi ya kuita kitendakazi kwa hoja:

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

Katika mfano huu, mistari miwili ya kwanza ni ufafanuzi wa kitendakazi, ambapo `a` na `b` ni vigezo. Mstari wa mwisho, `calculate_sum(3, 1)`, ni wito wa kitendakazi. Thamani `3` na `1` ni hoja. Zinapangwa kulingana na nafasi zao kwa vigezo, hivyo `3` imepewa thamani `a` na `1` imepewa thamani `b`.

Chini kabisa ya msimbo wako, fanya wito wa kitendakazi chako cha `use_eggs` na upitishe hoja mbili: kigezo cha kimataifa cha `available_eggs` na nambari `1` (inayowakilisha idadi ya mayai unayotaka kutumia).

# --hints--

Unapaswa kufanya wito wa kitendakazi cha `use_eggs` na kupitisha `available_eggs` na `1` kama hoja.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
