---
id: 6565a536ba1f9f25bd30e88b
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este proyecto, implementarás el Algoritmo de Luhn. Este algoritmo es una fórmula para validar una variedad de números de identificación.

Comienza declarando una función llamada `main`, esta servirá como el punto de entrada del programa. Usa la palabra clave `pass` para evitar un error.

# --hints--

Debes tener `def main():` en tu código.

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

Debes tener `pass` en tu código.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
