---
id: 656877c202f44995ed952608
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Reemplaza la sentencia `pass` con una variable llamada `card_number` y un valor de `'4111-1111-4555-1142'`.

# --hints--

Debes tener `card_number = '4111-1111-4555-1142'` dentro de la función `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

No deberías tener un `pass` en tu código.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
