---
id: 656878585631369a6b2d2191
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Definir la traducción no traduce la cadena por sí misma. Se debe llamar al método `translate` en la cadena que se desea traducir con la tabla de traducción como argumento:

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

Crea una variable llamada `translated_card_number` y asígnale el resultado de llamar al método `translate` sobre `card_number` con `card_translation` como argumento.

# --hints--

Debes crear una variable `translated_card_number` dentro de `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

Debes asignar a `translated_card_number` un valor de `card_number.translate(card_translation)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
