---
id: 65687c8d86e18cbd775a53c9
title: Paso 26
challengeType: 20
dashedName: step-26
---

# --description--

La siguiente parte del Algoritmo de Luhn es multiplicar todos los dígitos pares por `2`.

Dentro del ciclo `for` de dígitos pares, reemplaza la llamada a `print` con una variable llamada `number` y asígnale el valor de `digit` multiplicado por `2`.

# --hints--

Debes tener `number = int(digit) * 2` dentro del ciclo `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /number\s*=\s*int\(\s*digit\s*\)\s*\*\s*2/);
    }
})
```

No deberías tener `print(digit)` dentro del `for` loop.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\s*\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        print(digit)
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
