---
id: 65687db6a9fcf3c548a6c1b7
title: Paso 28
challengeType: 20
dashedName: step-28
---

# --description--

Parte del algoritmo es duplicar cada segundo dígito, comenzando desde la derecha. Si el resultado de duplicar el número es mayor o igual a `10`, suma los dos dígitos. Por ejemplo, si el dígito es `6`, duplícalo para obtener `12`. Suma `1` y `2` para obtener `3`. Puedes hacer esto usando *integer division* para obtener el primer dígito y el operador módulo (`%`) para obtener el segundo dígito:

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

La división entera da como resultado el cociente de la división, redondeado hacia abajo al entero más cercano.

Dentro de la sentencia `if`, asigna a `number` el resultado de `number // 10` (división entera) más el módulo de `number` y `10`.

# --hints--

Debes tener `number = number // 10 + number % 10` dentro de la sentencia `if`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
