---
id: 656879613a96aba21cbe80b9
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Dentro de tu función `main`, llama a la función `verify_card_number` y pasa la variable `translated_card_number` como argumento.

# --hints--

Debes tener `verify_card_number(translated_card_number)` dentro de la función `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +verify_card_number\(\s*translated_card_number\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

--fcc-editable-region--
main()
```
