---
id: 65687b48f2201ab32e06c37c
title: Paso 21
challengeType: 20
dashedName: step-21
---

# --description--

Actualmente, tu script lanza un `TypeError` porque estás intentando sumar una cadena a un entero. Puedes arreglar esto convirtiendo la variable `digit` a un entero antes de sumarla a `sum_of_odd_digits`, usando la función incorporada `int`:

```python
my_string = '123'
my_int = int(my_string)
```

Convierte la variable `digit` a un entero antes de sumarla a `sum_of_odd_digits`. Luego, mueve la llamada a `print` al final de la función `verify_card_number` para imprimir el valor de `sum_of_odd_digits`.

# --hints--

Debes tener `sum_of_odd_digits += int(digit)` dentro del ciclo `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

Debes tener `print(sum_of_odd_digits)` al final de la función `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

Ya no deberías tener `print(digit)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
