---
id: 656877f71bba2b97acedf9af
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Python viene con clases integradas que pueden ayudarnos con la manipulación de cadenas. Una de ellas es la clase `str`. Tiene un método llamado `maketrans` que puede ayudarnos a crear una tabla de traducción. Esta tabla puede usarse para reemplazar caracteres en una cadena:

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

Lo anterior, cuando se llama en una cadena, reemplazará todos los caracteres `t` con `c` y todos los caracteres `l` con `b`.

Crea una variable llamada `card_translation` y asígnale una tabla de traducción para reemplazar todos los caracteres `-` y ` ` con una cadena vacía.

# --hints--

Debes crear una variable `card_translation` dentro de `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

Debes asignar a `card_translation` un valor de `str.maketrans({'-': '', ' ': ''})`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
