---
id: 65687ad86596e0af38640a84
title: Paso 19
challengeType: 20
dashedName: step-19
---

# --description--

Usa un bucle `for` para iterar sobre cada dígito en la lista `odd_digits`. Mueve tu llamada a `print` del paso anterior dentro del bucle `for` y cámbiala para imprimir cada dígito.

# --hints--

Debes usar un ciclo `for` sobre `odd_digits`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /for\s+\w+\s+in\s+odd_digits/);
    }
})
```

Debes tener `--fcc-expected--` dentro del ciclo `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;
    
        // Get variable name used in for loop
        const for_loop_variable = function_body.match(/for\s+(\w+)\s+in\s+odd_digits/)?.[1];
        assert.exists(for_loop_variable);

        assert.equal(function_body.match(new RegExp(`print\\(\\s*${for_loop_variable}\\s*\\)`))?.[0], `print(${for_loop_variable})`);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    print(odd_digits)
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
