---
id: 656879a66338b2a461d5d307
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

El algoritmo de Luhn es el siguiente:

1. De derecha a izquierda, duplica el valor de cada segundo dígito; si el producto es mayor que 9, suma los dígitos del producto.
2. Suma todos los dígitos.
3. Si la suma de todos los dígitos es un múltiplo de 10, entonces el número es válido; de lo contrario, no es válido.

Asume un ejemplo de un número de cuenta "7992739871" al que se le añadirá un dígito de verificación, formando así 7992739871x:

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

Reemplaza la sentencia `pass` con una variable `sum_of_odd_digits` y un valor de `0`.

# --hints--

Debes tener `sum_of_odd_digits = 0` dentro de la función `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

No debes tener `pass` dentro de la función `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
