---
id: 69a04466917c66b132323a1b
title: Paso 18
challengeType: 1
dashedName: step-18
---

# --description--

Si `previous` no es `null`, esto significa que el elemento a eliminar no está al inicio sino en algún lugar dentro de la lista, por lo que el nodo actual debe ser omitido enlazando el nodo anterior directamente con el siguiente nodo.

Para hacer eso, crea otra sentencia `if` que verifique si previous es estrictamente distinto de `null`. Si ese es el caso, asigna `previous.next` a `current.next`.

# --hints--

Debes crear otra sentencia `if` con la condición `previous !== null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*/);
```

Debes asignar `previous.next` a `current.next`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*previous\.next\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--
    
--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
