---
id: 69a00eaa43458e618331873c
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Ahora trabajarás en la función que añadirá un nodo a la lista.

Crea una función `add` con los parámetros `list` y `element`.

# --hints--

Debes crear una función `add`.

```js
assert.isFunction(add)
```

Tu función `add` debe tener los parámetros `list` y `element`.

```js
const addParams = __helpers.getFunctionParams(add.toString());
const addParamsNames = addParams.map(param => param.name)
assert.include(addParamsNames, "list");
assert.include(addParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
