---
id: 69a0179a85b00b6127e0af45
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Dentro del ciclo `while`, actualiza `current` a `current.next` para que el ciclo avance a través de la lista un nodo a la vez, y se detenga cuando llegue al nodo con una propiedad `next` igual a `null`, es decir, el final de la lista.

# --hints--

Debes actualizar el valor de tu variable `current` a `current.next` dentro del ciclo `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
