---
id: 69a014bd9b1730e94136279f
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Si la lista no está vacía, entonces quieres encontrar su final, añadir el nuevo elemento allí y hacer que apunte a `null`. Esto significa que tienes que recorrer la lista desde `head`.

Para comenzar eso, abre una cláusula `else`. Dentro de ella, usa `let` para crear una variable `current` y asignarle `list.head`.

# --hints--

Deberías tener una cláusula `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

Deberías crear una variable `current` asignada a `list.head` dentro de la cláusula `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
