---
id: 69a04088e600bd8074a9852f
title: Paso 17
challengeType: 1
dashedName: step-17
---

# --description--

Si `current` es `null`, significa que ese es el final de la lista, así que la función debe dejar de ejecutarse.

Crea una sentencia `if` que verifique si `current` es estrictamente igual a `null`. Dentro de la sentencia `if`, usa `return` para salir de la función.

# --hints--

Debes crear una sentencia `if` con la condición `current === null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*/);
```

Debes usar la palabra clave `return` dentro de tu sentencia `if`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*return;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
