---
id: 69a0102e761a245bb1cbd8e7
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

En este punto, deberías probar la lista enlazada para entenderla mejor.

Crea una variable `myList` y asígnale `initList()`. Debajo, registra en la consola una llamada a `isEmpty()` usando `myList` como parámetro. Luego, después de eso, llama a `add` con `myList` y un número de tu elección, y registra `myList` en la consola. Debajo de eso, registra otra llamada a `isEmpty()` en la consola usando `myList` como parámetro nuevamente.

# --hints--

Deberías crear una variable `myList` asignada a `initList()`.

```js
assert.deepEqual(myList, initList());
```

Deberías registrar en la consola una llamada a `isEmpty(myList)`.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(isEmpty(myList))');
```

Deberías llamar a `add` con `myList` y un número de tu elección como parámetros.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.match(strippedCode, /add\(myList,\d+\)/);
```

Deberías registrar tu variable `myList` en la consola.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(myList)');
```

Deberías registrar `isEmpty(myList)` en la consola otra vez.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
const regexMatch = strippedCode.matchAll(/console\.log\(isEmpty\(myList\)\)/g);
assert.isAtLeast([...regexMatch].length, 2);
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
