---
id: 69a012c6224fac85832247dd
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

¿Notas los resultados `true`, `{ head: null, length: 0 }` y `true` en la consola? Esto sucede porque no has asignado nada a la lista `myList`, y no has incrementado la propiedad `length`.

Adelante, asigna `node` a `list.head` y, debajo de eso, incrementa la propiedad `length` de `list`.

# --hints--

Debes asignar `node` a `list.head`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'list.head=node');
```

Debes incrementar la propiedad `length` de `list`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'list.length++');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
