---
id: 69a00a051b9a18694d00c16b
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, implementarás una lista enlazada.

Para comenzar, crea una función llamada `initList`. Dentro de ella, devuelve un objeto con una propiedad `head` establecida en `null` y una propiedad `length` con valor `0`.

Esta es la función que llamarás para inicializar una lista enlazada. `head` representa el primer elemento de la lista; inicialmente es `null` porque no hay ningún elemento en la lista. `length` lleva la cuenta del número de nodos (elementos) en la lista, que incrementarás o disminuirás al agregar o eliminar un nodo de la lista.

# --hints--

Debes crear una función `initList`.

```js
assert.isFunction(initList)
```

Tu función `initList` debe devolver un objeto.

```js
assert.isObject(initList());
```

Tu función `initList` debe devolver un objeto con una propiedad `head` establecida en `null`.

```js
assert.propertyVal(initList(), "head", null);
```

Tu función `initList` debe devolver un objeto con una propiedad `length` establecida en `0`.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
