---
id: 69a00f673f171e29159e3ad8
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Este nodo necesitaría dos propiedades, el elemento a añadir y un puntero al siguiente nodo.

Dentro de la función add, crea un objeto llamado `node`, debe tener una propiedad llamada `element` y una propiedad `next` establecida en `null`, que es el puntero al siguiente elemento de la lista.

# --hints--

Debes crear un objeto llamado `node` dentro de tu función `add`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

Tu objeto `node` debe tener una propiedad llamada `element` establecida en `element`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

Tu objeto `node` debe tener una propiedad llamada `next` establecida en `null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
