---
id: 69a0102e761a245bb1cbd8e7
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

En este punto, deberías probar la lista enlazada para entenderla mejor.

Crea una variable `myList` y asígnale `initList()`. Debajo, registra en la consola una llamada a `isEmpty()` usando `myList` como parámetro. Luego, después de eso, llama a `add` con `myList` y un número de tu elección, y registra `myList` en la consola. Debajo de eso, registra otra llamada a `isEmpty()` en la consola usando `myList` como parámetro nuevamente.

# --hints--

Deberías crear una variable `myList` asignada a `initList()`.

```js
assert.deepEqual(myList, initList());
```

Deberías registrar en la consola una llamada a `isEmpty(myList)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\)/)
```

Deberías llamar a `add` con `myList` y un número de tu elección como parámetros.

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

Deberías registrar tu variable `myList` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\)/)
```

Deberías registrar `isEmpty(myList)` en la consola otra vez.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\)\s*console\.log\(isEmpty\(myList\)\)/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
