---
id: 69a00f673f171e29159e3ad8
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Este nodo necesitaría dos propiedades, el elemento a añadir y un puntero al siguiente nodo.

Dentro de la función add, crea un objeto llamado `node`, debe tener una propiedad llamada `element` y una propiedad `next` establecida en `null`, que es el puntero al siguiente elemento de la lista.

# --hints--

Debes crear un objeto llamado `node` dentro de tu función `add`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
assert.match(node?.value.toString(), /^\{[\s\S]*\}/);
assert.exists(node);
```

Tu objeto `node` debe tener una propiedad llamada `element` establecida en `element`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
const { element } = node?.objectProps ?? {};
assert.isTrue(element?.value.matches('element'));
```

Tu objeto `node` debe tener una propiedad llamada `next` establecida en `null`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
const { next } = node?.objectProps ?? {};
assert.isTrue(next?.value.matches('null'));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
