---
id: 69a03aaa1e6fe8152788ce6a
title: Paso 14
challengeType: 1
dashedName: step-14
---

# --description--

Dentro de la funciĂ³n `remove`, usa `let` para crear una variable `previous` asignada a `null`, y `let` nuevamente para crear una variable `current` asignada a `list.head`. Estas son las variables que usarĂ¡s en el ciclo que recorrerĂ¡ la lista.

`current` rastrea el nodo que se estĂ¡ revisando actualmente en la lista, comenzando por la cabeza y avanzando un nodo a la vez. `previous`, por otro lado, lleva el seguimiento del nodo justo antes del actual.

Esto es importante al eliminar un nodo porque debes actualizar el valor next del nodo anterior para que apunte al nodo despuĂ©s del que se estĂ¡ eliminando, con el fin de mantener la lista conectada.

# --hints--

Debes usar `let` para crear una variable `previous` asignada a `null` dentro de tu funciĂ³n `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

Debes usar `let` para crear una variable `current` asignada a `list.head` dentro de tu funciĂ³n `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
