---
id: 69a030fbe487f185f55b8627
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

Ahora, trabajarás en una función para eliminar un nodo de la lista.

Crea una función `remove` con los parámetros `list` y `element`.

# --hints--

Deberías crear una función `remove`.

```js
assert.isFunction(remove)
```

Tu función `remove` debe tener los parámetros `list` y `element`.

```js
const removeParams = __helpers.getFunctionParams(remove.toString());
const removeParamsNames = removeParams.map(param => param.name)
assert.include(removeParamsNames, "list");
assert.include(removeParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

--fcc-editable-region--

--fcc-editable-region--

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
