---
id: 69a01864d343dc2f2484d131
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

Después de encontrar el último nodo, necesitas conectar el nuevo nodo a él. Para hacer eso, asigna `current.next` a `node` dentro del ciclo while pero aún dentro de la cláusula `else`.

# --hints--

Debes asignar `current.next` a `node` dentro de tu ciclo `while`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'current.next=node');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
