---
id: 69a045fd236d0d973a59f103
title: Paso 19
challengeType: 1
dashedName: step-19
---

# --description--

Si `previous` es `null`, esto significa que el elemento a eliminar está al inicio de la lista, por lo que debes actualizar `head` para que apunte al siguiente nodo y así eliminar el primer elemento.

Para hacer eso, añade una cláusula `else`. Dentro de ella, asigna `list.head` a `current.next`.

# --hints--

Debes añadir una cláusula `else` a tu sentencia `if`.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*/);
```

Debes asignar `list.head` a `current.next` dentro de la cláusula `else`.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*list\.head\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    } 
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
