---
id: 69a00eaa43458e618331873c
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Ahora trabajarás en la función que añadirá un nodo a la lista.

Crea una función `add` con los parámetros `list` y `element`.

# --hints--

Debes crear una función `add`.

```js
assert.isFunction(add)
```

Tu función `add` debe tener los parámetros `list` y `element`.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.add.parameters;
assert.equal(parameters[0].toString(), "list");
assert.equal(parameters[1].toString(), "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
