---
id: 69a013a86c97dac3784b2636
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Ahora, el registro debería mostrar algo como `{ head: { element: 42, next: null }, length: 1 }` en la consola.

Sin embargo, esta no es la forma correcta de manejar un caso donde el nodo agregado es el primer elemento de la lista, es decir, la cabeza.

Para hacerlo correctamente, usa una sentencia if para verificar si la lista está vacía llamando a la función `isEmpty` con el `list`. Si eso es cierto, entonces mueve tu línea `list.head = node` dentro del cuerpo de la sentencia `if`.

# --hints--

Deberías tener una sentencia `if` con la condición `isEmpty(list)`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

Debes mover tu línea `list.head = node` dentro de la sentencia `if`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
