---
id: 69a00cd4ec513345d40e95ec
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Lo siguiente en lo que trabajarás es una función que te permita verificar si la lista está vacía.

Para eso, crea una función llamada `isEmpty` que reciba un parámetro `list`. Dentro de la función, devuelve una comprobación para ver si la propiedad `length` de `list` es estrictamente igual a `0`.

# --hints--

Debes crear una función `isEmpty`.

```js
assert.isFunction(isEmpty)
```

Tu función `isEmpty` debe tener un parámetro `list`.

```js
const isEmptyParams = __helpers.getFunctionParams(isEmpty.toString());
assert.deepEqual(isEmptyParams[0].name, "list");
```

Tu función `isEmpty` debe devolver `list.length === 0`.

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
