---
id: 62cc5b1779e4d313466f73c5
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

Nest a `link` element within the `head` element. Give it a `rel` attribute with a value of `stylesheet` and an `href` attribute with a value of `styles.css`.

# --hints--

代碼應該有一個 `link` 元素。

```js
assert.isNotNull(document.querySelector('link'));
```

不應更改現有的 `head` 標籤。 確保沒有刪除結束標籤。

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Your `link` element should be a void element, it should not have an end tag `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Your `link` element should be inside your `head` element.

```js
const headContentRegex = /(?<=<head\s*>)(?:.|\s*)*?(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 元素應該有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
