---
id: 62cc5b1779e4d313466f73c5
title: 步驟5
challengeType: 0
dashedName: step-5
---

# --description--

在 `head` 元素中嵌套一個 `link` 元素。 設置其 `rel` 屬性值爲 `stylesheet`，`href` 屬性值爲 `styles.css`。

# --hints--

代碼應該有一個 `link` 元素。

```js
assert.isNotNull(document.querySelector('link'));
```

不應更改現有的 `head` 標籤。 確保沒有刪除結束標籤。

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

你應該有一個 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

你的 `link` 元素應該是一個空元素，它不應該有結束標籤 `</link>`。

```js
assert.notMatch(code, /<\/link>/);
```

你的 `link` 元素應該在你的 `head` 元素中。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
