---
id: 655cd899f8de09431eabb40c
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste projeto, você aprenderá estruturas de dados construindo o algoritmo `merge sort`. 

Este é um algoritmo de ordenação que usa o princípio de dividir e conquistar para ordenar coleções de dados.Ou seja, ele 'divide' uma coleção em subpartes menores e 'conquista' as subpartes ordenando-as independentemente, depois mescla as subpartes ordenadas.

Crie uma função chamada `merge_sort`.Esta função irá lidar com a tarefa de ordenar uma lista de números. 

Use a palavra-chave `pass` no corpo da função.

# --hints--

Você deve declarar uma função chamada `merge_sort`.Não esqueça de usar a palavra-chave `pass` no corpo da função.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
