---
id: 6564aee9c077774ea49c3faf
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

O algoritmo de merge sort realiza principalmente três ações:

- Divida uma sequência não ordenada de itens em subpartes
- Ordene os itens nas subpartes
- Mesclar as subpartes ordenadas

O acima acontece recursivamente até que as subpartes sejam mescladas na sequência ordenada completa. Vamos começar dividindo a sequência.

Primeiro, substitua a palavra-chave `pass` por uma variável `middle_point`. Use o operador de divisão inteira (`//`) para dividir o comprimento da lista `array` por 2 e atribuir o resultado à sua nova variável `middle_point`. Lembre-se de indentar seu código.

# --hints--

Você não deve ter `pass` no corpo da função.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.notMatch(function_body, /^\s{4}pass/m);
  }
})
```

Você deve declarar uma variável `middle_point` e atribuir `len(array) // 2` a ela.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
