---
id: 6564a9fe51964c229d5b7f4c
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Você precisará de um parâmetro que denote a coleção de dados a ser ordenada.Crie um parâmetro chamado `array` na função `merge_sort`.

# --hints--

Sua função `merge_sort` deve receber um único parâmetro: `array`.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
