---
id: 656627b47bd2d2a4afbc945d
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

No passo anterior, você obteve o ponto médio.Você pode usá-lo para dividir `array` em duas partes e atribuir cada parte a novas variáveis. 

Use a sintaxe de slice para extrair a metade esquerda de `array` e atribuí-la a uma variável chamada `left_part`.


# --hints--

Você deve ter uma variável chamada `left_part` na sua função `merge_sort`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

Você deve atribuir `array[:middle_point]` à variável `left_part`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
