---
id: 656639771fed09b5c0e8fe71
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Agora que você dividiu a lista `array` em duas listas separadas, você continuará dividindo cada lista até que cada elemento esteja sozinho em sua própria lista.Uma lista com um único número está sempre ordenada.

Para fazer isso, chame recursivamente `merge_sort` dentro da sua função. Não passe nenhum argumento por enquanto.

# --hints--

Você deve chamar `merge_sort` no final do corpo da sua função.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
