---
id: 656680b0fc79f2c38a34d90e
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Agora é hora de ordenar e mesclar as listas (`left_part` e `right_part`) no `array` original. 

Você pode fazer isso comparando os elementos em ambas as listas e mesclando o elemento menor à lista principal.Você vai fazer essa comparação para todos os índices em `left_part` e `right_part`.

Crie três variáveis: `left_array_index`, `right_array_index` e `sorted_index` e defina seus valores como `0`.Essas variáveis vão ajudar você a acompanhar cada índice durante o processo de ordenação.

# --hints--

Você deve ter uma variável chamada `left_array_index` dentro da sua função `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index/);
  }
})
```

O valor de `left_array_index` deve ser 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index\s*=\s*0/);
  }
})
```

Você deve ter uma variável chamada `right_array_index` dentro da sua função `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index/);
  }
})
```

O valor de `right_array_index` deve ser 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index\s*=\s*0/);
  }
})

```

Você deve ter uma variável chamada `sorted_index` dentro da sua função `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index/);
  }
})
```

O valor de `sorted_index` deve ser 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index\s*=\s*0/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)

--fcc-editable-region--

--fcc-editable-region--
```
