---
id: 656706afd65547d22bee0b68
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Dentro do loop `while`, substitua `pass` por uma instrução `if` que verifica se o índice de `left_part` é menor que o índice de `right_part`.

Use a palavra-chave `pass` no corpo da declaração `if`.

# --hints--

Você deve substituir a palavra-chave `pass` por uma instrução `if` com esta condição: `left_part[left_array_index] < right_part[right_array_index]`.Adicione a palavra-chave `pass` ao corpo da declaração `if`.

```js
({
 test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_body().is_equivalent("if left_part[left_array_index] < right_part[right_array_index]:\\n  pass")`))
 })
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        pass
--fcc-editable-region--
```
