---
id: 656702f8b4cbd8cbf0a433c6
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Dentro da sua função, crie um loop `while` que compare um elemento em `left_part` com um elemento em `right_part` e una o elemento menor à lista principal `array`. 

Crie duas condições para o loop: uma que verifica se o `left_array_index` é menor que o comprimento de `left_part` e outra condição que verifica se o `right_array_index` é menor que o comprimento de `right_part`.

# --hints--

Seu loop `while` deve ter estas condições: `left_array_index < len(left_part) and right_array_index < len(right_part)`. Não esqueça de usar a palavra-chave `pass` no corpo do loop.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

Use a palavra-chave `pass` no corpo do loop

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
