---
id: 656709e50ed928da35260276
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

Em uma etapa anterior, você atribuiu o elemento em `left_part` à lista `array` porque ele era menor.Mas isso nem sempre será o caso.Em alguns casos de comparação, o elemento à direita pode ser menor. 

Crie uma cláusula `else` para executar quando o índice `left_part` não for menor que o índice `right_part`.

Dentro do bloco `else`, atribua `right_part[right_array_index]` a `array[sorted_index]`.

# --hints--

Você deve ter `array[sorted_index] = right_part[right_array_index]` no corpo da sua cláusula `else`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /else:\s*[^}]*array\s*\[\s*sorted_index\s*\]\s*=\s*right_part\s*\[\s*right_array_index\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
