---
id: 6569b743630ee592a65a7e2f
title: Passo 26
challengeType: 20
dashedName: step-26
---

# --description--

Antes de testar a função `merge_sort()`, você precisa criar um caso base que pare a execução da função quando o tamanho do `array` for menor ou igual a `1`.

Este caso base irá parar a chamada de recursão. Sem ele, a operação de merge sort continuaria a ser executada mesmo quando a lista já estivesse ordenada ou não tivesse nenhum elemento.

Logo após a declaração da função, crie uma instrução `if` com esta condição: `len(array) <= 1`. Use a palavra-chave `pass` no corpo da função.

# --hints--

Você deve criar uma declaração `if` para verificar se `len(array) <= 1`. Não esqueça de usar a palavra-chave `pass` no corpo da declaração `if`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):

--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
