---
id: 6577254891048c8f2c19e961
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Chame a função `merge_sort()` novamente. Não passe nenhum argumento por enquanto.

# --hints--

Você deve chamar `merge_sort()` novamente após a chamada da função anterior.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
