---
id: 68e45a2b4ab9b2f48b1cc6df
title: 가격 플랜 레이아웃 페이지 디자인하기
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리**

1. 페이지에 텍스트가 `h1`인 `Pricing Plans` 요소가 있어야 합니다.
2. `div` 요소 아래에 `pricing-container` 클래스를 가진 `h1` 요소가 있어야 합니다.
    - `.pricing-container` 선택자는 `display` 값을 가진 `flex` 속성과 `flex-wrap` 값을 가진 `wrap` 속성을 가져야 합니다.
3. `.pricing-container` 요소 안에 가격 플랜을 나타내기 위해 `div` 클래스를 가진 `pricing-card` 요소 세 개가 있어야 합니다.
    - `.pricing-card` 요소 중 하나는 `basic-plan` 클래스 외에 `pricing-card` 클래스를 가져야 합니다.
    - `.pricing-card` 요소 중 하나는 `pro-plan` 클래스 외에 `pricing-card` 클래스를 가져야 합니다.
    - `.pricing-card` 요소 중 하나는 `premium-plan` 클래스 외에 `pricing-card` 클래스를 가져야 합니다.
4. `.basic-plan` 요소에는 텍스트가 `h2`인 `Basic` 요소가 있어야 합니다.
5. `.basic-plan` 요소에는 텍스트가 `p`인 `$9/month` 요소가 있어야 합니다.
6. `.pro-plan` 요소에는 텍스트가 `h2`인 `Pro` 요소가 있어야 합니다.
7. `.pro-plan` 요소에는 텍스트가 `p`인 `$19/month` 요소가 있어야 합니다.
8. `.premium-plan` 요소에는 텍스트가 `h2`인 `Premium` 요소가 있어야 합니다.
9. `.premium-plan` 요소에는 텍스트가 `p`인 `$29/month` 요소가 있어야 합니다.
10. 각 `.pricing-card` 요소는 다음을 해야 합니다:
    - Flexbox를 사용해 내용을 열 방향으로 쌓고 자식 간 공간을 `space-between`로 정렬하세요.
    - `flex` 속성을 `0 0 200px`로 설정해 일정한 너비를 유지하고 레이아웃에서 크기가 변하지 않도록 하세요.
    - `border` 속성을 `2px solid black`로 설정해 카드들이 공간을 차지하는 방식을 확인하세요.
11. `.basic-plan` 요소가 레이아웃에서 가장 먼저 나타나야 합니다. 이를 위해 `order` 속성을 사용하세요.
12. `.pro-plan` 요소가 레이아웃에서 두 번째로 나타나야 합니다. `order` 속성을 사용하고 `flex-grow` 속성을 `2`로 설정해 다른 플랜보다 더 많은 공간을 차지하게 하세요.
13. `.premium-plan` 요소가 레이아웃에서 마지막에 와야 합니다. 이를 위해 `order` 속성을 사용하세요.

**참고:** HTML에 스타일시트를 연결하고 CSS를 적용하는 것을 잊지 마세요.

# --hints--

`h1` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h1'));
```

페이지에 텍스트가 `h1`인 `Pricing Plans` 요소가 있어야 합니다.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

`div` 요소 아래에 `h1` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h1 + div'));
```

`div` 요소는 `pricing-container` 클래스를 가져야 합니다.

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

`.pricing-container` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

`pricing-container` 선택자는 `display` 속성을 `flex` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

`.pricing-container` 선택자는 `flex-wrap` 속성을 `wrap` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

`div` 요소 안에 `pricing-card` 클래스를 가진 `.pricing-container` 요소 세 개가 있어야 합니다.

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

`.pricing-card` 요소 중 하나는 `basic-plan` 클래스를 가져야 합니다.

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

`.pricing-card` 요소 중 하나는 `premium-plan` 클래스를 가져야 합니다.

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

`.pricing-card` 요소 중 하나는 `pro-plan` 클래스를 가져야 합니다.

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

`.pricing-card` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

`.pricing-card` 선택자는 `display` 속성을 `flex` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

`.pricing-card` 선택자는 `flex-direction` 속성을 `column` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

`.pricing-card` 선택자는 `justify-content` 속성을 `space-between` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

`.pricing-card` 선택자는 `flex` 속성을 `0 0 200px` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

`.pricing-card` 선택자는 `border` 속성을 `2px solid black` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

`.basic-plan` 선택자는 `order` 속성을 `0` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

`.basic-plan` 요소에는 텍스트가 `h2`인 `Basic` 요소가 있어야 합니다.

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

`.basic-plan` 요소에는 텍스트가 `p`인 `$9/month` 요소가 있어야 합니다.

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

`.pro-plan` 요소에는 텍스트가 `h2`인 `Pro` 요소가 있어야 합니다.

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

`.pro-plan` 요소에는 텍스트가 `p`인 `$19/month` 요소가 있어야 합니다.

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

`.pro-plan` 요소는 `order` 속성을 `1` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

`.pro-plan` 요소는 `flex-grow` 속성을 `2` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

`.premium-plan` 요소에는 텍스트가 `h2`인 `Premium` 요소가 있어야 합니다.

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

`.premium-plan` 요소에는 텍스트가 `p`인 `$29/month` 요소가 있어야 합니다.

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

`premium-plan` 요소는 레이아웃에서 마지막 요소여야 하며 `order` 속성을 `2` 값으로 가져야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```
