---
id: 66b102ee0bdbad34a9f85ba0
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会上，你将通过构建教师聊天机器人继续学习字符串。 

首先，添加一个 `console` 语句，其中包含 `"Hi there!"` 的消息。

# --hints--

你应该有一个 `console` 语句。

```js
assert.match(code, /console\.log(.*)/);
```

你的控制台语句应该输出消息 `"Hi there!"`。

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
