---
id: 66b102ee0bdbad34a9f85ba0
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会上，你将通过构建教师聊天机器人继续学习字符串。

首先，添加一个 `console` 语句，其中包含 `"Hi there!"` 的消息。

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

你应该有一个 `console` 语句。

```js
assert.isAbove(logSpy.calls.length, 0);
```

你的控制台语句应该输出消息 `"Hi there!"`。

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
