---
id: 66b59be6ab830800c4df9146
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

现在是时候使用你在上一步中创建的变量了。

In the lectures, you learned how to work with template literals like this:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

首先创建一个 `sentence` 变量。

使用模板字面量语法，将字符串 `Today, you will learn about [topic variable goes here] in [subject variable goes here].` 分配给 `sentence` 变量。

你将用 `topic` 和 `subject` 变量和 `${}` 语法替换 `[topic variable goes here]` 和 `[subject variable goes here]` 占位符。

最后，将 `sentence` 变量记录到控制台。

# --hints--

你应该有一个 `sentence` 变量。

```js
assert.isNotNull(sentence);
```

你的 `sentence` 变量应该是一个字符串。

```js
assert.isString(sentence);
```

你应该使用模板字面量语法将字符串 `Today, you will learn about [topic variable goes here] in [subject variable goes here].` 分配给 `sentence` 变量。 如果你需要帮助，请参考示例。

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

你应该将 `sentence` 变量记录到控制台。

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
