---
id: 66b5928b86e907fdfbf34e56
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

现在是时候设置机器人的名字了。

创建一个名为 `botName` 的变量，并为其分配 `"teacherBot"` 的字符串值。

# --hints--

你应该有一个名为 `botName` 的变量。

```js
assert.isNotNull(botName);
```

你的 `botName` 变量应该是一个字符串。

```js
assert.isString(botName);
```

你的 `botName` 变量应具有 `"teacherBot"` 的值。

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
