---
id: 66b59b12a745e10011158f7b
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

下一步是创建一些将在未来的机器人消息中使用的变量。

创建一个名为 `subject` 的变量并为其分配字符串值 `"JavaScript"`。

然后创建一个名为 `topic` 的变量并为其分配字符串值 `"strings"`。

# --hints--

你应该有一个名为 `subject` 的变量。

```js
assert.isNotNull(subject);
```

你应该为 `subject` 变量分配一个字符串。

```js
assert.isString(subject);
```

你应该将字符串 `"JavaScript"` 分配给你的 `subject` 变量。

```js
assert.equal(subject, "JavaScript");
```

你应该有一个名为 `topic` 的变量。

```js
assert.isNotNull(topic);
```

你应该为 `topic` 变量分配一个字符串。

```js
assert.isString(topic);
```

你应该将字符串 `"strings"` 分配给你的 `topic` 变量。

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
