---
id: 66b59829dba144ff1351220f
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

现在是时候使用 `botName` 变量创建问候语了。

在之前的课程中，你学习了如何使用像这样的模板字面来连接字串：

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

首先创建一个名为 `greeting` 的变量。

接下来，使用模板字面量语法，分配一个字符串，该字符串表示 `My name is`，后跟 `botName` 变量，并以句点结尾（`.`）。

最后，将 `greeting` 变量记录到控制台。

# --hints--

你应该有一个名为 `greeting` 的变量。

```js
assert.isNotNull(greeting);
```

你应该为 `greeting` 变量分配一个字符串。

```js
assert.isString(greeting);
```

你应该使用模板字面量将字符串 `My name is` 与 `botName` 变量后跟句点 (`.`) 连接起来。

```js
assert.equal(greeting, "My name is teacherBot.");
```

你应该有一个 `console` 语句，将 `greeting` 变量记录到控制台。

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
