---
id: 66b6fdb76441c738719039fa
title: 步骤 13
challengeType: 1
dashedName: step-13
---

# --description--

在之前的课程中，你学习了如何像这样访问字串中的最后一个字符：

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` 将始终为你提供字符串的最后一个索引号。

创建一个名为 `lastCharacter` 的新变量，并将 `subject` 变量中最后一个字符的值分配给它。

然后，将 `lastCharacter` 变量的值记录到控制台。

# --hints--

你应该有一个名为 `lastCharacter` 的变量。

```js
assert.isNotNull(lastCharacter);
```

你应该将 `subject` 变量中的最后一个字符的值分配给 `lastCharacter` 变量。 如果你需要帮助，请参考示例。

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

你应该将 `lastCharacter` 变量的值记录到 `console`。

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
