---
id: 66b6e39031393e30f2c48d0f
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

要获取字符串的长度，可以使用 <dfn>length</dfn> 属性。 此属性返回字符串中的字符数。

这是一个示例：

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

首先使用 `length` 属性获取 `subject` 字符串的长度，并将该值记录到控制台。

# --hints--


你应该使用 `length` 属性来获取 `subject` 字符串的长度。

```js
assert.match(code, /subject\.length/);
```

你应该将 `subject` 字符串的长度记录到控制台。

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
