---
id: 66b6f80fca500635d1e8af8d
title: 步骤 12
challengeType: 1
dashedName: step-12
---

# --description--

现在是时候访问字符串的最后一个字符了。

首先添加另一个 `console` 语句，输出消息 `Here is an example of accessing the last letter in the word [subject].` 

记住用 `subject` 变量的实际值替换 `[subject]`，并使用正确的模板字面量语法。

# --hints--

你的 `console` 语句应该输出消息 `Here is an example of accessing the last letter in the word [subject].` 请记住使用正确的模板字面量语法。

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+last\s+letter\s+in\s+the\s+word\s+\${subject}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

--fcc-editable-region--

--fcc-editable-region--
```
