---
id: 66b6ffb42b88e33943788abf
title: 步骤14
challengeType: 1
dashedName: step-14
---

# --description--

在研讨会的最后一部分，你将回顾如何在字符串中找到子字符串的索引位置。 请记住，子字符串是字符串的一部分。

首先创建一个名为 `learningIsFunSentence` 的变量，并为其分配 `"Learning is fun."` 的字符串值。

# --hints--

你应该有一个名为 `learningIsFunSentence` 的变量。

```js
assert.isNotNull(learningIsFunSentence);
```

你的 `learningIsFunSentence` 应该保存一个字符串的值。

```js
assert.isString(learningIsFunSentence);
```

你应该将 `"Learning is fun."` 的值分配给变量 `learningIsFunSentence`。

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
