---
id: 66b6e80d6c3f0b329c360283
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

本研讨会的下一部分是回顾如何从字符串访问字符。

首先向控制台输出消息 `Here is an example of accessing the first letter in the word [subject].`。 

记住用 `subject` 变量替换 `[subject]`，并使用正确的模板字面量语法，就像在前面的步骤中所做的那样。

# --hints--

你的 `console` 语句应输出消息 `Here is an example of accessing the first letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+first\s+letter\s+in\s+the\s+word\s+\${subject}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

--fcc-editable-region--

--fcc-editable-region--
```
