---
id: 66b6e62423e8a031d6c1f03d
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

现在是时候获取 `topic` 字符串的长度了。

你可以在 `console` 语句中使用模板字面量，如下所示：

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

首先向控制台输出消息 `Here is an example of using the length property on the word [topic].`。

记住用 `topic` 变量替换 `[topic]`，并使用正确的模板字面量语法，就像在前面的步骤中所做的那样。

然后，添加第二个 `console.log` 语句，将 `topic` 字符串的长度输出到控制台。

# --hints--

你应该使用模板字面量语法将消息 `Here is an example of using the length property on the word [topic].` 输出到 `console`。

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

你应该有第二个 `console` 语句来将 `topic` 变量的长度输出到 `console`。

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
