---
id: 66b6f586767a1534f3097353
title: 步骤11
challengeType: 1
dashedName: step-11
---

# --description--

现在是时候访问 `subject` 变量的第二个字母了。

首先添加一个输出消息 `Here is an example of accessing the second letter in the word [subject].` 的 `console` 语句，记住用 `subject` 变量的实际值替换 `[subject]`，并使用正确的模板字面量语法。

然后添加另一个 `console` 语句，使用括号符号和正确的索引号输出 `subject` 变量的第二个字母。

# --hints--

你的 `console` 语句应输出消息 `Here is an example of accessing the second letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

你应该有另一个 `console` 语句，使用括号表示法索引 `1` 输出 `subject` 变量的第二个字母。

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
