---
id: 66b6efddeca35833cd6f0b03
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

在之前的课程中，你学习了如何像这样访问字串中的字符：

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

请记住，索引号从 `0` 开始，因此字符串中的第一个字母始终位于索引 `0`。

首先添加另一个 `console` 语句。 

在 `console` 语句中，使用括号表示法和正确的索引号输出 `subject` 变量的首字母。

# --hints--

你的 `console` 语句应使用括号表示法和索引 `0` 输出 `subject` 变量的首字母。

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
