---
id: 66b70b8a611cbf3bcc5c6c5f
title: 步骤 16
challengeType: 1
dashedName: step-16
---

# --description--

在之前的课程中，你学习了如何像这样使用 `indexOf` 方法：

```js
const sentence = "I love to learn.";

// returns index 2
console.log(sentence.indexOf("love"));

// returns -1
console.log(sentence.indexOf("hate"));
```

请记住，`indexOf` 方法返回字符串中子字符串第一次出现的索引位置。 如果未找到子字符串，则返回 `-1`。

添加新的 `console` 语句，输出对 `learningIsFunSentence` 变量使用 `indexOf` 方法查找子字符串 `"Learning"` 的索引位置的结果。

# --hints--

你的 `console` 应该在 `learningIsFunSentence` 变量上使用 `indexOf` 方法来查找子字符串 `"Learning"` 的索引位置。

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(Learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

--fcc-editable-region--

--fcc-editable-region--
```
