---
id: 66b7137348cfb53fd3ec6c73
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

接下来，你将添加一个新的 `console` 语句，该语句输出在 `learningIsFunSentence` 变量上使用 `indexOf` 方法查找子字符串 `"fun"` 的位置的结果。

在该 `console` 语句下方，添加一个新的 `console` 语句，输出使用 `indexOf` 方法查找子字符串 `"learning"` 的位置的结果。

注意最后一个 `console` 语句的输出。 

# --hints--

你的 `console` 应该从 `learningIsFunSentence` 变量输出子字符串 `"fun"` 的索引位置。

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(fun)\1\s*\)\s*\)/);
```

你的第二个 `console` 语句应该从 `learningIsFunSentence` 变量输出子字符串 `"learning"` 的索引位置。

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

console.log(learningIsFunSentence.indexOf("Learning"));

--fcc-editable-region--

--fcc-editable-region--
```
