---
id: 6735a7370e0ae93a4577c689
title: 设计报纸版型
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

在本实验中，你将使用 CSS Grid 设计一个报纸版型，包括网格行和网格列等概念。

**目标：**实现以下用户需求并通过所有测试以完成实验。 

**用户故事：**

1. 你的页面应包含一个带有类 `newspaper-layout` 的 `main` 元素。
1. `.newspaper-layout` 应包含一个带有类 `title` 的 `header`，其中包含一个 `h1` 用于显示报纸名称。
1. `.newspaper-layout` 应包含一个带有 `feature-article` 类的 `article`，用于主要新闻文章。
1. `.feature-article` 应包含一个用于文章标题的 `h2` 元素，后跟一个用于文章内容的 `p` 元素。
1. 你应该为较小的文章添加三个更多的 `article` 元素，类为 `small-article1`、`small-article2` 和 `small-article3`。
1. 每个较小的文章应包括一个用于文章标题的 `h3` 元素，后跟一个用于文章内容的 `p` 元素。
1. `.newspaper-layout` 应该包含一个带有类 `secondary-article` 的 `article` 元素，用于额外的新闻版块。
1. `.secondary-article` 应包含一个用于文章标题的 `h2` 元素，后跟一个用于文章内容的 `p` 元素。
1. `.newspaper-layout` 应该包含一个带有 `cover-image` 类的 `figure`，用于显示表现报纸内容的图像。
1. 具有 `title`、`feature-article`、`secondary-article`、`cover-image`、`small-article1`、`small-article2` 和 `small-article3` 类的元素应将 `grid-area` 属性设置为相同的类名。
1. 你的 `.newspaper-layout` 应该使用 CSS Grid，并带有 `grid-template-areas` 属性来定义各部分的排列：
   - `.title` 应跨越顶行。
   - `.feature-article` 应跨越第二行的两列，`.cover-image` 位于第三列。
   - `.secondary-article` 应跨越第三行的两列，`.cover-image` 位于第三列。
   - 这三篇小文章应出现在第四行。
1. `.newspaper-layout` 应该有一个 `grid-template-columns` 属性，将空间划分为三个相等的列。
1. 你应该将 `.newspaper-layout` 的 `grid-template-rows` 属性设置为第一行的 `auto`，并将剩余空间平均分配给其他行。
1. 你应该在网格项之间添加间隙。
1. 通过将 `max-width` 设置为 `100%`，确保 `.cover-image` 内的图像适合指定空间。

**注意：** 确保在 HTML 中链接样式表以应用 CSS。

# --hints--

你应该有一个 `main` 元素。

```js
assert.exists(document.querySelector('main'));
```

你的 `main` 元素应具有 `newspaper-layout` 的 `class`。

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

你应该在你的 `.newspaper-layout` 元素内有一个 `class` 为 `title` 的 `header` 元素。

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

你的 `header` 元素应包含一个 `h1` 元素。

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

你的 `.title` 元素应包含一些文本。

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

你应该为每篇文章使用一个 `article` 元素，分别带有类 `feature-article`、`small-article1`、`small-article2`、`small-article3` 和 `secondary-article`。

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

`.feature-article` 应该包含一个 `h2` 元素，后跟一个 `p` 元素。

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`、`.small-article2` 和 `.small-article3` 应包含一个 `h3` 元素，后跟一个 `p` 元素。

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

`.secondary-article` 应包含一个 `h2` 元素，后跟一个 `p` 元素。

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

你应该在你的 `.newspaper-layout` 中有一个带有 `cover-image` 类的 `figure` 元素。

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

你的 `.cover-image` 应该包含一个 `img` 元素。

```js
assert.exists(document.querySelector(".cover-image img"));
```

`.newspaper-layout` 元素的 `display` 属性应设置为 `grid`。

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

你的 `.newspaper-layout` 元素应具有一个 `grid-template-areas` 属性，其中 `title` 跨越第一行的所有三列。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

网格模板的第二行应包含一个跨越两列的 `feature-article` 和一个 `cover-image`。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

网格模板的第三行应包含一个跨越两列的 `secondary-article` 和一个 `cover-image`。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

网格模板的第四行应包含 `small-article1`、`small-article2` 和 `small-article3`。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

你的 `.newspaper-layout` 选择器应使用 `grid-template-columns` 属性将空间划分为三个相等部分。记住你可以使用 `repeat` 函数或手动定义每个分数单位。

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

你的 `.newspaper-layout` 选择器应具有设置为 `auto 1fr 1fr 1fr` 的 `grid-template-rows` 属性。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

你应该在网格项之间添加间隙。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title` 应该有一个 `grid-area` 为 `title`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` 应该有一个 `grid-area` 为 `feature-article`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` 应该有一个 `grid-area` 为 `cover-image`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article` 应该有一个 `grid-area` 为 `secondary-article`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1` 应该有一个 `grid-area` 为 `small-article1`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2` 应该有一个 `grid-area` 为 `small-article2`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3` 应该有一个 `grid-area` 为 `small-article3`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

`.cover-image` 内的 `img` 应该有 `max-width` 为 `100%`。

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
