---
id: 68ed661e9f463c155bb677a1
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Nesta oficina você praticará trabalhar com CSS flexbox projetando um conjunto de caixas coloridas.

Para começar, dentro do elemento `body`, adicione um elemento `header`. Aninhe um elemento `h1` dentro do elemento `header` com o texto `Colored Boxes Layout`.

# --hints--

Você deve ter um `header` dentro do `body`.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

O elemento `header` deve ter um elemento `h1` dentro de si.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

Seu elemento `h1` deve conter o texto `Colored Boxes Layout`.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
