---
id: 68ef0b3237fef991c92d9c19
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Agora adicione um elemento `div` com a classe `flex-container` abaixo do seu `header`. Depois crie seis elementos `div` com a classe `box` dentro do `div` que você acabou de criar.

Lembre-se de que um atributo class é frequentemente usado para apontar para um nome de classe em uma folha de estilos. Neste caso, todos os elementos `div` serão estilizados igualmente de acordo com a definição de estilo `.flex-container` e `.box`.

# --hints--

Você deve ter um `div` elemento abaixo do `header` elemento.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

Seu `div` elemento abaixo do elemento `header` deve ter a classe `flex-container`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

Você deve ter seis elementos `div` aninhados dentro do elemento `.flex-container`.

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

Todos os seis `div` elementos dentro do elemento `.flex-container` devem ter a classe `box`.

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
