---
id: 68f91df5776e693049fb8942
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Agora é hora de adicionar algumas regras de CSS ao `styles.css`. Comece criando um seletor para o elemento `h1`. 

Centralize o conteúdo do `h1` definindo sua propriedade `text-align` como `center`. Também crie uma propriedade `margin-bottom` com o valor `10px` para adicionar margem entre a parte inferior do elemento `h1` e qualquer elemento HTML que fique abaixo.

# --hints--

Você deve ter um seletor `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

O seletor `h1` deve definir a propriedade `text-align` para `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

Seu seletor `h1` deve definir a propriedade `margin-bottom` para `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```

