---
id: 692f22a2ad44ae887a24c8f9
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Dentro de cada elemento `.box`, adicione um elemento `h2` com o texto `Box` seguido de um espaço e um número sequencial começando em `1`.

# --hints--

Você deve ter um elemento `h2` com o texto `Box 1` dentro do primeiro `div` com a classe `box`.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

Você deve ter um elemento `h2` com o texto `Box 2` dentro do segundo `div` com a classe `box`.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

Você deve ter um elemento `h2` com o texto `Box 3` dentro do terceiro `div` com a classe `box`.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

Você deve ter um elemento `h2` com o texto `Box 4` dentro do quarto `div` com a classe `box`.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

Você deve ter um elemento `h2` com o texto `Box 5` dentro do quinto `div` com a classe `box`.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

Você deve ter um elemento `h2` com o texto `Box 6` dentro do último `div` com a classe `box`.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
