---
id: 692f314d449329ba8bcd5d2e
title: Passo 17
challengeType: 0
dashedName: step-17
---

# --description--

A propriedade `border-radius` arredonda os cantos da borda de um elemento. Você pode especificar:

- Um valor para aplicar a todos os quatro cantos.
- Dois valores, com o primeiro valor para topo-esquerda/baixo-direita e o segundo para topo-direita/baixo-esquerda.
- Três valores, correspondentes a top-left, top-right/bottom-left e bottom-right.
- Quatro valores, que definem o raio da borda no sentido horário começando pelo canto superior esquerdo.

Agora adicione a propriedade `border-radius` com o valor `5px` ao seletor `.box`.

# --hints--

Seu seletor `.box` deve ter a propriedade `border-radius` com o valor `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border-radius'), '5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
