---
id: 692f22933ba1cb86cd75a70c
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Agora é hora de criar o nome da cor para as caixas. Para fazer isso, adicione um elemento `p` abaixo de cada elemento `h2` com esses nomes de cores na ordem: `Red`, `Orange`, `Yellow`, `Green`, `Blue` e `Indigo`. Uma cor para cada caixa.

# --hints--

Você deve adicionar um elemento `p` com o texto `Red` abaixo do elemento `h2` dentro do primeiro elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

Você deve adicionar um elemento `p` com o texto `Orange` abaixo do elemento `h2` dentro do segundo elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

Você deve adicionar um elemento `p` com o texto `Yellow` abaixo do elemento `h2` dentro do terceiro elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

Você deve adicionar um elemento `p` com o texto `Green` abaixo do elemento `h2` dentro do quarto elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

Você deve adicionar um elemento `p` com o texto `Blue` abaixo do elemento `h2` dentro do quinto elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

Você deve adicionar um elemento `p` com o texto `Indigo` abaixo do elemento `h2` dentro do sexto elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
