---
id: 692f2c20cb7f9e5050b935a7
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Adicione a propriedade `border` com o valor `1px solid #000` ao seletor `.box`. A propriedade `border` é usada para criar um contorno visível ao redor de um elemento. Ela pode ser personalizada em termos de `border-width`, `border-style` e `border-color`. 

Neste caso, `1px` é `border-width`, que define a espessura da borda. `solid` é `border-style` e define a aparência da borda (solid, dashed, dotted e assim por diante). E `#000` é `border-color`, que define a cor da borda.

# --hints--

Seu seletor `.box` deve ter a propriedade `border` com o valor `1px solid #000`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border'), '1px solid rgb(0, 0, 0)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
