---
id: 5a24c314108439a4d4036147
title: Підключення Redux до React
challengeType: 6
forumTopicId: 301426
dashedName: connect-redux-to-react
---

# --description--

Тепер, коли ви написали обидві функції `mapStateToProps()` і `mapDispatchToProps()`, ви можете використати їх, щоб спроєктувати `state` і `dispatch` до `props` одного з ваших React-компонентів. Метод `connect` з React Redux може виконати це завдання. Цей метод приймає два необов’язкові аргументи, `mapStateToProps()` і `mapDispatchToProps()`. Вони необов’язкові, бо у вас може бути компонент, який потребує доступу лише до `state`, але не потребує виклику жодних дій, або навпаки.

Щоб використати цей метод, передайте функції як аргументи, а потім одразу викличте результат із вашим компонентом. Цей синтаксис трохи незвичний і виглядає так:

```js
connect(mapStateToProps, mapDispatchToProps)(MyComponent)
```

**Примітка:** Якщо ви хочете пропустити один із аргументів методу `connect`, передайте `null` замість нього.

# --instructions--

У редакторі коду є функції `mapStateToProps()` і `mapDispatchToProps()`, а також новий React-компонент `Presentational`. Підключіть цей компонент до Redux за допомогою методу `connect` з глобального об’єкта `ReactRedux` і одразу викличте його для компонента `Presentational`. Призначте результат новій константі `const` з ім’ям `ConnectedComponent`, яка представлятиме підключений компонент. Ось і все, тепер ви підключені до Redux! Спробуйте змінити будь-який із аргументів `connect` на `null` і спостерігайте за результатами тестів.

# --hints--

Компонент `Presentational` має відтворитися.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('Presentational').length === 1;
  })()
);
```

Компонент `Presentational` має отримати проп `messages` через `connect`.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return props.messages === '__INITIAL__STATE__';
  })()
);
```

Компонент `Presentational` має отримати проп `submitNewMessage` через `connect`.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return typeof props.submitNewMessage === 'function';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line

const ConnectedComponent = connect(mapStateToProps, mapDispatchToProps)(Presentational);
```
