---
id: 5a24c314108439a4d4036144
title: Використання Provider для підключення Redux до React
challengeType: 6
forumTopicId: 301435
dashedName: use-provider-to-connect-redux-to-react
---

# --description--

У попередньому завданні ви створили Redux store для обробки масиву повідомлень і створили дію для додавання нових повідомлень. Наступний крок — надати React доступ до Redux store і дій, які він має виконувати для оновлень. React Redux надає свій `react-redux` пакет, щоб допомогти виконати ці завдання.

React Redux пропонує невеликий API з двома ключовими функціями: `Provider` і `connect`. Інше завдання охоплює `connect`. `Provider` — це обгортковий компонент від React Redux, який обгортає ваш React застосунок. Ця обгортка дозволяє отримати доступ до Redux `store` і `dispatch` функцій по всьому дереву компонентів. `Provider` приймає два пропси: Redux store і дочірні елементи вашого застосунку. Визначення `Provider` для компонента App може виглядати так:

```jsx
<Provider store={store}>
  <App/>
</Provider>
```

# --instructions--

Редактор коду зараз показує весь ваш Redux і React код з попередніх кількох завдань. Він включає Redux store, дії та компонент `DisplayMessages`. Єдина нова частина — це компонент `AppWrapper` внизу. Використайте цей компонент верхнього рівня, щоб відтворити `Provider` з `ReactRedux`, і передайте Redux store як проп. Потім відтворіть компонент `DisplayMessages` як дочірній елемент. Коли закінчите, ви повинні побачити, що ваш React компонент відтворився на сторінці.

**Примітка:** React Redux доступний тут як глобальна змінна, тому ви можете отримати доступ до Provider через точкову нотацію. Код у редакторі використовує це і встановлює його в константу `Provider`, щоб ви могли використати її в методі відтворення `AppWrapper`.

# --hints--

Компонент `AppWrapper` має відтворюватися.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('AppWrapper').length === 1;
  })()
);
```

Обгортковий компонент `Provider` має отримати проп `store`, рівний Redux store.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
      return __helpers
        .removeWhiteSpace(code)
        .includes('<Providerstore={store}>');
    })()
  );
```

`DisplayMessages` має відтворюватися як дочірній елемент `AppWrapper`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('AppWrapper').find('DisplayMessages').length === 1
    );
  })()
);
```

Компонент `DisplayMessages` має відтворювати елементи `h2`, `input`, `button` і `ul`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h2').length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('ul').length === 1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};



const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
    this.handleChange = this.handleChange.bind(this);
    this.submitMessage = this.submitMessage.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {  
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Render the Provider below this line

  // Change code above this line
};
```

# --solutions--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
 this.handleChange = this.handleChange.bind(this);
 this.submitMessage = this.submitMessage.bind(this);
 }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };  
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Change code below this line
  render() {
    return (
      <Provider store = {store}>
        <DisplayMessages/>
      </Provider>
    );
  }
  // Change code above this line
};
```
