---
id: 5a24c314108439a4d4036143
title: Винесення логіки стану в Redux
challengeType: 6
forumTopicId: 301429
dashedName: extract-state-logic-to-redux
---

# --description--

Тепер, коли ви завершили компонент React, потрібно перенести логіку, яку він виконує локально у своєму `state`, у Redux. Це перший крок для підключення простого React-застосунку до Redux. Єдина функціональність вашого застосунку — додавати нові повідомлення від користувача до невпорядкованого списку. Приклад простий, щоб показати, як React і Redux працюють разом.

# --instructions--

Спочатку визначте тип дії `ADD` і встановіть його у константу `ADD`. Далі визначте творця дії `addMessage()`, який створює дію для додавання повідомлення. Вам потрібно передати `message` цьому творцю дії і включити повідомлення у повернений `action`.

Потім створіть редуктор з назвою `messageReducer()`, який обробляє стан повідомлень. Початковий стан має бути порожнім масивом. Цей редуктор повинен додавати повідомлення до масиву повідомлень у стані або повертати поточний стан. Нарешті, створіть свій Redux store і передайте йому цей редуктор.

# --hints--

Константа `ADD` має існувати і містити значення рядка `ADD`

```js
assert(ADD === 'ADD');
```

Творець дії `addMessage` має повертати об’єкт з `type`, рівним `ADD`, і `message`, рівним повідомленню, яке передається.

```js
assert(
  (function () {
    const addAction = addMessage('__TEST__MESSAGE__');
    return addAction.type === ADD && addAction.message === '__TEST__MESSAGE__';
  })()
);
```

`messageReducer` має бути функцією.

```js
assert(typeof messageReducer === 'function');
```

Store має існувати і мати початковий стан, встановлений у порожній масив.

```js
assert(
  (function () {
    const initialState = store.getState();
    return typeof store === 'object' && initialState.length === 0;
  })()
);
```

Виклик `addMessage` у store має імутабельно додавати нове повідомлення до масиву повідомлень у стані.

```js
assert(
  (function () {
    const initialState = store.getState();
    const isFrozen = DeepFreeze(initialState);
    store.dispatch(addMessage('__A__TEST__MESSAGE'));
    const addState = store.getState();
    return isFrozen && addState[0] === '__A__TEST__MESSAGE';
  })()
);
```

`messageReducer` має повертати поточний стан, якщо викликати його з будь-якими іншими діями.

```js
assert(
  (function () {
    const addState = store.getState();
    store.dispatch({ type: 'FAKE_ACTION' });
    const testState = store.getState();
    return addState === testState;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Define ADD, addMessage(), messageReducer(), and store here:
```

# --solutions--

```jsx
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);
```
