---
id: 5a24c314108439a4d403614a
title: Рухаємося далі звідси
challengeType: 6
forumTopicId: 301434
dashedName: moving-forward-from-here
---

# --description--

Вітаємо! Ви завершили уроки з React і Redux. Є ще один момент, на який варто звернути увагу перед тим, як рухатися далі. Зазвичай ви не писатимете React-застосунки у редакторі коду, як цей. Це завдання дає вам уявлення про те, як виглядає синтаксис, якщо ви працюєте з файловою системою на власному комп’ютері. Код має виглядати схоже, за винятком використання інструкцій `import` (вони підключають усі залежності, які були надані вам у завданнях).

Нарешті, написання коду на React і Redux зазвичай потребує певної конфігурації. Це може швидко ускладнитися. Якщо ви хочете експериментувати на власному комп’ютері, <a href="https://www.freecodecamp.org/news/install-react-with-create-react-app/" target="_blank" rel="noopener noreferrer nofollow">Create React App</a> вже налаштований і готовий до роботи.

Або ж ви можете увімкнути Babel як JavaScript-препроцесор у CodePen, додати React і ReactDOM як зовнішні ресурси JavaScript і працювати там.

# --instructions--

Виведіть повідомлення `'Now I know React and Redux!'` у консоль.

# --hints--

Повідомлення `Now I know React and Redux!` має бути виведене у консоль.

```js
assert(
    /console\s*\.\s*log\s*\(\s*('|"|`)Now I know React and Redux!\1\s*\)/.test(
      code)
  );
```

# --seed--

## --seed-contents--

```jsx
/*
import React from 'react'
import ReactDOM from 'react-dom'
import { Provider, connect } from 'react-redux'
import { createStore, combineReducers, applyMiddleware } from 'redux'
import thunk from 'redux-thunk'

import rootReducer from './redux/reducers'
import App from './components/App'

const store = createStore(
  rootReducer,
  applyMiddleware(thunk)
);

ReactDOM.render(
  <Provider store={store}>
    <App/>
  </Provider>,
  document.getElementById('root')
);
*/

// Only change code below this line
```

# --solutions--

```jsx
console.log('Now I know React and Redux!');
```
