---
id: 68e45a2b4ab9b2f48b1cc6df
title: 设计一个定价方案版型页面
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事**

1. 你的页面应有一个 `h1` 元素，文本为 `Pricing Plans`。
2. 你的页面应在 `h1` 元素下方有一个带有类 `pricing-container` 的 `div` 元素。
    - `.pricing-container` 选择器应具有值为 `flex` 的 `display` 属性和值为 `wrap` 的 `flex-wrap` 属性。
3. 在 `.pricing-container` 元素内，你应该有三个带有 `pricing-card` 类的 `div` 元素来表现定价方案。
    - 其中一个 `.pricing-card` 元素除了 `pricing-card` 类之外，还应具有 `basic-plan` 类。
    - 其中一个 `.pricing-card` 元素应除了 `pricing-card` 类之外，还应具有 `pro-plan` 类。
    - 其中一个 `.pricing-card` 元素除了 `pricing-card` 类之外，还应具有 `premium-plan` 类。
4. 你的 `.basic-plan` 元素应包含一个带有文本 `Basic` 的 `h2` 元素。
5. 你的 `.basic-plan` 元素应包含一个 `p` 元素，文本为 `$9/month`。
6. 你的 `.pro-plan` 元素应包含一个带有文本 `Pro` 的 `h2` 元素。
7. 你的 `.pro-plan` 元素应包含一个带有文本 `$19/month` 的 `p` 元素。
8. 你的 `.premium-plan` 元素应包含一个带有文本 `Premium` 的 `h2` 元素。
9. 你的 `.premium-plan` 元素应包含一个带有文本 `$29/month` 的 `p` 元素。
10. 你的每个 `.pricing-card` 元素应当：
    - 使用 Flexbox 将其内容堆叠成一列，并使用 `space-between` 来调整子元素之间的间距。
    - 将 `flex` 属性设置为 `0 0 200px`，以赋予其一致的宽度并防止其在版型中增长或缩小。
    - 将 `border` 属性设置为 `2px solid black`，以查看不同卡片占用的空间。
11. `.basic-plan` 元素应首先出现在版型中。你应使用 `order` 属性来实现这一点。
12. `.pro-plan` 元素应在版型中显示为第二个。你应使用 `order` 属性，并将其 `flex-grow` 属性设置为 `2`，以便它占用比其他方案更多的空间。
13. `.premium-plan` 元素应在版型中最后出现。你应该为此使用 `order` 属性。

**注意：** 确保在 HTML 中链接样式表以应用 CSS。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的页面应有一个 `h1` 元素，文本为 `Pricing Plans`。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

你应该在 `h1` 元素下方有一个 `div` 元素。

```js
assert.exists(document.querySelector('h1 + div'));
```

你的 `div` 元素应具有类 `pricing-container`。

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

你应该有一个 `.pricing-container` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

你的 `pricing-container` 选择器应具有属性 `display`，其值为 `flex`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

你的 `.pricing-container` 选择器应具有属性 `flex-wrap`，其值为 `wrap`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

你应该在 `.pricing-container` 元素内有三个带有 `pricing-card` 类的 `div` 元素。

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

你的一个 `.pricing-card` 元素应该有 `basic-plan` 类。

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

你的一个 `.pricing-card` 元素应具有 `premium-plan` 类。

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

你的一个 `.pricing-card` 元素应该有 `pro-plan` 类。

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

你应该有一个 `.pricing-card` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

你的 `.pricing-card` 选择器应具有属性 `display`，其值为 `flex`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

你的 `.pricing-card` 选择器应具有属性 `flex-direction`，其值为 `column`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

你的 `.pricing-card` 选择器应具有属性 `justify-content`，其值为 `space-between`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

你的 `.pricing-card` 选择器应具有值为 `0 0 200px` 的 `flex` 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

你的 `.pricing-card` 选择器应具有值为 `2px solid black` 的 `border` 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

你的 `.basic-plan` 选择器应具有属性 `order`，其值为 `0`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

你的 `.basic-plan` 元素应包含一个带有文本 `Basic` 的 `h2` 元素。

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

你的 `.basic-plan` 元素应包含一个 `p` 元素，文本为 `$9/month`。

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

你的 `.pro-plan` 元素应包含一个带有文本 `Pro` 的 `h2` 元素。

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

你的 `.pro-plan` 元素应包含一个带有文本 `$19/month` 的 `p` 元素。

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

你的 `.pro-plan` 元素应具有属性 `order`，其值为 `1`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

你的 `.pro-plan` 元素应具有属性 `flex-grow`，其值为 `2`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

你的 `.premium-plan` 元素应包含一个带有文本 `Premium` 的 `h2` 元素。

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

你的 `.premium-plan` 元素应包含一个带有文本 `$29/month` 的 `p` 元素。

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

你的 `premium-plan` 元素应是版型中的最后一个元素，并具有属性 `order` 和值 `2`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```

