---
id: 587d7fad367417b2b2512be2
title: Cambiare il testo con eventi click
challengeType: 6
forumTopicId: 301500
dashedName: change-text-with-click-events
---

# --description--

Quando si verifica l'evento click, puoi usare JavaScript per aggiornare un elemento HTML.

Ad esempio, quando un utente clicca il pulsante `Get Message`, cambia il testo dell'elemento con la classe `message` per mostrare `Here is the message`.

Questo funziona aggiungendo il seguente codice all'interno dell'evento click:

```js
document.getElementsByClassName('message')[0].textContent =
  'Here is the message';
```

# --instructions--

Aggiungi codice all'interno del gestore dell'evento `onclick` per cambiare il testo all'interno dell'elemento `message` in modo che dica `Here is the message`.

# --hints--

Il tuo codice dovrebbe usare il metodo `document.getElementsByClassName` per selezionare l'elemento con la classe `message` e impostare il suo `textContent` sulla stringa data.

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.textContent\s*?=\s*?('|")Here is the message\2/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      document.getElementsByClassName('message')[0].textContent =
        'Here is the message';
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
