---
id: 587d7fae367417b2b2512be4
title: "Consultare i dati JSON da un'API"
challengeType: 6
forumTopicId: 301499
dashedName: access-the-json-data-from-an-api
---

# --description--

Nella sfida precedente, hai visto come ottenere dati JSON dall'API freeCodeCamp Cat Photo.

Ora esaminerai più da vicino i dati restituiti per comprendere meglio il formato JSON. Ricorda alcune notazioni in JavaScript:

<blockquote>[ ] -> Le parentesi quadre rappresentano un array.<br>{ } -> Le parentesi graffe rappresentano un oggetto.<br>" " -> Le virgolette doppie rappresentano una stringa. Sono anche usate per i nomi delle chiavi in JSON.</blockquote>

Capire la struttura dei dati che un'API restituisce è importante perché influenza come recuperare i valori di cui hai bisogno.

A destra, clicca il pulsante `Get Message` per caricare il JSON dell'API freeCodeCamp Cat Photo nell'HTML.

Il primo e l'ultimo carattere che vedi nei dati JSON sono parentesi quadre `[ ]`. Questo significa che i dati restituiti sono un array. Il secondo carattere nei dati JSON è una parentesi graffa `{`, che apre un oggetto. Guardando bene, puoi vedere che ci sono tre oggetti separati. I dati JSON sono un array di tre oggetti, dove ogni oggetto contiene informazioni su una foto di un gatto.

Hai imparato prima che gli oggetti contengono "coppie chiave-valore" separate da virgole. Nell'esempio della foto del gatto, il primo oggetto ha `"id":0` dove `id` è una chiave e `0` è il suo valore corrispondente. Allo stesso modo, ci sono chiavi per `imageLink`, `altText` e `codeNames`. Ogni oggetto della foto del gatto ha queste stesse chiavi, ma con valori diversi.

Un'altra interessante "coppia chiave-valore" nel primo oggetto è `"codeNames":["Juggernaut","Mrs. Wallace","ButterCup"]`. Qui `codeNames` è la chiave e il suo valore è un array di tre stringhe. È possibile avere array di oggetti così come una chiave con un array come valore.

Ricorda come consultare i dati in array e oggetti. Gli array usano la notazione a parentesi per consultare un indice specifico di un elemento. Gli oggetti usano la notazione a parentesi o a punto per consultare il valore di una data proprietà. Ecco un esempio che stampa la proprietà `altText` della prima foto del gatto — nota che i dati JSON analizzati nell'editor sono salvati in una variabile chiamata `json`:

```js
console.log(json[0].altText);
```

La console mostrerebbe la stringa `A white cat wearing a green helmet shaped melon on its head.`.

# --instructions--

Per il gatto con `id` pari a 2, stampa nella console il secondo valore nell'array `codeNames`. Dovresti usare la notazione a parentesi e a punto sull'oggetto (che è salvato nella variabile `json`) per consultare il valore.

# --hints--

Il tuo codice dovrebbe usare la notazione a parentesi e a punto per consultare il nome codice corretto e stampare `Loki` nella console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*json\s*\[2\]\s*(\.\s*codeNames|\[\s*('|`|")codeNames\2\s*\])\s*\[\s*1\s*\]\s*\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line

        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line
        console.log(json[2].codeNames[1]);
        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
