---
id: 587d7faf367417b2b2512be8
title: Ottenere dati di geolocalizzazione per trovare le coordinate GPS di un utente
challengeType: 6
forumTopicId: 18188
dashedName: get-geolocation-data-to-find-a-users-gps-coordinates
---

# --description--

Un'altra cosa interessante che puoi fare è consultare la posizione attuale del tuo utente. Ogni browser ha un navigator integrato che può fornire queste informazioni.

Il navigator ottiene la longitudine e la latitudine correnti dell'utente.

Vedrai un prompt per consentire o bloccare a questo sito di conoscere la tua posizione attuale. La sfida può essere completata in entrambi i modi, purché il codice sia corretto.

Se selezioni consenti, vedrai il testo sul telefono di output cambiare con la tua latitudine e longitudine.

Ecco un codice che fa questo:

```js
if (navigator.geolocation) {
  navigator.geolocation.getCurrentPosition(function (position) {
    document.getElementById('data').innerHTML =
      'latitude: ' +
      position.coords.latitude +
      '<br>longitude: ' +
      position.coords.longitude;
  });
}
```

Prima verifica se l'oggetto `navigator.geolocation` esiste. Se esiste, viene chiamato il metodo `getCurrentPosition` su quell'oggetto, che avvia una richiesta asincrona per la posizione dell'utente. Se la richiesta ha successo, viene eseguita la funzione di callback nel metodo. Questa funzione consulta i valori dell'oggetto `position` per latitudine e longitudine usando la notazione a punto e aggiorna l'HTML.

# --instructions--

Aggiungi il codice di esempio all'interno dei tag `script` per controllare la posizione attuale di un utente e inserirla nell'HTML.

# --hints--

Il tuo codice dovrebbe usare `navigator.geolocation` per consultare la posizione attuale dell'utente.

```js
assert.match(code, /navigator\.geolocation\.getCurrentPosition/g);
```

Il tuo codice dovrebbe usare `position.coords.latitude` per mostrare la posizione latitudinale dell'utente.

```js
assert.match(code, /position\.coords\.latitude/g);
```

Il tuo codice dovrebbe usare `position.coords.longitude` per mostrare la posizione longitudinale dell'utente.

```js
assert.match(code, /position\.coords\.longitude/g);
```

Dovresti mostrare la posizione dell'utente all'interno dell'elemento `div` con `id="data"`.

```js
assert.match(
  code,
  /document\.getElementById\(\s*('|")data\1\s*\)\.innerHTML/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  // Add your code below this line

  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data"></div>
```

# --solutions--

```html
<script>
  // Add your code below this line
  if (navigator.geolocation) {
    navigator.geolocation.getCurrentPosition(function(position) {
      document.getElementById('data').innerHTML = "latitude: " + position.coords.latitude + "<br>longitude: " + position.coords.longitude;
    });
  }
  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data">

</div>

</section>
```
