---
id: 587d7fae367417b2b2512be3
title: Ottenere JSON con il metodo XMLHttpRequest di JavaScript
challengeType: 6
forumTopicId: 301502
dashedName: get-json-with-the-javascript-xmlhttprequest-method
---

# --description--

Puoi anche richiedere dati da una fonte esterna. Qui entrano in gioco le API.

Ricorda che le API - o interfacce di programmazione delle applicazioni - sono strumenti che i computer usano per comunicare tra loro. Imparerai come aggiornare l'HTML con i dati che otteniamo dalle API usando una tecnologia chiamata AJAX.

La maggior parte delle API web trasferisce dati in un formato chiamato JSON. JSON sta per JavaScript Object Notation.

La sintassi JSON è molto simile alla notazione letterale degli oggetti JavaScript. JSON ha proprietà dell'oggetto e i loro valori attuali, racchiusi tra un `{` e un `}`.

Queste proprietà e i loro valori sono spesso chiamati "coppie chiave-valore".

Tuttavia, il JSON trasmesso dalle API viene inviato come `bytes`, e la tua applicazione lo riceve come un `string`. Questi possono essere convertiti in oggetti JavaScript, ma per impostazione predefinita non sono oggetti JavaScript. Il metodo `JSON.parse` analizza la stringa e costruisce l'oggetto JavaScript descritto da essa.

Puoi richiedere il JSON dall'API Cat Photo di freeCodeCamp. Ecco il codice che puoi inserire nel tuo evento click per farlo:

```js
const req = new XMLHttpRequest();
req.open('GET', '/json/cats.json', true);
req.send();
req.onload = function () {
  const json = JSON.parse(req.responseText);
  document.getElementsByClassName('message')[0].innerHTML =
    JSON.stringify(json);
};
```

Ecco una revisione di cosa fa ogni parte. L'oggetto JavaScript `XMLHttpRequest` ha diverse proprietà e metodi usati per trasferire dati. Prima, viene creato un'istanza dell'oggetto `XMLHttpRequest` e salvata nella variabile `req`. Poi, il metodo `open` inizializza una richiesta - questo esempio richiede dati da un'API, quindi è una richiesta `GET`. Il secondo argomento per `open` è l'URL dell'API da cui stai richiedendo dati. Il terzo argomento è un valore booleano dove `true` la rende una richiesta asincrona. Il metodo `send` invia la richiesta. Infine, il gestore dell'evento `onload` analizza i dati restituiti e applica il metodo `JSON.stringify` per convertire l'oggetto JavaScript in una stringa. Questa stringa viene poi inserita come testo del messaggio.

# --instructions--

Aggiorna il codice per creare e inviare una richiesta `GET` all'API Cat Photo di freeCodeCamp. Poi clicca il pulsante `Get Message`. La tua funzione AJAX sostituirà il testo `The message will go here` con l'output JSON grezzo proveniente dall'API.

# --hints--

Il tuo codice dovrebbe creare un nuovo `XMLHttpRequest`.

```js
assert.match(code, /new\s+?XMLHttpRequest\(\s*?\)/g);
```

Il tuo codice dovrebbe usare il metodo `open` per inizializzare una richiesta `GET` all'API Cat Photo di freeCodeCamp.

```js
assert.match(
  code,
  /\.open\(\s*('|")GET\1\s*,\s*('|")\/json\/cats\.json\2\s*,\s*true\s*\)/g
);
```

Il tuo codice dovrebbe usare il metodo `send` per inviare la richiesta.

```js
assert.match(code, /\.send\(\s*\)/g);
```

Il tuo codice dovrebbe avere un gestore dell'evento `onload` impostato su una funzione.

```js
assert.match(
  code,
  /\.onload\s*=\s*(function|\(\s*\))\s*(\(\s*\)|\=\>)\s*{/g
);
```

Il tuo codice dovrebbe usare il metodo `JSON.parse` per analizzare il `responseText`.

```js
assert.match(code, /JSON\s*\.parse\(\s*.*\.responseText\s*\)/g);
```

Il tuo codice dovrebbe ottenere l'elemento con classe `message` e cambiare il suo inner HTML con la stringa dei dati JSON.

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.innerHTML\s*?=\s*?JSON\.stringify\(.+?\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = () => {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
