---
id: 5ccfad82bb2dc6c965a848e5
title: Ottenere JSON con il metodo fetch di JavaScript
challengeType: 6
forumTopicId: 301501
dashedName: get-json-with-the-javascript-fetch-method
---

# --description--

Un altro modo per richiedere dati esterni è usare il metodo JavaScript `fetch()`. È equivalente a `XMLHttpRequest`, ma la sintassi è considerata più facile da capire.

Ecco il codice per fare una richiesta GET a `/json/cats.json`

```js
fetch('/json/cats.json')
  .then(response => response.json())
  .then(data => {
    document.getElementById('message').innerHTML = JSON.stringify(data);
  });
```

Nota: Il metodo `fetch()` usa `GET` come metodo `HTTP` predefinito. Questo significa che non devi specificarlo esplicitamente per il recupero base dei dati.

Dai un’occhiata a ogni parte di questo codice.

La prima riga è quella che fa la richiesta. Quindi, `fetch(URL)` fa una richiesta `GET` all’URL specificato. Il metodo restituisce una Promise.

Dopo che viene restituita una Promise, se la richiesta ha successo, viene eseguito il metodo `then`, che prende la risposta e la converte in formato JSON.

Il metodo `then` restituisce anche una Promise, che viene gestita dal successivo metodo `then`. L’argomento nel secondo `then` è l’oggetto JSON che stai cercando!

Ora, seleziona l’elemento che riceverà i dati usando `document.getElementById()`. Poi modifica il codice HTML dell’elemento inserendo una stringa creata dall’oggetto JSON restituito dalla richiesta.

# --instructions--

Aggiorna il codice per creare e inviare una richiesta `GET` all’API Cat Photo di freeCodeCamp. Ma questa volta, usando il metodo `fetch` invece di `XMLHttpRequest`.

# --hints--

Il tuo codice dovrebbe usare i dati recuperati per sostituire l’HTML interno

```js
const catData = 'dummy data';
const ref = fetch;
fetch = () => Promise.resolve({ json: () => catData });

  try {
    document.getElementById('getMessage').click();
    await new Promise((resolve, reject) => setTimeout(() => resolve(), 250));
  } catch (error) {
    console.log(error);
  } finally {
    fetch = ref;
    assert.equal(
      document.getElementById('message').textContent,
      JSON.stringify(catData)
    );
  }

```

Il tuo codice dovrebbe fare una richiesta `GET` con `fetch`.

```js
assert.match(code, /fetch\s*\(\s*('|")\/json\/cats\.json\1\s*\)/g);
```

Il tuo codice dovrebbe usare `then` per convertire la risposta in JSON.

```js
assert.match(
  code,
  /\.then\s*\(\s*\(?(?<var>\w+)\)?\s*=>\s*\k<var>\s*\.json\s*\(\s*\)\s*\)/g
);
```

Il tuo codice dovrebbe usare `then` per gestire i dati convertiti in JSON dall’altro `then`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.then\(\(?\w+\)?=>{[^}]*}\)/g);
```

Il tuo codice dovrebbe prendere l’elemento con id `message` e cambiare il suo HTML interno con la stringa dei dati JSON.

```js
assert.match(
  __helpers.removeWhiteSpace(code),
  /document\.getElementById\(('|")message\1\)\.innerHTML=JSON\.stringify\(?\w+\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      fetch('/json/cats.json')
        .then(response => response.json())
        .then(data => {
          document.getElementById('message').innerHTML = JSON.stringify(data);
        });
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
