---
id: 5688e62ea601b2482ff8422b
title: プロフィール検索を作成する
challengeType:  26
dashedName: lab-profile-lookup
---

# --description--

このレッスンでは、連絡先リストの人物情報を検索するプロフィール検索を作成します。

この例では、連絡先にAkira Laineという名前の人がいると想定し、`lookUpProfile("Akira", "lastName")`は`Laine`を返すべきです。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してレッスンを完了してください。

**ユーザーストーリー:**

1. 名前とプロパティを引数に取る`lookUpProfile`という名前の関数を作成してください。 
2. 提供された`contacts`配列から連絡先情報を取得してください。 
3. 関数が連絡先の名前と、その連絡先に存在するプロパティを受け取った場合、そのプロパティの値を返すべきです。
4. 関数に渡された名前が連絡先配列のどの連絡先とも一致しない場合、関数は`"No such contact"`を返すべきです。
5. 見つかった連絡先にプロパティが存在しない場合、関数は`"No such property"`を返すべきです。

# --before-each--

```js
const _contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

contacts = structuredClone(_contacts); 
```

# --hints--

`contacts`配列はグローバルスコープに存在しているはずです。もし削除してしまった場合は、レッスンをリセットして復元してください。

```js
assert.isArray(contacts);
```

`lookUpProfile`という名前の関数を作成してください。

```js
assert.isFunction(lookUpProfile);
```

`lookUpProfile("Kristian", "lastName")`は文字列を返すべきです。

```js
assert.isString(lookUpProfile("Kristian", "lastName"));
```

`lookUpProfile("Kristian", "lastName")`は文字列`Vos`を返すべきです。

```js
assert.strictEqual(lookUpProfile('Kristian', 'lastName'), 'Vos');
```

`lookUpProfile("Sherlock", "likes")`は`["Intriguing Cases", "Violin"]`を返すべきです。

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")`は配列を返すべきです。

```js
assert.isArray(lookUpProfile('Harry', 'likes'));
```

`lookUpProfile("Bob", "number")`は文字列`No such contact`を返すべきです。

```js
assert.strictEqual(lookUpProfile('Bob', 'number'),'No such contact');
```

`lookUpProfile("Bob", "potato")`は文字列`No such contact`を返すべきです。

```js
assert.strictEqual(lookUpProfile('Bob', 'potato'), 'No such contact');
```

`lookUpProfile("Akira", "address")`は文字列`No such property`を返すべきです。

```js
assert.strictEqual(lookUpProfile('Akira', 'address'), 'No such property');
```

`contacts`配列は`lookUpProfile`を実行した後も変更されていないはずです。

```js
lookUpProfile("Sherlock", "likes");
assert.deepEqual(contacts,_contacts);
```

# --seed--

## --seed-contents--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


```

# --solutions--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
