---
id: 6716249b5405164036fd0b0d
title: Einen Sortier-Visualizer erstellen
challengeType: 25
dashedName: build-a-sorting-visualizer
demoType: onClick
---

# --description--

Der Bubble-Sort-Algorithmus sortiert eine Folge ganzer Zahlen, indem er Paare benachbarter Elemente vom Anfang der Folge an vergleicht. Wenn das erste Element größer als das zweite ist, tauscht er sie. Dann fährt er mit dem nächsten Paar fort. Wenn das letzte Element der Folge erreicht ist, beginnt er einen neuen Zyklus vom Anfang der Folge und wiederholt den Vorgang, bis die Elemente sortiert sind. Der Algorithmus stoppt, wenn ein Zyklus ohne Tausch abgeschlossen wird.

Für dieses Labor wurden Ihnen bereits das gesamte HTML und CSS bereitgestellt. Sie verwenden JavaScript, um den Bubble-Sort-Visualizer zu vervollständigen, sodass er jeden Schritt visualisiert, den der Bubble-Sort-Algorithmus benötigt, um ein Array aus fünf ganzen Zahlen zu sortieren.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `generateElement` haben, die eine zufällige ganze Zahl zwischen `1` und `100` (einschließlich) zurückgibt.
1. Sie sollten eine Funktion namens `generateArray` haben, die die Funktion `generateElement` verwendet, um ein Array mit fünf zufälligen ganzen Zahlen zurückzugeben.
1. Sie sollten eine Funktion namens `generateContainer` haben, die ein leeres `div`-Element erstellt und zurückgibt.
1. Sie sollten eine Funktion namens `fillArrContainer` haben, die ein HTML-Element als erstes Argument und ein Array als zweites Argument entgegennimmt.
1. Die Funktion `fillArrContainer` sollte ein Element als ersten Parameter und ein Array ganzer Zahlen als zweiten Parameter nehmen und das Element mit fünf `span`-Elementen füllen, wobei jedes `span` eine der Zahlen aus dem Array anzeigt.
1. Sie sollten eine Funktion namens `isOrdered` haben, die zwei ganze Zahlen entgegennimmt und einen Boolean zurückgibt, der angibt, ob die erste Zahl kleiner oder gleich der zweiten ist.
1. Sie sollten eine Funktion namens `swapElements` haben, die ein Array ganzer Zahlen und einen numerischen Index entgegennimmt.
1. Die Funktion `swapElements` sollte das Array an Ort und Stelle ändern, indem sie das Element am übergebenen Index und das folgende Element tauscht, falls `isOrdered` `false` zurückgibt.
1. Sie sollten eine Funktion namens `highlightCurrentEls` haben, die ein HTML-Element und einen numerischen Index entgegennimmt.
1. Die Funktion `highlightCurrentEls` sollte den `border` des Kindelements des gegebenen Elements am angegebenen Index und des unmittelbar nachfolgenden Kindes auf einen `dashed`-Stil, eine `red`-Farbe und eine von Ihnen gewählte Breite setzen.
1. Wenn Sie auf `#generate-btn` klicken, sollten Sie die Funktion `fillArrContainer` verwenden, um `#starting-array` mit fünf `span`-Elementen zu füllen, die jeweils eine zufällige Zahl als Text haben. Falls vorhanden, sollten andere Elemente aus `#array-container` entfernt werden.
1. Sie sollten den Bubble-Sort-Algorithmus implementieren, sodass nach einem Klick auf `#sort-btn` `#array-container` für jeden Schritt, den der Bubble-Sort-Algorithmus benötigt, um das Startarray zu sortieren, ein `div`-Element enthält, einschließlich des `div`, das das Startarray darstellt, und eines `div`, das das sortierte Array darstellt. Die bisher erstellten Funktionen können hier nützlich sein.
1. Jedes `div` sollte fünf `span`-Elemente enthalten, die das Array im aktuellen Sortierzustand darstellen.
1. Nach einem Klick auf `#sort-btn` sollte `#starting-array` den Startschritt mit dem Anfangsarray und den ersten zwei hervorgehobenen Zahlen darstellen.
1. Für jeden Sortierschritt sollten Sie `highlightCurrentEls` verwenden, um die zwei Zahlen hervorzuheben, die verglichen werden, und sie im nächsten Schritt mit `swapElements` tauschen.

# --hints--

Sie sollten eine Funktion namens `generateElement` haben.

```js
assert.isFunction(generateElement);
```

Ihre Funktion `generateElement` sollte eine zufällige ganze Zahl zwischen `1` und `100` (einschließlich) zurückgeben.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    assert.strictEqual(generateElement(), 26);
    assert.strictEqual(generateElement(), 9);
} finally {
    randomMocker.restore();
}
```

Sie sollten eine Funktion namens `generateArray` haben.

```js
assert.isFunction(generateArray)
```

Ihre Funktion `generateArray` sollte die Funktion `generateElement` verwenden.

```js
assert.match(generateArray.toString(), /\bgenerateElement\b/);
```

Ihre Funktion `generateArray` sollte ein Array mit fünf zufälligen ganzen Zahlen zwischen `1` und `100` zurückgeben.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    assert.deepEqual(generateArray(), [26, 9, 58, 23, 38])
} finally {
    randomMocker.restore();
}
```

Sie sollten eine Funktion namens `generateContainer` haben.

```js
assert.isFunction(generateContainer);
```

Ihre Funktion `generateContainer` sollte ein leeres `div`-Element zurückgeben.

```js
const div = generateContainer();
assert.equal(div.tagName, "DIV");
assert.isEmpty(div.children);
```

Sie sollten eine Funktion namens `fillArrContainer` haben.

```js
assert.isFunction(fillArrContainer);
```

Ihre Funktion `fillArrContainer()` sollte ein Element als ersten Parameter und ein Array ganzer Zahlen als zweiten Parameter nehmen und das Element mit fünf `span`-Elementen füllen, wobei jedes `span` eine der Zahlen aus dem Array anzeigt.

```js
const testDiv = document.createElement("div");
const testArr = [15, 98, 17, 5, 63]
fillArrContainer(testDiv, testArr);
const children = testDiv.children;
assert.lengthOf(children, 5);
Array.from(children).forEach((el, i) => {
    assert.equal(el.tagName, "SPAN");
    assert.equal(el.innerText.trim(), testArr[i])
})
```

Sie sollten eine Funktion namens `isOrdered` haben.

```js
assert.isFunction(isOrdered);
```

Ihre Funktion `isOrdered` sollte zwei ganze Zahlen entgegennehmen und einen Boolean zurückgeben, der angibt, ob die erste Zahl kleiner oder gleich der zweiten ist.

```js
assert.lengthOf(isOrdered, 2);
assert.isTrue(isOrdered(2, 60));
assert.isFalse(isOrdered(10, 3));
assert.isTrue(isOrdered(5, 5));
```

Sie sollten eine Funktion namens `swapElements` haben.

```js
assert.isFunction(swapElements);
```

Ihre Funktion `swapElements` sollte ein Array ganzer Zahlen und einen numerischen Index als Argumente nehmen. Sie sollte das übergebene Array an Ort und Stelle ändern, indem sie das Element am angegebenen Index und das folgende Element tauscht, falls das erste Element größer als das zweite ist.

```js
const testArr = [22, 4, 87, 47, 33];
swapElements(testArr, 0);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 1);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 2);
assert.deepEqual(testArr, [4, 22, 47, 87, 33]);
swapElements(testArr, 3);
assert.deepEqual(testArr, [4, 22, 47, 33, 87]);
```

Sie sollten eine Funktion namens `highlightCurrentEls` haben.

```js
assert.isFunction(highlightCurrentEls);
```

Ihre Funktion `highlightCurrentEls` sollte den Nachfahren des angegebenen Elements, die sich am angegebenen Index und am nächsten Index befinden, einen Rahmen geben, der `dashed`, `red` ist und eine von Ihnen gewählte Breite hat.

```js
const testDiv = document.createElement("div");
document.querySelector("body").appendChild(testDiv)
for (let i = 0; i < 5; i++) {
    testDiv.appendChild(document.createElement("span"));
}
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|#F00|hsl\(0,\s*100%,\s*50%\))/;
const revertBorder = () => {
    for (const el of children) {
        el.style.border = "revert";
    }
}
const children = testDiv.children;

for (let i = 0; i < 3; i++) {
    highlightCurrentEls(testDiv, i);
    for (let j = 0; j < 5; j++) {
        let b = getComputedStyle(children[j]).border;
        if (j == i || j == i + 1) {            
            assert.match(b, redBorderRegex);
            assert.isAbove(parseFloat(b), 0);
        } else {
            assert.notMatch(b, redBorderRegex);
        }
    }
    revertBorder();
}
testDiv.remove();
```

Wenn Sie auf `#generate-btn` klicken, sollten Sie `#starting-array` mit fünf `span`-Elementen füllen, die jeweils eine zufällige Zahl zwischen `1` und `100` als Text haben.

```js
const genBtn = document.querySelector("#generate-btn");
genBtn.click();
const children = document.querySelector("#starting-array").querySelectorAll("span");
assert.lengthOf(children, 5);
Array.from(children).forEach(el => {
    assert.equal(el.tagName, "SPAN");
    const num = Number(el.innerText.trim());
    assert.isAtMost(num, 100);
    assert.isAtLeast(num, 1);
})
```

Wenn `#starting-array` bereits ein generiertes Array enthält oder `#array-container` das sortierte Array enthält, sollte ein Klick auf `#generate-btn` andere Elemente in `#array-container` entfernen und nur `#starting-array` mit neu generierten Zahlen übriglassen.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
genBtn.dispatchEvent(new Event("click"));

const prevNumbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));

sortBtn.dispatchEvent(new Event("click"));
genBtn.dispatchEvent(new Event("click"));

const container = document.querySelector("#array-container");
assert.lengthOf(container.children, 1);
const numbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));
assert.lengthOf(numbers, 5);
assert.isTrue(prevNumbers.some((num, index) => num !== numbers[index]))
```

Nach einem Klick auf `#sort-btn` sollte `#array-container` so viele `div`-Elemente enthalten, wie Schritte erforderlich sind, damit der Bubble-Sort-Algorithmus das Startarray sortiert, einschließlich `div`, das das Startarray darstellt, und `div`, das das sortierte Array darstellt.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
// using randomMocker to be sure that the starting array requires 13 steps to be sorted
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.lengthOf(container.children, 13)
    Array.from(container.children).forEach(el => {assert.equal(el.tagName, "DIV")})
} finally {
    randomMocker.restore();
}
```

Nach einem Klick auf `#sort-btn` sollte jedes `div` innerhalb von `#array-container` fünf `span` enthalten, die jeweils eine Zahl als Text haben und so angeordnet sind, dass sie die Schritte darstellen, die der Bubble-Sort-Algorithmus benötigt, um das Startarray zu sortieren.

```js
const finalArr = [9, 23, 26, 38, 58];
const arrays = [
    [26, 9, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 23, 58, 38],
    [9, 26, 23, 38, 58],
    [9, 26, 23, 38, 58],
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr
]
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.isNotEmpty(container.children);
    Array.from(container.children).forEach((el, i) => {
        Array.from(el.children).forEach((j, k) => {
            assert.strictEqual(Number(j.innerText.trim()), arrays[i][k])
        })
    })
} finally {
    randomMocker.restore();
}
```

Wenn Sie auf `#sort-btn` klicken, sollten Sie die Funktion `highlightCurrentEls` verwenden, um die Elemente hervorzuheben, die in jedem Schritt verglichen werden.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
assert.match(code, /addEventListener\(\s*["']click["']\s*,[\s\S]*bubbleSort\s*\(/);
assert.match(bubbleSort.toString(), /highlightCurrentEls\s*\(/);
genBtn.dispatchEvent(new Event("click"));
sortBtn.dispatchEvent(new Event("click"));
```

Nach einem Klick auf `#sort-btn` sollte `#starting-array` den Startschritt mit dem Anfangsarray und den ersten zwei hervorgehobenen Zahlen mithilfe von `highlightCurrentEls` darstellen.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");   
genBtn.dispatchEvent(new Event("click"));
sortBtn.dispatchEvent(new Event("click"));
const firstContainer = document.querySelector("#starting-array");
const children = firstContainer.children
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|hsl\(0,\s*100%,\s*50%\))/;
assert.match(getComputedStyle(children[0]).border, redBorderRegex);
assert.match(getComputedStyle(children[1]).border, redBorderRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

#sort-btn {
   display: none 
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js
const arrayContainer = document.getElementById("array-container");
const generateArrayBtn = document.getElementById("generate-btn");
const sortArrayBtn = document.getElementById("sort-btn");
const arrayLength = 5;
const minVal = 1;
const maxVal = 100;
const isStart = () => arrayContainer.children.length === 1;
const clearArrayContainer = () => { arrayContainer.innerHTML = '<div id="starting-array"></div>' };
const showSortBtn = () => {sortArrayBtn.style.display = "inline-block"};
const hideSortBtn = () => {sortArrayBtn.style.display = "none"};
const generateElement = () => Math.floor(Math.random() * maxVal + minVal);
const generateArray = () => Array.from({ length: arrayLength }, generateElement);
const fillArrContainer = (container, arr) => {
    container.innerHTML = "";
    arr.forEach(i => {
        const el = document.createElement("span");
        el.innerText = i;
        el.id = `number-${i}`
        container.appendChild(el);
    })
}
const isOrdered = (el1, el2) => el1 <= el2;

const swapElements = (arr, n = 0) => {
    if (n < arr.length - 1 && !isOrdered(arr[n], arr[n + 1])) {
        const temp = arr[n];
        arr[n] = arr[n + 1];
        arr[n + 1] = temp;
    }
}
const getLastChildren = () => arrayContainer.lastElementChild;
const getLastArr = () => {
    const els = Array.from(getLastChildren().children);
    const arr = els.map(el => Number(el.id.replace("number-", "")));
    return arr;
}
const generateContainer = () => {
    const container = document.createElement("div");
    arrayContainer.appendChild(container);
    return container;
}
const highlightCurrentEls = (container, n = 0) => {
    const children = container.children;
    children[n].style.border = "2px dashed red";
    children[n + 1].style.border = "2px dashed red";
}

const highlightSorted = () => {
    getLastChildren().style.border = "4px solid green";
}

const bubbleSort = () => {
    let swapped = true;
    while (swapped) {
        const startingArr = getLastArr();
        startingArr.forEach((_, i) => {
            if (i + 1 < arrayLength) {
                highlightCurrentEls(getLastChildren(), i);
                const arr = getLastArr();
                swapElements(arr, i);
                fillArrContainer(generateContainer(), arr);
            }
        })
        const lastArr = getLastArr()
        if (startingArr.every((el, i) => el === lastArr[i])) swapped = false;
    }
}


generateArrayBtn.addEventListener("click", () => {
    if (!isStart()) {
        clearArrayContainer();
    }
    fillArrContainer(document.getElementById("starting-array"), generateArray());
    showSortBtn();
});

sortArrayBtn.addEventListener("click", () => {
    bubbleSort();
    highlightSorted();
    hideSortBtn();
})
```
