---
id: 66c6491fe4c8e0f16845425f
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Ora è il momento di aggiungere un altro frutto alla lista.

Usando lo stesso metodo dell'array di prima, aggiungi la stringa `"Grapes"` alla fine dell'array `shoppingList`.

Poi, aggiungi un `console.log` e chiama la funzione `getShoppingListMsg` con l'array `shoppingList` come argomento all'interno di `console.log` per vedere la lista aggiornata stampata nella console.

# --hints--

Il tuo array `shoppingList` dovrebbe ora contenere due elementi.

```js
assert.lengthOf(shoppingList, 2);
```

Il tuo array `shoppingList` dovrebbe contenere la stringa `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

Dovresti chiamare la funzione `getShoppingListMsg(shoppingList)` all'interno di `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
