---
id: 66c644c08ebcd2ef90c750bd
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Nelle lezioni precedenti, hai imparato come aggiungere elementi alla fine di un array usando il metodo `push` in questo modo:

```js
fruits.push('pear');
```

Usando il metodo `push`, aggiungi la stringa `"Apples"` all'array `shoppingList`.

# --hints--

Il tuo array `shoppingList` dovrebbe contenere un elemento.

```js
assert.lengthOf(shoppingList, 1);
```

Il tuo array `shoppingList` dovrebbe contenere la stringa `"Apples"`. Controlla bene che non ci siano errori di ortografia o maiuscole.

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
