---
id: 66c645b345a39bf04864dc50
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Per vedere la lista della spesa aggiornata, devi stampare nella console la lista attuale e un breve messaggio.

Poiché questo messaggio verrà usato più volte durante il workshop, è meglio creare una funzione riutilizzabile.

Crea una funzione chiamata `getShoppingListMsg` che prende un array come parametro e restituisce la stringa `"Current Shopping List: "` seguita dal contenuto dell'array fornito.

Puoi usare template literal o la concatenazione di stringhe con l'operatore `+` per unire la stringa con l'array passato alla funzione.

# --hints--

Dovresti avere una funzione chiamata `getShoppingListMsg` nel tuo codice.

```js
assert.isFunction(getShoppingListMsg);
```

La tua funzione `getShoppingListMsg` dovrebbe restituire una stringa.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

La tua funzione `getShoppingListMsg` dovrebbe accettare un array e restituire il messaggio `"Current Shopping List: "` seguito dal contenuto di quell'array. Controlla bene spazi e errori di ortografia.

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
