---
id: 66c63bfa4685e4e3843fa0dc
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop continuerai a imparare sugli array creando una lista della spesa.

Inizia aggiungendo un `console.log` che scriva la stringa `"Grocery shopping list"` nella console.

# --hints--

Dovresti avere un `console.log` nel tuo codice.

```js
assert.match(code, /console\.log(.*)/);
```

Il tuo `console.log` dovrebbe mostrare il messaggio `"Grocery shopping list"`. Controlla bene che non ci siano errori di ortografia o maiuscole.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
