---
id: 66c64095f61166ede6e39a84
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Per questa lista della spesa, userai un array per rappresentare gli elementi che devi comprare.

Nelle lezioni precedenti, hai imparato a creare array così:

```js
const fruits = ['apple', 'banana', 'orange'];
```

In questo passaggio, crea una variabile chiamata `shoppingList` e assegnale un array vuoto.

# --hints--

Dovresti avere una variabile chiamata `shoppingList`.

```js
assert.isNotNull(shoppingList);
```

Dovresti assegnare un array alla variabile `shoppingList`.

```js
assert.isArray(shoppingList);
```

Il tuo array dovrebbe essere vuoto.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
