---
id: 66c64210413532ee9d3bd342
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Per questa parte successiva del workshop, esercitati ad aggiungere elementi alimentari alla lista della spesa.

Inizia usando `console.log` per registrare il messaggio `"It will be nice to have some fruit to eat."`. 

# --hints--

Dovresti registrare il messaggio `"It will be nice to have some fruit to eat."` nella console.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
