---
id: 66c72f4d0528bd268a82107b
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Ora è il momento di registrare l'array `shoppingList` aggiornato nella console.

Aggiungi un altro `console.log` e chiama la funzione `getShoppingListMsg` con l'array `shoppingList` come argomento all'interno di `console.log` per vedere la lista aggiornata registrata nella console.

# --hints--

Dovresti chiamare `getShoppingListMsg(shoppingList)` all'interno di `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 4);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");

--fcc-editable-region--

--fcc-editable-region--
```
