---
id: 66c72a55418cc9247b710827
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Poi, aggiungi un `console.log` e chiama la funzione `getShoppingListMsg` con l’array `shoppingList` come argomento all’interno di `console.log` per vedere la lista aggiornata stampata nella console.

# --hints--

Dovresti chiamare `getShoppingListMsg(shoppingList)` all’interno di `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
