---
id: 66c72b0ffbc5522525768558
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Nei passaggi precedenti, hai visto come aggiungere un elemento alla fine dell'array usando il metodo `push`.

Ma il metodo `push` accetta argomenti multipli, quindi puoi aggiungere più elementi alla fine dell'array in questo modo:

```js
array.push(item1, item2, item3);
```

In questo passaggio, usa il metodo `push` per aggiungere le stringhe `"Popcorn", "Beef Jerky", "Potato Chips"` all'array `shoppingList`.

L'ordine è importante, quindi assicurati di aggiungere gli elementi nell'ordine in cui sono elencati.

# --hints--

Il tuo array `shoppingList` dovrebbe contenere un totale di `6` stringhe.

```js
assert.lengthOf(shoppingList, 6);
```

Il tuo `shoppingList` dovrebbe contenere la stringa `"Popcorn"`.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

Il tuo `shoppingList` dovrebbe contenere la stringa `"Beef Jerky"`.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

Il tuo `shoppingList` dovrebbe contenere la stringa `"Potato Chips"`.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
