---
id: 66c726c34ecf1e238aa9d7d5
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Nelle lezioni precedenti, hai imparato come aggiungere elementi all'inizio di un array usando il metodo `unshift()`. 

Ecco un promemoria su come usare il metodo `unshift()`:

```js
array.unshift(item1, item2, ..., itemX);
```

Usa il metodo `unshift()` per aggiungere la stringa `"Vegetable Oil"` all'inizio dell'array `shoppingList`.

# --hints--

Dovresti avere tre elementi nel tuo array `shoppingList`.

```js
assert.lengthOf(shoppingList, 3);
```

Dovresti usare il metodo `unshift()` per aggiungere `"Vegetable Oil"` all'inizio dell'array `shoppingList`.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
