---
id: 66c73a7798f6f62b2ae58f22
title: Passaggio 16
challengeType: 1
dashedName: step-16
---

# --description--

Ora è il momento di aggiungere altri elementi all'inizio della lista della spesa.

Inizia aggiungendo un'istruzione `console.log` che registra il messaggio `"It might be nice to get a dessert."`.

Sotto quell'istruzione `console`, usa il metodo array corretto per aggiungere la stringa `"Chocolate Cake"` all'inizio dell'array `shoppingList`.

Infine, aggiungi un'istruzione `console.log` e chiama la funzione `getShoppingListMsg` con l'array `shoppingList` come argomento all'interno di `console.log` per vedere la lista aggiornata stampata nella console.

# --hints--

Dovresti avere un'istruzione `console.log` che registra il messaggio `"It might be nice to get a dessert."`.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

Il tuo array `shoppingList` dovrebbe contenere un totale di `6` elementi.

```js
assert.lengthOf(shoppingList, 6);
```

Dovresti usare il metodo `unshift` per aggiungere la stringa `"Chocolate Cake"` all'inizio dell'array `shoppingList`.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

Dovresti chiamare `getShoppingListMsg(shoppingList)` all'interno di `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
