---
id: 66c730ee6ae076281721d0b9
title: Passaggio 14
challengeType: 1
dashedName: step-14
---

# --description--

Nelle lezioni precedenti, hai imparato come rimuovere elementi dalla fine di un array usando il metodo `pop`.

Ecco un promemoria su come usare il metodo `pop`:

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

Usa il metodo `pop` per rimuovere l’ultimo elemento dall’array `shoppingList`.

# --hints--

Il tuo array `shoppingList` dovrebbe ora contenere `5` elementi.

```js
assert.lengthOf(shoppingList, 5);
```

Il tuo array `shoppingList` non dovrebbe più contenere la stringa `"Potato Chips"`.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
