---
id: 66cbe2319d3845545a293a0b
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Ora è il momento di vedere il messaggio registrato nella console.

Aggiungi un `console.log` e chiama la funzione `getShoppingListMsg` con l'array `shoppingList` come argomento all'interno di `console.log` per vedere il messaggio registrato nella console.

# --hints--

Dovresti chiamare la funzione `getShoppingListMsg(shoppingList)` all'interno di `console.log`.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
