---
id: 587d7db1367417b2b2512b85
title: 將子物件的原型設定為父物件的實體
challengeType: 1
forumTopicId: 301325
dashedName: set-the-childs-prototype-to-an-instance-of-the-parent
---

# --description--

在前一個挑戰中，你看到了從超型別（或父類別）`Animal` 繼承行為的第一步：建立一個新的 `Animal` 實體。

這個挑戰涵蓋下一步驟：將子型別（或子類型）—在此案例中為 `Bird`—的 `prototype` 設定為 `Animal` 的一個實體。

```js
Bird.prototype = Object.create(Animal.prototype);
```

請記住，`prototype` 就像是建立物件的「配方」。在某種程度上，`Bird` 的配方現在包含了來自 `Animal` 的所有關鍵「成分」。

```js
let duck = new Bird("Donald");
duck.eat();
```

`duck` 繼承了 `Animal` 的所有屬性，包括 `eat` 方法。

# --instructions--

修改程式碼，使 `Dog` 的實體繼承自 `Animal`。

# --hints--

`Dog.prototype` 應該是 `Animal` 的一個實體。

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }

// Only change code below this line


let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);

let beagle = new Dog();
beagle.eat();
```
