---
id: 587d7db0367417b2b2512b81
title: 了解物件的原型來自何處
challengeType: 1
forumTopicId: 301330
dashedName: understand-where-an-objects-prototype-comes-from
---

# --description--

就像人類從父母繼承基因一樣，物件直接從建立它的建構式函式繼承它的 `prototype`。例如，這裡的 `Bird` 建構式建立了 `duck` 物件：

```js
function Bird(name) {
  this.name = name;
}

let duck = new Bird("Donald");
```

`duck` 從 `Bird` 建構式繼承它的 `prototype`。你可以用 `isPrototypeOf` 方法顯示這個關係：

```js
Bird.prototype.isPrototypeOf(duck);
```

這將傳回 `true`。

# --instructions--

使用 `isPrototypeOf` 來檢查 `beagle` 的 `prototype`。

# --hints--

你應該顯示 `Dog.prototype` 是 `beagle` 的 `prototype`。

```js
assert(/Dog\.prototype\.isPrototypeOf\(beagle\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
```
