---
id: 587d7daf367417b2b2512b7d
title: 迭代所有屬性
challengeType: 1
forumTopicId: 301320
dashedName: iterate-over-all-properties
---

# --description--

你現在已經看到兩種屬性：<dfn>自有屬性</dfn>和 `prototype` 屬性。自有屬性是直接定義在物件實體本身上。而 `prototype` 屬性是定義在 `prototype` 上。

```js
function Bird(name) {
  this.name = name;  //own property
}

Bird.prototype.numLegs = 2; // prototype property

let duck = new Bird("Donald");
```

以下是如何將 `duck` 自身的屬性添加到陣列 `ownProps`，以及將 `prototype` 屬性添加到陣列 `prototypeProps`：

```js
let ownProps = [];
let prototypeProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  } else {
    prototypeProps.push(property);
  }
}

console.log(ownProps);
console.log(prototypeProps);
```

`console.log(ownProps)` 會在主控台顯示 `["name"]`，而 `console.log(prototypeProps)` 會顯示 `["numLegs"]`。

# --instructions--

將 `beagle` 的所有自有屬性添加到陣列 `ownProps`。將 `Dog` 的所有 `prototype` 屬性添加到陣列 `prototypeProps`。

# --hints--

`ownProps` 陣列應該只包含 `name`。

```js
assert.deepEqual(ownProps, ['name']);
```

`prototypeProps` 陣列應該只包含 `numLegs`。

```js
assert.deepEqual(prototypeProps, ['numLegs']);
```

你應該在不使用內建方法 `Object.keys()` 的情況下解決此挑戰。

```js
assert(!/\Object.keys/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];
for (let prop in beagle) {
  if (beagle.hasOwnProperty(prop)) {
    ownProps.push(prop);
  } else {
    prototypeProps.push(prop);
  }
}
```
