---
id: 587d7db2367417b2b2512b8c
title: 使用 IIFE 來創建模組
challengeType: 1
forumTopicId: 301332
dashedName: use-an-iife-to-create-a-module
---

# --description--

立即執行的函式表達式（IIFE）常用於將相關功能群組到單一物件或<dfn>模組</dfn>中。例如，先前的挑戰定義了兩個 mixins：

```js
function glideMixin(obj) {
  obj.glide = function() {
    console.log("Gliding on the water");
  };
}
function flyMixin(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  };
}
```

我們可以將這些 mixin 群組到一個模組中，如下所示：

```js
let motionModule = (function () {
  return {
    glideMixin: function(obj) {
      obj.glide = function() {
        console.log("Gliding on the water");
      };
    },
    flyMixin: function(obj) {
      obj.fly = function() {
        console.log("Flying, wooosh!");
      };
    }
  }
})();
```

請注意，你有一個立即執行的函式表達式（IIFE），它會傳回一個物件 `motionModule`。這個傳回的物件包含所有混入行為作為物件的屬性。模組樣式的優點是所有的動作行為都可以包裝成一個單一物件，然後可以被你的程式碼的其他部分使用。以下是使用它的範例：

```js
motionModule.glideMixin(duck);
duck.glide();
```

# --instructions--

創建一個名為 `funModule` 的模組來包裝兩個 mixin `isCuteMixin` 和 `singMixin`。`funModule` 應該傳回一個物件。

# --hints--

`funModule` 應該被定義並傳回一個物件。

```js
assert(typeof funModule === 'object');
```

`funModule.isCuteMixin` 應該存取一個函式（程式）。

```js
assert(typeof funModule.isCuteMixin === 'function');
```

`funModule.singMixin` 應該存取一個函式（程式）。

```js
assert(typeof funModule.singMixin === 'function');
```

# --seed--

## --seed-contents--

```js
let isCuteMixin = function(obj) {
  obj.isCute = function() {
    return true;
  };
};
let singMixin = function(obj) {
  obj.sing = function() {
    console.log("Singing to an awesome tune");
  };
};
```

# --solutions--

```js
const funModule = (function () {
  return {
    isCuteMixin: obj => {
      obj.isCute = () => true;
    },
    singMixin: obj => {
      obj.sing = () => console.log("Singing to an awesome tune");
    }
  };
})();
```
