---
id: 587d7db2367417b2b2512b89
title: 使用 Mixin 為不相關的物件添加共通行為
challengeType: 1
forumTopicId: 301331
dashedName: use-a-mixin-to-add-common-behavior-between-unrelated-objects
---

# --description--

如你所見，行為是透過繼承來共享的。然而，在某些情況下，繼承並不是最佳的解決方案。繼承對於像 `Bird` 和 `Airplane` 這類無關的物件並不適用。它們都能飛行，但 `Bird` 並不是 `Airplane` 的一種型別，反之亦然。

對於不相關的物件，最好使用 <dfn>mixins</dfn>。mixin 允許其他物件使用一組函式（程式）。

```js
let flyMixin = function(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  }
};
```

`flyMixin` 接受任何物件並賦予它 `fly` 方法。

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let plane = {
  model: "777",
  numPassengers: 524
};

flyMixin(bird);
flyMixin(plane);
```

這裡將 `bird` 和 `plane` 傳入 `flyMixin`，然後它會將 `fly` 函式指定給每個物件。現在 `bird` 和 `plane` 都可以飛了：

```js
bird.fly();
plane.fly();
```

主控台會顯示字串 `Flying, wooosh!` 兩次，每次對應一個 `.fly()` 呼叫。

注意 mixin 如何允許不相關的物件 `bird` 和 `plane` 重複使用相同的 `fly` 方法。

# --instructions--

創建一個名為 `glideMixin` 的 mixin，該 mixin 定義一個名為 `glide` 的方法。然後使用 `glideMixin` 為 `bird` 和 `boat` 添加滑行的能力。

# --hints--

你的程式碼應該宣告一個 `glideMixin` 變數，該變數是一個函式（程式）。

```js
assert(typeof glideMixin === 'function');
```

你的程式碼應該在 `bird` 物件上使用 `glideMixin`，以賦予它 `glide` 方法。

```js
assert(typeof bird.glide === 'function');
```

你的程式碼應該在 `boat` 物件上使用 `glideMixin`，以賦予它 `glide` 方法。

```js
assert(typeof boat.glide === 'function');
```

# --seed--

## --seed-contents--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};

// Only change code below this line
```

# --solutions--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};
function glideMixin (obj) {
  obj.glide = () => 'Gliding!';
}

glideMixin(bird);
glideMixin(boat);
```
