---
id: 587d7db1367417b2b2512b87
title: 在繼承之後添加方法
challengeType: 1
forumTopicId: 301315
dashedName: add-methods-after-inheritance
---

# --description--

繼承其 `prototype` 物件自超型建構式的建構式函式，除了繼承的方法外，仍然可以擁有自己的方法。

例如，`Bird` 是一個建構式，它從 `Animal` 繼承其 `prototype`：

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
Bird.prototype.constructor = Bird;
```

除了從 `Animal` 繼承的內容之外，你還想為 `Bird` 物件添加獨有的行為。在這裡，`Bird` 將會有一個 `fly()` 函式。函式是以與任何建構式相同的方式添加到 `Bird` 的 `prototype`。

```js
Bird.prototype.fly = function() {
  console.log("I'm flying!");
};
```

現在 `Bird` 的實體將同時擁有 `eat()` 和 `fly()` 方法。

```js
let duck = new Bird();
duck.eat();
duck.fly();
```

`duck.eat()` 會在主控台顯示字串 `nom nom nom`，而 `duck.fly()` 會顯示字串 `I'm flying!`。

# --instructions--

將所有必要的程式碼添加到你的 `Dog` 物件，使其繼承自 `Animal`，並將 `Dog` 的 `prototype` 建構式設定為 `Dog`。然後為 `Dog` 物件添加一個 `bark()` 方法，使 `beagle` 可以同時執行 `eat()` 和 `bark()`。`bark()` 方法應該在主控台列印 `Woof!`。

# --hints--

`Animal` 不應該回應 `bark()` 方法。

```js
assert(typeof Animal.prototype.bark == 'undefined');
```

`Dog` 應該繼承自 `Animal` 的 `eat()` 方法。

```js
assert(typeof Dog.prototype.eat == 'function');
```

`Dog` 原型應該有一個 `bark()` 方法。

```js
assert('bark' in Dog.prototype);
```

`beagle` 應該是 `Animal` 的一個 `instanceof`。

```js
assert(beagle instanceof Animal);
```

`beagle` 的建構式應該設定為 `Dog`。

```js
assert(beagle.constructor === Dog);
```

`beagle.eat()` 應該會記錄字串 `nom nom nom`

```js
const spy = __helpers.spyOn(console, 'log');
beagle.eat();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'nom nom nom');
```

`beagle.bark()` 應該會記錄字串 `Woof!`

```js
const spy = __helpers.spyOn(console, 'log');
beagle.bark();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'Woof!');
```

# --seed--

## --seed-contents--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }

// Only change code below this line




// Only change code above this line

let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Dog.prototype.bark = function () {
  console.log('Woof!');
};
let beagle = new Dog();

beagle.eat();
beagle.bark();
```
