---
id: 587d7dac367417b2b2512b73
title: 創建一個基本的 JavaScript 物件
challengeType: 1
forumTopicId: 301317
dashedName: create-a-basic-javascript-object
---

# --description--

想想人們每天看到的事物，例如汽車、商店和鳥類。這些都是 <dfn>物件</dfn>：人們可以觀察和互動的有形事物。

這些物件有哪些特性？汽車有車輪。商店販售項目。鳥類有翅膀。

這些特性，或稱為 <dfn>properties</dfn>，定義了物件的組成。請注意，相似的物件具有相同的特性，但這些特性的值可能不同。 例如，所有汽車都有車輪，但並非所有汽車的車輪數量都相同。

JavaScript 中的物件用來模擬現實世界的物件，賦予它們屬性和行為，就像它們的現實世界對應物一樣。以下是一個使用這些概念來創建 `duck` 物件的範例：

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
```

這個 `duck` 物件有兩個屬性/值配對：`name` 為 `Aflac`，`numLegs` 為 2。

# --instructions--

創建一個具有 `name` 和 `numLegs` 屬性的 `dog` 物件，並分別將它們設為字串和數字。

# --hints--

`dog` 應該是一個物件。

```js
assert(typeof dog === 'object');
```

`dog` 應該有一個 `name` 屬性設定為字串。

```js
assert(typeof dog.name === 'string');
```

`dog` 應該有一個 `numLegs` 屬性，並設定為數字。

```js
assert(typeof dog.numLegs === 'number');
```

# --seed--

## --seed-contents--

```js
let dog = {

};
```

# --solutions--

```js
let dog = {
  name: '',
  numLegs: 4
};
```
