---
id: 587d7db1367417b2b2512b86
title: 重設繼承的建構式屬性
challengeType: 1
forumTopicId: 301324
dashedName: reset-an-inherited-constructor-property
---

# --description--

當一個物件從另一個物件繼承其 `prototype` 時，它也會繼承超型別的建構式屬性。

這裡有一個範例：

```js
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
let duck = new Bird();
duck.constructor
```

但是 `duck` 和所有 `Bird` 的實體應該顯示它們是由 `Bird` 建構，而不是由 `Animal` 建構。為此，你可以手動將 `Bird` 的建構式屬性設定為 `Bird` 物件：

```js
Bird.prototype.constructor = Bird;
duck.constructor
```

# --instructions--

修正程式碼，使 `duck.constructor` 和 `beagle.constructor` 傳回它們各自的建構式。

# --hints--

`Bird.prototype` 應該是 `Animal` 的一個實體。

```js
assert(Animal.prototype.isPrototypeOf(Bird.prototype));
```

`duck.constructor` 應該傳回 `Bird`。

```js
assert(duck.constructor === Bird);
```

`Dog.prototype` 應該是 `Animal` 的一個實體。

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

`beagle.constructor` 應該傳回 `Dog`。

```js
assert(beagle.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Animal() { }
function Bird() { }
function Dog() { }

Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);

// Only change code below this line



let duck = new Bird();
let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
function Bird() { }
function Dog() { }
Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Bird.prototype.constructor = Bird;
let duck = new Bird();
let beagle = new Dog();
```
