---
id: 587d7dae367417b2b2512b7c
title: 使用原型屬性來減少重複的程式碼
challengeType: 1
forumTopicId: 301336
dashedName: use-prototype-properties-to-reduce-duplicate-code
---

# --description--

由於 `numLegs` 在所有 `Bird` 實體中可能具有相同的值，你實際上在每個 `Bird` 實體內都有一個重複的變數 `numLegs`。

當只有兩個實體時，這可能不是問題，但試想如果有數百萬個實體，那將會有大量重複的變數。

更好的方法是使用 `Bird` 的 `prototype`。`prototype` 中的屬性會在所有 `Bird` 實體之間共享。以下是如何為 `Bird prototype` 添加 `numLegs`：

```js
Bird.prototype.numLegs = 2;
```

現在所有 `Bird` 的實體都有 `numLegs` 屬性。

```js
console.log(duck.numLegs);
console.log(canary.numLegs);
```

由於所有實體自動擁有 `prototype` 上的屬性，將 `prototype` 視為建立物件的「配方」。請注意，`duck` 和 `canary` 的 `prototype` 是作為 `Bird` 建構式的一部分，即 `Bird.prototype`。

# --instructions--

為你的 `Dog` 的 `prototype` 添加一個 `numLegs` 屬性

# --hints--

`beagle` 應該有一個 `numLegs` 屬性。

```js
assert(beagle.numLegs !== undefined);
```

`beagle.numLegs` 應該是一個數字。

```js
assert(typeof beagle.numLegs === 'number');
```

`numLegs` 應該是 `prototype` 屬性，而不是自身屬性。

```js
assert(beagle.hasOwnProperty('numLegs') === false);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}



// Only change code above this line
let beagle = new Dog("Snoopy");
```

# --solutions--

```js
function Dog (name) {
  this.name = name;
}
Dog.prototype.numLegs = 4;
let beagle = new Dog("Snoopy");
```
