---
id: 587d7dad367417b2b2512b76
title: 使用 this 關鍵字讓程式碼更具可重用性
challengeType: 1
forumTopicId: 301321
dashedName: make-code-more-reusable-with-the-this-keyword
---

# --description--

上一個挑戰為 `duck` 物件介紹了一個方法。它在傳回述語中使用 `duck.name` 點記法來存取 `name` 屬性的值：

```js
sayName: function() {return "The name of this duck is " + duck.name + ".";}
```

雖然這是存取物件屬性的一種有效方式，但這裡有一個陷阱。如果變數名稱改變，任何取用原始名稱的程式碼也需要更新。在簡短的物件定義中，這不是問題，但如果物件有許多對其屬性的引用，錯誤的機率會更高。

避免這些問題的一種方法是使用 `this` 關鍵字：

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + this.name + ".";}
};
```

`this` 是一個深奧的主題，上述範例只是使用它的一種方式。在目前的背景關係中，`this` 指的是該方法相關的物件：`duck`。如果物件的名稱改為 `mallard`，就不需要在程式碼中尋找所有對 `duck` 的引用。這使程式碼可重複使用且更容易閱讀。

# --instructions--

修改 `dog.sayLegs` 方法，移除任何對 `dog` 的引用。使用 `duck` 範例作為指引。

# --hints--

`dog.sayLegs()` 應該傳回給定的字串。

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

你的程式碼應該使用 `this` 關鍵字來存取 `dog` 的 `numLegs` 屬性。

```js
assert(__helpers.removeJSComments(code).match(/this\.numLegs/g));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs: function() {return "This dog has " + dog.numLegs + " legs.";}
};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
