---
id: 587d7daf367417b2b2512b7e
title: 了解建構式屬性
challengeType: 1
forumTopicId: 301327
dashedName: understand-the-constructor-property
---

# --description--

在先前挑戰中創建的物件實體 `duck` 和 `beagle` 上，有一個特殊的 `constructor` 屬性：

```js
let duck = new Bird();
let beagle = new Dog();

console.log(duck.constructor === Bird); 
console.log(beagle.constructor === Dog);
```

這兩個 `console.log` 呼叫都會在主控台顯示 `true`。

請注意，`constructor` 屬性是對建立該實體的建構式函式的引用。`constructor` 屬性的優點是可以透過檢查此屬性來判斷物件的種類。以下是一個如何使用此方法的範例：

```js
function joinBirdFraternity(candidate) {
  if (candidate.constructor === Bird) {
    return true;
  } else {
    return false;
  }
}
```

**注意：** 由於 `constructor` 屬性可以被覆寫（這將在接下來的兩個挑戰中涵蓋），通常最好使用 `instanceof` 方法來檢查物件的型別。

# --instructions--

撰寫一個 `joinDogFraternity` 函式（程式），該函式接受一個 `candidate` 參數，並使用 `constructor` 屬性，如果該候選者是 `Dog`，則傳回 `true`，否則傳回 `false`。

# --hints--

`joinDogFraternity` 應該被定義為一個函式（程式）。

```js
assert(typeof joinDogFraternity === 'function');
```

如果 `candidate` 是 `Dog` 的一個實體，`joinDogFraternity` 應該傳回 `true`。

```js
assert(joinDogFraternity(new Dog('')) === true);
```

`joinDogFraternity` 應該使用 `constructor` 屬性。

```js
assert(/\.constructor/.test(__helpers.removeJSComments(code)) && !/instanceof/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
function joinDogFraternity(candidate) {

}
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
function joinDogFraternity(candidate) {
  return candidate.constructor === Dog;
}
```
