---
id: 587d7dad367417b2b2512b75
title: 在物件上創建方法
challengeType: 1
forumTopicId: 301318
dashedName: create-a-method-on-an-object
---

# --description--

物件可以有一種特殊型別的屬性，稱為 <dfn>method</dfn>。

方法是屬於函式的屬性。這為物件添加不同的行為。以下是帶有方法的 `duck` 範例：

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + duck.name + ".";}
};
duck.sayName();
```

範例為 `sayName` 方法添加了一個函式，該函式傳回一個句子，說明 `duck` 的名稱。請注意，該方法在傳回述語中使用 `duck.name` 存取了 `name` 屬性。下一個挑戰將涵蓋另一種執行此操作的方法。

# --instructions--

使用 `dog` 物件，為它添加一個名為 `sayLegs` 的方法。該方法應該傳回句子 `This dog has 4 legs.`。

# --hints--

`dog.sayLegs()` 應該是一個函式（程式）。

```js
assert(typeof dog.sayLegs === 'function');
```

`dog.sayLegs()` 應該傳回給定的字串－請注意標點符號和空格很重要。

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,

};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
