---
id: 587d7dae367417b2b2512b7a
title: 使用 `instanceof` 驗證物件的建構式
challengeType: 1
forumTopicId: 301337
dashedName: verify-an-objects-constructor-with-instanceof
---

# --description--

每當建構式函式創建一個新的物件時，該物件即被稱為其建構式的<dfn>實體</dfn>。JavaScript 提供了一種方便的方法來驗證這一點，即使用 `instanceof` 運算子。`instanceof` 允許你將物件與建構式進行比較，根據該物件是否由該建構式創建，傳回 `true` 或 `false`。以下是一個範例：

```js
let Bird = function(name, color) {
  this.name = name;
  this.color = color;
  this.numLegs = 2;
}

let crow = new Bird("Alexis", "black");

crow instanceof Bird;
```

這個 `instanceof` 方法會傳回 `true`。

如果物件是在未使用建構式的情況下創建，`instanceof` 將驗證它不是該建構式的實體：

```js
let canary = {
  name: "Mildred",
  color: "Yellow",
  numLegs: 2
};

canary instanceof Bird;
```

這個 `instanceof` 方法會傳回 `false`。

# --instructions--

創建一個新的 `House` 建構式實體，將其命名為 `myHouse` 並傳入多個臥室數量。然後，使用 `instanceof` 驗證它是否為 `House` 的實體。

# --hints--

`myHouse` 應該有一個 `numBedrooms` 屬性設定為數字。

```js
assert(typeof myHouse.numBedrooms === 'number');
```

你應該使用 `instanceof` 運算子來驗證 `myHouse` 是否為 `House` 的實體。

```js
assert(/myHouse\s*instanceof\s*House/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}

// Only change code below this line
```

# --solutions--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}
const myHouse = new House(4);
console.log(myHouse instanceof House);
```
