---
id: 587d7db2367417b2b2512b8b
title: 了解立即呼叫的函式表達式（IIFE）
challengeType: 1
forumTopicId: 301328
dashedName: understand-the-immediately-invoked-function-expression-iife
---

# --description--

JavaScript 中一個常見的樣式是在函式宣告後立即執行該函式：

```js
(function () {
  console.log("Chirp, chirp!");
})();
```

這是一個立即執行的匿名函式（程式）表達式，並立即輸出 `Chirp, chirp!`。

請注意，該函式沒有名稱且未儲存在變數中。函式表達式末尾的兩個小括弧 () 會使其立即執行或喚起。這種樣式稱為 <dfn>立即執行函式表達式</dfn> 或 <dfn>IIFE</dfn>。

# --instructions--

重寫函式（程式）`makeNest` 並移除其呼叫，改為匿名立即執行函式（表達式、運算式）（IIFE）。

# --hints--

該函式（程式）應該是匿名的。

```js
assert(/\((function|\(\))(=>|\(\)){?/.test(__helpers.removeJSComments(code).replace(/\s/g, '')));
```

你的函式（程式）應該在表達式結尾有小括弧，以立即呼叫它。

```js
assert(/\(.*(\)\(|\}\(\))\)/.test(__helpers.removeJSComments(code).replace(/[\s;]/g, '')));
```

# --seed--

## --seed-contents--

```js
function makeNest() {
  console.log("A cozy nest is ready");
}

makeNest();
```

# --solutions--

```js
(function () {
  console.log("A cozy nest is ready");
})();
```

---

```js
(function () {
  console.log("A cozy nest is ready");
}());
```

---

```js
(() => {
  console.log("A cozy nest is ready");
})();
```

---

```js
(() =>
  console.log("A cozy nest is ready")
)();
```
