---
id: 587d7daf367417b2b2512b80
title: 更改原型時，請記得設定建構式屬性
challengeType: 1
forumTopicId: 301323
dashedName: remember-to-set-the-constructor-property-when-changing-the-prototype
---

# --description--

手動將原型設定為新物件有一個關鍵的副作用。它會清除 `constructor` 屬性！這個屬性可以用來檢查是哪個建構式函式創建了該實體，但由於該屬性已被覆寫，現在會給出錯誤的結果：

```js
duck.constructor === Bird;
duck.constructor === Object;
duck instanceof Bird;
```

依序，這些表達式會計算為 `false`、`true` 和 `true`。

為了解決此問題，每當手動將原型設定為新的物件時，請記得定義 `constructor` 屬性：

```js
Bird.prototype = {
  constructor: Bird,
  numLegs: 2,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name); 
  }
};
```

# --instructions--

在 `Dog` `prototype` 上定義 `constructor` 屬性。

# --hints--

`Dog.prototype` 應該設定 `constructor` 屬性。

```js
assert(Dog.prototype.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
Dog.prototype = {

  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
  constructor: Dog,
  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```
