---
id: 587d7db1367417b2b2512b88
title: 改寫繼承的方法
challengeType: 1
forumTopicId: 301322
dashedName: override-inherited-methods
---

# --description--

在之前的課程中，你學到物件可以透過取用它的 `prototype` 物件，從另一個物件繼承其行為（方法）。

```js
ChildObject.prototype = Object.create(ParentObject.prototype);
```

然後 `ChildObject` 透過將它們串鏈到自己的 `prototype` 上來接收自己的方法。

```js
ChildObject.prototype.methodName = function() {...};
```

可以改寫繼承的方法。做法相同－透過為 `ChildObject.prototype` 添加一個與要改寫的方法同名的方法來完成。以下是 `Bird` 改寫從 `Animal` 繼承的 `eat()` 方法的範例：

```js
function Animal() { }
Animal.prototype.eat = function() {
  return "nom nom nom";
};
function Bird() { }

Bird.prototype = Object.create(Animal.prototype);

Bird.prototype.eat = function() {
  return "peck peck peck";
};
```

如果你有一個實體 `let duck = new Bird();` 並且你呼叫 `duck.eat()`，JavaScript 就會在 `duck` 的 `prototype` 串鏈上尋找該方法：

1.  `duck` => 這裡有定義 `eat()` 嗎？沒有。
2.  `Bird` => 這裡有定義 `eat()` 嗎？ => 有。執行它並停止搜尋。
3.  `Animal` => `eat()` 也已定義，但 JavaScript 在到達此層次前已停止搜尋。
4.  物件 => JavaScript 在到達此層次前已停止搜尋。

# --instructions--

為 `Penguin` 改寫 `fly()` 方法，使其傳回字串 `Alas, this is a flightless bird.`

# --hints--

`penguin.fly()` 應該傳回字串 `Alas, this is a flightless bird.`

```js
assert(penguin.fly() === 'Alas, this is a flightless bird.');
```

`bird.fly()` 方法應該傳回字串 `I am flying!`

```js
assert(new Bird().fly() === 'I am flying!');
```

# --seed--

## --seed-contents--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;

// Only change code below this line



// Only change code above this line

let penguin = new Penguin();
console.log(penguin.fly());
```

# --solutions--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;
Penguin.prototype.fly = () => 'Alas, this is a flightless bird.';
let penguin = new Penguin();
console.log(penguin.fly());
```
