---
id: 587d7db0367417b2b2512b83
title: 使用繼承以避免重複自己
challengeType: 1
forumTopicId: 301334
dashedName: use-inheritance-so-you-dont-repeat-yourself
---

# --description--

程式設計中有一個原則稱為 <dfn>不要重複自己（DRY）</dfn>。重複的程式碼之所以是問題，是因為任何變更都需要在多個地方修正程式碼。這通常代表程式設計師需要更多工作量，且錯誤的機會也更多。

請注意以下範例中，`describe` 方法是由 `Bird` 和 `Dog` 共享的：

```js
Bird.prototype = {
  constructor: Bird,
  describe: function() {
    console.log("My name is " + this.name);
  }
};

Dog.prototype = {
  constructor: Dog,
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

`describe` 方法在兩個地方重複出現。可以透過建立一個名為 `Animal` 的 `supertype`（或父類別）來編輯程式碼，以遵循 DRY 原則：

```js
function Animal() { };

Animal.prototype = {
  constructor: Animal, 
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

由於 `Animal` 包含 `describe` 方法，你可以從 `Bird` 和 `Dog` 中移除它：

```js
Bird.prototype = {
  constructor: Bird
};

Dog.prototype = {
  constructor: Dog
};
```

# --instructions--

`eat` 方法在 `Cat` 和 `Bear` 中都重複出現。請依照 DRY 原則，將 `eat` 方法移動到 `Animal` `supertype` 中來編輯程式碼。

# --hints--

`Animal.prototype` 應該有 `eat` 屬性。

```js
assert(Animal.prototype.hasOwnProperty('eat'));
```

`Bear.prototype` 不應該有 `eat` 屬性。

```js
assert(!Bear.prototype.hasOwnProperty('eat'));
```

`Cat.prototype` 不應該有 `eat` 屬性。

```js
assert(!Cat.prototype.hasOwnProperty('eat'));
```

# --seed--

## --seed-contents--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,

};
```

# --solutions--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
```
