---
id: 587d7dad367417b2b2512b77
title: 定義建構式函式（Constructor Function）
challengeType: 1
forumTopicId: 16804
dashedName: define-a-constructor-function
---

# --description--

<dfn>建構式</dfn>是用來產生新物件的函式（程式）。它們定義將屬於新物件的屬性和行為。可以將它們視為新物件產生的藍圖。

這裡是一個建構式的範例：

```js
function Bird() {
  this.name = "Albert";
  this.color = "blue";
  this.numLegs = 2;
}
```

這個建構式定義了一個 `Bird` 物件，其屬性 `name`、`color` 和 `numLegs` 分別設定為 Albert、blue 和 2。建構式遵循一些慣例：

<ul><li>Constructors are defined with a capitalized name to distinguish them from other functions that are not <code>constructors</code>.</li><li>Constructors use the keyword <code>this</code> to set properties of the object they will create. Inside the constructor, <code>this</code> refers to the new object it will create.</li><li>Constructors define properties and behaviors instead of returning a value as other functions might.</li></ul>

# --instructions--

創建一個建構式 `Dog`，其屬性 `name`、`color` 和 `numLegs` 分別設定為字串、字串和數字。

# --hints--

`Dog` 應該有一個 `name` 屬性設定為字串。

```js
assert(typeof new Dog().name === 'string');
```

`Dog` 應該有一個 `color` 屬性設定為字串。

```js
assert(typeof new Dog().color === 'string');
```

`Dog` 應該有一個 `numLegs` 屬性，並設定為數字。

```js
assert(typeof new Dog().numLegs === 'number');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function Dog (name, color, numLegs) {
  this.name = 'name';
  this.color = 'color';
  this.numLegs = 4;
}
```
