---
id: 587d7dac367417b2b2512b74
title: 使用點記法存取物件的屬性
challengeType: 1
forumTopicId: 301333
dashedName: use-dot-notation-to-access-the-properties-of-an-object
---

# --description--

最後一個挑戰創建了一個具有各種屬性的物件。現在你將看到如何存取那些屬性的值。以下是一個範例：

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
console.log(duck.name);
```

點記法用於物件名稱 `duck`，後接屬性名稱 `name`，以存取 `Aflac` 的值。

# --instructions--

將 `dog` 物件的兩個屬性列印到你的主控台。

# --hints--

你的程式碼應該使用 `console.log` 來列印 `dog` 物件中 `name` 屬性的值。

```js
assert(/console.log\(.*dog\.name.*\)/g.test(__helpers.removeJSComments(code)));
```

你的程式碼應該使用 `console.log` 來列印 `dog` 物件的 `numLegs` 屬性值。

```js
assert(/console.log\(.*dog\.numLegs.*\)/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
// Only change code below this line
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
console.log(dog.name);
console.log(dog.numLegs);
```
