---
id: 587d7dae367417b2b2512b7b
title: 了解自身屬性
challengeType: 1
forumTopicId: 301326
dashedName: understand-own-properties
---

# --description--

在以下範例中，`Bird` 建構式定義了兩個屬性：`name` 和 `numLegs`：

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let duck = new Bird("Donald");
let canary = new Bird("Tweety");
```

`name` 和 `numLegs` 被稱為 <dfn>自身屬性</dfn>，因為它們是直接定義在實體物件上的。這表示 `duck` 和 `canary` 各自擁有這些屬性的獨立副本。事實上，每個 `Bird` 的實體都會有這些屬性的獨立副本。以下程式碼將 `duck` 的所有自身屬性添加到陣列 `ownProps`：

```js
let ownProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  }
}

console.log(ownProps);
```

主控台會顯示值 `["name", "numLegs"]`。

# --instructions--

將 `canary` 的 own properties 添加到陣列 `ownProps`。

# --hints--

`ownProps` 應該包含值 `numLegs` 和 `name`。

```js
assert(ownProps.indexOf('name') !== -1 && ownProps.indexOf('numLegs') !== -1);
```

你應該在不使用內建方法 `Object.keys()` 的情況下解決此挑戰。

```js
assert(!/Object(\.keys|\[(['"`])keys\2\])/.test(__helpers.removeJSComments(code)));
```

你應該在不硬編碼 `ownProps` 陣列的情況下解決這個挑戰。

```js
assert(
  !/\[\s*(?:'|")(?:name|numLegs)|(?:push|concat)\(\s*(?:'|")(?:name|numLegs)/.test(
    __helpers.removeJSComments(code)
  )
);
```

# --seed--

## --seed-contents--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
let ownProps = [];
// Only change code below this line
```

# --solutions--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
function getOwnProps (obj) {
  const props = [];

  for (let prop in obj) {
    if (obj.hasOwnProperty(prop)) {
      props.push(prop);
    }
  }

  return props;
}

const ownProps = getOwnProps(canary);
```
