---
id: 587d7db0367417b2b2512b84
title: 從超型別繼承行為
challengeType: 1
forumTopicId: 301319
dashedName: inherit-behaviors-from-a-supertype
---

# --description--

在前一個挑戰中，你創建了一個名為 `Animal` 的 `supertype`，該 `supertype` 定義了所有動物共有的行為：

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
```

這個和下一個挑戰將涵蓋如何在 `Bird` 和 `Dog` 中重複使用 `Animal` 的方法，而不需重新定義它們。它使用一種稱為繼承的技術。這個挑戰涵蓋第一步：產生 `supertype`（或父類別）的實體。你已經知道一種使用 `new` 運算子來創建 `Animal` 實體的方法：

```js
let animal = new Animal();
```

使用這種語法進行繼承有一些缺點，對於本挑戰的範圍來說過於複雜。相反地，這裡有一種沒有那些缺點的替代方法：

```js
let animal = Object.create(Animal.prototype);
```

`Object.create(obj)` 會創建一個新的物件，並將 `obj` 設定為新物件的 `prototype`。回想一下，`prototype` 就像是創建物件的「配方」。透過將 `animal` 的 `prototype` 設定為 `Animal` 的 `prototype`，你實際上是給予 `animal` 實體與其他任何 `Animal` 實體相同的「配方」。

```js
animal.eat();
animal instanceof Animal;
```

這裡的 `instanceof` 方法會傳回 `true`。

# --instructions--

使用 `Object.create` 來創建兩個名為 `duck` 和 `beagle` 的 `Animal` 實體。

# --hints--

`duck` 變數應該被定義。

```js
assert(typeof duck !== 'undefined');
```

`beagle` 變數應該被定義。

```js
assert(typeof beagle !== 'undefined');
```

`duck` 變數應該用 `Object.create` 來初始化。

```js
assert(
  /(let|const|var)\s{1,}duck\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

`beagle` 變數應該使用 `Object.create` 來初始化。

```js
assert(
  /(let|const|var)\s{1,}beagle\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

`duck` 應該有一個 `Animal` 的 `prototype`。

```js
assert(duck instanceof Animal);
```

`beagle` 應該有一個 `Animal` 的 `prototype`。

```js
assert(beagle instanceof Animal);
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

// Only change code below this line

let duck; // Change this line
let beagle; // Change this line
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
let duck = Object.create(Animal.prototype);
let beagle = Object.create(Animal.prototype);

duck.eat();
beagle.eat();
```
