---
id: 587d7dad367417b2b2512b78
title: 使用建構式來創建物件
challengeType: 1
forumTopicId: 18233
dashedName: use-a-constructor-to-create-objects
---

# --description--

這是前一個挑戰中的 `Bird` 建構式：

```js
function Bird() {
  this.name = "Albert";
  this.color  = "blue";
  this.numLegs = 2;
}

let blueBird = new Bird();
```

**注意：** 建構式內的 `this` 永遠取用正在被建立的物件。

請注意，當呼叫建構式時會使用 `new` 運算子。這告訴 JavaScript 生成一個名為 `blueBird` 的 `Bird` 新實體。若沒有 `new` 運算子，建構式內的 `this` 將不會指向新建立的物件，會導致意外的結果。現在 `blueBird` 擁有 `Bird` 建構式內定義的所有屬性：

```js
blueBird.name;
blueBird.color;
blueBird.numLegs;
```

就像任何其他物件一樣，其屬性可以被存取和修改：

```js
blueBird.name = 'Elvira';
blueBird.name;
```

# --instructions--

使用上一課的 `Dog` 建構式來創建一個新的 `Dog` 實體，並將它指定給變數 `hound`。

# --hints--

`hound` 應該使用 `Dog` 建構式來創建。

```js
assert(hound instanceof Dog);
```

你的程式碼應該使用 `new` 運算子來創建 `Dog` 的實體。

```js
assert(__helpers.removeJSComments(code).match(/new/g));
```

# --seed--

## --seed-contents--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
// Only change code below this line
```

# --solutions--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
const hound = new Dog();
```
