---
id: 587d7daf367417b2b2512b7f
title: 將原型更改為新的物件
challengeType: 1
forumTopicId: 301316
dashedName: change-the-prototype-to-a-new-object
---

# --description--

到目前為止，你一直在單獨為 `prototype` 添加屬性：

```js
Bird.prototype.numLegs = 2;
```

超過幾個屬性後，這會變得乏味。

```js
Bird.prototype.eat = function() {
  console.log("nom nom nom");
}

Bird.prototype.describe = function() {
  console.log("My name is " + this.name);
}
```

一種更高效的方法是將 `prototype` 設定為已經包含屬性的全新物件。這樣，屬性會一次性全部添加：

```js
Bird.prototype = {
  numLegs: 2, 
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --instructions--

透過將 `prototype` 設定為一個新物件，為 `Dog` 的 `prototype` 添加屬性 `numLegs` 以及兩個方法 `eat()` 和 `describe()`。

# --hints--

`Dog.prototype` 應該被設定為一個新的物件。

```js
assert(/Dog\.prototype\s*?=\s*?{/.test(__helpers.removeJSComments(code)));
```

`Dog.prototype` 應該有屬性 `numLegs`。

```js
assert(Dog.prototype.numLegs !== undefined);
```

`Dog.prototype` 應該有方法 `eat()`。

```js
assert(typeof Dog.prototype.eat === 'function');
```

`Dog.prototype` 應該有方法 `describe()`。

```js
assert(typeof Dog.prototype.describe === 'function');
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype = {
  // Only change code below this line

};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
numLegs: 4,
  eat () {
    console.log('nom nom nom');
  },
  describe () {
    console.log('My name is ' + this.name);
  }
};
```
