---
id: 66e2edb23013ca596f8d9568
title: 步骤6
challengeType: 1
dashedName: step-6
---

# --description--

现在是测试你的 `getPunctuationCount` 函数的时候了。

创建一个 `punctuationCount` 变量，并将调用 `getPunctuationCount` 函数且参数为 `"WHAT?!?!?!?!?"` 的结果赋值给它。

之后，将以下内容记录到控制台：`"Punctuation Count: [Punctuation count goes here]"`。将 `[Punctuation count goes here]` 替换为实际的变量名。你可以选择使用模板字串或使用 `+` 操作符进行字串串接。

# --hints--

你应该创建一个 `punctuationCount` 变量。

```js
assert.isNotNull(punctuationCount)
```

你的 `punctuationCount` 变量应设置为 `getPunctuationCount("WHAT?!?!?!?!?")` 的结果。

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

你应该将以下内容记录到控制台：`"Punctuation Count: [Punctuation count goes here]"`。将 `[Punctuation count goes here]` 替换为实际的变量名。确保在这里使用正确的字串串接语法。

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

