---
id: 66e2f06b191e305c00574e4d
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

最后，通过创建一个带有 `sentence` 参数的 `getWordCount` 函数来计算单词数。该函数在被调用时应返回传入句子中的单词总数。

# --hints--

你应该创建一个 `getWordCount` 函数。

```js
assert.isFunction(getWordCount);
```

你的 `getWordCount` 函数应有一个 `sentence` 参数。

```js
assert.match(getWordCount.toString(), /sentence/);
```

你的 `getWordCount` 函数应返回一个数字。

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

当句子是 `"When are you gonna start learning to code?"` 时，`getWordCount` 函数应返回 `8`。

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

当句子是 `"What's going on?"` 时，`getWordCount` 函数应返回 `3`。

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

你的单词计数应忽略分支大小写。

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

你的 `getWordCount` 函数应返回任何句子的正确字数。

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both frontend and backend development"), 7);
```

你的 `getWordCount` 函数应返回空字串或仅包含空格的字串的正确字数。

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
