---
id: 66e2e99caa0cb3570f4822b4
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

现在是测试你的 `getConsonantCount` 函数的时候了。

创建一个 `consonantCount` 变量，并将调用 `getConsonantCount` 函数且参数为 `"Coding is fun"` 的结果赋值给它。

之后，将以下内容记录到控制台：`"Consonant Count: [Consonant count goes here]"`。将 `[Consonant count goes here]` 替换为实际的变量名。你可以选择使用模板字串或使用 `+` 操作符进行字串串接。

# --hints--

你应该创建一个 `consonantCount` 变量。

```js
assert.isNotNull(consonantCount)
```

你的 `consonantCount` 变量应设置为 `getConsonantCount("Coding is fun")` 的结果。

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

你应该使用参数 `"Consonant Count: ${consonantCount}"` 将 `consonantCount` 变量记录到控制台日志。

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
