---
id: 66e2e5d890ac8a54c9d703bd
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

是时候统计辅音了。创建一个带有 `sentence` 参数的 `getConsonantCount` 函数。

在函数内部，使用循环来计算传入函数调用时的 `sentence` 中的辅音字母数量。辅音字母是指不属于以下字符的任何字母：`"aeiou"`。 

你的 `getConsonantCount` 函数必须返回一个数字。

# --hints--

你应该创建一个 `getConsonantCount` 函数。

```js
assert.isFunction(getConsonantCount);
```

你的 `getConsonantCount` 函数应有一个 `sentence` 参数。

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

你的 `getConsonantCount` 函数应返回一个数字。

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

当句子是 `"Coding is fun"` 时，`getConsonantCount` 函数应返回 `7`。

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

当句子是 `"Hello, World!"` 时，`getConsonantCount` 函数应返回 `7`。

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

你的辅音计数应忽略分支大小写。

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

你的 `getConsonantCount` 函数应返回任何句子的正确辅音数。

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
