---
id: 66e2df8c5676fe51cead8ec8
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

现在是测试你的 `getVowelCount` 函数的时候了。

创建一个 `vowelCount` 变量，并将调用 `getVowelCount` 函数且参数为 `"Apples are tasty fruits"` 的结果赋值给它。

之后，将以下内容记录到控制台：`"Vowel Count: [vowel count goes here]"`。将 `[vowel count goes here]` 替换为实际的变量名。你可以选择使用模板字串或使用 `+` 操作符进行字串串接。

# --hints--

你应该创建一个 `vowelCount` 变量。

```js
assert.isNotNull(vowelCount)
```

你的 `vowelCount` 变量应设置为 `getVowelCount("Apples are tasty fruits")` 的结果。

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

你应该使用参数 `"Vowel Count: ${vowelCount}"` 将 `vowelCount` 变量记录到控制台。

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
