---
id: 66e2d680e129e1423116a541
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会中，你将构建一个句子分析器，该分析器将接受一个句子并获取单词数、元音数、构造数和标点符号数。

首先，创建一个名为 `getVowelCount` 的函数，带有一个名为 `sentence` 的参数。你的函数应返回句子中元音的总数。

# --hints--

你应该创建一个 `getVowelCount` 函数。

```js
assert.isFunction(getVowelCount);
```

你的 `getVowelCount` 函数应有一个 `sentence` 参数。

```js
assert.match(getVowelCount.toString(), /sentence/);
```

你的 `getVowelCount` 函数应返回一个数字。

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

当句子是 `"Apples are tasty fruits"` 时，`getVowelCount` 函数应返回 `7`。

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

当句子是 `"Hello, World!"` 时，函数应返回 `3`。

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

你的元音计数应忽略分支大小写。

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

你的 `getVowelCount` 函数应返回任何句子的正确元音计数。

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
