---
id: 66e2eab8a5638f57b637b7cc
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

你现在应该计算标点符号的数量。

创建一个带有 `sentence` 参数的 `getPunctuationCount` 函数。

在函数内部，创建一个循环来计算传入函数调用时的 `sentence` 中的标点字符数量。对于我们的目的，标点字符是指任何不是空格（`" "`）或字母的字符。

你的 `getPunctuationCount` 函数必须返回一个数字。

# --hints--

你应该创建一个 `getPunctuationCount` 函数。

```js
assert.isFunction(getPunctuationCount);
```

你的 `getPunctuationCount` 函数应有一个 `sentence` 参数。

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

你的 `getPunctuationCount` 函数应返回一个数字。

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

当句子是 `"What's going on here?"` 时，`getPunctuationCount` 函数应返回 `2`。

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

当句子是 `"What????!"` 时，`getPunctuationCount` 函数应返回 `5`。

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

你的 `getPunctuationCount` 函数应返回任何句子的正确标点符号数量。

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
