---
id: 7a97b79c096ecb000d03bd36
title: Build an Nth Fibonacci Number Calculator
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**User Stories:**

1. Ви повинні створити функцію з назвою `fibonacci`.
2. Ви повинні визначити масив з назвою `sequence` всередині функції `fibonacci`, і він має бути ініціалізований значеннями `[0, 1]`.
3. Функція `fibonacci` повинна приймати один параметр — додатнє ціле число `n`.
4. Виклик `fibonacci(n)` має використовувати підхід динамічного програмування для обчислення та повернення `n`-го числа з послідовності Фібоначчі, де кожне число є сумою двох попередніх.
5. Кожне обчислене число на позиції `n` у послідовності Фібоначчі має зберігатися у масиві `sequence` за індексом `n - 1`.

# --hints--

Ви повинні мати функцію з назвою `fibonacci`.

```js
assert.isFunction(fibonacci);
```

Ваша функція `fibonacci` має приймати один параметр.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

Ви повинні мати масив з назвою `sequence` всередині функції `fibonacci`, ініціалізований `[0, 1]`.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` should return `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` should return `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` should return `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` should return `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` should return `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` should return `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` should return `610`.

```js
assert.equal(fibonacci(15), 610);
```

У вашому коді не слід використовувати рекурсію.

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
