---
id: 6733b5632fb854bc6c58dc0d
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Ahora crea otro objeto, con las mismas tres propiedades, para representar otro cultivar de lavanda española llamado `Pretty Polly`, y asígnalo a una variable llamada `prettyPolly`.

# --hints--

Debe tener una variable llamada `prettyPolly`.

```js
assert.exists(prettyPolly);
```

Debe asignar un objeto a su variable `prettyPolly`.

```js
assert.isObject(prettyPolly);
```

Su objeto `prettyPolly` debe tener una propiedad `commonName` establecida en `"Spanish lavender"`.

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

Su objeto `prettyPolly` debe tener una propiedad `scientificName` establ...

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

Su objeto `prettyPolly` debe tener una propiedad `cultivar` establecida en `"Pretty Polly"`.

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

Su objeto `prettyPolly` debe tener tres claves.

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```
