---
id: 6734f0ae9d1cf4dfd43a7805
title: Paso 15
challengeType: 1
dashedName: step-15
---

# --description--

Para eliminar un elemento de un objeto de mapa, puede usar el método `delete`, que toma una clave como argumento y elimina la entrada asociada con esa clave. Devuelve un valor booleano que indica si la eliminación ha tenido éxito.

Cree una función llamada `removePlant` que tome un objeto de planta como su argumento y lo elimine del catálogo. Hágala devolver `true` o `false` dependiendo de si la planta especificada ha sido eliminada del catálogo.

Después de eso, registre el resultado de llamar a su nueva función con la cadena `"ballerina"` como su argumento.

# --hints--

Debería tener una función llamada `removePlant`.

```js
assert.isFunction(removePlant);
```

Su función `removePlant` debería tomar un solo argumento.

```js
assert.lengthOf(removePlant, 1);
```

Su función `removePlant` debería llamar al método `delete` en `catalog`.

```js
let flag = false;
const temp = catalog.delete;
try {
    catalog.delete = (p) => {flag = true}
    removePlant("");
    assert.isTrue(flag);
} finally {
    catalog.delete = temp;
}
```

`removePlant(item)` debería eliminar `item` de su catálogo.

```js
const testArr = [];
const temp = catalog.delete;
try {
    catalog.delete = (p) => {testArr.push(p)}
    removePlant("test");
    assert.deepEqual(testArr, ["test"]);
} finally {
    catalog.delete = temp;
}
```

Su función `removePlant` debería devolver un valor booleano que indique si la eliminación ha tenido éxito.

```js
try {
    assert.isFalse(removePlant("test"));
    assert.isTrue(removePlant(royalCrown));
} finally {
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 })
}
```

Debería registrar `removePlant("ballerina")` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*removePlant\s*\(\s*("|')ballerina\1\s*\)\s*\)/);
```



# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

--fcc-editable-region--

--fcc-editable-region--

const sellPlants = (plant, size, potsNo) => {
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```
