---
id: 6734a2385c52c54d916a8b86
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Añade otro par clave-valor a su catálogo. Use su variable `prettyPolly` como clave y haga que el valor sea un objeto con las propiedades `small`, `medium` y `large` cuyos valores vuelvan a ser enteros.

# --hints--

Debe configurar su objeto `prettyPolly` como una clave de su mapa de `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

El valor de la clave `prettyPolly` en el mapa de `catalog` debe ser un objeto con las propiedades `small`, `medium` y `large`.

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

La propiedad `small` del objeto asociado con la clave `prettyPolly` en el mapa de `catalog` debe ser un entero positivo.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

La propiedad `medium` del objeto asociado con la clave `prettyPolly` en el mapa de `catalog` debe ser un entero positivo.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

La propiedad `large` del objeto asociado con la clave `prettyPolly` en el mapa de `catalog` debe ser un entero positivo.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Todavía debe tener la clave `ballerina` en su mapa de `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Su mapa de `catalog` debe contener dos elementos.

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
