---
id: 6734a606e5968a5a19b77171
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Continúe agregando las plantas restantes a su catálogo como lo hizo en los pasos anteriores.

# --hints--

Debe establecer su objeto `willowVale` como una clave de su mapa `catalog`.

```js
assert.isTrue(catalog.has(willowVale));
```

El valor de la clave `willowVale` en el mapa `catalog` debe ser un objeto con las propiedades `small`, `medium` y `large`, cada una con el valor de un número entero positivo.

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Debe establecer su objeto `hidcote` como una clave de su mapa `catalog`.

```js
assert.isTrue(catalog.has(hidcote));
```

El valor de la clave `hidcote` en el mapa `catalog` debe ser un objeto con las propiedades `small`, `medium` y `large`, cada una con el valor de un número entero positivo.

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Debe establecer su objeto `imperialGem` como una clave de su mapa `catalog`.

```js
assert.isTrue(catalog.has(imperialGem));
```

El valor de la clave `imperialGem` en el mapa `catalog` debe ser un objeto con las propiedades `small`, `medium` y `large`, cada una con el valor de un número entero positivo.

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Debe establecer su objeto `royalCrown` como una clave de su mapa `catalog`.

```js
assert.isTrue(catalog.has(royalCrown));
```

El valor de la clave `royalCrown` en el mapa `catalog` debe ser un objeto con las propiedades `small`, `medium` y `large`, cada una con el valor de un número entero positivo.

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Aún debe tener la clave `ballerina` en su mapa `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Aún debe tener la clave `prettyPolly` en su mapa `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Su mapa `catalog` debe contener seis elementos.

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
