---
id: 6734c10a327760a665b7d5b9
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

En los próximos pasos, trabajarás en una función que te permitirá actualizar el catálogo cuando vendas plantas.

Crea una función llamada `sellPlants` que tome tres argumentos: un objeto planta, el tamaño de la maceta y el número de macetas a vender, en este orden.

Cuando el número especificado de macetas sea mayor que las macetas disponibles, haz que tu función retorne `Not enough <pot-size> size pots for <item-name>. Only <pot-number> left.`. Reemplaza `<pot-size>` con el tamaño de la maceta, `<pot-number>` con las macetas restantes de ese tamaño, y `<item-name>` con el nombre científico de la planta seguido de un espacio y el nombre del cultivar entre comillas simples.

Para probar que todo funciona, registra el resultado de llamar a tu función `sellPlants` con el objeto `ballerina`, `"small"`, y `25` en la consola. 

# --hints--

No debes modificar tu mapa de `catalog`.

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 20, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

Debes tener una función llamada `sellPlants`.

```js
assert.isFunction(sellPlants);
```

Tu función `sellPlants` debería tomar tres argumentos.

```js
assert.lengthOf(sellPlants, 3);
```

<code>sellPlants(ballerina, "small", 25)` debería devolver `No hay suficientes macetas de tamaño pequeño para Lavandula stoechas 'Ballerina'. Solo quedan 20.</code>.

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")

assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

Debes registrar `sellPlants(ballerina, "small", 25)` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*25\s*\)\s*\)/);
```

`sellPlants(ballerina, "small", 10)` debería devolver `undefined`.

```js
assert.isUndefined(sellPlants(ballerina, "small", 10));
// prevent hardcoding
assert.isUndefined(sellPlants(hidcote, "large", 5));
assert.isUndefined(sellPlants(royalCrown, "medium", 5));
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog)
```
