---
id: 66069b992c1c5e3451f3deb0
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

La instanciación crea un objeto vacío. El método `__init__` es un método especial que te permite instanciar un objeto en un estado personalizado. Cuando una clase implementa un método `__init__`, `__init__` se llama automáticamente al instanciar.

Dentro de tu clase `Board`, elimina el método `spam` y reemplázalo con un método `__init__` que incluya un parámetro `self`.

# --hints--

No deberías tener un método `spam` en tu clase `Board`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

Debes definir un método `__init__` en tu clase `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

Tu método `__init__` debe tener un parámetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
