---
id: 66068fb0bfddba2b7977eb60
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este proyecto, aprenderás sobre clases y objetos construyendo un solucionador de puzzles de sudoku.

En Python, una clase es una estructura para crear objetos. Los objetos creados a partir de una clase son instancias de esa clase. Puedes crear una clase usando esta sintaxis:

```py
class ClassName:
    pass
```

Donde `class` es la palabra clave requerida para definir la clase y `ClassName` es el nombre de la clase, escrito por convención en *PascalCase*.

Comienza creando una clase `Board`.

# --hints--

Debes crear una clase llamada `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
