---
id: 6606933d6813a8308c962dd1
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Para ser un método de instancia, un método requiere un parámetro especial, llamado `self` por convención. Este parámetro es una referencia a la instancia de la clase y siempre debe ser el primer parámetro.

Agrega un parámetro `self` a tu método `spam`.

# --hints--

Tu método `spam` debe tener un parámetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
