---
id: 66069d65162e61357c793e0c
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

El rompecabezas de sudoku a resolver será una lista de listas, como la siguiente:

```py
[
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]
```

Ten en cuenta que las celdas vacías se llenan con un cero.

Declara una variable `puzzle` y asígnale la lista de listas del ejemplo anterior.

# --hints--

Debes declarar una variable `puzzle` y asignarle la lista 2D proporcionada.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("puzzle").is_equivalent("puzzle = [[0, 0, 2, 0, 0, 8, 0, 0, 0], [0, 0, 0, 0, 0, 3, 7, 6, 2], [4, 3, 0, 0, 0, 0, 8, 0, 0], [0, 5, 0, 0, 3, 0, 0, 9, 0], [0, 4, 0, 0, 0, 0, 0, 2, 6], [0, 0, 0, 4, 6, 7, 0, 0, 0], [0, 8, 6, 7, 0, 4, 0, 0, 0], [0, 0, 0, 5, 1, 9, 0, 0, 8], [1, 7, 0, 0, 0, 6, 0, 0, 5]]")`)) })
```

# --seed--

## --seed-contents--

```py

class Board:
    def __init__(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
gameboard = Board()
```
