---
id: 6606a4641ec48b3a9fe8c2fc
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

La función incorporada `enumerate` toma un iterable como argumento y devuelve un objeto enumerate sobre el cual puedes iterar.  
Proporciona el conteo (que por defecto comienza en cero) y el valor del iterable.

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

El ciclo del ejemplo anterior mostraría las tuplas `0, a`, `1, b` y `2, c`.

Dentro del método `find_empty_cell`, reemplaza `pass` con un bucle `for` que use la función `enumerate()` para iterar sobre cada fila en el tablero de sudoku. Usa `row` como el índice de la fila actual y `contents` para los elementos de la fila actual.


# --hints--

Debes reemplazar `pass` con un ciclo `for` que itere sobre `enumerate(self.board)`. Usa `row` y `contents` como las variables del ciclo.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
