---
id: 660699119472f332798860ad
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Ahora, reemplaza `pass` con una llamada a `print` y pásale la cadena `'Spam!'`.

# --hints--

Debes eliminar `pass` e imprimir la cadena `'Spam!'` dentro del método `spam`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
