---
id: 6606c14182435d4bab0de2ee
title: Paso 30
challengeType: 20
dashedName: step-30
---

# --description--

A continuación, vas a crear un método que verifica si un número puede ser insertado en una columna específica del tablero de sudoku comprobando que el número no esté ya presente en esa columna.

Dentro de la clase `Board`, crea un método llamado `valid_in_col` y dale tres parámetros: `self`, `col` y `num`.  
Donde `col` y `num` son el índice de la columna y el número que se debe verificar, respectivamente.

# --hints--

Debes crear un nuevo método llamado `valid_in_col` dentro de la clase `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_col")`)) })
```

Tu método `valid_in_col` debe tener tres parámetros: `self`, `col` y `num`, en este orden.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").has_args("self, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
