---
id: 6606927d010be4300a4e5330
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

La instanciación crea un objeto vacío. Pero las clases pueden tener métodos, que son como funciones locales para cada instancia. Dentro de una clase, los métodos se declaran de la siguiente manera:

```py
class ClassName:
    def method_name():
        pass
```

Dentro de la clase `Board`, reemplaza `pass` con un método vacío `spam`.

# --hints--

Debes definir un método llamado `spam` dentro de la clase `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
