---
id: 660699aabc59c532f2d556e5
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Para llamar a un método de instancia, necesitas usar la notación de punto:

```py
instance_name.method_name()
```

Donde `instance_name` es la instancia u objeto, y `method_name` es el método que quieres llamar.

Llama al método `spam` del objeto `gameboard`.

# --hints--

Debes llamar al método `spam` del objeto `gameboard` con `gameboard.spam()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
