---
id: 6606cc754a8834509cd0afb6
title: Paso 36
challengeType: 20
dashedName: step-36
---

# --description--

Otra cosa que debes verificar es si se puede insertar un número en un cuadrado 3x3.

Dentro de la clase `Board`, crea un método llamado `valid_in_square` con cuatro parámetros: `self`, `row`, `col` y `num`. Donde `row`, `col` y `num` representan el índice de la fila, el índice de la columna y el número que se debe verificar, respectivamente.

# --hints--

Debes crear un nuevo método llamado `valid_in_square` dentro de la clase `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_square")`)) })
```

Tu método `valid_in_square` debe tener cuatro parámetros: `self`, `row`, `col` y `num`, en este orden.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_square").has_args("self, row, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
