---
id: 6606bf4561f8794a0d345919
title: Paso 27
challengeType: 20
dashedName: step-27
---

# --description--

Es hora de probar el método `valid_in_row`. Llama a `valid_in_row` en `gameboard`. Pásale `0` y `8` como argumentos e imprime el resultado.

Nuevamente, observa cómo el método está definido con tres parámetros, pero se llama con solo dos argumentos porque `self` se pasa automáticamente como el objeto sobre el cual se llama el método.

# --hints--

Debes imprimir el resultado de llamar a `valid_in_row(0, 8)` en `gameboard`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.valid_in_row(0, 8))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
