---
id: 6606bc4e5535c0484990ccd5
title: Pasos 24
challengeType: 20
dashedName: step-24
---

# --description--

A continuación, vas a trabajar en un método que verifica si un número dado puede ser insertado en una fila específica del tablero de sudoku.

Dentro de la clase `Board`, crea un método llamado `valid_in_row` y dale tres parámetros: `self`, `row` y `num`. Donde `self` representa la instancia de la clase, y `row` y `num` son el índice de la fila y el número que se debe verificar, respectivamente.


# --hints--

Debes crear un nuevo método llamado `valid_in_row` dentro de la clase `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

Tu método `valid_in_row` debe tener tres parámetros: `self`, `row` y `num`, en este orden.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
