---
id: 6606a219f9efbf38ad496f67
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

También puedes usar la notación de puntos para acceder a un atributo de instancia.

Fuera de la clase `Board`, después de inicializar el objeto `gameboard`, usa `gameboard.board` para acceder al atributo `board` de tu objeto `gameboard` e imprimir el resultado en la pantalla.

# --hints--

Debes imprimir `gameboard.board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.board)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
