---
id: 66069f86f58f85371d47123e
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Un atributo es una variable asociada con un objeto, que se usa para almacenar datos como variables normales.

Dentro del método `__init__`, asigna el parámetro `board` (que se pasa al crear una instancia de la clase `Board`) a un atributo de instancia `board` usando `self.board`.

`self.board` se refiere al atributo `board` de la instancia de la clase. Es una variable que pertenece al objeto creado a partir de la clase `Board`.

# --hints--

Debes eliminar `pass` y asignar el parámetro `board` a `self.board` dentro del método `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
