---
id: 66069167b3307b2f4067b22b
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Una nueva instancia de una clase se crea usando la notación de función, que consiste en añadir un par de paréntesis al nombre de la clase.

Fuera de la definición de la clase, crea una instancia de la clase `Board` y asígnala a una variable llamada `gameboard`.

# --hints--

Debes declarar una variable `gameboard` fuera de la clase `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

Tu variable `gameboard` debe tener el valor de `Board()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
