---
id: 6606cd69f56e27516583b0cc
title: Paso 37
challengeType: 20
dashedName: step-37
---

# --description--

Ahora necesitas calcular el índice de fila inicial para el cuadrado 3x3 dentro de la cuadrícula del tablero y asegurarte de que el índice de fila inicial para cada cuadrado 3x3 sea un múltiplo de 3.

Esto se puede lograr tomando el resultado de la división entera `row // 3` multiplicado por `3`. Reemplaza `pass` con una variable `row_start` y asígnale `(row // 3) * 3`.


# --hints--

Debes eliminar `pass` y declarar una variable `row_start` con el valor de `(row // 3) * 3`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_square").find_body().is_equivalent("row_start = (row//3)*3")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))
--fcc-editable-region--
    def valid_in_square(self, row, col, num):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
