---
id: 6606b224a69a293f98f8db8f
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

El método `.index()` lanza una excepción `ValueError` cuando no se encuentra el valor. Para evitar que el programa detenga la ejecución, anidarás esta línea de código dentro de un bloque `try`. La sentencia `try` se usa para encapsular código que podría lanzar una excepción. La cláusula `except`, por otro lado, ofrece código alternativo para ejecutar si ocurre una excepción:

```py
try:
    <code>
except:
    <code>
```

Pon la asignación de `col` dentro de un bloque `try`. Luego, crea una cláusula `except` y llena su cuerpo con `pass`.

# --hints--

Debes poner la asignación de `col` dentro de un bloque `try` y crear una cláusula `except`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
