---
id: a8e512fbe388ac2f9198f0fa
title: 實作匹配物件篩選器
challengeType: 26
dashedName: implement-a-matching-object-filter
---

# --description--

在本實驗中，你將創建一個函式（程式），用來篩選一個物件陣列，並且只傳回符合給定原始碼物件中所有鍵值對組的物件。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `whatIsInAName` 函式（程式），它接受兩個引數，一個物件陣列和一個原始碼物件。
1. `whatIsInAName` 函式（程式）應該傳回一個新的陣列，該陣列只包含群集中所有擁有原始碼物件中所有鍵值配對的物件。
1. 如果沒有任何物件符合來自 `source` 的所有鍵–值配對，該函式應該傳回一個空的陣列。 例如：

```js
whatIsInAName(
  [
    { first: "Romeo", last: "Montague" },
    { first: "Mercutio", last: null },
    { first: "Tybalt", last: "Capulet" }
  ],
  { last: "Capulet" }
);
// should return [{ first: "Tybalt", last: "Capulet" }]
```

# --hints--

你應該有一個 `whatIsInAName` 函式（程式）。

```js
assert.isFunction(whatIsInAName);
```

`whatIsInAName([{ first: "Romeo", last: "Montague" }, { first: "Mercutio", last: null }, { first: "Tybalt", last: "Capulet" }], { last: "Capulet" })` 應該傳回 `[{ first: "Tybalt", last: "Capulet" }]`。

```js
assert.deepEqual(
  whatIsInAName(
    [{ first: "Romeo", last: "Montague" }, { first: "Mercutio", last: null }, { first: "Tybalt", last: "Capulet" }],
    { last: "Capulet" }
  ),
  [{ first: "Tybalt", last: "Capulet" }]
);
```

`whatIsInAName([{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }], { "apple": 1 })` 應該傳回 `[{"apple": 1}, {"apple": 1}, {"apple": 1, "bat": 2}]`。

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }],
    { "apple": 1 }
  ),
  [{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }], { "apple": 1, "bat": 2 })` 應該傳回 `[{"apple": 1, "bat": 2}, {"apple": 1, "bat": 2, "cookie": 2}]`。

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }],
    { "apple": 1, "bat": 2 }
  ),
  [{ "apple": 1, "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }], { "apple": 1, "cookie": 2 })` 應該傳回 `[{"apple": 1, "bat": 2, "cookie": 2}]`。

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }],
    { "apple": 1, "cookie": 2 }
  ),
  [{ "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }, { "bat": 2 }], { "apple": 1, "bat": 2 })` 應該傳回 `[{"apple": 1, "bat": 2}, {"apple": 1, "bat": 2, "cookie": 2}]`

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }, { "bat": 2 }],
    { "apple": 1, "bat": 2 }
  ),
  [{ "apple": 1, "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{"a": 1, "b": 2, "c": 3}], {"a": 1, "b": 9999, "c": 3})` 應該傳回 `[]`。

```js
assert.deepEqual(
  whatIsInAName([{ "a": 1, "b": 2, "c": 3 }], { "a": 1, "b": 9999, "c": 3 }),
  []
);
```

`whatIsInAName([{"a": 1, "b": 2, "c": 3, "d": 9999}], {"a": 1, "b": 9999, "c": 3})` 應該傳回 `[]`。

```js
assert.deepEqual(
  whatIsInAName([{ "a": 1, "b": 2, "c": 3, "d": 9999 }], { "a": 1, "b": 9999, "c": 3 }),
  []
);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function whatIsInAName(collection, source) {
  const sourceKeys = Object.keys(source);

  return collection.filter(obj =>
    sourceKeys.every(key => obj.hasOwnProperty(key) && obj[key] === source[key])
  );
}
```
