---
id: 68507cadbf36aa089a84ce2d
title: Luhnアルゴリズムを実装する
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

Luhnアルゴリズムは、「モジュラス10」または「mod 10」アルゴリズムとも呼ばれ、クレジットカード番号のようなさまざまな識別番号の検証に使われるシンプルなチェックサムの計算式です。Luhnアルゴリズムを使って番号を検証する手順は以下の通りです：

- 右端の数字（チェックディジット）を除き、右から始めて1つおきの数字の値を2倍にします。
- 2倍した結果が`9`より大きい場合は、各桁の数字を足して1桁の数字にします。あるいは、結果から`9`を引いてもかまいません。
- チェックディジットを含むすべての数字の合計を求めます。
- すべての数字の合計が`10`の倍数であれば、その番号は有効です。そうでなければ無効です。

例えば、番号`453914881`があるとします。Luhnアルゴリズムで検証する手順は以下のようになります：

 ```md 
 口座番号           4   5   3   9   1   4   8   8   1 
 1つおきに2倍      4  10   3  18   1   8   8  16   1 
 2桁の数字を合計    4   1   3   9   1   8   8   7   1 
 ``` 

- その後、すべての数字を合計し、`4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`となります。
- `42`は`10`の倍数ではないため、この番号は無効です。

このラボでは、Luhnアルゴリズムを使ったクレジットカードの検証機能を作成します。

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `verify_card_number`という名前の関数を定義し、数字の文字列（カード番号を表す）を受け取り、Luhnアルゴリズムに従って有効かどうかを検証してください。

2. `verify_card_number`関数内で：

   - 渡されたカード番号に含まれるダッシュやスペースを処理してください。
   - カード番号が有効な場合は`VALID!`を返し、そうでなければ`INVALID!`を返してください。

プロジェクトを完成させると、入力に応じて以下のメッセージが表示されます：

| カード番号           | メッセージ  |
| ------------------- | -------- |
| 453914889           | VALID!   |
| 4111-1111-1111-1111 | VALID!   |
| 1234 5678 9012 3456 | INVALID! |


# --hints--

`verify_card_number`という名前の関数を用意してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')`は`VALID!`を返すべきです。

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')`は`VALID!`を返すべきです。

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')`は`INVALID!`を返すべきです。

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')`は`INVALID!`を返すべきです。

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

`verify_card_number`は有効なクレジットカード番号で呼び出されたときに`VALID!`を返すべきです。

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

`verify_card_number`は無効なクレジットカード番号で呼び出されたときに`INVALID!`を返すべきです。

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
