---
id: 68e6a99a0fd7386fb66e156f
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, practicarás el estilo de casillas de verificación construyendo un panel de configuración del juego.

Todo el boilerplate de HTML (`DOCTYPE`, `html`, `head` y `body`) ha sido proporcionado para ti.

Dentro del `body`, crea un elemento `div` con un elemento `h1` anidado dentro. Dale al elemento `h1` el texto `Game Settings`

# --hints--

Debe tener un elemento `div`.

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

Debería tener un elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Tu elemento `h1` debe estar dentro de tu elemento `div`.

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

Tu elemento `h1` debe tener el texto `Game Settings`.

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
