---
id: 68e9a834dfef752dc35f6f90
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Ahora que tus casillas de verificación cambian de color cuando el usuario hace clic en ellas, vas a aplicar algo de estilo para que sea un poco más obvio que la casilla ha sido seleccionada.

Las pseudo-clases de CSS pueden encadenarse. Puedes agregar una pseudo-clase o pseudo-elemento después de otra para dirigirte a elementos que cumplan múltiples condiciones.

Crea un nuevo selector para `input[type="checkbox"]`. Agrega la pseudo-clase `:checked` seguida del pseudo-elemento `::after`. Esto te permite añadir un indicador visual cuando el checkbox está marcado.

A continuación, agrega `content` con el valor de `"✓"`.

# --hints--

Debes tener un selector de tipo combinado con un selector de pseudo-clase y un selector de pseudo-elemento de `input[type="checkbox"]:checked::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

Debes agregar una propiedad `content` a tu selector de tipo combinado con un selector de pseudo-clase y un selector de pseudo-elemento de `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

Debes tener una propiedad `content` con un valor de `"✓"`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```

