---
id: 68e9a3194160f8f5620ddf21
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Inmediatamente después de tu elemento `h1`, crea cuatro elementos `label`.

Dentro de cada uno de los elementos `label`, crea un elemento `input` con el atributo `type` establecido en `"checkbox"`.

Después de cada elemento `input`, da a tus elementos `label` los siguientes textos, en orden: `Sound Effects`, `Background Music`, `Hard Mode` y `Haptic Feedback`.

# --hints--

Debes tener cuatro elementos `label`.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Tu primer elemento `label` debe tener un elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

Tu primer elemento `input` debe tener el atributo `type` configurado en `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

Tu primer elemento `label` debe tener el texto `Sound Effects`.

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

El texto de tus `Sound Effects` debe ir después del elemento `input`.

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

Tu segundo elemento `label` debe tener un elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

Tu segundo elemento `input` debe tener el atributo `type` configurado como `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

Tu segundo elemento `label` debe tener el texto `Background Music`.

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

Tu texto de `Background Music` debe ir después del elemento `input`.

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

Tu tercer elemento `label` debe tener un elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

Tu tercer elemento `input` debe tener el atributo `type` configurado como `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

Tu tercer elemento `label` debe tener el texto `Hard Mode`.

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

Tu texto de `Hard Mode` debe ir después del elemento `input`.

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

Tu cuarto elemento `label` debe tener un elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

Tu cuarto elemento `input` debe tener el atributo `type` configurado como `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

Tu cuarto elemento `label` debe tener el texto `Haptic Feedback`.

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

Tu texto de `Haptic Feedback` debe ir después del elemento `input`.

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
        <h1>Game Settings</h1>
--fcc-editable-region--

--fcc-editable-region--
    </div>
  </body>
</html>
```
