---
id: 68e9a6201e2efa142d62bfe1
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

A continuación, vas a agrandar las casillas de verificación para una mejor visibilidad.

Comienza configurando un selector para `input`, pero apuntando específicamente a tu `[type="checkbox"]`.

Dentro de tu selector, establece el `width` y el `height` en `20px`. Esto lo hace más grande que antes.

Y para ajustarte a tu configuración de `cursor` que se estableció en las etiquetas, agrega `cursor` y asígnale el valor de `pointer`. Después de eso, cuando pases el cursor sobre las casillas de verificación, mostrará un puntero.

# --hints--

Debes tener un selector en tu CSS de `input[type="checkbox"]`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

Debes agregar una propiedad `width` a tu selector `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

Debes tener una propiedad `width` con un valor de `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

Debes agregar una propiedad `height` a tu selector `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

Debes tener una propiedad `height` con un valor de `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

Debes agregar una propiedad `cursor` a tu selector `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

Debes tener una propiedad `cursor` con un valor de `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
