---
id: 68e9a7bd60c971278d2d4ffc
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

En una lección anterior, aprendiste sobre pseudo-clases y pseudo-elementos en CSS. Ahora vas a aplicar ese conocimiento creando un selector de tipo combinado con un selector de pseudo-clase.

Primero, comienza con el formato para el selector de tipo configurándolo como `input[type="checkbox"]` y añadiendo `:checked` al final.

Luego, dale un `background-color` con el valor de `#f1be32`.

Y por último, establece el `border-color` para que tenga un valor de `#e2a60d`.

# --hints--

Debes tener un selector combinado de tipo y pseudo-clase de `input[type="checkbox"]:checked`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked'));
```

Debes agregar una propiedad `background-color` a tu selector combinado de tipo y pseudo-clase `input[type="checkbox"]:checked`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor);
```

Debes tener una propiedad `background-color` con un valor de `#f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor, "rgb(241, 190, 50)");
```

Debes agregar una propiedad `border-color` a tu selector combinado de tipo y pseudo-clase `input[type="checkbox"]:checked`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor);
```

Debes tener una propiedad `border-color` con un valor de `#e2a60d`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor, "rgb(226, 166, 13)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

--fcc-editable-region--

--fcc-editable-region--

```
