---
id: 68ee86e3ee07b01626a392cc
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Dentro de tu selector `.settings-card` establece la propiedad `margin` en `auto`.

Configurar la propiedad `margin` en `auto` ajusta automáticamente los márgenes de un elemento para distribuir de manera uniforme el espacio restante en su contenedor, comúnmente usado para centrar elementos de nivel bloque horizontalmente.

Y por último, establece una propiedad `text-align` con el valor `left`. Esto alineará el contenido en línea, como el texto, al lado izquierdo de su elemento contenedor.

# --hints--

Debes agregar una propiedad `margin` al selector de clase `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

Debes tener una propiedad `margin` con un valor de `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

Debes agregar una propiedad `text-align` al selector de clase `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

Debes tener una propiedad `text-align` con un valor de `left`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
