---
id: 68f15ef0446464325ce1aa45
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Necesitas un poco de espacio entre las casillas de verificación y las etiquetas. Comienza usando el selector para `label` en tu CSS y establece una propiedad `display` de `block`.

A continuación, establece una propiedad `margin` de `8px auto` para añadir espacio vertical entre los elementos. Esto creará un espacio consistente arriba y abajo de cada elemento, ayudando a separar el contenido y mejorar la legibilidad.

Finalmente, establece la propiedad `cursor` en `pointer` en los elementos `label`. Esto cambiará el cursor a un ícono de mano cuando un usuario pase el cursor sobre un `label`, señalando que el elemento es clickeable y mejorando la experiencia general del usuario.

# --hints--

Debes tener un selector de tipo en tu CSS para `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

Debes agregar una propiedad `display` a tu selector `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

Debes tener una propiedad `display` con un valor de `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

Debes agregar la propiedad `margin` a tu selector `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

Debes tener una propiedad `margin` con un valor de `8px auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

Debes agregar una propiedad `cursor` a tu selector `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

Debes tener una propiedad `cursor` con un valor de `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
