---
id: 68f1560aff1214103ab6b59b
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Ahora comenzarás a embellecer la página con algo de estilo CSS. Empieza creando el selector `body`.

Establece el `body` para que tenga una propiedad `height` con un valor de `100vh` y una propiedad `background-color` con un valor de `#f0f0f0`.

La `height` de `100vh` hace que el `body` ocupe toda la altura de la ventana del navegador, mientras que el color de fondo gris claro proporciona una base sutil para la página.

Por último, establece una propiedad `text-align` con el valor `center`. Esto centrará todo el contenido en línea contenido dentro de la página a menos que un elemento hijo lo sobrescriba con su propia alineación.

# --hints--

Debes usar el selector de tipo `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Debes agregar la propiedad `height` dentro del selector de tipo `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

Debes tener una propiedad `height` con el valor de `100vh`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

Debes agregar la propiedad `background-color` dentro del selector de tipo `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

Debes tener una propiedad `background-color` con un valor de `#f0f0f0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

Debes agregar la propiedad `text-align` dentro del selector de tipo `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

Debes tener una propiedad `text-align` con el valor `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
