---
id: 68f15cd397503429f0345262
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

En la declaración para el selector `input[type="checkbox"]:checked::after`, establece el `display` en `block`. Establecer tu propiedad `display` como `block` hace que el elemento sea un elemento de nivel bloque, lo que significa que ocupa todo el ancho de su contenedor y comienza en una nueva línea, permitiéndote controlar su ancho, alto y espaciado más fácilmente.

A continuación, configura `text-align` a `center`. Esto centrará el contenido en línea (como texto o elementos en línea) horizontalmente dentro del bloque. El bloque seguirá ocupando todo el ancho de su contenedor, pero todo lo que esté dentro estará alineado al centro.

# --hints--

Debes agregar una propiedad `display` a tu selector de tipo combinado con un selector de pseudo-clase y un selector de pseudo-elemento de `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

Debes tener una propiedad `display` con un valor de `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

Debes agregar una propiedad `text-align` a tu selector de tipo combinado con un selector de pseudo-clase y un selector de pseudo-elemento de `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

Debes tener una propiedad `text-align` con un valor de `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
